---
id: 596fd036dc1ab896c5db98b1
title: تحويل الثواني إلى مدة مركبة
challengeType: 1
forumTopicId: 302236
dashedName: convert-seconds-to-compound-duration
---

# --description--

نفذ دالة تقوم بما يلي:

<ul>
<li>تأخذ عددًا صحيحًا موجبًا يمثل مدة بالثواني كمُدخَل (مثلًا، <code>100</code>)، و</li>
  <li>تعيد سلسلة نصية تُظهر نفس المدة مقسمة إلى أسابيع وأيام وساعات ودقائق وثوانٍ كما هو موضح أدناه (مثلًا، <code>1 min, 40 sec</code>).</li>
</ul>

أثبت أنها تجتاز الحالات الثلاث التالية للاختبار:

<div style='font-size:115%; font-weight: bold;'>حالات الاختبار</div>

| الرقم المُدخل | الرقم الناتج                         |
| ------------ | ------------------------------------- |
| 7259         | <code>2 hr, 59 sec</code>             |
| 86400        | <code>1 d</code>                      |
| 6000000      | <code>9 wk, 6 d, 10 hr, 40 min</code> |

<div style="font-size:115%; font-weight: bold;">التفاصيل</div>
<ul>
  <li>
يجب استخدام الوحدات الخمس التالية:

| الوحدة  | اللاحقة المستخدمة في الناتج | التحويل              |
| ------ | --------------------------- | --------------------- |
| week   | <code>wk</code>             | 1 أسبوع = 7 أيام      |
| day    | <code>d</code>              | 1 يوم = 24 ساعة       |
| hour   | <code>hr</code>             | 1 ساعة = 60 دقيقة     |
| minute | <code>min</code>            | 1 دقيقة = 60 ثانية    |
| second | <code>sec</code>            | ---                   |

  </li>
  <li>
مع ذلك، أدرج <strong>فقط</strong> الكميات التي قيمتها غير صفرية في الناتج (مثلًا، أعد <code>1 d</code> وليس <code>0 wk, 1 d, 0 hr, 0 min, 0 sec</code>).
  </li>
  <li>
أعطِ الأولوية للوحدات الأكبر على الأصغر قدر الإمكان (مثلًا، أعد <code>2 min, 10 sec</code> وليس <code>1 min, 70 sec</code> أو <code>130 sec</code>).
  </li>
  <li>
قلّد تنسيق الحالات الاختبارية المعروضة (ترتيب الكميات من الأكبر إلى الأصغر مفصولة بفاصلة ومسافة؛ والقيمة والوحدة لكل كمية مفصولان بمسافة).
  </li>
</ul>

# --before-each--

```js
const testCases = [7259, 86400, 6000000];
const results = ['2 hr, 59 sec', '1 d', '9 wk, 6 d, 10 hr, 40 min'];
```

# --hints--

يجب أن تكون `convertSeconds` دالة.

```js
assert(typeof convertSeconds === 'function');
```

يجب أن تعيد `convertSeconds(7259)` القيمة `2 hr, 59 sec`.

```js
assert.equal(convertSeconds(testCases[0]), results[0]);
```

يجب أن تعيد `convertSeconds(86400)` القيمة `1 d`.

```js
assert.equal(convertSeconds(testCases[1]), results[1]);
```

يجب أن تعيد `convertSeconds(6000000)` القيمة `9 wk, 6 d, 10 hr, 40 min`.

```js
assert.equal(convertSeconds(testCases[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function convertSeconds(sec) {

  return true;
}
```

# --solutions--

```js
function convertSeconds(sec) {
  const localNames = ['wk', 'd', 'hr', 'min', 'sec'];
  // compoundDuration :: [String] -> Int -> String
  const compoundDuration = (labels, intSeconds) =>
    weekParts(intSeconds)
    .map((v, i) => [v, labels[i]])
    .reduce((a, x) =>
      a.concat(x[0] ? [`${x[0]} ${x[1] || '?'}`] : []), []
    )
    .join(', ');

    // weekParts :: Int -> [Int]
  const weekParts = intSeconds => [0, 7, 24, 60, 60]
    .reduceRight((a, x) => {
      const r = a.rem;
      const mod = x !== 0 ? r % x : r;

      return {
        rem: (r - mod) / (x || 1),
        parts: [mod].concat(a.parts)
      };
    }, {
      rem: intSeconds,
      parts: []
    })
    .parts;

  return compoundDuration(localNames, sec);
}
```
