---
id: 597089c87eec450c68aa1643
title: تعريف نوع بيانات بدائي
challengeType: 1
forumTopicId: 302248
dashedName: define-a-primitive-data-type
---

# --description--

عرّف نوعًا يتصرف مثل عدد صحيح لكن له أقل قيمة صالحة هي 1 وأعلى قيمة صالحة هي 10.

معالجة الأخطاء:

<ul>
<li>إذا حاولت إنشاء كائن من <code>Num</code> بقيمة خارج النطاق 1 - 10، يجب أن يرمي <code>TypeError</code> مع رسالة خطأ <code>'Out of range'</code>.</li>
  <li>إذا حاولت إنشاء كائن من <code>Num</code> بقيمة ليست رقمًا، يجب أن يرمي <code>TypeError</code> مع رسالة خطأ <code>'Not a Number'</code>.</li>
</ul>

# --hints--

يجب أن تكون `Num` دالة.

```js
assert(typeof Num === 'function');
```

يجب أن تُرجع `new Num(4)` كائنًا.

```js
assert(typeof new Num(4) === 'object');
```

يجب أن يرمي `new Num('test')` TypeError مع رسالة 'Not a Number'.

```js
assert.throws(() => new Num('test'), TypeError);
```

يجب أن يرمي `new Num(0)` TypeError مع رسالة 'Out of range'.

```js
assert.throws(() => new Num(0), TypeError);
```

يجب أن يرمي `new Num(-5)` TypeError مع رسالة 'Out of range'.

```js
assert.throws(() => new Num(-5), TypeError);
```

يجب أن يرمي `new Num(10)` TypeError مع رسالة 'Out of range'.

```js
assert.throws(() => new Num(11), TypeError);
```

يجب أن يرمي `new Num(20)` TypeError مع رسالة 'Out of range'.

```js
assert.throws(() => new Num(20), TypeError);
```

يجب أن يساوي `new Num(3) + new Num(4)` القيمة 7.

```js
assert.equal(new Num(3) + new Num(4), 7);
```

يجب أن يساوي `new Num(3) - new Num(4)` القيمة -1.

```js
assert.equal(new Num(3) - new Num(4), -1);
```

يجب أن يساوي `new Num(3) * new Num(4)` القيمة 12.

```js
assert.equal(new Num(3) * new Num(4), 12);
```

يجب أن يساوي `new Num(3) / new Num(4)` القيمة 0.75.

```js
assert.equal(new Num(3) / new Num(4), 0.75);
```

يجب أن تكون نتيجة `new Num(3) < new Num(4)` صحيحة.

```js
assert(new Num(3) < new Num(4));
```

يجب أن تكون نتيجة `new Num(3) > new Num(4)` خاطئة.

```js
assert(!(new Num(3) > new Num(4)));
```

يجب أن تُرجع `(new Num(5)).toString()` السلسلة النصية '5'.

```js
assert.equal(new Num(5).toString(), '5');
```

# --seed--

## --seed-contents--

```js
function Num(n) {

  return n;
}
```

# --solutions--

```js
function Num(n) {
  if (isNaN(n)) {
    throw new TypeError('Not a Number');
  }
  if (n < 1 || n > 10) {
    throw new TypeError('Out of range');
  }

  this._value = +n;
}
Num.prototype.valueOf = function() { return this._value; };
Num.prototype.toString = function () { return this._value.toString(); };
```
