---
id: 59713bd26bdeb8a594fb9413
title: عد العملات المعدنية
challengeType: 1
forumTopicId: 302238
dashedName: count-the-coins
---

# --description--

هناك أربعة أنواع من العملات المعدنية الشائعة في العملة الأمريكية:

<ul>
<li>ربع دولار (25 سنتًا)</li>
  <li>عشر سنتات (10 سنتات)</li>
  <li>نيكل (5 سنتات)، و</li>
  <li>بنس (1 سنت)</li>
</ul>

<p>هناك ست طرق لصرف 15 سنتًا:</p>

<ul>
<li>عشر سنتات ونيكل</li>
  <li>عشر سنتات و5 بنسات</li>
  <li>3 نيكلات</li>
  <li>2 نيكل و5 بنسات</li>
  <li>نيكل و10 بنسات</li>
  <li>15 بنسًا</li>
</ul>

# --instructions--

نفذ دالة لتحديد عدد الطرق لصرف مبلغ مُدخَل، `cents`، الذي يمثل كمية من البنسات الأمريكية باستخدام هذه العملات المعدنية الشائعة.

# --hints--

يجب أن تكون `countCoins` دالة.

```js
assert(typeof countCoins === 'function');
```

يجب أن تُرجع `countCoins(15)` القيمة `6`.

```js
assert.equal(countCoins(15), 6);
```

يجب أن تُرجع `countCoins(85)` القيمة `163`.

```js
assert.equal(countCoins(85), 163);
```

يجب أن تُرجع `countCoins(100)` القيمة `242`.

```js
assert.equal(countCoins(100), 242);
```

# --seed--

## --seed-contents--

```js
function countCoins(cents) {

  return true;
}
```

# --solutions--

```js
function countCoins(cents) {
  const operands = [1, 5, 10, 25];
  const targetsLength = cents + 1;
  const operandsLength = operands.length;
  const t = [1];

  for (let a = 0; a < operandsLength; a++) {
    for (let b = 1; b < targetsLength; b++) {
      // initialise undefined target
      t[b] = t[b] ? t[b] : 0;

      // accumulate target + operand ways
      t[b] += (b < operands[a]) ? 0 : t[b - operands[a]];
    }
  }

  return t[targetsLength - 1];
}
```
