---
id: 597b2b2a2702b44414742771
title: مضروب
challengeType: 1
forumTopicId: 302263
dashedName: factorial
---

# --description--

اكتب دالة تُرجع مضروب رقم.

يُعطى مضروب الرقم بالعلاقة التالية:

<pre><big>n! = n * (n-1) * (n-2) * ..... * 1</big>
</pre>

على سبيل المثال:

<ul>
  <li><code>3! = 3 * 2 * 1 = 6</code></li>
  <li><code>4! = 4 * 3 * 2 * 1 = 24</code></li>
</ul>

**ملاحظة:** `0! = 1`

# --hints--

يجب أن تكون `factorial` دالة.

```js
assert(typeof factorial === 'function');
```

يجب أن تُرجع `factorial(2)` رقمًا.

```js
assert(typeof factorial(2) === 'number');
```

يجب أن تُرجع `factorial(3)` القيمة 6.

```js
assert.equal(factorial(3), 6);
```

يجب أن تُرجع `factorial(5)` القيمة 120.

```js
assert.equal(factorial(5), 120);
```

يجب أن تُرجع `factorial(10)` القيمة 3,628,800.

```js
assert.equal(factorial(10), 3628800);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

}
```

# --solutions--

```js
function factorial(n) {
  let sum = 1;
  while (n > 1) {
    sum *= n;
    n--;
  }
  return sum;
}
```
