---
id: 597f1e7fbc206f0e9ba95dc4
title: عوامل عدد صحيح
challengeType: 1
forumTopicId: 302265
dashedName: factors-of-an-integer
---

# --description--

اكتب دالة تُرجع عوامل عدد صحيح موجب كمصفوفة.

هذه العوامل هي الأعداد الصحيحة الموجبة التي يُمكن قسمة العدد المراد تحليل عوامله عليها للحصول على نتيجة عدد صحيح موجب.

# --before-each--

```js
const ans=[[1,3,5,9,15,45],[1,53],[1,2,4,8,16,32,64]];
```

# --hints--

`factors` يجب أن تكون دالة.

```js
assert(typeof factors === 'function');
```

`factors(45)` يجب أن تُرجع `[1,3,5,9,15,45]`.

```js
assert.deepEqual(factors(45), ans[0]);
```

`factors(53)` يجب أن تُرجع `[1,53]`.

```js
assert.deepEqual(factors(53), ans[1]);
```

`factors(64)` يجب أن تُرجع `[1,2,4,8,16,32,64]`.

```js
assert.deepEqual(factors(64), ans[2]);
```

# --seed--

## --seed-contents--

```js
function factors(num) {

}
```

# --solutions--

```js
function factors(num)
{
 let n_factors = [], i, sqr=Math.floor(Math.sqrt(num));

 for (i = 1; i <=sqr ; i += 1)
  if (num % i === 0)
  {
   n_factors.push(i);
   if (num / i !== i)
    n_factors.push(num / i);
  }
 n_factors.sort(function(a, b){return a - b;});
 return n_factors;
}
```
