---
id: 597f24c1dda4e70f53c79c81
title: متتالية فيبوناتشي
challengeType: 1
forumTopicId: 302268
dashedName: fibonacci-sequence
---

# --description--

اكتب دالة لتوليد العدد <code>n<sup>th</sup></code> في متتالية فيبوناتشي.

العدد <code>n<sup>th</sup></code> في متتالية فيبوناتشي يُعطى بالعلاقة:

<code>F<sub>n</sub> = F<sub>n-1</sub> + F<sub>n-2</sub></code>

أول حدين في المتتالية هما 0 و 1.

لذا، المتتالية هي: 0، 1، 1، 2، 3، 5، 8، 13...

# --hints--

يجب أن تكون `fibonacci` دالة.

```js
assert(typeof fibonacci === 'function');
```

يجب أن تُرجع `fibonacci(2)` رقمًا.

```js
assert(typeof fibonacci(2) == 'number');
```

يجب أن تُرجع `fibonacci(3)` القيمة 2.

```js
assert.equal(fibonacci(3), 2);
```

يجب أن تُرجع `fibonacci(5)` القيمة 5.

```js
assert.equal(fibonacci(5), 5);
```

يجب أن تُرجع `fibonacci(10)` القيمة 55.

```js
assert.equal(fibonacci(10), 55);
```

# --seed--

## --seed-contents--

```js
function fibonacci(n) {

}
```

# --solutions--

```js
function fibonacci(n) {
  let a = 0, b = 1, t;
  while (--n >= 0) {
    t = a;
    a = b;
    b += t;
  }
  return a;
}
```
