---
id: 59880443fb36441083c6c20e
title: طريقة أويلر
challengeType: 1
forumTopicId: 302258
dashedName: euler-method
---

# --description--

تُقرب طريقة أويلر عدديًا حلول المعادلات التفاضلية العادية من الدرجة الأولى (ODEs) بقيمة ابتدائية معطاة. هي طريقة صريحة لحل مشكلات القيمة الابتدائية (IVPs)، كما هو موضح في <a href="https://www.freecodecamp.org/news/eulers-method-explained-with-examples/" title="Euler's Method Explained with Examples" target="_blank" rel="noopener noreferrer nofollow">هذا المقال</a>.

يجب تقديم المعادلة التفاضلية العادية بالشكل التالي:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt} = f(t,y(t))$</big></li>
</ul>

بقيمة ابتدائية

<ul style='list-style: none;'>
  <li><big>$y(t_0) = y_0$</big></li>
</ul>

للحصول على حل عددي، نستبدل المشتقة في الجانب الأيسر بتقريب الفرق المحدود:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt}  \approx \frac{y(t+h)-y(t)}{h}$</big></li>
</ul>

ثم نحل لـ $y(t+h)$:

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, \frac{dy(t)}{dt}$</big></li>
</ul>

وهو نفسه

<ul style='list-style: none;'>
<li><big>$y(t+h) \approx y(t) + h \, f(t,y(t))$</big></li>
</ul>

قاعدة الحل التكراري هي:

<ul style='list-style: none;'>
  <li><big>$y_{n+1} = y_n + h \, f(t_n, y_n)$</big></li>
</ul>

حيث $h$ هو حجم الخطوة، وهو المعلمة الأكثر أهمية لدقة الحل. حجم خطوة أصغر يزيد الدقة لكنه يرفع تكلفة الحساب، لذلك يجب اختياره يدويًا حسب المشكلة المطروحة.

**مثال: قانون تبريد نيوتن**

يصف قانون تبريد نيوتن كيف يبرد جسم ذو درجة حرارة ابتدائية $T(t_0) = T_0$ في بيئة ذات درجة حرارة $T_R$:

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, \Delta T$</big></li>
</ul>

أو

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, (T(t) - T_R)$</big></li>
</ul>

يقول إن معدل تبريد الجسم $\\frac{dT(t)}{dt}$ يتناسب مع فرق درجة الحرارة الحالي $\\Delta T = (T(t) - T_R)$ مع البيئة المحيطة.

الحل التحليلي، الذي سنقارنه بالتقريب العددي، هو

<ul style='list-style: none;'>
  <li><big>$T(t) = T_R + (T_0 - T_R) \; e^{-k t}$</big></li>
</ul>

# --instructions--

نفذ روتين طريقة أويلر ثم استخدمه لحل المثال المعطى لقانون تبريد نيوتن بثلاثة أحجام خطوات مختلفة:

<ul>
<li><code>2 s</code></li>
  <li><code>5 s</code> و</li>
  <li><code>10 s</code></li>
</ul>

وقارن مع الحل التحليلي.

**القيم الابتدائية:**

<ul>
<li>درجة الحرارة الابتدائية <big>$T_0$</big> يجب أن تكون <code>100 °C</code></li>
  <li>درجة حرارة الغرفة <big>$T_R$</big> يجب أن تكون <code>20 °C</code></li>
  <li>ثابت التبريد <big>$k$</big> يجب أن يكون <code>0.07</code></li>
  <li>فترة الزمن للحساب يجب أن تكون من <code>0 s</code> إلى <code>100 s</code></li>
</ul>  

المعلمة الأولى للدالة هي الوقت الابتدائي، الثانية هي درجة الحرارة الابتدائية، الثالثة هي الزمن المنقضي والرابعة هي حجم الخطوة.

# --hints--

`eulersMethod` يجب أن تكون دالة.

```js
assert(typeof eulersMethod === 'function');
```

`eulersMethod(0, 100, 100, 2)` يجب أن تُرجع رقمًا.

```js
assert(typeof eulersMethod(0, 100, 100, 2) === 'number');
```

`eulersMethod(0, 100, 100, 2)` يجب أن تُرجع 20.0424631833732.

```js
assert.equal(eulersMethod(0, 100, 100, 2), 20.0424631833732);
```

`eulersMethod(0, 100, 100, 5)` يجب أن تُرجع 20.01449963666907.

```js
assert.equal(eulersMethod(0, 100, 100, 5), 20.01449963666907);
```

`eulersMethod(0, 100, 100, 10)` يجب أن تُرجع 20.000472392.

```js
assert.equal(eulersMethod(0, 100, 100, 10), 20.000472392);
```

# --seed--

## --seed-contents--

```js
function eulersMethod(x1, y1, x2, h) {

}
```

# --solutions--

```js
function eulersMethod(x1, y1, x2, h) {
  let x = x1;
  let y = y1;

  while ((x < x2 && x1 < x2) || (x > x2 && x1 > x2)) {
    y += h * (-0.07 * (y - 20));
    x += h;
  }

  return y;
}
```
