---
id: 598de241872ef8353c58a7a2
title: حساب معاملات ذات الحدين
challengeType: 1
forumTopicId: 302259
dashedName: evaluate-binomial-coefficients
---

# --description--

اكتب دالة لحساب معامل ذات الحدين للقيم المعطاة n و k.

الصيغة التالية موصى بها:

$\\binom{n}{k} = \\frac{n!}{(n-k)!k!} = \\frac{n(n-1)(n-2)\\ldots(n-k+1)}{k(k-1)(k-2)\\ldots 1}$

# --hints--

يجب أن تكون `binom` دالة.

```js
assert(typeof binom === 'function');
```

يجب أن تُرجع `binom(5,3)` القيمة 10.

```js
assert.equal(binom(5, 3), 10);
```

يجب أن تُرجع `binom(7,2)` القيمة 21.

```js
assert.equal(binom(7, 2), 21);
```

يجب أن تُرجع `binom(10,4)` القيمة 210.

```js
assert.equal(binom(10, 4), 210);
```

يجب أن تُرجع `binom(6,1)` القيمة 6.

```js
assert.equal(binom(6, 1), 6);
```

يجب أن تُرجع `binom(12,8)` القيمة 495.

```js
assert.equal(binom(12, 8), 495);
```

# --seed--

## --seed-contents--

```js
function binom(n, k) {

}
```

# --solutions--

```js
function binom(n, k) {
  let coeff = 1;
  for (let i = n - k + 1; i <= n; i++) coeff *= i;
  for (let i = 1; i <= k; i++) coeff /= i;
  return coeff;
}
```
