---
id: 598eea87e5cf4b116c3ff81a
title: عوامل عدد ميرسين
challengeType: 1
forumTopicId: 302264
dashedName: factors-of-a-mersenne-number
---

# --description--

عدد ميرسين هو عدد على شكل <code>2<sup>P</sup>-1</code>.

إذا كان `P` عددًا أوليًا، فقد يكون عدد ميرسين عددًا أوليًا ميرسينيًا. (إذا لم يكن `P` أوليًا، فإن عدد ميرسين أيضًا ليس أوليًا.)

في البحث عن أعداد ميرسين الأولية، من المفيد استبعاد الأسس بإيجاد عامل صغير قبل بدء اختبار <a href="https://rosettacode.org/wiki/Lucas-Lehmer test" target="_blank" rel="noopener noreferrer nofollow">لوكاس-ليمر</a> الذي قد يكون طويلاً.

هناك خوارزميات فعالة جدًا لتحديد ما إذا كان عدد ما يقسم <code>2<sup>P</sup>-1</code> (أو بالمثل، إذا كان <code>2<sup>P</sup> mod (العدد) = 1</code>).

بعض اللغات لديها بالفعل تنفيذات مدمجة لهذه العملية الأسية مع باقي القسمة (تسمى modPow أو ما شابه).

فيما يلي كيفية تنفيذ هذه العملية modPow بنفسك:

على سبيل المثال، لنحسب <code>2<sup>23</sup> mod 47</code>.

حوّل الأس 23 إلى ثنائي، ستحصل على 10111. بدءًا من <code><tt>square</tt> = 1</code>، قم بتربيعها مرارًا.

أزل أعلى بت من الأس، وإذا كان 1 اضرب `square` في أساس الأس (2)، ثم احسب <code><tt>square</tt> modulo 47</code>.

استخدم نتيجة باقي القسمة من الخطوة السابقة كقيمة ابتدائية لـ `square` في الخطوة التالية:

<pre>إزالة   اختياري
square        أعلى بت  الضرب في 2  mod 47
------------  -------  -------------  ------
1*1 = 1       1  0111  1*2 = 2           2
2*2 = 4       0   111     لا             4
4*4 = 16      1    11  16*2 = 32        32
32*32 = 1024  1     1  1024*2 = 2048    27
27*27 = 729   1        729*2 = 1458      1
</pre>

بما أن <code>2<sup>23</sup> mod 47 = 1</code>، فإن 47 هو عامل لـ <code>2<sup>P</sup>-1</code>.

(لرؤية ذلك، اطرح 1 من كلا الطرفين: <code>2<sup>23</sup>-1 = 0 mod 47</code>.)

بما أننا أثبتنا أن 47 عامل، فإن <code>2<sup>23</sup>-1</code> ليس عددًا أوليًا.

خصائص أخرى لأعداد ميرسين تسمح لنا بتحسين العملية أكثر.

أي عامل `q` لـ <code>2<sup>P</sup>-1</code> يجب أن يكون على شكل `2kP+1`، حيث `k` عدد صحيح موجب أو صفر. علاوة على ذلك، يجب أن يكون `q` إما `1` أو `7 mod 8`.

أخيرًا، يجب أن يكون أي عامل محتمل `q` <a href="https://rosettacode.org/wiki/Primality_by_trial_division" target="_blank" rel="noopener noreferrer nofollow">عددًا أوليًا</a>.

كما في خوارزميات القسمة التجريبية الأخرى، تتوقف الخوارزمية عندما يكون `2kP+1 > sqrt(N)`. هذه الاختبارات تعمل أساسًا فقط على أعداد ميرسين حيث يكون `P` أوليًا. على سبيل المثال، <code>M<sub>4</sub>=15</code> لا يعطي عوامل باستخدام هذه التقنيات، لكنه يتحلل إلى 3 و5، ولا ينتميان إلى الشكل `2kP+1`.

# --instructions--

باستخدام الطريقة أعلاه، ابحث عن عامل لـ <code>2<sup>p</sup>-1</code>.

# --hints--

يجب أن تكون `check_mersenne` دالة.

```js
assert(typeof check_mersenne === 'function');
```

يجب أن تُرجع `check_mersenne(3)` سلسلة نصية.

```js
assert(typeof check_mersenne(3) == 'string');
```

يجب أن تُرجع `check_mersenne(3)` السلسلة النصية `M3 = 2^3-1 is prime`.

```js
assert.equal(check_mersenne(3), 'M3 = 2^3-1 is prime');
```

يجب أن تُرجع `check_mersenne(23)` السلسلة النصية `M23 = 2^23-1 is composite with factor 47`.

```js
assert.equal(check_mersenne(23), 'M23 = 2^23-1 is composite with factor 47');
```

يجب أن تُرجع `check_mersenne(929)` السلسلة النصية `M929 = 2^929-1 is composite with factor 13007`.

```js
assert.equal(
  check_mersenne(929),
  'M929 = 2^929-1 is composite with factor 13007'
);
```

# --seed--

## --seed-contents--

```js
function check_mersenne(p) {

}
```

# --solutions--

```js
function check_mersenne(p){
    function isPrime(value){
      for (let i=2; i < value; i++){
        if (value % i == 0){
          return false;
        }
        if (value % i != 0){
          return true;
         }
      }
    }

    function trial_factor(base, exp, mod){
      let square, bits;
      square = 1;
      bits = exp.toString(2).split('');
      for (let i=0,ln=bits.length; i<ln; i++){
        square = Math.pow(square, 2) * (bits[i] == 1 ? base : 1) % mod;
      }
      return (square == 1);
    }

    function mersenne_factor(p){
      let limit, k, q;
      limit = Math.sqrt(Math.pow(2,p) - 1);
      k = 1;
      while ((2*k*p - 1) < limit){
        q = 2*k*p + 1;
        if (isPrime(q) && (q % 8 == 1 || q % 8 == 7) && trial_factor(2,p,q)){
          return q; // q is a factor of 2**p-1
        }
        k++;
      }
      return null;
    }
  let f, result;
  result="M"+p+" = 2^"+p+"-1 is ";
  f = mersenne_factor(p);
  result+=f == null ? "prime" : "composite with factor "+f;
  return result;
}
```
