---
id: 5992e222d397f00d21122931
title: كلمة فيبوناتشي
challengeType: 1
forumTopicId: 302269
dashedName: fibonacci-word
---

# --description--

يمكن إنشاء متتالية كلمة فيبوناتشي بطريقة مشابهة لمتتالية فيبوناتشي، لكنها تركز على تكرار عملية الربط.

<pre>عرّف  F_Word<sub>1</sub>  كـ  <strong>1</strong>
عرّف  F_Word<sub>2</sub>  كـ  <strong>0</strong>
كوّن   F_Word<sub>3</sub>  بربط  F_Word<sub>2</sub>   مع  F_Word<sub>1</sub>   أي:  <strong>01</strong>
كوّن   F_Word<sub>n</sub>  بربط  F_Word<sub>n-1</sub>  مع  F_word<sub>n-2</sub>
</pre>

يتطلب هذا التحدٍّ حساب الإنتروبيا، <a href="https://www.freecodecamp.org/learn/coding-interview-prep/rosetta-code/entropy" target="_blank" rel="noopener noreferrer nofollow">كما هو موضح في تحدي Rosetta Code هذا</a>

# --instructions--

اكتب دالة تُرجع أول `n` كلمات فيبوناتشي. يُعطى عدد `n` كمعلمة للدالة. يجب أن تُرجع الدالة مصفوفة من الكائنات. يجب أن تكون الكائنات على الشكل: `{ N: 1, Length: 1, Entropy: 0, Word: '1' }`. تُحسب قيمة `Entropy` للسلسلة النصية `Word` وتُقرب إلى 8 أرقام عشرية. لاحظ أن فهارس هذه المتتالية تبدأ من `1`.

# --hints--

يجب أن تكون `fibWord` دالة.

```js
assert(typeof fibWord === 'function');
```

يجب أن تُرجع `fibWord(5)` مصفوفة.

```js
assert(Array.isArray(fibWord(5)));
```

يجب أن تُرجع `fibWord(5)` القيمة `[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }]`.

```js
assert.deepEqual(fibWord(5), globalThis.words5);
```

يجب أن تُرجع `fibWord(7)` القيمة `[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }, { N:6, Length:8, Entropy:0.954434, Word:'01001010' }, { N:7, Length:13, Entropy:0.9612366, Word:'0100101001001' }]`.

```js
assert.deepEqual(fibWord(7), globalThis.words7);
```

# --before-all--

```js
globalThis.words5 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' }
];

globalThis.words7 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' },
  { N: 6, Length: 8, Entropy: 0.954434, Word: '01001010' },
  { N: 7, Length: 13, Entropy: 0.9612366, Word: '0100101001001' }
];
```

# --seed--

## --seed-contents--

```js
function fibWord(n) {

}
```

# --solutions--

```js
// Round to digits
function roundFloat(num, digits) {
  return Math.round(num * 10.0**digits) / (10.0**digits);
}

// Entropy calculation for string with only 0 and 1
function entropy(word) {
  function digitEntropy(count) {
    return count < 1 ? 0
      : - count / word.length * Math.log2(count / word.length);
  }
  const numZeros = word.split('').filter(e => e === '0').length;
  const numOnes  = word.length - numZeros;
  return roundFloat(digitEntropy(numZeros) + digitEntropy(numOnes), 8);
}

// Compute array of Fibonacci words
function fibWord(n) {
  return [...Array(n).keys()]
    .reduce((words, i) => {
      const word = i === 0 ? "1"
                 : i === 1 ? "0"
                 : words[i - 1].Word + words[i - 2].Word;
      words.push(
        { N: i + 1, Length: word.length, Entropy: entropy(word), Word: word }
      );
      return words;
    }, []);
}
```
