---
id: 599c333915e0ea32d04d4bec
title: العمليات عنصر بعنصر
challengeType: 1
forumTopicId: 302252
dashedName: element-wise-operations
---

# --description--

نفذ عمليات أساسية عنصر بعنصر بين مصفوفة ومصفوفة أخرى، وبين عدد ثابت ومصفوفة.

**نفذ:**

<ul>
<li>الجمع</li>
  <li>الطرح</li>
  <li>الضرب</li>
  <li>القسمة</li>
  <li>الأس</li>
</ul>

المعلمة الأولى ستكون العملية التي ستُنفذ، مثل "m_add" للجمع بين المصفوفات و"s_add" للجمع مع عدد ثابت. المعلمتان الثانية والثالثة ستكونان المصفوفات التي ستُجرى عليها العمليات.

# --hints--

يجب أن تكون `operation` دالة.

```js
assert(typeof operation === 'function');
```

يجب أن تُرجع `operation("m_add",[[1,2],[3,4]],[[1,2],[3,4]])` القيمة `[[2,4],[6,8]]`.

```js
assert.deepEqual(
  operation(
    'm_add',
    [
      [1, 2],
      [3, 4]
    ],
    [
      [1, 2],
      [3, 4]
    ]
  ),
  [
    [2, 4],
    [6, 8]
  ]
);
```

يجب أن تُرجع `operation("s_add",[[1,2],[3,4]],2)` القيمة `[[3,4],[5,6]]`.

```js
assert.deepEqual(
  operation(
    's_add',
    [
      [1, 2],
      [3, 4]
    ],
    2
  ),
  [
    [3, 4],
    [5, 6]
  ]
);
```

يجب أن تُرجع `operation("m_sub",[[1,2],[3,4]],[[1,2],[3,4]])` القيمة `[[0,0],[0,0]]`.

```js
assert.deepEqual(
  operation(
    'm_sub',
    [
      [1, 2],
      [3, 4]
    ],
    [
      [1, 2],
      [3, 4]
    ]
  ),
  [
    [0, 0],
    [0, 0]
  ]
);
```

يجب أن تُرجع `operation("m_mult",[[1,2],[3,4]],[[1,2],[3,4]])` القيمة `[[1,4],[9,16]]`.

```js
assert.deepEqual(
  operation(
    'm_mult',
    [
      [1, 2],
      [3, 4]
    ],
    [
      [1, 2],
      [3, 4]
    ]
  ),
  [
    [1, 4],
    [9, 16]
  ]
);
```

يجب أن تُرجع `operation("m_div",[[1,2],[3,4]],[[1,2],[3,4]])` القيمة `[[1,1],[1,1]]`.

```js
assert.deepEqual(
  operation(
    'm_div',
    [
      [1, 2],
      [3, 4]
    ],
    [
      [1, 2],
      [3, 4]
    ]
  ),
  [
    [1, 1],
    [1, 1]
  ]
);
```

يجب أن تُرجع `operation("m_exp",[[1,2],[3,4]],[[1,2],[3,4]])` القيمة `[[1,4],[27,256]]`.

```js
assert.deepEqual(
  operation(
    'm_exp',
    [
      [1, 2],
      [3, 4]
    ],
    [
      [1, 2],
      [3, 4]
    ]
  ),
  [
    [1, 4],
    [27, 256]
  ]
);
```

يجب أن تُرجع `operation("m_add",[[1,2,3,4],[5,6,7,8]],[[9,10,11,12],[13,14,15,16]])` القيمة `[[10,12,14,16],[18,20,22,24]]`.

```js
assert.deepEqual(
  operation(
    'm_add',
    [
      [1, 2, 3, 4],
      [5, 6, 7, 8]
    ],
    [
      [9, 10, 11, 12],
      [13, 14, 15, 16]
    ]
  ),
  [
    [10, 12, 14, 16],
    [18, 20, 22, 24]
  ]
);
```

# --seed--

## --seed-contents--

```js
function operation(op, arr1, arr2) {

}
```

# --solutions--

```js
function operation(op, arr1, arr2) {
  const ops = {
    add: ((a, b) => a + b),
    sub: ((a, b) => a - b),
    mult: ((a, b) => a * b),
    div: ((a, b) => a / b),
    exp: ((a, b) => Math.pow(a, b))
  };
  const ifm = op.startsWith('m');
  const doOp = ops[op.substring(2)];
  for (let i = 0; i < arr1.length; i++) {
    for (let j = 0; j < arr1[0].length; j++) {
      arr1[i][j] = doOp(arr1[i][j], (ifm) ? (arr2[i][j]) : (arr2));
    }
  }
  return arr1;
}
```
