---
id: 599d1566a02b571412643b84
title: الضرب الإثيوبي
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

الضرب الإثيوبي هو طريقة لضرب الأعداد الصحيحة باستخدام الجمع فقط، والضعف، والنصف.

**الطريقة:**

<ol>
<li>اختر رقمين تريد ضربهما واكتبهما في أعلى عمودين</li>
  <li>في العمود الأيسر، قم بتقسيم الرقم الأخير على اثنين مرارًا مع تجاهل الباقي، واكتب النتيجة أسفل الرقم الأخير في نفس العمود، حتى تصل إلى القيمة <code>1</code></li>
  <li>في العمود الأيمن، قم بمضاعفة الرقم الأخير مرارًا واكتب النتيجة أسفله. توقف عندما تضيف نتيجة في نفس الصف الذي يظهر فيه <code>1</code> في العمود الأيسر</li>
  <li>افحص الجدول الناتج وتجاهل أي صف يكون فيه الرقم في العمود الأيسر زوجيًا</li>
  <li>اجمع القيم المتبقية في العمود الأيمن لتحصل على نتيجة ضرب الرقمين الأصليين</li>
</ol>

**على سبيل المثال:** `17 × 34`

<pre>17   34
</pre>

تقسيم العمود الأول على اثنين:

<pre>17   34
8
4
2
1
</pre>

مضاعفة العمود الثاني:

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

شطب الصفوف التي يكون فيها أول خلية زوجية:

<pre>17   34
8    <strike>68</strike>
4   <strike>136</strike>
2   <strike>272</strike>
1   544
</pre>

اجمع الأرقام المتبقية في العمود الأيمن:

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

إذاً `17` مضروبة في `34`، بالطريقة الإثيوبية تساوي `578`.

# --instructions--

المهمة هي تعريف ثلاث دوال/طرق/إجراءات مسماة:

<ol>
<li>واحدة لتقسيم عدد صحيح على اثنين،</li>
  <li>واحدة لمضاعفة عدد صحيح، و</li>
  <li>واحدة لتحديد ما إذا كان العدد الصحيح زوجيًا</li>
</ol>

استخدم هذه الدوال لإنشاء دالة تقوم بالضرب الإثيوبي.

<!-- markdownlint-disable MD046-->

# --hints--

يجب أن تكون `eth_mult` دالة.

```js
assert(typeof eth_mult === 'function');
```

يجب أن تُرجع `eth_mult(17,34)` القيمة `578`.

```js
assert.equal(eth_mult(17, 34), 578);
```

يجب أن تُرجع `eth_mult(23,46)` القيمة `1058`.

```js
assert.equal(eth_mult(23, 46), 1058);
```

يجب أن تُرجع `eth_mult(12,27)` القيمة `324`.

```js
assert.equal(eth_mult(12, 27), 324);
```

يجب أن تُرجع `eth_mult(56,98)` القيمة `5488`.

```js
assert.equal(eth_mult(56, 98), 5488);
```

يجب أن تُرجع `eth_mult(63,74)` القيمة `4662`.

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
