---
id: 599d0ba974141b0f508b37d5
title: أعداد إميرب الأولية
challengeType: 1
forumTopicId: 302253
dashedName: emirp-primes
---

# --description--

الإميرب هو عدد أولي (**prime** مكتوب بالعكس) وعندما يُعكس (في تمثيله العشري) يكون عددًا أوليًا مختلفًا.

# --instructions--

اكتب دالة تقوم بما يلي:

<ul>
<li>تعرض أول <code>n</code> عدد إميرب.</li>
  <li>تعرض أعداد الإميرب في مدى معين.</li>
  <li>تعرض عدد أعداد الإميرب في مدى معين.</li>
  <li>تعرض العدد الإميرب <code>n<sup>th</sup></code>.</li>
</ul>

يجب أن تستقبل الدالة معلمتين. الأولى تستقبل `n` أو المدى كمصفوفة. الثانية تستقبل قيمة منطقية تحدد ما إذا كانت الدالة تُرجع أعداد الإميرب كمصفوفة أو رقمًا واحدًا (عدد الأعداد الأولية في المدى أو العدد الإميرب <code>n<sup>th</sup></code>). وفقًا للمعلمات يجب أن تُرجع الدالة مصفوفة أو رقمًا.

# --hints--

يجب أن تكون `emirps` دالة.

```js
assert(typeof emirps === 'function');
```

`emirps(20,true)` يجب أن تُرجع `[13,17,31,37,71,73,79,97,107,113,149,157,167,179,199,311,337,347,359,389]`

```js
assert.deepEqual(emirps(20, true), [
  13,
  17,
  31,
  37,
  71,
  73,
  79,
  97,
  107,
  113,
  149,
  157,
  167,
  179,
  199,
  311,
  337,
  347,
  359,
  389
]);
```

`emirps(1000)` يجب أن تُرجع `70529`

```js
assert.deepEqual(emirps(1000), 70529);
```

`emirps([7700,8000],true)` يجب أن تُرجع `[7717,7757,7817,7841,7867,7879,7901,7927,7949,7951,7963]`

```js
assert.deepEqual(emirps([7700, 8000], true), [
  7717,
  7757,
  7817,
  7841,
  7867,
  7879,
  7901,
  7927,
  7949,
  7951,
  7963
]);
```

`emirps([7700,8000],false)` يجب أن تُرجع `11`

```js
assert.deepEqual(emirps([7700, 8000], false), 11);
```

# --seed--

## --seed-contents--

```js
function emirps(n) {

}
```

# --solutions--

```js
function emirps(num, showEmirps)
{
  const is_prime = function(n)
    {
    if (!(n % 2) || !(n % 3)) return false;
    let p = 1;
    while (p * p < n)
                    { if (n % (p += 4) == 0 || n % (p += 2) == 0)
                            { return false; } }
    return true;
  };
  const is_emirp = function(n) {
    const r = parseInt(n.toString().split('').reverse().join(''));
    return r != n && is_prime(n) && is_prime(r);
  };

  let i,
    arr = [];
  if (typeof num === 'number') {
    for (i = 0; arr.length < num; i++) if (is_emirp(i)) arr.push(i);
    // first x emirps
    if (showEmirps) return arr;
    // xth emirp
    return arr.pop();
  }

  if (Array.isArray(num)) {
    for (i = num[0]; i <= num[1]; i++) if (is_emirp(i)) arr.push(i);
    // emirps between x .. y
    if (showEmirps) return arr;
    // number of emirps between x .. y
    return arr.length;
  }
}
```
