---
id: 599d15309e88c813a40baf58
title: الإنتروبيا
challengeType: 1
forumTopicId: 302254
dashedName: entropy
---

# --description--

احسب إنتروبيا شانون H لسلسلة نصية مُدخَل معينة.

بالنظر إلى المتغير العشوائي المنفصل $X$ الذي هو سلسلة من $N$ "رموز" (إجمالي الأحرف) تتكون من $n$ أحرف مختلفة (n=2 للنظام الثنائي)، فإن إنتروبيا شانون لـ X بوحدة البت/رمز هي:

$H_2(X) = -\\sum\_{i=1}^n \\frac{count_i}{N} \\log_2 \\left(\\frac{count_i}{N}\\right)$

حيث $count_i$ هو عدد تكرار الحرف $n_i$.

# --hints--

يجب أن تكون `entropy` دالة.

```js
assert(typeof entropy === 'function');
```

يجب أن تُرجع `entropy("0")` القيمة `0`

```js
assert.equal(entropy('0'), 0);
```

يجب أن تُرجع `entropy("01")` القيمة `1`

```js
assert.equal(entropy('01'), 1);
```

يجب أن تُرجع `entropy("0123")` القيمة `2`

```js
assert.equal(entropy('0123'), 2);
```

يجب أن تُرجع `entropy("01234567")` القيمة `3`

```js
assert.equal(entropy('01234567'), 3);
```

يجب أن تُرجع `entropy("0123456789abcdef")` القيمة `4`

```js
assert.equal(entropy('0123456789abcdef'), 4);
```

يجب أن تُرجع `entropy("1223334444")` القيمة `1.8464393446710154`

```js
assert.equal(entropy('1223334444'), 1.8464393446710154);
```

# --seed--

## --seed-contents--

```js
function entropy(s) {

}
```

# --solutions--

```js
function entropy(s) {
    // Create a dictionary of character frequencies and iterate over it.
  function process(s, evaluator) {
    let h = Object.create(null),
      k;
    s.split('').forEach(c => {
      h[c] && h[c]++ || (h[c] = 1); });
    if (evaluator) for (k in h) evaluator(k, h[k]);
    return h;
  }
    // Measure the entropy of a string in bits per symbol.

  let sum = 0,
    len = s.length;
  process(s, (k, f) => {
    const p = f / len;
    sum -= p * Math.log(p) / Math.log(2);
  });
  return sum;
}
```
