---
id: 59c3ec9f15068017c96eb8a3
title: متتالية فاري
challengeType: 1
forumTopicId: 302266
dashedName: farey-sequence
---

# --description--

متتالية فاري <code>F<sub>n</sub></code> من الرتبة `n` هي متتالية الكسور المبسطة تمامًا بين `0` و `1` والتي، عندما تكون في أبسط صورة، يكون مقامها أقل من أو يساوي `n`، مرتبة بترتيب تصاعدي.

أحيانًا يُخطئ في تسمية *متتالية فاري* بـ *سلسلة فاري*.

كل متتالية فاري:

<ul>
<li>تبدأ بالقيمة 0، الممثلة بالكسر $ \frac{0}{1} $</li>
  <li>تنتهي بالقيمة 1، الممثلة بالكسر $ \frac{1}{1}$.</li>
</ul>

متتاليات فاري للرتب من `1` إلى `5` هي:

<ul>
  <li style='list-style: none;'>${\bf\it{F}}_1 = \frac{0}{1}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_2 = \frac{0}{1}, \frac{1}{2}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_3 = \frac{0}{1}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_4 = \frac{0}{1}, \frac{1}{4}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{3}{4}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_5 = \frac{0}{1}, \frac{1}{5}, \frac{1}{4}, \frac{1}{3}, \frac{2}{5}, \frac{1}{2}, \frac{3}{5}, \frac{2}{3}, \frac{3}{4}, \frac{4}{5}, \frac{1}{1}$</li>
</ul>

# --instructions--

اكتب دالة تُرجع متتالية فاري من الرتبة `n`. يجب أن تحتوي الدالة على معلمة واحدة هي `n`. يجب أن تُرجع المتتالية كمصفوفة.

# --hints--

`farey` يجب أن تكون دالة.

```js
assert(typeof farey === 'function');
```

`farey(3)` يجب أن تُرجع مصفوفة

```js
assert(Array.isArray(farey(3)));
```

`farey(3)` يجب أن تُرجع `['0/1','1/3','1/2','2/3','1/1']`

```js
assert.deepEqual(farey(3),['0/1', '1/3', '1/2', '2/3', '1/1']);
```

`farey(4)` يجب أن تُرجع `['0/1','1/4','1/3','1/2','2/3','3/4','1/1']`

```js
assert.deepEqual(farey(4), ['0/1', '1/4', '1/3', '1/2', '2/3', '3/4', '1/1']);
```

`farey(5)` يجب أن تُرجع `['0/1','1/5','1/4','1/3','2/5','1/2','3/5','2/3','3/4','4/5','1/1']`

```js
assert.deepEqual(farey(5), [
  '0/1',
  '1/5',
  '1/4',
  '1/3',
  '2/5',
  '1/2',
  '3/5',
  '2/3',
  '3/4',
  '4/5',
  '1/1'
]);
```

# --seed--

## --seed-contents--

```js
function farey(n) {

}
```

# --solutions--

```js
function farey(n) {
  const sequence = [{ string: "0/1", float: 0.0 }];
  for (let i = 1; i < n; i++) {
    for (let j = n; j >= i; j--) {
      if (i === 1 || j % i > 0) {
        sequence.push({ string: `${i}/${j}`, float: i / j });
      }
    }
  }
  return sequence
    .sort((a, b) => a.float - b.float)
    .map(e => e.string)
}
```
