---
id: 59d9c6bc214c613ba73ff012
title: أرقام SEDOL
challengeType: 1
forumTopicId: 302305
dashedName: sedols
---

# --description--

<abbr title="Stock Exchange Daily Official List">SEDOL</abbr> هي قائمة بأرقام تعريف الأوراق المالية الصادرة عن بورصة لندن.

لكل قائمة أرقام مكونة من 6 أرقام <abbr title="Stock Exchange Daily Official List">SEDOL</abbr>، احسب وأضف رقم التحقق. أي، بالنظر إلى سلسلة المُدخَل على اليسار، يجب أن تُرجع دالتك السلسلة المقابلة على اليمين:

<pre>
710889 => 7108899
B0YBKJ => B0YBKJ7
406566 => 4065663
B0YBLH => B0YBLH2
228276 => 2282765
B0YBKL => B0YBKL9
557910 => 5579107
B0YBKR => B0YBKR5
585284 => 5852842
B0YBKT => B0YBKT7
B00030 => B000300
</pre>

تحقق من أن كل مُدخَل مُشكّل بشكل صحيح، خصوصًا بالنسبة للأحرف الصالحة المسموح بها في سلسلة SEDOL. يجب أن تُرجع دالتك `null` عند وجود مُدخَل غير صالح.

# --hints--

يجب أن تكون `sedol` دالة.

```js
assert(typeof sedol === 'function');
```

يجب أن تُرجع `sedol('a')` القيمة null.

```js
assert(sedol('a') === null);
```

يجب أن تُرجع `sedol('710889')` القيمة '7108899'.

```js
assert(sedol('710889') === '7108899');
```

يجب أن تُرجع `sedol('BOATER')` القيمة null.

```js
assert(sedol('BOATER') === null);
```

يجب أن تُرجع `sedol('228276')` القيمة '2282765'.

```js
assert(sedol('228276') === '2282765');
```

# --seed--

## --seed-contents--

```js
function sedol(input) {
  const checkSum = 0
  
  return checkSum;
}
```

# --solutions--

```js
function sedol(input) {
  const checkDigit = sedolCheckDigit(input);
  if (checkDigit !== null) {
    return input + checkDigit;
  }
  return null;
}

const weight = [1, 3, 1, 7, 3, 9, 1];
function sedolCheckDigit(char6) {
  if (char6.search(/^[0-9BCDFGHJKLMNPQRSTVWXYZ]{6}$/) === -1) {
    return null;
  }

  let sum = 0;
  for (let i = 0; i < char6.length; i++) {
    sum += weight[i] * parseInt(char6.charAt(i), 36);
  }
  const checkSum = (10 - (sum % 10)) % 10;
  return checkSum.toString();
}
```
