---
id: 59da22823d04c95919d46269
title: "بحارة وجوز الهند ومشكلة القرد"
challengeType: 1
forumTopicId: 302304
dashedName: sailors-coconuts-and-a-monkey-problem
---

# --description--

تحطمت سفينة خمسة بحارة على جزيرة وجمعوا كومة كبيرة من جوز الهند خلال النهار.

في تلك الليلة، استيقظ البحار الأول وقرر أن يأخذ حصته مبكرًا، فحاول تقسيم كومة جوز الهند إلى خمس أكوام متساوية لكنه وجد أن هناك جوز هند واحد متبقٍ، فألقاه إلى القرد ثم أخفى "حصته" من الأكوام الخمسة المتساوية الحجم ودفع الأكوام الأربعة الأخرى معًا لتشكيل كومة واحدة مرئية من جوز الهند ثم عاد إلى النوم.

لتقصير القصة، استيقظ كل بحار بدوره مرة واحدة خلال الليل وقام بنفس الإجراءات من تقسيم كومة جوز الهند إلى خمسة، ووجد أن هناك جوز هند واحد متبقٍ فأعطى هذا الباقي الوحيد للقرد، ثم دفع الأكوام الأربعة الأخرى معًا لتشكيل كومة واحدة.

في الصباح (بعد الفعل السري والمنفصل لكل من البحارة الخمسة خلال الليل)، تُقسم جوز الهند المتبقية إلى خمس أكوام متساوية لكل بحار، حيث يُكتشف أن كومة جوز الهند تُقسم بالتساوي بين البحارة دون وجود باقي. (لا شيء للقرد في الصباح.)

# --instructions--

أنشئ دالة تُرجع الحد الأدنى الممكن لحجم كومة جوز الهند الأولية التي جُمعت خلال النهار لـ `N` بحارًا. **ملاحظة:** بالطبع تُروى القصة في عالم يمكن فيه جمع أي عدد من جوز الهند خلال النهار وتقسيم الكومة عدة مرات، وما إلى ذلك، في وقت يتناسب مع سياق القصة، بحيث لا يؤثر ذلك على الرياضيات.

# --hints--

يجب أن تكون `splitCoconuts` دالة.

```js
assert(typeof splitCoconuts === 'function');
```

يجب أن تُرجع `splitCoconuts(5)` القيمة 3121.

```js
assert(splitCoconuts(5) === 3121);
```

يجب أن تُرجع `splitCoconuts(6)` القيمة 233275.

```js
assert(splitCoconuts(6) === 233275);
```

يجب أن تُرجع `splitCoconuts(7)` القيمة 823537.

```js
assert(splitCoconuts(7) === 823537);
```

# --seed--

## --seed-contents--

```js
function splitCoconuts(intSailors) {

  return true;
}
```

# --solutions--

```js
function splitCoconuts(intSailors) {
  let intNuts = intSailors;
  let result = splitCoconutsHelper(intNuts, intSailors);
  while (!result) {
    intNuts += 1;
    result = splitCoconutsHelper(intNuts, intSailors);
  }

  return intNuts;
}

function splitCoconutsHelper(intNuts, intSailors, intDepth) {
  const nDepth = intDepth !== undefined ? intDepth : intSailors;
  const portion = Math.floor(intNuts / intSailors);
  const remain = intNuts % intSailors;

  if (portion <= 0 || remain !== (nDepth ? 1 : 0)) {
    return null;
  }

  if (nDepth) {
    return splitCoconutsHelper(
      intNuts - portion - remain, intSailors, nDepth - 1
    );
  }

  return intNuts;
}
```
