---
id: 59e0a8df964e4540d5abe599
title: تنفيذ Brain****
challengeType: 1
forumTopicId: 302261
dashedName: execute-brain
---

# --description--

اكتب دالة لتنفيذ مفسر Brain****. ستأخذ الدالة سلسلة نصية كمعلمة ويجب أن تُرجع سلسلة نصية كنتيجة. التفاصيل أكثر موضحة أدناه:

RCBF هي مجموعة من <a href="https://rosettacode.org/wiki/Brainf***" target="_blank" rel="noopener noreferrer nofollow">مترجمات ومفسرات Brainf***</a> مكتوبة لموقع Rosetta Code بعدة لغات.

فيما يلي روابط لكل إصدارات RCBF.

يكفي أن تنفذ التعليمات التالية بشكل صحيح:

| الأمر           | الوصف                                                                        |
| ----------------- | ---------------------------------------------------------------------------------- |
| <code>></code>    | حرك المؤشر إلى اليمين                                                      |
| <code>&lt;</code> | حرك المؤشر إلى اليسار                                                       |
| <code>+</code>    | زد قيمة خلية الذاكرة تحت المؤشر                                        |
| <code>-</code>    | انقص قيمة خلية الذاكرة تحت المؤشر                                        |
| <code>.</code>    | أخرج الحرف الذي تمثله الخلية تحت المؤشر                          |
| <code>,</code>    | أدخل حرفًا وخزنه في الخلية تحت المؤشر                          |
| <code>\[</code>   | اقفز بعد <code>]</code> المطابقة إذا كانت الخلية تحت المؤشر تساوي 0           |
| <code>]</code>    | اقفز إلى <code>\[</code> المطابقة إذا كانت الخلية تحت المؤشر غير صفرية |

أي حجم للخلية مسموح به، ودعم نهاية الملف (EOF) اختياري، وكذلك وجود ذاكرة محدودة أو غير محدودة.

# --hints--

`brain(globalThis.hello)` يجب أن تُرجع سلسلة نصية

```js
assert(typeof brain(globalThis.hello) === 'string');
```

`brain("++++++[>++++++++++<-]>+++++.")` يجب أن تُرجع "A"

```js
assert.equal(brain('++++++[>++++++++++<-]>+++++.'), 'A');
```

`brain(globalThis.hello)` يجب أن تُرجع `Hello World!\n`

```js
assert.equal(brain(globalThis.hello), 'Hello World!\n');
```

`brain(globalThis.fib)` يجب أن تُرجع `1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89`

```js
assert.equal(brain(globalThis.fib), '1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89');
```

# --before-all--

```js
globalThis.fib=`+

++

+++

++++

+>+>>

>>++++

+++++++

++++++++

+++++++++

++++++++++

++++++>++++

++++++++++++

+++++++++++++

+++<<<<<<[>[>>

>>>>+>+<<<<<<<-

]>>>>>>>[<<<<<<<

+>>>>>>>-]<[>++++

++++++[-<-[>>+>+<<

<-]>>>[<<<+>>>-]+<[

>[-]<[-]]>[<<[>>>+<<

<-]>>[-]]<<]>>>[>>+>+

<<<-]>>>[<<<+>>>-]+<[>

[-]<[-]]>[<<+>>[-]]<<<<

<<<]>>>>>[++++++++++++++

+++++++++++++++++++++++++

+++++++++.[-]]++++++++++<[

->-<]>+++++++++++++++++++++

+++++++++++++++++++++++++++.

[-]<<<<<<<<<<<<[>>>+>+<<<<-]>

>>>[<<<<+>>>>-]<-[>>.>.<<<[-]]

<<[>>+>+<<<-]>>>[<<<+>>>-]<<[<+

>-]>[<+>-]<<<-]`;
globalThis.hello='++++++++[>++++[>++>+++>+++>+<<<<-]>+>+>->>+[<]<-]>>.>---.+++++++..+++.>>.<-.<.+++.------.--------.>>+.>++.'
```

# --seed--

## --seed-contents--

```js
function brain(prog) {

}
```

# --solutions--

```js
function brain(prog){
  var output="";
    var code; // formatted code
  var ip = 0; // current instruction within code
  var nest = 0; // current bracket nesting (for Out button)
  var ahead = []; // locations of matching brackets

  var data = [0]; // data array (mod by +, -)
  var dp = 0; // index into data (mod by <, >)

  var inp = 0; // current input character (fetch with ,)
  var quit = 0;
    var commands = {
    '>':function() { if (++dp >= data.length) data[dp]=0 },
    '<':function() { if (--dp < 0) quit++ },
    '+':function() { ++data[dp] },
    '-':function() { --data[dp] },
    '[':function() { if (!data[dp]) ip = ahead[ip]; else ++nest },
    ']':function() { if ( data[dp]) ip = ahead[ip]; else --nest },
    ',':function() {
        var c = document.getElementById("input").value.charCodeAt(inp++);
        data[dp] = isNaN(c) ? 0 : c; // EOF: other options are -1 or no change
    },
    '.':function() {
            output+=String.fromCharCode(data[dp]);
            /*var s = document.getElementById("output").innerHTML)
             + String.fromCharCode(data[dp]);
            s = s.replace(/\n/g,"<br>").replace(/ /g,"&amp;nbsp;");
            document.getElementById("output").innerHTML = s;*/
        },
    };

    let ar=prog.split('');
    var st = [], back, error = -1;
    for (ip=0; ip<ar.length; ip++) {
        switch(ar[ip]) {
        case '[':
            st.push(ip);
            break;
        case ']':
            if (st.length == 0) error = ip;
            back = st.pop();
            ahead[ip] = back;
            ahead[back] = ip;
            break;
        }
    }

    for(ip=0;ip<ar.length;ip++){
    if(commands.hasOwnProperty(ar[ip]))
          commands[ar[ip]]();
    }

    return output;
}
```
