---
id: 59f40b17e79dbf1ab720ed7a
title: أرقام الأقسام
challengeType: 1
forumTopicId: 302249
dashedName: department-numbers
---

# --description--

هناك مدينة منظمة للغاية قررت تعيين رقم لكل من أقسامها:

<ul>
<li>قسم الشرطة</li>
  <li>قسم النظافة</li>
  <li>قسم الإطفاء</li>
</ul>

يمكن أن يكون لكل قسم رقم بين 1 و7 (شامل).

يجب أن تكون أرقام الأقسام الثلاثة فريدة (مختلفة عن بعضها البعض) ويجب أن يكون مجموعها 12.

رئيس الشرطة لا يحب الأرقام الفردية ويريد أن يكون لقسمه رقم زوجي.

# --instructions--

اكتب برنامجًا يعرض جميع التركيبات الصالحة كمصفوفة.

```js
[2, 3, 7] [2, 4, 6] [2, 6, 4]
[2, 7, 3] [4, 1, 7] [4, 2, 6]
[4, 3, 5] [4, 5, 3] [4, 6, 2]
[4, 7, 1] [6, 1, 5] [6, 2, 4]
[6, 4, 2] [6, 5, 1]
```

# --before-each--

```js
const nums = [1, 2, 3, 4, 5, 6, 7];
const total = 12;
const len = 14;
const result = [
  [2, 3, 7],
  [2, 4, 6],
  [2, 6, 4],
  [2, 7, 3],
  [4, 1, 7],
  [4, 2, 6],
  [4, 3, 5],
  [4, 5, 3],
  [4, 6, 2],
  [4, 7, 1],
  [6, 1, 5],
  [6, 2, 4],
  [6, 4, 2],
  [6, 5, 1]
];
```

# --hints--

يجب أن تكون `combinations` دالة.

```js
assert(typeof combinations === 'function');
```

يجب أن تُرجع `combinations([1, 2, 3], 6)` مصفوفة.

```js
assert(Array.isArray(combinations([1, 2, 3], 6)));
```

يجب أن تُرجع `combinations([1, 2, 3, 4, 5, 6, 7], 12)` مصفوفة بطول 14.

```js
assert(combinations(nums, total).length === len);
```

يجب أن تُرجع `combinations([1, 2, 3, 4, 5, 6, 7], 12)` جميع التركيبات الصالحة.

```js
assert.deepEqual(combinations(nums, total), result);
```

# --seed--

## --seed-contents--

```js
function combinations(possibleNumbers, total) {

  return true;
}
```

# --solutions--

```js
function combinations(possibleNumbers, total) {
  let firstNumber;
  let secondNumber;
  let thirdNumber;
  const allCombinations = [];

  for (let i = 0; i < possibleNumbers.length; i += 1) {
    firstNumber = possibleNumbers[i];

    if (firstNumber % 2 === 0) {
      for (let j = 0; j < possibleNumbers.length; j += 1) {
        secondNumber = possibleNumbers[j];

        if (j !== i && firstNumber + secondNumber <= total) {
          thirdNumber = total - firstNumber - secondNumber;

          if (thirdNumber !== firstNumber && thirdNumber !== secondNumber && possibleNumbers.includes(thirdNumber)) {
            allCombinations.push([firstNumber, secondNumber, thirdNumber]);
          }
        }
      }
    }
  }
  return allCombinations;
}
```
