---
id: 5a23c84252665b21eecc8041
title: مجموع متتالية
challengeType: 1
forumTopicId: 302333
dashedName: sum-of-a-series
---

# --description--

احسب الحد **n**<sup>th</sup> من <em>متتالية</em>، أي مجموع أول **n** حدود من <em>المتتالية</em> المقابلة. بشكل غير رسمي، يُسمى هذا القيمة، أو حدها عندما يقترب **n** من اللانهاية، *مجموع المتتالية*، ومن هنا جاء عنوان هذا التمرين. استخدم لهذه المهمة: $S_n = \displaystyle\sum_{k=1}^n \frac{1}{k^2}$.

# --instructions--

اكتب دالة تأخذ $a$ و $b$ كمعلمتين وتُرجع مجموع الأعضاء من المرتبة $a^{th}$ إلى $b^{th}$ في المتتالية.

# --hints--

يجب أن تكون `sum` دالة.

```js
assert(typeof sum == 'function');
```

يجب أن تُرجع `sum(1, 100)` رقمًا.

```js
assert(typeof sum(1, 100) == 'number');
```

يجب أن تُرجع `sum(1, 100)` القيمة `1.6349839001848923`.

```js
assert.equal(sum(1, 100), 1.6349839001848923);
```

يجب أن تُرجع `sum(33, 46)` القيمة `0.009262256361481223`.

```js
assert.equal(sum(33, 46), 0.009262256361481223);
```

يجب أن تُرجع `sum(21, 213)` القيمة `0.044086990748706555`.

```js
assert.equal(sum(21, 213), 0.044086990748706555);
```

يجب أن تُرجع `sum(11, 111)` القيمة `0.08619778593108679`.

```js
assert.equal(sum(11, 111), 0.08619778593108679);
```

يجب أن تُرجع `sum(1, 10)` القيمة `1.5497677311665408`.

```js
assert.equal(sum(1, 10), 1.5497677311665408);
```

# --seed--

## --seed-contents--

```js
function sum(a, b) {

}
```

# --solutions--

```js
function sum(a, b) {
  function fn(x) {
    return 1 / (x * x);
  }
  var s = 0;
  for (; a <= b; a++) s += fn(a);
  return s;
}
```
