---
id: 5a23c84252665b21eecc8017
title: Soundex
challengeType: 1
forumTopicId: 302320
dashedName: soundex
---

# --description--

**خوارزمية Soundex** تتعامل مع *نية* الكلمات. تُنشئ تمثيلًا للكلمات التي تتشابه في النطق.

تُستخدم في البحث عن <em>الأسماء</em> و<em>العناوين</em>. هذا يعني أن الشخص الذي أدخل <em>الاسم</em> يمكنه التركيز على طريقة نطقه بدلًا من تصحيح تهجئة <em>الأسماء</em>.

على سبيل المثال:

إذا سمعت الاسم `Quenci` للمرة الأولى وكتبت تهجئته خطأ، ستحصل على رمز **Soundex** `Q520`.

عندما تكتب الاسم `Quincy` بشكل صحيح في المرة التالية، ستظل تحصل على نفس الرمز `Q520`، مما يعني أنه يمكنك ربط نطق الأسماء المتعددة بنفس <em>الشخص</em> دون الحاجة لإضافة كل تهجئة.

إليك القواعد:

<ul>
<li>إذا فصل حرف متحرك (A, E, I, O, U) بين حرفين ساكنين لهما نفس رمز soundex، يُرمز للحرف الساكن على يمين المتحرك. يُرمز لكلمة Tymczak كـ T-522 (T، 5 للـ M، 2 للـ C، Z مهملة (انظر قاعدة "جنبًا إلى جنب" أعلاه)، 2 للـ K). بما أن المتحرك "A" يفصل بين Z وK، يُرمز للـ K.</li>
  <li>إذا فصل "H" أو "W" بين حرفين ساكنين لهما نفس رمز soundex، لا يُرمز للحرف الساكن على يمين المتحرك. مثال: يُرمز لكلمة Ashcraft كـ A-261 (A، 2 للـ S، C مهملة، 6 للـ R، 1 للـ F). لا يُرمز لها كـ A-226.</li>
</ul>

# --instructions--

اكتب دالة تأخذ سلسلة نصية كمعلمة وتُرجع السلسلة المشفرة.

# --hints--

`soundex` يجب أن تكون دالة.

```js
assert(typeof soundex == 'function');
```

`soundex("Soundex")` يجب أن تُرجع سلسلة نصية.

```js
assert(typeof soundex('Soundex') == 'string');
```

`soundex("Soundex")` يجب أن تُرجع `"S532"`.

```js
assert.equal(soundex('Soundex'), 'S532');
```

`soundex("Example")` يجب أن تُرجع `"E251"`.

```js
assert.equal(soundex('Example'), 'E251');
```

`soundex("Sownteks")` يجب أن تُرجع `"S532"`.

```js
assert.equal(soundex('Sownteks'), 'S532');
```

`soundex("Ekzampul")` يجب أن تُرجع `"E251"`.

```js
assert.equal(soundex('Ekzampul'), 'E251');
```

`soundex("Euler")` يجب أن تُرجع `"E460"`.

```js
assert.equal(soundex('Euler'), 'E460');
```

`soundex("Gauss")` يجب أن تُرجع `"G200"`.

```js
assert.equal(soundex('Gauss'), 'G200');
```

`soundex("Hilbert")` يجب أن تُرجع `"H416"`.

```js
assert.equal(soundex('Hilbert'), 'H416');
```

`soundex("Knuth")` يجب أن تُرجع `"K530"`.

```js
assert.equal(soundex('Knuth'), 'K530');
```

`soundex("Lloyd")` يجب أن تُرجع `"L300"`.

```js
assert.equal(soundex('Lloyd'), 'L300');
```

`soundex("Lukasiewicz")` يجب أن تُرجع `"L222"`.

```js
assert.equal(soundex('Lukasiewicz'), 'L222');
```

# --seed--

## --seed-contents--

```js
function soundex(s) {

}
```

# --solutions--

```js
function soundex(s) {
  var a = s.toLowerCase().split('');
  var f = a.shift(),
    r = '',
    codes = {
      a: '',
      e: '',
      i: '',
      o: '',
      u: '',
      b: 1,
      f: 1,
      p: 1,
      v: 1,
      c: 2,
      g: 2,
      j: 2,
      k: 2,
      q: 2,
      s: 2,
      x: 2,
      z: 2,
      d: 3,
      t: 3,
      l: 4,
      m: 5,
      n: 5,
      r: 6
    };
  r =
    f +
    a
      .map(function(v, i, a) {
        return codes[v];
      })
      .filter(function(v, i, a) {
        return i === 0 ? v !== codes[f] : v !== a[i - 1];
      })
      .join('');

  return (r + '000').slice(0, 4).toUpperCase();
}
```
