---
id: 5a23c84252665b21eecc7e80
title: شفرة جراي
challengeType: 1
forumTopicId: 302276
dashedName: gray-code
---

# --description--

شفرة جراي هي شكل من أشكال الترميز الثنائي حيث تختلف التحولات بين الأعداد المتتالية بمقدار بت واحد فقط.

هذا الترميز مفيد لتقليل مخاطر بيانات الأجهزة مع القيم التي تتغير بسرعة و/أو تتصل بأجهزة أبطأ كمُدخَلات.

كما أنه مفيد لتوليد مُدخَلات لخرائط كارنو بالترتيب من اليسار إلى اليمين أو من الأعلى إلى الأسفل.

# --instructions--

أنشئ دالة لترميز رقم إلى شفرة جراي وفك ترميز رقم منها. يجب أن تحتوي الدالة على معلمتين.

الأولى ستكون قيمة منطقية. يجب أن ترمز الدالة عند true وتفك الترميز عند false. المعلمة الثانية ستكون الرقم المراد ترميزه/فك ترميزه.

اعرض التمثيلات الثنائية العادية، وتمثيلات شفرة جراي، وقيم شفرة جراي المفكوكة لجميع الأعداد الثنائية ذات 5 بت (من 0 إلى 31 شاملة، لا حاجة للأصفار البادئة).

هناك العديد من شفرات جراي الممكنة. التالي يرمز ما يسمى "شفرة جراي المنعكسة الثنائية".

الترميز (MSB هو البت 0، b هو الثنائي، g هو شفرة جراي):

<pre>إذا كان b[i-1] = 1
  g[i] = ليس b[i]
وإلا
  g[i] = b[i]
</pre>

أو:

<pre>g = b xor (b تُزاح منطقيًا إلى اليمين مرة واحدة)
</pre>

فك الترميز (MSB هو البت 0، b هو الثنائي، g هو شفرة جراي):

<pre>b[0] = g[0]<br>
لبقية البتات:
b[i] = g[i] xor b[i-1]</pre>
</pre>

# --hints--

`gray` يجب أن تكون دالة.

```js
assert(typeof gray == 'function');
```

`gray(true,177)` يجب أن تُرجع رقمًا.

```js
assert(typeof gray(true, 177) == 'number');
```

`gray(true,177)` يجب أن تُرجع `233`.

```js
assert.equal(gray(true, 177), 233);
```

`gray(true,425)` يجب أن تُرجع `381`.

```js
assert.equal(gray(true, 425), 381);
```

`gray(true,870)` يجب أن تُرجع `725`.

```js
assert.equal(gray(true, 870), 725);
```

`gray(false,233)` يجب أن تُرجع `177`.

```js
assert.equal(gray(false, 233), 177);
```

`gray(false,381)` يجب أن تُرجع `425`.

```js
assert.equal(gray(false, 381), 425);
```

`gray(false,725)` يجب أن تُرجع `870`.

```js
assert.equal(gray(false, 725), 870);
```

# --seed--

## --seed-contents--

```js
function gray(enc, number) {

}
```

# --solutions--

```js
function gray(enc, number){
  if(enc){
      return number ^ (number >> 1);
  }else{
      let n = number;

      while (number >>= 1) {
          n ^= number;
      }
      return n;
  }
}
```
