---
id: 5a23c84252665b21eecc8001
title: خوارزميات الترتيب/ترتيب الخرز
challengeType: 1
forumTopicId: 302310
dashedName: sorting-algorithmsbead-sort
---

# --description--

يبدأ *ترتيب الخرز* بإنشاء مصفوفة من الأصفار طولها يساوي قيمة أكبر عنصر في المصفوفة المُدخَلة. تُحوَّل المصفوفة بإضافة واحد إلى جميع العناصر بين الفهرس الصفري والفهرس المحدد بالعنصر الحالي. يُكرر هذا الإجراء حتى تملأ المصفوفة.

عند التكرار على المصفوفة، وجمع عدد العناصر الأكبر من صفر، ثم تقليل قيمة كل عنصر بواحد، تحصل على المصفوفة المرتبة.

**ملاحظة:** كل عنصر في المصفوفة المُدخَلة فريد.

رتب مصفوفة من الأعداد الصحيحة الموجبة باستخدام خوارزمية ترتيب الخرز.

# --hints--

يجب أن تكون `beadSort` دالة.

```js
assert(typeof beadSort == 'function');
```

يجب أن تُرجع `beadSort([25, 32, 12, 7, 20])` مصفوفة.

```js
assert(Array.isArray(beadSort([25, 32, 12, 7, 20])));
```

يجب أن تُرجع `beadSort([25, 32, 12, 7, 20])` `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(beadSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

يجب أن تُرجع `beadSort([38, 45, 35, 8, 13])` `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(beadSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

يجب أن تُرجع `beadSort([43, 36, 20, 34, 24])` `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(beadSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

يجب أن تُرجع `beadSort([12, 33, 26, 18, 1, 16, 38])` `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(beadSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

يجب أن تُرجع `beadSort([3, 39, 48, 16, 1, 4, 29])` `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(beadSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function beadSort(arr) {

}
```

# --solutions--

```js
function beadSort(arr) {
  let max = 0;
  for (let i = 0; i < arr.length; i++) if (arr[i] > max) max = arr[i];
  const grid = new Array(arr.length);
  for (let i = 0; i < grid.length; i++) {
    grid[i] = new Array(max);
  }
  const levelcount = new Array(max);
  levelcount.fill(0);
  for (let i = 0; i < max; i++) {
    levelcount[i] = 0;
    for (let j = 0; j < arr.length; j++) grid[j][i] = '_';
  }
  for (let i = 0; i < arr.length; i++) {
    let num = arr[i];
    for (let j = 0; num > 0; j++) {
      grid[levelcount[j]++][j] = '*';
      num--;
    }
  }
  const sorted = new Array(arr.length);
  sorted.fill(0);
  for (let i = 0; i < arr.length; i++) {
    let putt = 0;
    for (
      let j = 0;
      j < max &&
      (function(c) {
        return c.charCodeAt == null ? c : c.charCodeAt(0);
      })(grid[arr.length - 1 - i][j]) == '*'.charCodeAt(0);
      j++
    )
      putt++;
    sorted[i] = putt;
  }
  return sorted;
}
```
