---
id: 5a23c84252665b21eecc7e06
title: قص مستطيل
challengeType: 1
forumTopicId: 302242
dashedName: cut-a-rectangle
---

# --description--

المستطيل المعطى مكوّن من مربعات بأبعاد *m* × *n*. إذا لم يكن *m* و *n* كلاهما فرديين، فيمكن قطع مسار عبر المستطيل على طول حواف المربعات بحيث ينقسم المستطيل إلى قطعتين متصلتين لهما نفس الشكل (بعد تدوير إحدى القطعتين 180°). جميع هذه المسارات للمستطيلات 2 × 2 و 4 × 3 موضحة أدناه.

<img src="https://cdn.freecodecamp.org/curriculum/coding-interview-prep/rosetta-cut-a-rectangle.png" alt="Diagram showing the possible paths for 2 by 2 and 4 by 3 rectangles" />

# --instructions--

اكتب دالة تحسب عدد الطرق المختلفة لقص مستطيل بأبعاد *m* × *n*.

# --hints--

يجب أن تكون `cutRectangle` دالة.

```js
assert(typeof cutRectangle == 'function');
```

يجب أن تُرجع `cutRectangle(2, 2)` رقمًا.

```js
assert(typeof cutRectangle(2, 2) == 'number');
```

يجب أن تُرجع `cutRectangle(2, 2)` القيمة `2`.

```js
assert.equal(cutRectangle(2, 2), 2);
```

يجب أن تُرجع `cutRectangle(4, 3)` القيمة `9`.

```js
assert.equal(cutRectangle(4, 3), 9);
```

يجب أن تُرجع `cutRectangle(4, 4)` القيمة `22`.

```js
assert.equal(cutRectangle(4, 4), 22);
```

يجب أن تُرجع `cutRectangle(8, 3)` القيمة `53`.

```js
assert.equal(cutRectangle(8, 3), 53);
```

يجب أن تُرجع `cutRectangle(7, 4)` القيمة `151`.

```js
assert.equal(cutRectangle(7, 4), 151);
```

# --seed--

## --seed-contents--

```js
function cutRectangle(w, h) {

}
```

# --solutions--

```js
function cutRectangle(w, h) {
  if (w % 2 == 1 && h % 2 == 1) return;

  var dirs = [[0, -1], [-1, 0], [0, 1], [1, 0]];

  var grid = new Array(h);
  for (var i = 0; i < grid.length; i++) grid[i] = new Array(w);
  var stack = [];

  var half = Math.floor((w * h) / 2);
  var bits = Math.pow(2, half) - 1;
  var result = 0;
  for (; bits > 0; bits -= 2) {
    for (var i = 0; i < half; i++) {
      var r = Math.floor(i / w);
      var c = i % w;
      grid[r][c] = (bits & (1 << i)) != 0 ? 1 : 0;
      grid[h - r - 1][w - c - 1] = 1 - grid[r][c];
    }

    stack.push(0);
    grid[0][0] = 2;
    var count = 1;
    while (stack.length != 0) {
      var pos = stack.pop();
      var r = Math.floor(pos / w);
      var c = pos % w;

      for (var dir of dirs) {
        var nextR = r + dir[0];
        var nextC = c + dir[1];

        if (nextR >= 0 && nextR < h && nextC >= 0 && nextC < w) {
          if (grid[nextR][nextC] == 1) {
            stack.push(nextR * w + nextC);
            grid[nextR][nextC] = 2;
            count++;
          }
        }
      }
    }

    if (count == half) {
      result++;
    }
  }

  return result;
}
```
