---
id: 5a23c84252665b21eecc7e84
title: أكبر مجموع لتسلسل فرعي
challengeType: 1
forumTopicId: 302278
dashedName: greatest-subsequential-sum
---

# --description--

بالنظر إلى متتالية من الأعداد الصحيحة، ابحث عن تسلسل فرعي مستمر يُعظّم مجموع عناصره، أي أن عناصر أي تسلسل فرعي آخر لا تجمع إلى قيمة أكبر من هذه.

يُعتبر التسلسل الفرعي الفارغ ذا مجموع \( 0 \)؛ لذا إذا كانت كل العناصر سالبة، يجب أن تكون النتيجة التسلسل الفارغ.

# --hints--

يجب أن تكون `maximumSubsequence` دالة.

```js
assert(typeof maximumSubsequence == 'function');
```

يجب أن تُرجع `maximumSubsequence([ 1, 2, -1, 3, 10, -10 ])` مصفوفة.

```js
assert(Array.isArray(maximumSubsequence([1, 2, -1, 3, 10, -10])));
```

يجب أن تُرجع `maximumSubsequence([ 1, 2, -1, 3, 10, -10 ])` `[ 1, 2, -1, 3, 10 ]`.

```js
assert.deepEqual(maximumSubsequence([1, 2, -1, 3, 10, -10]), [1, 2, -1, 3, 10]);
```

يجب أن تُرجع `maximumSubsequence([ 0, 8, 10, -2, -4, -1, -5, -3 ])` `[ 0, 8, 10 ]`.

```js
assert.deepEqual(maximumSubsequence([0, 8, 10, -2, -4, -1, -5, -3]), [
  0,
  8,
  10
]);
```

يجب أن تُرجع `maximumSubsequence([ 9, 9, -10, 1 ])` `[ 9, 9 ]`.

```js
assert.deepEqual(maximumSubsequence([9, 9, -10, 1]), [9, 9]);
```

يجب أن تُرجع `maximumSubsequence([ 7, 1, -5, -3, -8, 1 ])` `[ 7, 1 ]`.

```js
assert.deepEqual(maximumSubsequence([7, 1, -5, -3, -8, 1]), [7, 1]);
```

يجب أن تُرجع `maximumSubsequence([ -3, 6, -1, 4, -4, -6 ])` `[ 6, -1, 4 ]`.

```js
assert.deepEqual(maximumSubsequence([-3, 6, -1, 4, -4, -6]), [6, -1, 4]);
```

يجب أن تُرجع `maximumSubsequence([ -1, -2, 3, 5, 6, -2, -1, 4, -4, 2, -1 ])` `[ 3, 5, 6, -2, -1, 4 ]`.

```js
assert.deepEqual(maximumSubsequence([-1, -2, 3, 5, 6, -2, -1, 4, -4, 2, -1]), [
  3,
  5,
  6,
  -2,
  -1,
  4
]);
```

# --seed--

## --seed-contents--

```js
function maximumSubsequence(population) {

}
```

# --solutions--

```js
function maximumSubsequence(population) {
  function sumValues(arr) {
      var result = 0;
      for (var i = 0, len = arr.length; i < len; i++) {
          result += arr[i];
      }
      return result;
  }
  var greatest;
  var maxValue = 0;

  for (var i = 0, len = population.length; i < len; i++) {
      for (var j = i; j <= len; j++) {
          var subsequence = population.slice(i, j);
          var value = sumValues(subsequence);
          if (value > maxValue) {
              maxValue = value;
              greatest = subsequence;
          };
      }
  }

  return greatest;
}
```
