---
id: 5a23c84252665b21eecc8046
title: الفرق المتماثل
challengeType: 1
forumTopicId: 16086
dashedName: symmetric-difference
---

# --description--

معطى مجموعتان *A* و *B*، احسب $(A \\setminus B) \\cup (B \\setminus A).$ أي، عد العناصر التي توجد في *A* أو *B* ولكن ليست في كلتيهما. تُسمى هذه المجموعة الفرق المتماثل بين *A* و *B*. بعبارة أخرى: $(A \\cup B) \\setminus (A \\cap B)$ (مجموعة العناصر التي توجد في واحدة على الأقل من *A* أو *B* ناقص مجموعة العناصر التي توجد في كلتيهما).

مثال:

للمجموعتين `A = [1, 2, 3]` و `B = [1, 3, 4]`، الفرق المتماثل بين *A* و *B* هو `[2, 4]`.

# --instructions--

اكتب دالة تأخذ مصفوفتين كمعلمات وتُرجع الفرق المتماثل. رتب المصفوفة الناتجة قبل إرجاعها.

# --hints--

يجب أن تكون `symmetricDifference` دالة.

```js
assert(typeof symmetricDifference == 'function');
```

يجب أن تُرجع `symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])` مصفوفة.

```js
assert(
  Array.isArray(
    symmetricDifference(
      ['John', 'Bob', 'Mary', 'Serena'],
      ['Jim', 'Mary', 'John', 'Bob']
    )
  )
);
```

يجب أن تُرجع `symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])` `["Jim", "Serena"]`.

```js
assert.deepEqual(
  symmetricDifference(
    ['John', 'Bob', 'Mary', 'Serena'],
    ['Jim', 'Mary', 'John', 'Bob']
  ),
  ['Jim', 'Serena']
);
```

يجب أن تُرجع `symmetricDifference([1, 2, 3], [3, 4])` `[1, 2, 4]`.

```js
assert.deepEqual(symmetricDifference([1, 2, 3], [3, 4]), [1, 2, 4]);
```

يجب أن تُرجع `symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7])` `[1, 2, 5, 7, 8]`.

```js
assert.deepEqual(symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7]), [
  1,
  2,
  5,
  7,
  8
]);
```

يجب أن تُرجع `symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9])` `[2, 4, 9]`.

```js
assert.deepEqual(
  symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9]),
  [2, 4, 9]
);
```

يجب أن تُرجع `symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9])` `[1, 3, 4, 8]`.

```js
assert.deepEqual(symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9]), [
  1,
  3,
  4,
  8
]);
```

# --seed--

## --seed-contents--

```js
function symmetricDifference(A, B) {

}
```

# --solutions--

```js
function symmetricDifference(A, B) {
  function relative_complement(A, B) {
    return A.filter(function(elem) {
      return B.indexOf(elem) == -1;
    });
  }

  function unique(ary) {
    var u = ary.concat().sort();
    for (var i = 1; i < u.length; ) {
      if (u[i - 1] === u[i]) u.splice(i, 1);
      else i++;
    }
    return u;
  }

  return unique(
    relative_complement(A, B).concat(relative_complement(B, A))
  ).sort();
}
```
