---
id: 5a23c84252665b21eecc7ec2
title: مسافة Jaro
challengeType: 1
forumTopicId: 302292
dashedName: jaro-distance
---

# --description--

مسافة Jaro هي مقياس للتشابه بين سلسلتين نصيتين. كلما زادت مسافة Jaro بين سلسلتين، زاد تشابههما. تُطبع النتيجة بحيث تعادل القيمة `0` عدم وجود تشابه والقيمة `1` تعني تطابقًا تامًا.

**التعريف**

مسافة Jaro \( d_j \) لسلسلتين معطاة \(s_1\) و \(s_2\) هي

\\begin{align}d_j = \\begin{cases}0& & \\text{إذا كان }m=0 \\\\\\\\{\\frac {1}{3}}\\left({\\frac {m}{|s\_{1}|}}+{\\frac {m}{|s\_{2}|}}+{\\frac {m-t}{m}}\\right)& & \\text{وإلا}\\end{cases}\\end{align}

حيث:

<ul>
<li>\(m\) هو عدد <i>الأحرف المتطابقة</i>؛</li>
  <li> \(t\) هو نصف عدد <i>التبديلات</i>.</li>
</ul>

يُعتبر حرفان من \(s_1\) و \(s_2\) على التوالي *متطابقين* فقط إذا كانا متماثلين ولم يبعدا أكثر من \(\\left\\lfloor\\frac{\\max(|s_1|,|s_2|)}{2}\\right\\rfloor-1\\).

يُقارن كل حرف من \(s_1\) مع جميع أحرفه المتطابقة في \(s_2\). عدد الأحرف المتطابقة التي تختلف في ترتيب المتتالية مقسومًا على 2 يحدد عدد *التبديلات*.

**مثال**

بالنظر إلى السلسلتين \(s_1\) *DWAYNE* و \(s_2\) *DUANE* نجد:

<ul>
  <li>\(m = 4\)</li>
  <li>\(|s_1| = 6\)</li>
  <li>\(|s_2| = 5\)</li>
  <li>\(t = 0\)</li>
</ul>

نجد نتيجة Jaro: \(d_j = \frac{1}{3}\left(\frac{4}{6} + \frac{4}{5} + \frac{4-0}{4}\right) = 0.822\).

# --instructions--

اكتب دالة تأخذ سلسلتين نصيتين كمعلمات وتُرجع مسافة Jaro المرتبطة بهما.

# --hints--

`jaro` يجب أن تكون دالة.

```js
assert(typeof jaro == 'function');
```

`jaro("MARTHA", "MARHTA")` يجب أن تُرجع رقمًا.

```js
assert(typeof jaro('MARTHA', 'MARHTA') == 'number');
```

`jaro("MARTHA", "MARHTA")` يجب أن تُرجع `0.9444444444444445`.

```js
assert.equal(jaro('MARTHA', 'MARHTA'), 0.9444444444444445);
```

`jaro("DIXON", "DICKSONX")` يجب أن تُرجع `0.7666666666666666`.

```js
assert.equal(jaro('DIXON', 'DICKSONX'), 0.7666666666666666);
```

`jaro("JELLYFISH", "SMELLYFISH")` يجب أن تُرجع `0.8962962962962964`.

```js
assert.equal(jaro('JELLYFISH', 'SMELLYFISH'), 0.8962962962962964);
```

`jaro("HELLOS", "CHELLO")` يجب أن تُرجع `0.888888888888889`.

```js
assert.equal(jaro('HELLOS', 'CHELLO'), 0.888888888888889);
```

`jaro("ABCD", "BCDA")` يجب أن تُرجع `0.8333333333333334`.

```js
assert.equal(jaro('ABCD', 'BCDA'), 0.8333333333333334);
```

# --seed--

## --seed-contents--

```js
function jaro(s, t) {

}
```

# --solutions--

```js
function jaro(s, t) {
  var s_len = s.length;
  var t_len = t.length;

  if (s_len == 0 && t_len == 0) return 1;

  var match_distance = Math.max(s_len, t_len) / 2 - 1;

  var s_matches = new Array(s_len);
  var t_matches = new Array(t_len);

  var matches = 0;
  var transpositions = 0;

  for (var i = 0; i < s_len; i++) {
    var start = Math.max(0, i - match_distance);
    var end = Math.min(i + match_distance + 1, t_len);

    for (var j = start; j < end; j++) {
      if (t_matches[j]) continue;
      if (s.charAt(i) != t.charAt(j)) continue;
      s_matches[i] = true;
      t_matches[j] = true;
      matches++;
      break;
    }
  }

  if (matches == 0) return 0;

  var k = 0;
  for (var i = 0; i < s_len; i++) {
    if (!s_matches[i]) continue;
    while (!t_matches[k]) k++;
    if (s.charAt(i) != t.charAt(k)) transpositions++;
    k++;
  }

  return ((matches / s_len) +
    (matches / t_len) +
    ((matches - transpositions / 2.0) / matches)) / 3.0;
}
```
