---
id: 5a23c84252665b21eecc8000
title: فرز قائمة فرعية متفرقة
challengeType: 1
forumTopicId: 302307
dashedName: sort-disjoint-sublist
---

# --description--

معطى قائمة من القيم ومجموعة من الفهارس الصحيحة داخل تلك القائمة، المطلوب هو فرز القيم عند الفهارس المعطاة، مع الحفاظ على القيم عند الفهارس خارج مجموعة الفهارس التي يجب فرزها.

اجعل دالتك تعمل مع قائمة القيم ومجموعة الفهارس التالية:

<code>values: [7, <b>6</b>, 5, 4, 3, 2, <b>1</b>, <b>0</b>]</code>

```js
indices(0-based): {6, 1, 7}
```

والنتيجة الصحيحة ستكون:

<code>[7, <b>0</b>, 5, 4, 3, 2, <b>1</b>, <b>6</b>]</code>.

# --hints--

يجب أن تكون `sortDisjoint` دالة.

```js
assert(typeof sortDisjoint == 'function');
```

يجب أن تُرجع `sortDisjoint([7, 6, 5, 4, 3, 2, 1, 0], [6, 1, 7])` مصفوفة.

```js
assert(Array.isArray(sortDisjoint([7, 6, 5, 4, 3, 2, 1, 0], [6, 1, 7])));
```

يجب أن تُرجع `sortDisjoint([7, 6, 5, 4, 3, 2, 1, 0], [6, 1, 7])` القيمة `[7, 0, 5, 4, 3, 2, 1, 6]`.

```js
assert.deepEqual(sortDisjoint([7, 6, 5, 4, 3, 2, 1, 0], [6, 1, 7]), [
  7,
  0,
  5,
  4,
  3,
  2,
  1,
  6
]);
```

يجب أن تُرجع `sortDisjoint([7, 6, 5, 4, 3, 2, 1, 0], [1, 2, 5, 6])` القيمة `[7, 1, 2, 4, 3, 5, 6, 0]`.

```js
assert.deepEqual(sortDisjoint([7, 6, 5, 4, 3, 2, 1, 0], [1, 2, 5, 6]), [
  7,
  1,
  2,
  4,
  3,
  5,
  6,
  0
]);
```

يجب أن تُرجع `sortDisjoint([8, 7, 6, 5, 4, 3, 2, 1], [6, 1, 7])` القيمة `[8, 1, 6, 5, 4, 3, 2, 7]`.

```js
assert.deepEqual(sortDisjoint([8, 7, 6, 5, 4, 3, 2, 1], [6, 1, 7]), [
  8,
  1,
  6,
  5,
  4,
  3,
  2,
  7
]);
```

يجب أن تُرجع `sortDisjoint([8, 7, 6, 5, 4, 3, 2, 1], [1, 3, 5, 6])` القيمة `[8, 2, 6, 3, 4, 5, 7, 1]`.

```js
assert.deepEqual(sortDisjoint([8, 7, 6, 5, 4, 3, 2, 1], [1, 3, 5, 6]), [
  8,
  2,
  6,
  3,
  4,
  5,
  7,
  1
]);
```

يجب أن تُرجع `sortDisjoint([6, 1, 7, 1, 3, 5, 6], [6, 1, 5, 4])` القيمة `[6, 1, 7, 1, 3, 5, 6]`.

```js
assert.deepEqual(sortDisjoint([6, 1, 7, 1, 3, 5, 6], [6, 1, 5, 4]), [
  6,
  1,
  7,
  1,
  3,
  5,
  6
]);
```

# --seed--

## --seed-contents--

```js
function sortDisjoint(values, indices) {

}
```

# --solutions--

```js
function sortDisjoint(values, indices) {
  let sublist = [];

  indices.sort(function(a, b) {
    return a - b;
  });

  for (let i = 0; i < indices.length; i++) {
    sublist.push(values[indices[i]]);
  }

  sublist.sort((a, b) => {
    return a - b;
  });

  for (let i = 0; i < indices.length; i++) {
    values[indices[i]] = sublist[i];
  }

  return values;
}
```
