---
id: 5a23c84252665b21eecc7e05
title: CUSIP
challengeType: 1
forumTopicId: 302241
dashedName: cusip
---

# --description--

الـ **CUSIP** هو رمز أبجدي رقمي مكون من تسعة أحرف يحدد ورقة مالية في أمريكا الشمالية بهدف تسهيل مسح وتسوية الصفقات. تم اعتماد CUSIP كمعيار وطني أمريكي تحت المعايير المعتمدة X9.6.

# --instructions--

اكتب دالة تأخذ سلسلة نصية كمعلمة وتتحقق مما إذا كانت السلسلة تمثل CUSIP صالحًا.

# --hints--

يجب أن تكون `isCusip` دالة.

```js
assert(typeof isCusip == 'function');
```

يجب أن تُرجع `isCusip("037833100")` قيمة منطقية.

```js
assert(typeof isCusip('037833100') == 'boolean');
```

يجب أن تُرجع `isCusip("037833100")` القيمة `true`.

```js
assert.equal(isCusip('037833100'), true);
```

يجب أن تُرجع `isCusip("17275R102")` القيمة `true`.

```js
assert.equal(isCusip('17275R102'), true);
```

يجب أن تُرجع `isCusip("38259P50a")` القيمة `false`.

```js
assert.equal(isCusip('38259P50a'), false);
```

يجب أن تُرجع `isCusip("38259P508")` القيمة `true`.

```js
assert.equal(isCusip('38259P508'), true);
```

يجب أن تُرجع `isCusip("38259P50#")` القيمة `false`.

```js
assert.equal(isCusip('38259P50#'), false);
```

يجب أن تُرجع `isCusip("68389X105")` القيمة `true`.

```js
assert.equal(isCusip('68389X105'), true);
```

يجب أن تُرجع `isCusip("68389X106")` القيمة `false`.

```js
assert.equal(isCusip('68389X106'), false);
```

يجب أن تُرجع `isCusip("5949181")` القيمة `false`.

```js
assert.equal(isCusip('5949181'), false);
```

# --seed--

## --seed-contents--

```js
function isCusip(s) {

}
```

# --solutions--

```js
function isCusip(s) {
  if (s.length != 9) return false;
  var sum = 0;
  var ASCII = x => x.charCodeAt(0);
  for (var i = 0; i < 7; i++) {
    var c = s.charCodeAt(i);

    var v;
    if (c >= ASCII('0') && c <= ASCII('9')) {
      v = c - 48;
    } else if (c >= ASCII('A') && c <= ASCII('Z')) {
      v = c - 64; // lower case letters apparently invalid
    } else if (c == ASCII('*')) {
      v = 36;
    } else if (c == ASCII('@')) {
      v = 37;
    } else if (c == ASCII('#')) {
      v = 38;
    } else {
      return false;
    }
    if (i % 2 == 1) v *= 2; // check if odd as using 0-based indexing
    sum += Math.floor(v / 10) + (v % 10);
  }
  return s.charCodeAt(8) - 48 == (10 - (sum % 10)) % 10;
}
```
