---
id: 5a23c84252665b21eecc8040
title: مجموع مضاعفات 3 و 5
challengeType: 1
forumTopicId: 302332
dashedName: sum-multiples-of-3-and-5
---

# --description--

الهدف هو كتابة دالة تجد مجموع كل المضاعفات الموجبة لـ 3 أو 5 التي تقل عن *n*.

# --hints--

يجب أن تكون `sumMults` دالة.

```js
assert(typeof sumMults == 'function');
```

يجب أن تُرجع `sumMults(10)` رقمًا.

```js
assert(typeof sumMults(10) == 'number');
```

يجب أن تُرجع `sumMults(10)` القيمة `23`.

```js
assert.equal(sumMults(10), 23);
```

يجب أن تُرجع `sumMults(100)` القيمة `2318`.

```js
assert.equal(sumMults(100), 2318);
```

يجب أن تُرجع `sumMults(1000)` القيمة `233168`.

```js
assert.equal(sumMults(1000), 233168);
```

يجب أن تُرجع `sumMults(10000)` القيمة `23331668`.

```js
assert.equal(sumMults(10000), 23331668);
```

يجب أن تُرجع `sumMults(100000)` القيمة `2333316668`.

```js
assert.equal(sumMults(100000), 2333316668);
```

# --seed--

## --seed-contents--

```js
function sumMults(n) {

}
```

# --solutions--

```js
function sumMults(n) {
  var sum = 0;
  for (var i = 1; i < n; i++) {
    if (i % 3 == 0 || i % 5 == 0) sum += i;
  }
  return sum;
}
```
