---
id: 5a23c84252665b21eecc803f
title: مجموع أرقام عدد صحيح
challengeType: 1
forumTopicId: 302331
dashedName: sum-digits-of-an-integer
---

# --description--

اكتب دالة تأخذ سلسلة نصية كمعلمة. تمثل هذه السلسلة رقمًا يمكن أن يكون بأي أساس (أقل من 37) وأرجع مجموع أرقامه.

<ul>
<li><b>1</b><sub>10</sub> مجموعها <b>1</b></li>
  <li><b>1234</b><sub>10</sub> مجموعها <b>10</b></li>
  <li><b>fe</b><sub>16</sub> مجموعها <b>29</b></li>
  <li><b>f0e</b><sub>16</sub> مجموعها <b>29</b></li>
</ul>

# --hints--

`sumDigits` يجب أن تكون دالة.

```js
assert(typeof sumDigits == 'function');
```

`sumDigits("1")` يجب أن تُرجع رقمًا.

```js
assert(typeof sumDigits('1') == 'number');
```

`sumDigits("1")` يجب أن تُرجع `1`.

```js
assert.equal(sumDigits('1'), 1);
```

`sumDigits("12345")` يجب أن تُرجع `15`.

```js
assert.equal(sumDigits('12345'), 15);
```

`sumDigits("254")` يجب أن تُرجع `11`.

```js
assert.equal(sumDigits('254'), 11);
```

`sumDigits("fe")` يجب أن تُرجع `29`.

```js
assert.equal(sumDigits('fe'), 29);
```

`sumDigits("f0e")` يجب أن تُرجع `29`.

```js
assert.equal(sumDigits('f0e'), 29);
```

`sumDigits("999ABCXYZ")` يجب أن تُرجع `162`.

```js
assert.equal(sumDigits('999ABCXYZ'), 162);
```

# --seed--

## --seed-contents--

```js
function sumDigits(n) {

}
```

# --solutions--

```js
function sumDigits(n) {
  n += '';
  for (var s = 0, i = 0, e = n.length; i < e; i += 1)
    s += parseInt(n.charAt(i), 36);
  return s;
}
```
