---
id: 5a23c84252665b21eecc8036
title: إزالة رموز التحكم والحروف الموسعة من سلسلة نصية
challengeType: 1
forumTopicId: 302327
dashedName: strip-control-codes-and-extended-characters-from-a-string
---

# --description--

المهمة هي إزالة رموز التحكم والحروف الموسعة من سلسلة نصية. يجب أن يوضح الحل كيفية تحقيق كل من النتائج التالية: سلسلة نصية تمت إزالة رموز التحكم والحروف الموسعة منها. في نظام ASCII، رموز التحكم لها رموز عشرية من 0 إلى 31 و127. في نظام يعتمد على ASCII، إذا أُزيلت رموز التحكم، ستكون جميع حروف السلسلة الناتجة ضمن النطاق من 32 إلى 126 عشرية في جدول ASCII. في نظام غير معتمد على ASCII، نعتبر الحروف التي لا تمتلك رمزًا مطابقًا في جدول ASCII (ضمن النطاق من 32 إلى 126 عشرية) حروفًا موسعة لغرض هذه المهمة.

# --hints--

`strip` يجب أن تكون دالة.

```js
assert(typeof strip == 'function');
```

`strip("abc")` يجب أن تُرجع سلسلة نصية.

```js
assert(typeof strip('abc') == 'string');
```

`strip("\ba\x00b\n\rc\fd\xc3")` يجب أن تُرجع `"abcd"`.

```js
assert.equal(strip('\ba\x00b\n\rc\fd\xc3'), 'abcd');
```

`strip("\u0000\n abc\u00E9def\u007F")` يجب أن تُرجع `" abcdef"`.

```js
assert.equal(strip('\u0000\n abc\u00E9def\u007F'), ' abcdef');
```

`strip("a\n\tb\u2102d\u2147f")` يجب أن تُرجع `"abdf"`.

```js
assert.equal(strip('a\n\tb\u2102d\u2147f'), 'abdf');
```

`strip("Français.")` يجب أن تُرجع `"Franais."`.

```js
assert.equal(strip('Français.'), 'Franais.');
```

`strip("123\tabc\u0007DEF\u007F+-*/€æŧðłþ")` يجب أن تُرجع `"123abcDEF+-*/"`.

```js
assert.equal(strip('123\tabc\u0007DEF\u007F+-*/€æŧðłþ'), '123abcDEF+-*/');
```

# --seed--

## --seed-contents--

```js
function strip(s) {

}
```

# --solutions--

```js
function strip(s) {
  return s
    .split('')
    .filter(function(x) {
      var n = x.charCodeAt(0);

      return 31 < n && 127 > n;
    })
    .join('');
}
```
