---
id: 5a23c84252665b21eecc7e82
title: القاسم المشترك الأكبر
challengeType: 1
forumTopicId: 302277
dashedName: greatest-common-divisor
---

# --description--

اكتب دالة تُرجع القاسم المشترك الأكبر لعددين صحيحين.

# --hints--

يجب أن تكون `gcd` دالة.

```js
assert(typeof gcd == 'function');
```

يجب أن تُرجع `gcd(24,36)` عددًا.

```js
assert(typeof gcd(24, 36) == 'number');
```

يجب أن تُرجع `gcd(24,36)` القيمة `12`.

```js
assert.equal(gcd(24, 36), 12);
```

يجب أن تُرجع `gcd(30,48)` القيمة `6`.

```js
assert.equal(gcd(30, 48), 6);
```

يجب أن تُرجع `gcd(10,15)` القيمة `5`.

```js
assert.equal(gcd(10, 15), 5);
```

يجب أن تُرجع `gcd(100,25)` القيمة `25`.

```js
assert.equal(gcd(100, 25), 25);
```

يجب أن تُرجع `gcd(13,250)` القيمة `1`.

```js
assert.equal(gcd(13, 250), 1);
```

يجب أن تُرجع `gcd(1300,250)` القيمة `50`.

```js
assert.equal(gcd(1300, 250), 50);
```

# --seed--

## --seed-contents--

```js
function gcd(a, b) {

}
```

# --solutions--

```js
function gcd(a, b) {
  return b==0 ? Math.abs(a):gcd(b, a % b);
}
```
