---
id: 5a23c84252665b21eecc8038
title: Subleq
challengeType: 1
forumTopicId: 302328
dashedName: subleq
---

# --description--

Subleq هو مثال على حاسوب بمجموعة تعليمات واحدة (One-Instruction Set Computer - OISC).

سُمي بهذا الاسم نسبةً إلى تعليمة واحدة فقط فيه، وهي **SU**btract و**B**ranch إذا كان أقل من أو **EQ**ual للصفر.

مهمتك هي إنشاء مفسر يحاكي هذه الآلة.

ذاكرة الآلة تتكون من مصفوفة من الأعداد الصحيحة ذات إشارة. أي حجم كلمة معقول مقبول، لكن يجب أن تستوعب الذاكرة أعدادًا سالبة وموجبة على حد سواء.

يبدأ التنفيذ بمؤشر التعليمات موجهًا إلى الكلمة الأولى، وهي العنوان 0. ويتبع الخطوات التالية:

<ol>
<li>لتكن A وB وC هي القيم المخزنة في ثلاث كلمات متتالية في الذاكرة تبدأ من مؤشر التعليمات.</li>
  <li>حرك مؤشر التعليمات 3 كلمات ليشير إلى العنوان الذي يلي العنوان الذي يحتوي على C.</li>
  <li>إذا كانت A تساوي -1، يُقرأ حرف من الإدخال القياسي ويُخزن رمز الحرف في العنوان المحدد بواسطة B. لا يُستخدم C.</li>
  <li>إذا كانت B تساوي -1، يُفسر الرقم الموجود في العنوان المحدد بواسطة A كرمز حرف ويُطبع الحرف المقابل. لا يُستخدم C أيضًا.</li>
  <li>خلاف ذلك، يُعامل كل من A وB كعناوين مواقع في الذاكرة. يُطرح الرقم الموجود في العنوان المحدد بواسطة A من الرقم في العنوان المحدد بواسطة B (ويُخزن الناتج في العنوان B). إذا كان الناتج صفرًا أو سالبًا، يصبح C هو مؤشر التعليمات الجديد.</li>
  <li>إذا أصبح مؤشر التعليمات سالبًا، يتوقف التنفيذ.</li>
</ol>

يمكنك التعامل مع العناوين السالبة الأخرى غير -1 على أنها تعادل -1، أو توليد خطأ، حسب ما تراه مناسبًا.

يجب أن يقبل الحل برنامجًا للتنفيذ على الآلة، منفصلًا عن المُدخَل المقدم للبرنامج نفسه.

يجب أن يكون هذا البرنامج في "شفرة آلة" subleq الخام - أرقام عشرية مفصولة بمسافات بيضاء، بدون أسماء رمزية أو امتدادات على مستوى التجميع، ليتم تحميلها في الذاكرة بدءًا من العنوان 0. اعرض ناتج الحل عند تغذيته بهذا البرنامج "Hello, world!". (لاحظ أن المثال يفترض ASCII أو مجموعة تحتويها، مثل أي من مجموعات Latin-N أو Unicode. يمكنك ترجمته إلى مجموعة أحرف أخرى إذا كان تنفيذك على بيئة غير متوافقة مع ASCII.)

<pre>15 17 -1 17 -1 -1 16 1 -1 16 3 -1 15 15 0 0 -1 72 101 108 108 111 44 32 119 111 114 108 100 33 10 0</pre>

وهذا يعادل شيئًا مثل هذا في لغة تجميع افتراضية:

<pre>start:
    صفر، رسالة، -1
    رسالة، -1، -1
    سالب1، start+1، -1
    سالب1، start+3، -1
    صفر، صفر، start
صفر: 0
سالب1: -1
رسالة: "Hello, world!\n\0"
</pre>

# --instructions--

اكتب دالة تأخذ مصفوفة من الأعداد الصحيحة كمعلمة. هذه تمثل عناصر الذاكرة. يجب أن تفسر الدالة المتتالية وتُرجع سلسلة النص الناتجة. في هذه المهمة، افترض أنه لا يوجد إدخال قياسي.

# --hints--

`Subleq` يجب أن تكون دالة.

```js
assert(typeof Subleq == 'function');
```

`Subleq([15, 17, -1, 17, -1, -1, 16, 1, -1, 16, 3, -1, 15, 15, 0, 0, -1, 72, 101, 108, 108, 111, 44, 32, 119, 111, 114, 108, 100, 33, 0])` يجب أن تُرجع سلسلة نصية.

```js
assert(
  typeof Subleq([
    15,
    17,
    -1,
    17,
    -1,
    -1,
    16,
    1,
    -1,
    16,
    3,
    -1,
    15,
    15,
    0,
    0,
    -1,
    72,
    101,
    108,
    108,
    111,
    44,
    32,
    119,
    111,
    114,
    108,
    100,
    33,
    0
  ]) == 'string'
);
```

`Subleq([15, 17, -1, 17, -1, -1, 16, 1, -1, 16, 3, -1, 15, 15, 0, 0, -1, 72, 101, 108, 108, 111, 44, 32, 119, 111, 114, 108, 100, 33, 0])` يجب أن تُرجع `"Hello, world!"`.

```js
assert.equal(
  Subleq([
    15,
    17,
    -1,
    17,
    -1,
    -1,
    16,
    1,
    -1,
    16,
    3,
    -1,
    15,
    15,
    0,
    0,
    -1,
    72,
    101,
    108,
    108,
    111,
    44,
    32,
    119,
    111,
    114,
    108,
    100,
    33,
    0
  ]),
  'Hello, world!'
);
```

# --seed--

## --seed-contents--

```js
function Subleq(mem) {

}
```

# --solutions--

```js
function Subleq(mem) {
  var out = '';
  var instructionPointer = 0;
  do {
    var a = mem[instructionPointer];
    var b = mem[instructionPointer + 1];
    if (a === -1) {
    } else if (b === -1) {
      out += String.fromCharCode(mem[a]);
    } else {
      mem[b] -= mem[a];
      if (mem[b] < 1) {
        instructionPointer = mem[instructionPointer + 2];
        continue;
      }
    }
    instructionPointer += 3;
  } while (instructionPointer >= 0);

  return out;
}
```
