---
id: 5a23c84252665b21eecc7edf
title: المضاعف المشترك الأصغر
challengeType: 1
forumTopicId: 302301
dashedName: least-common-multiple
---

# --description--

المضاعف المشترك الأصغر للعددين 12 و18 هو 36، لأن 12 هو عامل (12 × 3 = 36)، و18 هو عامل (18 × 2 = 36)، ولا يوجد عدد صحيح موجب أقل من 36 يحتوي على كلا العاملين. كحالة خاصة، إذا كان أي من $m$ أو $n$ يساوي صفرًا، فإن المضاعف المشترك الأصغر يكون صفرًا. إحدى الطرق لحساب المضاعف المشترك الأصغر هي تكرار جميع مضاعفات $m$ حتى تجد مضاعفًا يكون أيضًا من مضاعفات $n$. إذا كان لديك بالفعل $gcd$ لـ<a href="https://rosettacode.org/wiki/Greatest_common_divisor" target="_blank" rel="noopener noreferrer nofollow">القاسم المشترك الأكبر</a>، فإن هذه الصيغة تحسب $lcm$.

$$
\\operatorname{lcm}(m, n) = \\frac{|m \\times n|}{\\operatorname{gcd}(m, n)}
$$

# --instructions--

احسب المضاعف المشترك الأصغر لمصفوفة من الأعداد الصحيحة. بالنظر إلى *m* و *n*، المضاعف المشترك الأصغر هو أصغر عدد صحيح موجب يحتوي على كل من *m* و *n* كعوامل.

# --hints--

`LCM` يجب أن تكون دالة.

```js
assert(typeof LCM == 'function');
```

`LCM([2, 4, 8])` يجب أن تُرجع رقمًا.

```js
assert(typeof LCM([2, 4, 8]) == 'number');
```

`LCM([2, 4, 8])` يجب أن تُرجع `8`.

```js
assert.equal(LCM([2, 4, 8]), 8);
```

`LCM([4, 8, 12])` يجب أن تُرجع `24`.

```js
assert.equal(LCM([4, 8, 12]), 24);
```

`LCM([3, 4, 5, 12, 40])` يجب أن تُرجع `120`.

```js
assert.equal(LCM([3, 4, 5, 12, 40]), 120);
```

`LCM([11, 33, 90])` يجب أن تُرجع `990`.

```js
assert.equal(LCM([11, 33, 90]), 990);
```

`LCM([-50, 25, -45, -18, 90, 447])` يجب أن تُرجع `67050`.

```js
assert.equal(LCM([-50, 25, -45, -18, 90, 447]), 67050);
```

# --seed--

## --seed-contents--

```js
function LCM(A) {

}
```

# --solutions--

```js
function LCM(A) {
  var n = A.length,
    a = Math.abs(A[0]);
  for (var i = 1; i < n; i++) {
    var b = Math.abs(A[i]),
      c = a;
    while (a && b) {
      a > b ? (a %= b) : (b %= a);
    }
    a = Math.abs(c * A[i]) / (a + b);
  }
  return a;
}
```
