---
id: 5a23c84252665b21eecc8010
title: خوارزميات الترتيب/ترتيب شيل
challengeType: 1
forumTopicId: 302317
dashedName: sorting-algorithmsshell-sort
---

# --description--

اكتب دالة لترتيب مصفوفة من العناصر باستخدام خوارزمية ترتيب شيل، وهي ترتيب بتناقص الزيادة. يجب أن تُرجع الدالة المصفوفة المرتبة.

ترتيب شيل (المعروف أيضًا باسم Shellsort أو طريقة شيل) سُمي على اسم مخترعه، دونالد شيل، الذي نشر الخوارزمية في عام 1959.

ترتيب شيل هو متتالية من عمليات الإدراج المتداخلة تعتمد على متتالية زيادة. يتم تقليل حجم الزيادة بعد كل تمريرة حتى يصبح حجم الزيادة 1.

عندما يكون حجم الزيادة 1، يكون الترتيب هو ترتيب إدراج أساسي، لكن في هذه المرحلة تكون البيانات مضمونة تقريبًا أن تكون مرتبة، وهو "أفضل حالة" لترتيب الإدراج.

أي متتالية ستُرتب البيانات طالما تنتهي بالعدد 1، لكن بعضها يعمل بشكل أفضل من الآخر.

أظهرت الدراسات التجريبية أن متتالية زيادة هندسية بنسبة حوالي 2.2 تعمل جيدًا في التطبيق العملي.

# --hints--

يجب أن تكون `shellSort` دالة.

```js
assert(typeof shellSort == 'function');
```

يجب أن تُرجع `shellSort([25, 32, 12, 7, 20])` مصفوفة.

```js
assert(Array.isArray(shellSort([25, 32, 12, 7, 20])));
```

يجب أن تُرجع `shellSort([25, 32, 12, 7, 20])` `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(shellSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

يجب أن تُرجع `shellSort([38, 45, 35, 8, 13])` `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(shellSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

يجب أن تُرجع `shellSort([43, 36, 20, 34, 24])` `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(shellSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

يجب أن تُرجع `shellSort([12, 33, 26, 18, 1, 16, 38])` `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(shellSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

يجب أن تُرجع `shellSort([3, 39, 48, 16, 1, 4, 29])` `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(shellSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function shellSort(a) {

}
```

# --solutions--

```js
function shellSort(a) {
  for (var h = a.length; h > 0; h = parseInt(h / 2)) {
    for (var i = h; i < a.length; i++) {
      var k = a[i];
      for (var j = i; j >= h && k < a[j - h]; j -= h) a[j] = a[j - h];
      a[j] = k;
    }
  }
  return a;
}
```
