---
id: 5a23c84252665b21eecc7ec5
title: مشكلة جوزيفوس
challengeType: 1
forumTopicId: 302294
dashedName: josephus-problem
---

# --description--

مشكلة جوزيفوس هي لغز رياضي بوصف قاتم: يقف $n$ سجينًا في دائرة، مرقمين بالتتابع من $0$ إلى $n-1$.

يمشي جلاد حول الدائرة، بدءًا من السجين رقم $0$، ويزيل كل سجين $k$-th ويقتله.

مع استمرار العملية، تصبح الدائرة أصغر وأصغر، حتى يبقى سجين واحد فقط، ثم يُطلق سراحه.

على سبيل المثال، إذا كان هناك $n=5$ سجناء و $k=2$، فإن ترتيب قتل السجناء (دعنا نسميه "متتالية القتل") سيكون 1، 3، 0، و4، والناجي سيكون رقم 2.

معطى أي $n, k > 0$، اكتشف أي سجين سيكون الناجي النهائي.

في حادثة من هذا النوع، كان هناك 41 سجينًا وكان يُقتل كل سجين ثالث ($k=3$).

من بينهم كان هناك رجل ذكي اسمه جوزيفوس حل المشكلة، وقف في موقع الناجي، وعاش ليحكي القصة.

ما هو رقمه؟

# --instructions--

اكتب دالة تأخذ عدد السجناء الابتدائي و`k` كمعلمات وتُرجع رقم السجين الناجي.

# --hints--

يجب أن تكون `josephus` دالة.

```js
assert(typeof josephus == 'function');
```

يجب أن تُرجع `josephus(30,3)` رقمًا.

```js
assert(typeof josephus(30, 3) == 'number');
```

يجب أن تُرجع `josephus(30,3)` القيمة `28`.

```js
assert.equal(josephus(30, 3), 28);
```

يجب أن تُرجع `josephus(30,5)` القيمة `2`.

```js
assert.equal(josephus(30, 5), 2);
```

يجب أن تُرجع `josephus(20,2)` القيمة `8`.

```js
assert.equal(josephus(20, 2), 8);
```

يجب أن تُرجع `josephus(17,6)` القيمة `1`.

```js
assert.equal(josephus(17, 6), 1);
```

يجب أن تُرجع `josephus(29,4)` القيمة `1`.

```js
assert.equal(josephus(29, 4), 1);
```

# --seed--

## --seed-contents--

```js
function josephus(init, kill) {

}
```

# --solutions--

```js
function josephus(init, kill) {
  const arr = Array.from(Array(init).keys());
  let curr = -1
  while (arr.length > 1) {
    curr = (curr + kill) % arr.length;
    arr.splice(curr, 1);
    curr--;
  }
  return arr[0];
}
```
