---
id: 5a23c84252665b21eecc7ec4
title: JortSort
challengeType: 1
forumTopicId: 302293
dashedName: jortsort
---

# --description--

jortSort هي مجموعة أدوات للترتيب تجعل المستخدم يقوم بالعمل وتضمن الكفاءة لأنك لن تحتاج إلى الترتيب مرة أخرى. قدمتها في الأصل Jenn "Moneydollars" Schiffer في مؤتمر JSConf2014 المرموق.

يجب أن تكون jortSort دالة تأخذ مصفوفة واحدة من الكائنات القابلة للمقارنة كمعلمتها. ثم ترتب المصفوفة بترتيب تصاعدي وتقارن المصفوفة المرتبة بالمصفوفة الأصلية المقدمة. إذا تطابقت المصفوفات (أي أن المصفوفة الأصلية كانت مرتبة بالفعل)، تعيد الدالة true. إذا لم تتطابق المصفوفات (أي أن المصفوفة الأصلية لم تكن مرتبة)، تعيد الدالة false.

# --hints--

`jortsort` يجب أن تكون دالة.

```js
assert(typeof jortsort == 'function');
```

`jortsort([1,2,3,4,5])` يجب أن تعيد قيمة منطقية.

```js
assert(typeof jortsort([1, 2, 3, 4, 5]) == 'boolean');
```

`jortsort([1,2,3,4,5])` يجب أن تعيد `true`.

```js
assert.equal(jortsort([1, 2, 3, 4, 5]), true);
```

`jortsort([1,2,13,4,5])` يجب أن تعيد `false`.

```js
assert.equal(jortsort([1, 2, 13, 4, 5]), false);
```

`jortsort([12,4,51,2,4])` يجب أن تعيد `false`.

```js
assert.equal(jortsort([12, 4, 51, 2, 4]), false);
```

`jortsort([1,2])` يجب أن تعيد `true`.

```js
assert.equal(jortsort([1, 2]), true);
```

`jortsort([5,4,3,2,1])` يجب أن تعيد `false`.

```js
assert.equal(jortsort([5, 4, 3, 2, 1]), false);
```

`jortsort([1,1,1,1,1])` يجب أن تعيد `true`.

```js
assert.equal(jortsort([1, 1, 1, 1, 1]), true);
```

# --seed--

## --seed-contents--

```js
function jortsort(array) {

}
```

# --solutions--

```js
function jortsort(array) {
  // sort the array
  var originalArray = array.slice(0);
  array.sort( function(a,b){return a - b} );

  // compare to see if it was originally sorted
  for (var i = 0; i < originalArray.length; ++i) {
    if (originalArray[i] !== array[i]) return false;
  }

  return true;
};
```
