---
id: 5a23c84252665b21eecc800b
title: خوارزميات الترتيب/ترتيب الفطائر
challengeType: 1
forumTopicId: 302315
dashedName: sorting-algorithmspancake-sort
---

# --description--

اكتب دالة لترتيب مصفوفة من الأعداد الصحيحة (بحجم مناسب) بترتيب تصاعدي باستخدام ترتيب الفطائر. يجب أن تُرجع الدالة المصفوفة المرتبة.

باختصار، بدلًا من ترتيب العناصر الفردية، العملية الوحيدة المسموح بها هي "قلب" أحد طرفي القائمة، كما يلي:

<pre>قبل:
<b>6 7 8 9</b> 2 5 3 4 1<br>
بعد:
<b>9 8 7 6</b> 2 5 3 4 1
</pre>

يمكن قلب طرف واحد فقط من القائمة؛ ويفضل أن يكون الطرف الأدنى، لكن الطرف الأعلى مقبول إذا كان أسهل في الشفرة أو يعمل بشكل أفضل، لكن يجب أن يكون نفس الطرف طوال الحل بأكمله. (لا يمكن تغيير الطرف المقلب بشكل عشوائي.)

# --hints--

يجب أن تكون `pancakeSort` دالة.

```js
assert(typeof pancakeSort == 'function');
```

يجب أن تُرجع `pancakeSort([25, 32, 12, 7, 20])` مصفوفة.

```js
assert(Array.isArray(pancakeSort([25, 32, 12, 7, 20])));
```

يجب أن تُرجع `pancakeSort([25, 32, 12, 7, 20])` `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(pancakeSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

يجب أن تُرجع `pancakeSort([38, 45, 35, 8, 13])` `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(pancakeSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

يجب أن تُرجع `pancakeSort([43, 36, 20, 34, 24])` `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(pancakeSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

يجب أن تُرجع `pancakeSort([12, 33, 26, 18, 1, 16, 38])` `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(pancakeSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

يجب أن تُرجع `pancakeSort([3, 39, 48, 16, 1, 4, 29])` `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(pancakeSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function pancakeSort(arr) {

}
```

# --solutions--

```js
function pancakeSort(arr) {
  for (var i = arr.length - 1; i >= 1; i--) {
    // find the index of the largest element not yet sorted
    var max_idx = 0;
    var max = arr[0];
    for (var j = 1; j <= i; j++) {
      if (arr[j] > max) {
        max = arr[j];
        max_idx = j;
      }
    }

    if (max_idx == i) continue; // element already in place

    var new_slice;

    // flip arr max element to index 0
    if (max_idx > 0) {
      new_slice = arr.slice(0, max_idx + 1).reverse();
      for (var j = 0; j <= max_idx; j++) arr[j] = new_slice[j];
    }

    // then flip the max element to its place
    new_slice = arr.slice(0, i + 1).reverse();
    for (var j = 0; j <= i; j++) arr[j] = new_slice[j];
  }
  return arr;
}
```
