---
id: 5a23c84252665b21eecc8028
title: متتالية ستيرن-بروكوت
challengeType: 1
forumTopicId: 302324
dashedName: stern-brocot-sequence
---

# --description--

لهذه المهمة، تُولد متتالية ستيرن-بروكوت بواسطة خوارزمية مشابهة لتلك المستخدمة في توليد <a href="https://rosettacode.org/wiki/Fibonacci_sequence" target="_blank" rel="noopener noreferrer nofollow">متتالية فيبوناتشي</a>.

<ol>
<li>العنصران الأول والثاني من المتتالية كلاهما 1:</li>
    <ul><li>1, 1</li></ul>
<li>ابدأ بالنظر إلى العضو الثاني في المتتالية</li>
  <li>اجمع العضو الجاري النظر فيه مع العضو السابق له، (1 + 1) = 2، وأضفه إلى نهاية
    المتتالية:</li>
    <ul><li>1, 1, 2</li></ul>
<li>أضف العنصر الجاري النظر إليه إلى نهاية المتتالية:</li>
    <ul><li>1, 1, 2, 1</li></ul>
<li>انظر إلى العنصر التالي في السلسلة، (العنصر الثالث أي 2)</li>
  <li>اذهب إلى 3 </li>
    <ul>
      <li></li>
<li> ─── بتوسيع حلقة أخرى نحصل على: ───</li>
      <li></li>
    </ul>
<li>اجمع العضو الحالي من المتتالية والعضو السابق له، (2 + 1) = 3، وأضفه إلى نهاية
    المتتالية:</li>
    <ul><li>1, 1, 2, 1, 3</li></ul>
<li>أضف العنصر الجاري النظر إليه إلى نهاية المتتالية:</li>
    <ul><li>1, 1, 2, 1, 3, 2</li></ul>
<li>انظر إلى العنصر التالي في السلسلة، (العنصر الرابع أي 1)</li>
</ol>

# --instructions--

أنشئ دالة تُرجع الموقع في متتالية ستيرن-بروكوت حيث يُقابل $ n $ لأول مرة، مع توليد المتتالية بالطريقة الموضحة أعلاه. لاحظ أن هذه المتتالية تستخدم فهرسة تبدأ من 1.

# --hints--

يجب أن تكون `sternBrocot` دالة.

```js
assert(typeof sternBrocot == 'function');
```

يجب أن تُرجع `sternBrocot(2)` رقمًا.

```js
assert(typeof sternBrocot(2) == 'number');
```

يجب أن تُرجع `sternBrocot(2)` القيمة `3`.

```js
assert.equal(sternBrocot(2), 3);
```

يجب أن تُرجع `sternBrocot(3)` القيمة `5`.

```js
assert.equal(sternBrocot(3), 5);
```

يجب أن تُرجع `sternBrocot(5)` القيمة `11`.

```js
assert.equal(sternBrocot(5), 11);
```

يجب أن تُرجع `sternBrocot(7)` القيمة `19`.

```js
assert.equal(sternBrocot(7), 19);
```

يجب أن تُرجع `sternBrocot(10)` القيمة `39`.

```js
assert.equal(sternBrocot(10), 39);
```

# --seed--

## --seed-contents--

```js
function sternBrocot(num) {

}
```

# --solutions--

```js
function sternBrocot(num) {
  function f(n) {
    return n < 2
      ? n
      : n & 1
      ? f(Math.floor(n / 2)) + f(Math.floor(n / 2 + 1))
      : f(Math.floor(n / 2));
  }

  function gcd(a, b) {
    return a ? (a < b ? gcd(b % a, a) : gcd(a % b, b)) : b;
  }
  var n;
  for (n = 1; f(n) != num; n++);
  return n;
}
```
