---
id: 5a23c84252665b21eecc7eb0
title: I قبل E إلا بعد C
challengeType: 1
forumTopicId: 302288
dashedName: i-before-e-except-after-c
---

# --description--

**<a href="http://www.rosettacode.org/wiki/I_before_E_except_after_C" target="_blank" rel="noopener noreferrer nofollow">"I قبل E، إلا بعد C"</a>** قاعدة عامة في تهجئة اللغة الإنجليزية. إذا لم تكن متأكدًا مما إذا كانت الكلمة تُهجى بالثنائي `ei` أو `ie`، فإن القافية تشير إلى أن الترتيب الصحيح هو `ie` ما لم يسبق الحرف `c`، وفي هذه الحالة قد يكون `ei`.

باستخدام الكلمات المقدمة، تحقق مما إذا كان البندان الفرعيان في العبارة معقولين بشكل فردي:

<ol>
  <li>
<i>"I قبل E عندما لا يسبقه C".</i>
  </li>
  <li>
<i>"E قبل I عندما يسبقه C".</i>
  </li>
</ol>

إذا كان كلا الجزأين الفرعيين معقولين، فيمكن القول إن العبارة الأصلية معقولة.

# --instructions--

اكتب دالة تقبل كلمة وتتحقق مما إذا كانت الكلمة تتبع هذه القاعدة. يجب أن تُرجع الدالة true إذا اتبعت الكلمة القاعدة وfalse إذا لم تفعل.

# --hints--

`IBeforeExceptC` يجب أن تكون دالة.

```js
assert(typeof IBeforeExceptC == 'function');
```

`IBeforeExceptC("receive")` يجب أن تُرجع قيمة منطقية.

```js
assert(typeof IBeforeExceptC('receive') == 'boolean');
```

`IBeforeExceptC("receive")` يجب أن تُرجع `true`.

```js
assert.equal(IBeforeExceptC('receive'), true);
```

`IBeforeExceptC("science")` يجب أن تُرجع `false`.

```js
assert.equal(IBeforeExceptC('science'), false);
```

`IBeforeExceptC("imperceivable")` يجب أن تُرجع `true`.

```js
assert.equal(IBeforeExceptC('imperceivable'), true);
```

`IBeforeExceptC("inconceivable")` يجب أن تُرجع `true`.

```js
assert.equal(IBeforeExceptC('inconceivable'), true);
```

`IBeforeExceptC("insufficient")` يجب أن تُرجع `false`.

```js
assert.equal(IBeforeExceptC('insufficient'), false);
```

`IBeforeExceptC("omniscient")` يجب أن تُرجع `false`.

```js
assert.equal(IBeforeExceptC('omniscient'), false);
```

# --seed--

## --seed-contents--

```js
function IBeforeExceptC(word) {

}
```

# --solutions--

```js
function IBeforeExceptC(word)
{
    if(word.indexOf("c")==-1 && word.indexOf("ie")!=-1)
        return true;
    else if(word.indexOf("cei")!=-1)
        return true;
    return false;
}
```
