---
id: 5a23c84252665b21eecc7eca
title: أعداد كابريكار
challengeType: 1
forumTopicId: 302296
dashedName: kaprekar-numbers
---

# --description--

العدد الصحيح الموجب هو عدد كابريكار إذا:

<ul>
<li>كان 1، أو،</li>
  <li>يمكن تقسيم التمثيل العشري لمربعه مرة واحدة إلى جزأين يتكون كل منهما من أعداد صحيحة موجبة مجموعها يساوي العدد الأصلي.</li>
</ul>

لاحظ أن التقسيم الذي ينتج جزءًا يتكون فقط من أصفار غير صالح، لأن 0 لا يُعتبر عددًا موجبًا.

أمثلة على أعداد كابريكار:

<ul>
<li><code>2223</code> هو عدد كابريكار، حيث <code>2223 * 2223 = 4941729</code>، ويمكن تقسيم <code>4941729</code> إلى <code>494</code> و<code>1729</code>، و<code>494 + 1729 = 2223</code></li>
  <li>سلسلة أعداد كابريكار معروفة باسم A006886، وتبدأ بـ <code>1, 9, 45, 55, ...</code></li>
</ul>

# --instructions--

اكتب دالة تأخذ رقمًا $n$، وقاعدة $bs$، وتُرجع true إذا كان الرقم عدد كابريكار للقاعدة المعطاة. وإلا، تُرجع الدالة false.

# --hints--

يجب أن تكون `isKaprekar` دالة.

```js
assert(typeof isKaprekar == 'function');
```

يجب أن تُرجع `isKaprekar(1, 10)` قيمة منطقية.

```js
assert(typeof isKaprekar(1, 10) == 'boolean');
```

يجب أن تُرجع `isKaprekar(1, 10)` القيمة `true`.

```js
assert.equal(isKaprekar(1, 10), true);
```

يجب أن تُرجع `isKaprekar(9, 10)` القيمة `true`.

```js
assert.equal(isKaprekar(9, 10), true);
```

يجب أن تُرجع `isKaprekar(2223, 10)` القيمة `true`.

```js
assert.equal(isKaprekar(2223, 10), true);
```

يجب أن تُرجع `isKaprekar(22823, 10)` القيمة `false`.

```js
assert.equal(isKaprekar(22823, 10), false);
```

يجب أن تُرجع `isKaprekar(9, 17)` القيمة `false`.

```js
assert.equal(isKaprekar(9, 17), false);
```

يجب أن تُرجع `isKaprekar(225, 17)` القيمة `true`.

```js
assert.equal(isKaprekar(225, 17), true);
```

يجب أن تُرجع `isKaprekar(999, 17)` القيمة `false`.

```js
assert.equal(isKaprekar(999, 17), false);
```

# --seed--

## --seed-contents--

```js
function isKaprekar(n, bs) {

}
```

# --solutions--

```js
function isKaprekar(n, bs) {
  if (n < 1) return false;
  if (n == 1) return true;
  for (var a = n * n, b = 0, s = 1; a; s *= bs) {
    b += (a % bs) * s;
    a = Math.floor(a / bs);
    if (b && a + b == n) return true;
  }
  return false;
}
```
