---
id: 5a23c84252665b21eecc8002
title: خوارزميات الترتيب / Bogosort
challengeType: 1
forumTopicId: 302311
dashedName: sorting-algorithmsbogosort
---

# --description--

رتب قائمة أرقام باستخدام Bogosort.

Bogosort يعيد خلط مجموعة عشوائيًا حتى تصبح مرتبة.

"Bogosort" خوارزمية غير فعالة بشكل ساخر تُستخدم فقط كدعابة داخلية.

متوسط زمن تشغيلها هو O(n!) لأن احتمال أن ينتهي أي خلط معين لمجموعة بترتيب مرتب هو حوالي واحد من مضروب *n*، وأسوأ حالة هي لا نهائية إذ لا يوجد ضمان أن خلطًا عشوائيًا سينتج متتالية مرتبة.

أفضل حالة لها هي O(n) لأن مرورًا واحدًا عبر العناصر قد يكفي لترتيبها.

رمز زائف:

<pre><b>طالما لم تكن</b> InOrder(list) <b>افعل</b>
  Shuffle(list)
<b>انتهى</b>
</pre>

# --hints--

يجب أن تكون `bogosort` دالة.

```js
assert(typeof bogosort == 'function');
```

يجب أن تُرجع `bogosort([25, 32, 12, 7, 20])` مصفوفة.

```js
assert(Array.isArray(bogosort([25, 32, 12, 7, 20])));
```

يجب أن تُرجع `bogosort([25, 32, 12, 7, 20])` `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(bogosort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

يجب أن تُرجع `bogosort([38, 45, 35, 8, 13])` `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(bogosort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

يجب أن تُرجع `bogosort([43, 36, 20, 34, 24])` `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(bogosort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

يجب أن تُرجع `bogosort([12, 33, 26, 18, 1, 16, 38])` `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(bogosort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

يجب أن تُرجع `bogosort([3, 39, 48, 16, 1, 4, 29])` `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(bogosort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function bogosort(v) {

}
```

# --solutions--

```js
function bogosort(v) {
  function shuffle(v) {
    for (
      var j, x, i = v.length;
      i;
      j = Math.floor(Math.random() * i), x = v[--i], v[i] = v[j], v[j] = x
    );
    return v;
  }

  function isSorted(v) {
    for (var i = 1; i < v.length; i++) {
      if (v[i - 1] > v[i]) {
        return false;
      }
    }
    return true;
  }
  var sorted = false;
  while (sorted == false) {
    v = shuffle(v);
    sorted = isSorted(v);
  }
  return v;
}
```
