---
id: 5a23c84252665b21eecc8014
title: ثبات الترتيب
challengeType: 1
forumTopicId: 302308
dashedName: sort-stability
---

# --description--

عند فرز السجلات في جدول حسب عمود أو حقل معين، فإن <a href="https://www.freecodecamp.org/news/stability-in-sorting-algorithms-a-treatment-of-equality-fa3140a5a539/" target="_blank" rel="noopener noreferrer nofollow">الفرز المستقر</a> يحافظ دائمًا على الترتيب النسبي للسجلات التي لها نفس المفتاح.

على سبيل المثال، في هذا الجدول الخاص بالدول والمدن، سيحافظ الفرز المستقر على **العمود الثاني**، المدن، على بقاء Birmingham في الولايات المتحدة فوق Birmingham في المملكة المتحدة. (مع أن الفرز غير المستقر *قد* يضع Birmingham في الولايات المتحدة فوق Birmingham في المملكة المتحدة في هذه الحالة، فإن روتين الفرز المستقر *يضمن* ذلك).

<pre>المملكة المتحدة  لندن
الولايات المتحدة  نيويورك
الولايات المتحدة  برمنغهام
المملكة المتحدة  برمنغهام
</pre>

بنفس الطريقة، سيولد الفرز المستقر على العمود الأول فقط "UK London" كأول عنصر و"US Birmingham" كآخر عنصر (لأن ترتيب العناصر التي لها نفس الكلمة الأولى – "UK" أو "US" – سيُحفظ).

# --instructions--

اكتب دالة تأخذ مصفوفة ثنائية الأبعاد كمعلمة. كل عنصر يحتوي على عنصرين مشابهين للمثال أعلاه. يجب أن ترتب الدالة المصفوفة كما ذُكر سابقًا وتُرجع المصفوفة المرتبة.

# --hints--

يجب أن تكون `stableSort` دالة.

```js
assert(typeof stableSort == 'function');
```

يجب أن تُرجع `stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` مصفوفة.

```js
assert(
  Array.isArray(
    stableSort([
      ['UK', 'London'],
      ['US', 'New York'],
      ['US', 'Birmingham'],
      ['UK', 'Birmingham']
    ])
  )
);
```

يجب أن تُرجع `stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` القيمة `[["US", "Birmingham"], ["UK", "Birmingham"], ["UK", "London"], ["US", "New York"]]`.

```js
assert.deepEqual(
  stableSort([
    ['UK', 'London'],
    ['US', 'New York'],
    ['US', 'Birmingham'],
    ['UK', 'Birmingham']
  ]),
  [
    ['US', 'Birmingham'],
    ['UK', 'Birmingham'],
    ['UK', 'London'],
    ['US', 'New York']
  ]
);
```

يجب أن تُرجع `stableSort([[2, 2], [1, 2], [1, 4], [1, 5]])` القيمة `[[2, 2], [1, 2], [1, 4], [1, 5]]`.

```js
assert.deepEqual(
  stableSort([
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]),
  [
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]
);
```

يجب أن تُرجع `stableSort([[11, 55], [12, 45], [11, 45], [32, 45]])` القيمة `[[12, 45], [11, 45], [32, 45], [11, 55]]`.

```js
assert.deepEqual(
  stableSort([
    [11, 55],
    [12, 45],
    [11, 45],
    [32, 45]
  ]),
  [
    [12, 45],
    [11, 45],
    [32, 45],
    [11, 55]
  ]
);
```

يجب أن تُرجع `stableSort([[10, 22], [1, 2], [1, 4], [1, 5], [10, 9]])` القيمة `[[1, 2], [1, 4], [1, 5], [10, 9], [10, 22]]`.

```js
assert.deepEqual(
  stableSort([
    [10, 22],
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9]
  ]),
  [
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9],
    [10, 22]
  ]
);
```

يجب أن تُرجع `stableSort([[55, 54], [12, 22], [31, 43], [31, 54], [10, 49]])` القيمة `[[12, 22], [31, 43], [10, 49], [55, 54], [31, 54]]`.

```js
assert.deepEqual(
  stableSort([
    [55, 54],
    [12, 22],
    [31, 43],
    [31, 54],
    [10, 49]
  ]),
  [
    [12, 22],
    [31, 43],
    [10, 49],
    [55, 54],
    [31, 54]
  ]
);
```

# --seed--

## --seed-contents--

```js
function stableSort(arr) {

}
```

# --solutions--

```js
function stableSort(arr) {
  arr.sort(function(a, b) {
    return a[1] < b[1] ? -1 : a[1] > b[1] ? 1 : 0;
  });
  return arr;
}
```
