---
id: 5a23c84252665b21eecc7edc
title: الجمعة الأخيرة من كل شهر
challengeType: 1
forumTopicId: 302299
dashedName: last-friday-of-each-month
---

# --description--

اكتب دالة تُرجع تاريخ الجمعة الأخيرة من شهر معين في سنة معينة.

# --hints--

يجب أن تكون `lastFriday` دالة.

```js
assert(typeof lastFriday == 'function');
```

يجب أن تُرجع `lastFriday(2018, 1)` رقمًا.

```js
assert(typeof lastFriday(2018, 1) == 'number');
```

يجب أن تُرجع `lastFriday(2018, 1)` القيمة `26`.

```js
assert.equal(lastFriday(2018, 1), 26);
```

يجب أن تُرجع `lastFriday(2017, 2)` القيمة `24`.

```js
assert.equal(lastFriday(2017, 2), 24);
```

يجب أن تُرجع `lastFriday(2012, 3)` القيمة `30`.

```js
assert.equal(lastFriday(2012, 3), 30);
```

يجب أن تُرجع `lastFriday(1900, 4)` القيمة `27`.

```js
assert.equal(lastFriday(1900, 4), 27);
```

يجب أن تُرجع `lastFriday(2000, 5)` القيمة `26`.

```js
assert.equal(lastFriday(2000, 5), 26);
```

يجب أن تُرجع `lastFriday(2006, 6)` القيمة `30`.

```js
assert.equal(lastFriday(2006, 6), 30);
```

يجب أن تُرجع `lastFriday(2010, 7)` القيمة `30`.

```js
assert.equal(lastFriday(2010, 7), 30);
```

يجب أن تُرجع `lastFriday(2005, 8)` القيمة `26`.

```js
assert.equal(lastFriday(2005, 8), 26);
```

# --seed--

## --seed-contents--

```js
function lastFriday(year, month) {

}
```

# --solutions--

```js
function lastFriday(year, month) {
  var i, last_day;
  i = 0;
  while (true) {
    last_day = new Date(year, month, i);
    if (last_day.getDay() === 5) {
      return last_day.getDate();
    }
    i -= 1;
  }
}
```
