---
id: 5a23c84252665b21eecc801c
title: مصفوفة لولبية
challengeType: 1
forumTopicId: 302321
dashedName: spiral-matrix
---

# --description--

أنشئ مصفوفة لولبية. *المصفوفة اللولبية* هي ترتيب مربع لأعداد طبيعية أولى N<sup>2</sup>، حيث تزداد الأعداد بالتتابع أثناء الدوران حول حواف المصفوفة بشكل لولبي نحو الداخل. على سبيل المثال، إذا كانت القيمة **5**، أنشئ هذه المصفوفة:

<pre>
0  1  2  3  4
15 16 17 18 5
14 23 24 19 6
13 22 21 20 7
12 11 10  9 8
</pre>

# --hints--

يجب أن تكون `spiralArray` دالة.

```js
assert(typeof spiralArray == 'function');
```

يجب أن تُرجع `spiralArray(3)` مصفوفة.

```js
assert(Array.isArray(spiralArray(3)));
```

يجب أن تُرجع `spiralArray(3)` القيمة `[[0, 1, 2],[7, 8, 3],[6, 5, 4]]`.

```js
assert.deepEqual(spiralArray(3), [
  [0, 1, 2],
  [7, 8, 3],
  [6, 5, 4]
]);
```

يجب أن تُرجع `spiralArray(4)` القيمة `[[0, 1, 2, 3],[11, 12, 13, 4],[10, 15, 14, 5],[9, 8, 7, 6]]`.

```js
assert.deepEqual(spiralArray(4), [
  [0, 1, 2, 3],
  [11, 12, 13, 4],
  [10, 15, 14, 5],
  [9, 8, 7, 6]
]);
```

يجب أن تُرجع `spiralArray(5)` القيمة `[[0, 1, 2, 3, 4],[15, 16, 17, 18, 5],[14, 23, 24, 19, 6],[13, 22, 21, 20, 7],[12, 11, 10, 9, 8]]`.

```js
assert.deepEqual(spiralArray(5), [
  [0, 1, 2, 3, 4],
  [15, 16, 17, 18, 5],
  [14, 23, 24, 19, 6],
  [13, 22, 21, 20, 7],
  [12, 11, 10, 9, 8]
]);
```

# --seed--

## --seed-contents--

```js
function spiralArray(n) {

}
```

# --solutions--

```js
function spiralArray(n) {
  var arr = Array(n),
    x = 0,
    y = n,
    total = n * n--,
    dx = 1,
    dy = 0,
    i = 0,
    j = 0;
  while (y) arr[--y] = [];
  while (i < total) {
    arr[y][x] = i++;
    x += dx;
    y += dy;
    if (++j == n) {
      if (dy < 0) {
        x++;
        y++;
        n -= 2;
      }
      j = dx;
      dx = -dy;
      dy = j;
      j = 0;
    }
  }
  return arr;
}
```
