---
id: 5a23c84252665b21eecc7e1e
title: الضرب الداخلي
challengeType: 1
forumTopicId: 302251
dashedName: dot-product
---

# --description--

أنشئ دالة لحساب **الضرب الداخلي**، المعروف أيضًا باسم **الضرب القياسي** بين متجهين.

# --hints--

يجب أن تكون `dotProduct` دالة.

```js
assert(typeof dotProduct == 'function');
```

يجب أن تُرجع `dotProduct([1, 3, -5], [4, -2, -1])` قيمة عددية.

```js
assert(typeof dotProduct([1, 3, -5], [4, -2, -1]) == 'number');
```

يجب أن تُرجع `dotProduct([1, 3, -5], [4, -2, -1])` القيمة `3`.

```js
assert.equal(dotProduct([1, 3, -5], [4, -2, -1]), 3);
```

يجب أن تُرجع `dotProduct([1, 2, 3, 4, 5], [6, 7, 8, 9, 10])` القيمة `130`.

```js
assert.equal(dotProduct([1, 2, 3, 4, 5], [6, 7, 8, 9, 10]), 130);
```

يجب أن تُرجع `dotProduct([5, 4, 3, 2], [7, 8, 9, 6])` القيمة `106`.

```js
assert.equal(dotProduct([5, 4, 3, 2], [7, 8, 9, 6]), 106);
```

يجب أن تُرجع `dotProduct([-5, 4, -3, 2], [-7, -8, 9, -6])` القيمة `-36`.

```js
assert.equal(dotProduct([-5, 4, -3, 2], [-7, -8, 9, -6]), -36);
```

يجب أن تُرجع `dotProduct([17, 27, 34, 43, 15], [62, 73, 48, 95, 110])` القيمة `10392`.

```js
assert.equal(dotProduct([17, 27, 34, 43, 15], [62, 73, 48, 95, 110]), 10392);
```

# --seed--

## --seed-contents--

```js
function dotProduct(ary1, ary2) {

}
```

# --solutions--

```js
function dotProduct(ary1, ary2) {
  var dotprod = 0;
  for (var i = 0; i < ary1.length; i++) dotprod += ary1[i] * ary2[i];
  return dotprod;
}
```
