---
id: 5a23c84252665b21eecc8013
title: خوارزميات الفرز/فرز الخيوط
challengeType: 1
forumTopicId: 302319
dashedName: sorting-algorithmsstrand-sort
---

# --description--

يُنشئ **فرز الخيوط** مجموعات جزئية مرتبة تُدمج لإنشاء النتيجة النهائية.

تأمل في `unsortedArray = [3, 1, 4, 2]`. اختر العنصر الأول `3` وانسخه إلى مصفوفة منفصلة. ابحث عن أي عنصر أكبر يتبع هذا العنصر. عندما تجد عنصرًا أكبر، في هذه الحالة `4`، انسخه إلى المصفوفة المنفصلة، `[3, 4]`، وقارن العناصر التالية مع هذه القيمة الجديدة، `4`.

بعد أن تصل إلى نهاية المصفوفة، احذف العناصر التي نسختها، `[3, 4]`، وابدأ مجددًا بالعنصر الأول المتبقي في `unsortedArray`، في هذه الحالة `1`.

اتباع هذه العملية يؤدي إلى مصفوفتين مرتبتين، `[3, 4]` و `[1, 2]`. ادمج هاتين المصفوفتين لإنشاء `strandSortedArray`.

```js
const unsortedArray = [3, 1, 4, 2];
const strandsortedArray = [1, 2, 3, 4];
```

اكتب دالة لفرز مصفوفة باستخدام **فرز الخيوط**. يجب أن تُرجع الدالة المصفوفة المرتبة.


# --hints--

يجب أن تكون `strandSort` دالة.

```js
assert(typeof strandSort == 'function');
```

يجب أن تُرجع `strandSort([25, 32, 12, 7, 20])` مصفوفة.

```js
assert(Array.isArray(strandSort([25, 32, 12, 7, 20])));
```

يجب أن تُرجع `strandSort([25, 32, 12, 7, 20])` `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(strandSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

يجب أن تُرجع `strandSort([38, 45, 35, 8, 13])` `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(strandSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

يجب أن تُرجع `strandSort([43, 36, 20, 34, 24])` `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(strandSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

يجب أن تُرجع `strandSort([12, 33, 26, 18, 1, 16, 38])` `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(strandSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

يجب أن تُرجع `strandSort([3, 39, 48, 16, 1, 4, 29])` `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(strandSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function strandSort(list) {

}
```

# --solutions--

```js
function strandSort(list) {
  function merge(left, right) {
    var result = [];
    while (left.length != 0 && right.length != 0) {
      if (left[0] <= right[0]) result.push(left.shift());
      else result.push(right.shift());
    }
    result.push.apply(result, left);
    result.push.apply(result, right);
    return result;
  }

  if (list.length <= 1) return list;
  var result = [];
  while (list.length > 0) {
    var sorted = [];
    sorted.push(list.shift());
    var len = list.length;
    for (var i = 1; i < len; i++) {
      var elem = list[i];
      if (sorted[i - 1] <= elem) {
        sorted.push(elem);
        sorted.splice(i, 1);
      }
    }

    result = merge(result, sorted);
  }
  return result;
}
```
