---
id: 5a23c84252665b21eecc7eb1
title: مصفوفة الوحدة
challengeType: 1
forumTopicId: 302290
dashedName: identity-matrix
---

# --description--

مصفوفة *الوحدة* هي مصفوفة مربعة بحجم \\( n \\times n \\)، حيث تكون كل العناصر القطرية `1`s (واحدات)، وكل العناصر الأخرى `0`s (أصفار).

<ul>
<li style='list-style: none;'>\(\displaystyle I_{n}=\begin{bmatrix} 1 &#x26; 0 &#x26; 0 \cr 0 &#x26; 1 &#x26; 0 \cr 0 &#x26; 0 &#x26; 1 \cr \end{bmatrix}\) مصفوفة الوحدة</li>
</ul>

# --instructions--

اكتب دالة تأخذ رقمًا `n` كمعلمة وتُرجع مصفوفة الوحدة من الرتبة \\( n \\times n \\).

# --before-each--

```js
let results=[[ [ 1 ] ],
[ [ 1, 0 ], [ 0, 1 ] ],
[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ],
[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]]
```

# --hints--

يجب أن تكون `idMatrix` دالة.

```js
assert(typeof idMatrix == 'function');
```

يجب أن تُرجع `idMatrix(1)` مصفوفة.

```js
assert(Array.isArray(idMatrix(1)));
```

يجب أن تُرجع `idMatrix(1)` `[ [ 1 ] ]`.

```js
assert.deepEqual(idMatrix(1), results[0]);
```

يجب أن تُرجع `idMatrix(2)` `[ [ 1, 0 ], [ 0, 1 ] ]`.

```js
assert.deepEqual(idMatrix(2), results[1]);
```

يجب أن تُرجع `idMatrix(3)` `[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ]`.

```js
assert.deepEqual(idMatrix(3), results[2]);
```

يجب أن تُرجع `idMatrix(4)` `[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]`.

```js
assert.deepEqual(idMatrix(4), results[3]);
```

# --seed--

## --seed-contents--

```js
function idMatrix(n) {

}
```

# --solutions--

```js
function idMatrix(n) {
    return Array.apply(null, new Array(n)).map(function (x, i, xs) {
        return xs.map(function (_, k) {
            return i === k ? 1 : 0;
        })
    });
}
```
