---
id: 5a23c84252665b21eecc8005
title: خوارزميات الترتيب/ترتيب المشط
challengeType: 1
forumTopicId: 302313
dashedName: sorting-algorithmscomb-sort
---

# --description--

نفذ *ترتيب المشط*.

**ترتيب المشط** هو نوع من أنواع ترتيب الفقاعات.

مثل ترتيب شيل، يزيد ترتيب المشط الفجوة المستخدمة في المقارنات والتبادلات.

أفضل طريقة هي قسمة الفجوة على $(1-e^{-\\varphi})^{-1} \\approx 1.247330950103979$، لكن 1.3 قد تكون أكثر عملية.

بعض التنفيذات تستخدم ترتيب الإدراج بمجرد أن تصبح الفجوة أقل من مقدار معين.

الأنواع:

<ul>
<li>يضمن Combsort11 أن تنتهي الفجوة بالقيم (11، 8، 6، 4، 3، 2، 1)، وهو أسرع بكثير من النهايات الأخرى الممكنة.</li>
  <li>يغير Combsort مع نهايات مختلفة إلى ترتيب أكثر كفاءة عندما تكون البيانات شبه مرتبة (عندما تكون الفجوة صغيرة). ترتيب المشط مع فجوة صغيرة ليس أفضل بكثير من ترتيب الفقاعات.</li>
</ul>

رمز زائف:

<pre><b>function</b> combsort(<b>array</b> input)
  gap := input<b>.size</b> <i>// تهيئة حجم الفجوة</i>
  <b>loop until</b> gap = 1 <b>and</b> swaps = 0
    <i>// تحديث قيمة الفجوة للتمشيط التالي. المثال أدناه</i>
    gap := int(gap / 1.25)
    <b>if</b> gap &#x3C; 1 
      <i>// الحد الأدنى للفجوة هو 1</i>
      gap := 1
    <b>end if</b>
    i := 0
    swaps := 0 <i>// انظر <a href='https://rosettacode.org/wiki/Sorting_algorithms/Bubble_sort' target='_blank'>فرز الفقاعات</a> للتوضيح</i>
    <i>// تمريرة "تمشيط" واحدة على قائمة المُدخَل</i>
    <b>loop until</b> i + gap >= input<b>.size</b> <i>// انظر <a href='https://rosettacode.org/wiki/Sorting_algorithms/Shell_sort' target='_blank'>فرز شيل</a> لفكرة مشابهة</i>
      <b>if</b> input[i] > input[i+gap]
        <b>swap</b>(input[i], input[i+gap])
        swaps := 1 <i>// علامة على حدوث تبادل، إذًا</i>
            <i>// القائمة ليست مضمونة أن تكون مرتبة</i>
      <b>end if</b>
      i := i + 1
    <b>end loop</b>
  <b>end loop</b>
<b>end function</b>
</pre>

# --instructions--

اكتب دالة ترتب مصفوفة معينة باستخدام ترتيب المشط.

# --hints--

`combSort` يجب أن تكون دالة.

```js
assert(typeof combSort == 'function');
```

`combSort([25, 32, 12, 7, 20])` يجب أن تُرجع مصفوفة.

```js
assert(Array.isArray(combSort([25, 32, 12, 7, 20])));
```

`combSort([25, 32, 12, 7, 20])` يجب أن تُرجع `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(combSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`combSort([38, 45, 35, 8, 13])` يجب أن تُرجع `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(combSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`combSort([43, 36, 20, 34, 24])` يجب أن تُرجع `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(combSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`combSort([12, 33, 26, 18, 1, 16, 38])` يجب أن تُرجع `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(combSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`combSort([3, 39, 48, 16, 1, 4, 29])` يجب أن تُرجع `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(combSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function combSort(arr) {

}
```

# --solutions--

```js
function combSort(arr) {
  function is_array_sorted(arr) {
    var sorted = true;
    for (var i = 0; i < arr.length - 1; i++) {
      if (arr[i] > arr[i + 1]) {
        sorted = false;
        break;
      }
    }
    return sorted;
  }
  var iteration_count = 0;
  var gap = arr.length - 2;
  var decrease_factor = 1.25;

  // Until array is not sorted, repeat iterations
  while (!is_array_sorted(arr)) {
    // If not first gap
    if (iteration_count > 0)
      // Calculate gap
      gap = gap == 1 ? gap : Math.floor(gap / decrease_factor);

    // Set front and back elements and increment to a gap
    var front = 0;
    var back = gap;
    while (back <= arr.length - 1) {
      // If elements are not ordered swap them
      if (arr[front] > arr[back]) {
        var temp = arr[front];
        arr[front] = arr[back];
        arr[back] = temp;
      }

      // Increment and re-run swapping
      front += 1;
      back += 1;
    }
    iteration_count += 1;
  }

  return arr;
}
```
