---
id: 5a23c84252665b21eecc7ee0
title: مضروبات اليسار
challengeType: 1
forumTopicId: 302302
dashedName: left-factorials
---

# --description--

**مضروبات اليسار**، $ !n $، قد تشير إلى إما *المضروبات الجزئية* أو إلى *مجاميع المضروب*. يمكن رؤية نفس التدوين مستخدمًا بشكل مربك لتعريفين مختلفين. أحيانًا، قد تستخدم *المضروبات الجزئية* (المعروفة أيضًا باسم *الترتيبات الخاطئة*) أيًا من التدوينات التالية:

<ul>
  <li>$!n`$</li>
  <li>$!n$</li>
  <li>$n¡$</li>
</ul>

(قد لا يكون واضحًا بصريًا، لكن المثال الأخير يستخدم علامة تعجب مقلوبة رأسًا على عقب.) ستستخدم هذه المهمة الصيغة التالية لـ **مضروب اليسار**:

$ !n = \\sum\_{k=0}^{n-1} k! $

حيث $!0 = 0$

# --instructions--

اكتب دالة لحساب مضروب اليسار لعدد معين.

# --hints--

يجب أن تكون `leftFactorial` دالة.

```js
assert(typeof leftFactorial == 'function');
```

يجب أن تُرجع `leftFactorial(0)` رقمًا.

```js
assert(typeof leftFactorial(0) == 'number');
```

يجب أن تُرجع `leftFactorial(0)` القيمة `0`.

```js
assert.equal(leftFactorial(0), 0);
```

يجب أن تُرجع `leftFactorial(1)` القيمة `1`.

```js
assert.equal(leftFactorial(1), 1);
```

يجب أن تُرجع `leftFactorial(2)` القيمة `2`.

```js
assert.equal(leftFactorial(2), 2);
```

يجب أن تُرجع `leftFactorial(3)` القيمة `4`.

```js
assert.equal(leftFactorial(3), 4);
```

يجب أن تُرجع `leftFactorial(10)` القيمة `409114`.

```js
assert.equal(leftFactorial(10), 409114);
```

يجب أن تُرجع `leftFactorial(17)` القيمة `22324392524314`.

```js
assert.equal(leftFactorial(17), 22324392524314);
```

يجب أن تُرجع `leftFactorial(19)` القيمة `6780385526348314`.

```js
assert.equal(leftFactorial(19), 6780385526348314);
```

# --seed--

## --seed-contents--

```js
function leftFactorial(n) {

}
```

# --solutions--

```js
function leftFactorial(n) {
  if (n == 0) return 0;
  if (n == 1) return 1;

  // Note: for n>=20, the result may not be correct.
  // This is because JavaScript uses 53 bit integers and
  // for n>=20 result becomes too large.

  let res = 2,
    fact = 2;
  for (var i = 2; i < n; i++) {
    res += fact;
    fact *= i + 1;
  }

  return res;
}
```
