---
id: 5a23c84252665b21eecc8004
title: خوارزميات الفرز/فرز الكوكتيل
challengeType: 1
forumTopicId: 302312
dashedName: sorting-algorithmscocktail-sort
---

# --description--

<a href="https://rosettacode.org/wiki/Sorting_algorithms/Cocktail_sort" target="_blank" rel="noopener noreferrer nofollow">فرز الكوكتيل</a> هو تحسين على فرز الفقاعات. عند إعطاء مصفوفة من الأرقام، يمر فرز الكوكتيل عبر المصفوفة من البداية إلى النهاية، ناقلاً أكبر رقم إلى النهاية. ثم يعبر المصفوفة بالعكس وينقل أصغر رقم إلى البداية. يكرر هذين المرورين، ناقلاً الرقم الأكبر/الأصغر التالي إلى موقعه الصحيح في المصفوفة حتى تكتمل عملية الفرز.

# --instructions--

اكتب دالة تفرز مصفوفة معطاة باستخدام فرز الكوكتيل.

# --hints--

يجب أن تكون `cocktailSort` دالة.

```js
assert(typeof cocktailSort == 'function');
```

يجب أن تُرجع `cocktailSort([25, 32, 12, 7, 20])` مصفوفة.

```js
assert(Array.isArray(cocktailSort([25, 32, 12, 7, 20])));
```

يجب أن تُرجع `cocktailSort([25, 32, 12, 7, 20])` القيمة `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(cocktailSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

يجب أن تُرجع `cocktailSort([38, 45, 35, 8, 13])` القيمة `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(cocktailSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

يجب أن تُرجع `cocktailSort([43, 36, 20, 34, 24])` القيمة `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(cocktailSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

يجب أن تُرجع `cocktailSort([12, 33, 26, 18, 1, 16, 38])` القيمة `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(cocktailSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

يجب أن تُرجع `cocktailSort([3, 39, 48, 16, 1, 4, 29])` القيمة `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(cocktailSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function cocktailSort(arr) {

}
```

# --solutions--

```js
function cocktailSort(arr) {
  let isSorted = true;
  while (isSorted) {
    for (let i = 0; i < arr.length - 1; i++) {
      if (arr[i] > arr[i + 1]) {
        let temp = arr[i];
        arr[i] = arr[i + 1];
        arr[i + 1] = temp;
        isSorted = true;
      }
    }

    if (!isSorted) break;

    isSorted = false;

    for (let j = arr.length - 1; j > 0; j--) {
      if (arr[j - 1] > arr[j]) {
        let temp = arr[j];
        arr[j] = arr[j - 1];
        arr[j - 1] = temp;
        isSorted = true;
      }
    }
  }

  return arr;
}
```
