---
id: 5a23c84252665b21eecc7e03
title: الانحراف المعياري التراكمي
challengeType: 1
forumTopicId: 302240
dashedName: cumulative-standard-deviation
---

# --description--

افترض أن المجتمع الإحصائي الكامل المثير للاهتمام هو ثمانية طلاب في صف معين. لمجموعة محدودة من الأعداد، يُحسب الانحراف المعياري للمجتمع بأخذ الجذر التربيعي لمتوسط مربعات الانحرافات للقيم بعد طرح متوسطها. علامات صف مكون من ثمانية طلاب (أي مجتمع إحصائي) هي القيم الثمانية التالية:

$2, 4, 4, 4, 5, 5, 7, 9$

هذه النقاط الثمانية لها المتوسط (المعدل) 5:

$$\mu ={\frac {2+4+4+4+5+5+7+9}{8}}={\frac {40}{8}}=5$$

أولًا، احسب انحراف كل نقطة بيانات عن المتوسط، ثم اربع النتيجة لكل منها:

| انحرافات كل بيانات | اربع النتيجة |
| ---------------------- | --------------------|
| $(2-5)^{2}=(-3)^{2}=9$ | $(5-5)^{2}=0^{2}=0$ |
| $(4-5)^{2}=(-1)^{2}=1$ | $(5-5)^{2}=0^{2}=0$ |
| $(4-5)^{2}=(-1)^{2}=1$ | $(7-5)^{2}=2^{2}=4$ |
| $(4-5)^{2}=(-1)^{2}=1$ | $(9-5)^{2}=4^{2}=16$|

التباين هو متوسط هذه القيم:

$$\sigma ^{2}={\frac {9+1+1+1+0+0+4+16}{8}}={\frac {32}{8}}=4$$

والانحراف المعياري للمجتمع يساوي الجذر التربيعي للتباين:

$$\sigma ={\sqrt {4}}=2$$

اكتب دالة تأخذ مصفوفة من الأعداد كمعلمة وتُرجع <a href="https://rosettacode.org/wiki/Standard_deviation" target="_blank" rel="noopener noreferrer nofollow">الانحراف المعياري</a> للسلسلة.

# --hints--

`standardDeviation` يجب أن تكون دالة.

```js
assert(typeof standardDeviation == 'function');
```

`standardDeviation([2, 4, 4, 4, 5, 5, 7, 9])` يجب أن تُرجع رقمًا.

```js
assert(typeof standardDeviation([2, 4, 4, 4, 5, 5, 7, 9]) == 'number');
```

`standardDeviation([2, 4, 4, 4, 5, 5, 7, 9])` يجب أن تُرجع `2`.

```js
assert.equal(standardDeviation([2, 4, 4, 4, 5, 5, 7, 9]), 2);
```

`standardDeviation([600, 470, 170, 430, 300])` يجب أن تُرجع `147.323`.

```js
assert.equal(standardDeviation([600, 470, 170, 430, 300]), 147.323);
```

`standardDeviation([75, 83, 96, 100, 121, 125])` يجب أن تُرجع `18.239`.

```js
assert.equal(standardDeviation([75, 83, 96, 100, 121, 125]), 18.239);
```

`standardDeviation([23, 37, 45, 49, 56, 63, 63, 70, 72, 82])` يجب أن تُرجع `16.87`.

```js
assert.equal(
  standardDeviation([23, 37, 45, 49, 56, 63, 63, 70, 72, 82]),
  16.87
);
```

`standardDeviation([271, 354, 296, 301, 333, 326, 285, 298, 327, 316, 287, 314])` يجب أن تُرجع `22.631`.

```js
assert.equal(
  standardDeviation([
    271,
    354,
    296,
    301,
    333,
    326,
    285,
    298,
    327,
    316,
    287,
    314
  ]),
  22.631
);
```

# --seed--

## --seed-contents--

```js
function standardDeviation(arr) {

}
```

# --solutions--

```js
function standardDeviation(arr) {
  var sum = 0,
    sum_sq = 0,
    n = arr.length;
  arr.forEach(function(e) {
    sum += e;
    sum_sq += e * e;
  });

  var std_dev = Math.sqrt(sum_sq / n - Math.pow(sum / n, 2));
  return Math.round(std_dev * 1000) / 1000;
}
```
