---
id: 5a23c84252665b21eecc7ede
title: سنة كبيسة
challengeType: 1
forumTopicId: 302300
dashedName: leap-year
---

# --description--

حدد ما إذا كانت السنة المعطاة سنة كبيسة في التقويم الغريغوري.

# --hints--

يجب أن تكون `isLeapYear` دالة.

```js
assert(typeof isLeapYear == 'function');
```

يجب أن تُرجع `isLeapYear()` قيمة منطقية.

```js
assert(typeof isLeapYear(2018) == 'boolean');
```

يجب أن تُرجع `isLeapYear(2018)` القيمة `false`.

```js
assert.equal(isLeapYear(2018), false);
```

يجب أن تُرجع `isLeapYear(2016)` القيمة `true`.

```js
assert.equal(isLeapYear(2016), true);
```

يجب أن تُرجع `isLeapYear(2000)` القيمة `true`.

```js
assert.equal(isLeapYear(2000), true);
```

يجب أن تُرجع `isLeapYear(1900)` القيمة `false`.

```js
assert.equal(isLeapYear(1900), false);
```

يجب أن تُرجع `isLeapYear(1996)` القيمة `true`.

```js
assert.equal(isLeapYear(1996), true);
```

يجب أن تُرجع `isLeapYear(1800)` القيمة `false`.

```js
assert.equal(isLeapYear(1800), false);
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

}
```

# --solutions--

```js
function isLeapYear(year) {
  return year % 100 === 0 ? year % 400 === 0 : year % 4 === 0;
}
```
