---
id: 5a23c84252665b21eecc7e7b
title: مولد/أسّي
challengeType: 1
forumTopicId: 302275
dashedName: generatorexponential
---

# --description--

المولد هو كيان قابل للتنفيذ (مثل دالة أو إجراء) يحتوي على شفرة تُنتج متتالية من القيم، واحدة تلو الأخرى، بحيث في كل مرة تستدعي فيها المولد، تُزوّد بالقيمة التالية في المتتالية.

غالبًا ما تُبنى المولدات على قمة الـ coroutines أو الكائنات بحيث تُدار الحالة الداخلية للكائن "بطريقة طبيعية".

تُستخدم المولدات غالبًا في الحالات التي تكون فيها المتتالية محتملة اللانهاية، وحيث يمكن إنشاء القيمة التالية في المتتالية باستخدام أقل حالة ممكنة.

# --instructions--

اكتب دالة تستخدم المولدات لتوليد المربعات والمكعبات. أنشئ مولدًا جديدًا يقوم بتصفية كل المكعبات من مولد المربعات.

يجب أن تُرجع الدالة القيمة \( n^{th} \) من المولد المُصفى.

على سبيل المثال، عندما \(n=7\)، يجب أن تُرجع الدالة 81 لأن المتتالية ستكون 4، 9، 16، 25، 36، 49، 81. هنا تم تصفية 64 لأنه مكعب.

# --hints--

`exponentialGenerator` يجب أن تكون دالة.

```js
assert(typeof exponentialGenerator == 'function');
```

`exponentialGenerator()` يجب أن تُرجع رقمًا.

```js
assert(typeof exponentialGenerator(10) == 'number');
```

`exponentialGenerator(10)` يجب أن تُرجع `144`.

```js
assert.equal(exponentialGenerator(10), 144);
```

`exponentialGenerator(12)` يجب أن تُرجع `196`.

```js
assert.equal(exponentialGenerator(12), 196);
```

`exponentialGenerator(14)` يجب أن تُرجع `256`.

```js
assert.equal(exponentialGenerator(14), 256);
```

`exponentialGenerator(20)` يجب أن تُرجع `484`.

```js
assert.equal(exponentialGenerator(20), 484);
```

`exponentialGenerator(25)` يجب أن تُرجع `784`.

```js
assert.equal(exponentialGenerator(25), 784);
```

# --seed--

## --seed-contents--

```js
function exponentialGenerator(n) {

}
```

# --solutions--

```js
function exponentialGenerator(n){
  function* PowersGenerator(m) {
    var n=0;
    while(1) {
        yield Math.pow(n, m);
        n += 1;
    }
  }

  function* FilteredGenerator(g, f){
    var value = g.next().value;
    var filter = f.next().value;
    while(1) {
        if( value < filter ) {
            yield value;
            value = g.next().value;
        } else if ( value > filter ) {
            filter = f.next().value;
        } else {
            value = g.next().value;
            filter = f.next().value;
        }
    }
  }

  var squares = PowersGenerator(2);
  var cubes = PowersGenerator(3);

  var filtered = FilteredGenerator(squares, cubes);

  var curr=0;
  for(var i=0;i<n;i++) curr=filtered.next();

  return curr.value;
}
```
