---
id: 5a23c84252665b21eecc7e78
title: فيزباز العامة
challengeType: 1
forumTopicId: 302273
dashedName: general-fizzbuzz
---

# --description--

اكتب إصدارًا عامًا من <a href="https://rosettacode.org/wiki/FizzBuzz" target="_blank" rel="noopener noreferrer nofollow">FizzBuzz</a> يعمل مع أي قائمة من العوامل، مع كلماتها.

هذا في الأساس تنفيذ لـ "fizzbuzz" حيث تُزوَّد قواعد اللعبة للمستخدم. أنشئ دالة لتنفيذ ذلك. يجب أن تأخذ الدالة معلمتين.

الأولى ستكون مصفوفة تحتوي على قواعد FizzBuzz. على سبيل المثال: `[ [3, "Fizz"] , [5, "Buzz"] ]`.

هذا يعني أنه يجب طباعة `Fizz` إذا كان الرقم من مضاعفات 3 و`Buzz` إذا كان من مضاعفات 5. إذا كان من مضاعفات الاثنين، تُدمج السلاسل النصية حسب الترتيب المحدد في المصفوفة. في هذه الحالة، `FizzBuzz` إذا كان الرقم من مضاعفات 3 و5.

المعلمة الثانية هي الرقم الذي يجب أن تُرجع الدالة له سلسلة نصية كما ذُكر أعلاه.

# --hints--

يجب أن تكون `genFizzBuzz` دالة.

```js
assert(typeof genFizzBuzz == 'function');
```

يجب أن تُرجع `genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 6)` سلسلة نصية.

```js
assert(
  typeof genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    6
  ) == 'string'
);
```

يجب أن تُرجع `genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 6)` القيمة `"Fizz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    6
  ),
  'Fizz'
);
```

يجب أن تُرجع `genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 10)` القيمة `"Buzz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    10
  ),
  'Buzz'
);
```

يجب أن تُرجع `genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 12)` القيمة `"Buzz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    12
  ),
  'Buzz'
);
```

يجب أن تُرجع `genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 13)` القيمة `"13"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    13
  ),
  '13'
);
```

يجب أن تُرجع `genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 15)` القيمة `"BuzzFizz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    15
  ),
  'BuzzFizz'
);
```

يجب أن تُرجع `genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 15)` القيمة `"FizzBuzz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    15
  ),
  'FizzBuzz'
);
```

يجب أن تُرجع `genFizzBuzz([[3, "Fizz"],[5, "Buzz"],[7, "Baxx"]], 105)` القيمة `"FizzBuzzBaxx"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz'],
      [7, 'Baxx']
    ],
    105
  ),
  'FizzBuzzBaxx'
);
```

# --seed--

## --seed-contents--

```js
function genFizzBuzz(rules, num) {

}
```

# --solutions--

```js
function genFizzBuzz(rules, num) {
  let res='';
  rules.forEach(function (e) {
    if(num % e[0] == 0)
      res+=e[1];
  })

  if(res==''){
    res=num.toString();
  }

  return res;
}
```
