---
id: 5a7dad05be01840e1778a0d1
title: فركتران
challengeType: 1
forumTopicId: 302270
dashedName: fractran
---

# --description--

فركتران هي لغة برمجة غريبة مكتملة من حيث تورينغ اخترعها عالم الرياضيات جون هورتون كونواي.

برنامج فركتران هو قائمة مرتبة من الكسور الموجبة $P = (f_1, f_2, \\ldots, f_m)$، مع مُدخَل أولي موجب من نوع عدد صحيح $n$.

يُنفذ البرنامج بتحديث العدد الصحيح $n$ كما يلي:

<ul>
<li>لأول كسر، $f_i$، في القائمة حيث يكون $nf_i$ عددًا صحيحًا، استبدل $n$ بـ $nf_i$ ;</li>
  <li>كرر هذه القاعدة حتى لا ينتج أي كسر في القائمة عددًا صحيحًا عند ضربه في $n$، ثم توقف.</li>
</ul>

قدم كونواي برنامجًا للأعداد الأولية في فركتران:

$\\dfrac{17}{91}$, $\\dfrac{78}{85}$, $\\dfrac{19}{51}$, $\\dfrac{23}{38}$, $\\dfrac{29}{33}$, $\\dfrac{77}{29}$, $\\dfrac{95}{23}$, $\\dfrac{77}{19}$, $\\dfrac{1}{17}$, $\\dfrac{11}{13}$, $\\dfrac{13}{11}$, $\\dfrac{15}{14}$, $\\dfrac{15}{2}$, $\\dfrac{55}{1}$

بدءًا من $n=2$، سيغير هذا البرنامج في فركتران $n$ إلى $15=2\\times (\\frac{15}{2})$، ثم إلى $825=15\\times (\\frac{55}{1})$، مولدًا متتالية الأعداد الصحيحة التالية:

$2$, $15$, $825$, $725$, $1925$, $2275$, $425$, $390$, $330$, $290$, $770$, $\\ldots$

بعد 2، تحتوي هذه المتتالية على القوى التالية للعدد 2:

$2^2=4$, $2^3=8$, $2^5=32$, $2^7=128$, $2^{11}=2048$, $2^{13}=8192$, $2^{17}=131072$, $2^{19}=524288$, $\\ldots$

وهي قوى العدد الأولي 2.

# --instructions--

اكتب دالة تأخذ برنامج فركتران كمعلمة سلسلة نصية وتُرجع أول 10 أعداد من البرنامج كمصفوفة. إذا لم يكن الناتج يحتوي على 10 أعداد، فأرجع الأعداد كما هي.

# --hints--

`fractran` يجب أن تكون دالة.

```js
assert(typeof fractran == 'function');
```

`fractran("3/2, 1/3")` يجب أن تُرجع مصفوفة.

```js
assert(Array.isArray(fractran('3/2, 1/3')));
```

`fractran("3/2, 1/3")` يجب أن تُرجع `[ 2, 3, 1 ]`.

```js
assert.deepEqual(fractran('3/2, 1/3'), [2, 3, 1]);
```

`fractran("3/2, 5/3, 1/5")` يجب أن تُرجع `[ 2, 3, 5, 1 ]`.

```js
assert.deepEqual(fractran('3/2, 5/3, 1/5'), [2, 3, 5, 1]);
```

`fractran("3/2, 6/3")` يجب أن تُرجع `[ 2, 3, 6, 9, 18, 27, 54, 81, 162, 243 ]`.

```js
assert.deepEqual(fractran('3/2, 6/3'), [2, 3, 6, 9, 18, 27, 54, 81, 162, 243]);
```

`fractran("2/7, 7/2")` يجب أن تُرجع `[ 2, 7, 2, 7, 2, 7, 2, 7, 2, 7 ]`.

```js
assert.deepEqual(fractran('2/7, 7/2'), [2, 7, 2, 7, 2, 7, 2, 7, 2, 7]);
```

`fractran("17/91, 78/85, 19/51, 23/38, 29/33, 77/29, 95/23, 77/19, 1/17, 11/13, 13/11, 15/14, 15/2, 55/1")` يجب أن تُرجع `[ 2, 15, 825, 725, 1925, 2275, 425, 390, 330, 290 ]`.

```js
assert.deepEqual(
  fractran(
    '17/91, 78/85, 19/51, 23/38, 29/33, 77/29, 95/23, 77/19, 1/17, 11/13, 13/11, 15/14, 15/2, 55/1'
  ),
  [2, 15, 825, 725, 1925, 2275, 425, 390, 330, 290]
);
```

# --seed--

## --seed-contents--

```js
function fractran(progStr) {

}
```

# --solutions--

```js
function fractran(progStr){
  var num = new Array();
  var den = new Array();
  var val ;
  var out="";
  function compile(prog){
    var regex = /\s*(\d*)\s*\/\s*(\d*)\s*(.*)/m;
    while(regex.test(prog)){
      num.push(regex.exec(prog)[1]);
      den.push(regex.exec(prog)[2]);
      prog = regex.exec(prog)[3];
    }
  }

  function step(val){
    var i=0;
    while(i<den.length && val%den[i] != 0) i++;
    return num[i]*val/den[i];
  }

  var seq=[]

  function exec(val){
    var i = 0;
    while(val && i<limit){
      seq.push(val)
      val = step(val);
      i ++;
    }
  }

  // Main
  compile(progStr);
  var limit = 10;
  exec(2);
  return seq;
}
```
