---
id: 5e4ce2f5ac708cc68c1df261
title: مولد التوافق الخطي
challengeType: 1
forumTopicId: 385266
dashedName: linear-congruential-generator
---

# --description--

مولد التوافق الخطي (LCG) هو <em>خوارزمية</em> تُنتج متتالية من الأعداد شبه العشوائية محسوبة بمعادلة خطية قطعة قطعة غير متصلة. تستخدم جميع مولدات التوافق الخطي هذه الصيغة:

$$r_{n + 1} = (a \times r_n + c) \bmod m$$

حيث:

<ul>
<li>$ r_0 $ هو البذرة.</li>
<li>$r_1$، $r_2$، $r_3$، ...، هي الأعداد العشوائية.</li>
<li>$a$، $c$، $m$ هي ثوابت.</li>
</ul>

إذا اختيرت قيم $a$ و $c$ و $m$ بعناية، فإن المولد ينتج توزيعًا منتظمًا من الأعداد الصحيحة من $0$ إلى $m - 1$.

أعداد <abbr title="linear congruential generator">LCG</abbr> ذات جودة ضعيفة. $r_n$ و $r\_{n + 1}$ ليست مستقلة كما هي الأعداد العشوائية الحقيقية. أي شخص يعرف $r_n$ يمكنه التنبؤ بـ $r\_{n + 1}$، لذلك <abbr title="linear congruential generator">LCG</abbr> ليست آمنة تشفيرياً. مع ذلك، <abbr title="linear congruential generator">LCG</abbr> جيدة بما يكفي للمهام البسيطة مثل اختبار أولية Miller-Rabin أو توزيع أوراق FreeCell. من فوائد <abbr title="linear congruential generator">LCG</abbr> أنه يمكن بسهولة إعادة إنتاج متتالية الأعداد نفسها من نفس $r_0$. كما يمكن إعادة إنتاج هذه المتتالية باستخدام لغة برمجة مختلفة، لأن الصيغة بسيطة جدًا.

# --instructions--

اكتب دالة تأخذ $r_0,a,c,m,n$ كمعلمات وتُرجع $r_n$.

# --hints--

`linearCongGenerator` يجب أن تكون دالة.

```js
assert(typeof linearCongGenerator == 'function');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` يجب أن تُرجع عددًا.

```js
assert(typeof linearCongGenerator(324, 1145, 177, 2148, 3) == 'number');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` يجب أن تُرجع `855`.

```js
assert.equal(linearCongGenerator(324, 1145, 177, 2148, 3), 855);
```

`linearCongGenerator(234, 11245, 145, 83648, 4)` يجب أن تُرجع `1110`.

```js
assert.equal(linearCongGenerator(234, 11245, 145, 83648, 4), 1110);
```

`linearCongGenerator(85, 11, 1234, 214748, 5)` يجب أن تُرجع `62217`.

```js
assert.equal(linearCongGenerator(85, 11, 1234, 214748, 5), 62217);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 1)` يجب أن تُرجع `12345`.

```js
assert.equal(linearCongGenerator(0, 1103515245, 12345, 2147483648, 1), 12345);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 2)` يجب أن تُرجع `1406932606`.

```js
assert.equal(
  linearCongGenerator(0, 1103515245, 12345, 2147483648, 2),
  1406932606
);
```

# --seed--

## --seed-contents--

```js
function linearCongGenerator(r0, a, c, m, n) {

}
```

# --solutions--

```js
function linearCongGenerator(r0, a, c, m, n) {
    for (let i = 0; i < n; i++) {
        r0 = (a * r0 + c) % m;
    }
    return r0;
}
```
