---
id: 5e4ce2eaac708cc68c1df260
title: مسافة ليفنشتاين
challengeType: 1
forumTopicId: 385264
dashedName: levenshtein-distance
---

# --description--

في نظرية المعلومات وعلوم الحاسوب، تُعد **مسافة ليفنشتاين** مقياسًا لقياس مقدار الاختلاف بين متتاليتين (أي مسافة تعديل). تُعرف مسافة ليفنشتاين بين سلسلتين نصيتين بأنها الحد الأدنى لعدد التعديلات اللازمة لتحويل سلسلة إلى أخرى، مع السماح بعمليات التعديل التالية: الإدراج، الحذف، أو الاستبدال بحرف واحد.

مثال:

مسافة ليفنشتاين بين "**kitten**" و"**sitting**" هي 3، لأن التعديلات الثلاثة التالية تحول واحدة إلى الأخرى، ولا توجد طريقة للقيام بذلك بأقل من ثلاث تعديلات:

<ul>
<li><strong>k</strong>itten   <strong>s</strong>itten    (استبدال 'k' بـ 's')</li>
  <li>sitt<strong>e</strong>n   sitt<strong>i</strong>n    (استبدال 'e' بـ 'i')</li>
  <li>sittin   sittin<strong>g</strong>    (إدراج 'g' في النهاية).</li>
</ul>

*مسافة ليفنشتاين بين "**rosettacode**" و"**raisethysword**" هي **8**.*

*المسافة بين سلسلتين هي نفسها عندما تُعكس كلتا السلسلتين.*

# --instructions--

اكتب دالة تُرجع مسافة ليفنشتاين بين سلسلتين نصيتين معطاة كمعلمات.

# --hints--

`levenshtein` يجب أن تكون دالة.

```js
assert(typeof levenshtein == 'function');
```

`levenshtein("mist", "dist")` يجب أن تُرجع رقمًا.

```js
assert(typeof levenshtein('mist', 'dist') == 'number');
```

`levenshtein("mist", "dist")` يجب أن تُرجع `1`.

```js
assert.equal(levenshtein('mist', 'dist'), 1);
```

`levenshtein("tier", "tor")` يجب أن تُرجع `2`.

```js
assert.equal(levenshtein('tier', 'tor'), 2);
```

`levenshtein("kitten", "sitting")` يجب أن تُرجع `3`.

```js
assert.equal(levenshtein('kitten', 'sitting'), 3);
```

`levenshtein("stop", "tops")` يجب أن تُرجع `2`.

```js
assert.equal(levenshtein('stop', 'tops'), 2);
```

`levenshtein("rosettacode", "raisethysword")` يجب أن تُرجع `8`.

```js
assert.equal(levenshtein('rosettacode', 'raisethysword'), 8);
```

`levenshtein("mississippi", "swiss miss")` يجب أن تُرجع `8`.

```js
assert.equal(levenshtein('mississippi', 'swiss miss'), 8);
```

# --seed--

## --seed-contents--

```js
function levenshtein(a, b) {

}
```

# --solutions--

```js
function levenshtein(a, b) {
  var t = [], u, i, j, m = a.length, n = b.length;
  if (!m) { return n; }
  if (!n) { return m; }
  for (j = 0; j <= n; j++) { t[j] = j; }
  for (i = 1; i <= m; i++) {
    for (u = [i], j = 1; j <= n; j++) {
      u[j] = a[i - 1] === b[j - 1] ? t[j - 1] : Math.min(t[j - 1], t[j], u[j - 1]) + 1;
    } t = u;
  } return u[n];
}
```
