---
id: 5e6dd1278e6ca105cde40ea9
title: أطول تسلسل فرعي مشترك
challengeType: 1
forumTopicId: 385271
dashedName: longest-common-subsequence
---

# --description--

**أطول تسلسل فرعي مشترك** (أو **LCS**) للمجموعتين A و B هو أطول مجموعة من العناصر المشتركة بين المجموعتين وبنفس الترتيب في كل مجموعة. على سبيل المثال، المتتاليات `1234` و `1224533324` لها LCS هو `1234`:

<u>1234</u>

<u>12</u>245<u>3</u>332<u>4</u>

كمثال على السلاسل النصية، اعتبر المتتاليات `thisisatest` و `testing123testing`. سيكون LCS هو `tsitest`:

<u>t</u>hi<u>si</u>sa<u>test</u>

<u>t</u>e<u>s</u>t<u>i</u>ng123<u>test</u>ing.

شفرتك تحتاج فقط إلى التعامل مع السلاسل النصية.

# --instructions--

اكتب دالة حساسة لحالة الأحرف تُرجع LCS لسلسلتين نصيتين. لا تحتاج إلى إظهار عدة LCS.

# --hints--

`lcs` يجب أن تكون دالة.

```js
assert(typeof lcs == 'function');
```

`lcs("thisisatest", "testing123testing")` يجب أن تُرجع سلسلة نصية.

```js
assert(typeof lcs('thisisatest', 'testing123testing') == 'string');
```

`lcs("thisisatest", "testing123testing")` يجب أن تُرجع `"tsitest"`.

```js
assert.equal(lcs('thisisatest', 'testing123testing'), 'tsitest');
```

`lcs("ABCDGH", "AEDFHR")` يجب أن تُرجع `"ADH"`.

```js
assert.equal(lcs('ABCDGH', 'AEDFHR'), 'ADH');
```

`lcs("AGGTAB", "GXTXAYB")` يجب أن تُرجع `"GTAB"`.

```js
assert.equal(lcs('AGGTAB', 'GXTXAYB'), 'GTAB');
```

`lcs("BDACDB", "BDCB")` يجب أن تُرجع `"BDCB"`.

```js
assert.equal(lcs('BDACDB', 'BDCB'), 'BDCB');
```

`lcs("ABAZDC", "BACBAD")` يجب أن تُرجع `"ABAD"`.

```js
assert.equal(lcs('ABAZDC', 'BACBAD'), 'ABAD');
```

# --seed--

## --seed-contents--

```js
function lcs(a, b) {

}
```

# --solutions--

```js
function lcs(a, b) {
  var aSub = a.substring(0, a.length - 1);
  var bSub = b.substring(0, b.length - 1);

  if (a.length === 0 || b.length === 0) {
    return '';
  } else if (a.charAt(a.length - 1) === b.charAt(b.length - 1)) {
    return lcs(aSub, bSub) + a.charAt(a.length - 1);
  } else {
    var x = lcs(a, bSub);
    var y = lcs(aSub, b);
    return (x.length > y.length) ? x : y;
  }
}
```
