---
id: 5e6dd14797f5ce267c2f19d0
title: متتالية النظر والقول
challengeType: 1
forumTopicId: 385277
dashedName: look-and-say-sequence
---

# --description--

متتالية النظر والقول هي متتالية أعداد معرفة بشكل تكراري.

تعريف المتتالية

<ul><li>اختر عددًا عشريًا</li>
<li><span>انظر</span> إلى العدد، مع تجميع بصري للتتابعات المتتالية من نفس الرقم.</li>
<li><span>قل</span> العدد، من اليسار إلى اليمين، مجموعة بمجموعة؛ بعدد تلك الأرقام المتتالية - متبوعًا بالرقم نفسه.</li></ul><span> هذا يصبح العدد التالي في المتتالية.</span>

مثال:

<ul><li>ابدأ بالعدد 1، لديك <span>واحد</span> 1 مما ينتج 11</li>
<li>ابدأ بالعدد 11، لديك <span>اثنان</span> 1. أي: 21</li>
<li>ابدأ بالعدد 21، لديك <span>واحد</span> 2، ثم <span>واحد</span> 1. أي: (12)(11) التي تصبح 1211</li>
<li>ابدأ بالعدد 1211، لديك <span>واحد</span> 1، <span>واحد</span> 2، ثم <span>اثنان</span> 1. أي: (11)(12)(21) التي تصبح 111221</li></ul>

# --instructions--

اكتب دالة تقبل سلسلة نصية كمعلمة، تعالجها، وتُرجع السلسلة الناتجة.

# --hints--

`lookAndSay` يجب أن تكون دالة.

```js
assert(typeof lookAndSay == 'function');
```

`lookAndSay("1")` يجب أن تُرجع سلسلة نصية.

```js
assert(typeof lookAndSay('1') == 'string');
```

`lookAndSay("1")` يجب أن تُرجع `"11"`.

```js
assert.equal(lookAndSay('1'), '11');
```

`lookAndSay("11")` يجب أن تُرجع `"21"`.

```js
assert.equal(lookAndSay('11'), '21');
```

`lookAndSay("21")` يجب أن تُرجع `"1211"`.

```js
assert.equal(lookAndSay('21'), '1211');
```

`lookAndSay("1211")` يجب أن تُرجع `"111221"`.

```js
assert.equal(lookAndSay('1211'), '111221');
```

`lookAndSay("3542")` يجب أن تُرجع `"13151412"`.

```js
assert.equal(lookAndSay('3542'), '13151412');
```

# --seed--

## --seed-contents--

```js
function lookAndSay(str) {

}
```

# --solutions--

```js
function lookAndSay(str) {
    return str.replace(/(.)\1*/g, function(seq, p1) {
      return seq.length.toString() + p1;
    });
}
```
