---
id: 5e6decd8ec8d7db960950d1c
title: تحليل LU
challengeType: 1
forumTopicId: 385280
dashedName: lu-decomposition
---

# --description--

يمكن تحليل كل مصفوفة مربعة $A$ إلى حاصل ضرب مصفوفة مثلثية دنيا $L$ ومصفوفة مثلثية عُليا $U$. يُعرف هذا بتحليل LU.

$A = LU$

هو شكل معدل من الحذف الغاوسي.

بينما يعمل تحليل تشولسكي فقط للمصفوفات المتماثلة والموجبة التعريف، يعمل تحليل LU العام لأي مصفوفة مربعة.

هناك عدة خوارزميات لحساب $L$ و $U$.

لاشتقاق *خوارزمية كروت* لمثال 3x3، علينا حل النظام التالي:

\\begin{align}A = \\begin{pmatrix} a\_{11} & a\_{12} & a\_{13}\\\\ a\_{21} & a\_{22} & a\_{23}\\\\ a\_{31} & a\_{32} & a\_{33}\\\\ \\end{pmatrix}= \\begin{pmatrix} l\_{11} & 0 & 0 \\\\ l\_{21} & l\_{22} & 0 \\\\ l\_{31} & l\_{32} & l\_{33}\\\\ \\end{pmatrix} \\begin{pmatrix} u\_{11} & u\_{12} & u\_{13} \\\\ 0 & u\_{22} & u\_{23} \\\\ 0 & 0 & u\_{33} \\end{pmatrix} = LU\\end{align}  

الآن علينا حل 9 معادلات مع 12 مجهولًا. لجعل النظام قابلًا للحل بشكل فريد، عادةً ما تُضبط العناصر القطرية في $L$ على 1

$l\_{11}=1$

$l\_{22}=1$

$l\_{33}=1$

بحيث نحصل على نظام قابل للحل مكوّن من 9 مجهولات و9 معادلات.

\\begin{align}A = \\begin{pmatrix} a\_{11} & a\_{12} & a\_{13}\\\\ a\_{21} & a\_{22} & a\_{23}\\\\ a\_{31} & a\_{32} & a\_{33}\\\\ \\end{pmatrix} = \\begin{pmatrix} 1 & 0 & 0 \\\\ l\_{21} & 1 & 0 \\\\ l\_{31} & l\_{32} & 1\\\\ \\end{pmatrix} \\begin{pmatrix} u\_{11} & u\_{12} & u\_{13} \\\\ 0 & u\_{22} & u\_{23} \\\\ 0 & 0 & u\_{33} \\end{pmatrix} = \\begin{pmatrix} u\_{11} & u\_{12} & u\_{13} \\\\ u\_{11}l\_{21} & u\_{12}l\_{21}+u\_{22} & u\_{13}l\_{21}+u\_{23} \\\\ u\_{11}l\_{31} & u\_{12}l\_{31}+u\_{22}l\_{32} & u\_{13}l\_{31} + u\_{23}l\_{32}+u\_{33} \\end{pmatrix} = LU\\end{align}

بحل باقي $l$ و $u$، نحصل على المعادلات التالية:

$u\_{11}=a\_{11}$

$u\_{12}=a\_{12}$

$u\_{13}=a\_{13}$

$u\_{22}=a\_{22} - u\_{12}l\_{21}$

$u\_{23}=a\_{23} - u\_{13}l\_{21}$

$u\_{33}=a\_{33} - (u\_{13}l\_{31} + u\_{23}l\_{32})$

ولـ $l$:

$l\_{21}=\\frac{1}{u\_{11}} a\_{21}$

$l\_{31}=\\frac{1}{u\_{11}} a\_{31}$

$l\_{32}=\\frac{1}{u\_{22}} (a\_{32} - u\_{12}l\_{31})$

نلاحظ وجود نمط حسابي يمكن التعبير عنه بالمعادلات التالية، أولًا لـ $U$

$u\_{ij} = a\_{ij} - \\sum\_{k=1}^{i-1} u\_{kj}l\_{ik}$

ثم لـ $L$

$l\_{ij} = \\frac{1}{u\_{jj}} (a\_{ij} - \\sum\_{k=1}^{j-1} u\_{kj}l\_{ik})$

نرى في المعادلة الثانية أنه للحصول على $l\_{ij}$ تحت القطر، يجب القسمة على العنصر القطري (المحور) $u\_{jj}$، مما يسبب مشاكل عندما يكون $u\_{jj}$ صفرًا أو صغيرًا جدًا، وهذا يؤدي إلى عدم استقرار عددي.

الحل لهذه المشكلة هو *تدوير المحور* $A$، أي إعادة ترتيب صفوف $A$ قبل تحليل $LU$ بحيث يكون أكبر عنصر في كل عمود على القطر. إعادة ترتيب الصفوف تعني ضرب $A$ بمصفوفة تبديل $P$:

$PA \\Rightarrow A'$

مثال:

\\begin{align} \\begin{pmatrix} 0 & 1 \\\\ 1 & 0 \\end{pmatrix} \\begin{pmatrix} 1 & 4 \\\\ 2 & 3 \\end{pmatrix} \\Rightarrow \\begin{pmatrix} 2 & 3 \\\\ 1 & 4 \\end{pmatrix} \\end{align}

ثم يُطبق خوارزمية التحليل على المصفوفة المعاد ترتيبها بحيث

$PA = LU$

# --instructions--

المهمة هي تنفيذ روتين يأخذ مصفوفة مربعة $n \times n$ $A$ ويُرجع مصفوفة مثلثية دنيا $L$، ومصفوفة مثلثية عُليا $U$، ومصفوفة تبديل $P$، بحيث تتحقق المعادلة أعلاه. يجب أن تكون القيمة المرجعة في شكل `[L, U, P]`.

# --before-each--

```js
const matrix1 = [
  [1, 3, 5],
  [2, 4, 7],
  [1, 1, 0]
];
const expected1 = [
  [
    [1, 0, 0],
    [0.5, 1, 0],
    [0.5, -1, 1]
  ],
  [
    [2, 4, 7],
    [0, 1, 1.5],
    [0, 0, -2]
  ],
  [
    [0, 1, 0],
    [1, 0, 0],
    [0, 0, 1]
  ]
];

const matrix2 = [
  [11, 9, 24, 2],
  [1, 5, 2, 6],
  [3, 17, 18, 1],
  [2, 5, 7, 1]
];
const expected2 = [
  [
    [1, 0, 0, 0],
    [0.2727272727272727, 1, 0, 0],
    [0.09090909090909091, 0.2875, 1, 0],
    [0.18181818181818182, 0.23124999999999996, 0.0035971223021580693, 1]
  ],
  [
    [11, 9, 24, 2],
    [0, 14.545454545454547, 11.454545454545455, 0.4545454545454546],
    [0, 0, -3.4749999999999996, 5.6875],
    [0, 0, 0, 0.510791366906476]
  ],
  [
    [1, 0, 0, 0],
    [0, 0, 1, 0],
    [0, 1, 0, 0],
    [0, 0, 0, 1]
  ]
];

const matrix3 = [
  [1, 1, 1],
  [4, 3, -1],
  [3, 5, 3]
];
const expected3 = [
  [
    [1, 0, 0],
    [0.75, 1, 0],
    [0.25, 0.09090909090909091, 1]
  ],
  [
    [4, 3, -1],
    [0, 2.75, 3.75],
    [0, 0, 0.9090909090909091]
  ],
  [
    [0, 1, 0],
    [0, 0, 1],
    [1, 0, 0]
  ]
];

const matrix4 = [
  [1, -2, 3],
  [2, -5, 12],
  [0, 2, -10]
];
const expected4 = [
  [
    [1, 0, 0],
    [0, 1, 0],
    [0.5, 0.25, 1]
  ],
  [
    [2, -5, 12],
    [0, 2, -10],
    [0, 0, -0.5]
  ],
  [
    [0, 1, 0],
    [0, 0, 1],
    [1, 0, 0]
  ]
];
```

# --hints--

`luDecomposition` يجب أن تكون دالة.

```js
assert(typeof luDecomposition == 'function');
```

`luDecomposition([[1, 3, 5], [2, 4, 7], [1, 1, 0]])` يجب أن تُرجع مصفوفة.

```js
assert(Array.isArray(luDecomposition(matrix1)));
```

`luDecomposition([[1, 3, 5], [2, 4, 7], [1, 1, 0]])` يجب أن تُرجع `[[[1, 0, 0], [0.5, 1, 0], [0.5, -1, 1]], [[2, 4, 7], [0, 1, 1.5], [0, 0, -2]], [[0, 1, 0], [1, 0, 0], [0, 0, 1]]]`.

```js
assert.deepEqual(luDecomposition(matrix1), expected1);
```

`luDecomposition([[11, 9, 24, 2], [1, 5, 2, 6], [3, 17, 18, 1], [2, 5, 7, 1]])` يجب أن تُرجع `[[[1, 0, 0, 0], [0.2727272727272727, 1, 0, 0], [0.09090909090909091, 0.2875, 1, 0], [0.18181818181818182, 0.23124999999999996, 0.0035971223021580693, 1]], [[11, 9, 24, 2], [0, 14.545454545454547, 11.454545454545455, 0.4545454545454546], [0, 0, -3.4749999999999996, 5.6875], [0, 0, 0, 0.510791366906476]], [[1, 0, 0, 0], [0, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 1]]]`.

```js
assert.deepEqual(luDecomposition(matrix2), expected2);
```

`luDecomposition([[1, 1, 1], [4, 3, -1], [3, 5, 3]])` يجب أن تُرجع `[[[1, 0, 0], [0.75, 1, 0], [0.25, 0.09090909090909091, 1]], [[4, 3, -1], [0, 2.75, 3.75], [0, 0, 0.9090909090909091]], [[0, 1, 0], [0, 0, 1], [1, 0, 0]]]`.

```js
assert.deepEqual(luDecomposition(matrix3), expected3);
```

`luDecomposition([[1, -2, 3], [2, -5, 12], [0, 2, -10]])` يجب أن تُرجع `[[[1, 0, 0], [0, 1, 0], [0.5, 0.25, 1]], [[2, -5, 12], [0, 2, -10], [0, 0, -0.5]], [[0, 1, 0], [0, 0, 1], [1, 0, 0]]]`.

```js
assert.deepEqual(luDecomposition(matrix4), expected4);
```

# --seed--

## --seed-contents--

```js
function luDecomposition(A) {

}
```

# --solutions--

```js
function luDecomposition(A) {

    function dotProduct(a, b) {
        var sum = 0;
        for (var i = 0; i < a.length; i++)
            sum += a[i] * b[i]
        return sum;
    }

    function matrixMul(A, B) {
        var result = new Array(A.length);
        for (var i = 0; i < A.length; i++)
            result[i] = new Array(B[0].length)
        var aux = new Array(B.length);

        for (var j = 0; j < B[0].length; j++) {

            for (var k = 0; k < B.length; k++)
                aux[k] = B[k][j];

            for (var i = 0; i < A.length; i++)
                result[i][j] = dotProduct(A[i], aux);
        }
        return result;
    }

    function pivotize(m) {
        var n = m.length;
        var id = new Array(n);
        for (var i = 0; i < n; i++) {
            id[i] = new Array(n);
            id[i].fill(0)
            id[i][i] = 1;
        }

        for (var i = 0; i < n; i++) {
            var maxm = m[i][i];
            var row = i;
            for (var j = i; j < n; j++)
                if (m[j][i] > maxm) {
                    maxm = m[j][i];
                    row = j;
                }

            if (i != row) {
                var tmp = id[i];
                id[i] = id[row];
                id[row] = tmp;
            }
        }
        return id;
    }

    var n = A.length;
    var L = new Array(n);
    for (var i = 0; i < n; i++) { L[i] = new Array(n); L[i].fill(0) }
    var U = new Array(n);
    for (var i = 0; i < n; i++) { U[i] = new Array(n); U[i].fill(0) }
    var P = pivotize(A);
    var A2 = matrixMul(P, A);

    for (var j = 0; j < n; j++) {
        L[j][j] = 1;
        for (var i = 0; i < j + 1; i++) {
            var s1 = 0;
            for (var k = 0; k < i; k++)
                s1 += U[k][j] * L[i][k];
            U[i][j] = A2[i][j] - s1;
        }
        for (var i = j; i < n; i++) {
            var s2 = 0;
            for (var k = 0; k < j; k++)
                s2 += U[k][j] * L[i][k];
            L[i][j] = (A2[i][j] - s2) / U[j][j];
        }
    }
    return [L, U, P];
}
```
