---
id: 5e6dd15004c88cf00d2a78b3
title: التكرار على عدة مصفوفات في نفس الوقت
challengeType: 1
forumTopicId: 385279
dashedName: loop-over-multiple-arrays-simultaneously
---

# --description--

كرر التكرار على عدة مصفوفات وأنشئ مصفوفة جديدة يكون العنصر $i^{th}$ فيها هو دمج العنصر $i^{th}$ من كل مصفوفة من المصفوفات المعطاة.

لهذا المثال، إذا أعطيت هذه المصفوفة من المصفوفات:

```js
[ ["a", "b", "c"], ["A", "B", "C"], [1, 2, 3] ]
```

يجب أن يكون الناتج:

```js
["aA1","bB2","cC3"]
```

# --instructions--

اكتب دالة تأخذ مصفوفة من المصفوفات كمعلمة وتُرجع مصفوفة من سلاسل نصية تحقق الوصف المعطى.

# --hints--

يجب أن تكون `loopSimult` دالة.

```js
assert(typeof loopSimult == 'function');
```

يجب أن تُرجع `loopSimult([["a", "b", "c"], ["A", "B", "C"], [1, 2, 3]])` مصفوفة.

```js
assert(
  Array.isArray(
    loopSimult([
      ['a', 'b', 'c'],
      ['A', 'B', 'C'],
      [1, 2, 3]
    ])
  )
);
```

يجب أن تُرجع `loopSimult([["a", "b", "c"], ["A", "B", "C"], [1, 2, 3]])` القيمة `["aA1", "bB2", "cC3"]`.

```js
assert.deepEqual(
  loopSimult([
    ['a', 'b', 'c'],
    ['A', 'B', 'C'],
    [1, 2, 3]
  ]),
  ['aA1', 'bB2', 'cC3']
);
```

يجب أن تُرجع `loopSimult([["c", "b", "c"], ["4", "5", "C"], [7, 7, 3]])` القيمة `["c47", "b57", "cC3"]`.

```js
assert.deepEqual(
  loopSimult([
    ['c', 'b', 'c'],
    ['4', '5', 'C'],
    [7, 7, 3]
  ]),
  ['c47', 'b57', 'cC3']
);
```

يجب أن تُرجع `loopSimult([["a", "b", "c", "d"], ["A", "B", "C", "d"], [1, 2, 3, 4]])` القيمة `["aA1", "bB2", "cC3", "dd4"]`.

```js
assert.deepEqual(
  loopSimult([
    ['a', 'b', 'c', 'd'],
    ['A', 'B', 'C', 'd'],
    [1, 2, 3, 4]
  ]),
  ['aA1', 'bB2', 'cC3', 'dd4']
);
```

يجب أن تُرجع `loopSimult([["a", "b"], ["A", "B"], [1, 2]])` القيمة `["aA1", "bB2"]`.

```js
assert.deepEqual(
  loopSimult([
    ['a', 'b'],
    ['A', 'B'],
    [1, 2]
  ]),
  ['aA1', 'bB2']
);
```

يجب أن تُرجع `loopSimult([["b", "c"], ["B", "C"], [2, 3]])` القيمة `["bB2", "cC3"]`.

```js
assert.deepEqual(
  loopSimult([
    ['b', 'c'],
    ['B', 'C'],
    [2, 3]
  ]),
  ['bB2', 'cC3']
);
```

# --seed--

## --seed-contents--

```js
function loopSimult(A) {

}
```

# --solutions--

```js
function loopSimult(A) {
    var res = [], output;
    for (var i = 0; i < A[0].length; i += 1) {
        output = "";
        for (var j = 0; j < A.length; j++)
            output += A[j][i];
        res.push(output);
    }
    return res;
}
```
