---
id: 5e94a54cc7b022105bf0fd2c
title: تكرار الكلمات
challengeType: 1
forumTopicId: 393913
dashedName: word-frequency
---

# --description--

بالنظر إلى سلسلة نصية وعدد صحيح n، أعد أكثر n كلمة شيوعًا في الملف (مع عدد مرات ظهورها) بترتيب تنازلي للتكرار.

# --instructions--

اكتب دالة لحساب عدد مرات ظهور كل كلمة وأعد أكثر n كلمة شيوعًا مع عدد مرات ظهورها بترتيب تنازلي.

يجب أن تُعيد الدالة مصفوفة ثنائية الأبعاد بحيث يكون كل عنصر منها بالشكل التالي: `[word, freq]`. يجب أن تكون `word` هي النسخة ذات الأحرف الصغيرة من الكلمة و`freq` هو العدد الذي يدل على التكرار.

يجب أن تُعيد الدالة مصفوفة فارغة إذا لم تُزوّد بسلسلة نصية.

يجب أن تكون الدالة غير حساسة لحالة الأحرف، على سبيل المثال، يجب معاملة السلسلتين "Hello" و"hello" على أنهما نفس الكلمة.

يمكنك اعتبار الكلمات التي تحتوي على أحرف خاصة مثل الشرطة السفلية، الشرطات، الفواصل العليا، الفواصل، وغيرها، كلمات مميزة منفصلة.

على سبيل المثال، إذا كانت السلسلة "Hello hello goodbye"، يجب أن تُعيد دالتك `[['hello', 2], ['goodbye', 1]]`.

# --hints--

`wordFrequency` يجب أن تكون دالة.

```js
assert.isFunction(wordFrequency);
```

`wordFrequency` يجب أن تُعيد مصفوفة.

```js
assert.isArray(wordFrequency('test'));
```

`wordFrequency("Hello hello world", 2)` يجب أن تُعيد `[['hello', 2], ['world', 1]]`

```js
const example_1 = 'Hello hello world';
const example_1_solution = [['hello', 2], ['world', 1]];

assert.deepEqual(wordFrequency(example_1, 2), example_1_solution);
```

`wordFrequency("The quick brown fox jumped over the lazy dog", 1)` يجب أن تُعيد `[['the', 2]]`

```js
const example_2 = 'The quick brown fox jumped over the lazy dog';
const example_2_solution = [['the', 2]];

assert.deepEqual(wordFrequency(example_2, 1), example_2_solution);
```

`wordFrequency("Opensource opensource open-source open source", 1)` يجب أن تُعيد `[['opensource', 2]]`

```js
const example_3 = 'Opensource opensource open-source open source';
const example_3_solution = [['opensource', 2]];

assert.deepEqual(wordFrequency(example_3, 1), example_3_solution);
```

`wordFrequency("Apple App apply aPP aPPlE", 3)` يجب أن تُعيد `[['app', 2], ['apple', 2], ['apply', 1]]` أو `[['apple', 2], ['app', 2], ['apply', 1]]`

```js
const example_4 = 'Apple App apply aPP aPPlE';
const example_4_solution = [["app", 2],["apple", 2],["apply", 1]];

const arr = wordFrequency(example_4, 3);
assert.includeDeepMembers(arr, example_4_solution);
```

`wordFrequency("c d a d c a b d d c", 4)` يجب أن تُعيد `[['d', 4], ['c', 3], ['a', 2], ['b', 1]]`

```js
const example_5 = 'c d a d c a b d d c';
const example_5_solution = [['d', 4], ['c', 3], ['a', 2], ['b', 1]];

assert.deepEqual(wordFrequency(example_5, 4), example_5_solution);
```

`wordFrequency("", 5)` يجب أن تُعيد `[]`

```js
const example_6 = '';
const example_6_solution = [];

assert.deepEqual(wordFrequency(example_6, 5), example_6_solution);
```

# --seed--

## --seed-contents--

```js
function wordFrequency(txt, n) {

}
```

# --solutions--

```js
function wordFrequency(txt, n) {
  var words = txt.split(/\s+/);
  var wordCount = {};
  words.forEach(word => {
    if (word == '') {
      return;
    }
    const lowerWord = word.toLowerCase();
    wordCount[lowerWord] =
      lowerWord in wordCount ? wordCount[lowerWord] + 1 : 1;
  });

  var wordsArray = [];
  for (let [word, count] of Object.entries(wordCount)) {
    wordsArray.push([word, count]);
  }

  wordsArray.sort((a, b) => {
    if (a[1] !== b[1]) {
      return b[1] - a[1];
    } else if (a[0] !== b[0]) {
      return a[0] < b[0] ? -1 : 1;
    }
    return 0;
  });
  return wordsArray.slice(0, n);
}
```
