---
id: 5e9ddb06ec35240f39657419
title: فيزباز
challengeType: 1
forumTopicId: 385370
dashedName: fizzbuzz
---

# --description--

اكتب برنامجًا يُولد مصفوفة من الأعداد الصحيحة من 1 إلى 100 (شاملة). لكن:

<ul>
<li>لأعداد القابلة للقسمة على 3، أضف <code>"Fizz"</code> إلى المصفوفة بدلًا من الرقم</li>
    <li>لأعداد القابلة للقسمة على 5، أضف <code>"Buzz"</code> إلى المصفوفة بدلًا من الرقم</li>
    <li>لأعداد القابلة للقسمة على 3 و5 معًا، أضف <code>"FizzBuzz"</code> إلى المصفوفة بدلًا من الرقم</li>
</ul>

# --instructions--

يجب أن يُرجع برنامجك مصفوفة تحتوي النتائج بناءً على القواعد السابقة.

# --hints--

يجب أن تكون `fizzBuzz` دالة.

```js
assert(typeof fizzBuzz == 'function');
```

يجب أن تُرجع `fizzBuzz()` مصفوفة.

```js
assert(Array.isArray(fizzBuzz()) == true);
```

الأعداد القابلة للقسمة على 3 فقط يجب أن تُرجع `"Fizz"`.

```js
assert.equal(fizzBuzz()[2], 'Fizz');
```

الأعداد القابلة للقسمة على 5 فقط يجب أن تُرجع `"Buzz"`.

```js
assert.equal(fizzBuzz()[99], 'Buzz');
```

الأعداد القابلة للقسمة على كل من 3 و5 يجب أن تُرجع `"FizzBuzz"`.

```js
assert.equal(fizzBuzz()[89], 'FizzBuzz');
```

الأعداد غير القابلة للقسمة على 3 أو 5 يجب أن تُرجع الرقم نفسه.

```js
assert.equal(fizzBuzz()[12], 13);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz() {

}
```

# --solutions--

```js
function fizzBuzz() {
    let res=[];
    for (let i =1; i < 101; i++) {
        if (i % 3 === 0  && i % 5 === 0) {
            res.push("FizzBuzz");
        }
        else if (i % 3 === 0) {
            res.push("Fizz");
        }
        else if (i % 5 === 0) {
            res.push("Buzz");
        } 
        else {
            res.push(i);
        }
    }
    return res;
}
```
