---
id: 5ea2815a8640bcc6cb7dab3c
title: أعداد ليشريل
challengeType: 1
forumTopicId: 385287
dashedName: lychrel-numbers
---

# --description--

<ol>
<li>خذ عددًا صحيحًا <code>n₀</code> أكبر من الصفر.</li>
  <li>كوّن العدد التالي <code>n</code> في المتتالية بعكس <code>n₀</code> وجمعه مع <code>n₀</code></li>
  <li>توقف عندما يصبح <code>n</code> متناظرًا - أي أن أرقام <code>n</code> بالعكس تساوي <code>n</code>.</li>
</ol>

العلاقة التكرارية أعلاه عند تطبيقها على معظم الأعداد الابتدائية `n` = 1، 2، ... تنتهي بسرعة إلى عدد متناظر.

على سبيل المثال إذا كانت `n₀` = 12 نحصل على:

```md
12
12 + 21 = 33,  a palindrome!
```

وإذا كانت `n₀` = 55 نحصل على:

```md
55
55 + 55 = 110
110 + 011 = 121,  a palindrome!
```

لاحظ أن التحقق من التناظر يحدث *بعد* الجمع.

بعض الأعداد الابتدائية تبدو وكأنها تستمر إلى الأبد؛ العلاقة التكرارية للعدد 196 حسبت لملايين التكرارات مكونة أعدادًا بملايين الأرقام، دون أن تتكون عدد متناظر. هذه الأعداد التي لا تنتهي بعدد متناظر تسمى **أعداد ليشريل**.

لأغراض هذه المهمة، عدد ليشريل هو أي عدد ابتدائي لا يشكل عددًا متناظرًا خلال 500 تكرار (أو أكثر).

**أعداد ليشريل البذرة والمرتبطة:**

أي عدد صحيح ينتج في متتالية عدد ليشريل هو أيضًا عدد ليشريل.

بشكل عام، أي متتالية من عدد ليشريل *قد* تتقارب لتلتحق بالمتتالية من عدد ليشريل مرشح سابق؛ على سبيل المثال تبدأ المتتاليات للأعداد 196 ثم 689 بـ:

```md
    196
    196 + 691 = 887
    887 + 788 = 1675
    1675 + 5761 = 7436
    7436 + 6347 = 13783
    13783 + 38731 = 52514
    52514 + 41525 = 94039
    ...
    689
    689 + 986 = 1675
    1675 + 5761 = 7436
    ...
```

نرى إذًا أن المتتالية التي تبدأ بـ 689 تتقارب وتستمر بنفس الأعداد كما في متتالية 196.

بسبب هذا يمكننا تقسيم أعداد ليشريل إلى مرشحين حقيقيين لـ **البذرة**، وأعداد **مرتبطة** لا تنتج أعدادًا متناظرة لكنها تحتوي على أعداد صحيحة في متتاليتها تظهر كجزء من المتتالية الناتجة من عدد ليشريل أقل.

# --instructions--

اكتب دالة تأخذ عددًا كمعلمة. أرجع true إذا كان العدد عدد ليشريل. وإلا، أرجع false. تذكر أن حد التكرار هو 500.

# --hints--

`isLychrel` يجب أن تكون دالة.

```js
assert(typeof isLychrel === 'function');
```

`isLychrel(12)` يجب أن ترجع قيمة منطقية.

```js
assert(typeof isLychrel(12) === 'boolean');
```

`isLychrel(12)` يجب أن ترجع `false`.

```js
assert.equal(isLychrel(12), false);
```

`isLychrel(55)` يجب أن ترجع `false`.

```js
assert.equal(isLychrel(55), false);
```

`isLychrel(196)` يجب أن ترجع `true`.

```js
assert.equal(isLychrel(196), true);
```

`isLychrel(879)` يجب أن ترجع `true`.

```js
assert.equal(isLychrel(879), true);
```

`isLychrel(44987)` يجب أن ترجع `false`.

```js
assert.equal(isLychrel(44987), false);
```

`isLychrel(7059)` يجب أن ترجع `true`.

```js
assert.equal(isLychrel(7059), true);
```

# --seed--

## --seed-contents--

```js
function isLychrel(n) {

}
```

# --solutions--

```js
function isLychrel(n) {
  function reverse(num) {
    return parseInt(
      num
        .toString()
        .split('')
        .reverse()
        .join('')
    );
  }

  var i;
  for (i = 0; i < 500; i++) {
    n = n + reverse(n);
    if (n == reverse(n)) break;
  }

  return i == 500;
}
```
