---
id: 5ea28156e79528a9ab248f27
title: اختبار Luhn لأرقام بطاقات الائتمان
challengeType: 1
forumTopicId: 385284
dashedName: luhn-test-of-credit-card-numbers
---

# --description--

يستخدم اختبار Luhn من قبل بعض شركات بطاقات الائتمان لتمييز أرقام بطاقات الائتمان الصالحة عن ما قد يكون اختيارًا عشوائيًا من الأرقام.

تستخدم تلك الشركات أرقام بطاقات ائتمان يمكن التحقق من صحتها بواسطة اختبار Luhn، حيث تمر الأرقام بالاختبار التالي:

<ol>
<li> عكس ترتيب الأرقام في الرقم.</li>
  <li> أخذ الرقم الأول، الثالث، ... وكل رقم فردي آخر في الأرقام المعكوسة وجمعها لتشكيل المجموع الجزئي s1</li>
  <li> أخذ الرقم الثاني، الرابع ... وكل رقم زوجي آخر في الأرقام المعكوسة:</li>
    <ol>
<li>ضرب كل رقم في اثنين وجمع أرقامه إذا كانت النتيجة أكبر من تسعة لتشكيل المجاميع الجزئية للأرقام الزوجية.</li>
      <li>جمع المجاميع الجزئية للأرقام الزوجية لتشكيل s2.</li>
    </ol>
<li>إذا كان مجموع s1 + s2 ينتهي بصفر، فإن الرقم الأصلي يكون على شكل رقم بطاقة ائتمان صالح كما يتحقق من ذلك اختبار Luhn.</li>
</ol>

على سبيل المثال، إذا كان الرقم التجريبي هو 49927398716:

```md
Reverse the digits:
  61789372994
Sum the odd digits:
  6 + 7 + 9 + 7 + 9 + 4 = 42 = s1
The even digits:
    1, 8, 3, 2, 9
  Two times each even digit:
    2, 16, 6, 4, 18
  Sum the digits of each multiplication:
    2, 7, 6, 4, 9
  Sum the last:
    2 + 7 + 6 + 4 + 9 = 28 = s2

s1 + s2 = 70 which ends in zero which means that 49927398716 passes the Luhn test.
```

# --instructions--

اكتب دالة تتحقق من صحة رقم باستخدام اختبار Luhn. أرجع true إذا كان الرقم صالحًا، وإلا أرجع false.

# --hints--

يجب أن تكون `luhnTest` دالة.

```js
assert(typeof luhnTest === 'function');
```

يجب أن تُرجع `luhnTest("4111111111111111")` قيمة منطقية.

```js
assert(typeof luhnTest('4111111111111111') === 'boolean');
```

يجب أن تُرجع `luhnTest("4111111111111111")` القيمة `true`.

```js
assert.equal(luhnTest('4111111111111111'), true);
```

يجب أن تُرجع `luhnTest("4111111111111112")` القيمة `false`.

```js
assert.equal(luhnTest('4111111111111112'), false);
```

يجب أن تُرجع `luhnTest("49927398716")` القيمة `true`.

```js
assert.equal(luhnTest('49927398716'), true);
```

يجب أن تُرجع `luhnTest("49927398717")` القيمة `false`.

```js
assert.equal(luhnTest('49927398717'), false);
```

يجب أن تُرجع `luhnTest("1234567812345678")` القيمة `false`.

```js
assert.equal(luhnTest('1234567812345678'), false);
```

يجب أن تُرجع `luhnTest("1234567812345670")` القيمة `true`.

```js
assert.equal(luhnTest('1234567812345670'), true);
```

# --seed--

## --seed-contents--

```js
function luhnTest(str) {

}
```

# --solutions--

```js
function luhnTest(str) {
  var luhnArr = [0, 2, 4, 6, 8, 1, 3, 5, 7, 9];
  var counter = 0;
  var incNum;
  var odd = false;
  var temp = String(str).replace(/[^\d]/g, '');
  if (temp.length == 0) return false;
  for (var i = temp.length - 1; i >= 0; --i) {
    incNum = parseInt(temp.charAt(i), 10);
    counter += (odd = !odd) ? incNum : luhnArr[incNum];
  }
  return counter % 10 == 0;
}
```
