---
id: 5eaf48389ee512d4d103684b
title: أعداد ذاتية الوصف
challengeType: 1
forumTopicId: 385289
dashedName: self-describing-numbers
---

# --description--

هناك عدة أعداد صحيحة تُسمى "ذاتية الوصف" أو "ذاتية التوصيف".

يُقال عن العدد الصحيح إنه "ذاتي الوصف" إذا كان لديه الخاصية التي، عند تسمية مواقع الأرقام من 0 إلى N-1، تكون قيمة الرقم في كل موقع مساوية لعدد مرات ظهور ذلك الرقم في العدد.

على سبيل المثال، **2020** هو عدد ذاتي الوصف مكون من أربعة أرقام:

<ul>
<li> الموقع 0 له القيمة 2 ويوجد صفران 0 في العدد؛ </li>
    <li> الموقع 1 له القيمة 0 ولا يوجد أي 1 في العدد؛ </li>
    <li> الموقع 2 له القيمة 2 ويوجد صفران 2؛ </li>
    <li> الموقع 3 له القيمة 0 ولا يوجد أي 3 في العدد؛ </li>
</ul>

الأعداد ذاتية الوصف الأقل من 100,000,000 هي: 1210، 2020، 21200، 3211000، 42101000.

# --instructions--

اكتب دالة تأخذ عددًا صحيحًا موجبًا كمعلمة. إذا كان ذاتي الوصف، أرجع true. وإلا، أرجع false.

# --hints--

يجب أن تكون `isSelfDescribing` دالة.

```js
assert(typeof isSelfDescribing == 'function');
```

يجب أن تُرجع `isSelfDescribing()` قيمة منطقية.

```js
assert(typeof isSelfDescribing(2020) == 'boolean');
```

يجب أن تُرجع `isSelfDescribing(2020)` القيمة `true`.

```js
assert.equal(isSelfDescribing(2020), true);
```

يجب أن تُرجع `isSelfDescribing(3021)` القيمة `false`.

```js
assert.equal(isSelfDescribing(3021), false);
```

يجب أن تُرجع `isSelfDescribing(3211000)` القيمة `true`.

```js
assert.equal(isSelfDescribing(3211000), true);
```

# --seed--

## --seed-contents--

```js
function isSelfDescribing(n) {

}
```

# --solutions--

```js
function isSelfDescribing(n) {
    let digits = String(n).split("");
    digits = digits.map(function(e) {return parseInt(e)});
    let count = digits.map((x) => {return 0})
    digits.forEach((d) =>{
        if (d >= count.length) {
            return false
        }
        count[d] += 1;
    });

     if (digits === count) {
        return true;
    }
    if (digits.length != count.length) {
        return false;
    }
    
    for (let i=0; i< digits.length; i++){
      if (digits[i] !== count[i]) {
        return false;
      }
    }
    return true;
}
```
