---
id: 5eb3e497b8d6d7f63c5517ea
title: البحث في قائمة السجلات
challengeType: 1
forumTopicId: 385315
dashedName: search-a-list-of-records
---

# --description--

يتكون السجل من خصائص تصف كيانًا. لكل خاصية اسم وقيمة. على سبيل المثال، يمكن أن يكون للشخص خاصية `age` بقيمة 25. عملية مهمة على قائمة السجلات هي إيجاد سجل بقيمة خاصية معينة.

# --instructions--

اكتب دالة تأخذ سلسلة نصية كمعلمة. يجب أن تُرجع الدالة فهرس العنصر في `list` الذي تتطابق فيه قيمة الخاصية `name` مع السلسلة النصية المعطاة.

# --hints--

يجب أن تكون `searchCity` دالة.

```js
assert(typeof searchCity === 'function');
```

يجب أن تُرجع `searchCity("Dar Es Salaam")` رقمًا.

```js
assert(typeof searchCity('Dar Es Salaam') === 'number');
```

يجب أن تُرجع `searchCity("Dar Es Salaam")` القيمة `6`.

```js
assert.equal(searchCity('Dar Es Salaam'), 6);
```

يجب أن تُرجع `searchCity("Casablanca")` القيمة `9`.

```js
assert.equal(searchCity('Casablanca'), 9);
```

يجب أن تُرجع `searchCity("Cairo")` القيمة `1`.

```js
assert.equal(searchCity('Cairo'), 1);
```

يجب أن تُرجع `searchCity("Mogadishu")` القيمة `4`.

```js
assert.equal(searchCity('Mogadishu'), 4);
```

يجب أن تُرجع `searchCity("Lagos")` القيمة `0`.

```js
assert.equal(searchCity('Lagos'), 0);
```

# --seed--

## --seed-contents--

```js
function searchCity(name) {

}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```

# --solutions--

```js
function searchCity(name) {
  return list.findIndex(item => item.name === name);
}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```
