---
id: 5eb3e4a21f462f409d656c73
title: متتالية ذاتية الإشارة
challengeType: 1
forumTopicId: 385317
dashedName: self-referential-sequence
---

# --description--

هناك عدة طرق لتوليد متتالية ذاتية الإشارة. واحدة شائعة جدًا (متتالية <a href="https://rosettacode.org/wiki/Look-and-say_sequence" target="_blank" rel="noopener noreferrer nofollow">النظر والقول</a>) تبدأ برقم صحيح موجب، ثم تولد الحد التالي بربط مجموعات مرقمة من الأرقام المتجاورة المتشابهة:

<pre>0, 10, 1110, 3110, 132110, 1113122110, 311311222110 ...</pre>

الحدود المولدة تزداد طولًا بشكل هندسي ولا تتقارب أبدًا.

طريقة أخرى لتوليد متتالية ذاتية الإشارة هي تلخيص الحد السابق.

احسب عدد كل رقم متشابه، ثم اربط المجموع والرقم لكل رقم مرقم ومرتب. لاحظ أن الحدود الخمسة الأولى هي نفسها كما في المتتالية السابقة.

<pre>0, 10, 1110, 3110, 132110, 13123110, 23124110 ...</pre>

رتب الأرقام من الأكبر إلى الأصغر. لا تشمل عدد الأرقام التي لا تظهر في الحد السابق.

اعتمادًا على قيمة البذرة، السلاسل المولدة بهذه الطريقة إما تتقارب دائمًا إلى قيمة مستقرة أو إلى نمط دوري قصير. (لأغراضنا، التقارب يعني أن عنصرًا يطابق عنصرًا سبق رؤيته.) المتتالية المعروضة، بقيمة بذرة 0، تتقارب إلى قيمة مستقرة 1433223110 بعد 11 تكرارًا. قيمة البذرة التي تتقارب بأسرع وقت هي 22. تصبح مستقرة بعد العنصر الأول. (العنصر التالي هو 22، الذي سبق رؤيته.)

# --instructions--

اكتب دالة تأخذ قيمة البذرة كمعلمة، تولد متتالية ذاتية الإشارة حتى تتقارب، وتعيدها كمصفوفة.

# --hints--

`selfReferential` يجب أن تكون دالة.

```js
assert(typeof selfReferential === 'function');
```

`selfReferential(40)` يجب أن تعيد مصفوفة.

```js
assert(Array.isArray(selfReferential(40)));
```

`selfReferential(40)` يجب أن تعيد `["40", "1410", "142110", "14123110", "1413124110", "2413125110", "151413224110", "152413225110", "251413324110", "152423224110", "152413423110"]`.

```js
assert.deepEqual(selfReferential(40), [
  '40',
  '1410',
  '142110',
  '14123110',
  '1413124110',
  '2413125110',
  '151413224110',
  '152413225110',
  '251413324110',
  '152423224110',
  '152413423110'
]);
```

`selfReferential(132110)` يجب أن تعيد `["132110", "13123110", "23124110", "1413223110", "1423224110", "2413323110", "1433223110"]`.

```js
assert.deepEqual(selfReferential(132110), [
  '132110',
  '13123110',
  '23124110',
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(132211)` يجب أن تعيد `["132211", "132231", "232221", "134211", "14131231", "14231241", "24132231", "14233221"]`.

```js
assert.deepEqual(selfReferential(132211), [
  '132211',
  '132231',
  '232221',
  '134211',
  '14131231',
  '14231241',
  '24132231',
  '14233221'
]);
```

`selfReferential(1413223110)` يجب أن تعيد `["1413223110", "1423224110", "2413323110", "1433223110"]`.

```js
assert.deepEqual(selfReferential(1413223110), [
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(251413126110)` يجب أن تعيد `["251413126110", "16151413225110", "16251413226110", "26151413325110", "16251423225110", "16251413424110", "16153413225110"]`.

```js
assert.deepEqual(selfReferential(251413126110), [
  '251413126110',
  '16151413225110',
  '16251413226110',
  '26151413325110',
  '16251423225110',
  '16251413424110',
  '16153413225110'
]);
```

# --seed--

## --seed-contents--

```js
function selfReferential(n) {

}
```

# --solutions--

```js
function selfReferential(n) {
  var descending,
    i,
    incr,
    j,
    max_i,
    max_len,
    max_seq,
    seq,
    sequence,
    indexOf =
      [].indexOf ||
      function(item) {
        for (var i = 0, l = this.length; i < l; i++) {
          if (i in this && this[i] === item) return i;
        }
        return -1;
      };

  sequence = function(n) {
    var c, cnt, cnts, d, digit, i, j, l, len, new_cnts, ref, s, seq;
    cnts = {};
    ref = n.toString();
    for (j = 0, len = ref.length; j < len; j++) {
      c = ref[j];
      d = parseInt(c);
      incr(cnts, d);
    }
    seq = [ref];
    while (true) {
      s = '';
      for (i = l = 9; l >= 0; i = --l) {
        if (cnts[i]) {
          s += '' + cnts[i] + i;
        }
      }
      if (indexOf.call(seq, s) >= 0) {
        break;
      }
      seq.push(s);
      new_cnts = {};
      for (digit in cnts) {
        cnt = cnts[digit];
        incr(new_cnts, cnt);
        incr(new_cnts, digit);
      }
      cnts = new_cnts;
    }
    return seq;
  };

  incr = function(h, k) {
    if (h[k] == null) {
      h[k] = 0;
    }
    return (h[k] += 1);
  };

  descending = function(n) {
    var tens;
    if (n < 10) {
      return true;
    }
    tens = n / 10;
    if (n % 10 > tens % 10) {
      return false;
    }
    return descending(tens);
  };

  return sequence(n);
}
```
