---
id: 5eb3e4b20aa93c437f9e9717
title: مجموعة الأعداد الحقيقية
challengeType: 1
forumTopicId: 385322
dashedName: set-of-real-numbers
---

# --description--

تشكل جميع الأعداد الحقيقية المجموعة غير القابلة للعد ℝ. من بين مجموعاتها الجزئية، المجموعات المحدبة هي أبسط نسبيًا، حيث يُعبَّر عن كل منها كنطاق بين عددين حقيقيين *a* و *b* بحيث *a* ≤ *b*. هناك في الواقع أربع حالات لمعنى "بين"، حسب كون الحدود مفتوحة أو مغلقة:

<ul>
<li>[<i>a</i>, <i>b</i>]: {<i>x</i> | <i>a</i> ≤ <i>x</i> و <i>x</i> ≤ <i>b</i> }</li>
  <li>(<i>a</i>, <i>b</i>): {<i>x</i> | <i>a</i> < <i>x</i> و <i>x</i> < <i>b</i> }</li>
  <li>[<i>a</i>, <i>b</i>): {<i>x</i> | <i>a</i> ≤ <i>x</i> و <i>x</i> < <i>b</i> }</li>
  <li>(<i>a</i>, <i>b</i>]: {<i>x</i> | <i>a</i> < <i>x</i> و <i>x</i> ≤ <i>b</i> }</li>
</ul>

لاحظ أنه إذا كان *a* = *b*، فإن من بين الأربع حالات فقط \[*a*, *a*] ستكون غير فارغة.

**المهمة**

<ul>
<li>ابتكر طريقة لتمثيل أي مجموعة من الأعداد الحقيقية، حسب تعريف "أي" في ملاحظات التنفيذ أدناه.</li>
  <li>وفر طرقًا لهذه العمليات الشائعة على المجموعات (<i>x</i> عدد حقيقي؛ <i>A</i> و <i>B</i> مجموعتان):</li>
  <ul>
    <li>
<i>x</i> ∈ <i>A</i>: حدد ما إذا كان <i>x</i> عنصرًا في <i>A</i><br>
      مثال: 1 موجود في [1, 2)، بينما 2، 3، ... غير موجودين.
    </li>
    <li>
<i>A</i> ∪ <i>B</i>: اتحاد <i>A</i> و <i>B</i>، أي {<i>x</i> | <i>x</i> ∈ <i>A</i> أو <i>x</i> ∈ <i>B</i>}<br>
      مثال: [0, 2) ∪ (1, 3) = [0, 3)؛ [0, 1) ∪ (2, 3] = حسنًا، [0, 1) ∪ (2, 3]
    </li>
    <li>
<i>A</i> ∩ <i>B</i>: تقاطع <i>A</i> و <i>B</i>، أي {<i>x</i> | <i>x</i> ∈ <i>A</i> و <i>x</i> ∈ <i>B</i>}<br>
      مثال: [0, 2) ∩ (1, 3) = (1, 2)؛ [0, 1) ∩ (2, 3] = مجموعة فارغة
    </li>
    <li>
<i>A</i> - <i>B</i>: الفرق بين <i>A</i> و <i>B</i>، ويُكتب أيضًا <i>A</i> \ <i>B</i>، أي {<i>x</i> | <i>x</i> ∈ <i>A</i> و <i>x</i> ∉ <i>B</i>}<br>
      مثال: [0, 2) − (1, 3) = [0, 1]
    </li>
  </ul>
</ul>

# --instructions--

اكتب دالة تأخذ كمعلمات كائنين، وسلسلة نصية، ومصفوفة. تمثل الكائنات المجموعة ولها خصائص: `low`، `high` و `rangeType`.

يمكن أن تكون قيمة `rangeType` 0، 1، 2، أو 3 لـ `CLOSED`، `BOTH_OPEN`، `LEFT_OPEN` و `RIGHT_OPEN`، على التوالي. يجب أن تنفذ الدالة مجموعة باستخدام هذه المعلومات.

تمثل السلسلة النصية العملية التي ستُجرى على المجموعات. يمكن أن تكون: `"union"`، `"intersect"` و `"subtract"` (الفرق).

بعد تنفيذ العملية، يجب أن تتحقق الدالة مما إذا كانت القيم في المصفوفة موجودة في المجموعة الناتجة وتخزن قيمة منطقية مقابلة في مصفوفة. يجب أن تُرجع الدالة هذه المصفوفة.

# --hints--

يجب أن تكون `realSet` دالة.

```js
assert(typeof realSet == 'function');
```

`realSet({"low":0, "high":1, "rangeType":2}, {"low":0, "high":2, "rangeType":3}, "union", [1, 2, 3])` يجب أن تُرجع مصفوفة.

```js
assert(
  Array.isArray(
    realSet(
      { low: 0, high: 1, rangeType: 2 },
      { low: 0, high: 2, rangeType: 3 },
      'union',
      [1, 2, 3]
    )
  )
);
```

`realSet({"low":0, "high":1, "rangeType":2}, {"low":0, "high":2, "rangeType":3}, "union", [1, 2, 3])` يجب أن تُرجع `[true, false, false]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 1, rangeType: 2 },
    { low: 0, high: 2, rangeType: 3 },
    'union',
    [1, 2, 3]
  ),
  [true, false, false]
);
```

`realSet({"low":0, "high":2, "rangeType":3}, {"low":1, "high":2, "rangeType":2}, "intersect", [0, 1, 2])` يجب أن تُرجع `[false, false, false]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 2, rangeType: 3 },
    { low: 1, high: 2, rangeType: 2 },
    'intersect',
    [0, 1, 2]
  ),
  [false, false, false]
);
```

`realSet({"low":0, "high":3, "rangeType":3}, {"low":0, "high":1, "rangeType":1}, "subtract", [0, 1, 2])` يجب أن تُرجع `[true, true, true]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 3, rangeType: 3 },
    { low: 0, high: 1, rangeType: 1 },
    'subtract',
    [0, 1, 2]
  ),
  [true, true, true]
);
```

`realSet({"low":0, "high":3, "rangeType":3}, {"low":0, "high":1, "rangeType":0}, "subtract", [0, 1, 2])` يجب أن تُرجع `[false, false, true]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 3, rangeType: 3 },
    { low: 0, high: 1, rangeType: 0 },
    'subtract',
    [0, 1, 2]
  ),
  [false, false, true]
);
```

`realSet({"low":0, "high":33, "rangeType":1}, {"low":30, "high":31, "rangeType":0}, "intersect", [30, 31, 32])` يجب أن تُرجع `[true, true, false]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 33, rangeType: 1 },
    { low: 30, high: 31, rangeType: 0 },
    'intersect',
    [30, 31, 32]
  ),
  [true, true, false]
);
```

# --seed--

## --seed-contents--

```js
function realSet(set1, set2, operation, values) {

}
```

# --solutions--

```js
function realSet(set1, set2, operation, values) {
  const RangeType = {
    CLOSED: 0,
    BOTH_OPEN: 1,
    LEFT_OPEN: 2,
    RIGHT_OPEN: 3
  };

  function Predicate(test) {
    this.test = test;
    this.or = function(other) {
      return new Predicate(t => this.test(t) || other.test(t));
    };
    this.and = function(other) {
      return new Predicate(t => this.test(t) && other.test(t));
    };
    this.negate = function() {
      return new Predicate(t => !this.test(t));
    };
  }

  function RealSet(start, end, rangeType, predF) {
    this.low = start;
    this.high = end;

    if (predF) {
      this.predicate = new Predicate(predF);
    } else {
      this.predicate = new Predicate(d => {
        switch (rangeType) {
          case RangeType.CLOSED:
            return start <= d && d <= end;
          case RangeType.BOTH_OPEN:
            return start < d && d < end;
          case RangeType.LEFT_OPEN:
            return start < d && d <= end;
          case RangeType.RIGHT_OPEN:
            return start <= d && d < end;
        }
      });
    }

    this.contains = function(d) {
      return this.predicate.test(d);
    };

    this.union = function(other) {
      var low2 = Math.min(this.low, other.low);
      var high2 = Math.max(this.high, other.high);
      return new RealSet(low2, high2, null, d =>
        this.predicate.or(other.predicate).test(d)
      );
    };

    this.intersect = function(other) {
      var low2 = Math.min(this.low, other.low);
      var high2 = Math.max(this.high, other.high);
      return new RealSet(low2, high2, null, d =>
        this.predicate.and(other.predicate).test(d)
      );
    };

    this.subtract = function(other) {
      return new RealSet(this.low, this.high, null, d =>
        this.predicate.and(other.predicate.negate()).test(d)
      );
    };
  }
  set1 = new RealSet(set1.low, set1.high, set1.rangeType);
  set2 = new RealSet(set2.low, set2.high, set2.rangeType);
  var result = [];
  values.forEach(function(value) {
    result.push(set1[operation](set2).contains(value));
  });
  return result;
}
```
