---
id: 68e5293bd00d2fe134f58988
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Tekeleza njia ya `is_applicable` kwa darasa lako la `FixedAmountDiscount`. Punguzo la kiasi kilichowekwa lazima litumike tu ikiwa lina maana kiuchumi.

Rudisha `True` ikiwa `product.price * 0.9` ni kubwa kuliko `self.amount`, vinginevyo rudisha `False`. Hii inahakikisha punguzo ni chini ya asilimia 90 ya bei ya awali, hivyo si halisi kupita kiasi.

# --hints--

Darasa lako la `FixedAmountDiscount` linapaswa kuwa na njia ya `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("is_applicable")`)
})
```

Njia yako ya `is_applicable` inapaswa kuwa na vigezo `self`, `product: Product`, na `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_args("self, product:Product, user_tier:str")`)
})
```

Njia yako ya `is_applicable` inapaswa kuwa na aina ya kurudisha `-> bool`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

Njia yako ya `is_applicable` inapaswa kurudisha `True` wakati `product.price * 0.9` ni kubwa kuliko `self.amount`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(50)
    assert discount.is_applicable(test_product, "regular") == True
  `)
})
```

Njia yako ya `is_applicable` inapaswa kurudisha `False` wakati `product.price * 0.9` ni ndogo au sawa na `self.amount`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(95)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
