---
id: 68e524c94e2e4ab7b4bfc7e2
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda kalkuleta ya punguzo inayoweza kutumia mikakati tofauti ya punguzo kwa bidhaa. Mfumo utaamua bei bora kwa mteja kulingana na sheria nyingi za punguzo.

Anza kwa kufafanua darasa la vitu linaloitwa `Product`. Mpe darasa lako njia `__init__` yenye vigezo vifuatavyo: `self`, `name`, na `price`. Ndani ya njia `__init__`, weka thamani za vigezo `name` na `price` kwa sifa za kipengele za mfano zenye majina sawa.

# --hints--

Unapaswa kuunda darasa la vitu linaloitwa `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

Darasa lako la `Product` linapaswa kuwa na njia `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

Njia yako ya `__init__` inapaswa kuwa na vigezo vifuatavyo: `self`, `name`, na `price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

Unapaswa kuweka thamani ya `name` kwa `self.name` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

Unapaswa kuweka thamani ya `price` kwa `self.price` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
