---
id: 68e529226b0defdff26c4a89
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Sasa hivi, ukitengeneza mfano wa `Product` na kuuchapisha, utaona uwakilishi wa chaguo-msingi wa kitu ambacho hakutoa taarifa nyingi. Kwa mfano, inaweza kuonekana kama `<__main__.Product object at 0x7f9b8c2d1d30>`.

Ili kubadilisha uwakilishi wa chaguo-msingi wa kitu kuwa mfuatano wa herufi unaosomeka, tengeneza njia ya `__str__` katika darasa lako la `Product` ambayo inarudisha mfuatano wa herufi wenye muundo wa `name - $price` (ambapo `name` na `price` zinapaswa kubadilishwa na thamani za sifa za kipengele zinazolingana).

# --hints--

Darasa lako la `Product` linapaswa kuwa na njia ya `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

Njia yako ya `__str__` inapaswa kurudisha mfuatano wa herufi wenye muundo wa `name - $price`.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
        
--fcc-editable-region--
    
--fcc-editable-region--
```
