---
id: 68e52994bc7ea2e3dec530a1
title: Hatua 23
challengeType: 20
dashedName: step-23
---

# --description--

Sasa unahitaji njia ya kusimamia mikakati mingi ya punguzo na kupata bei bora zaidi.

Tengeneza darasa la vitu linaloitwa `DiscountEngine`. Darasa hili litasimamia mikakati yote ya punguzo na kuhesabu bei bora kwa bidhaa.

Mpe njia ya `__init__` inayopokea `self` na `strategies` kama vigezo, ambapo `strategies` ina aina ya `list[DiscountStrategy]`. Kutumia `list[DiscountStrategy]` kunaonyesha kuwa kigezo ni orodha yenye vitu vya `DiscountStrategy`. Aina ya kurudisha ya njia hiyo inapaswa kuwa `None`. Ndani ya njia ya `__init__`, hifadhi `strategies` kama sifa ya kipengele cha mfano. 

# --hints--

Unapaswa kutengeneza darasa la vitu linaloitwa `DiscountEngine`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountEngine")`)
})
```

Darasa lako la `DiscountEngine` linapaswa kuwa na njia ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("__init__")`)
})
```

Njia yako ya `__init__` inapaswa kuwa na vigezo `self` na `strategies: list[DiscountStrategy]`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_args("self, strategies: list[DiscountStrategy]")`)
})
```

Njia yako ya `__init__` inapaswa kuwa na aina ya kurudisha `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_returns("None")`)
})
```

Unapaswa kuweka `strategies` kwa `self.strategies` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_stmt("self.strategies = strategies")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
