---
id: 68e529226b0defdff26c4a8b
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Sasa utaanza kujenga kiolesura cha mfumo wa punguzo. Python hutoa njia ya kufafanua madarasa ya vitu ya abstrakti ambayo hutumika kama ramani za madarasa mengine. Hii hufanyika kwa kutumia darasa la `ABC` (Abstract Base Class) la moduli ya `abc`.

Juu ya msimbo wako, ongeza tamko lifuatalo la import:

```py
from abc import ABC, abstractmethod
```

Hii inaleta darasa la `ABC` na dekoreta wa `abstractmethod`, ambao utatumia kufafanua njia ambazo lazima zitekelezwe na darasa lolote linalorithi kutoka kwa darasa lako la abstrakti.

# --hints--

Unapaswa kuwa na tamko la import lililotajwa juu ya msimbo wako.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
