---
id: 68e5293bd00d2fe134f58985
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Kwa kuwa `PercentageDiscount` inarithi kutoka kwa darasa la vitu la abstract, lazima utekeleze njia zote za abstract. Anza kwa kutekeleza njia ya `is_applicable`. Njia ya `is_applicable` inapaswa kuchukua `self`, `product: Product`, na `user_tier: str` kama vigezo. Inapaswa kuwa na aina ya kurudisha yenye alama ya `bool`.

Njia hii inapaswa kubaini kama punguzo la asilimia ni halali. Rudisha `True` ikiwa `self.percent` ni ndogo au sawa na `70`, vinginevyo rudisha `False`. Hii inazuia punguzo zisizo halisi kama punguzo la asilimia 90.

# --hints--

Darasa lako la `PercentageDiscount` linapaswa kuwa na njia ya `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

Njia yako ya `is_applicable` inapaswa kuwa na vigezo `self`, `product: Product`, na `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

Njia yako ya `is_applicable` inapaswa kurudisha `True` wakati `self.percent` ni ndogo au sawa na `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

Njia yako ya `is_applicable` inapaswa kurudisha `False` wakati `self.percent` ni kubwa kuliko `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

Njia yako ya `is_applicable` inapaswa kuwa na alama ya aina ya kurudisha ya `bool`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
