---
id: 68e5293bd00d2fe134f58981
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Tengeneza darasa la vitu linaloitwa `DiscountStrategy` ambalo linarithi kutoka `ABC`. Hii itakuwa darasa la vitu la msingi la muhtasari linaloeleza kiolesura cha mikakati yote ya punguzo.

Kwa sasa, tumia tu `pass` ndani ya mwili wa darasa ili kuunda darasa tupu.

# --hints--

Unapaswa kuunda darasa la vitu linaloitwa `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

Darasa lako la `DiscountStrategy` linapaswa kurithi kutoka `ABC`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
