---
id: 68e5293bd00d2fe134f58987
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Sasa utaunda mkakati mwingine wa punguzo. Unda darasa la vitu linaloitwa `FixedAmountDiscount` ambalo linarithi kutoka `DiscountStrategy`.

Mpe njia ya `__init__` inayochukua kigezo `self` na `amount`. Mpe kigezo `amount` aina ya `int` na aina ya kurudisha `None` kwa njia hiyo. Ndani ya njia ya `__init__`, hifadhi `amount` kama sifa ya kipengele cha mfano.

# --hints--

Unapaswa kuunda darasa la vitu linaloitwa `FixedAmountDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

Darasa lako la `FixedAmountDiscount` linapaswa kurithi kutoka `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

Darasa lako la `FixedAmountDiscount` linapaswa kuwa na njia ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

Njia yako ya `__init__` inapaswa kuwa na vigezo `self` na `amount: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

Njia yako ya `__init__` inapaswa kuwa na aina ya kurudisha `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

Unapaswa kuweka thamani ya `amount` kwa `self.amount` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
