---
id: 68e5293bd00d2fe134f58984
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Sasa ni wakati wa kuunda utekelezaji halisi wa darasa lako la vitu la muhtasari. Unda darasa la vitu linaloitwa `PercentageDiscount` ambalo linarithi kutoka `DiscountStrategy`.

Mpe njia ya `__init__` inayopokea `self` na `percent` kama vigezo. Mpe dalili ya aina ya `int` kwa `percent` na weka aina ya kurudi ya njia hiyo kuwa `None`. Ndani ya njia ya `__init__`, hifadhi `percent` kama sifa ya kipengele cha mfano.

# --hints--

Unapaswa kuunda darasa la vitu linaloitwa `PercentageDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

Darasa lako la `PercentageDiscount` linapaswa kurithi kutoka `DiscountStrategy`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

Darasa lako la `PercentageDiscount` linapaswa kuwa na njia ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

Njia yako ya `__init__` inapaswa kuwa na vigezo `self` na `percent: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

Njia yako ya `__init__` inapaswa kuwa na kidokezo cha aina ya kurudisha cha `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

Unapaswa kuweka thamani ya `percent` kwa `self.percent` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

--fcc-editable-region--

--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
print(product)
```
