---
id: 68e5293bd00d2fe134f58983
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Sasa ongeza njia ya pili ya abstract iitwayo `apply_discount` katika darasa lako la `DiscountStrategy`. Tumia dekoreta ya `@abstractmethod` na uipe vigezo `self` na `product`.

Ongeza vidokezo vya aina: `product` inapaswa kuwa ya aina ya `Product`, na njia hiyo inapaswa kurudisha `float`. Tumia `pass` ndani ya mwili wa njia hiyo.

# --hints--

Darasa lako la `DiscountStrategy` linapaswa kuwa na njia iitwayo `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

Njia yako ya `apply_discount` inapaswa kuwa na dekoreta ya `@abstractmethod`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

Njia yako ya `apply_discount` inapaswa kuwa na vigezo `self` na `product: Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

Njia yako ya `apply_discount` inapaswa kuwa na kidokezo cha aina ya kurudisha `-> float`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
