---
id: 68e52994bc7ea2e3dec530a6
title: Hatua 29
challengeType: 20
dashedName: step-29
---

# --description--

Sasa ni wakati wa kujaribu mfumo kamili wa punguzo! Ongeza tamko la `if` linalokagua kama `__name__ == '__main__'` na hamisha `product = Product('Wireless Mouse', 50.0)` ndani yake.

Baada ya hapo, ndani ya blokki ya `if`, tengeneza kigezo `user_tier` na uweke thamani yake kuwa mfuatano wa herufi `Premium`. Kisha tengeneza kigezo `strategies` kilicho na orodha ya mikakati inayojumuisha `PercentageDiscount(10)`, `FixedAmountDiscount(5)`, na `PremiumUserDiscount()`.

# --hints--

Unapaswa kuwa na tamko la `if` linalokagua kama `__name__ == '__main__'`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent('__name__ == "__main__"')`)
})
```

Unapaswa kuhama `product = Product('Wireless Mouse', 50.0)` ndani ya blokki ya `if`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt("product = Product('Wireless Mouse', 50.0)")`)
})
```

Unapaswa kuweka `user_tier` kuwa mfuatano wa herufi `Premium` ndani ya blokki ya `if`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('user_tier = "Premium"')`)
})
```

Unapaswa kutengeneza orodha iitwayo `strategies` inayojumuisha `PercentageDiscount(10)`, `FixedAmountDiscount(5)`, na `PremiumUserDiscount()`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('strategies = [PercentageDiscount(10), FixedAmountDiscount(5), PremiumUserDiscount()]')`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

class DiscountEngine:
    def __init__(self, strategies: list[DiscountStrategy]) -> None:
        self.strategies = strategies

    def calculate_best_price(self, product: Product, user_tier: str) -> float:
        prices = [product.price]

        for strategy in self.strategies:
            if strategy.is_applicable(product, user_tier):
                discounted = strategy.apply_discount(product)
                prices.append(discounted)

        return min(prices)
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
--fcc-editable-region--
```
