---
id: 68e79d91daf8fb4808f93c47
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Baada ya kutekeleza njia ya `apply_discount`, jaribu mkakati wako wa punguzo. Unda mfano wa `PercentageDiscount` ukiwa na asilimia `10` na weka thamani yake katika kigezo kinachoitwa `discount`. Kisha tumia punguzo kwa kutumia `discount.apply_discount(product)`, na chapisha matokeo kuona punguzo likifanya kazi.

# --hints--

Unapaswa kuunda kigezo kinachoitwa `discount` chenye mfano wa `PercentageDiscount`.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

Mfano wako wa `PercentageDiscount` unapaswa kuanzishwa ukiwa na asilimia `10`.

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

Unapaswa kuita `discount.apply_discount(product)` na kuchapisha matokeo.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
