---
id: 68e9207ddcba053c5cb24413
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Katika Python, vidokezo vya aina ni ishara za hiari zinazowaambia wasanidi programu wengine ni aina gani ya data inayotarajiwa kwa kigezo au kitendakazi. Hapa kuna mfano unaoongeza kidokezo cha aina `str` kwa kigezo `name`.

```py
def greet(name: str):
    return 'Hello, ' + name
```

Kwa vigezo vya njia `__init__`, sasisha kigezo `name` ili kiwe na kidokezo cha aina `str` na kigezo `price` kiwe na kidokezo cha aina `float`.

# --hints--

Unapaswa kuongeza kidokezo cha aina `str` kwa kigezo `name` na kidokezo cha aina `float` kwa kigezo `price` katika njia `__init__` ya darasa la vitu `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
--fcc-editable-region--
    def __init__(self, name, price):
--fcc-editable-region--
        self.name = name
        self.price = price
```
