---
id: 68f119ad905b29ff59d858c8
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Kama ilivyo katika hatua iliyopita, katika njia yako ya `__str__`, ongeza aina ya kurudi ya `str` kwa sababu njia hii inarudisha uwakilishi wa mfuatano wa herufi wa kitu.

# --hints--

Njia yako ya `__str__` inapaswa kuwa na aina ya kurudi ya `str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
