---
id: bad87fee1348bd9aedf08805
title: "CSS-Selektoren verwenden, um Elemente zu gestalten"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKMBT2'
forumTopicId: 18349
dashedName: use-css-selectors-to-style-elements
---

# --description--

Mit CSS gibt es Hunderte von CSS-Eigenschaften, mit denen Sie das Aussehen eines Elements auf Ihrer Seite ändern können.

Als Sie `<h2 style="color: red;">CatPhotoApp</h2>` eingegeben haben, haben Sie dieses einzelne `h2`-Element mit inline CSS gestaltet, was für Cascading Style Sheets steht.

Das ist eine Möglichkeit, den Stil eines Elements festzulegen, aber es gibt eine bessere Methode, CSS anzuwenden.

Erstellen Sie oben in Ihrem Code einen `style`-Block wie diesen:

```html
<style>
</style>
```

Innerhalb dieses Style-Blocks können Sie einen <dfn>CSS-Selektor</dfn> für alle `h2`-Elemente erstellen. Wenn Sie beispielsweise alle `h2`-Elemente rot färben möchten, würden Sie eine Stilregel hinzufügen, die so aussieht:

```html
<style>
  h2 {
    color: red;
  }
</style>
```

Beachten Sie, dass es wichtig ist, sowohl öffnende als auch schließende geschweifte Klammern (`{` und `}`) um jede Stilregel des Elements zu haben. Außerdem müssen Sie sicherstellen, dass die Stildefinition Ihres Elements zwischen den öffnenden und schließenden Style-Tags steht. Schließlich fügen Sie am Ende jeder Stilregel Ihres Elements ein Semikolon hinzu.

# --instructions--

Löschen Sie das Style-Attribut Ihres `h2`-Elements und erstellen Sie stattdessen einen CSS-`style`-Block. Fügen Sie das notwendige CSS hinzu, um alle `h2`-Elemente blau zu färben.

# --hints--

Das `style`-Attribut sollte von Ihrem `h2`-Element entfernt werden.

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

Sie sollten ein `style`-Element erstellen.

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
assert.isAtLeast(document.querySelectorAll('style:not(.fcc-hide-header)').length, 1);
```

Ihr `h2`-Element sollte blau sein.

```js
const h2Element = document.querySelector('h2');
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

Die Deklaration Ihres Stylesheet-`h2` sollte mit einem Semikolon und einer schließenden Klammer gültig sein.

```js
assert.match(__helpers.removeCssComments(code), /h2\s*\{\s*color\s*:.*;\s*\}/g);
```

Alle Ihre `style`-Elemente sollten gültig sein und schließende Tags haben.

```js
assert.match(__helpers.removeHtmlComments(code), /<\/style>/g);

const closingElementLength = __helpers.removeHtmlComments(code).match(/<\/style>/g).length;
const openingElementsLength = __helpers.removeHtmlComments(code).match(/<style((\s)*((type|media|scoped|title|disabled)="[^"]*")?(\s)*)*>/g).length;
assert.strictEqual(closingElementLength, openingElementsLength);
```

# --seed--

## --seed-contents--

```html
<h2 style="color: red;">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  h2 {
    color: blue;
  }
</style>
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
