---
id: bad87fee1348bd8aedf06756
title: Deklarationen von Klassen durch Styling von ID-Attributen überschreiben
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpDhB'
forumTopicId: 18251
dashedName: override-class-declarations-by-styling-id-attributes
---

# --description--

Wir haben gerade bewiesen, dass Browser CSS von oben nach unten in der Reihenfolge ihrer Deklaration lesen. Das bedeutet, dass der Browser im Falle eines Konflikts die CSS-Deklaration verwendet, die zuletzt kam. Beachten Sie, dass selbst wenn wir `blue-text` vor `pink-text` in den Klassen unseres `h1`-Elements gesetzt hätten, der Browser trotzdem die Deklarationsreihenfolge und nicht die Reihenfolge ihrer Verwendung betrachtet!

Aber wir sind noch nicht fertig. Es gibt noch andere Möglichkeiten, CSS zu überschreiben. Erinnern Sie sich an id-Attribute?

Lassen Sie uns Ihre `pink-text`- und `blue-text`-Klassen überschreiben und Ihr `h1`-Element orange machen, indem Sie dem `h1`-Element eine id geben und diese id dann stylen.

# --instructions--

Geben Sie Ihrem `h1`-Element das `id`-Attribut mit dem Wert `orange-text`. Denken Sie daran, id-Stile sehen so aus:

```html
<h1 id="orange-text">
```

Lassen Sie die `blue-text`- und `pink-text`-Klassen an Ihrem `h1`-Element.

Erstellen Sie eine CSS-Deklaration für Ihre `orange-text`-id in Ihrem `style`-Element. So sieht das zum Beispiel aus:

```css
#brown-text {
  color: brown;
}
```

**Hinweis:** Es spielt keine Rolle, ob Sie dieses CSS über oder unter der `pink-text`-Klasse deklarieren, da das `id`-Attribut immer Vorrang hat.

# --hints--

Ihr `h1`-Element sollte die Klasse `pink-text` haben.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ihr `h1`-Element sollte die Klasse `blue-text` haben.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Ihr `h1`-Element sollte die id `orange-text` haben.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'),'orange-text');
```

Es sollte nur ein `h1`-Element geben.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Ihre `orange-text`-id sollte eine CSS-Deklaration haben.

```js
assert.match(__helpers.removeCssComments(code), /#orange-text\s*{/gi);
```

Ihr `h1` sollte keine `style`-Attribute haben.

```js
assert.notMatch(__helpers.removeHtmlComments(code), /<h1.*style.*>/gi);
```

Ihr `h1`-Element sollte orange sein.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 165, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
  #orange-text {
    color: orange;
  }  
</style>
<h1 id="orange-text"  class="pink-text blue-text">Hello World!</h1>
```
