---
id: bad87fee1348bd9aedf08746
title: Stile vom Body-Element erben
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bmdtR'
forumTopicId: 18204
dashedName: inherit-styles-from-the-body-element
---

# --description--

Wir haben nun bewiesen, dass jede HTML-Seite ein `body`-Element hat und dass sein `body`-Element ebenfalls mit CSS gestylt werden kann.

Denken Sie daran, dass Sie Ihr `body`-Element genauso wie jedes andere HTML-Element stylen können, und alle anderen Elemente erben die Stile Ihres `body`-Elements.

# --instructions--

Erstellen Sie zuerst ein `h1`-Element mit dem Text `Hello World`.

Dann geben Sie allen Elementen auf Ihrer Seite die Farbe `green`, indem Sie `color: green;` zur Stil-Deklaration Ihres `body`-Elements hinzufügen.

Geben Sie abschließend Ihrem `body`-Element die Schriftfamilie `monospace`, indem Sie `font-family: monospace;` zur Stil-Deklaration Ihres `body`-Elements hinzufügen.

# --hints--

Sie sollten ein `h1`-Element erstellen.

```js
assert.isNotEmpty(document.querySelectorAll('h1'));
```

Ihr `h1`-Element sollte den Text `Hello World` enthalten.

```js
assert.match(
  document.querySelector('h1').textContent,
  /hello world/i
);
```

Ihr `h1`-Element sollte ein schließendes Tag haben.

```js
const commentlessCode = __helpers.removeHtmlComments(code);
assert.match(commentlessCode, /<\/h1>/g);
assert.match(commentlessCode, /<h1/g);
assert.lengthOf(commentlessCode.match(/<\/h1>/g), commentlessCode.match(/<h1/g).length);
```

Ihr `body`-Element sollte die Eigenschaft `color` mit dem Wert `green` haben.

```js
const bodyElement = document.querySelector('body');
const color = window.getComputedStyle(bodyElement)['color']; 
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Ihr `body`-Element sollte die Eigenschaft `font-family` mit dem Wert `monospace` haben.

```js
const bodyElement = document.querySelector('body');
const fontFamily = window.getComputedStyle(bodyElement)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

Ihr `h1`-Element sollte die Schriftart `monospace` von Ihrem `body`-Element erben.

```js
const h1Element = document.querySelector('h1');
const fontFamily = window.getComputedStyle(h1Element)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

Ihr `h1`-Element sollte die Farbe `green` von Ihrem `body`-Element erben.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }

</style>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }

</style>
<h1>Hello World!</h1>
```
