---
id: bad87fee1348bd9aedf08756
title: Einen Stil gegenüber einem anderen priorisieren
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ8wnHv'
forumTopicId: 18258
dashedName: prioritize-one-style-over-another
---

# --description--

Manchmal erhalten Ihre HTML-Elemente mehrere sich widersprechende Stile.

Zum Beispiel kann Ihr `h1`-Element nicht gleichzeitig grün und rosa sein.

Sehen wir uns an, was passiert, wenn wir eine Klasse erstellen, die Text rosa macht, und diese dann auf ein Element anwenden. Wird unsere Klasse die `body`-Eigenschaft `color: green;` des Elements überschreiben?

# --instructions--

Erstellen Sie eine CSS-Klasse namens `pink-text`, die einem Element die Farbe Rosa gibt.

Geben Sie Ihrem `h1`-Element die Klasse `pink-text`.

# --hints--

Ihr `h1`-Element sollte die Klasse `pink-text` haben.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ihr `<style>` sollte eine `pink-text`-CSS-Klasse haben, die die `color` ändert.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*\{\s*color\s*:\s*.+\s*;?\s*\}/g);
```

Ihr `h1`-Element sollte rosa sein.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
</style>
<h1>Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```
