---
id: bad87fee1348bd9aedf08726
title: Hex-Code für bestimmte Farben verwenden
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

Wussten Sie, dass es in CSS noch andere Möglichkeiten gibt, Farben darzustellen? Eine dieser Möglichkeiten heißt hexadezimaler Code, oder kurz Hex-Code.

Wir verwenden normalerweise <dfn>Dezimalzahlen</dfn> oder Zahlen im Dezimalsystem, die die Symbole 0 bis 9 für jede Ziffer nutzen. <dfn>Hexadezimale Zahlen</dfn> (oder <dfn>Hex</dfn>) sind Zahlen im Hexadezimalsystem mit der Basis 16. Das bedeutet, dass sie sechzehn verschiedene Symbole verwenden. Wie bei Dezimalzahlen repräsentieren die Symbole 0-9 die Werte von null bis neun. Dann stehen A, B, C, D, E, F für die Werte zehn bis fünfzehn. Insgesamt können 0 bis F eine Ziffer im Hexadezimalsystem darstellen, was uns 16 mögliche Werte gibt. Mehr Informationen zu <a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">hexadezimalen Zahlen finden Sie hier</a>.

In CSS können wir 6 hexadezimale Ziffern verwenden, um Farben darzustellen, jeweils zwei für die roten (R), grünen (G) und blauen (B) Komponenten. Zum Beispiel ist `#000000` schwarz und stellt auch den niedrigsten möglichen Wert dar. Mehr Informationen zum <a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">RGB-Farbsystem finden Sie hier</a>.

```css
body {
  color: #000000;
}
```

# --instructions--

Ersetzen Sie das Wort `black` in der Hintergrundfarbe unseres `body`-Elements durch seine hexadezimale Darstellung `#000000`.

# --hints--

Ihr `body`-Element sollte die `background-color` schwarz haben.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Der Hex-Code für die Farbe Schwarz sollte anstelle des Wortes `black` verwendet werden.

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
