---
id: bad87fee1348bd9aedf06756
title: Klasse-Deklarationen mit inline-Styles überschreiben
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDRha'
forumTopicId: 18252
dashedName: override-class-declarations-with-inline-styles
---

# --description--

Wir haben also bewiesen, dass id-Deklarationen Klasse-Deklarationen überschreiben, egal wo sie in Ihrem `style` Element-CSS deklariert sind.

Es gibt noch andere Möglichkeiten, CSS zu überschreiben. Erinnern Sie sich an inline-Styles?

# --instructions--

Verwenden Sie einen inline-Style, um zu versuchen, unser `h1` Element weiß zu machen. Denken Sie daran, inline-Styles sehen so aus:

```html
<h1 style="color: green;">
```

Lassen Sie die `blue-text` und `pink-text` Klassen an Ihrem `h1` Element.

# --hints--

Ihr `h1` Element sollte die Klasse `pink-text` haben.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ihr `h1` Element sollte die Klasse `blue-text` haben.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Ihr `h1` Element sollte die id `orange-text` haben.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

Ihr `h1` Element sollte einen inline-Style haben.

```js
assert.exists(document.querySelector('h1[style]'));
```

Ihr `h1` Element sollte weiß sein.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
