---
id: bad82fee1322bd9aedf08721
title: Absolute und relative Einheiten verstehen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN66JSL'
forumTopicId: 301089
dashedName: understand-absolute-versus-relative-units
---

# --description--

In den letzten Aufgaben wurde der Außenabstand oder der Innenabstand eines Elements immer mit Pixeln (`px`) festgelegt. Pixel sind eine Art Längeneinheit, die dem Browser angibt, wie ein Element dimensioniert oder positioniert wird. Neben `px` bietet CSS verschiedene Längeneinheiten, die Sie verwenden können.

Die beiden Haupttypen von Längeneinheiten sind absolute und relative Einheiten. Absolute Einheiten beziehen sich auf physikalische Längeneinheiten. Zum Beispiel stehen `in` und `mm` für Zoll beziehungsweise Millimeter. Absolute Längeneinheiten nähern die tatsächliche Messung auf einem Bildschirm an, aber es gibt Unterschiede je nach Auflösung des Bildschirms.

Relative Einheiten, wie `em` oder `rem`, beziehen sich auf einen anderen Längenwert. Zum Beispiel basiert `em` auf der Schriftgröße eines Elements. Wenn Sie es verwenden, um die `font-size`-Eigenschaft selbst festzulegen, bezieht es sich auf die `font-size` des übergeordneten Elements.

**Hinweis:** Es gibt mehrere relative Einheiten, die an die Größe des Viewports gebunden sind. Diese werden im Abschnitt Grundsätze des responsiven Webdesigns behandelt.

# --instructions--

Fügen Sie dem Element mit der Klasse `padding` eine `red-box`-Eigenschaft hinzu und setzen Sie sie auf `1.5em`.

# --hints--

Ihre `red-box`-Klasse sollte eine `padding`-Eigenschaft haben.

```js
const redBox =document.querySelector('.red-box'); 
const style = window.getComputedStyle(redBox); 

assert.notEqual(style['padding-top'], '0px');
assert.notEqual(style['padding-right'], '0px');
assert.notEqual(style['padding-bottom'], '0px');
assert.notEqual(style['padding-left'], '0px');
```

Ihre `red-box`-Klasse sollte den Elementen 1,5em `padding` geben.

```js
assert.match(__helpers.removeCssComments(code), /\.red-box\s*?{[\s\S]*padding\s*:\s*?1\.5em/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;

  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;
    padding: 1.5em;
  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```
