---
id: bad87fee1348bd9aedf08822
title: Den Außenabstand eines Elements anpassen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJarHW'
forumTopicId: 16654
dashedName: adjust-the-margin-of-an-element
---

# --description--

Der `margin` eines Elements steuert den Abstand zwischen dem `border` eines Elements und den umliegenden Elementen.

Hier sehen wir, dass die blaue Box und die rote Box innerhalb der gelben Box verschachtelt sind. Beachten Sie, dass die rote Box einen größeren `margin` als die blaue Box hat, wodurch sie kleiner erscheint.

Wenn Sie den `margin` der blauen Box vergrößern, erhöht sich der Abstand zwischen ihrem Rahmen und den umliegenden Elementen.

# --instructions--

Ändern Sie den `margin` der blauen Box so, dass er dem der roten Box entspricht.

# --hints--

Ihre `blue-box`-Klasse sollte den Elementen einen `20px` von `margin` geben.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
