---
id: bad87fee1348bd9aecf08806
title: "Eine CSS-Klasse verwenden, um ein Element zu gestalten"
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MvDtV'
forumTopicId: 18337
dashedName: use-a-css-class-to-style-an-element
---

# --description--

Klassen sind wiederverwendbare Stile, die HTML-Elementen hinzugefügt werden können.

Hier ist ein Beispiel für eine CSS-Klassendeklaration:

```html
<style>
  .blue-text {
    color: blue;
  }
</style>
```

Sie sehen, dass wir eine CSS-Klasse namens `blue-text` innerhalb des `<style>`-Tags erstellt haben. Sie können eine Klasse einem HTML-Element so zuweisen: `<h2 class="blue-text">CatPhotoApp</h2>`. Beachten Sie, dass in Ihrem CSS-`style`-Element Klassennamen mit einem Punkt beginnen. Im class-Attribut Ihrer HTML-Elemente enthält der Klassenname keinen Punkt.

# --instructions--

Ändern Sie innerhalb Ihres `style`-Elements den `h2`-Selektor in `.red-text` und aktualisieren Sie den Farbwert von `blue` auf `red`.

Geben Sie Ihrem `h2`-Element das `class`-Attribut mit dem Wert `red-text`.

# --hints--

Ihr `h2`-Element sollte rot sein.

```js
const h2Element = document.querySelector('h2'); 
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Ihr `h2`-Element sollte die Klasse `red-text` haben.

```js
assert.isTrue(document.querySelector('h2').classList.contains('red-text'));
```

Ihr Stylesheet sollte eine `red-text`-Klasse deklarieren und deren Farbe auf `red` setzen.

```js
assert.match(__helpers.removeCssComments(code), /\.red-text\s*\{\s*color\s*:\s*red;?\s*\}/g);
```

Sie sollten keine inline-Stil-Deklarationen wie `style="color: red"` in Ihrem `h2`-Element verwenden.

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

# --seed--

## --seed-contents--

```html
<style>
  h2 {
    color: blue;
  }
</style>

<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  .red-text {
    color: red;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
