---
id: bad87fee1348bd9aedf08808
title: "Angeben, wie Schriftarten herabgestuft werden sollen"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cpVKBfQ'
forumTopicId: 18304
dashedName: specify-how-fonts-should-degrade
---

# --description--

Es gibt mehrere Standardschriftarten, die in allen Browsern verfügbar sind. Diese generischen Schriftfamilien umfassen `monospace`, `serif` und `sans-serif`.

Wenn eine Schriftart nicht verfügbar ist, können Sie dem Browser mitteilen, auf eine andere Schriftart „herabzustufen“.

Zum Beispiel, wenn ein Element die Schriftart `Helvetica` verwenden soll, aber auf die Schriftart `sans-serif` herabgestuft werden soll, wenn `Helvetica` nicht verfügbar ist, geben Sie dies wie folgt an:

```css
p {
  font-family: Helvetica, sans-serif;
}
```

Generische Schriftfamiliennamen sind nicht case-sensitiv. Außerdem benötigen sie keine Anführungszeichen, da sie CSS-Schlüsselwörter sind.

# --instructions--

Beginnen Sie damit, die Schriftart `monospace` auf das `h2`-Element anzuwenden, sodass es jetzt zwei Schriftarten hat – `Lobster` und `monospace`.

Im letzten Challenge haben Sie die Schriftart `Lobster` mit dem `link`-Tag importiert. Kommentieren Sie nun diesen Import der Schriftart `Lobster` (mithilfe der HTML-Kommentare, die Sie zuvor gelernt haben) von Google Fonts aus, sodass sie nicht mehr verfügbar ist. Beobachten Sie, wie Ihr `h2`-Element auf die Schriftart `monospace` herabgestuft wird.

**Hinweis:** Wenn Sie die Schriftart `Lobster` auf Ihrem Computer installiert haben, sehen Sie die Herabstufung nicht, da Ihr Browser die Schriftart finden kann.

# --hints--

Ihr h2-Element sollte die Schriftart `Lobster` verwenden.

```js
const h2Element = document.querySelector('h2');
const fontFamily = window.getComputedStyle(h2Element)['font-family']; 
assert.match(fontFamily, /^"?lobster/i);
```

Ihr h2-Element sollte auf die Schriftart `monospace` herabgestuft werden, wenn `Lobster` nicht verfügbar ist.

```js
assert.match(__helpers.removeCssComments(code), /\s*h2\s*\{\s*font-family\s*\:\s*(\'|"|)Lobster\1\s*,\s*monospace\s*;?\s*\}/gi);
```

Sie sollten Ihren Aufruf an Google für die Schriftart `Lobster` auskommentieren, indem Sie `<!--` davor setzen.

```js
assert.match(code, /<!--[^fc]/gi);
```

Sie sollten Ihren Kommentar schließen, indem Sie `-->` hinzufügen.

```js
assert.match(code, /[^fc]-->/gi);
```

# --seed--

## --seed-contents--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<!--<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">-->
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster, monospace;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
