---
id: bad87fee1348bd9aedf08807
title: Google-Schriftart importieren
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9MRsJ'
forumTopicId: 18200
dashedName: import-a-google-font
---

# --description--

Zusätzlich zur Angabe gängiger Schriftarten, die auf den meisten Betriebssystemen vorhanden sind, können wir auch nicht standardmäßige, benutzerdefinierte Web-Schriftarten für unsere Webseite festlegen. Im Internet gibt es viele Quellen für Web-Schriftarten. Für dieses Beispiel konzentrieren wir uns auf die Google Fonts-Bibliothek.

Google Fonts ist eine kostenlose Bibliothek von Web-Schriftarten, die Sie in Ihrem CSS verwenden können, indem Sie die URL der Schriftart referenzieren.

Also importieren und verwenden wir eine Google-Schriftart (beachten Sie, dass Sie diese Aufgabe überspringen müssen, wenn Google in Ihrem Land blockiert ist).

Um eine Google-Schriftart zu importieren, können Sie die URL der Schriftart aus der Google Fonts-Bibliothek kopieren und dann in Ihr HTML einfügen. Für diese Aufgabe importieren wir die `Lobster`-Schriftart. Kopieren Sie dazu den folgenden Code-Ausschnitt und fügen Sie ihn oben in Ihren Code-Editor ein (vor dem öffnenden `style`-Element):

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
```

Jetzt können Sie die `Lobster`-Schriftart in Ihrem CSS verwenden, indem Sie `Lobster` als FAMILY_NAME verwenden, wie im folgenden Beispiel:

```css
font-family: FAMILY_NAME, GENERIC_NAME;
```

Der GENERIC_NAME ist optional und dient als Fallback-Schriftart, falls die andere angegebene Schriftart nicht verfügbar ist. Dies wird in der nächsten Aufgabe behandelt.

Schriftfamiliennamen sind case-sensitiv und müssen in Anführungszeichen gesetzt werden, wenn der Name ein Leerzeichen enthält. Zum Beispiel benötigen Sie Anführungszeichen, um die `"Open Sans"`-Schriftart zu verwenden, aber nicht für die `Lobster`-Schriftart.

# --instructions--

Importieren Sie die `Lobster`-Schriftart in Ihre Webseite. Verwenden Sie dann einen Elementselektor, um `Lobster` als `font-family` für Ihr `h2`-Element festzulegen.

# --hints--

Sie sollten die `Lobster`-Schriftart importieren.

```js
assert.exists(document.querySelector('link[href*="googleapis" i]'));
```

Ihr `h2`-Element sollte die Schriftart `Lobster` verwenden.

```js
const h2 = document.querySelector('h2'); 
const fontFamily = window.getComputedStyle(h2)['font-family']; 
assert.match(fontFamily, /lobster/i);
```

Sie sollten nur einen `h2`-Elementselektor verwenden, um die Schriftart zu ändern.

```js
assert.match(__helpers.removeHtmlComments(code), /\s*[^\.]h2\s*\{\s*font-family\s*:\s*('|"|)Lobster\1\s*(,\s*('|"|)[a-z -]+\3\s*)?(;\s*\}|\})/gi);
```

Ihr `p`-Element sollte weiterhin die Schriftart `monospace` verwenden.

```js
const paragraphElement = document.querySelector('p');
const fontFamily = window.getComputedStyle(paragraphElement)['font-family']; 
assert.match(fontFamily, /monospace/i);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }

  h2 {
    font-family: Lobster;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
