---
id: bad87fee1348bd9aede08718
title: "RGB-Werte verwenden, um Elemente zu färben"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

Eine weitere Möglichkeit, Farben in CSS darzustellen, ist die Verwendung von `RGB`-Werten.

Der `RGB`-Wert für Schwarz sieht so aus:

```css
rgb(0, 0, 0)
```

Der `RGB`-Wert für Weiß sieht so aus:

```css
rgb(255, 255, 255)
```

Anstatt sechs hexadezimale Ziffern wie beim Hex-Code zu verwenden, geben Sie bei `RGB` die Helligkeit jeder Farbe mit einer Zahl zwischen 0 und 255 an.

Wenn Sie nachrechnen, entsprechen die zwei Ziffern für eine Farbe 16 mal 16, was insgesamt 256 Werte ergibt. `RGB`, das bei null zu zählen beginnt, hat also genau dieselbe Anzahl möglicher Werte wie der Hex-Code.

Hier ist ein Beispiel, wie Sie den `body`-Hintergrund mit seinem RGB-Code in Orange ändern.

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

Ersetzen Sie den Hex-Code im Hintergrundfarbwert unseres `body`-Elements durch den RGB-Wert für Schwarz: `rgb(0, 0, 0)`

# --hints--

Ihr `body`-Element sollte einen schwarzen Hintergrund haben.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Sie sollten `rgb` verwenden, um Ihrem `body`-Element einen schwarzen Hintergrund zu geben.

```js
assert.match(__helpers.removeCssComments(code), /rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
