---
id: bad87fee1348bd9aedf08736
title: Das HTML-Body-Element gestalten
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB77PHW'
forumTopicId: 18313
dashedName: style-the-html-body-element
---

# --description--

Beginnen wir jetzt von vorne und sprechen über CSS-Vererbung.

Jede HTML-Seite hat ein `body`-Element.

# --instructions--

Wir können beweisen, dass das `body`-Element hier existiert, indem wir ihm eine `background-color` in Schwarz geben.

Das erreichen wir, indem wir Folgendes zu unserem `style`-Element hinzufügen:

```css
body {
  background-color: black;
}
```

# --hints--

Ihr `body`-Element sollte die `background-color` Schwarz haben.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color'];

assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Ihre CSS-Regel sollte korrekt formatiert sein mit sowohl öffnenden als auch schließenden geschweiften Klammern.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

Ihre CSS-Regel sollte mit einem Semikolon enden.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

# --seed--

## --seed-contents--

```html
<style>

</style>
```

# --solutions--

```html
<style>
body {
  background-color: black;
}
</style>
```
