---
id: bad87fee1348bd9afdf08726
title: Außenabstand eines Elements mit Uhrzeigernotation angeben
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpybAd'
forumTopicId: 18345
dashedName: use-clockwise-notation-to-specify-the-margin-of-an-element
---

# --description--

Versuchen wir es noch einmal, diesmal aber mit `margin`.

Anstatt die `margin-top`, `margin-right`, `margin-bottom` und `margin-left` eines Elements einzeln anzugeben, können Sie sie alle in einer Zeile angeben, so:

```css
margin: 10px 20px 10px 20px;
```

Diese vier Werte funktionieren wie eine Uhr: oben, rechts, unten, links, und erzeugen genau dasselbe Ergebnis wie die Verwendung der seitenbezogenen Außenabstand-Anweisungen.

# --instructions--

Verwenden Sie die Uhrzeigernotation, um dem Element mit der `blue-box`-Klasse einen Außenabstand von `40px` oben und links zu geben, aber nur `20px` unten und rechts.

# --hints--

Ihre `blue-box`-Klasse sollte den oberen Außenabstand der Elemente auf `40px` von `margin` setzen.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '40px');
```

Ihre `blue-box`-Klasse sollte den rechten Außenabstand der Elemente auf `20px` von `margin` setzen.

```js
const blueBox = document.querySelector('.blue-box');
const marginRight = window.getComputedStyle(blueBox)['margin-right'];
assert.strictEqual(marginRight, '20px');
```

Ihre `blue-box`-Klasse sollte den unteren Außenabstand der Elemente auf `20px` von `margin` setzen.

```js
const blueBox = document.querySelector('.blue-box');
const marginBottom = window.getComputedStyle(blueBox)['margin-bottom'];
assert.strictEqual(marginBottom, '20px');
```

Ihre `blue-box`-Klasse sollte den linken Außenabstand der Elemente auf `40px` von `margin` setzen.

```js
const blueBox = document.querySelector('.blue-box');
const marginLeft = window.getComputedStyle(blueBox)['margin-left'];
assert.strictEqual(marginLeft, '40px');
```

Sie sollten die Uhrzeigernotation verwenden, um den Außenabstand der `blue-box`-Klasse festzulegen.

```js
const cssCode = __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(cssCode, /\.blue-box\s*{[\s\S]*margin[\s]*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    margin: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
