---
id: bad87fee1348bd9aedf04756
title: Überschreiben von Styles in nachfolgendem CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDQug'
forumTopicId: 18253
dashedName: override-styles-in-subsequent-css
---

# --description--

Unsere `pink-text`-Klasse hat die CSS-Deklaration unseres `body`-Elements überschrieben!

Wir haben gerade bewiesen, dass unsere Klassen die CSS des `body`-Elements überschreiben. Die nächste logische Frage lautet also: Was können wir tun, um unsere `pink-text`-Klasse zu überschreiben?

# --instructions--

Erstellen Sie eine zusätzliche CSS-Klasse namens `blue-text`, die einem Element die Farbe Blau gibt. Stellen Sie sicher, dass sie unterhalb Ihrer `pink-text`-Klassendeklaration steht.

Wenden Sie die `blue-text`-Klasse zusätzlich zu Ihrer `h1`-Klasse auf Ihr `pink-text`-Element an, und sehen wir, welche gewinnt.

Mehrere Klassenattribute auf ein HTML-Element anzuwenden, erfolgt mit einem Leerzeichen dazwischen, so:

```html
class="class1 class2"
```

**Hinweis:** Die Reihenfolge, in der die Klassen im HTML-Element aufgelistet sind, spielt keine Rolle.

Wichtig ist jedoch die Reihenfolge der `class`-Deklarationen im `<style>`-Abschnitt. Die zweite Deklaration hat immer Vorrang vor der ersten. Da `.blue-text` als zweite deklariert wird, überschreibt sie die Attribute von `.pink-text`.

# --hints--

Ihr `h1`-Element sollte die Klasse `pink-text` haben.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ihr `h1`-Element sollte die Klasse `blue-text` haben.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Sowohl `blue-text` als auch `pink-text` sollten zum selben `h1`-Element gehören.

```js
assert.isTrue(document.querySelector('.pink-text').classList.contains('blue-text'));
```

Ihr `h1`-Element sollte blau sein.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }

  .blue-text {
    color: blue;
  }  
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```
