---
id: bad87fee1348bd9aedf08721
title: "Hex-Code verwenden, um Farben zu mischen"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cK89PhP'
forumTopicId: 18359
dashedName: use-hex-code-to-mix-colors
---

# --description--

Zur Wiederholung: Hex-Codes verwenden 6 hexadezimale Ziffern, um Farben darzustellen, jeweils zwei für die roten (R), grünen (G) und blauen (B) Komponenten.

Aus diesen drei reinen Farben (Rot, Grün und Blau) können wir die Mengen jeder einzelnen variieren, um über 16 Millionen weitere Farben zu erzeugen!

Zum Beispiel ist Orange reines Rot, gemischt mit etwas Grün und ohne Blau. Im Hex-Code wird das zu `#FFA500`.

Die Ziffer `0` ist die niedrigste Zahl im Hex-Code und steht für das vollständige Fehlen von Farbe.

Die Ziffer `F` ist die höchste Zahl im Hex-Code und steht für die maximale Helligkeit.

# --instructions--

Ersetzen Sie die Farbnamen in unserem `style`-Element durch ihre korrekten Hex-Codes.

<table><tbody><tr><th>Farbe</th><th>Hex-Code</th></tr><tr><td>Dodger Blue</td><td><code>#1E90FF</code></td></tr><tr><td>Grün</td><td><code>#00FF00</code></td></tr><tr><td>Orange</td><td><code>#FFA500</code></td></tr><tr><td>Rot</td><td><code>#FF0000</code></td></tr></tbody></table>

# --hints--

Ihr `h1`-Element mit dem Text `I am red!` sollte das `color`-Rot erhalten.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Der `hex code` für die Farbe Rot sollte anstelle des Wortes `red` verwendet werden.

```js
assert.match(code, /\.red-text\s*?{\s*?color\s*:\s*?(#FF0000|#F00)\s*?;?\s*?}/gi);
```

Ihr `h1`-Element mit dem Text `I am green!` sollte das `color`-Grün erhalten.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

Der `hex code` für die Farbe Grün sollte anstelle des Wortes `green` verwendet werden.

```js
assert.match(code, /\.green-text\s*?{\s*?color\s*:\s*?(#00FF00|#0F0)\s*?;?\s*?}/gi);
```

Ihr `h1`-Element mit dem Text `I am dodger blue!` sollte das `color`-Dodger Blue erhalten.

```js
const blueText = document.querySelector('.dodger-blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(30, 144, 255)');
```

Der `hex code` für die Farbe Dodger Blue sollte anstelle des Wortes `dodgerblue` verwendet werden.

```js
assert.match(code, /\.dodger-blue-text\s*?{\s*?color\s*:\s*?#1E90FF\s*?;?\s*?}/gi);
```

Ihr `h1`-Element mit dem Text `I am orange!` sollte das `color`-Orange erhalten.

```js
const orangeText = document.querySelector('.orange-text');
const color = window.getComputedStyle(orangeText)['color']; 

assert.strictEqual(color, 'rgb(255, 165, 0)');
```

Der `hex code` für die Farbe Orange sollte anstelle des Wortes `orange` verwendet werden.

```js
assert.match(code, /\.orange-text\s*?{\s*?color\s*:\s*?#FFA500\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: black;
  }
  .green-text {
    color: black;
  }
  .dodger-blue-text {
    color: black;
  }
  .orange-text {
    color: black;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #FF0000;
  }
  .green-text {
    color: #00FF00;
  }
  .dodger-blue-text {
    color: #1E90FF;
  }
  .orange-text {
    color: #FFA500;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```
