---
id: bad87fee1348bd9aedf08719
title: Abgekürzten Hex-Code verwenden
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpKAm'
forumTopicId: 18338
dashedName: use-abbreviated-hex-code
---

# --description--

Viele Menschen fühlen sich von den Möglichkeiten von mehr als 16 Millionen Farben überwältigt. Und es ist schwierig, sich den Hex-Code zu merken. Zum Glück können Sie ihn abkürzen.

Zum Beispiel kann der Hex-Code von Rot `#FF0000` auf `#F00` verkürzt werden. Diese abgekürzte Form gibt jeweils eine Ziffer für Rot, Grün und Blau an.

Dadurch reduziert sich die Gesamtzahl der möglichen Farben auf etwa 4.000. Aber Browser interpretieren `#FF0000` und `#F00` als genau dieselbe Farbe.

# --instructions--

Probieren Sie es aus und verwenden Sie die abgekürzten Hex-Codes, um die richtigen Elemente zu färben.

<table><tbody><tr><th>Farbe</th><th>Abgekürzter Hex-Code</th></tr><tr><td>Cyan</td><td><code>#0FF</code></td></tr><tr><td>Grün</td><td><code>#0F0</code></td></tr><tr><td>Rot</td><td><code>#F00</code></td></tr><tr><td>Fuchsia</td><td><code>#F0F</code></td></tr></tbody></table>

# --hints--

Ihr `h1`-Element mit dem Text `I am red!` sollte das `color`-Rot erhalten.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Der abgekürzte Hex-Code für die Farbe Rot sollte anstelle des Hex-Codes `#FF0000` verwendet werden.

```js
assert.match(code,/\.red-text\s*?{\s*?color\s*:\s*?#F00\s*?;?\s*?}/gi);
```

Ihr `h1`-Element mit dem Text `I am green!` sollte das `color`-Grün erhalten.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

Der abgekürzte Hex-Code für die Farbe Grün sollte anstelle des Hex-Codes `#00FF00` verwendet werden.

```js
assert.match(__helpers.removeCssComments(code), /\.green-text\s*?{\s*?color\s*:\s*?#0F0\s*?;?\s*?}/gi);
```

Ihr `h1`-Element mit dem Text `I am cyan!` sollte das `color`-Cyan erhalten.

```js
const cyanText = document.querySelector('.cyan-text');
const color = window.getComputedStyle(cyanText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 255)');
```

Der abgekürzte Hex-Code für die Farbe Cyan sollte anstelle des Hex-Codes `#00FFFF` verwendet werden.

```js
assert.match(__helpers.removeCssComments(code), /\.cyan-text\s*?{\s*?color\s*:\s*?#0FF\s*?;?\s*?}/gi);
```

Ihr `h1`-Element mit dem Text `I am fuchsia!` sollte das `color`-Fuchsia erhalten.

```js
const fuchsiaText = document.querySelector('.fuchsia-text');
const color = window.getComputedStyle(fuchsiaText)['color']; 
assert.strictEqual(color,'rgb(255, 0, 255)');
```

Der abgekürzte Hex-Code für die Farbe Fuchsia sollte anstelle des Hex-Codes `#FF00FF` verwendet werden.

```js
assert.match(__helpers.removeCssComments(code), /\.fuchsia-text\s*?{\s*?color\s*:\s*?#F0F\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .fuchsia-text {
    color: #000000;
  }
  .cyan-text {
    color: #000000;
  }
  .green-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #F00;
  }
  .fuchsia-text {
    color: #F0F;
  }
  .cyan-text {
    color: #0FF;
  }
  .green-text {
    color: #0F0;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```
