---
id: bad87fee1348bd9aedf07756
title: Alle anderen Styles mit Important überschreiben
challengeType: 0
videoUrl: 'https://scrimba.com/c/cm24rcp'
forumTopicId: 18249
dashedName: override-all-other-styles-by-using-important
---

# --description--

Juhu! Wir haben gerade bewiesen, dass inline Styles alle CSS-Deklarationen in Ihrem `style`-Element überschreiben.

Aber warten Sie. Es gibt noch eine letzte Möglichkeit, CSS zu überschreiben. Das ist die mächtigste Methode von allen. Aber bevor wir dazu kommen, sprechen wir darüber, warum Sie CSS überhaupt überschreiben wollen.

In vielen Situationen verwenden Sie CSS-Bibliotheken. Diese können versehentlich Ihr eigenes CSS überschreiben. Wenn Sie also absolut sicher sein müssen, dass ein Element ein bestimmtes CSS hat, können Sie `!important` verwenden.

Gehen wir ganz zurück zu unserer `pink-text`-Klassendeklaration. Erinnern Sie sich, dass unsere `pink-text`-Klasse von nachfolgenden Klassendeklarationen, ID-Deklarationen und inline Styles überschrieben wurde.

# --instructions--

Fügen Sie dem Farbdeklaration Ihres pinkfarbenen Text-Elements das Schlüsselwort `!important` hinzu, um 100 % sicherzustellen, dass Ihr `h1`-Element pink ist.

Ein Beispiel, wie das geht, ist:

```css
color: red !important;
```

# --hints--

Ihr `h1`-Element sollte die Klasse `pink-text` haben.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Ihr `h1`-Element sollte die Klasse `blue-text` haben.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Ihr `h1`-Element sollte die `id` von `orange-text` haben.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

Ihr `h1`-Element sollte den inline Style `color: white` haben.

```js
const commentessCode = __helpers.removeHtmlComments(code);
assert.match(commentessCode, /<h1.*style/gi);
assert.match(commentessCode, /<h1.*style.*color\s*?:/gi);
```

Ihre `pink-text`-Klassendeklaration sollte das Schlüsselwort `!important` enthalten, um alle anderen Deklarationen zu überschreiben.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*?\{[\s\S]*?color:.*pink.*!important\s*;?[^\.]*\}/g);
```

Ihr `h1`-Element sollte pink sein.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink !important;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
