---
id: bad87fee1348bd9aedf08823
title: Einen negativen Außenabstand zu einem Element hinzufügen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpyGs3'
forumTopicId: 16166
dashedName: add-a-negative-margin-to-an-element
---

# --description--

Der `margin` eines Elements steuert den Abstand zwischen dem `border` eines Elements und den umliegenden Elementen.

Wenn Sie den `margin` eines Elements auf einen negativen Wert setzen, wird das Element größer.

# --instructions--

Versuchen Sie, den `margin` auf einen negativen Wert wie bei der roten Box zu setzen.

Ändern Sie den `margin` der blauen Box in `-15px`, damit sie die gesamte horizontale Breite der gelben Box um sie herum ausfüllt.

# --hints--

Ihre `blue-box`-Klasse sollte den Elementen `-15px` von `margin` geben.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)["margin-top"]; 
assert.strictEqual(marginTop, "-15px"); 
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: -15px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }
</style>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: -15px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
    margin-top: -15px;
  }
</style>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
