---
id: bad87fee1348bd9aedf08826
title: Innenabstand eines Elements mit der Uhrzeigernotation angeben
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB7mBS9'
forumTopicId: 18346
dashedName: use-clockwise-notation-to-specify-the-padding-of-an-element
---

# --description--

Anstatt die `padding-top`, `padding-right`, `padding-bottom` und `padding-left` Eigenschaften eines Elements einzeln anzugeben, können Sie sie alle in einer Zeile angeben, so:

```css
padding: 10px 20px 10px 20px;
```

Diese vier Werte funktionieren wie eine Uhr: oben, rechts, unten, links, und ergeben genau dasselbe Ergebnis wie die Verwendung der seitenbezogenen Innenabstand-Anweisungen.

# --instructions--

Verwenden Sie die Uhrzeigernotation, um der `.blue-box` Klasse einen `padding` von `40px` an der oberen und linken Seite zu geben, aber nur `20px` an der unteren und rechten Seite.

# --hints--

Ihre `blue-box` Klasse sollte dem oberen Bereich der Elemente einen `40px` von `padding` geben.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '40px');
```

Ihre `blue-box` Klasse sollte dem rechten Bereich der Elemente einen `20px` von `padding` geben.

```js
const blueBox = document.querySelector('.blue-box');
const paddingRight = window.getComputedStyle(blueBox)['padding-right'];
assert.strictEqual(paddingRight, '20px');
```

Ihre `blue-box` Klasse sollte dem unteren Bereich der Elemente einen `20px` von `padding` geben.

```js
const blueBox = document.querySelector('.blue-box');
const paddingBottom = window.getComputedStyle(blueBox)['padding-bottom'];
assert.strictEqual(paddingBottom, '20px');
```

Ihre `blue-box` Klasse sollte dem linken Bereich der Elemente einen `40px` von `padding` geben.

```js
const blueBox = document.querySelector('.blue-box');
const paddingLeft = window.getComputedStyle(blueBox)['padding-left'];
assert.strictEqual(paddingLeft, '40px');
```

Sie sollten die Uhrzeigernotation verwenden, um den Innenabstand der `blue-box` Klasse festzulegen.

```js
const css =  __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(css, /\.blue-box\s*{[\s\S]*padding\s*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
