---
id: 5895f70cf9fc0f352b528e65
title: Налаштування Passport
challengeType: 2
forumTopicId: 301565
dashedName: set-up-passport
---

# --description--

Настав час налаштувати *Passport*, щоб ви нарешті могли дозволити користувачам реєструватися або входити в акаунт. Окрім Passport, ви будете використовувати Express-session для обробки сесій. Express-session має багато розширених функцій, які можна застосовувати, але наразі ви використовуватимете лише базові. Це підпрограмне забезпечення зберігає ідентифікатор сесії як cookie у клієнта і дозволяє нам отримувати доступ до даних сесії за цим ідентифікатором на сервері. Таким чином, ви тримаєте особисту інформацію акаунта поза cookie, яке клієнт використовує, щоб повідомити серверу, що клієнти автентифіковані, і зберігаєте *ключ* для доступу до даних, що зберігаються на сервері.

`passport@~0.4.1` і `express-session@~1.17.1` вже встановлені і обидва вказані як залежності у вашому `package.json` файлі.

Вам потрібно налаштувати параметри сесії та ініціалізувати Passport. Спочатку створіть змінні `session` і `passport`, щоб вимагати `express-session` і `passport` відповідно.

Потім налаштуйте ваш Express-застосунок для використання сесії, визначивши такі опції:

```javascript
app.use(session({
  secret: process.env.SESSION_SECRET,
  resave: true,
  saveUninitialized: true,
  cookie: { secure: false }
}));
```

Обов’язково додайте `SESSION_SECRET` до вашого `.env` файлу і надайте йому випадкове значення. Це використовується для обчислення хешу, який шифрує ваше cookie!

Після цього скажіть вашому express-застосунку **використовувати** `passport.initialize()` і `passport.session()`.

Відправте вашу сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-passport-3" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

Passport і Express-session мають бути залежностями.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport',
    'Your project should list "passport" as a dependency'
  );
  assert.property(
    packJson.dependencies,
    'express-session',
    'Your project should list "express-session" as a dependency'
  );
```

Залежності мають бути правильно підключені.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport("|')/,
    'You should have required passport'
  );
  assert.match(
    data,
    /require.*("|')express-session("|')/,
    'You should have required express-session'
  );
```

Express-застосунок має використовувати нові залежності.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /passport\.initialize/, 'Your express app should use "passport.initialize()"');
  assert.match(data, /passport\.session/, 'Your express app should use "passport.session()"');
```

Сесія та секрет сесії мають бути налаштовані правильно.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /secret *:\s*process\.env(\.SESSION_SECRET|\[(?<q>"|')SESSION_SECRET\k<q>\])/,
    'Your express app should have express-session set up with your secret as process.env.SESSION_SECRET'
  );
```
