---
id: 5895f700f9fc0f352b528e63
title: Налаштування шаблонного рушія
challengeType: 2
forumTopicId: 301564
dashedName: set-up-a-template-engine
---

# --description--

Робота над цими завданнями передбачає написання коду одним із таких способів:

- Клонуйте <a href="https://github.com/freeCodeCamp/boilerplate-advancednode/" target="_blank" rel="noopener noreferrer nofollow">цей репозиторій GitHub</a> і виконайте завдання локально.
- Використайте будь-який конструктор сайтів на свій вибір для виконання проєкту. Обов’язково включіть усі файли з нашого репозиторію GitHub.

Шаблонний рушій дозволяє використовувати статичні шаблонні файли (наприклад, написані на *Pug*) у вашому застосунку. Під час виконання шаблонний рушій замінює змінні у шаблонному файлі на фактичні значення, які може надати ваш сервер. Потім він перетворює шаблон у статичний HTML-файл, який надсилається клієнту. Такий підхід полегшує проєктування HTML-сторінки і дозволяє відображати змінні на сторінці без необхідності робити API-виклик з клієнта.

`pug@~3.0.0` вже встановлено і він зазначений як залежність у вашому `package.json` файлі.

Express має знати, який шаблонний рушій ви використовуєте. Використайте метод `set`, щоб призначити `pug` значенням властивості `view engine`:

```javascript
app.set('view engine', 'pug');
```

Після цього додайте ще один метод `set`, який встановить властивість `views` вашого `app`, щоб вказати на директорію `./views/pug`. Це повідомить Express відтворювати всі подання відносно цієї директорії.

Нарешті, використайте `res.render()` у маршруті для вашої домашньої сторінки, передаючи `index` як перший аргумент. Це відтворить шаблон `pug`.

Якщо все пройшло успішно, домашня сторінка вашого застосунку більше не буде порожньою. Натомість вона відобразить повідомлення, що ви успішно відтворили шаблон Pug!

Відправте свою сторінку, коли вважатимете, що все зроблено правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-a-template-engine-1" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

Pug має бути залежністю.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'pug',
    'Your project should list "pug" as a dependency'
  );
```

View engine має бути Pug.

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.['view engine'], "pug");
```

Ви маєте встановити властивість `views` застосунку в `./views/pug`.

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.views, "./views/pug");
```

Використайте правильний метод ExpressJS, щоб відтворити сторінку index з відповіді.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /FCC Advanced Node and Express/gi,
        'You successfully rendered the Pug template!'
      );
```

Pug має працювати.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /pug-success-message/gi,
        'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
      );
```
