---
id: 5895f70ef9fc0f352b528e6b
title: як зібрати профіль
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

Тепер, коли ви можете переконатися, що користувач, який отримує доступ до `/profile`, є автентифікованим, ви можете використовувати інформацію, що міститься в `req.user`, на своїй сторінці.

Передайте об’єкт, який містить властивість `username` і значення `req.user.username`, як другий аргумент для методу `render` у вигляді профілю.

Потім перейдіть до свого вигляду `profile.pug` і додайте наступний рядок під існуючим елементом `h1` на тому ж рівні відступу:

```pug
h2.center#welcome Welcome, #{username}!
```

Це створює елемент `h2` з класом `center` і ідентифікатором `welcome`, що містить текст `Welcome, `, за яким слідує ім’я користувача.

Також у `profile.pug` додайте посилання, яке веде до маршруту `/logout`, де буде розміщена логіка для деавтентифікації користувача:

```pug
a(href='/logout') Logout
```

Відправте свою сторінку, коли вважатимете, що все зроблено правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

Ви повинні правильно додати змінну для відтворення Pug у `/profile`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );
```
