---
id: 5895f70bf9fc0f352b528e64
title: Використання можливостей шаблонного рушія
challengeType: 2
forumTopicId: 301567
dashedName: use-a-template-engines-powers
---

# --description--

Одна з найкращих функцій використання шаблонного рушія — це можливість передавати змінні з сервера у шаблонний файл перед його відтворенням у HTML.

У вашому Pug-файлі ви можете використовувати змінну, посилаючись на її ім’я як `#{variable_name}` рядково разом з іншим текстом у елементі або використовуючи знак «дорівнює» на елементі без пробілу, наприклад `p=variable_name`, що призначає значення змінної тексту елемента p.

Pug базується на використанні пробілів і табуляцій для показу вкладених елементів і зменшенні кількості коду, необхідного для створення гарного сайту.

Візьмемо для прикладу наступний код Pug: 
 
```pug
head
  script(type='text/javascript').
    if (foo) bar(1 + 5);
body
  if youAreUsingPug
      p You are amazing
    else
      p Get on it!
```

Вищенаведений код дає такий HTML:

```html
<head>
  <script type="text/javascript">
    if (foo) bar(1 + 5);
  </script>
</head>
<body>
  <p>You are amazing</p>
</body>
```

Ваш файл `index.pug`, включений у проєкт, використовує змінні `title` і `message`.

Передайте їх із сервера у Pug-файл, додавши об’єкт як другий аргумент до виклику `res.render` зі змінними та їхніми значеннями. Призначте `title` значення `Hello`, а `message` — значення `Please log in`.

Це має виглядати так:

```javascript
res.render('index', { title: 'Hello', message: 'Please log in' });
```

Тепер оновіть сторінку, і ви повинні побачити ці значення, відтворені у вашому вигляді у правильному місці, як це задано у вашому файлі `index.pug`!

Відправте сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете переглянути <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#use-a-template-engines-power-2" target="_blank" rel="noopener noreferrer nofollow">проєкт, виконаний до цього моменту</a>.

# --hints--

Pug має правильно відтворювати змінні.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /pug-variable("|')>Please log in/gi,
    'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
  );
```
