---
id: 5895f70cf9fc0f352b528e67
title: Реалізувати серіалізацію користувача Passport
challengeType: 2
forumTopicId: 301556
dashedName: implement-the-serialization-of-a-passport-user
---

# --description--

Ви не завантажуєте реальний об’єкт користувача, оскільки база даних не налаштована. Підключіться до бази даних один раз, коли запускаєте сервер, і підтримуйте постійне з’єднання протягом усього життєвого циклу застосунку. Для цього додайте рядок підключення вашої бази даних (наприклад: `mongodb+srv://<username>:<password>@cluster0-jvwxi.mongodb.net/?retryWrites=true&w=majority`) у змінну середовища `MONGO_URI`. Це використовується у файлі `connection.js`.

*Якщо у вас виникають проблеми з налаштуванням безкоштовної бази даних на MongoDB Atlas, перегляньте цей <a href="https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/" target="_blank" rel="noopener noreferrer nofollow">підручник</a>.*

Тепер ви хочете підключитися до бази даних, а потім почати слухати запити. Мета цього — не дозволяти запити до підключення бази даних або у разі помилки бази даних. Щоб це зробити, обгорніть вашу серіалізацію та маршрути застосунку у наступний код:

```javascript
myDB(async client => {
  const myDataBase = await client.db('database').collection('users');

  // Be sure to change the title
  app.route('/').get((req, res) => {
    // Change the response to render the Pug template
    res.render('index', {
      title: 'Connected to Database',
      message: 'Please login'
    });
  });

  // Serialization and deserialization here...

  // Be sure to add this...
}).catch(e => {
  app.route('/').get((req, res) => {
    res.render('index', { title: e, message: 'Unable to connect to database' });
  });
});
// app.listen out here...
```

Обов’язково розкоментуйте код `myDataBase` у `deserializeUser` і відредагуйте ваш `done(null, null)`, щоб включити `doc`.

Відправте вашу сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implement-the-serialization-of-a-passport-user-5" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

Підключення до бази даних має бути присутнє.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Connected to Database/gi,
    'You successfully connected to the database!'
  );
```

Десеріалізація тепер має правильно використовувати базу даних, і `done(null, null)` має викликатися з `doc`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /null,\s*doc/gi,
    'The callback in deserializeUser of (null, null) should be altered to (null, doc)'
  );
```
