---
id: 5895f70df9fc0f352b528e6a
title: створити нове підпрограмне забезпечення
challengeType: 2
forumTopicId: 301551
dashedName: create-new-middleware
---

# --description--

Як є, будь-який користувач може просто перейти до `/profile`, незалежно від того, автентифікований він чи ні, ввівши URL. Ви хочете запобігти цьому, перевіряючи, чи користувач автентифікований, перш ніж відтворити сторінку профілю. Це ідеальний приклад, коли варто створити підпрограмне забезпечення.

Завдання тут — створити функцію підпрограмного забезпечення `ensureAuthenticated(req, res, next)`, яка перевірятиме, чи користувач автентифікований, викликаючи метод `isAuthenticated` Passport на `request`, який перевіряє, чи визначено `req.user`. Якщо так, тоді слід викликати `next()`. Інакше можна просто відповісти на запит перенаправленням на домашню сторінку для входу.

Приклад реалізації цього підпрограмного забезпечення:

```javascript
function ensureAuthenticated(req, res, next) {
  if (req.isAuthenticated()) {
    return next();
  }
  res.redirect('/');
};
```

Створіть вказану функцію підпрограмного забезпечення, а потім передайте `ensureAuthenticated` як підпрограмне забезпечення для запитів до сторінки профілю перед аргументом GET-запиту:

```javascript
app
 .route('/profile')
 .get(ensureAuthenticated, (req,res) => {
    res.render('profile');
 });
```

Відправте свою сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#create-new-middleware-8" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

Підпрограмне забезпечення `ensureAuthenticated` слід реалізувати та приєднати до маршруту `/profile`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /ensureAuthenticated[^]*req.isAuthenticated/,
    'Your ensureAuthenticated middleware should be defined and utilize the req.isAuthenticated function'
  );
  assert.match(
    data,
    /profile[^]*get[^]*ensureAuthenticated/,
    'Your ensureAuthenticated middleware should be attached to the /profile route'
  );
```

Неавтентифікований GET-запит до `/profile` має правильно перенаправляти на `/`.

```js
  const url = new URL("/profile", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/,
    'An attempt to go to the profile at this point should redirect to the homepage since we are not logged in'
  );
```
