---
id: 5895f70cf9fc0f352b528e66
title: Серіалізація об’єкта користувача
challengeType: 2
forumTopicId: 301563
dashedName: serialization-of-a-user-object
---

# --description--

Серіалізація та десеріалізація — це важливі поняття в контексті автентифікації. Серіалізувати об’єкт означає перетворити його вміст у невеликий *ключ*, який потім можна десеріалізувати назад у початковий об’єкт. Саме це дозволяє нам знати, хто спілкувався з сервером, без необхідності надсилати дані автентифікації, як-от ім’я користувача та пароль, при кожному запиті нової сторінки.

Щоб налаштувати це правильно, потрібно мати функцію серіалізації та функцію десеріалізації. У Passport їх можна створити так:

```javascript
passport.serializeUser(cb);
passport.deserializeUser(cb);
```

Функція зворотного виклику, передана в `serializeUser`, викликається з двома аргументами: повним об’єктом користувача та функцією зворотного виклику, яку використовує passport.

Ця функція зворотного виклику очікує два аргументи: помилку, якщо вона є, і унікальний ключ для ідентифікації користувача, який має бути повернений у функції зворотного виклику. Ви використаєте `_id` користувача в об’єкті. Він гарантовано унікальний, оскільки генерується MongoDB.

Аналогічно, `deserializeUser` викликається з двома аргументами: унікальним ключем і функцією зворотного виклику.

Ця функція зворотного виклику очікує два аргументи: помилку, якщо вона є, і повний об’єкт користувача. Щоб отримати повний об’єкт користувача, зробіть запит пошуку за Mongo `_id`, як показано нижче:

```javascript
passport.serializeUser((user, done) => {
  done(null, user._id);
});

passport.deserializeUser((id, done) => {
  myDataBase.findOne({ _id: new ObjectID(id) }, (err, doc) => {
    done(null, null);
  });
});
```

Додайте дві наведені вище функції до свого сервера. Клас `ObjectID` походить із пакету `mongodb`. `mongodb@~3.6.0` вже додано як залежність. Оголосіть цей клас так:

```javascript
const { ObjectID } = require('mongodb');
```

`deserializeUser` видасть помилку, доки ви не налаштуєте підключення до бази даних. Тож наразі перетворіть на коментар виклик `myDatabase.findOne` і просто викличте `done(null, null)` у функції зворотного виклику `deserializeUser`.

Відправте свою сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#serialization-of-a-user-object-4" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

Ви повинні правильно серіалізувати об’єкт користувача.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.serializeUser/gi,
    'You should have created your passport.serializeUser function'
  );
  assert.match(
    data,
    /null,\s*user._id/gi,
    'There should be a callback in your serializeUser with (null, user._id)'
  );
```

Ви повинні правильно десеріалізувати об’єкт користувача.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.deserializeUser/gi,
    'You should have created your passport.deserializeUser function'
  );
  assert.match(
    data,
    /null,\s*null/gi,
    'There should be a callback in your deserializeUser with (null, null) for now'
  );
```

MongoDB має бути залежністю.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'mongodb',
    'Your project should list "mongodb" as a dependency'
  );
```

MongoDB має бути правильно імпортований, включно з ObjectId.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')mongodb\1/gi,
    'You should have required mongodb'
  );
  assert.match(
    data,
    /new ObjectID.*id/gi,
    'Even though the block is commented out, you should use new ObjectID(id) for when we add the database'
  );
```
