---
id: 58965611f9fc0f352b528e6c
title: Вихід користувача з системи
challengeType: 2
forumTopicId: 301560
dashedName: logging-a-user-out
---

# --description--

Створити логіку виходу з системи дуже просто. Маршрут має лише скасувати автентифікацію користувача та зробити перенаправлення на головну сторінку замість відтворення будь-якого вигляду.

У passport скасувати автентифікацію користувача так само просто, як викликати `req.logout()` перед перенаправленням. Додайте цей маршрут `/logout`, щоб це зробити:

```js
app.route('/logout')
  .get((req, res) => {
    req.logout();
    res.redirect('/');
});
```

Ви могли помітити, що ви не обробляєте відсутні сторінки (404). Загальний спосіб обробки цього у Node — за допомогою такого підпрограмного забезпечення. Додайте його після всіх інших маршрутів:

```js
app.use((req, res, next) => {
  res.status(404)
    .type('text')
    .send('Not Found');
});
```

Відправте свою сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#logging-a-user-out-10" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

`req.logout()` має бути викликаний у вашому маршруті `/logout`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /req.logout/gi,
    'You should be calling req.logout() in your /logout route'
  );
```

`/logout` має зробити перенаправлення на головну сторінку.

```js
  const url = new URL("/logout", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/gi,
    'When a user logs out they should be redirected to the homepage'
  );
```
