---
id: 589690e6f9fc0f352b528e6e
title: Очистіть свій проєкт за допомогою модулів
challengeType: 2
forumTopicId: 301549
dashedName: clean-up-your-project-with-modules
---

# --description--

Зараз у вас усе знаходиться у файлі `server.js`. Це може призвести до складного для керування коду, який важко розширювати. Створіть 2 нові файли: `routes.js` та `auth.js`.

Обидва повинні починатися з такого коду:

```js
module.exports = function (app, myDataBase) {

}
```

Тепер у верхній частині вашого серверного файлу підключіть ці файли так: `const routes = require('./routes.js');` Одразу після встановлення успішного з’єднання з базою даних створіть їхні екземпляри так: `routes(app, myDataBase)`

Нарешті, візьміть усі маршрути у вашому сервері та вставте їх у нові файли, а з серверного файлу видаліть. Також візьміть функцію `ensureAuthenticated`, оскільки вона була створена спеціально для маршрутизації. Тепер вам потрібно правильно додати залежності, які використовуються, наприклад `const passport = require('passport');`, у самому верху, над рядком експорту у файлі `routes.js`.

Продовжуйте додавати їх, доки не зникнуть усі помилки, і у вашому серверному файлі більше не буде маршрутизації (**окрім маршруту у блоці catch**)!

Зробіть те саме у файлі `auth.js` з усім, що стосується автентифікації, наприклад серіалізації та налаштування локальної стратегії, і видаліть це з серверного файлу. Обов’язково додайте залежності і викличте `auth(app, myDataBase)` у сервері в тому ж місці.

Відправте свою сторінку, коли вважатимете, що все зроблено правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#clean-up-your-project-with-modules-2" target="_blank" rel="noopener noreferrer nofollow">переглянути приклад завершеного проєкту</a>.

# --hints--

Модулі мають бути присутні.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\s*\(('|")\.\/routes(\.js)?\1\)/gi,
    'You should have required your new files'
  );
  assert.match(
    data,
    /client\s*\.db[^]*routes/gi,
    'Your new modules should be called after your connection to the database'
  );
```
