---
id: 589a69f5f9fc0f352b528e70
title: Реалізація соціальної автентифікації
challengeType: 2
forumTopicId: 301559
dashedName: implementation-of-social-authentication
---

# --description--

Основний шлях, який цей тип автентифікації пройде у вашому застосунку, такий:

1. Користувач натискає кнопку або посилання, що відправляє його на ваш маршрут для автентифікації за допомогою конкретної стратегії (наприклад, GitHub).
2. Ваш маршрут виконує `passport.authenticate('github')`, який перенаправляє користувача на GitHub.
3. Сторінка, на яку потрапляє користувач на GitHub, дозволяє йому увійти, якщо він ще не увійшов. Потім вона запитує дозвіл на доступ до його профілю з вашого застосунку.
4. Користувач повертається до вашого застосунку за конкретним URL функції зворотного виклику з його профілем, якщо доступ було схвалено.
5. Тепер він автентифікований, і ваш застосунок має перевірити, чи це повернений профіль, або зберегти його у вашій базі даних, якщо ні.

Стратегії з OAuth вимагають, щоб у вас був принаймні *Client ID* і *Client Secret* — це спосіб для сервісу перевірити, звідки надходить запит автентифікації і чи він дійсний. Їх отримують на сайті, з яким ви намагаєтеся реалізувати автентифікацію, наприклад, GitHub, і вони унікальні для вашого застосунку — **ЇХ НЕ МОЖНА ПОДІЛИТИСЯ** і ніколи не слід завантажувати у публічний репозиторій або писати безпосередньо у вашому коді. Зазвичай їх кладуть у ваш `.env` файл і посилаються на них так: `process.env.GITHUB_CLIENT_ID`. Для цього завдання ви будете використовувати стратегію GitHub.

Дотримуйтесь <a href="https://www.freecodecamp.org/news/how-to-set-up-a-github-oauth-application/" target="_blank" rel="noopener noreferrer nofollow">цих інструкцій</a>, щоб отримати *Client ID і Secret* від GitHub. Встановіть URL домашньої сторінки на вашу домашню сторінку (**не URL коду проєкту**), а URL функції зворотного виклику — на той самий URL домашньої сторінки з доданим `/auth/github/callback` в кінці. Збережіть client ID і client secret у файлі `.env` вашого проєкту як `GITHUB_CLIENT_ID` і `GITHUB_CLIENT_SECRET`.

У вашому файлі `routes.js` додайте `showSocialAuth: true` до маршруту домашньої сторінки, після `showRegistration: true`. Тепер створіть 2 маршрути, які приймають GET-запити: `/auth/github` і `/auth/github/callback`. Перший має лише викликати passport для автентифікації `'github'`. Другий має викликати passport для автентифікації `'github'` з перенаправленням на `/` у разі невдачі, а якщо успішно — перенаправляти на `/profile` (подібно до вашого останнього проєкту).

Приклад того, як має виглядати `/auth/github/callback`, схожий на те, як ви обробляли звичайний вхід:

```js
app.route('/login')
  .post(passport.authenticate('local', { failureRedirect: '/' }), (req,res) => {
    res.redirect('/profile');
  });
```

Відправте вашу сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-3" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт до цього моменту</a>.

# --hints--

Маршрут `/auth/github` має бути правильним.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
          data.replace(/\s/g, ''),
          /passport.authenticate.*?github/g,
          'Route auth/github should only call passport.authenticate with github'
        );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```

Маршрут `/auth/github/callback` має бути правильним.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
        data.replace(/\s/g, ''),
        /failureRedirect:("|')\/\1/g,
        'Route auth/github/callback should accept a get request and call passport.authenticate for github with a failure redirect to home'
      );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github/callback');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github/callback"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```
