---
id: 589a8eb3f9fc0f352b528e72
title: Реалізація автентифікації через соціальні мережі III
challengeType: 2
forumTopicId: 301558
dashedName: implementation-of-social-authentication-iii
---

# --description--

Фінальна частина стратегії — це обробка профілю, який повертає GitHub. Нам потрібно завантажити об’єкт користувача з бази даних, якщо він існує, або створити його, якщо ні, і заповнити поля з профілю, а потім повернути об’єкт користувача. GitHub надає унікальний *ідентифікатор* у кожному профілі, який ми можемо використати для пошуку та серіалізації користувача (вже реалізовано). Нижче наведено приклад реалізації, який ви можете використати у своєму проєкті — він розміщується у функції, що є другим аргументом для нової стратегії, безпосередньо під тим місцем, де зараз знаходиться `console.log(profile);`:

```js
myDataBase.findOneAndUpdate(
  { id: profile.id },
  {
    $setOnInsert: {
      id: profile.id,
      username: profile.username,
      name: profile.displayName || 'John Doe',
      photo: profile.photos[0].value || '',
      email: Array.isArray(profile.emails)
        ? profile.emails[0].value
        : 'No public email',
      created_on: new Date(),
      provider: profile.provider || ''
    },
    $set: {
      last_login: new Date()
    },
    $inc: {
      login_count: 1
    }
  },
  { upsert: true, new: true },
  (err, doc) => {
    return cb(null, doc.value);
  }
);
```

`findOneAndUpdate` дозволяє шукати об’єкт і оновлювати його. Якщо об’єкт не існує, він буде вставлений і доступний для функції зворотного виклику. У цьому прикладі ми завжди встановлюємо `last_login`, збільшуємо `login_count` на `1` і заповнюємо більшість полів лише при вставці нового об’єкта (нового користувача). Зверніть увагу на використання значень за замовчуванням. Іноді повернутий профіль не містить усієї інформації або користувач залишає її приватною. У такому випадку ви маєте це обробити, щоб уникнути помилки.

Тепер ви маєте змогу увійти у свій застосунок. Спробуйте!

Відправте свою сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-iii-5" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

Налаштування стратегії GitHub має бути завершене.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /GitHubStrategy[^]*myDataBase/gi,
    'Strategy should use now use the database to search for the user'
  );
  assert.match(
    data,
    /GitHubStrategy[^]*cb/gi,
    'Strategy should return the callback function "cb"'
  );
```
