---
id: 589a69f5f9fc0f352b528e71
title: Реалізація соціальної автентифікації II
challengeType: 2
forumTopicId: 301557
dashedName: implementation-of-social-authentication-ii
---

# --description--

Остання частина налаштування автентифікації через GitHub — це створення самої стратегії. `passport-github@~1.1.0` вже додано як залежність, тому підключіть її у вашому `auth.js` файлі як `GitHubStrategy` ось так: `const GitHubStrategy = require('passport-github').Strategy;`. Не забудьте також підключити і налаштувати `dotenv` для використання ваших змінних середовища.

Щоб налаштувати стратегію GitHub, потрібно повідомити Passport використовувати створений екземпляр `GitHubStrategy`, який приймає 2 аргументи: об’єкт (що містить `clientID`, `clientSecret` і `callbackURL`) та функцію, яку викликають, коли користувач успішно автентифікований. Ця функція визначає, чи є користувач новим, і які поля зберігати ініціально в об’єкті користувача в базі даних. Це загальна практика для багатьох стратегій, але деякі можуть вимагати більше інформації, як описано у README конкретної стратегії на GitHub. Наприклад, Google також вимагає *scope*, який визначає, яку інформацію ваш запит хоче отримати, і просить користувача підтвердити такий доступ.

Поточна стратегія, яку ви реалізуєте, автентифікує користувачів за допомогою облікового запису GitHub і токенів OAuth 2.0. Ідентифікатор клієнта та секрет, отримані під час створення застосунку, передаються як опції при створенні стратегії. Стратегія також вимагає функції зворотного виклику `verify`, яка отримує токен доступу та необов’язковий токен оновлення, а також `profile`, що містить автентифікований профіль користувача GitHub. Функція зворотного виклику `verify` має викликати `cb`, передаючи користувача для завершення автентифікації.

Ось як має виглядати ваша нова стратегія на цьому етапі:

```js
passport.use(new GitHubStrategy({
  clientID: process.env.GITHUB_CLIENT_ID,
  clientSecret: process.env.GITHUB_CLIENT_SECRET,
  callbackURL: /*INSERT CALLBACK URL ENTERED INTO GITHUB HERE*/
},
  function(accessToken, refreshToken, profile, cb) {
    console.log(profile);
    //Database logic here with callback containing your user object
  }
));
```

Ваша автентифікація поки що не буде успішною, і фактично вона викличе помилку без логіки бази даних і функції зворотного виклику, але якщо спробувати, вона має вивести ваш профіль GitHub у консоль!

Відправте свою сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-ii-4" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

`passport-github` залежність має бути додана.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport-github',
    'Your project should list "passport-github" as a dependency'
  );
```

`passport-github` має бути підключена.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport-github("|')/gi,
    'You should have required passport-github'
  );
```

Стратегія GitHub має бути налаштована правильно на цей момент.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport\.use.*new GitHubStrategy/gis,
    'Passport should use a new GitHubStrategy'
  );
  assert.match(
    data,
    /callbackURL:\s*("|').*("|')/gi,
    'You should have a callbackURL'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_SECRET|\[(?<q>"|')GITHUB_CLIENT_SECRET\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_SECRET'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_ID|\[(?<q>"|')GITHUB_CLIENT_ID\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_ID'
  );
```
