---
id: 589fc831f9fc0f352b528e76
title: Обробка відключення
challengeType: 2
forumTopicId: 301552
dashedName: handle-a-disconnect
---

# --description--

Ви могли помітити, що до цього моменту ви лише збільшували кількість користувачів. Обробка відключення користувача так само проста, як і обробка початкового підключення, тільки вам потрібно слухати цю подію на кожному сокеті, а не на всьому сервері.

Для цього додайте ще одного слухача всередині існуючого `'connect'` слухача, який слухатиме `'disconnect'` на сокеті без передачі даних. Ви можете перевірити цю функціональність, просто вивівши в консоль повідомлення про те, що користувач відключився.

```js
socket.on('disconnect', () => {
  /*anything you want to do on disconnect*/
});
```

Щоб клієнти безперервно отримували оновлену кількість поточних користувачів, слід зменшити `currentUsers` на 1, коли відбувається відключення, а потім надіслати подію `'user count'` з оновленою кількістю.

**Примітка:** Так само, як і `'disconnect'`, всі інші події, які сокет може надсилати серверу, слід обробляти всередині слухача підключення, де визначено 'socket'.

Відправте свою сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#handle-a-disconnect-8" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

Сервер має обробляти подію disconnect від сокета.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /socket.on.*('|")disconnect('|")/s, '');
```

Ваш клієнт має слухати подію `'user count'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
