---
id: 589fc832f9fc0f352b528e78
title: Оголошення нових користувачів
challengeType: 2
forumTopicId: 301546
dashedName: announce-new-users
---

# --description--

Багато чатів можуть оголошувати, коли користувач підключається або відключається, і показувати це всім підключеним користувачам у чаті. Оскільки ви вже відправляєте подію при підключенні та відключенні, вам потрібно лише змінити цю подію, щоб підтримати таку функцію. Найлогічніший спосіб зробити це — надіслати 3 дані з подією: ім’я користувача, який підключився/відключився, поточну кількість користувачів і чи цей користувач підключився чи відключився.

Змініть назву події на `'user'` і передайте об’єкт, що містить поля `username`, `currentUsers` та `connected` (яке буде `true` у разі підключення або `false` у разі відключення користувача). Обов’язково змініть обидві події `'user count'` і встановіть для події відключення відправку `false` у полі `connected` замість `true`, як це робиться у події при підключенні.

```js
io.emit('user', {
  username: socket.request.user.username,
  currentUsers,
  connected: true
});
```

Тепер ваш клієнт матиме всю необхідну інформацію, щоб правильно відобразити поточну кількість користувачів і оголосити, коли користувач підключається або відключається! Щоб обробити цю подію на стороні клієнта, слід слухати `'user'`, потім оновити поточну кількість користувачів, використовуючи jQuery для зміни тексту `#num-users` на `'{NUMBER} users online'`, а також додати `<li>` до невпорядкованого списку з ідентифікатором `messages` з `'{NAME} has {joined/left} the chat.'`.

Приклад реалізації може виглядати так:

```js
socket.on('user', data => {
  $('#num-users').text(data.currentUsers + ' users online');
  let message =
    data.username +
    (data.connected ? ' has joined the chat.' : ' has left the chat.');
  $('#messages').append($('<li>').html('<b>' + message + '</b>'));
});
```

Відправте свою сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете переглянути <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135/3#announce-new-users-10" target="_blank" rel="noopener noreferrer nofollow">проєкт, виконаний до цього моменту</a>.

# --hints--

Подія `'user'` має відправлятися з `username`, `currentUsers` та `connected`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  // Regex is lenient to match both `username` and `name` as the key on purpose.
  assert.match(
    data,
    /io.emit.*('|")user\1.*name.*currentUsers.*connected/s,
    'You should have an event emitted named user sending name, currentUsers, and connected'
  );
```

Клієнт має правильно обробляти та відображати нові дані з події `'user'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*num-users/s,
    'You should change the text of "#num-users" within on your client within the "user" event listener to show the current users connected'
  );
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*messages.*li/s,
    'You should append a list item to "#messages" on your client within the "user" event listener to announce a user came or went'
  );
```
