---
id: 589fc831f9fc0f352b528e75
title: Спілкування за допомогою емісії
challengeType: 2
forumTopicId: 301550
dashedName: communicate-by-emitting
---

# --description--

<dfn>Emit</dfn> — це найпоширеніший спосіб спілкування, який ви будете використовувати. Коли ви емінуєте щось із сервера до 'io', ви надсилаєте назву події та дані всім підключеним сокетам. Хорошим прикладом цієї концепції буде емісія поточного лічильника підключених користувачів щоразу, коли підключається новий користувач!

Почніть із додавання змінної для відстеження користувачів, безпосередньо перед тим місцем, де ви зараз слухаєте підключення.

```js
let currentUsers = 0;
```

Тепер, коли хтось підключається, слід збільшити лічильник перед емісією цього лічильника. Тож додайте збільшувач всередину слухача підключення.

```js
++currentUsers;
```

Нарешті, після збільшення лічильника, слід емінувати подію (ще всередині слухача підключення). Подія має називатися 'user count', а дані — це просто `currentUsers`.

```js
io.emit('user count', currentUsers);
```

Тепер ви можете реалізувати спосіб, щоб ваш клієнт слухав цю подію! Подібно до того, як сервер слухає підключення, ви використаєте ключове слово `on`.

```js
socket.on('user count', function(data) {
  console.log(data);
});
```

Тепер спробуйте завантажити свій застосунок, автентифікуватися, і ви повинні побачити у консолі клієнта '1', що представляє поточну кількість користувачів! Спробуйте завантажити більше клієнтів і автентифікуватися, щоб побачити, як число зростає.

Відправте свою сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#communicate-by-emitting-7" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

`currentUsers` має бути визначений.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /currentUsers/s,
    'You should have variable currentUsers defined'
  );
```

Сервер має емінувати поточну кількість користувачів при кожному новому підключенні.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.emit.*('|")user count('|").*currentUsers/s,
    'You should emit "user count" with data currentUsers'
  );
```

Ваш клієнт має слухати подію `'user count'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
