---
id: 589fc831f9fc0f352b528e77
title: Автентифікація з Socket.IO
challengeType: 2
forumTopicId: 301548
dashedName: authentication-with-socket-io
---

# --description--

Наразі ви не можете визначити, хто підключений до вашого веб-сокета. Хоча `req.user` містить об’єкт користувача, це лише тоді, коли ваш користувач взаємодіє з вебсервером, а з веб-сокетами у вас немає `req` (запиту) і, відповідно, немає даних користувача. Один зі способів розв’язати проблему визначення, хто підключений до вашого веб-сокета — це розпарсити та декодувати cookie, яке містить сесію passport, а потім десеріалізувати його, щоб отримати об’єкт користувача. На щастя, існує пакет на npm саме для цього, який перетворює колись складне завдання на просте!

`passport.socketio@~3.7.0`, `connect-mongo@~3.2.0` і `cookie-parser@~1.4.5` вже додані як залежності. Імпортуйте їх як `passportSocketIo`, `MongoStore` і `cookieParser` відповідно. Також нам потрібно ініціалізувати нове сховище в пам’яті з `express-session`, який ми раніше імпортували. Це має виглядати так:

```js
const MongoStore = require('connect-mongo')(session);
const URI = process.env.MONGO_URI;
const store = new MongoStore({ url: URI });
```

Тепер нам залишилось лише сказати Socket.IO використовувати це сховище і встановити опції. Переконайтеся, що це додано перед існуючим кодом сокета і не всередині існуючого слухача підключення. Для вашого сервера це має виглядати так:

```js
io.use(
  passportSocketIo.authorize({
    cookieParser: cookieParser,
    key: 'express.sid',
    secret: process.env.SESSION_SECRET,
    store: store,
    success: onAuthorizeSuccess,
    fail: onAuthorizeFail
  })
);
```

Зверніть увагу, що налаштування автентифікації Passport для Socket.IO дуже схоже на те, як ми налаштовували підпрограмне забезпечення `session` для API. Це тому, що вони призначені для використання одного й того самого методу автентифікації — отримання ідентифікатора сесії з cookie і його перевірки.

Раніше, коли ми налаштовували підпрограмне забезпечення `session`, ми явно не вказували ім’я cookie для сесії (`key`). Це тому, що пакет `session` використовував значення за замовчуванням. Тепер, коли ми додали ще один пакет, який потребує доступу до того самого значення з cookie, нам потрібно явно встановити значення `key` в обох об’єктах конфігурації.

Обов’язково додайте `key` з ім’ям cookie до підпрограмного забезпечення `session`, яке відповідає ключу Socket.IO. Також додайте референс `store` до опцій, поруч із тим, де ми встановлюємо `saveUninitialized: true`. Це необхідно, щоб сказати Socket.IO, до якої сесії відноситися.

<hr>

Тепер визначте функції зворотного виклику `success` і `fail`:

```js
function onAuthorizeSuccess(data, accept) {
  console.log('successful connection to socket.io');

  accept(null, true);
}

function onAuthorizeFail(data, message, error, accept) {
  if (error) throw new Error(message);
  console.log('failed connection to socket.io:', message);
  accept(null, false);
}
```

Об’єкт користувача тепер доступний у вашому об’єкті сокета як `socket.request.user`. Наприклад, тепер ви можете додати таке:

```js
console.log('user ' + socket.request.user.username + ' connected');
```

Це виведе в консоль сервера, хто підключився!

Відправте свою сторінку, коли вважатимете, що зробили все правильно. Якщо виникають помилки, ви можете  <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#authentication-with-socketio-9" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт до цього моменту</a>.

# --hints--

`passport.socketio` має бути залежністю.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport.socketio',
    'Your project should list "passport.socketio" as a dependency'
  );
```

`cookie-parser` має бути залежністю.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'cookie-parser',
    'Your project should list "cookie-parser" as a dependency'
  );
```

passportSocketIo має бути правильно імпортований.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\((['"])passport\.socketio\1\)/gi,
    'You should correctly require and instantiate "passport.socketio"'
  );
```

passportSocketIo має бути правильно налаштований.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io\.use\(\s*\w+\.authorize\(/,
    'You should register "passport.socketio" as socket.io middleware and provide it correct options'
  );
```
