---
id: 589fc830f9fc0f352b528e74
title: Налаштування середовища
challengeType: 2
forumTopicId: 301566
dashedName: set-up-the-environment
---

# --description--

Наступні завдання використовуватимуть файл `chat.pug`. Тож у вашому файлі `routes.js` додайте маршрут GET, який вказує на `/chat`, що використовує `ensureAuthenticated` і відтворює `chat.pug`, передаючи `{ user: req.user }` як аргумент у відповідь. Тепер змініть існуючий маршрут `/auth/github/callback`, щоб встановити `req.session.user_id = req.user.id` і зробити перенаправлення на `/chat`.

`socket.io@~2.3.0` вже додано як залежність, тож підключіть/ініціалізуйте його у вашому сервері так, як показано з `http` (вбудований у Node.js):

```javascript
const http = require('http').createServer(app);
const io = require('socket.io')(http);
```

Тепер, коли сервер *http* змонтовано на *express-застосунку*, потрібно слухати події від сервера *http*. Змініть лінію з `app.listen` на `http.listen`.

Перше, що потрібно обробити — це прослуховування нового підключення від клієнта. Ключове слово <dfn>on</dfn> саме це і робить — слухає конкретну подію. Воно вимагає 2 аргументи: рядок із назвою події, яка викликається, і функцію, через яку передаються дані. У випадку нашого слухача підключень використайте `socket` для визначення даних у другому аргументі. Сокет — це окремий клієнт, який підключений.

Щоб слухати підключення до вашого сервера, додайте наступне всередині підключення до бази даних:

```javascript
io.on('connection', socket => {
  console.log('A user has connected');
});
```

Тепер, щоб клієнт підключився, потрібно додати наступне у ваш `client.js`, який завантажується сторінкою після автентифікації:

```js
/*global io*/
let socket = io();
```

Коментар пригнічує помилку, яку ви зазвичай побачите, оскільки 'io' не визначено у файлі. Ви вже додали надійний CDN для бібліотеки Socket.IO на сторінці у `chat.pug`.

Тепер спробуйте завантажити свій застосунок і автентифікуватися, і ви повинні побачити у консолі сервера `A user has connected`.

**Примітка:** `io()` працює лише при підключенні до сокета, розміщеного на тому ж url/сервері. Для підключення до зовнішнього сокета, розміщеного в іншому місці, слід використовувати `io.connect('URL');`.

Відправте свою сторінку, коли вважатимете, що все зроблено правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-the-environment-6" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

`socket.io` має бути залежністю.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'socket.io',
    'Your project should list "socket.io" as a dependency'
  );
```

Ви маєте правильно підключити та ініціалізувати `http` як `http`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /http.*=.*require.*('|")http\1/s,
    'Your project should list "http" as a dependency'
  );
```

Ви маєте правильно підключити та ініціалізувати `socket.io` як `io`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.*=.*require.*('|")socket.io\1.*http/s,
    'You should correctly require and instantiate socket.io as io.'
  );
```

Socket.IO має слухати підключення.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.on.*('|")connection\1.*socket/s,
    'io should listen for "connection" and socket should be the 2nd arguments variable'
  );
```

Ваш клієнт має підключатися до вашого сервера.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.*=.*io/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
