---
id: 589fc832f9fc0f352b528e79
title: Надсилання та відображення повідомлень чату
challengeType: 2
forumTopicId: 301562
dashedName: send-and-display-chat-messages
---

# --description--

Настав час дозволити клієнтам надсилати повідомлення чату на сервер, щоб він міг розсилати їх усім клієнтам! У вашому файлі `client.js` ви вже маєте блок коду, який обробляє відправлення форми повідомлення.

```js
$('form').submit(function() {
  /*logic*/
});
```

У коді обробки відправлення форми ви повинні викликати подію після того, як визначите `messageToSend`, але до того, як очистите текстовий блок `#m`. Подія має називатися `'chat message'`, а дані — просто `messageToSend`.

```js
socket.emit('chat message', messageToSend);
```

Тепер на сервері ви повинні слухати сокет на подію `'chat message'`, де дані мають називатися `message`. Коли подія буде отримана, сервер має розіслати подію `'chat message'` усім сокетам за допомогою `io.emit`, передаючи об’єкт даних, що містить `username` і `message`.

У `client.js` тепер слід слухати подію `'chat message'` і, коли її отримано, додати елемент списку до `#messages` з ім’ям користувача, двокрапкою та повідомленням!

На цьому етапі чат має повністю працювати і надсилати повідомлення всім клієнтам!

Відправте свою сторінку, коли вважатимете, що все зроблено правильно. Якщо виникають помилки, ви можете <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#send-and-display-chat-messages-11" target="_blank" rel="noopener noreferrer nofollow">переглянути проєкт, виконаний до цього моменту</a>.

# --hints--

Сервер має слухати `'chat message'` і правильно його відправляти.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*io.emit.*('|")chat message('|").*username.*message/s,
    'Your server should listen to the socket for "chat message" then emit to all users "chat message" with name and message in the data object'
  );
```

Клієнт має правильно обробляти та відображати нові дані з події `'chat message'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*messages.*li/s,
    'You should append a list item to #messages on your client within the "chat message" event listener to display the new message'
  );
```
