---
id: a5229172f011153519423690
title: بناء حاسبة مجموع فيبوناتشي الفردية
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

في هذا المختبر ستبني حاسبة مجموع فيبوناتشي الفردية التي تحسب مجموع كل أعداد فيبوناتشي الفردية التي تكون أقل من أو تساوي عددًا صحيحًا موجبًا معينًا.

**الهدف**: حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة `sumFibs` تقبل رقمًا كمعلمة.
1. يجب أن تُرجع دالة `sumFibs` مجموع كل أعداد فيبوناتشي الفردية التي تكون أقل من أو تساوي الرقم المعطى.
1. تبدأ متتالية فيبوناتشي بـ `0` و `1`، وكل رقم لاحق هو مجموع الرقمين السابقين.
1. يجب إضافة أعداد فيبوناتشي الفردية فقط إلى المجموع.

# --hints--

يجب أن تملك دالة `sumFibs`.

```js
assert.isFunction(sumFibs);
```

يجب أن تُرجع `sumFibs(1)` رقمًا.

```js
assert.isNumber(sumFibs(1));
```

يجب أن تُرجع `sumFibs(1000)` القيمة `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

يجب أن تُرجع `sumFibs(4000000)` القيمة `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

يجب أن تُرجع `sumFibs(4)` القيمة `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

يجب أن تُرجع `sumFibs(75024)` القيمة `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

يجب أن تُرجع `sumFibs(75025)` القيمة `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
