---
id: 66be24cb4144f955b6bcc550
title: بناء صفحة من أوراق اللعب
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تبني صفحة ويب تعرض ثلاثة أوراق لعب على الأقل.
1. يجب أن تحتوي على عنصر `main` بمعرف `playing-cards`.
1. داخل عنصر `#playing-cards` يجب أن يحتوي على ثلاثة عناصر `div` على الأقل، كل منها بصنف `card`.
1. داخل كل عنصر `.card` يجب أن يحتوي على ثلاثة عناصر `div`، الأول بصنف `left`، والثاني بصنف `middle`، والثالث بصنف `right`.
1. يجب أن يستخدم عنصر `#playing-cards` flexbox لمركزة أطفاله أفقيًا، والسماح لهم بالالتفاف، ووضع فراغ `20px` بينهم.
1. يجب أن يستخدم كل عنصر من عناصر `.card` flexbox لتوزيع أطفاله باستخدام `space-between`.
1. يجب أن تستخدم كل عناصر `.left` خصائص عنصر flexbox لمحاذاة نفسها عند بداية الحاوية الأصلية.
1. يجب أن تستخدم كل عناصر `.middle` خصائص عنصر flexbox لمحاذاة نفسها في وسط الحاوية الأصلية.
1. يجب أن تستخدم كل عناصر `.right` خصائص عنصر flexbox لمحاذاة نفسها عند نهاية الحاوية الأصلية.
1. يجب أن تستخدم كل عناصر `.middle` flexbox لعرض أطفاله في عمود.

إليك بعض الرموز التي يمكنك نسخها ولصقها لاستخدامها في أوراق اللعب إذا أردت: `♠`، `♣`، `♥`، و`♦`.

# --hints--

يجب أن تحتوي على عنصر `main` مع `id` يساوي `playing-cards`.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

يجب أن يكون لديك ثلاثة عناصر `div` على الأقل بصنف `card` داخل عنصر `#playing-cards`.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

يجب أن يكون لكل عنصر من `.card` عرض وارتفاع.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

يجب أن يحتوي كل عنصر من `.card` على ثلاثة عناصر `div` فقط كأطفال.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

يجب أن يكون لديك عنصر `div` بصنف `left` داخل كل عنصر `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

يجب أن يكون لديك عنصر `div` بصنف `middle` داخل كل عنصر `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

يجب أن يكون لديك عنصر `div` بصنف `right` داخل كل عنصر `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

يجب أن يكون لديك محدد `#playing-cards` يضبط خاصية `display` لعناصره إلى `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

يجب أن يكون لديك محدد `#playing-cards` يضبط خاصية `justify-content` لعناصره إلى `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

يجب أن يكون لديك محدد `#playing-cards` يضبط خاصية `flex-wrap` لعناصره إلى `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

يجب أن يكون لديك محدد `#playing-cards` يضبط خاصية `gap` لعناصره إلى `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

يجب أن يكون لديك محدد `.card` يضبط خاصية `display` لعناصره إلى `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

يجب أن يكون لديك محدد `.card` يضبط خاصية `justify-content` لعناصره إلى `space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

يجب أن يكون لديك محدد `.left` يضبط خاصية `align-self` لعناصره إلى `flex-start`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

يجب أن يكون لديك محدد `.middle` يضبط خاصية `align-self` لعناصره إلى `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

يجب أن يكون لديك محدد `.right` يضبط خاصية `align-self` لعناصره إلى `flex-end`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

يجب أن يكون لديك محدد `.middle` يضبط خاصية `display` لعناصره إلى `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

يجب أن تضبط خاصية `flex-direction` إلى `column` في صنف `.middle`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
