---
id: 587d825a367417b2b2512c89
title: تنفيذ خوارزمية Quicksort
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

في هذا المختبر ستنفذ خوارزمية الفرز: quicksort.

quicksort هي طريقة فعالة تعتمد على التقسيم والتغلب بشكل تكراري لفرز مصفوفة. في هذه الطريقة، يُختار قيمة محورية في المصفوفة الأصلية. ثم تُقسّم المصفوفة إلى مصفوفتين فرعيتين تحتويان على القيم الأقل والأعلى من القيمة المحورية. بعد ذلك يُدمج ناتج الاستدعاء التكراري لخوارزمية quicksort على المصفوفتين الفرعيتين. تستمر هذه العملية حتى الوصول إلى الحالة الأساسية وهي مصفوفة فارغة أو تحتوي على عنصر واحد فقط، والتي تُعاد كما هي. فك الاستدعاءات التكرارية يعيد المصفوفة مرتبة.

مع أن اختيار القيمة المحورية مهم، أي قيمة محورية تصلح لأغراضنا هنا. للبساطة، يمكن استخدام العنصر الأول أو الأخير.

quicksort طريقة فرز فعالة جدًا، تقدم أداء *O(nlog(n))* في المتوسط. كما أنها سهلة التنفيذ نسبيًا. هذه الصفات تجعلها طريقة فرز شائعة ومفيدة.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تمتلك دالة `quicksort` تأخذ مصفوفة من الأعداد الصحيحة كمُدخَل.
1. يجب أن تُعيد دالة `quicksort` مصفوفة تحتوي نفس الأعداد الصحيحة لكن مرتبة من الأصغر إلى الأكبر.
1. يجب أن تستدعي دالة `quicksort` نفسها بشكل تكراري لفرز المصفوفة.

# --hints--

يجب أن تمتلك دالة `quicksort`.

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` يجب أن تُعيد مصفوفة لم تتغير إلا في الترتيب.

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` يجب أن تُعيد مصفوفة مرتبة (من الأقل إلى الأعلى).

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` يجب ألا تستخدم الطريقة المدمجة `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
