---
id: 6995d0325fd7da40854fac40
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

前のレッスンで、二分探索と線形探索の違いについて学びました。

このワークショップでは、それを実践して二分探索アルゴリズムを作成します。

まず、`binarySearch` 関数を `searchList` と `value` のパラメータで作成してください。

`searchList` は関数が探索する要素の配列で、`value` は関数が `searchList` の中から見つけようとするターゲットです。

# --hints--

`binarySearch` 関数を作成してください。

```js
assert.isFunction(binarySearch);
```

`binarySearch` 関数は `searchList` と `value` のパラメータを持つべきです。

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
