---
id: 6995d0325fd7da40854fac40
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

前のレッスンで、二分探索と線形探索の違いについて学びました。

このワークショップでは、二分探索アルゴリズムを作成して実践します。

まず、`binarySearch` 関数を `searchList` と `value` のパラメータで作成してください。

`searchList` は関数が探索する要素の配列で、`value` は関数が `searchList` の中から見つけようとするターゲットです。

# --hints--

`binarySearch` 関数を作成してください。

```js
assert.isFunction(binarySearch);
```

`binarySearch` 関数は `searchList` と `value` のパラメータを持つべきです。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
assert.lengthOf(binarySearch?.parameters, 2);
assert.equal(binarySearch?.parameters[0].toString(), "searchList");
assert.equal(binarySearch?.parameters[1].toString(), "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
