---
id: 6995d0325fd7da40854fac40
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

前のレッスンで、二分探索と線形探索の違いについて学びました。

このワークショップでは、二分探索アルゴリズムを作成して実践します。

まず、`binarySearch` 関数を `searchList` と `value` のパラメータで作成してください。

`searchList` は関数が探索する要素の配列で、`value` は関数が `searchList` の中から見つけようとするターゲットです。

# --hints--

`binarySearch` 関数を作成してください。

```js
assert.isFunction(binarySearch);
```

`binarySearch` 関数は `searchList` と `value` のパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('binarySearch', ['searchList', 'value']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
