---
id: 6996c60281feca6718cc15c1
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

`binarySearch` 関数の中で、まず `pathToTarget` 変数を作成し、空の配列に初期化してください。

# --hints--

`pathToTarget` 関数の中に、空の配列 `[]` に設定された `binarySearch` 変数があるはずです。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
