---
id: 6996c8375f4301004459ee85
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

`mid` インデックスを計算したので、次はそのインデックスにある `searchList` から実際に値を取得する必要があります。この値が、検索対象である `value` パラメータと比較する値になります。

`valueAtMiddle` 変数を作成し、`searchList` の `mid` インデックスにある要素を代入してください。

# --hints--

`valueAtMiddle` 変数を作成し、`searchList[mid]` に設定してください。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
