---
id: 6996cb5869acf58de751ae81
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

関数は `searchList` という配列と、探している `value` を受け取ることを思い出してください。

次に、`binarySearch([1, 2, 3, 4, 5], 3)` と `binarySearch([1, 2, 3, 4, 5, 9], 4)` を使って関数を呼び出し、両方の文をコンソールに出力してください。

# --hints--

`binarySearch([1, 2, 3, 4, 5], 3)` を使って関数を呼び出し、コンソールに出力してください。

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*\]\s*,\s*3\s*\)\s*\)/
);
```

`binarySearch([1, 2, 3, 4, 5, 9], 4)` を使って関数を呼び出し、コンソールに出力してください。

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*,\s*9\s*\]\s*,\s*4\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--  

--fcc-editable-region--  
```
