---
id: 6996c60281feca6718cc15c1
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

`binarySearch` 関数の中で、まず `pathToTarget` 変数を作成し、空の配列に初期化してください。

# --hints--

`binarySearch` 関数内で、空の配列に設定された `pathToTarget` 変数を宣言してください。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
