---
id: 6996c8f94fe6b250f11a2a1e
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

今度は、ターゲットかどうかに関わらず、検索中に通ったステップを追跡するために `valueAtMiddle` を `pathToTarget` 配列に含めてください。これは `push()` メソッドで行えます。

# --hints--

`push()` メソッドを使って `valueAtMiddle` を `pathToTarget` 配列に追加してください。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
