---
id: 6996c675b599baa7b103192d
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

二分探索では、最も低いインデックスと最も高いインデックスがあります。

最も低いインデックスは、現在探索している値の範囲の左端の境界を表し、この場合は `searchList` です。一方、最も高いインデックスは、探索している値の右端の境界を表します。

これら2つを考慮するために、探索している配列全体を考える必要があります。そこで、`binarySearch` 関数の中で、値が `low` の `0` 変数と、探索している配列の最後のインデックスを表す値の `high` 変数を定義してください。

# --hints--

`low` 関数の中に、`0` という名前の変数を `binarySearch` に設定しているはずです。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

`high` 関数の中に、`searchList.length - 1` という名前の変数を `binarySearch` に設定しているはずです。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
