---
id: 6996d0a2c2ac55041145a13f
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

2回目の関数呼び出しでリストに `3, 5, 4` の値が表示されたことに気づきましたか？これは、二分探索が意図した通りに動作していることを示しています。

こうして動作しました：アルゴリズムはまず初期リストの中央として `3` を確認しました。`4` が `3` より大きいため、探索は右半分に移動しました。次に新しい中央として `5` を調べました。`4` が `5` より小さいため、探索は左に移動し、最終的に `4` を最終範囲の中央として特定しました。

関数をもう一度テストするには、`[1, 3, 5, 9, 14, 22], 10` を使って呼び出し、コンソールにログを出力してください。この場合、値は見つかりません。

# --hints--

関数を `binarySearch([1, 3, 5, 9, 14, 22], 10)` で呼び出し、コンソールにログを出力してください。

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
