---
id: 6996ccadf9048c79b857bb15
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

最初の呼び出しが `[3]` を返すのがわかります。これは `3` が `[1, 2, 3, 4, 5]` の中央の値だからです。

一方、2回目の呼び出しは空の配列を返します。これは、中央の値を確認した後、ループが現在 `low` または `high` の変数を更新して検索範囲を狭めるロジックを持っていないために起こります。単に中央の値を一度チェックして、それ以上の検索をせずに終了しています。

2分探索が検索範囲を狭め続けられるように、`else if` が `value` より大きいかどうかをチェックする `valueAtMiddle` 文を追加してください。

# --hints--

条件 `else if` を持つ `value > valueAtMiddle` 文を追加してください。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
