---
id: 6996ca21a74104721f2c7e5d
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

アルゴリズムの流れをこの初期段階で理解するために、動作をテストする必要があります。

そのためには、まずループから抜け出す必要があります。現在の実装では1回の繰り返ししか許されないため、条件が満たされないと無限ループになってしまいます。

`if` 文の直後に、`break` ループから抜け出すために `while` キーワードを使ってください。そして、`while` ループの後に、値が見つからなかったことを示すために空の配列を返してください。

# --hints--

`break`文の後に`while`ループから抜け出すために`if`キーワードを使うべきです。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

`while` ループの後に空の配列を返すべきです。

```js
assert.deepEqual(binarySearch([], NaN), []);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
