---
id: 6996cb5869acf58de751ae81
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

関数は `searchList` という配列と、探している `value` を受け取ることを思い出してください。

次に、`binarySearch([1, 2, 3, 4, 5], 3)` と `binarySearch([1, 2, 3, 4, 5, 9], 4)` を使って関数を呼び出し、両方の文をコンソールに出力してください。

# --hints--

`binarySearch([1, 2, 3, 4, 5], 3)` を使って関数を呼び出し、コンソールに出力してください。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5],3))";
assert.include(strippedCode, expectedStatement);
```

`binarySearch([1, 2, 3, 4, 5, 9], 4)` を使って関数を呼び出し、コンソールに出力してください。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5,9],4))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--

--fcc-editable-region--
```
