---
id: 6996c675b599baa7b103192d
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

二分探索では、最も低いインデックスと最も高いインデックスがあります。

最も低いインデックスは、現在探索している値の範囲の左端の境界を表し、この場合は `searchList` です。一方、最も高いインデックスは、探索している値の右端の境界を表します。

これら2つを考慮するために、検索対象の配列全体を考える必要があります。そこで、`binarySearch` 関数の中で、値が `0` の `low` 変数と、検索対象の配列の最後のインデックスを表す値の `high` 変数を宣言してください。

# --hints--

`binarySearch` 関数内で、`low` という名前の変数を `0` に設定して宣言してください。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

`binarySearch` 関数内で、`high` という名前の変数を `searchList.length - 1` に設定して宣言してください。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
