---
id: 6996c7363a4a437839f1f15d
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

検索範囲の境界を定義したので、次は二分探索アルゴリズムを実行するループを作成しましょう。

二分探索は、検索範囲を繰り返し狭めていくことで動作します。この処理は、チェックすべき有効な範囲がある限り続きます。

これは、`while` ループで表現でき、`low` ポインターが `high` ポインター以下である限り続きます。

もし `low` が `high` より大きくなった場合、検索範囲が空になったことを意味し、`value` は配列に存在しません。

`binarySearch` 関数の中で、`while` ループを作成し、条件として `low` が high 以下であるかをチェックしてください。

# --hints--

`while` ループを作成し、`low` が `high` 以下の間実行されるようにしてください。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
