---
id: 6996cfb6c01e9a802add38be
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

最後に、`value` が中央に見つからず、かつ `valueAtMiddle` より大きくもない場合は、`value` は `valueAtMiddle` より小さく、左側にある必要があります。この処理は最終的に `else` ブロックで行います。

それでは、`break` キーワードを削除し、`else` ブロックを追加してループを完了させてください。その中で、`high` 変数を `1` から `mid` を引いた値に更新します。

# --hints--

`break` キーワードを `else` ブロックに置き換え、`high` 変数を `mid - 1` に更新してください。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    }

--fcc-editable-region--      
    break;
--fcc-editable-region--

  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
