---
id: 6996c796385b6ea4e222d5e8
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

ループを始めるには、現在の探索範囲の中央のインデックスを見つける必要があります。

中央のインデックスは、`low` と `high` の変数の平均を計算することで求められます。2つの値を足した後に床除算を使って平均を出す必要があるため、答えは常に切り捨てられた整数になります。インデックスは整数でなければならないからです。

`while` のループ内で、`mid` と `low` の平均を設定した `high` 変数を作成してください。

# --hints--

while ループ内で `mid` 変数を `Math.floor((low + high) / 2)` に設定してください。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
--fcc-editable-region--
      
--fcc-editable-region--
  }
}
```
