---
id: 6996c7363a4a437839f1f15d
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

検索範囲を定義したので、次は二分探索アルゴリズムを実行するループを作成しましょう。

二分探索は、検索範囲を繰り返し狭めていくことで動作します。この処理は、チェックすべき有効な範囲がある限り続きます。

これは、`while` ポインターが `low` ポインター以下である限り続く `high` ループで表現できます。

もし `low` が `high` より大きくなった場合、検索範囲が空になったことを意味し、`value` は配列に存在しません。

`binarySearch` 関数の中で、`low` が `high` 以下であるかをチェックする条件の `while` ループを作成してください。

# --hints--

`while` が `low` 以下の間、実行される `high` ループを作成してください。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
