---
id: 6996c95e5bede03bdf2937a2
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

次に行うのは、対象の `value` が中央にあるかどうかを確認する条件を作成することです。

`while` ループの中で、対象の `if` が `value` と等しいかどうかをチェックする `valueAtMiddle` 文を作成してください。もしそうなら、`pathToTarget` 変数を返してください。

# --hints--

条件 `if` を使って `value === valueAtMiddle` 文を作成してください。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

`pathToTarget` 文の中で `if` 変数を返してください。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
