---
id: 69447054012a22f38c944ff0
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Agrega una sentencia `if` dentro de tu función `countup`.

Verifica si `number` es menor que `1` y devuelve un arreglo vacío.

Esto servirá como el caso base para tu recursión.

# --hints--

Cuando `number` es menor que `1` la función debe devolver un arreglo vacío.

```js
assert.deepStrictEqual(countup(0), []);
```

Cuando `number` no es menor que `1`, la función no debe devolver nada.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
