---
id: 69446f18fdcb11aa516f7c5c
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, crearás una función `countup`, que crea un arreglo con números del 1 hasta un número dado usando recursión.

Comienza creando una función vacía llamada `countup` que tome un solo parámetro.

# --hints--

Debes tener una función llamada `countup`.

```js
assert.isFunction(countup);
```

Tu función `countup` debe aceptar un solo parámetro.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
