---
id: 6948705160b66cf93a916e29
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Declara un arreglo vacío llamado `countArray` dentro de la función `countup` justo encima de tu declaración `if`.

Este arreglo almacenará el valor de `number` en cada llamada a la función.

# --hints--

Debes declarar un arreglo vacío llamado `countArray`.

```js
const explorer = await __helpers.Explorer(code);
const { countup } = explorer.functions;
assert.exists(countup);
const { countArray } = countup.variables;
assert.exists(countArray);
assert.isTrue(countArray.value.matches('[]'));
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
