---
id: 694870bcf3c6874112b6f951
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Para todos los números mayores o iguales a `1`, crea un bloque `else` y dentro de él, llama a la función `countup` con `number - 1` y asigna el resultado a `countArray`. 

Luego usa `.push()` para agregar `number` a `countArray`, y finalmente devuelve `countArray`. 

Esto asegura que la recursión eventualmente alcance el caso base y termine.

# --hints--

Debes asignar el resultado de llamar a `countup` con `number - 1` a `countArray`.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

Debes insertar `number` en `countArray`.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

Debes devolver el `countArray`.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
