---
id: 68e4c81577243e4db1676491
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 의미론적 HTML을 사용해 개발자 직업을 얻기 위한 Quincy Larson의 팁을 포함하는 웹 페이지를 만드는 연습을 합니다. 기본 HTML 보일러플레이트가 준비되어 있습니다.

먼저 텍스트가 `h1`인 `Quincy's Tips for Getting a Developer Job` 요소를 만드세요.

# --hints--

`h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 요소의 텍스트는 `Quincy's Tips for Getting a Developer Job`여야 합니다.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
