---
id: 68e65ec5348e40a739903e6d
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

이제 `h1` 요소 아래에 단락 요소를 만드세요. 이 단락 안에 `Learning to code is hard, but as Quincy Larson says, You can become a developer.` 텍스트를 추가하세요.

# --hints--

`p` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('p'));
```

단락에 `Learning to code is hard, but as Quincy Larson says, You can become a developer.` 텍스트가 있어야 합니다. 철자를 다시 확인하세요.

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

단락은 `h1` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
