---
id: 68e90335bd53a901c41dfc13
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이전 수업에서 `q` 요소가 인용된 텍스트를 주변 내용과 구분하는 데 사용된다는 것을 배웠습니다.

예를 들어 보겠습니다:

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

대부분 최신 브라우저는 `q` 요소를 사용할 때 인라인 인용문 주위에 자동으로 따옴표를 추가합니다.

단락 안에서 `You can become a developer.`를 인라인 인용 요소로 감싸고, 나머지 단락은 변경하지 마세요.

# --hints--

`q` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('q'));
```

`q` 요소는 `h1` 요소 아래에 있는 단락 안에 중첩되어야 합니다.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

`q` 요소는 `You can become a developer.` 텍스트를 포함해야 합니다.

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

`q` 요소 안의 텍스트는 앞뒤에 공백(스페이스, 탭, 줄 바꿈)이 없어야 합니다.

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

원본 텍스트를 변경하지 않고 문장의 적절한 부분에만 `q` 태그를 추가해야 합니다. `p` 요소는 여전히 `Learning to code is hard, but as Quincy Larson says, You can become a developer.`을 읽어야 합니다.

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
