---
id: 68e96ce52087043b0e999296
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

단락 요소 아래에 `main` 요소를 추가하고 그 안에 세 개의 `section` 요소를 중첩하세요.

# --hints--

단락 요소 아래에 `main` 요소가 있어야 합니다.

```js
const mainEl = document.querySelector('p + main');
assert.exists(mainEl);
```

`section` 요소 안에 세 개의 `main` 요소가 있어야 합니다.

```js
const sections = document.querySelectorAll('main > section');
assert.lengthOf(sections, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
