---
id: 68e964b84f954c322376e592
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

인용문의 출처가 웹사이트인 경우 `cite` 속성으로 출처를 표시할 수 있습니다. 이 속성의 값은 유효한 URL이어야 합니다. 이 속성은 블록 인용문의 표시 방식을 바꾸지 않지만, 화면 읽기 프로그램과 검색 엔진에 인용문에 대한 추가 정보를 제공하는 데 매우 유용합니다.

다음은 `cite` 속성이 포함된 인라인 인용 요소의 예입니다:

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

이 URL로 인라인 인용 요소에 `cite` 속성을 추가하세요:

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

`q` 요소에 `cite` 속성이 있어야 합니다.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

`cite` 요소의 `q` 속성 값은 `https://www.freecodecamp.org/news/learn-to-code-book/`이어야 합니다.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
