---
id: 68ea5ccc1229b85eca38a3e8
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

첫 번째 `section` 요소 안에 텍스트 `h2`를 가진 `Envisioning Success` 요소를 추가하세요.

# --hints--

첫 번째 `section` 요소 안에 `h2` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('section:first-of-type > h2'));
```

`h2`에는 텍스트 `Envisioning Success`가 있어야 합니다.

```js
const firstH2El = document.querySelector('section:first-of-type > h2');
assert.equal(firstH2El?.innerText.trim(), `Envisioning Success`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
