---
id: 68eaa3593b6ffd019d772a1f
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

인용 블록 요소 아래에 `&mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]` 텍스트가 들어간 단락 요소를 추가하세요.

`&mdash;`은 대시 `—`를 나타내는 HTML 엔티티입니다.

# --hints--

`p` 요소 아래에 `blockquote` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('section:first-of-type > blockquote + p'));
```

`p` 요소에는 `&mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]` 텍스트가 있어야 합니다.

```js
const pEl = document.querySelector('section:first-of-type > blockquote + p');
const html = pEl?.innerHTML.trim();
assert.oneOf(html, [
  '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]',
  '&mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]' 
]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
