---
id: 68eaa9ef8a0867071f195f59
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

지금까지는 `cite` 속성을 사용해 인용문의 출처를 표시했지만, 이 속성은 사용자에게 출처를 시각적으로 보여주지 않습니다.

출처를 시각적으로 표시하려면, 인용문 블록 요소 밖에 `cite` 인용 요소를 추가할 수 있습니다. 인용 요소는 책, 기사, 노래, 영화, 웹사이트, 연구 논문 같은 참조된 창작물의 제목을 표시하는 데 사용하는 HTML 요소입니다.

예를 들어 다음과 같습니다:

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>&mdash;Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

인용문 블록 요소 아래의 `p` 요소 안에 `How to Learn to Code and Get a Developer Job [Full Book]`을 `cite` 요소로 감싸세요.

# --hints--

`cite` 요소 아래 `p` 요소 안에 `blockquote` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

`cite` 요소에는 `How to Learn to Code and Get a Developer Job [Full Book]` 텍스트가 있어야 합니다.

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

방금 추가한 `cite` 요소를 무시하면, `p` 요소의 텍스트는 여전히 `—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]` 여야 합니다.

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
      --fcc-editable-region--
          &mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
      --fcc-editable-region--
        </p>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
