---
id: 68ea7e7b916e8270483689b6
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`h2` 요소 아래에 Quincy의 또 다른 인용문을 추가하세요. 이번에는 인용문이 더 큰 단락의 일부가 아니고, 단락 전체가 인용문이 될 것입니다. 이렇게 인용된 텍스트를 구분하려면 블록 인용 요소인 `blockquote`를 사용하세요. 브라우저에서 보면 텍스트가 약간 들여쓰기된 것을 확인할 수 있습니다.

다음은 인용된 텍스트가 포함된 블록 인용 요소의 예입니다:

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

이제 첫 번째 섹션 안에서 `h2` 요소 아래에 `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?` 텍스트가 있는 블록 인용 요소를 추가하세요.

# --hints--

첫 번째 `section` 요소 안에 `blockquote` 요소 아래에 `h2` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

`blockquote` 요소에는 `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?` 텍스트가 있어야 합니다.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
