---
id: 68ea9d23318b6481ac3de777
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

인라인 인용 요소와 마찬가지로, 블록 인용 요소에도 `cite` 속성을 추가할 수 있습니다.

다음은 `cite` 속성이 있는 블록 인용 요소 예시입니다:

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

이제 URL `cite`을 가진 블록 인용 요소에 `https://www.freecodecamp.org/news/learn-to-code-book/` 속성을 추가하세요.

# --hints--

`blockquote` 요소에 `cite` 속성이 있어야 합니다.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.exists(firstBlockquoteEl?.getAttribute('cite'));
```

`cite` 속성은 `https://www.freecodecamp.org/news/learn-to-code-book/` 값을 가져야 합니다.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
      --fcc-editable-region--
        <blockquote>
      --fcc-editable-region--
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
