---
id: 68eab6a2b37d2f13602d0c3f
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

인용 블록 요소 안에 직접 인용문을 작성할 수 있습니다. 첫 번째 `section` 요소에서 했던 것처럼요. 또는 하나 이상의 단락 요소로 감쌀 수도 있습니다. 텍스트가 여러 단락으로 구성되어 있지만 같은 인용 블록 요소 안에 유지하고 싶을 때 유용합니다.

두 단락이 있는 예시는 다음과 같습니다:

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  <p>So many people ask me each week: is college still worth it? In this 1-hour video I answer this question and other commonly asked questions about university.</p>
  <p>I've been in adult education for two decades at this point, and even though I'm not a labor market economist, I do feel confident enough to answer these questions.</p>
</blockquote>
```

두 번째 섹션에서, 기존 인용 블록 요소 안에 다음 텍스트를 가진 네 개의 `p` 요소를 순서대로 추가하세요:

- `So much of getting a job is who you know.`
- `It's OK to be an introvert, but you do need to push your boundaries.`
- `Create GitHub, Twitter, LinkedIn, and Discord accounts.`
- `Go to tech meetups and conferences. Travel if you have to.`

# --hints--

두 번째 섹션 안의 `p` 요소 안에 네 개의 `blockquote` 요소가 중첩되어 있어야 합니다.

```js
const pElements = document.querySelectorAll('main > section:nth-of-type(2) > blockquote > p');
assert.lengthOf(pElements, 4);
```

첫 번째 `p` 요소는 `So much of getting a job is who you know.` 텍스트를 가져야 합니다.

```js
const firstPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(1)');
assert.equal(firstPEl?.innerText.trim(), 'So much of getting a job is who you know.');
```

두 번째 `p` 요소는 `It's OK to be an introvert, but you do need to push your boundaries.` 텍스트를 가져야 합니다.

```js
const secondPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(2)');
assert.equal(secondPEl?.innerText.trim(), "It's OK to be an introvert, but you do need to push your boundaries.");
```

세 번째 `p` 요소는 `Create GitHub, Twitter, LinkedIn, and Discord accounts.` 텍스트를 가져야 합니다.

```js
const thirdPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(3)');
assert.equal(thirdPEl?.innerText.trim(), 'Create GitHub, Twitter, LinkedIn, and Discord accounts.');
```

네 번째 `p` 요소는 `Go to tech meetups and conferences. Travel if you have to.` 텍스트를 가져야 합니다.

```js
const forthPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(4)');
assert.equal(forthPEl?.innerText.trim(), 'Go to tech meetups and conferences. Travel if you have to.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>
      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
        --fcc-editable-region--
          
        --fcc-editable-region--
        </blockquote>
      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
