---
id: 68ebe9cd7523d340c0b343d3
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

세 번째이자 마지막 `section` 요소에서 이전 섹션과 동일한 HTML 구조를 만드세요.

먼저, 마지막 섹션 안에 텍스트가 `h2`인 `Importance of Building a Reputation` 요소를 추가하세요.

그다음, `h2` 요소 아래에 `cite` 속성이 `https://www.freecodecamp.org/news/learn-to-code-book/`로 설정된 인용문 요소를 추가하세요.

마지막으로, 인용문 요소 안에 다음 텍스트를 순서대로 가진 세 개의 단락 요소를 중첩하세요:

1. `Share short video demos of your projects.`
2. `Keep applying to speak at bigger and bigger conferences.`
3. `Hang out at hackerspaces and help people who are even newer to coding than you.`

축하합니다! 이로써 이 워크숍을 완료했습니다.

# --hints--

세 번째 섹션 안에 `h2` 요소가 중첩되어 있어야 합니다.

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > h2'));
```

`h2` 요소에는 텍스트 `Importance of Building a Reputation`이 있어야 합니다.

```js
const h2El = document.querySelector('main > section:nth-of-type(3) > h2');
assert.equal(h2El?.innerText.trim(), 'Importance of Building a Reputation');
```

세 번째 섹션 안에 `blockquote` 요소가 중첩되어 있어야 합니다.

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > blockquote'));
```

`blockquote` 요소는 `h2` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > h2 + blockquote'));
```

`blockquote` 요소에는 `cite` 속성이 있어야 합니다.

```js
const blockquoteEl = document.querySelector('main > section:nth-of-type(3) > blockquote');
assert.exists(blockquoteEl?.getAttribute('cite'));
```

`cite` 요소의 `blockquote` 속성은 `https://www.freecodecamp.org/news/learn-to-code-book/` 값을 가져야 합니다.

```js
const blockquoteEl = document.querySelector('main > section:nth-of-type(3) > blockquote');
assert.equal(blockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

세 번째 섹션의 `p` 요소 안에 세 개의 `blockquote` 요소가 중첩되어 있어야 합니다.

```js
const pElements = document.querySelectorAll('main > section:nth-of-type(3) > blockquote > p');
assert.lengthOf(pElements, 3);
```

첫 번째 `p` 요소에는 텍스트 `Share short video demos of your projects.`이 있어야 합니다.

```js
const firstPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(1)');
assert.equal(firstPEl?.innerText.trim(), 'Share short video demos of your projects.');
```

두 번째 `p` 요소에는 텍스트 `Keep applying to speak at bigger and bigger conferences.`가 있어야 합니다.

```js
const secondPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(2)');
assert.equal(secondPEl?.innerText.trim(), "Keep applying to speak at bigger and bigger conferences.");
```

세 번째 `p` 요소에는 텍스트 `Hang out at hackerspaces and help people who are even newer to coding than you.`이 있어야 합니다.

```js
const thirdPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(3)');
assert.equal(thirdPEl?.innerText.trim(), 'Hang out at hackerspaces and help people who are even newer to coding than you.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>
      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>So much of getting a job is who you know.</p>
          <p>It's OK to be an introvert, but you do need to push your boundaries.</p>
          <p>Create GitHub, Twitter, LinkedIn, and Discord accounts.</p>
          <p>Go to tech meetups and conferences. Travel if you have to.</p>
        </blockquote>
      </section>
      <section>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>

      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>So much of getting a job is who you know.</p>
          <p>It's OK to be an introvert, but you do need to push your boundaries.</p>
          <p>Create GitHub, Twitter, LinkedIn, and Discord accounts.</p>
          <p>Go to tech meetups and conferences. Travel if you have to.</p>
        </blockquote>
      </section>

      <section>
        <h2>Importance of Building a Reputation</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>Share short video demos of your projects.</p>
          <p>Keep applying to speak at bigger and bigger conferences.</p>
          <p>Hang out at hackerspaces and help people who are even newer to coding than you.</p>
        </blockquote>
      </section>
    </main>
  </body>
</html>
```
