---
id: 68f1eb9ad0a7575e4dbc31d5
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

次に、`video`の下に`h1`要素を作成してください。これからの数ステップで、ビデオプレーヤーを機能させるために必要な属性を追加していきます。

# --hints--

`video`要素があるはずです。

```js
assert.exists(document.querySelector('video'));
```

新しい`video`要素は`h1`の下に配置してください。

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
