---
id: 68f1f6f97af9749a42260a43
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`controls` 属性は、`video` 要素の再生、巻き戻し、音量調整などの再生コントロールを提供します。

`controls` 属性はブール値の属性で、値を指定する必要はありません。

`controls` 要素に `video` 属性を追加してください。

これでページに `video` 要素が表示されるはずです。

# --hints--

`video` 要素には `controls` 属性が必要です。

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
