---
id: 68f1f57723debf9383c75947
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`loop` 属性は、再生が完了すると動画を再起動します。再生を繰り返すインターネットミームのように考えてください。`loop` 属性を省略すると、動画は一度だけ再生されます。

`loop` 属性はブール値の属性で、値を指定する必要はありません。

`loop` 属性を `video` 要素に追加してください。

# --hints--

`video` 要素には `loop` 属性が必要です。

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
