---
id: 68f1f371fefd448b4f70b898
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

前のレッスンで、`video` 要素に使えるさまざまな属性について学びました。`width` 属性は動画の幅をピクセル単位で決めます。

`width` 要素に `video` 属性を追加し、その値を `640` にしてください。

# --hints--

`video` 要素には、属性 `width` を値 `640` で設定してください。

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
