---
id: 68f1f993460f06a64c033061
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`poster` 属性は動画のサムネイル画像です。YouTubeで見る動画をイメージしてください。動画がダウンロードされている間に表示されます。もしこの属性が省略されると、ダウンロード中は動画の最初のフレームが表示されます。

では、`poster` 要素に `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` 属性を値 `video` で追加してください。

# --hints--

`video` 要素には、値 `poster` の `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` 属性が必要です。

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
