---
id: 68f18724d71db1dda8d6fec7
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、HTMLビデオプレーヤーを作成します。HTMLのボイラープレートはすでに用意されています。

`h1` 要素を作成し、テキストに `Working with the HTML Video Element` を設定してください。

# --hints--

`h1` 要素があるはずです。

```js
assert.exists(document.querySelector('h1'));
```

`h1` 要素にはテキスト `Working with the HTML Video Element` が含まれている必要があります。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Working with the HTML Video Element');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
