---
id: 68f1fac259a4afabb478b960
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

実際のビデオにリンクしていないことに気づいたかもしれません。次のステップでそれを行います。ビデオファイルの種類にはブラウザのサポートに違いがあります。これに対応するために、`source` 要素の中に `video` 要素を使うことができ、ブラウザは最初に互換性のある `source` を選択します。

こちらは `source` 要素の例です：

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

`source` 要素は空要素なので終了タグはありません。

`source` 要素の中に `video` 要素を追加してください。


# --hints--

`source` 要素の中に `video` 要素があるはずです。

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
