---
id: 68f1f8c248306da204655117
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`muted` 属性は、最初の再生時に音声をミュートにします。`controls` が有効になっている場合、ユーザーは音声のミュートを解除できます。`muted` 属性を省略すると、最初の再生時に音声が再生されます。

`muted` 属性はブール値の属性で、値を指定する必要はありません。

`muted` 属性を `video` 要素に追加してください。

# --hints--

`video` 要素には `muted` 属性が必要です。

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
