---
id: 68f27139bddf9f6bf2821acf
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

最後に追加する`source`要素は、`video/quicktime` MIMEタイプ用です。

QuickTimeはAppleが作成した拡張可能なマルチメディアアーキテクチャで、さまざまなデジタルメディア形式の再生、ストリーミング、エンコード、トランスコードをサポートしています。MP4形式ほど一般的ではありませんが、レガシーアプリケーションのサポートに必要な場合があります。

3つ目の`source`要素の下に、4つ目の`source`要素を追加し、`src`属性に`https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`の値を、`type`属性に`video/quicktime`の値を設定してください。

おめでとうございます！HTMLビデオプレイヤーワークショップを完了しました。

# --hints--

4つ目の`source`要素があるはずです。

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

4つ目の`source`要素には`src`属性が必要です。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

4つ目の`source`要素には、`src`属性があり、その値は`https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`です。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

4つ目の`source`要素には`type`属性があります。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

4つ目の`source`要素には、`type`属性があり、その値は`video/quicktime`です。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
