---
id: 68f2701b2160c765f23b60f8
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

もう一つよく使われる MIME タイプは `video/ogg` MIME タイプです。

Ogg は、デジタルマルチメディアの効率的なストリーミングと操作を目的に設計されたデジタルマルチメディアのコンテナ形式です。Xiph.Org Foundation によって管理されており、ソフトウェア特許の制限を受けない無料でオープンな形式です。その名前はコンピューターゲーム Netrek の専門用語「ogging」に由来しています。

2つ目の `source` 要素の下に、3つ目の `source` 要素を追加し、`src` 属性に `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` の値を、`type` 属性に `video/ogg` の値を設定してください。

# --hints--

3つ目の `source` 要素があるはずです。

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

3つ目の `source` 要素には `src` 属性が必要です。

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

3つ目の `source` 要素には `src` 属性があり、その値は `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` であるべきです。

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

3つ目の `source` 要素には `type` 属性が必要です。

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

3つ目の `source` 要素には `type` 属性があり、その値は `video/ogg` であるべきです。

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
