---
id: 68f26a5f79e97c5032c10642
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`mp4` ファイル拡張子のビデオファイルを使っていて、ブラウザーにそのファイルの読み方を認識させる必要があります。

`type` 属性を使って `video/mp4` MIMEタイプを指定します。

MIME（Multipurpose Internet Mail Extensions）は、ASCIIテキスト以外の形式のドキュメント、例えば音声、動画、画像などを記述するための標準です。

MP4は正式にはMPEG-4 Part 14として知られるデジタルマルチメディアコンテナ形式です。動画や音声の保存に広く使われていますが、字幕や静止画像など他のデータタイプも含めることができます。MP4ファイルはインターネットでのストリーミング用に設計されており、多くのデバイスやプラットフォームと互換性があります。

では、`type` 属性と値 `video/mp4` を追加してください。

# --hints--

`source` 要素には `type` 属性が必要です。

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

`source` 要素には値 `type` の `video/mp4` 属性が必要です。

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
