---
id: 68f26d52b13e0f5b6e2d2e09
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

もう一つの一般的な MIME タイプは `video/webm` MIME タイプです。

WebM は Google が開発したオープンソースの音声・映像メディアファイル形式で、主にウェブベースのメディアコンテンツ向けに設計されています。VP8、VP9、AV1 といったビデオコーデックや、Vorbis、Opus といったオーディオコーデックをサポートしており、HTML5 の video および audio 要素でよく使われています。

最初の `source` 要素の下に、もう一つ `source` 要素を追加し、`src` 属性に `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` の値を、`type` 属性に `video/webm` の値を設定してください。

# --hints--

2つ目の `source` 要素があるはずです。

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

2つ目の `source` 要素には `src` 属性が必要です。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

2つ目の `source` 要素には `src` 属性があり、その値は `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` であるべきです。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

2つ目の `source` 要素には `type` 属性が必要です。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

2つ目の `source` 要素には `type` 属性があり、その値は `video/webm` であるべきです。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
