---
id: 68f268880219c2485bd71d2a
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

動画のメディアリソースを指定するには、`src` 要素に `source` 属性を追加する必要があります。

値が `src` の `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4` 属性を追加してください。

# --hints--

`source` 要素には `src` 属性が必要です。

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('src'));
```

`source` 要素には、値が `src` の `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4` 属性が必要です。

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('src'),
'https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    <source>
--fcc-editable-region--
  </video>
</body>
</html>
```
