---
id: 66b205e6eacba4c4e54ea434
title: بناء أداة إخفاء البريد الإلكتروني
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

في هذا المختبر، ستقوم بإخفاء جزء اسم المستخدم من عنوان البريد الإلكتروني باستخدام النجوم. الإخفاء هو مصطلح يُستخدم لإخفاء أو استبدال المعلومات الحساسة بالنجوم أو أحرف أخرى.

على سبيل المثال، إذا كان عنوان البريد الإلكتروني هو `myEmail@email.com`، فسيكون عنوان البريد الإلكتروني المخفي هو `m*****l@email.com`.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. أنشئ دالة باسم `maskEmail` تأخذ `email` كمعلمة.
2. داخل الدالة، عليك إخفاء جزء اسم المستخدم في `email` وإلحاق اسم النطاق به. تذكر أنه يمكنك استخدام طرق مثل `slice`، `repeat`، `indexOf` أو حتى `replace` لمساعدتك.
3. خارج الدالة، أعلن عن متغير باسم `email` لتخزين عنوان البريد الإلكتروني الذي تريد إخفاءه.
4. استدعِ دالة `maskEmail` مع المتغير `email` واطبع النتيجة في وحدة التحكم.
5. يجب أن تُرجع `maskEmail("apple.pie@example.com")` القيمة `"a*******e@example.com"`.
6. يجب أن تُرجع `maskEmail("freecodecamp@example.com")` القيمة `"f**********p@example.com"`.
7. يجب أن تُرجع `maskEmail("info@test.dev")` القيمة `"i**o@test.dev"`.
8. يجب أن تُرجع `maskEmail("user@domain.org")` القيمة `"u**r@domain.org"`.

# --hints--

يجب أن تعرف دالة باسم `maskEmail`.

```js
assert.isFunction(maskEmail);
```

يجب أن تأخذ دالة `maskEmail` سلسلة نصية، `email` كمعلمة.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

يجب أن يكون لديك متغير `email` خارج الدالة.

```js
assert.isDefined(email);
```

يجب أن تعيّن عنوان بريد إلكتروني صالح إلى متغير `email` الخاص بك.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

يجب أن تُعيد الدالة `maskEmail("apple.pie@example.com")` القيمة `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

يجب أن تُعيد الدالة `maskEmail("freecodecamp@example.com")` القيمة `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

يجب أن تُعيد الدالة `maskEmail("info@test.dev")` القيمة `"i**o@test.dev"`.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

يجب أن تُعيد الدالة `maskEmail("user@domain.org")` القيمة `"u**r@domain.org"`.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

يجب أن تُنتج دالة `maskEmail` الخاصة بك النتيجة الصحيحة.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

يجب أن تطبع ناتج استدعاء `maskEmail` مع `email` كمعلمة.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
