---
id: 67457ff2dc531fbed928a0f0
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

حان الوقت الآن لبناء مكوِّن `Navbar` باستخدام JSX.

تعلمت في الدروس السابقة أن JSX هي امتداد صيغة نحوية لـ JavaScript يسمح لك بكتابة شفرة تشبه HTML داخل JavaScript، مما يسهل إنشاء وإدارة مكوِّنات React.

إليك تذكير بما تبدو عليه JSX:

```jsx
const App = () => {
  return (
    <div>
      <h1>freeCodeCamp React Curriculum</h1>
      <p>I love learning with freeCodeCamp!</p>
    </div>
  );
}
```


داخل مكوِّن `Navbar`، أعد عنصر `nav` يحتوي على خاصية `className` بقيمة `navbar`.

السبب في استخدام `className` هنا بدلاً من `class` هو أن `class` كلمة محجوزة في JavaScript. لذلك، يستخدم React `className` كطريقة لتطبيق الأصناف على العناصر.

# --hints--

يجب أن تُعيد عنصر `nav`.

```js
assert.exists(document.querySelector('nav'));
```

يجب أن يحتوي عنصر `nav` الخاص بك على خاصية `className` بقيمة `navbar`.

```js
assert.equal(document.querySelector('nav')?.getAttribute('class'), 'navbar');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Navbar</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Navbar } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <Navbar />
      );
    </script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --white: #fff;
  --light-grey: #e1e0e0;
  --dark-purple: #7c0e7c;
  --black: #000;
}

body {
  background-color: var(--light-grey);
}

.navbar {
  background-color: var(--white);
}

.navbar ul {
  display: flex;
  justify-content: space-around;
}

.navbar ul li {
  list-style: none;
  border-radius: 4px;
}

.navbar ul li a {
  text-decoration: none;
  color: var(--black);
  padding: 10px;
  display: inline-block;
  width: 100%;
}

button {
  background: transparent;
  border: none;
  font-family: 'Times New Roman', Times, serif;
  padding: 10px;
  font-size: 1rem;
}

.navbar ul .nav-item a:hover {
  background-color: var(--dark-purple);
  color: var(--white);
}

button:hover {
  background-color: var(--dark-purple);
  color: var(--white);
}

.navbar ul .nav-item .sub-menu {
  visibility: hidden;
  opacity: 0;
  position: absolute;
  right: 5%;
  transition: opacity 0.5s ease;
  display: block;
  background-color: var(--white);
}

@media (min-width: 768px) {
  .navbar ul .nav-item .sub-menu {
    right: 15%;
  }
}

@media (min-width: 1024px) {
  .navbar ul .nav-item .sub-menu {
    right: 13%;
  }
}

.navbar ul .nav-item:hover .sub-menu,
.navbar ul .nav-item:focus-within .sub-menu {
  visibility: visible;
  opacity: 1;
}
```

```jsx
export const Navbar = () => {
  --fcc-editable-region--
  
  --fcc-editable-region--
}
```
