---
id: 587d8259367417b2b2512c85
title: 選択ソートアルゴリズムを実装する
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

このラボでは、選択ソートアルゴリズムを実装します。選択ソートは、リストの中で最小値を選び、それをリストの最初の値と入れ替えることで動作します。次に2番目の位置から始め、残りのリストの中で最も小さい値を選び、2番目の要素と入れ替えます。リストの最後までこの操作を繰り返し、要素を入れ替え続けます。これでリストはソートされます。選択ソートはすべての場合で二次の時間計算量を持ちます。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `selectionSort` は関数である必要があります。
1. `selectionSort` 関数は配列を受け取り、同じ要素を持つが最小値から最大値の順に並んだ配列を返す必要があります。

# --hints--

`selectionSort` は関数である必要があります。

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` は順序以外は変更されていない配列を返す必要があります。

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` は最小値から最大値にソートされた配列を返す必要があります。

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` は組み込みの `.sort()` メソッドを使わないでください。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
