---
id: 655771d889132f9ccd341060
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

當圖的邊帶有權重，表述距離、時間或其他量化值時，該圖稱為*加權*圖。

在你的情況中，這些權重將是每個節點或空間中點之間的距離。為了表述加權圖，你可以修改你的字典，對每個值使用一個元組的列表。

元組中的第一個元素將是連接的節點，第二個元素將是一個整數，表示距離。

將 `my_graph['A']` 修改為元組的列表，並考慮 `A-B` 距離為 `3` 以及 `A-D` 距離為 `1`。

# --hints--

`my_graph["A"]` 應該是一個包含元組 `('B', 3)` 和 `('D', 1)` 的列表。

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` 應該有 4 個鍵，名稱分別為 `'A'`、`'B'`、`'C'` 和 `'D'`。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
