---
id: 655776db1eeae0a620e42a0d
title: 步驟 31
challengeType: 20
dashedName: step-31
---

# --description--

現在，呼叫你的函式，並傳入 `my_graph` 和 `'A'` 作為引數。

# --hints--

你應該呼叫 `shortest_path`，並傳入 `my_graph` 和 `'A'` 作為引數。

```js
({ test: () => assert.match(code, /^shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
--fcc-editable-region--
```
