---
id: 655775221059f5a20493d5d7
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

在你的 `for` 迴圈之後，為你的程式碼添加一個 `print()` 呼叫，並傳入以下字串以查看你已創建的變數值：`f'Unvisited: {unvisited}\nDistances: {distances}'`。

# --hints--

你應該在你的 `for` 迴圈之後列印 `f'Unvisited: {unvisited}\nDistances: {distances}'`。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)for.*:.*^\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
--fcc-editable-region--
```
