---
id: 6557712d77ce2d9bd7e63afd
title: 步驟18
challengeType: 20
dashedName: step-18
---

# --description--

在你的圖中添加另一個節點連接 `B` 節點，並將其命名爲 `C` 。

修改你現有的字典來表示這種排列方式：在 `my_graph` 中添加另一個鍵 `'C'` 併爲其賦值爲字符串 `'B'` 。

另外，更改現有的 `'B'` 鍵的值爲列表 `['A', 'C']` 來表示 `'B'` 節點的多個連接。

# --hints--

你的字典應該包含三個鍵 - `'A'`、 `'B'`、 和 `'C'` 。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` 的值應該是字符串 `'B'` 。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` 應該是一個列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']` 的值應該是包含 `'A'` 和 `'C'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']` 的值應該是字符串 `'B'` 。

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
