---
id: 65577236b056379d5dbc7000
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

現在你將開始開發演算法，以計算你新圖中每個節點之間的最短路徑。

宣告一個名為 `shortest_path` 的空函式。使用 `pass` 關鍵字填充函式主體。 

# --hints--

你應該定義一個名為 `shortest_path` 的函式（程式）。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
