---
id: 6557924d47c325bf27afbe51
title: 步驟 53
challengeType: 20
dashedName: step-53
---

# --description--

演算法已完成，但你可以改進輸出。此外，你可以為 `function` 提供額外的引數，以傳回僅在兩個節點之間的路徑。

為你的函式宣告式將 `target` 添加為第三個參數，並給予其空字串的預設值。

# --hints--

你的函式（程式）應該接受三個參數：`graph`、`start` 和 `target`。順序很重要。

```js
({ test: () => assert.match(code, /^def\s+shortest_path\s*\(\s*graph\s*,\s*start\s*,\s*target\s*=?\s*.*\s*\)\s*:/m) })
```

`target` 參數應該有空字串的預設值。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == "(graph, start, target='')"
  `))
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
--fcc-editable-region--
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
