---
id: 655774955b097ea14897db12
title: 步驟 28
challengeType: 20
dashedName: step-28
---

# --description--

起始節點的距離為零，因為演算法從那裡開始評估。

在你的迴圈中將 `node` 附加到 `unvisited` 後，創建一個 `if` 述語，當該節點相等於起始節點時觸發。然後在 `distances` 字典中為該節點指定 `0`。

# --hints--

你應該創建一個 `if` 述語，當 `node` 與 `start` 相等時執行。

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

在你的新 `if` 述語中，你應該將 `0` 指定給 `distances` 字典中的節點。

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
