---
id: 65578d4fc3afc3b8f554c882
title: 步驟 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` 也需要一個只使用關鍵字的參數。 將一個函數作爲附加參數傳遞到 `min()`, 你可以修改列表項的比較方式。

你剛剛在上一步寫下的行的結果是按字母順序排列的節點。 相反，你想選擇距離起始節點最小的未訪問節點。

將 `key=distances.get` 作爲第二個參數傳遞給你的 `min()` 調用。 這樣，比較將根據每個 `unvisited` 列表項目在 `distances` 字典中的值進行。

# --hints--

您應該將`key=distances.get`作爲調用`min()`的第二個參數傳遞。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')

```
