---
id: 65578d0f6c78a0b868a43b9c
title: 步驟 40
challengeType: 20
dashedName: step-40
---

# --description--

在 `while` 迴圈中，第一件要做的事是定義目前要造訪的節點。為此你可以使用 `min()` 函式。它會傳回作為引數傳入的可疊代物件中最小的項目。

移除 `pass`，然後創建一個名為 `current` 的變數並指定其為 `min(unvisited)`。

# --hints--

你應該在你的 `while` 迴圈中創建一個 `current` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

你應該將 `min(unvisited)` 指定到你的 `current` 變數。記得刪除 `pass`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
