---
id: 655774d01daeeaa1978b99d5
title: 第 29 步
challengeType: 20
dashedName: step-29
---

# --description--

一開始，圖中所有其他節點都被視為與 `source` 節點的距離為無限大，因為距離尚未確定。

創建一個 `else` 語句，並將無限值指定給 `distances` 字典中的節點。為此，使用帶有字串 `'inf'` 作為引數的 `float()` 函式（程式）來產生表述正無限大的浮點數。

# --hints--

你應該有一個 `else` 語句。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

你應該在你的新 `else` 語句中將 `float('inf')` 指定給 `distances[node]`。

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
