---
id: 6557746aad2844a0cd864e12
title: 步驟 27
challengeType: 20
dashedName: step-27
---

# --description--

當演算法探索圖形時，應該追蹤起始節點與其他節點之間目前已知的最短距離。

在你的 `for` 迴圈之前，創建一個名為 `distances` 的新變數，並將它指定為一個空字典。

# --hints--

你應該有一個名為 `distances` 的變數。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

你的 `distances` 變數應該是一個空字典。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
