---
id: 6557709b0aee699a6a00528c
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

例如，圖形可以用來表述空間中的兩個點 `A` 和 `B`，並由一條路徑連接。這樣的圖形將由兩個節點和一條邊所組成。

將現有的 `'species'` 鍵替換為字串 `'A'`。然後，將對應的值替換為字串 `'B'`，以表述 `'A'` 和 `'B'` 節點之間的連接。

# --hints--

你的字典應該有一個 `'A'` 鍵。

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

你的 `'A'` 鍵應該有 `'B'` 作為值。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

你的字典應該有兩個鍵。

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
