---
id: 655773b0591c5f9f4045883e
title: 第 24 步
challengeType: 20
dashedName: step-24
---

# --description--

演算法將從指定的節點開始。然後它會探索圖形，以尋找起始節點或 *source* 與所有其他節點之間的最短路徑。

為你的函式（程式）添加兩個參數：`graph` 和 `start`。將它們添加到你的函式（程式）宣告式中。

# --hints--

你的函式（程式）應該依此順序以 `graph` 和 `start` 作為參數。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
