---
id: 65577739f57ecca6c39bb4e9
title: 步驟 33
challengeType: 20
dashedName: step-33
---

# --description--

`list()` 型別建構式可讓你從可迭代物件建置列表。

修改你 `unvisited` 變數的指派以使用 `list()`，並將 `graph` 作為可迭代物件傳入。

# --hints--

你應該使用 `list()` 從 `graph` 字典中產生一個列表。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/list\s*\(\s*graph\s*\)/));
  }
})
```

你應該將 `list(graph)` 指定給 `unvisited`。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\s*=\s*list\s*\(\s*graph\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
