---
id: 65578cb031cd93b77a285db2
title: 第 38 步
challengeType: 20
dashedName: step-38
---

# --description--

你的函式（程式）將會探索所有與起始節點相連的節點。它會計算所有節點的最短路徑。然後，它會從未造訪節點中移除起始節點。 

接著，將會造訪最近的鄰居節點，並重複該處理程序直到所有節點皆被造訪。

從現在開始，你將著手處理探索圖中節點的主要迴圈。為避免在演算法開發期間發生無窮迴圈的問題，將你的函式呼叫改成註解。

# --hints--

你應該將你的函式呼叫改成註解。

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
