---
id: 65577a17564ce8a8e06c1460
title: 步驟 35
challengeType: 20
dashedName: step-35
---

# --description--

字典推導式也支援條件式 `if`/`else` 語法：

```py
{key: val_1 if condition else val_2 for key in dict}
```

在上述範例中，`dict` 是現有的字典。當 `condition` 計算結果為 `True` 時，`key` 將會有值 `val_1`，否則為 `val_2`。

使用字典推導式根據 `graph` 創建字典並將其指定給 `distances` 變數。如果節點與起始節點相等，則給鍵指定值為零，否則為無限大。使用 `float('inf')` 來達成後者。

# --hints--

你應該使用字典推導語法為你的 `distances` 變數指定值。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
