---
id: 65579228c669fcbebffd01d5
title: 步驟 52
challengeType: 20
dashedName: step-52
---

# --description--

另一個錯誤比較微妙。當找到鄰近節點的較短距離時，`paths[current]` 會被指定為鄰近節點路徑 `paths[node]`。 

這表示兩個變數都指向相同的列表。由於列表是可變的，當你將鄰居節點附加到它的路徑時，`paths[node]` 和 `paths[current]` 都會被修改，因為它們是相同的列表。這會導致錯誤的路徑，儘管距離是正確的。

你可以透過將 `paths[current]` 的複本指定給鄰居節點路徑來修正該錯誤。為此你可以使用切片語法：

```py
my_list[:]
```

其中 `my_list` 是你想要複製的列表。在你的 `if` 區段內，透過切片 `paths[current]` 來修改現有的 `paths[node] = paths[current]` 指派。

# --hints--

你應該將 `paths[current][:]` 指定到鄰近節點路徑。

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
