---
id: 6557709b0aee699a6a00528c
title: 步驟 16
challengeType: 20
dashedName: step-16
---

# --description--

例如，圖形可用於表示空間中由路徑連接的兩個點 `A` 和 `B`。 像這樣的圖將由兩個通過邊連接的節點組成。

用字符串 `'A'` 替換現有的 `'species'` 鍵。 然後，將 correspondent 的值替換爲字符串 `'B'` 來表示 `'A'` 和 `'B'` 兩個節點之間的連接。

# --hints--

Your dictionary should have an `'A'` key.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Your `'A'` key should have `'B'` as the value.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Your dictionary should have two keys.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
