---
id: 6557712d77ce2d9bd7e63afd
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

為你的圖形添加另一個連接到 `B` 的節點，並將它命名為 `C`。

修改你現有的字典以表述此排列：為 `my_graph` 添加另一個鍵 `'C'` 並給它字串 `'B'` 的值。

此外，將現有的 `'B'` 鍵的值更改為列表 `['A', 'C']`，以表述你的 `'B'` 節點的多個連接。

# --hints--

你的字典應該有 3 個鍵 — `'A'`、`'B'` 和 `'C'`。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` 的值應該是字串 `'B'`。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` 應該是一個列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']` 的值應該是一個包含 `'A'` 和 `'C'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']` 的值應該是字串 `'B'`。

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
