---
id: 655790d113d14dbb727eaf41
title: 步驟 45
challengeType: 20
dashedName: step-45
---

# --description--

一旦在 `distances` 字典中設定了到節點的距離，你也需要追蹤到該節點的路徑。如果已更新處理過的元組中節點的距離，其路徑中的最後一個項目即為該節點本身。

在你的條件式中，巢狀另一個 `if` 述語，當 `paths[node]` 的最後一個元素與 `node` 相等時觸發。使用 `pass` 來填充 `if` 述語的主體。

# --hints--

你應該創建巢狀的 `if` 述語來檢查 `paths[node][-1]` 是否與 `node` 相等。別忘了使用 `pass`。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1(\s{4})if\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\2\2pass/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')
```
