---
id: 65576ff7888f9e96f52a4be1
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

字典以 *key*-*value* 配對的表單形式儲存資料。鍵與對應的值以冒號分隔。每個鍵值配對與下一個配對以逗號分隔：

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

為你的字典添加一個新的鍵值對。使用字串 `'species'` 作為鍵，並使用字串 `'guinea pig'` 作為值。

# --hints--

你應該在你的 `copper` 字典中有一個鍵為 `'species'`，值為 `'guinea pig'`。

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

你的 `copper` 字典應該只有一個鍵值配對。

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
