---
id: 6557716aadbd2d9c42c0e69a
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

為你的節點添加最後一個節點，`'D'`，它與 `'A'` 和 `'C'` 相連。

修改你的字典以表現此結構。再次，使用列表來表現多個連接。

# --hints--

你的字典應該有 4 個鍵，分別稱為 `'A'`、`'B'`、`'C'` 和 `'D'`。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` 應該是一個列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["A"]) is list
  `))
})
```

`my_graph['A']` 應該是包含 `'B'` 和 `'D'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["A"]) == 2 and "B" in my_graph["A"] and "D" in my_graph["A"]
  `))
})
```

`my_graph['B']` 應該是一個列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']` 應該是包含 `'A'` 和 `'C'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']` 應該是一個列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["C"]) is list
  `))
})
```

`my_graph['C']` 應該是包含 `'B'` 和 `'D'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["C"]) == 2 and "B" in my_graph["C"] and "D" in my_graph["C"]
  `))
})
```

`my_graph['D']` 應該是一個列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["D"]) is list
  `))
})
```

`my_graph['D']` 應該是包含 `'A'` 和 `'C'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["D"]) == 2 and "A" in my_graph["D"] and "C" in my_graph["D"]
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': ['A', 'C'],
    'C': 'B'
}
--fcc-editable-region--
```
