---
id: 65578f895f2a65ba7a916804
title: 步驟 42
challengeType: 20
dashedName: step-42
---

# --description--

在 `current` 變數指派之後，創建一個 `for` 迴圈來迭代 `graph[current]` 列表中的元組。你將需要兩個迭代變數來完成這個操作。記得使用 `pass` 來填充迴圈主體。

# --hints--

你應該創建一個 `for` 迴圈來迭代 `graph[current]` 列表中的 tuples 項目。使用兩個迭代變數，並且不要忘記添加 `pass` 關鍵字。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)while\s+unvisited\s*:.*^\1\1for\s+\w+\s*,\s*\w+\s+in\s+graph\s*\[\s*current\s*\]\s*:\s*^\1\1\1pass/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited, key=distances.get)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
