---
id: 655773f8b8b5db9fc6d0ae76
title: 第 25 步
challengeType: 20
dashedName: step-25
---

# --description--

為了追蹤已訪問的節點，你需要一個圖中所有節點的列表。一旦節點被訪問，它將從該列表中移除。

現在，將 `pass` 關鍵字取代為名為 `unvisited` 的變數，並指定一個空的列表。

# --hints--

你應該在 `shortest_path` 函式（程式）內有一個名為 `unvisited` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

你應該為你的 `unvisited` 變數指定一個空列表。記得刪除 `pass`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
