---
id: 65577791ad8c26a7705e2919
title: 步驟 34
challengeType: 20
dashedName: step-34
---

# --description--

使用字典推導式，你可以從現有的字典創建一個字典：

```py
{key: val for key in dict}
```

在上述範例中，`val` 是 `key` 在新字典中將擁有的值，而 `dict` 是現有的字典。

你想要追蹤起始節點與每個其他節點之間的路徑。

在 `distances` 變數之後，創建一個 `paths` 變數並指定它為一個包含 `graph` 所有鍵的字典。為每個鍵指定一個空列表，並使用字典推導式來建置你的字典。

# --hints--

你應該有一個名為 `paths` 的變數。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

你的 `paths` 變數應該使用字典推導語法，為圖中每個節點指定一個空列表。

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
