---
id: 655771d889132f9ccd341060
title: 步驟 20
challengeType: 20
dashedName: step-20
---

# --description--

當一個圖的每一邊都有定量比如距離，時間或者其他的定量值，那這個圖叫做 *weighted* graph。

在此例子裏定量值將是兩個點的距離。 您可以把字典改成一串系的元組來代表一個定量值圖。

每一個元組裏第一個變量是連接的點，第二個變量是兩點的距離。

將 `my_graph['A']` 修改爲元組列表，考慮 `A-B` 距離爲 `3` ，`A-D` 距離爲 `1` 。

# --hints--

`my_graph["A"]` 應該是一個包含元組 `('B', 3)` 和 `('D', 1)` 的列表。

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` 應該有4個鍵值名爲 `'A'`、 `'B'`、 `'C'` 和 `'D'`。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
