---
id: 65579228c669fcbebffd01d5
title: 第 52 步
challengeType: 20
dashedName: step-52
---

# --description--

另一個錯誤則比較隱蔽。 當找到一條距離相鄰節點比較短的路徑時，`paths[current]` 將被分配給相鄰節點路徑，即 `paths[node]` 。

這意味着兩個變量指向相同的列表。 因爲列表是可變的，當你將鄰接節點附加到其路徑時， `paths[node]` 和 `paths[current]` 都被修改，因爲它們是相同的列表。 這導致路徑錯誤，儘管距離是正確的。

你可以通過將一個 `paths[current]` 分配給鄰居節點路徑來修復這個錯誤。 爲此，你可以使用切片語法：

```py
my_list[:]
```

`my_list` 是你想要複製的列表。 通過切分 `paths[current]` 來修改 `if` 塊內的現有 `paths[node] = paths[current]` 賦值。

# --hints--

你應該將 `paths[current][:]` 分配給相鄰節點路徑。

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)

    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)

    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

shortest_path(my_graph, 'A')

```
