---
id: 65578d4fc3afc3b8f554c882
title: 步驟 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` 也接受一個僅限關鍵字的引數。將函式作為額外引數傳遞給 `min()`，你可以修改列表項目的比較方式。

你剛剛在前一步驟中寫下的那一行結果是字母順序中第一個的節點。相反地，你想選擇距離起始節點距離最小且尚未訪問的節點。 

將 `key=distances.get` 作為第二個引數傳遞給你的 `min()` 呼叫。這樣，比較將根據每個 `unvisited` 列表項目在 `distances` 字典中的值進行。

# --hints--

你應該將 `key=distances.get` 作為第二個引數傳遞給你的 `min()` 呼叫。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
