---
id: 65578c17d54dfab65cd54b95
title: 第 36 步
challengeType: 20
dashedName: step-36
---

# --description--

由於演算法從起始節點開始評估，在建立 `paths` 字典後，你需要將起始節點添加到 `paths` 字典中它自己的列表。

使用 `.append()` 方法將 `start` 附加到 `paths[start]` 列表。


# --hints--

你應該使用 `.append()` 方法將 `start` 附加到 `paths[start]`。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}paths\s*\[\s*start\s*\]\s*\.append\s*\(\s*start\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
