---
id: 6559d70c5161b16ff1d6530d
title: 步驟 54
challengeType: 20
dashedName: step-54
---

# --description--

Python 提供了一種使用三元語法撰寫 `if`/`else` 條件判斷的簡潔方式：

```py
val_1 if condition else val_2
```

上述表達式在 `condition` 為 true 時計算為 `val_1`，否則計算為 `val_2`。

刪除你的 `print` 呼叫，並在你的 `while` 迴圈之後創建一個名為 `targets_to_print` 的變數。當 `target` 為真值時，使用三元語法將其指定為 `[target]`，否則指定為 `graph`。


# --hints--

你應該刪除你的 `print` 呼叫。

```js
({ test: () => assert.isFalse( /print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\\nPaths:\s\{\s*paths\s*\}\1\s*\)/.test(code)) })
```

你應該在你的 `while` 迴圈之後創建一個名為 `targets_to_print` 的變數。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\.remove\(\s*current\s*\).*^\s{4}targets_to_print\s*=/ms));
  }
})
```

當 `target` 為真值時，你應該使用三元語法將 `[target]` 指定給你的 `targets_to_print` 變數，否則指定 `graph`。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\.remove\(\s*current\s*\).*^\s{4}targets_to_print\s*=\s*\[\s*target\s*\]\s+if\s+target\s+else\s+graph/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start, target = ''):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
