---
id: 6559d70c5161b16ff1d6530d
title: 步驟54
challengeType: 20
dashedName: step-54
---

# --description--

Python 提供了三元語法來簡潔地書寫 `if`/`else` 條件語句:

```py
val_1 if condition else val_2
```

如果 `condition` 爲 true，則上述表達式的計算結果爲 `val_1`，否則爲 `val_2`。

刪除您的 `print` 調用，並在您的`while`循環語句後創建一個名爲`targets_to_print`的變量。 當 `target` 爲真時，使用三元語法賦值爲 `[target]`，否則爲 `graph`。


# --hints--

您應該刪除您的 `print` 調用。

```js
({ test: () => assert.isFalse( /print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\\nPaths:\s\{\s*paths\s*\}\1\s*\)/.test(code)) })
```

您應該在您的 `while`循環之後創建名爲 `targets_to_print`的變量。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\.remove\(\s*current\s*\).*^\s{4}targets_to_print\s*=/ms));
  }
})
```

當target爲真時，你應該使用三元語法將`[target]`賦值給`target`，否則`graph`賦值給你的`targets_to_print`變量。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\.remove\(\s*current\s*\).*^\s{4}targets_to_print\s*=\s*\[\s*target\s*\]\s+if\s+target\s+else\s+graph/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start, target = ''):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)

    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

shortest_path(my_graph, 'A')
--fcc-editable-region--
```
