---
id: 6566195b0a021bb660b2b4b1
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

現在將 `my_graph['B']` 修改為一個元組的列表，元組中的第一個元素是連接的節點，第二個元素是距離。`B-C` 的距離是 `4`。

# --hints--

`my_graph['B']` 應該是一個包含元組 `('A', 3)` 和 `('C', 4)` 的列表。

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph` 應該有 4 個鍵，名稱分別為 `'A'`、`'B'`、`'C'` 和 `'D'`。

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
