---
id: 65661b72d6745ebec6a96923
title: 步驟 22
challengeType: 20
dashedName: 步驟-22
---

# --description--

同樣，考慮到 `C-D` 距離是 `7` ，修改剩餘的兩個列表。

# --hints--

`my_graph['C']` 應該是一個包含元組 `('B', 4)` 和 `('D', 7)` 的列表。

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

`my_graph['D']` 應該是一個包含元組 `('A', 1)` 和 `('C', 7)` 的列表。

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

`my_graph` 應該有4個鍵值名爲 `'A'`、 `'B'`、 `'C'` 和 `'D'`。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` 應該是一個包含元組 `('B', 3)` 和 `('D', 1)` 的列表。

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph['B']` 應該是一個包含元組 `('A', 3)` 和 `('C', 4)` 的列表。

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
