---
id: 65789506b30453080f77470c
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

到目前爲止，你已經遇到了不同的數據類型：

- 不可變的數據類型，如整數、字符串、元組和布爾值。
- 可變數據類型，例如列表和字典。

字典由一對花括號`{}` 標識。

首先創建一個變量 `copper` ，並使用一對花括號爲其分配一個空字典，就像使用一對方括號創建一個空列表一樣。

# --hints--

你應該有一個名爲 `copper` 的變量。

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

你的 `copper` 變量應該有一個空字典的值。 爲此使用一對花括號。

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
