---
id: 6578b57361f2f132a02e2a18
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

你可以透過字典的鍵來存取儲存在其中的資料：

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

在你的字典之後，依照上述範例存取 `copper` 的 `'species'` 鍵並列印結果。

# --hints--

你不應該修改你的字典的指派。

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

你應該使用 `copper['species']` 來存取 `'species'` 鍵的值。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

你應該呼叫 `print()`，並傳入 `copper['species']` 作為引數。

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
