---
id: 6578b13757611e2825beb8a5
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

鍵必須在字典中唯一，且只能是不可變的資料型別。這表示你不能使用列表或另一個字典作為鍵。

為你的字典添加另一個鍵 `'age'` 並將整數數字 `2` 作為值。

# --hints--

你應該在你的 `copper` 字典中，位於 `'species': 'guinea pig'` 之後，加入一個 `'age'` 鍵，其值為 `2`。別忘了逗號。

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

你的 `copper` 字典中應該有兩個鍵值配對。

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
