---
id: 65789506b30453080f77470c
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

到目前為止，你已經遇到不同的資料型別：

- 不可變的資料型別，例如整數、字串、元組和布林值。
- 可變的資料型別，例如列表和字典。

字典由一對大括號 `{}` 所標示。

首先創建一個名為 `copper` 的變數，並使用一對大括號指定一個空字典，就像你使用一對方括號創建一個空列表一樣。

# --hints--

你應該有一個名為 `copper` 的變數。

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

你的 `copper` 變數應該具有空字典的值。為此請使用一對大括號。

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
