---
id: 6579cbab9825b8170974c69a
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

你可以使用 `del` 關鍵字從字典中移除一個鍵值配對：

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

就在你的 `for` 迴圈之前，使用 `del` 關鍵字從 `copper` 中刪除 `'age'` 鍵及其值。

# --hints--

你應該在 `for` 迴圈之前使用 `del` 關鍵字刪除 `copper['age']`。

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
