---
id: 6579ca0923cfa7162089d2f0
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

相同的語法可用於更改現有鍵的值。

就在 `print()` 呼叫之前，存取 `'species'` 鍵並將其值重新指定為 `'Cavia porcellus'`。

# --hints--

你應該在 `print()` 呼叫之前，將 `copper['species']` 重新指定為字串 `'Cavia porcellus'`。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
