---
id: 65796fac81f983127558f3f4
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

現在，修改你現有的 `print()` 呼叫以列印 `age` 鍵的值。


# --hints--

你應該使用 `copper['age']` 來存取 `age` 鍵的值。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

你應該呼叫 `print()`，並傳入 `copper['age']` 作為引數。

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
