---
id: 6579717f0920131304286804
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

要在宣告字典後新增一個新的鍵值對，你可以用存取現有鍵的相同方式來指出鍵，並使用指定運算子設值新的鍵的值：

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

刪除你的 `print()` 呼叫。然後，在 `copper` 宣告式之後，為你的字典添加鍵 `'food'` 並將其值設為 `'hay'`。

# --hints--

你不應該在你的程式碼中有 `print(copper['age'])`。

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

你應該在宣告字典後，為你的 `copper` 添加鍵 `'food'`。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

你應該在宣告字典後將 `copper['food']` 設定為 `'hay'`。

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
