---
id: 6579cd5f6dd62c189e53ddbb
title: 步驟14
challengeType: 20
dashedName: step-14
---

# --description--

現在你已經複習了字典的基本方面，可以繼續構建最短路徑算法。

刪除`copper`字典聲明後的每一行代碼。

# --hints--

你應該刪除字典聲明後的所有代碼行。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
