---
id: 65796fac81f983127558f3f4
title: 步驟5
challengeType: 20
dashedName: step-5
---

# --description--

現在，修改你現有的 `print()` 調用來打印 `age` 鍵值。


# --hints--

你應該使用 `copper['age']` 來訪問鍵 `age` 的值。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

你應該調用 `print()`，將`copper['age']`作爲參數傳遞。

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
