---
id: 6579cd5f6dd62c189e53ddbb
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

現在你已經檢視了字典的基本面向，就可以繼續建置最短路徑演算法。

刪除在 `copper` 字典宣告之後的每一行程式碼。

# --hints--

你應該刪除字典宣告後的所有程式碼行。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
