---
id: 6579717f0920131304286804
title: 步驟6
challengeType: 20
dashedName: step-6
---

# --description--

想要在字典聲明之後添加新的鍵值對，你可以通過和訪問已有的鍵值相同的方式聲明鍵，並使用賦值操作符設置新鍵的值：

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

刪除你的 `print()` 調用。 然後，在 `copper` 聲明之後，添加鍵 `'food'` 到字典，並賦值爲 `'hay'` 。

# --hints--

你不應該在你的代碼裏出現 `print(copper['age'])` 。

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

你應該在字典聲明之後添加鍵 `'food'` 到 `copper` 。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

你應該在字典聲明之後設置 `copper['food']` 的值爲 `'hay'` 。

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
