---
id: 6579dd49fa8a8e1fd06b85a9
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

圖形是表述元素配對之間關係的資料結構。這些元素稱為*節點*，可以是真實世界的物件、實體、空間中的點或其他。節點之間的連線稱為*邊*。

這是一個圖形的視覺表述：

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


將 `copper` 字典重新命名為 `my_graph`。這將表述用來測試你的演算法的圖形。

# --hints--

你應該將你的 `copper` 字典重新命名為 `my_graph`。

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

你的 `my_graph` 變數應該是一個字典。

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
