---
id: 65affe602a2a828b0382a058
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

將 `'age'` 鍵取代為字串 `'B'`，並將其值設定為字串 `'A'`，以表述節點之間雙向的連接。

# --hints--

你的字典應該有一個 `'B'` 鍵。

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

你的 `'B'` 鍵應該有 `'A'` 作為值。

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

你的字典應該有一個 `'A'` 鍵。

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

你的 `'A'` 鍵應該有 `'B'` 作為值。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

你的字典應該有兩個鍵。

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
