---
id: 65b7d25992879698180e6a71
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

最後，如果你想要能夠遍歷鍵值對組，你可以使用 `.items()` 方法。

將你的 `for` 迴圈修改為迭代 `copper.items()`，而非 `copper.values()`。

# --hints--

你應該修改你的 `for` 迴圈，以迭代 `copper.items()` 而不是 `copper.values()`。

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
