---
id: 65b7cf140d34058d7ea8935f
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

如果你想要迭代字典鍵的值，一種方法是使用 `.values()` 方法。

修改你的 `for` 迴圈以迭代 `copper.values()` 而不是 `copper`，並查看輸出。

# --hints--

你應該修改你的 `for` 迴圈，以迭代 `copper.values()` 而不是 `copper`。

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
