---
id: 65b7d3319ebec69b983fb91b
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

如你從輸出中所見，`.items()` 會創建一個資料結構，將每個鍵值配對儲存在不同的元組中。若要迭代這些元組中的元素，你可以為你的第二個迴圈變數添加。

```py
for i, j in dict.items():
    print(i, j)
```

修改你的 `for` 迴圈，使其接受兩個迴圈變數，並在迴圈主體內列印它們兩個。

# --hints--

你應該修改你的 `for` 迴圈以使用兩個變數。

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

你應該在你的 `for` 迴圈中列印兩個迴圈變數。

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
