---
id: 65b7cd2b7bd9a684ccf1dc16
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

要對字典的鍵進行迭代，你可以直接將字典放入 `for` 迴圈。以下程式碼會列印字典 `dict` 中的每個鍵：

```py
for i in dict:
   print(i)
```

將 `print()` 呼叫取代為一個 `for` 迴圈，該迴圈迭代 `copper` 並列印每個鍵。

# --hints--

你的程式碼中不應該有 `print(copper)`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

你應該創建一個 `for` 迴圈來迭代 `copper` 字典。

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

你應該在你的 `for` 迴圈中列印 `copper` 的每個鍵。

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
