---
id: 6565a536ba1f9f25bd30e88b
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika mradi huu, utautekeleza algoriti ya Luhn. Algoriti hii ni fomula ya kuthibitisha nambari mbalimbali za utambulisho.

Anza kwa kutangaza kitendakazi kinachoitwa `main`, hiki kitakuwa kama sehemu ya kuanzisha programu. Tumia neno kuu la `pass` ili kuepuka kosa.

# --hints--

Unapaswa kuwa na `def main():` katika msimbo wako.

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

Unapaswa kuwa na `pass` katika msimbo wako.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
