---
id: 65687db6a9fcf3c548a6c1b7
title: Hatua 28
challengeType: 20
dashedName: step-28
---

# --description--

Sehemu ya algoriti ni kuzidisha mara mbili kila nambari ya pili, kuanzia upande wa kulia. Ikiwa matokeo ya kuzidisha mara mbili nambari ni kubwa au sawa na `10`, ongeza nambari hizo mbili pamoja. Kwa mfano, ikiwa nambari ni `6`, zidisha mara mbili kupata `12`. Ongeza `1` na `2` pamoja kupata `3`. Unaweza kufanya hivi kwa kutumia _integer division_ kupata nambari ya kwanza na operator wa moduli (`%`) kupata nambari ya pili:

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

Mgawanyiko wa nambari kamili (integer division) hutoa gawio la mgawanyiko, ukizungushwa chini hadi nambari kamili iliyo karibu.

Ndani ya tamko la `if`, weka thamani ya `number` kuwa matokeo ya `number // 10` (mgawanyiko wa nambari kamili) pamoja na moduli ya `number` na `10`.

# --hints--

Unapaswa kuwa na `number = number // 10 + number % 10` ndani ya tamko la `if`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
