---
id: 656878585631369a6b2d2191
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Kufafanua tafsiri yenyewe hakutafsiri mfuatano wa herufi. Njia ya `translate` lazima iitwe kwenye mfuatano wa herufi unaotakiwa kutafsiriwa kwa kutumia jedwali la tafsiri kama hoja:

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

Tengeneza kigezo kinachoitwa `translated_card_number` na weka thamani yake kuwa matokeo ya wito wa njia ya `translate` kwenye `card_number` ukiwa na `card_translation` kama hoja.

# --hints--

Unapaswa kutengeneza kigezo cha `translated_card_number` ndani ya `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

Unapaswa kuweka `translated_card_number` thamani ya `card_number.translate(card_translation)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
