---
id: 65687b48f2201ab32e06c37c
title: Hatua 21
challengeType: 20
dashedName: step-21
---

# --description--

Kwa sasa, skiripti yako inatupa `TypeError` kwa sababu unajaribu kuongeza mfuatano wa herufi kwenye nambari kamili. Unaweza kurekebisha hili kwa kubadilisha thamani ya `digit` kuwa nambari kamili kabla ya kuiongeza kwenye `sum_of_odd_digits`, ukitumia kitendakazi iliyo ndani cha `int`:

```python
my_string = '123'
my_int = int(my_string)
```

Badilisha thamani ya `digit` kuwa nambari kamili kabla ya kuiongeza kwenye `sum_of_odd_digits`. Kisha, hamisha wito wa `print` mwishoni mwa kitendakazi cha `verify_card_number` ili kuchapisha thamani ya `sum_of_odd_digits`.

# --hints--

Unapaswa kuwa na `sum_of_odd_digits += int(digit)` ndani ya mzunguko wa `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

Unapaswa kuwa na `print(sum_of_odd_digits)` mwishoni mwa kitendakazi cha `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

Haupaswi kuwa na `print(digit)` tena.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
