---
id: 656879c6f35a85a59c06b3a7
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Umefikiliwa vipengele (herufi) vya mfuatano wa herufi hapo awali, ukitumia kielezo cha kielezo `[]`. Pia unaweza kutumia kielezo cha kielezo kufikilia safu ya herufi katika mfuatano wa herufi kwa `string[start:stop:step]`:

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

Ambapo `start` ni kielezo cha kuanzia (kinachojumuisha), `stop` ni kielezo cha mwisho (kisichojumuisha), na `step` ni idadi ya herufi za kuruka. Ikiwa haijatajwa, `step` ni chaguo-msingi kuwa 1.

Tengeneza kigezo kinachoitwa `card_number_reversed` na weka thamani ya herufi 4 za kwanza za `card_number`.

# --hints--

Unapaswa kuwa na `card_number_reversed = card_number[0:4]` ndani ya kitendakazi cha `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
