---
id: 656877c202f44995ed952608
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Badilisha usemi wa `pass` na kigezo kinachoitwa `card_number` na thamani ya `'4111-1111-4555-1142'`.

# --hints--

Unapaswa kuwa na `card_number = '4111-1111-4555-1142'` ndani ya kitendakazi cha `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

Haupasi kuwa na `pass` katika msimbo wako.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
