---
id: 656879a66338b2a461d5d307
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

Algoriti ya Luhn ni kama ifuatavyo:

1. Kuanzia kulia kwenda kushoto, zidisha mara mbili thamani ya kila nambari ya pili; ikiwa bidhaa ni kubwa zaidi ya 9, jumlisha nambari za bidhaa hizo.
2. Chukua jumla ya nambari zote.
3. Ikiwa jumla ya nambari zote ni mara kumi, basi nambari hiyo ni halali; vinginevyo si halali.

Chukulia mfano wa nambari ya akaunti "7992739871" ambayo itakuwa na nambari ya ukaguzi iliyoongezwa, ikifanya iwe katika muundo wa 7992739871x:

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

Badilisha tamko la `pass` na kigezo `sum_of_odd_digits` na thamani ya `0`.

# --hints--

Unapaswa kuwa na `sum_of_odd_digits = 0` ndani ya kitendakazi cha `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

Hupaswi kuwa na `pass` ndani ya kitendakazi cha `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
