---
id: 6569b831a07d7154c793301b
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Tafsiri kitendakazi `verify_card_number` chenye kigezo `card_number`, na tumia neno kuu `pass` kufanya kitendakazi hicho kisifanye chochote.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `verify_card_number` chenye kigezo kinachoitwa `card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_parameters } = verify_card_number;
        assert.equal(function_parameters.trim(), "card_number");
    }
})
```

Unapaswa kutumia neno kuu `pass` ndani ya mwili wa kitendakazi `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_body } = verify_card_number;
        assert.match(function_body, / +pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

main()
```
