---
id: 67116d7584d0b469b14579c2
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Um eine Liste von Buchzusammenfassungen zu erhalten, können Sie die `about`-Eigenschaft jedes Buchs verwenden.

Erstellen Sie eine `getBookSummaries`-Funktion mit einem einzelnen Parameter, die ein Array mit Buchobjekten entgegennimmt.

Die `getBookSummaries`-Funktion soll ein Array von Strings zurückgeben, die die Zusammenfassung für jedes Buch darstellen.

# --hints--

Sie sollten eine `getBookSummaries`-Funktion erstellen.

```js
assert.isFunction(getBookSummaries);
```

Ihre `getBookSummaries`-Funktion sollte einen Parameter haben.

```js
assert.lengthOf(getBookSummaries, 1);
```

Ihre `getBookSummaries`-Funktion sollte ein Array zurückgeben.

```js
assert.isArray(getBookSummaries(library));
```

Ihre `getBookSummaries`-Funktion sollte ein Array mit der gleichen Länge wie das übergebene Array zurückgeben.

```js
assert.lengthOf(getBookSummaries(library), 8);

const _differentLibrary = [
  {
    title: 'Title A',
    author: 'Author A',
    about: 'About A',
    pages: 320,
  },
  {
    title: 'Title B',
    author: 'Author B',
    about: 'About B',
    pages: 320,
  },
  {
    title: 'Title C',
    author: 'Author C',
    about: 'About C',
    pages: 304,
  },
];
assert.lengthOf(getBookSummaries(_differentLibrary), 3);
```

Ihre `getBookSummaries`-Funktion sollte ein Array mit der `about`-Eigenschaft jedes Buchs aus dem übergebenen Array zurückgeben.

```js
const summaries = getBookSummaries(library);
library.forEach((book) => {
  assert.include(summaries, book.about);
});

const _differentLibrary = [
  {
    title: 'Title A',
    author: 'Author A',
    about: 'About A',
    pages: 320,
  },
  {
    title: 'Title B',
    author: 'Author B',
    about: 'About B',
    pages: 320,
  },
];
const differentSummaries = getBookSummaries(_differentLibrary);
_differentLibrary.forEach((book) => {
  assert.include(differentSummaries, book.about);
});
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\nList of book summaries:\n");

--fcc-editable-region--

--fcc-editable-region--
```
