---
id: 67116d7584d0b469b14579bf
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Erstellen Sie innerhalb des `library`-Arrays ein Objekt mit den folgenden Eigenschaften und Werten:

| Eigenschaft | Wert   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

Sie sollten ein Objekt innerhalb des `library`-Arrays haben.

```js
assert.isObject(library[0])
```

Ihr Objekt sollte eine `title`-Eigenschaft besitzen.

```js
assert.property(library[0], "title")
```

Der Wert Ihrer `title`-Eigenschaft sollte `"Your Next Five Moves: Master the Art of Business Strategy"` sein.

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

Ihr Objekt sollte eine `author`-Eigenschaft besitzen.

```js
assert.property(library[0], "author")
```

Ihre `author`-Eigenschaft sollte den Wert `"Patrick Bet-David and Greg Dinkin"` haben.

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

Ihr Objekt sollte eine `about`-Eigenschaft besitzen.

```js
assert.property(library[0], "about")
```

Ihre `about`-Eigenschaft sollte den Wert `"A book on how to plan ahead"` haben.

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

Ihr Objekt sollte eine `pages`-Eigenschaft besitzen.

```js
assert.property(library[0], "pages")
```

Der Wert Ihrer `pages`-Eigenschaft sollte eine Zahl sein.

```js
assert.isNumber(library[0]?.pages)
```

Ihre `pages`-Eigenschaft sollte auf `320` gesetzt sein.

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
