---
id: 67116d7584d0b469b14579c0
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Erstellen Sie ein weiteres Objekt innerhalb des `library`-Arrays mit den folgenden Eigenschaften und Werten:

| Eigenschaft | Wert   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


Sie sollten ein Objekt als zweiten Eintrag Ihres `library`-Arrays haben.

```js
assert.isObject(library[1]);
```

Ihr Objekt sollte eine `title`-Eigenschaft besitzen.

```js
assert.property(library[1], "title")
```

Die `title`-Eigenschaft sollte den Wert `"Atomic Habits"` haben.

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

Ihr Objekt sollte eine `author`-Eigenschaft besitzen.

```js
assert.property(library[1], "author")
```

Die `author`-Eigenschaft sollte den Wert `"James Clear"` haben.

```js
assert.propertyVal(library[1], "author", "James Clear");
```

Ihr Objekt sollte eine `about`-Eigenschaft besitzen.

```js
assert.property(library[1], "about")
```

Die `about`-Eigenschaft sollte den Wert `"A practical book about discarding bad habits and building good ones"` haben.

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

Ihr Objekt sollte eine `pages`-Eigenschaft besitzen.

```js
assert.property(library[1], "pages")
```

Der Wert der `pages`-Eigenschaft sollte eine Zahl sein.

```js
assert.isNumber(library[1]?.pages)
```

Die `pages`-Eigenschaft sollte auf `320` gesetzt sein.

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
