---
id: 67116c175a77cb64c5c10d49
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man mit Higher-Order-Funktionen wie `map`, `filter` und `reduce` arbeitet.

In diesem Workshop bauen Sie eine Bibliotheksverwaltungs-App, die Ihnen die Möglichkeit gibt, den Umgang mit verschiedenen Higher-Order-Funktionen zu üben.

Um zu beginnen, erstellen Sie eine Variable namens `library` und weisen Sie ihr ein leeres Array zu.

# --hints--

Sie sollten ein `library`-Array erstellen.

```js
assert.isArray(library)
```

Ihr `library`-Array sollte leer sein.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
