---
id: 69008059d43d6751385883c1
title: معالجة أخطاء معاملات التزايد والتقليل في تطبيق به أخطاء برمجية
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

طلب منك صديقك مساعدته في معالجة الأخطاء في شفرتهم. هم يتدربون على معاملات التزايد والتقليل لكنهم يحصلون على نتائج غير متوقعة في وحدة التحكم.

مهمتك هي مساعدتهم في تحديد الأخطاء وتحديث الشفرة بحيث تتطابق مخرجات وحدة التحكم مع النتائج المتوقعة.

يمكنك الرجوع إلى الدروس السابقة حول معاملات التزايد والتقليل وكذلك تدوين البادئة واللاحقة.

**قصص المستخدم:**

1. يجب أن يكون لديك متغير باسم `updatedCounter` ينتج القيمة `11` بشكل صحيح باستخدام معامل التزايد. حدّث التعبير لاستخدام التدوين الصحيح للبادئة أو اللاحقة بدلاً من ترميز القيمة مباشرة.
2. يجب أن يكون لديك متغير باسم `finalScore` ينتج القيمة `8` بشكل صحيح باستخدام معامل التزايد. حدّث التعبير لاستخدام التدوين الصحيح للبادئة أو اللاحقة بدلاً من ترميز القيمة مباشرة.
3. يجب أن يكون لديك متغير باسم `updatedCoins` ينتج القيمة `2` بشكل صحيح باستخدام معامل التقليل / معامل التخفيض. حدّث التعبير لاستخدام التدوين الصحيح للبادئة أو اللاحقة بدلاً من ترميز القيمة مباشرة.
4. يجب أن يكون لديك متغير باسم `newHealth` ينتج القيمة `7` بشكل صحيح باستخدام معامل التقليل / معامل التخفيض. حدّث التعبير لاستخدام التدوين الصحيح للبادئة أو اللاحقة بدلاً من ترميز القيمة مباشرة.

# --hints--

يجب أن يكون لديك متغير باسم `updatedCounter` يستخدم تدوين التزايد بالبادئة (`++counter`).

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

يجب أن يساوي `updatedCounter` القيمة `11`.

```js
assert.strictEqual(updatedCounter, 11);
```

يجب أن يكون لديك متغير باسم `finalScore` يستخدم تدوين التزايد باللاحقة (`score++`).

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

يجب أن يساوي `finalScore` القيمة `8`.

```js
assert.strictEqual(finalScore, 8);
```

يجب أن يكون لديك متغير باسم `updatedCoins` يستخدم تدوين التقليل / التخفيض بالبادئة (`--coins`).

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

يجب أن يساوي `updatedCoins` القيمة `2`.

```js
assert.strictEqual(updatedCoins, 2);
```

يجب أن يكون لديك متغير باسم `newHealth` يستخدم تدوين التقليل / التخفيض باللاحقة (`health--`).

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

يجب أن يساوي `newHealth` القيمة `7`.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
