---
id: 587d7dbe367417b2b2512bb9
title: استخدام @for لإنشاء حلقة Sass
challengeType: 0
forumTopicId: 301462
dashedName: use-for-to-create-a-sass-loop
---

# --description--

تضيف توجيه `@for` أنماطًا في حلقة، مشابهة جدًا لحلقة `for` في JavaScript.

يُستخدم `@for` بطريقتين: "من البداية حتى النهاية" أو "من البداية إلى النهاية". الفرق الرئيسي هو أن "من البداية **إلى** النهاية" *يستبعد* الرقم النهائي من العد، و"من البداية **حتى** النهاية" *يتضمن* الرقم النهائي ضمن العد.

إليك مثالًا على "من البداية **حتى** النهاية":

```scss
@for $i from 1 through 12 {
  .col-#{$i} { width: 100%/12 * $i; }
}
```

جزء `#{$i}` هو الصيغة النحوية لدمج متغير (`i`) مع نص لصنع سلسلة نصية. عندما يُحوّل ملف Sass إلى CSS، يظهر بهذا الشكل:

```scss
.col-1 {
  width: 8.33333%;
}

.col-2 {
  width: 16.66667%;
}

...

.col-12 {
  width: 100%;
}
```

هذه طريقة قوية لإنشاء تخطيط شبكة. الآن لديك اثنا عشر خيارًا لأعرض الأعمدة متاحة كـ CSS classes.

# --instructions--

اكتب توجيه `@for` يأخذ متغير `$j` يتدرج من 1 **إلى** 6.

يجب أن ينشئ 5 أصناف تسمى `.text-1` إلى `.text-5` حيث يكون لكل منها `font-size` مضروبًا في 15px حسب الفهرس.

# --hints--

يجب أن يستخدم شفرتك توجيه `@for`.

```js
assert(code.match(/@for /g));
```

يجب أن يكون لـ `.text-1` حجم خط `font-size` مقداره 15px.

```js
assert($('.text-1').css('font-size') == '15px');
```

يجب أن يكون لـ `.text-2` حجم خط `font-size` مقداره 30px.

```js
assert($('.text-2').css('font-size') == '30px');
```

يجب أن يكون لـ `.text-3` حجم خط `font-size` مقداره 45px.

```js
assert($('.text-3').css('font-size') == '45px');
```

يجب أن يكون لـ `.text-4` حجم خط `font-size` مقداره 60px.

```js
assert($('.text-4').css('font-size') == '60px');
```

يجب أن يكون لـ `.text-5` حجم خط `font-size` مقداره 75px.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>

@for $i from 1 through 5 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

---

```html
<style type='text/scss'>

@for $i from 1 to 6 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
