---
id: 587d7dbf367417b2b2512bbb
title: تطبيق نمط حتى يتحقق شرط باستخدام @while
challengeType: 0
forumTopicId: 301454
dashedName: apply-a-style-until-a-condition-is-met-with-while
---

# --description--

تُعد توجيه `@while` خيارًا بوظيفة مشابهة لحلقة `while` في JavaScript. ينشئ قواعد CSS حتى يتحقق شرط معين.

قدم تحدي `@for` مثالًا لإنشاء نظام شبكة بسيط. يمكن أن يعمل هذا أيضًا مع `@while`.

```scss
$x: 1;
@while $x < 13 {
  .col-#{$x} { width: 100%/12 * $x;}
  $x: $x + 1;
}
```

أولًا، عرّف متغيرًا `$x` وضع قيمته 1. بعد ذلك، استخدم توجيه `@while` لإنشاء نظام الشبكة *طالما* أن `$x` أقل من 13. بعد تعيين قاعدة CSS للخاصية `width`، زد قيمة `$x` بمقدار 1 لتجنب حلقة لا نهائية.

# --instructions--

استخدم `@while` لإنشاء سلسلة من الأصناف ذات `font-sizes` مختلفة.

يجب أن يكون هناك 5 أصناف مختلفة من `text-1` إلى `text-5`. ثم عيّن `font-size` إلى `15px` مضروبًا في رقم الفهرس الحالي. تأكد من تجنب الحلقة اللانهائية!

# --hints--

يجب أن يستخدم شفرتك توجيه `@while`.

```js
assert(code.match(/@while /g));
```

يجب أن تستخدم شفرتك متغير فهرس يبدأ من الفهرس 1.

```js
assert(code.match(/\$.*:\s*?1;/gi));
```

يجب أن تزيد شفرتك قيمة متغير العداد.

```js
assert(code.match(/\$(.*)\s*?:\s*\$\1\s*\+\s*1\s*;/gi));
```

يجب أن يحتوي الصنف `.text-1` على `font-size` بقيمة `15px`.

```js
assert($('.text-1').css('font-size') == '15px');
```

يجب أن يحتوي الصنف `.text-2` على `font-size` بقيمة `30px`.

```js
assert($('.text-2').css('font-size') == '30px');
```

يجب أن يحتوي الصنف `.text-3` على `font-size` بقيمة `45px`.

```js
assert($('.text-3').css('font-size') == '45px');
```

يجب أن يحتوي الصنف `.text-4` على `font-size` بقيمة `60px`.

```js
assert($('.text-4').css('font-size') == '60px');
```

يجب أن يحتوي الصنف `.text-5` على `font-size` بقيمة `75px`.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>
  $x: 1;
  @while $x < 6 {
    .text-#{$x}{
      font-size: 15px * $x;
    }
    $x: $x + 1;
  }
</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
