---
id: 587d7dbd367417b2b2512bb5
title: التسكين في CSS باستخدام Sass
challengeType: 0
forumTopicId: 301457
dashedName: nest-css-with-sass
---

# --description--

يتيح Sass التسكين لقواعد CSS، وهي طريقة مفيدة لتنظيم ورقة الأنماط.

عادةً، يُستهدف كل عنصر في سطر مختلف لتنسيقه، كما يلي:

```scss
article {
  height: 200px;
}

article p {
  color: white;
}

article ul {
  color: blue;
}
```

في مشروع كبير، يحتوي ملف CSS على العديد من الأسطر والقواعد. هنا يساعد التسكين في تنظيم الشفرة بوضع قواعد التنسيق الفرعية داخل العناصر الأصلية الخاصة بها:

```scss
article {
  height: 200px;

  p {
    color: white;
  }

  ul {
    color: blue;
  }
}

```

# --instructions--

استخدم تقنية التسكين الموضحة أعلاه لإعادة تنظيم قواعد CSS لكلا الطفلين في عنصر `.blog-post`. لأغراض الاختبار، يجب أن يأتي `h1` قبل عنصر `p`.

# --hints--

يجب أن تعيد الشفرة تنظيم قواعد CSS بحيث يكون `h1` و `p` مسكنين داخل عنصر `.blog-post` الأصل.

```js
assert(
  code.match(
    /\.blog-post\s*?{\s*?h1\s*?{\s*?text-align:\s*?center;\s*?color:\s*?blue;\s*?}\s*?p\s*?{\s*?font-size:\s*?20px;\s*?}\s*?}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  .blog-post {

  }
  h1 {
    text-align: center;
    color: blue;
  }
  p {
    font-size: 20px;
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  .blog-post {
    h1 {
      text-align: center;
      color: blue;
    }
    p {
      font-size: 20px;
    }
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```
