---
id: 587d7dbd367417b2b2512bb6
title: إنشاء CSS قابلة لإعادة الاستخدام باستخدام الخلائط
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

في Sass، <dfn>الخليط</dfn> هو مجموعة من إعلانات CSS يمكن إعادة استخدامها في جميع أنحاء ورقة الأنماط. يبدأ التعريف بقاعدة عند `@mixin`، تليها اسم مخصص. تطبق الخليط باستخدام قاعدة عند `@include`.

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

يُحوَّل إلى:

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

يمكن لخليطك أيضًا أن يأخذ معلمات، مما يسمح بتخصيص سلوكه. المعلمات مطلوبة عند استخدام الخليط.

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

يمكنك جعل المعلمات اختيارية بإعطاء المعلمات قيمًا افتراضية.

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

اكتب خليطًا باسم `shape` وأعطه 3 معلمات: `$w`، `$h`، و`$bg-color`.

استخدم خليط `shape` لإعطاء العنصر `#square` عرضًا وارتفاعًا بقيمة `50px`، ولون خلفية `red`. للعنصر `#rect-a` أضف عرضًا `100px`، وارتفاعًا `50px`، ولون خلفية `blue`. وأخيرًا، للعنصر `#rect-b` أضف عرضًا `50px`، وارتفاعًا `100px`، ولون خلفية `orange`.

# --hints--

يجب إعلان خليط باسم `shape` بثلاث معلمات: `$w`، `$h`، و`$bg-color`.

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

يجب أن يتضمن خليطك خاصية `width` تستخدم المعلمة `$w`.

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

يجب أن يتضمن خليطك خاصية `height` تستخدم المعلمة `$h`.

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

يجب أن يتضمن خليطك خاصية `background-color` تستخدم المعلمة `$bg-color`.

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

يجب استبدال الأنماط داخل المحدد `#square` باستدعاء لخليط `shape` باستخدام كلمة `@include`. ضبط العرض والارتفاع إلى `50px`، ولون الخلفية إلى `red`.

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

يجب استبدال الأنماط داخل المحدد `#rect-a` باستدعاء لخليط `shape` باستخدام كلمة `@include`. ضبط العرض إلى `100px`، والارتفاع إلى `50px`، ولون الخلفية إلى `blue`.

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

يجب استبدال الأنماط داخل المحدد `#rect-b` باستدعاء لخليط `shape` باستخدام كلمة `@include`. ضبط العرض إلى `50px`، والارتفاع إلى `100px`، ولون الخلفية إلى `orange`.

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
