---
id: 587d7dbf367417b2b2512bbc
title: قسم أنماطك إلى أجزاء أصغر باستخدام الأجزاء الجزئية
challengeType: 0
forumTopicId: 301459
dashedName: split-your-styles-into-smaller-chunks-with-partials
---

# --description--

<dfn>الأجزاء الجزئية</dfn> في Sass هي ملفات منفصلة تحتوي على مقاطع من شفرة CSS. تُستورد وتُستخدم في ملفات Sass أخرى. هذه طريقة رائعة لتجميع الشفرة المتشابهة في وحدة واحدة للحفاظ على التنظيم.

تبدأ أسماء الأجزاء الجزئية بالحرف السفلية (`_`)، مما يخبر Sass أنها مقطع صغير من CSS ولا يحولها إلى ملف CSS. كما تنتهي ملفات Sass بالامتداد `.scss`. لاستيراد الشفرة في الجزء الجزئي إلى ملف Sass آخر، استخدم التوجيه `@import`.

على سبيل المثال، إذا كانت كل الخلائط محفوظة في جزء جزئي اسمه "\_mixins.scss"، وتحتاجها في ملف "main.scss"، فهذه هي الطريقة لاستخدامها في الملف الرئيسي:

```scss
@import 'mixins'
```

لاحظ أن الحرف السفلية وامتداد الملف غير مطلوبين في عبارة `import` — Sass يفهم أنها جزء جزئي. بمجرد استيراد جزء جزئي إلى ملف، تصبح كل المتغيرات والخلائط والشفرة الأخرى متاحة للاستخدام.

# --instructions--

اكتب عبارة `@import` لاستيراد جزء جزئي اسمه `_variables.scss` إلى ملف main.scss.

# --hints--

يجب أن تستخدم شفرتك التوجيه `@import`، ولا يجب أن تتضمن الحرف السفلية في اسم الملف.

```js
assert(code.match(/@import\s+?('|")variables\1/gi));
```

# --seed--

## --seed-contents--

```html
<!-- The main.scss file -->
```

# --solutions--

```html
@import 'variables'
```
