---
id: 587d7dbf367417b2b2512bbc
title: قسم أنماطك إلى أجزاء أصغر باستخدام الأجزاء الجزئية
challengeType: 0
forumTopicId: 301459
dashedName: split-your-styles-into-smaller-chunks-with-partials
---

# --description--

<dfn>الملفات الجزئية</dfn> في Sass هي ملفات منفصلة تحتوي على أجزاء من كود CSS. تُستورد وتُستخدم في ملفات Sass أخرى. هذه طريقة رائعة لتجميع الكود المتشابه في وحدة واحدة للحفاظ على التنظيم.

تبدأ أسماء الملفات الجزئية بالحرف السفلي (`_`)، مما يخبر Sass أنها جزء صغير من CSS ولا يحولها إلى ملف CSS. كما تنتهي ملفات Sass بالامتداد `.scss`. لاستيراد الكود الموجود في الملف الجزئي إلى ملف Sass آخر، استخدم التوجيه `@import`.

على سبيل المثال، إذا كانت كل الخلائط محفوظة في جزء جزئي اسمه "\_mixins.scss"، وتحتاجها في ملف "main.scss"، فهذه هي الطريقة لاستخدامها في الملف الرئيسي:

```scss
@import 'mixins'
```

لاحظ أن الحرف السفلي وامتداد الملف غير مطلوبين في جملة `import` — يفهم Sass أنها ملف جزئي. بمجرد استيراد ملف جزئي إلى ملف، تصبح كل المتغيرات، والخليطات، والكود الآخر متاحة للاستخدام.

# --instructions--

اكتب عبارة `@import` لاستيراد جزء جزئي اسمه `_variables.scss` إلى ملف main.scss.

# --hints--

يجب أن يستخدم كودك التوجيه `@import`، ولا يجب أن يتضمن الحرف السفلي في اسم الملف.

```js
assert(code.match(/@import\s+?('|")variables\1/gi));
```

# --seed--

## --seed-contents--

```html
<!-- The main.scss file -->
```

# --solutions--

```html
@import 'variables'
```
