---
id: 587d7dbd367417b2b2512bb4
title: تخزين البيانات باستخدام متغيرات Sass
challengeType: 0
forumTopicId: 301460
dashedName: store-data-with-sass-variables
---

# --description--

ميزة واحدة في Sass تختلف عن CSS هي استخدام المتغيرات. تُعلن وتُحدد لتخزين البيانات، مشابهة لـ JavaScript.

في JavaScript، تُعرف المتغيرات باستخدام الكلمات المفتاحية `let` و `const`. في Sass، تبدأ المتغيرات بعلامة `$` تليها اسم المتغير.

إليك بعض الأمثلة:

```scss
$main-fonts: Arial, sans-serif;
$headings-color: green;
```

ولاستخدام المتغيرات:

```scss
h1 {
  font-family: $main-fonts;
  color: $headings-color;
}
```

مثال على فائدة المتغيرات هو عندما تحتاج عدة عناصر إلى نفس اللون. إذا تغير هذا اللون، يكون مكان تعديل الشفرة الوحيد هو قيمة المتغير.

# --instructions--

أنشئ متغيرًا `$text-color` وحدد قيمته إلى `red`. ثم غيّر قيمة خاصية `color` لكل من `.blog-post` و `h2` إلى المتغير `$text-color`.

# --hints--

يجب أن يحتوي شفرتك على إعلان متغير Sass باسم `$text-color` بقيمة `red`.

```js
assert(code.match(/\$text-color\s*:\s*?red\s*;/g));
```

يجب أن تستخدم شفرتك المتغير `$text-color` لتغيير خاصية `color` لعناصر `.blog-post` و `h2`.

```js
assert(code.match(/color\s*:\s*\$text-color\s*;?/g));
```

يجب أن يكون لعنصر `.blog-post` خاصية `color` بقيمة red.

```js
assert($('.blog-post').css('color') == 'rgb(255, 0, 0)');
```

يجب أن تكون لعناصر `h2` خاصية `color` بقيمة red.

```js
assert($('h2').css('color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>


  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: red;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  $text-color: red;

  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: $text-color;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```
