---
id: 587d7dbe367417b2b2512bb8
title: استخدام @if و @else لإضافة منطق إلى أنماطك
challengeType: 0
forumTopicId: 301463
dashedName: use-if-and-else-to-add-logic-to-your-styles
---

# --description--

تُستخدم توجيهات `@if` في Sass لاختبار حالة محددة - تعمل تمامًا مثل جملة `if` في JavaScript.

```scss
@mixin make-bold($bool) {
  @if $bool == true {
    font-weight: bold;
  }
}
```

وكما في JavaScript، تختبر توجيهات `@else if` و `@else` حالات إضافية:

```scss
@mixin text-effect($val) {
  @if $val == danger {
    color: red;
  }
  @else if $val == alert {
    color: yellow;
  }
  @else if $val == success {
    color: green;
  }
  @else {
    color: black;
  }
}
```

# --instructions--

أنشئ خليطًا يسمى `border-stroke` يأخذ معلمة `$val`. يجب أن يتحقق الخليط من الحالات التالية باستخدام توجيهات `@if` و `@else if` و `@else`:

```scss
light - 1px solid black
medium - 3px solid black
heavy - 6px solid black
```

إذا لم تكن قيمة المعلمة `$val` هي `light` أو `medium` أو `heavy`، يجب تعيين خاصية `border` إلى `none`.

# --hints--

يجب أن تعلن شفرتك عن خليط باسم `border-stroke` يحتوي على معلمة باسم `$val`.

```js
assert(code.match(/@mixin\s+?border-stroke\s*?\(\s*?\$val\s*?\)\s*?{/gi));
```

يجب أن يحتوي الخليط على جملة `@if` للتحقق مما إذا كانت `$val` تساوي `light`، ولتعيين خاصية `border` إلى `1px solid black`.

```js
assert(
  code.match(
    /@if\s+?\$val\s*?===?\s*?light\s*?{\s*?border\s*?:\s*?1px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

يجب أن يحتوي الخليط على جملة `@else if` للتحقق مما إذا كانت `$val` تساوي `medium`، ولتعيين خاصية `border` إلى `3px solid black`.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?medium\s*?{\s*?border\s*?:\s*?3px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

يجب أن يحتوي الخليط على جملة `@else if` للتحقق مما إذا كانت `$val` تساوي `heavy`، ولتعيين خاصية `border` إلى `6px solid black`.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?heavy\s*?{\s*?border\s*?:\s*?6px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

يجب أن يحتوي الخليط على جملة `@else` لتعيين خاصية `border` إلى `none`.

```js
assert(code.match(/@else\s*?{\s*?border\s*?:\s*?none\s*?;\s*?}/gi));
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```

# --solutions--

```html
<style type='text/scss'>
  @mixin border-stroke($val) {
    @if $val == light {
      border: 1px solid black;
    }
    @else if $val == medium {
      border: 3px solid black;
    }
    @else if $val == heavy {
      border: 6px solid black;
    }
    @else {
      border: none;
    }
  }


  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```
