---
id: 587d7fa5367417b2b2512bbd
title: تمديد مجموعة واحدة من أنماط CSS إلى عنصر آخر
challengeType: 0
forumTopicId: 301456
dashedName: extend-one-set-of-css-styles-to-another-element
---

# --description--

يحتوي Sass على ميزة تسمى `extend` تسهل استعارة قواعد CSS من عنصر وبناؤها على عنصر آخر.

على سبيل المثال، كتلة قواعد CSS أدناه تُنسق صنف `.panel`. لديها `background-color` و`height` و`border`.

```scss
.panel{
  background-color: red;
  height: 70px;
  border: 2px solid green;
}
```

الآن تريد لوحة أخرى تسمى `.big-panel`. لها نفس الخصائص الأساسية مثل `.panel`، لكنها تحتاج أيضًا إلى `width` و`font-size`. يمكن نسخ ولصق قواعد CSS الأولية من `.panel`، لكن الشفرة تصبح متكررة مع إضافة أنواع أكثر من اللوحات. توجيه `extend` هو طريقة بسيطة لإعادة استخدام القواعد المكتوبة لعنصر واحد، ثم إضافة المزيد لعنصر آخر:

```scss
.big-panel{
  @extend .panel;
  width: 150px;
  font-size: 2em;
}
```

سيكون لـ `.big-panel` نفس خصائص `.panel` بالإضافة إلى الأنماط الجديدة.

# --instructions--

أنشئ صنفًا `.info-important` يمدد `.info` ويحتوي أيضًا على `background-color` مضبوط على magenta.

# --hints--

يجب أن يحتوي صنف `info-important` على `background-color` مضبوط على `magenta`.

```js
assert(
  code.match(
    /\.info-important\s*?{[\s\S]*background-color\s*?:\s*?magenta\s*?;[\s\S]*}/gi
  )
);
```

يجب أن يستخدم صنف `info-important` توجيه `@extend` لوراثة التنسيق من صنف `info`.

```js
assert(
  code.match(/\.info-important\s*?{[\s\S]*@extend\s*?.info\s*?;[\s\S]*/gi)
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }




</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }
  .info-important{
    @extend .info;
    background-color: magenta;
  }



</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```
