---
id: 587d7dbf367417b2b2512bba
title: استخدام @each لتكرار العناصر في قائمة
challengeType: 0
forumTopicId: 301461
dashedName: use-each-to-map-over-items-in-a-list
---

# --description--

أظهر التحدٍّ الأخير كيف يستخدم التوجيه `@for` قيمة بداية وقيمة نهاية للتكرار عدد معين من المرات. كما يوفر Sass التوجيه `@each` الذي يكرر على كل عنصر في قائمة أو خريطة. في كل تكرار، يُعيَّن المتغير إلى القيمة الحالية من القائمة أو الخريطة.

```scss
@each $color in blue, red, green {
  .#{$color}-text {color: $color;}
}
```

لدى الخريطة صيغة نحوية مختلفة قليلاً. إليك مثالاً:

```scss
$colors: (color1: blue, color2: red, color3: green);

@each $key, $color in $colors {
  .#{$color}-text {color: $color;}
}
```

لاحظ أن المتغير `$key` ضروري للإشارة إلى المفاتيح في الخريطة. وإلا، ستحتوي CSS المجمعة على `color1`، `color2`... فيها. كلا المثالين السابقين من الشفرة يتحولان إلى CSS التالية:

```scss
.blue-text {
  color: blue;
}

.red-text {
  color: red;
}

.green-text {
  color: green;
}
```

# --instructions--

اكتب توجيه `@each` يمر عبر قائمة: `blue, black, red` ويُعيِّن كل متغير إلى صنف `.color-bg`، حيث يتغير جزء `color` لكل عنصر إلى اللون المناسب. يجب أن يضبط كل صنف خاصية `background-color` إلى اللون المناسب أيضًا.

# --hints--

يجب أن يستخدم شفرتك التوجيه `@each`.

```js
assert(code.match(/@each /g));
```

يجب أن يحتوي صنف `.blue-bg` على خاصية `background-color` باللون الأزرق.

```js
assert($('.blue-bg').css('background-color') == 'rgb(0, 0, 255)');
```

يجب أن يحتوي صنف `.black-bg` على خاصية `background-color` باللون الأسود.

```js
assert($('.black-bg').css('background-color') == 'rgb(0, 0, 0)');
```

يجب أن يحتوي صنف `.red-bg` على خاصية `background-color` باللون الأحمر.

```js
assert($('.red-bg').css('background-color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

# --solutions--

```html
<style type='text/scss'>

  @each $color in blue, black, red {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

---

```html
<style type='text/scss'>

  $colors: (color1: blue, color2: black, color3: red);

  @each $key, $color in $colors {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```
