---
id: 655cd899f8de09431eabb40c
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви вивчите структури даних, створюючи алгоритм сортування злиттям.

Це алгоритм сортування, який використовує принцип розділяй і володарюй для сортування колекцій даних. Тобто він «ділить» колекцію на менші підчастини, а потім «володарює» над ними, сортує їх окремо, а потім об’єднує відсортовані підчастини.

Створіть функцію з назвою `merge_sort`. Ця функція виконуватиме завдання сортування списку чисел.

Використайте ключове слово `pass` у тілі функції.

# --hints--

Ви повинні оголосити функцію з ім’ям `merge_sort`. Не забудьте використати ключове слово `pass` у тілі функції.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
