---
id: 6564a9fe51964c229d5b7f4c
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Вам потрібен параметр, який позначає колекцію даних для сортування. Створіть параметр з назвою `array` у функції `merge_sort`.

# --hints--

Ваша функція `merge_sort` має приймати один параметр: `array`.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
