---
id: 6564aee9c077774ea49c3faf
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Алгоритм сортування злиттям виконує головним чином три дії:

- Розділити несортовану послідовність елементів на підчастини
- Відсортувати елементи в підчастинах
- Об’єднати відсортовані підчастини

Вищезазначене відбувається рекурсивно, доки підчастини не об’єднаються у повну відсортовану послідовність. Почнемо з розділення послідовності.

Спочатку замініть ключове слово `pass` на змінну `middle_point`. Використайте оператор цілочисельного ділення (`//`) для ділення довжини списку `array` на 2 і призначте результат новій змінній `middle_point`. Не забудьте відступити код.

# --hints--

У тілі функції не повинно бути `pass`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

Потрібно оголосити змінну `middle_point` і призначити їй `len(array) // 2`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
