---
id: 656680b0fc79f2c38a34d90e
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Тепер настав час відсортувати та об’єднати списки (`left_part` і `right_part`) у початковий `array`.

Ви можете зробити це, порівнюючи елементи обох списків і додаючи менший елемент до основного списку. Ви будете робити це порівняння для всіх індексів у `left_part` і `right_part`.

Створіть три змінні: `left_array_index`, `right_array_index` і `sorted_index` та встановіть їхні значення в `0`. Ці змінні допоможуть вам відстежувати кожен індекс під час процесу сортування.

# --hints--

У вас має бути змінна з іменем `left_array_index` всередині вашої функції `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index/);
  }
})
```

Значення `left_array_index` має бути 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index\s*=\s*0/);
  }
})
```

У вас має бути змінна з іменем `right_array_index` всередині вашої функції `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index/);
  }
})
```

Значення `right_array_index` має бути 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index\s*=\s*0/);
  }
})

```

У вас має бути змінна з іменем `sorted_index` всередині вашої функції `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index/);
  }
})
```

Значення `sorted_index` має бути 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index\s*=\s*0/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)

--fcc-editable-region--

--fcc-editable-region--
```
