---
id: 656639771fed09b5c0e8fe71
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер, коли ви розділили список `array` на два окремі списки, продовжуйте ділити кожен список, доки кожен елемент не опиниться у власному списку. Список з одним числом завжди відсортований.

Для цього рекурсивно викликайте `merge_sort` у своїй функції. Поки що не передавайте жодних аргументів.

# --hints--

Ви повинні викликати `merge_sort` в кінці тіла функції.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
