---
id: 656627b47bd2d2a4afbc945d
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

У попередньому кроці ви отримали середню точку. Ви можете використати її, щоб розділити `array` на дві частини та призначити кожну частину новим змінним.

Використайте синтаксис slice, щоб витягти ліву половину `array` і призначити її змінній з ім’ям `left_part`.


# --hints--

У вашій функції `left_part` має бути змінна з ім’ям `merge_sort`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

Ви маєте призначити `array[:middle_point]` змінній `left_part`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
