---
id: 656710d1e0ec62253426db24
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Цикл `while`, який ви створили, порівнює один елемент із `left_part` з іншим у `right_part`, а потім додає менший елемент до основного списку `array`.

Він продовжуватиме цю операцію, доки не залишиться елементів для порівняння. Але в `left_part` можуть залишатися елементи, тоді як у `right_part` їх немає, і навпаки.

Створіть ще один цикл `while`, щоб скопіювати залишкові елементи з `left_part` у список `array`. Використовуйте `left_array_index < len(left_part)` для умови `while`.

# --hints--

Ваш цикл `while` має мати таку умову: `left_array_index < len(left_part)`. Не забудьте використати ключове слово `pass` у тілі циклу.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):/m);
  }
})
```

У тілі вашого циклу `pass` має бути ключове слово `while`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
