---
id: 657f59751d5a5c9b069ae0f3
title: Крок 28
challengeType: 20
dashedName: step-28
---

# --description--

Ви можете використати змінну `__name__`, щоб визначити, чи виконується Python-скрипт як основна програма, чи він імпортується як модуль (код, написаний в іншому Python-файлі).

Якщо значення `__name__` встановлено в `'__main__'`, це означає, що поточний скрипт є основною програмою, а не модулем.

У цьому проєкті ви використаєте поточний скрипт як основну програму.

Створіть інструкцію `if`, яка перевіряє, чи є значення `__name__` рівним `'__main__'`.

Використайте ключове слово `pass` у тілі інструкції `if`.

# --hints--

Ваша інструкція `if` має перевіряти, чи має `__name__` значення `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
