---
id: 5d822fd413a79914d39e98d2
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Verschachteln Sie ein `div` mit einer Klasse von `bb1` im Container für Hintergrundgebäude. Öffnen Sie Ihre `styles.css`-Datei und geben Sie `.bb1` eine `width` von `10%` und eine `height` von `70%`. „bb“ steht für „background building“; dies wird Ihr erstes Gebäude sein.

# --hints--

Sie sollten ein neues `div`-Element erstellen.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

Sie sollten dem neuen `div` eine Klasse von `bb1` geben.

```js
assert.exists(document.querySelector('div.bb1'));
```

Sie sollten einen `.bb1`-Klassen-Selektor verwenden, um das Element zu stylen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

Sie sollten dem `.bb1`-Element eine `width` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

Sie sollten dem `.bb1`-Element eine `height` von `70%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
    
```

