---
id: 5d822fd413a79914d39e98ca
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie unter dem `html` ein öffnendes und ein schließendes `DOCTYPE` Tag hinzu, damit Sie einen Platz haben, um Code einzufügen. Stellen Sie sicher, dass die Sprache auf Englisch eingestellt ist.

# --hints--

Ihre `DOCTYPE` Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ihr `html` Element sollte ein öffnendes Tag und ein `lang` Attribut mit dem Wert `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr `html` Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Ihre `html` Tags sollten in der richtigen Reihenfolge sein.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Sie sollten nur ein `html` Element haben.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
