---
id: 5d822fd413a79914d39e98cd
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

In CSS können Sie mit einem Sternchen alles ansprechen. Fügen Sie allen Elementen einen Rahmen hinzu, indem Sie den `*`-Selektor verwenden und ihm eine `border`-Eigenschaft mit dem Wert `1px solid black` geben. Das ist ein Trick, der hilft, zu sehen, wo sich Elemente befinden und wie groß sie sind. Sie werden das später wieder entfernen.

# --hints--

Sie sollten den `*`-Selektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Sie sollten die `border`-Eigenschaft verwenden, um alle Elemente zu gestalten.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Alle Elemente sollten einen `1px solid black`-Rahmen haben.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
