---
id: 5d822fd413a79914d39e98cb
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Fügen Sie als Nächstes öffnende und schließende `head`- und `body`-Tags innerhalb des `html`-Elements hinzu.

# --hints--

Sie sollten ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Sie sollten ein schließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Sie sollten ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Sie sollten ein schließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Die `head`- und `body`-Elemente sollten Geschwister sein.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

Das `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

Das `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
  
</html>
--fcc-editable-region--

```

