---
id: 5d822fd413a79914d39e98d7
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Um eine Variable zu verwenden, setzen Sie den Variablennamen in runde Klammern mit `var` davor, so: `var(--variable-name)`. Welchen Wert Sie der Variable auch immer gegeben haben, dieser wird auf jede Eigenschaft angewendet, bei der Sie sie verwenden.

Fügen Sie die Variable `--building-color1`, die Sie im vorherigen Schritt erstellt haben, als Wert der Eigenschaft `background-color` der Klasse `.bb1a` hinzu.

# --hints--

Die `background-color` des `.bb1a` Elements sollte gesetzt werden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a')?.backgroundColor)
```

Sie sollten `var(--building-color1)` verwenden, um die `background-color` des `.bb1a` Elements zu setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
