---
id: 5d822fd413a79914d39e98d4
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Geben Sie den Teilen Ihres Gebäudes `width` und `height` Eigenschaften mit diesen Werten: `70%` und `10%` zu `.bb1a`, `80%` und `10%` zu `.bb1b`, `90%` und `10%` zu `.bb1c` und `100%` und `70%` zu `.bb1d`. Denken Sie daran, dass diese Prozentsätze relativ zum übergeordneten Element sind, und beachten Sie, dass sich die Höhen zu 100 % addieren – wodurch der Container vertikal ausgefüllt wird.

# --hints--

Sie sollten einen Klassenselektor verwenden, um `.bb1a` zu stylen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Sie sollten `.bb1a` eine `width` von `70%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Sie sollten `.bb1a` eine `height` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Sie sollten einen Klassenselektor verwenden, um `.bb1b` zu stylen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Sie sollten `.bb1b` eine `width` von `80%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Sie sollten `.bb1b` eine `height` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Sie sollten einen Klassenselektor verwenden, um `.bb1c` zu stylen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Sie sollten `.bb1c` eine `width` von `90%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Sie sollten `.bb1c` eine `height` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Sie sollten einen Klassenselektor verwenden, um `.bb1d` zu stylen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Sie sollten `.bb1d` eine `width` von `100%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Sie sollten `.bb1d` eine `height` von `70%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

