---
id: 5d822fd413a79914d39e98ce
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Fügen Sie außerdem einen `box-sizing` von `border-box` zu allem hinzu. Dadurch wird verhindert, dass der von Ihnen hinzugefügte Rahmen die Größe Ihrer Elemente verändert.

# --hints--

Sie sollten die `box-sizing`-Eigenschaft verwenden.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

Sie sollten den vorhandenen `*`-Selektor nutzen.

```js
// Two selectors create two CSSStyleRule objects
assert.equal(new __helpers.CSSHelp(document).getStyleDeclarations('*').length, 1);
```

Alle Elemente sollten eine `box-sizing` von `border-box` haben.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```

