---
id: 6733b5e6f210c1bd1afdf8b3
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

剩餘的植物物件已為你添加。現在是時候創建你的目錄。

像一般的 JavaScript 物件，`Map` 物件包含一個鍵值對組的群集。要創建一個 `Map` 物件，你可以使用 `Map` 建構式：

```js
const myMap = new Map();
```

宣告一個 `catalog` 變數並指定一個 `Map` 物件。

# --hints--

你應該有一個名為 `catalog` 的變數。

```js
assert.exists(catalog);
```

你應該將一個 `Map` 物件指定到你的 `catalog` 變數。

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
