---
id: 673339d96c56d36e963e888e
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將建置一個植物苗圃目錄，同時練習 JavaScript 的 Map 和 Set 物件。

為了表述你目錄中的每一株植物，你將使用物件，每個物件有三個屬性：

- `commonName`，植物的通用名稱
- `scientificName`，植物的學名（或拉丁名）
- `cultivar`，植物的品種

首先建立一個具有上述三個屬性的物件，該物件是名為 `Ballerina` 的品種的 `Spanish lavender`（學名為 `Lavandula stoechas`），並將它指定給名為 `ballerina` 的變數。

# --hints--

你應該有一個名為 `ballerina` 的變數。

```js
assert.exists(ballerina);
```

你應該將一個物件指定到你的 `ballerina` 變數。

```js
assert.isObject(ballerina);
```

你的 `ballerina` 物件應該有一個 `commonName` 屬性設定為 `"Spanish lavender"`。

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

你的 `ballerina` 物件應該有一個 `scientificName` 屬性，設定為 `"Lavandula stoechas"`。

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

你的 `ballerina` 物件應該有一個 `cultivar` 屬性設定為 `"Ballerina"`。

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

你的 `ballerina` 物件應該有三個鍵。

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
