---
id: 6733b5632fb854bc6c58dc0d
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

現在創建另一個物件，具有相同的三個屬性，以表述另一種名為 `Pretty Polly` 的西班牙薰衣草品種，並將它指定給名為 `prettyPolly` 的變數。

# --hints--

你應該有一個名為 `prettyPolly` 的變數。

```js
assert.exists(prettyPolly);
```

你應該將一個物件指定到你的 `prettyPolly` 變數。

```js
assert.isObject(prettyPolly);
```

你的 `prettyPolly` 物件應該有一個 `commonName` 屬性設定為 `"Spanish lavender"`。

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

你的 `prettyPolly` 物件應該有一個 `scientificName` 屬性，設定為 `"Lavandula stoechas"`。

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

你的 `prettyPolly` 物件應該有一個 `cultivar` 屬性，設定為 `"Pretty Polly"`。

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

你的 `prettyPolly` 物件應該有三個鍵。

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
