---
id: 6734c10a327760a665b7d5b9
title: 步驟 13
challengeType: 1
dashedName: step-13
---

# --description--

在接下來的幾個步驟中，你將會撰寫一個函式（程式），讓你在銷售植物時更新目錄。

創建一個名為 `sellPlants` 的函式（程式），該函式接受三個引數：一個植物物件、花盆大小和要銷售的花盆數量，順序如上。

當指定的盆栽數量大於可用的盆栽數量時，讓你的函式傳回 `Not enough <pot-size> size pots for <item-name>. Only <pot-number> left.`。將 `<pot-size>` 替換為盆栽尺寸，將 `<pot-number>` 替換為該尺寸剩餘的盆栽數量，並將 `<item-name>` 替換為植物的學名，後面接一個空格及用單引號包住的栽培品種名稱。

為了測試一切是否正常，將使用 `ballerina` 物件、`"small"` 和 `25` 呼叫你的 `sellPlants` 函式（程式），並將結果記錄到主控台。 

# --hints--

你不應該修改你的 `catalog` 地圖。

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

你應該有一個名為 `sellPlants` 的函式（程式）。

```js
assert.isFunction(sellPlants);
```

你的 `sellPlants` 函式應該接受三個引數。

```js
assert.lengthOf(sellPlants, 3);
```

`sellPlants(ballerina, "small", 25)` 應該傳回 `Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.`

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

你應該將 `sellPlants(ballerina, "small", 25)` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)` 應該傳回 `undefined`。

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog)
```
