---
id: 6734a606e5968a5a19b77171
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

繼續將剩餘的植物添加到你的目錄中，就像你在前面的步驟中所做的那樣。

# --hints--

你應該將你的 `willowVale` 物件設為你的 `catalog` 映射的鍵。

```js
assert.isTrue(catalog.has(willowVale));
```

`catalog` 映射中 `willowVale` 鍵的值應該是一個具有 `small`、`medium` 和 `large` 屬性的物件，每個屬性的值都是正整數。

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你應該將你的 `hidcote` 物件設為你的 `catalog` 映射的鍵。

```js
assert.isTrue(catalog.has(hidcote));
```

`catalog` 映射中 `hidcote` 鍵的值應該是一個具有 `small`、`medium` 和 `large` 屬性的物件，每個屬性的值都是正整數。

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你應該將你的 `imperialGem` 物件設為你的 `catalog` 映射的鍵。

```js
assert.isTrue(catalog.has(imperialGem));
```

`catalog` 映射中 `imperialGem` 鍵的值應該是一個具有 `small`、`medium` 和 `large` 屬性的物件，每個屬性的值都是正整數。

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你應該將你的 `royalCrown` 物件設為你的 `catalog` 映射的鍵。

```js
assert.isTrue(catalog.has(royalCrown));
```

`catalog` 映射中 `royalCrown` 鍵的值應該是一個具有 `small`、`medium` 和 `large` 屬性的物件，每個屬性的值都是正整數。

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你應該仍然在你的 `catalog` 映射中擁有 `ballerina` 鍵。

```js
assert.isTrue(catalog.has(ballerina));
```

你仍然應該在你的 `catalog` 映射中保留 `prettyPolly` 鍵。

```js
assert.isTrue(catalog.has(prettyPolly));
```

你的 `catalog` 地圖應該包含六個元素。

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
