---
id: 6734a13a29acdf47f2b8e161
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

與一般物件不同，鍵必須始終是字串，map 物件允許你使用任何資料型別作為鍵。在你的目錄中，你將使用物件-物件的鍵值對組，將你的每個植物物件與一個包含不同尺寸可用花盆數量的物件關聯起來。

修改你的 `set` 呼叫以使用你的 `ballerina` 物件作為鍵。對於值，使用包含三個屬性 `small`、`medium` 和 `large` 的物件，每個屬性的值都是正整數。 

# --hints--

你應該將你的 `ballerina` 物件設為你的 `catalog` 映射的鍵。

```js
assert.isTrue(catalog.has(ballerina));
```

`catalog` 映射中 `ballerina` 鍵的值應該是一個具有 `small`、`medium` 和 `large` 屬性的物件。

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`catalog` 映射中與 `ballerina` 鍵相關的物件的 `small` 屬性應該是正整數。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`catalog` 映射中與 `ballerina` 鍵相關的物件的 `medium` 屬性應該是正整數。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`catalog` 映射中與 `ballerina` 鍵相關的物件的 `large` 屬性應該是正整數。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你的 `catalog` 地圖應該只包含一個鍵值配對。

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
