---
id: 6734a2385c52c54d916a8b86
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

為你的目錄添加另一個鍵值配對。使用你的 `prettyPolly` 變數作為鍵，並將值設為一個物件，該物件具有 `small`、`medium` 和 `large` 屬性，這些屬性的值再次為整數。

# --hints--

你應該將你的 `prettyPolly` 物件設為你的 `catalog` 映射的鍵。

```js
assert.isTrue(catalog.has(prettyPolly));
```

`catalog` 映射中 `prettyPolly` 鍵的值應該是一個具有 `small`、`medium` 和 `large` 屬性的物件。

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`catalog` 映射中與 `prettyPolly` 鍵相關的物件的 `small` 屬性應該是正整數。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`catalog` 映射中與 `prettyPolly` 鍵相關的物件的 `medium` 屬性應該是正整數。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`catalog` 映射中與 `prettyPolly` 鍵相關的物件的 `large` 屬性應該是正整數。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你應該仍然在你的 `catalog` 映射中擁有 `ballerina` 鍵。

```js
assert.isTrue(catalog.has(ballerina));
```

你的 `catalog` 地圖應該包含兩個元素。

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
