---
id: 67349ba230cff237c0f2bdf2
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

Map 物件會記住鍵被插入到 map 中的順序。要在你的 map 物件中設定一個鍵值配對，你可以使用 `set` 方法。以下範例將鍵 `"name"` 設定為值 `"Jerry"`，在 `myMap` map 物件中。

```js
myMap.set("name", "Jerry");
```

將你的 `catalog` 映射中的 `"ballerina"` 鍵設定為你的 `ballerina` 物件的值。然後，將你的 catalog 記錄到主控台。

請注意，要正確查看你的地圖物件，你必須打開瀏覽器主控台。

# --hints--

你應該將 `"ballerina"` 鍵設為你在 `catalog` 映射中的 `ballerina` 物件的值。

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

你的 `catalog` 地圖應該只包含一個鍵值配對。

```js
assert.lengthOf(catalog, 1);
```

你應該將你的 `catalog` 映射記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
