---
id: a26cbbe9ad8655a977e1ceb5
title: Tengeneza programu ya kutafuta neno refu zaidi
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

Katika maabara hii, utatengeneza kitendakazi kinachorudisha urefu wa neno refu zaidi katika sentensi iliyotolewa.

Kwa mfano, katika sentensi `"The quick brown fox jumped over the lazy dog"`, neno refu zaidi ni `"jumped"`, ambalo lina urefu wa 6.

**Lengo:** Timiza hadithi za mtumizi zilizo hapa chini na upite mitihani yote ili kumaliza maabara.

**Hadithi za mtumizi:**

1. Unapaswa kutengeneza kitendakazi kinachoitwa `findLongestWordLength` kinachopokea mfuatano wa herufi kama hoja.
2. Kitendakazi kinapaswa kurudisha urefu wa neno refu zaidi katika mfuatano wa herufi. 

# --hints--

Unapaswa kutengeneza kitendakazi kinachoitwa `findLongestWordLength`.

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength` inapaswa kuwa na kigezo kimoja.

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` inapaswa kurudisha nambari.

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` inapaswa kurudisha `6`.

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")` inapaswa kurudisha `5`.

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")` inapaswa kurudisha `6`.

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("Googling do a barrel roll")` inapaswa kurudisha `8`.

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` inapaswa kurudisha `8`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` inapaswa kurudisha `19`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
