---
id: 56533eb9ac21ba0edf2244b3
title: Construye un conversor de Celsius a Fahrenheit
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

En este laboratorio, escribirás una función que convierta la temperatura de Celsius a Fahrenheit. La fórmula para convertir de Celsius a Fahrenheit es:

```js
fahrenheit = celsius * (9/5) + 32
```

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `convertCtoF`.
2. `convertCtoF` debe tomar un solo argumento numérico, que es la temperatura en Celsius.
3. `convertCtoF` debe devolver un número.

# --hints--

Debes crear una función llamada `convertCtoF`.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` debe tomar un solo parámetro.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` debería devolver un número.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` debería devolver un valor de `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` debería devolver un valor de `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` debería devolver un valor de `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` debería devolver un valor de `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` debería devolver un valor de `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
