---
id: 67ec2f7924ade2450c06bcd8
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이전 수업에서 `useEffect` 훅을 사용해 데이터 가져오기 같은 부수 효과를 처리하는 방법을 배웠습니다. 이번 워크숍에서는 사용자 입력에 따라 API에서 과일 데이터를 가져와 결과를 동적으로 표시하는 애플리케이션을 만들 것입니다.

먼저, `FruitsSearch` 컴포넌트가 `div`의 `id`를 가진 `search-container` 요소를 반환하도록 만드세요.

# --hints--

`div` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("div"), 2);
```

`div` 요소에 `id`라는 `search-container`가 있어야 합니다.

```js
const el = document.getElementById("search-container");
assert.exists(el);
assert.strictEqual(el.tagName, "DIV");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

--fcc-editable-region--
export function FruitsSearch() {
  
}
--fcc-editable-region--
```
