---
id: 67f8ec758a740e34432222db
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

사용자가 입력하는 내용을 캡처하기 위해 input에 `onChange` 속성을 추가하세요. `e`을 인수로 받는 화살표 함수로 설정하세요. 함수 내부에서, 입력의 현재 값을 가진 `query`을 setter 함수에 전달하여 `e.target.value` 상태를 업데이트하세요. 예시는 다음과 같습니다:

```js
<input
  onChange={(e) => setExample(e.target.value)}
/>
```

# --hints--

`onChange` 핸들러가 정의되어 있는지 확인하세요.

```js
const searchInput = document.getElementById("search-input");
assert.exists(searchInput);

const reactKey = Object.keys(searchInput).find(key => key.startsWith("__reactProps"));
const reactProps = searchInput[reactKey];

assert.exists(reactProps);
assert.isFunction(reactProps.onChange);
```

`onChange` 핸들러는 `e` 매개변수를 가져야 합니다.

```js
const searchInput = document.getElementById("search-input");
assert.exists(searchInput);

const reactKey = Object.keys(searchInput).find(key => key.startsWith("__reactProps"));
const reactProps = searchInput[reactKey];

assert.match(reactProps.onChange.toString(), /onChange\s*\(\s*e\s*\)/);
```

`onChange` 핸들러는 `setQuery()`와 함께 `e.target.value`을 호출해야 합니다.

```js
const searchInput = document.getElementById("search-input");
assert.exists(searchInput);

const reactKey = Object.keys(searchInput).find(key => key.startsWith("__reactProps"));
const reactProps = searchInput[reactKey];

assert.match(reactProps.onChange.toString(), /setQuery\s*\(\s*e\.target\.value\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

export function FruitsSearch() {
  const [query, setQuery] = useState('');
  const [results, setResults] = useState([]);

  return (
    <div id="search-container">
      <form>
        <label htmlFor="search-input">Search for fruits:</label>
        --fcc-editable-region--
        <input
          id="search-input"
          type="search"
          value={query}
          
        />
        --fcc-editable-region--
      </form>
    </div>
  );
}
```
