---
id: 67f8ef57b9974052e5ef6b30
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

`#results` 요소 안에서 삼항 연산자를 사용해 `results` 배열의 길이가 0보다 큰지 확인하세요. 크다면 항목들을 map으로 순회하며 각 항목을 `p` 요소에 `result-item` 클래스를 붙여 표시하세요. 각 문단에는 `key` 속성도 반드시 포함하세요.

`results`가 비어 있다면 `p` 요소를 렌더링하고 메시지 `No results found`를 표시하세요.

# --hints--

`#results` 컨테이너에는 쿼리를 기준으로 일치하는 과일 목록이 포함되어야 합니다. 각 결과는 `p` 클래스를 가진 `result-item` 요소여야 하며 과일 이름을 보여야 합니다.

```js
  {
    const mockResults = ["Item 1", "Item 2"];
    React.useState = (arg) => {
      if (Array.isArray(arg)) return [mockResults, () => {}]; 
      return [null, () => {}];
    };
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.FruitsSearch);
    const items = div.querySelectorAll("p.result-item");

    assert.equal(items.length, mockResults.length);
    assert.equal(items[0].textContent, "Item 1");
    assert.equal(items[1].textContent, "Item 2");

    assert.isTrue(items[0].classList.contains("result-item"));
    assert.isTrue(items[1].classList.contains("result-item"));
  }
```

`#results` 요소에는 `p` 배열이 비어 있을 때 텍스트 `No results found`이 있는 `results` 요소가 포함되어야 합니다.

```js
  {
    React.useState = (arg) => {
      if (Array.isArray(arg)) return [[], () => {}];
      return [null, () => {}];
    };
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.FruitsSearch);
    const noResultsEl = div.querySelector("#results p");
    assert.exists(noResultsEl);
    assert.equal(noResultsEl.textContent, "No results found");
  }
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

export function FruitsSearch() {
  const [query, setQuery] = useState('');
  const [results, setResults] = useState([]);

  function handleSubmit(e) {
    e.preventDefault();
  }

  return (
    <div id="search-container">
      <form onSubmit={handleSubmit}>
        <label htmlFor="search-input">Search for fruits:</label>
        <input
          id="search-input"
          type="search"
          value={query}
          onChange={(e) => setQuery(e.target.value)}
        />
      </form>
      --fcc-editable-region--
      <div id="results">

      </div>
      --fcc-editable-region--
    </div>
  );
}
```
