---
id: 66f26c32ec6f90df01a44f60
title: Ein Kontaktformular gestalten
challengeType: 25
dashedName: design-a-contact-form
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `div` mit einer Klasse `form-container` haben, das Ihr Formular enthält.

2. Sie sollten ein `form`-Element innerhalb des `.form-container`-Elements haben.

3. Innerhalb des `form`-Elements sollten Sie die folgenden Elemente und Eingabefelder haben:

   - Ein `h2`-Element mit etwas Text.
   - Ein Texteingabefeld für den Namen mit dem `type`, das auf `text` gesetzt ist, sowie den Attributen `id`, `name`, `required`.
   - Ein E-Mail-Eingabefeld mit dem `type`, das auf `email` gesetzt ist, sowie den Attributen `id`, `name`, `required`.
   - Ein Textbereich für die Nachricht mit `id`, `name` und dem Attribut `required`.
   - Ein `button`-Element mit dem Text `"Submit"` und einem `type`-Attribut von `submit`.

4. Jedes `input`- und `textarea`-Element sollte sein entsprechendes `label`-Element mit einem `for`-Attribut haben, das auf den `id` des Elements gesetzt ist.

5. Das `.form-container`-Element sollte eine Hintergrundfarbe haben und Werte für `border-radius`, `padding` und `width` besitzen.

6. Die `label`-Elemente sollten einen Außenabstand und eine Schriftfarbe haben.

7. Die `input`- und `textarea`-Elemente sollten Werte für folgende Eigenschaften haben:
   
   - `width`.
   - `padding`.
   - Außenabstand unten.

8. Das `button`-Element sollte eine:
   
   - Hintergrundfarbe
   - Schriftgröße abweichend von der Standardgröße haben.

9. Der Button sollte einen Hover-Effekt haben, der die Hintergrundfarbe ändert.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verlinken und Ihr CSS anwenden.

# --hints--

Sie sollten ein `div` mit einer Klasse `form-container` haben.

```js
const el = document.querySelector('div.form-container');
assert.isNotNull(el);
```

Sie sollten ein `form`-Element innerhalb des `.form-container`-Elements haben.

```js
const el = document.querySelector('.form-container form');
assert.isNotNull(el);
```

Innerhalb des `form` sollten Sie ein `h2`-Element mit etwas Text haben.

```js
const el = document.querySelector('.form-container form h2');
assert.isNotEmpty(el.innerText);
// assert.isAtLeast(document.querySelector('h1').innerText.length, 1);

```

Innerhalb des `form` sollten Sie ein Texteingabefeld mit dem Typ `text` haben.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.isNotNull(el);
```

Ihr `type=text`-Eingabefeld sollte ein `id`, ein `name` und ein `required`-Attribut haben.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.exists(el);
assert.isTrue(el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

Innerhalb des `form` sollten Sie ein E-Mail-Eingabefeld mit dem `type`, das auf `email` gesetzt ist, haben.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isNotNull(el);
```

Ihr `type=email`-Eingabefeld sollte ein `id`, ein `name` und ein `required`-Attribut haben.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isTrue(el?.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

Innerhalb des `form` sollten Sie ein `textarea` mit einem `id`, `name` und einem `required`-Attribut haben.

```js
const el = document.querySelector('.form-container form textarea');
assert.isNotNull(el);
assert.isNotNull(document.querySelector('.form-container form textarea[required]'));
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

Innerhalb des `form` sollten Sie ein `button`-Element mit einem `type`-Attribut von `submit` haben.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(el?.getAttribute('type'), 'submit');
```

Ihr Button sollte den Text `"Submit"` haben.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(document.querySelector('.form-container form button').textContent, 'Submit');
```

Sie sollten ein `label`-Element für das `#name`-Element mit einem `for`-Attribut haben, das mit dem `id`-Wert des entsprechenden `input` übereinstimmt.

```js
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value
assert.isNotEmpty(nameInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === nameInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

Sie sollten ein `label`-Element für das `#email`-Element mit einem `for`-Attribut haben, das mit dem `id`-Wert des entsprechenden `input` übereinstimmt.

```js
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);
// check if id has a value
assert.isNotEmpty(emailInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === emailInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

Sie sollten ein `label`-Element für das `#message`-Element mit einem `for`-Attribut haben, das mit dem `id`-Wert des `textarea` übereinstimmt.

```js
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value

assert.isNotEmpty(messageInput);

assert.exists(document.querySelector(`label[for="${messageInput}"]`));
```

Die Bezeichnungen sollten nicht leer sein.

```js
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

for (let label of labels) {
    const labelText = label.textContent.trim();
    assert.isNotEmpty(labelText);
}

```

Das `.form-container`-Element sollte eine Hintergrundfarbe haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

Das `.form-container`-Element sollte einen `border-radius` haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

Das `.form-container`-Element sollte Innenabstand haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

Das `.form-container`-Element sollte eine Breite haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

Ihre Bezeichnungen sollten einen Außenabstand und eine Schriftfarbe haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

Ihre `input`-Elemente sollten eine Breite, Innenabstand und einen Außenabstand unten haben.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('input'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);

```

Ihr `textarea`-Element sollte eine Breite, Innenabstand und Außenabstand unten haben.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('textarea'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);
```

Ihr `button`-Element sollte eine Hintergrundfarbe haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

Ihr `button`-Element sollte eine Schriftgröße haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

Ihr `button`-Element sollte einen Hover-Effekt haben, der die Hintergrundfarbe ändert.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const btn_hover = gs('button:hover')?.backgroundColor;
assert.isNotEmpty(btn_hover);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled Form Project</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="form-container">
        <form>
            <h2>Contact Us</h2>
            <label for="name">Name:</label>
            <input type="text" id="name" name="name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" required>

            <label for="message">Message:</label>
            <textarea id="message" name="message" required></textarea>

            <button type="submit">Submit</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background: linear-gradient(45deg, #83a4d4, #b6fbff);
    font-family: Arial, sans-serif;
}

.form-container {
    background-color: white;
    border-radius: 10px;
    padding: 20px;
    width: 300px;
    text-align: center;
}

form h2 {
    margin-bottom: 20px;
    color: #333;
}

label {
    display: block;
    margin: 10px 0 5px;
    color: #555;
}

input{
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

textarea {
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

button {
    background-color: #4CAF50;
    color: white;
    padding: 10px 15px;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    font-size: 1em;
}

button:hover {
    background-color: #45a049;
}
```
