---
id: a0b5010f579e69b815e7c5d6
title: Tafuta na Badilisha
challengeType: 1
forumTopicId: 16045
dashedName: search-and-replace
---

# --description--

Fanya utafutaji na ubadilishaji kwenye sentensi ukitumia hoja zilizotolewa na urudishe sentensi mpya.

Hoja ya kwanza ni sentensi ambayo utafanya utafutaji na ubadilishaji.

Hoja ya pili ni neno ambalo utabadilisha (kabla).

Hoja ya tatu ni kile utakachobadilisha neno la pili nacho (baada).

**Kumbuka:** Hifadhi herufi ya kwanza ya neno la awali unalobadilisha. Kwa mfano, ikiwa unakusudia kubadilisha neno `Book` na neno `dog`, linapaswa kubadilishwa kuwa `Dog`

# --hints--

`myReplace("Let us go to the store", "store", "mall")` inapaswa kurudisha mfuatano wa herufi `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` inapaswa kurudisha mfuatano wa herufi `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` inapaswa kurudisha mfuatano wa herufi `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` inapaswa kurudisha mfuatano wa herufi `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` inapaswa kurudisha mfuatano wa herufi `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` inapaswa kurudisha mfuatano wa herufi `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js
function myReplace(str, before, after) {
  return str;
}

myReplace("A quick brown fox jumped over the lazy dog", "jumped", "leaped");
```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
