---
id: a103376db3ba46b2d50db289
title: Kesi ya Spinal Tap
challengeType: 1
forumTopicId: 16078
dashedName: spinal-tap-case
---

# --description--

Badilisha mfuatano wa herufi kuwa kesi ya spinal. Kesi ya spinal ni maneno yote yenye herufi ndogo yaliyounganishwa kwa alama za dashi.

# --hints--

`spinalCase("This Is Spinal Tap")` inapaswa kurudisha mfuatano wa herufi `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` inapaswa kurudisha mfuatano wa herufi `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` inapaswa kurudisha mfuatano wa herufi `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` inapaswa kurudisha mfuatano wa herufi `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` inapaswa kurudisha mfuatano wa herufi `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js
function spinalCase(str) {
  return str;
}

spinalCase('This Is Spinal Tap');
```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
