---
id: a105e963526e7de52b219be9
title: Muungano uliopangwa
challengeType: 1
forumTopicId: 16077
dashedName: sorted-union
---

# --description--

Andika kitendakazi kinachopokea safu za data mbili au zaidi na kurudisha safu mpya yenye thamani za kipekee kwa mpangilio wa safu za data za awali zilizotolewa.

Kwa maneno mengine, thamani zote zilizopo katika safu zote za data zinapaswa kujumuishwa kwa mpangilio wao wa awali, lakini bila kurudia katika safu ya mwisho.

Nambari za kipekee zinapaswa kupangwa kulingana na mpangilio wao wa awali, lakini safu ya mwisho haipaswi kupangwa kwa mpangilio wa nambari.

Kagua majaribio ya uthibitisho kwa mifano.

# --hints--

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` inapaswa kurudisha `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` inapaswa kurudisha `[1, 2, 3, 5]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` inapaswa kurudisha `[1, 2, 3, 5, 4, 6, 7, 8]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` inapaswa kurudisha `[1, 3, 2, 5, 4, 6]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` inapaswa kurudisha `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js
function uniteUnique(arr) {
  return arr;
}

uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]);
```

# --solutions--

```js
function uniteUnique(arr) {
  return [].slice.call(arguments).reduce(function(a, b) {
    return [].concat(
      a, 
      b.filter(function(e, currentIndex) {
        return b.indexOf(e) === currentIndex && a.indexOf(e) === -1;
      }));
  }, []);
}
```
