---
id: a3566b1109230028080c9345
title: Jumlisha nambari zote katika safu ya nambari
challengeType: 1
forumTopicId: 16083
dashedName: sum-all-numbers-in-a-range
---

# --description--

Tutakutumia safu ya data yenye nambari mbili. Rudisha jumla ya nambari hizo mbili pamoja na jumla ya nambari zote zilizo kati yao. Nambari ndogo haitakuwa daima ya kwanza.

Kwa mfano, `sumAll([4,1])` inapaswa kurudisha `10` kwa sababu jumla ya nambari zote kati ya 1 na 4 (zikiwa zote ni pamoja) ni `10`.

# --hints--

`sumAll([1, 4])` inapaswa kurudisha nambari.

```js
assert(typeof sumAll([1, 4]) === 'number');
```

`sumAll([1, 4])` inapaswa kurudisha 10.

```js
assert.deepEqual(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` inapaswa kurudisha 10.

```js
assert.deepEqual(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` inapaswa kurudisha 45.

```js
assert.deepEqual(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` inapaswa kurudisha 45.

```js
assert.deepEqual(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
function sumAll(arr) {
  return 1;
}

sumAll([1, 4]);
```

# --solutions--

```js
function sumAll(arr) {
  var sum = 0;
  arr.sort(function(a,b) {return a-b;});
  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }
  return sum;
}
```
