---
id: a39963a4c10bc8b4d4f06d7e
title: Tafuta na Uondoe
challengeType: 1
forumTopicId: 16046
dashedName: seek-and-destroy
---

# --description--

Utapewa safu ya data ya awali kama hoja ya kwanza kwa kitendakazi cha `destroyer`, ikifuatiwa na hoja moja au zaidi. Ondoa vipengele vyote kutoka kwenye safu ya data ya awali ambavyo vina thamani sawa na hizi hoja.

Kitendakazi lazima kikubali idadi isiyojulikana ya hoja, pia kinajulikana kama kitendakazi cha variadic. Unaweza kupata hoja za ziada kwa kuongeza kigezo cha rest kwenye ufafanuzi wa kitendakazi au kwa kutumia kitu cha `arguments`.

# --hints--

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` inapaswa kurudisha `[1, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` inapaswa kurudisha `[1, 5, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` inapaswa kurudisha `[1]`.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` inapaswa kurudisha `[]`.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` inapaswa kurudisha `["hamburger"]`.

```js
assert.deepEqual(destroyer(['tree', 'hamburger', 53], 'tree', 53), [
  'hamburger'
]);
```

`destroyer(["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan")` inapaswa kurudisha `[12,92,65]`.

```js
assert.deepEqual(
  destroyer(
    [
      'possum',
      'trollo',
      12,
      'safari',
      'hotdog',
      92,
      65,
      'grandma',
      'bugati',
      'trojan',
      'yacht'
    ],
    'yacht',
    'possum',
    'trollo',
    'safari',
    'hotdog',
    'grandma',
    'bugati',
    'trojan'
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js
function destroyer(arr) {
  return arr;
}

destroyer([1, 2, 3, 1, 2, 3], 2, 3);
```

# --solutions--

```js
function destroyer(arr) {
  var hash = Object.create(null);
  [].slice.call(arguments, 1).forEach(function(e) {
    hash[e] = true;
  });
  return arr.filter(function(e) { return !(e in hash);});
}

destroyer([1, 2, 3, 1, 2, 3], 2, 3);
```
