---
id: a3bfc1673c0526e06d3ac698
title: Jumlisha namba msingi zote
challengeType: 1
forumTopicId: 16085
dashedName: sum-all-primes
---

# --description--

Namba msingi ni nambari kamili kubwa kuliko 1 yenye mgawanyiko halisi miwili tu: 1 na yenyewe. Kwa mfano, 2 ni namba msingi kwa sababu inagawanyika tu na 1 na 2. Kinyume chake, 4 si namba msingi kwa sababu inagawanyika na 1, 2 na 4.

Andika upya `sumPrimes` ili irudishe jumla ya namba msingi zote ambazo ni ndogo au sawa na num.

# --hints--

`sumPrimes(10)` inapaswa kurudisha nambari.

```js
assert.deepEqual(typeof sumPrimes(10), 'number');
```

`sumPrimes(10)` inapaswa kurudisha 17.

```js
assert.deepEqual(sumPrimes(10), 17);
```

`sumPrimes(977)` inapaswa kurudisha 73156.

```js
assert.deepEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js
function sumPrimes(num) {
  return num;
}

sumPrimes(10);
```

# --solutions--

```js
class PrimeSeive {
  constructor(num) {
    const seive = Array(Math.floor((num - 1) / 2)).fill(true);
    const upper = Math.floor((num - 1) / 2);
    const sqrtUpper = Math.floor((Math.sqrt(num) - 1) / 2);

    for (let i = 0; i <= sqrtUpper; i++) {
      if (seive[i]) {
        // Mark value in seive array
        const prime = 2 * i + 3;
        // Mark all multiples of this number as false (not prime)
        const primeSquaredIndex = 2 * i ** 2 + 6 * i + 3;
        for (let j = primeSquaredIndex; j < upper; j += prime) {
          seive[j] = false;
        }
      }
    }

    this._seive = seive;
  }

  isPrime(num) {
    return num === 2
      ? true
      : num % 2 === 0
        ? false
        : this.isOddPrime(num);
  }

  isOddPrime(num) {
    return this._seive[(num - 3) / 2];
  }
};

function sumPrimes(num) {
  const primeSeive = new PrimeSeive(num);

  let sum = 2;
  for (let i = 3; i <= num; i += 2) {
    if (primeSeive.isOddPrime(i)) sum += i;
  }
  return sum;
}

sumPrimes(10);
```
