---
id: a5deed1811a43193f9f1c841
title: Acha
challengeType: 1
forumTopicId: 16010
dashedName: drop-it
---

# --description--

Ikiwa umepewa safu ya data `arr`, rudia kupitia na ondoa kila kipengele kuanzia kipengele cha kwanza (kielezo 0) hadi kitendakazi `func` kitakaporudisha `true` wakati kipengele kinaporudiwa kupitia kwake.

Kisha rudisha sehemu iliyobaki ya safu ya data mara tu hali itakapokidhiwa, vinginevyo, `arr` inapaswa kurudishwa kama safu ya data tupu.

# --hints--

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` inapaswa kurudisha `[3, 4]`.

```js
assert.deepEqual(
  dropElements([1, 2, 3, 4], function (n) {
    return n >= 3;
  }),
  [3, 4]
);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` inapaswa kurudisha `[1, 0, 1]`.

```js
assert.deepEqual(
  dropElements([0, 1, 0, 1], function (n) {
    return n === 1;
  }),
  [1, 0, 1]
);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` inapaswa kurudisha `[1, 2, 3]`.

```js
assert.deepEqual(
  dropElements([1, 2, 3], function (n) {
    return n > 0;
  }),
  [1, 2, 3]
);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` inapaswa kurudisha `[]`.

```js
assert.deepEqual(
  dropElements([1, 2, 3, 4], function (n) {
    return n > 5;
  }),
  []
);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` inapaswa kurudisha `[7, 4]`.

```js
assert.deepEqual(
  dropElements([1, 2, 3, 7, 4], function (n) {
    return n > 3;
  }),
  [7, 4]
);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` inapaswa kurudisha `[3, 9, 2]`.

```js
assert.deepEqual(
  dropElements([1, 2, 3, 9, 2], function (n) {
    return n > 2;
  }),
  [3, 9, 2]
);
```

# --seed--

## --seed-contents--

```js
function dropElements(arr, func) {
  return arr;
}

dropElements([1, 2, 3], function(n) {return n < 3; });
```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
