---
id: a97fd23d9b809dac9921074f
title: Hoja si za lazima
challengeType: 1
forumTopicId: 14271
dashedName: arguments-optional
---

# --description--

Tengeneza kitendakazi kinachojumlisha hoja mbili pamoja. Ikiwa hoja moja tu itatolewa, basi rudisha kitendakazi kinachotarajia hoja moja na kurudisha jumla.

Kwa mfano, `addTogether(2, 3)` inapaswa kurudisha `5`, na `addTogether(2)` inapaswa kurudisha kitendakazi.

Kupiga wito kwa kitendakazi kilichorudishwa na hoja moja kutaongeza na kurudisha jumla:

```js
var sumTwoAnd = addTogether(2);
```

`sumTwoAnd(3)` inarudisha `5`.

Ikiwa hoja yoyote si nambari halali, rudisha undefined.

# --hints--

`addTogether(2, 3)` inapaswa kurudisha 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` inapaswa kurudisha 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` inapaswa kurudisha `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` inapaswa kurudisha `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` inapaswa kurudisha `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` inapaswa kurudisha kitendakazi.

```js
assert.deepEqual(typeof(addTogether(5)), 'function');
```

`addTogether(5)(7)` inapaswa kurudisha 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` inapaswa kurudisha `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` inapaswa kurudisha `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js
function addTogether() {
  return false;
}

addTogether(2,3);
```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
