---
id: aa7697ea2477d1316795783b
title: Pig Latin
challengeType: 1
forumTopicId: 16039
dashedName: pig-latin
---

# --description--

Pig Latin ni njia ya kubadilisha maneno ya Kiingereza. Sheria ni kama ifuatavyo:

\- Ikiwa neno linaanza na konsonanti, chukua konsonanti ya kwanza au kundi la konsonanti, ihamishe mwishoni mwa neno, kisha ongeza `ay`.

\- Ikiwa neno linaanza na vokali, ongeza tu `way` mwishoni.

# --instructions--

Tafsiri mfuatano wa herufi uliotolewa kuwa Pig Latin. Mfuatano wa herufi wa ingizo unahakikishiwa kuwa maneno ya Kiingereza yenye herufi ndogo zote.

# --hints--

`translatePigLatin("california")` inapaswa kurudisha mfuatano wa herufi `aliforniacay`.

```js
assert.deepEqual(translatePigLatin('california'), 'aliforniacay');
```

`translatePigLatin("paragraphs")` inapaswa kurudisha mfuatano wa herufi `aragraphspay`.

```js
assert.deepEqual(translatePigLatin('paragraphs'), 'aragraphspay');
```

`translatePigLatin("glove")` inapaswa kurudisha mfuatano wa herufi `oveglay`.

```js
assert.deepEqual(translatePigLatin('glove'), 'oveglay');
```

`translatePigLatin("algorithm")` inapaswa kurudisha mfuatano wa herufi `algorithmway`.

```js
assert.deepEqual(translatePigLatin('algorithm'), 'algorithmway');
```

`translatePigLatin("eight")` inapaswa kurudisha mfuatano wa herufi `eightway`.

```js
assert.deepEqual(translatePigLatin('eight'), 'eightway');
```

Inapaswa kushughulikia maneno ambapo vokali ya kwanza ipo katikati ya neno. `translatePigLatin("schwartz")` inapaswa kurudisha mfuatano wa herufi `artzschway`.

```js
assert.deepEqual(translatePigLatin('schwartz'), 'artzschway');
```

Inapaswa kushughulikia maneno yasiyo na vokali. `translatePigLatin("rhythm")` inapaswa kurudisha mfuatano wa herufi `rhythmay`.

```js
assert.deepEqual(translatePigLatin('rhythm'), 'rhythmay');
```

# --seed--

## --seed-contents--

```js
function translatePigLatin(str) {
  return str;
}

translatePigLatin("consonant");
```

# --solutions--

```js
function translatePigLatin(str) {
  if (isVowel(str.charAt(0))) return str + "way";
  var front = [];
  str = str.split('');
  while (str.length && !isVowel(str[0])) {
    front.push(str.shift());
  }
  return [].concat(str, front).join('') + 'ay';
}

function isVowel(c) {
  return ['a', 'e', 'i', 'o', 'u'].indexOf(c.toLowerCase()) !== -1;
}
```
