---
id: af4afb223120f7348cdfc9fd
title: Panga taka za anga
challengeType: 1
forumTopicId: 16021
dashedName: map-the-debris
---

# --description--

Kulingana na Sheria ya Tatu ya Kepler, kipindi cha mzunguko $T$ cha vitu viwili vinavyozungukana katika mzunguko wa mduara au elipti ni:

$$
T = 2 \pi \sqrt{\frac{a^{3}}{\mu}}
$$

- $a$ ni nusu-mhimili mkuu wa mzunguko
- $μ = GM$ ni kigezo cha mvutano wa kawaida
- $G$ ni thabiti ya mvutano,
- $M$ ni uzito wa mwili wenye uzito mkubwa zaidi.

Rudisha safu mpya ya data inayobadilisha urefu wa wastani wa vitu kuwa vipindi vyao vya mzunguko (katika sekunde).

Safu ya data itakuwa na vitu katika muundo `{name: 'name', avgAlt: avgAlt}`.

Thamani zinapaswa kuzungushwa hadi nambari kamili iliyo karibu zaidi. Mwili unaozungukwa ni Dunia.

Mduara wa Dunia ni kilomita 6367.4447, na thamani ya GM ya Dunia ni 398600.4418 km<sup>3</sup>s<sup>-2</sup>.

# --hints--

`orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}])` inapaswa kurudisha `[{name: "sputnik", orbitalPeriod: 86400}]`.

```js
assert.deepEqual(orbitalPeriod([{ name: 'sputnik', avgAlt: 35873.5553 }]), [
  { name: 'sputnik', orbitalPeriod: 86400 }
]);
```

`orbitalPeriod([{name: "iss", avgAlt: 413.6}, {name: "hubble", avgAlt: 556.7}, {name: "moon", avgAlt: 378632.553}])` inapaswa kurudisha `[{name : "iss", orbitalPeriod: 5557}, {name: "hubble", orbitalPeriod: 5734}, {name: "moon", orbitalPeriod: 2377399}]`.

```js
assert.deepEqual(
  orbitalPeriod([
    { name: 'iss', avgAlt: 413.6 },
    { name: 'hubble', avgAlt: 556.7 },
    { name: 'moon', avgAlt: 378632.553 }
  ]),
  [
    { name: 'iss', orbitalPeriod: 5557 },
    { name: 'hubble', orbitalPeriod: 5734 },
    { name: 'moon', orbitalPeriod: 2377399 }
  ]
);
```

# --seed--

## --seed-contents--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  return arr;
}

orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}]);
```

# --solutions--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  const TAU = 2 * Math.PI;
  return arr.map(function(obj) {
    return {
      name: obj.name,
      orbitalPeriod: Math.round(TAU * Math.sqrt(Math.pow(obj.avgAlt+earthRadius, 3)/GM))
    };
  });
}
```
