---
id: af7588ade1100bde429baf20
title: Herufi zilizokosekana
challengeType: 1
forumTopicId: 16023
dashedName: missing-letters
---

# --description--

Tafuta herufi iliyokosekana katika safu ya herufi iliyotolewa na urudishe.

Kama herufi zote zipo katika safu, rudisha `undefined`.

# --hints--

`fearNotLetter("abce")` inapaswa kurudisha mfuatano wa herufi `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` inapaswa kurudisha mfuatano wa herufi `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` inapaswa kurudisha mfuatano wa herufi `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` inapaswa kurudisha mfuatano wa herufi `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` inapaswa kurudisha `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js
function fearNotLetter(str) {
  return str;
}

fearNotLetter("abce");
```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    var letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
