---
id: 69bab8292750d84cac5b52a0
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Katika warsha hii, throughput inahusu jumla ya idadi ya washiriki waliopokelewa na lango katika alama zote za kupigia katika simulizi moja (asubuhi au usiku). Kitendakazi chako cha `initializeThroughput` kitaandaa kitu cha kufuatilia ni washiriki wangapi kila lango limepokea.

Kwanza, tengeneza kitu tupu kinachoitwa `summary`.

# --hints--

Kitendakazi chako cha `initializeThroughput` kinapaswa kuunda kitu tupu na kuweka thamani kwa kigezo kinachoitwa `summary`.

```js
const explorer = await __helpers.Explorer(code);
const { initializeThroughput } = explorer.allFunctions;
const { summary } = initializeThroughput?.variables ?? {};
const validExpressions = ["{}", "new Object()"];
assert.isTrue(
  validExpressions.some((expression) => summary?.value.matches(expression))
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
