---
id: 69babdf3f896711c11516373
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Kila mara simulation inapopita:

- Idadi fulani ya washiriki huwasili langoni (kutoka kwa `queue` yake).
- Lango linashughulikia washiriki kulingana na `capacity` yake.
- Ikiwa washiriki wengi zaidi huwasili kuliko lango linavyoweza kushughulikia, baadhi wataendelea kubaki (kuzidi).

Sasa utaunda kitendakazi kinachoshughulikia mantiki hii kwa lango moja katika tick moja.

Tengeneza kitendakazi tupu kinachoitwa `processGateFlow` kinachokubali vigezo viwili:

- `gate`: Kitu cha lango kinachoshughulikiwa.
- `tickIndex`: Nafasi ya tick ya sasa ya simulation katika safu ya `queue`.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `processGateFlow`.

```js
assert.isFunction(processGateFlow);
```

Kitendakazi chako `processGateFlow` kinapaswa kuwa na vigezo `gate` na `tickIndex`.

```js
const regex = __helpers.functionRegex('processGateFlow', ['gate', 'tickIndex']);
assert.match(__helpers.removeJSComments(processGateFlow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
