---
id: 69baafd6d683dd9921b75db3
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utatumia JavaScript kuiga mtiririko wa watu wanaohudhuria tamasha la muziki.

Utafanya kazi na seti mbili za data zinazofafanua milango ya tamasha wakati wa zamu za asubuhi na usiku:

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

Kila kitu kinachohusiana na mlango kina vigezo vifuatavyo:

- `id`: Mfuatano wa herufi unaotambulisha mlango.
- `capacity`: Idadi ya watu mlango unaweza kuwahudumia kwa kila kipindi.
- `queue`: Safu ya data ya nambari zinazowakilisha idadi ya watu wanaofika mlango wakati wa kipindi fulani.

Ili kuanza, tengeneza kitendakazi tupu kinachoitwa `initializeThroughput` kinachopokea kigezo kinachoitwa `gates`.

# --hints--

Unapaswa kutengeneza kitendakazi kinachoitwa `initializeThroughput`.

```js
assert.isFunction(initializeThroughput);
```

Kitendakazi chako cha `initializeThroughput` kinapaswa kuwa na kigezo `gates`.

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
