---
id: 69cfdc80443b715a8aa452a7
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Ili ukamilishe kitendakazi chako cha `initializeThroughput`:

- Rudia kupitia kila lango katika safu ya data ya ingizo ya `gates`. Kwa kila lango, ongeza sifa kwenye `summary` ambapo:
  - Ufunguo ni `id` ya lango hilo.
  - Thamani ni `0`.
- Baada ya mzunguko, rudisha kitu cha `summary`.

Kitendakazi chako cha `initializeThroughput` hakipaswi kubadilisha safu ya data ya ingizo ya asili `gates`.

# --hints--

Kupiga wito `initializeThroughput()` na `[]` kunapaswa kurudisha `{}` bila kubadilisha safu ya data ya ingizo ya asili.

```js
const input = [];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  {},
  "Calling `initializeThroughput([])` should return an empty object."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given an empty array."
);
```

Kupiga wito `initializeThroughput()` na `[{ id: "North", capacity: 5, queue: [1, 2, 3] }]` kunapaswa kurudisha `{ North: 0 }` bila kubadilisha safu ya data ya ingizo ya asili.

```js
const input = [{ id: "North", capacity: 5, queue: [1, 2, 3] }];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0 },
  "Calling `initializeThroughput([{ id: 'North', ... }])` should return `{ North: 0 }`."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given one gate."
);
```

Kupiga wito `initializeThroughput()` na `[{ id: "North", capacity: 5, queue: [] }, { id: "East", capacity: 3, queue: [] }, { id: "South", capacity: 4, queue: [] }]` kunapaswa kurudisha `{ North: 0, East: 0, South: 0}` bila kubadilisha safu ya data ya ingizo ya asili.

```js
const input = [
  { id: "North", capacity: 5, queue: [] },
  { id: "East", capacity: 3, queue: [] },
  { id: "South", capacity: 4, queue: [] }
];

const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0, East: 0, South: 0 },
  "Calling `initializeThroughput` with multiple gates should return an object with an entry for each gate id set to 0."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given multiple gates."
);
```

Unapaswa kutumia mzunguko wa `for` kurudia kupitia kila kitu cha lango ndani ya safu ya data ya ingizo ya `gates`.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));
assert.match(cleaned, /for\s*\(/);
```

Baada ya mzunguko, unapaswa kurudisha `summary`. Hii inapaswa kuwa mstari wa mwisho katika kitendakazi.

```js
const funcStr = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));

assert.match(
  funcStr,
  /returnsummary;?}$/,
  "Your `initializeThroughput` function should end with `return summary`."
);
```

Kitendakazi chako cha `initializeThroughput` hakipaswi kubadilisha safu ya data ya ingizo ya asili `gates`.

```js
const morningCopy = JSON.parse(JSON.stringify(morningGates));
const nightCopy = JSON.parse(JSON.stringify(nightGates));

initializeThroughput(morningGates);
initializeThroughput(nightGates);

assert.deepEqual(
  morningGates,
  morningCopy,
  "`initializeThroughput` should not mutate the `morningGates` seed array."
);

assert.deepEqual(
  nightGates,
  nightCopy,
  "`initializeThroughput` should not mutate the `nightGates` seed array."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
--fcc-editable-region--

--fcc-editable-region--
}
```
