---
id: 68c31a762717d583d90ad31e
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Auf der Blogseite sind insgesamt drei Blogbeiträge in generischen `div`-Elementen eingehüllt. Es wäre jedoch besser, semantische Elemente zu verwenden, um jeden Beitrag einzuschließen.

Ändern Sie jedes `div`-Element, das jeden Blogbeitrag umschließt, in ein `article`-Element.

# --hints--

Ihr erstes `div`-Element für den ersten Beitrag sollte in ein `article`-Element geändert werden.

```js
const firstPost = document.querySelector("nav + article");
assert.isNotNull(firstPost);
```

Ihr zweites `div`-Element für den zweiten Beitrag sollte in ein `article`-Element geändert werden.

```js
const secondPost = document.querySelector("article + article");
assert.isNotNull(secondPost);
```

Ihr drittes `div`-Element für den dritten Beitrag sollte in ein `article`-Element geändert werden.

```js
const thirdPost = document.querySelector("article + article + article");
assert.isNotNull(thirdPost);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <nav>
    <h2>Navigation</h2>
    <ul>
      <li><a href="#post1">My Journey</a></li>
      <li><a href="#post2">Accessibility</a></li>
      <li><a href="#post3">Next Steps</a></li>
    </ul>
  </nav>

--fcc-editable-region--
  <div>
    <h2 id="post1">My Journey Learning to Code</h2>
    <p>I started learning to code a few months ago and it's been a wild ride!</p>
    
    <h3>Early Challenges</h3>
    <p>At first, syntax was really confusing.</p>
    
    <h3>Breakthroughs</h3>
    <p>Eventually things started to click.</p>
  </div>

  <div>
    <h2 id="post2">Accessibility Matters</h2>
    <p>Today I learned that not everyone uses the web the same way I do.</p>

    <h3>Screen Readers</h3>
    <p>These tools help visually impaired users browse websites.</p>
  </div>

  <div>
    <h2 id="post3">What's Next?</h2>
    <p>I'm excited to dive into JavaScript and build interactive features!</p>

    <h3>Coming soon: My first JavaScript project!</h3>
    <p>Stay tuned for some exciting interactive blog features.</p>
  </div>
--fcc-editable-region--
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
</body>
</html>
```
