---
id: 587d78af367417b2b2512b03
title: Costruisci un modulo di sondaggio
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301145
dashedName: build-a-survey-form
---

# --description--

**Obiettivo:** Costruisci un'app che sia funzionalmente simile a <a href="https://survey-form.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://survey-form.freecodecamp.rocks</a>. **Non copiare questo progetto dimostrativo**.

**Storie utente:**

1. Dovresti avere un titolo della pagina in un elemento `h1` con un `id` di `title`
1. Dovresti avere una breve spiegazione in un elemento `p` con un `id` di `description`
1. Dovresti avere un elemento `form` con un `id` di `survey-form`
1. All'interno dell'elemento modulo, sei **tenuto** a inserire il tuo nome in un campo `input` che ha un `id` di `name` e un `type` di `text`
1. All'interno dell'elemento modulo, sei **tenuto** a inserire la tua email in un campo `input` che ha un `id` di `email`
1. Se inserisci un'email non formattata correttamente, vedrai un errore di validazione HTML5
1. All'interno del modulo, puoi inserire un numero in un campo `input` che ha un `id` di `number`
1. Il campo di input numerico non dovrebbe accettare valori non numerici, sia impedendoti di digitarli sia mostrando un errore di validazione HTML5 (a seconda del browser).
1. Se inserisci numeri fuori dall'intervallo definito dagli attributi `min` e `max` del campo numerico, vedrai un errore di validazione HTML5
1. Per i campi di input nome, email e numero, puoi vedere elementi `label` corrispondenti nel modulo, che descrivono lo scopo di ogni campo con gli id seguenti: `id="name-label"`, `id="email-label"` e `id="number-label"`
1. Per i campi di input nome, email e numero, puoi vedere un testo segnaposto che fornisce una descrizione o istruzioni per ogni campo
1. All'interno dell'elemento modulo, dovresti avere un menu a discesa `select` con un `id` di `dropdown` e almeno due opzioni tra cui scegliere
1. All'interno dell'elemento modulo, puoi selezionare un'opzione da un gruppo di almeno due pulsanti di opzione raggruppati usando l'attributo `name`
1. All'interno dell'elemento modulo, puoi selezionare diversi campi da una serie di caselle di spunta, ognuna delle quali deve avere un attributo `value`
1. All'interno dell'elemento modulo, ti viene presentato un `textarea` per commenti aggiuntivi
1. All'interno dell'elemento modulo, ti viene presentato un pulsante con `id` di `submit` per inviare tutti gli input

Completa le storie utente e supera tutti i test qui sotto per completare questo progetto. Dagli il tuo stile personale. Buona programmazione!

**Nota:** Assicurati di aggiungere `<link rel="stylesheet" href="styles.css">` nel tuo HTML per collegare il foglio di stile e applicare il CSS

# --hints--

Dovresti avere un elemento `h1` con un `id` di `title`.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

Il tuo `#title` non dovrebbe essere vuoto.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Dovresti avere un elemento `p` con un `id` di `description`.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

Il tuo `#description` non dovrebbe essere vuoto.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Dovresti avere un elemento `form` con un `id` di `survey-form`.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

Dovresti avere un elemento `input` con un `id` di `name`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Il tuo `#name` dovrebbe avere un `type` di `text`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

Il tuo `#name` dovrebbe richiedere un input.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Il tuo `#name` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

Dovresti avere un elemento `input` con un `id` di `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Il tuo `#email` dovrebbe avere un `type` di `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

Il tuo `#email` dovrebbe richiedere un input.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Il tuo `#email` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

Dovresti avere un elemento `input` con un `id` di `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Il tuo `#number` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

Il tuo `#number` dovrebbe avere un `type` di `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

Il tuo `#number` dovrebbe avere un attributo `min` con un valore numerico.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

Il tuo `#number` dovrebbe avere un attributo `max` con un valore numerico.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

Dovresti avere un elemento `label` con un `id` di `name-label`.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Dovresti avere un elemento `label` con un `id` di `email-label`.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Dovresti avere un elemento `label` con un `id` di `number-label`.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Il tuo `#name-label` dovrebbe contenere un testo che descrive l'input.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Il tuo `#email-label` dovrebbe contenere un testo che descrive l'input.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Il tuo `#number-label` dovrebbe contenere un testo che descrive l'input.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Il tuo `#name-label` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

Il tuo `#email-label` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

Il tuo `#number-label` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

Il tuo `#name` dovrebbe avere un attributo `placeholder` e un valore.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Il tuo `#email` dovrebbe avere un attributo `placeholder` e un valore.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Il tuo `#number` dovrebbe avere un attributo `placeholder` e un valore.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Dovresti avere un campo `select` con un `id` di `dropdown`.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

Il tuo `#dropdown` dovrebbe avere almeno due elementi `option` selezionabili (non disabilitati).

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

Il tuo `#dropdown` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

Dovresti avere almeno due elementi `input` con un `type` di `radio` (pulsanti di opzione).

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Dovresti avere almeno due pulsanti di opzione che sono discendenti di `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Tutti i tuoi pulsanti di opzione dovrebbero avere un attributo `value` e un valore.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Tutti i tuoi pulsanti di opzione dovrebbero avere un attributo `name` e un valore.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Ogni gruppo di pulsanti di opzione dovrebbe avere almeno 2 pulsanti di opzione.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
});

assert.isAbove(groupKeys.length, 0);
```

Dovresti avere almeno due elementi `input` con un `type` di `checkbox` (caselle di spunta) che sono discendenti di `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

Tutte le tue caselle di spunta dentro `#survey-form` dovrebbero avere un attributo `value` e un valore.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Dovresti avere almeno un elemento `textarea` che è un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

Dovresti avere un elemento `input` o `button` con un `id` di `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

Il tuo `#submit` dovrebbe avere un `type` di `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

Il tuo `#submit` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Survey Form</title>
  </head>
  <body>
    <h1>Survey Form</h1>
    <p>The card below was built as a sample survey form for freeCodeCamp.</p>
    <main id="main">
      <h1 id="title">Join the Togepi Fan Club!</h1>
      <p id="description">
        Enter your information here to receive updates about club activities,
        our monthly newsletter, and other email communications.
      </p>
      <form id="survey-form" action="#">
        <label for="name" id="name-label"
          ><p>Name:</p>
          <input type="text" id="name" placeholder="e.g. John Smith" required />
        </label>
        <label for="email" id="email-label"
          ><p>Email:</p>
          <input
            type="email"
            id="email"
            placeholder="e.g. john.smith@email.com"
            required
          />
        </label>
        <label for="age" id="number-label"
          ><p>Age<em>(optional)</em>:</p>
          <input
            type="number"
            id="number"
            placeholder="e.g. 19"
            min="18"
            max="99"
          />
        </label>
        <label for="interest" id="interest-label"
          ><p>What are you most interested in?</p>
          <select id="dropdown">
            <option selected disabled hidden></option>
            <option id="battles">Battling</option>
            <option id="breeding">Breeding</option>
            <option id="catching">Completing my Pokedex</option>
            <option id="exploring">Exploring new regions</option>
          </select>
        </label>
        <p>Who is your favourite Pokemon?</p>
        <label for="togepi">
          <input
            id="togepi"
            type="radio"
            name="favorite"
            value="togepi"
          />Togepi!
        </label>
        <label for="pikachu">
          <input
            id="pikachu"
            type="radio"
            name="favorite"
            value="pikachu"
          />Pikachu
        </label>
        <label for="other">
          <input id="other" type="radio" name="favorite" value="other" />Other
        </label>
        <p>Which communications do you want to receive?</p>
        <label for="newsletter">
          <input
            id="newsletter"
            type="checkbox"
            name="communications"
            value="newsletter"
          />Newsletter
        </label>
        <label for="events">
          <input
            id="events"
            type="checkbox"
            name="communications"
            value="events"
          />Event updates
        </label>
        <label for="updates">
          <input
            id="updates"
            type="checkbox"
            name="communications"
            value="updates"
          />Club updates
        </label>
        <p>Any other information you would like to share?</p>
        <textarea id="additional-information" rows="4" cols="50">
You can provide comments, suggestions, or feedback here.</textarea
        >
        <p>
          <em
            >Please note: This form is a proof of concept. Submitting the form
            will not actually transmit your data.</em
          >
        </p>
        <button type="Submit" id="submit">Submit</button>
      </form>
    </main>
  </body>
  <footer>
    <a href="../">Return to Project List</a> |
    <a href="https://www.nhcarrigan.com">Return to HomePage</a>
  </footer>
</html>
```

```css
main {
  text-align: center;
  background-color: #92869c;
  background-blend-mode: lighten;
  max-width: 500px;
  margin: 20px auto;
  border-radius: 50px;
  box-shadow: 10px 10px rgba(0, 0, 0, 0.5);
  color: black;
}
body {
  text-align: center;
  background: #3a3240;
  color: white;
}
input,
textarea,
select,
button {
  background: #3a3240;
  color: white;
}
a {
  color: white;
}
```
