---
id: 687ad7be1ff45032ccf1d92f
title: 构建主题切换器
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

在本实验中，你将构建一个在不同主题之间切换的应用。当你切换到不同主题时，背景颜色将改变，并显示不同的消息。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个带有文本 `Switch Theme` 的 `button` 元素。
2. 你的 `button` 元素应具有以下属性：
   - 一个 `id` 属性设置为 `"theme-switcher-button"`。
   - 设置为 `"true"` 的 `aria-haspopup` 属性。
   - 设置为 `"false"` 的 `aria-expanded` 属性。
   - 设置为 `"theme-dropdown"` 的 `aria-controls` 属性。
3. 你应该有一个带有以下属性的 `ul` 元素：
   - 一个 `id` 属性设置为 `"theme-dropdown"`。
   - 设置为 `"menu"` 的 `role` 属性。
   - 设置为 `"theme-switcher-button"` 的 `aria-labelledby` 属性。
   - 一个 `hidden` 属性。
4. 你的 `ul` 元素应至少包含两个 `li` 元素，这些元素的 `role` 属性设置为 `"menuitem"`，并且文本内容表现不同的主题。
5. 你的每个 `li` 元素都应有一个以 `theme-` 开头并以你为该 `li` 元素文本设置的主题结尾的 `id` 属性。例如，如果你的某个主题是 `Light`，那么你的 `id` 应该是 `theme-light`。
6. 你应该有一个 `aria-live` 属性设置为 `"polite"` 的元素。
7. 你应该有一个包含至少两个对象的 `themes` 数组，每个对象都包含 `name` 和 `message` 属性。`name` 将表现为不同的主题，`message` 会在主题切换时显示。你可以自由选择 `name` 和 `message` 的值，但 `name` 的值应与 `li` 项中设置的选项之一匹配。
8. 当用户点击 `#theme-switcher-button` 时，`#theme-dropdown` 元素应显示可用的主题。
9. 当显示 `#theme-dropdown` 元素时，`button` 元素的 `aria-expanded` 属性应设置为 `"true"`。
10. 当用户点击 `#theme-switcher-button` 且 `#theme-dropdown` 元素显示时，菜单应关闭，且应添加 `hidden` 属性，并将 `aria-expanded` 属性设置为 `"false"`。
11. 当用户点击 `#theme-switcher-button` 并选择一个主题时，应将对应的 `theme-<name>` 类添加到 `body` 元素中，并且你在 `themes` 数组中设置的相关主题消息应显示在 `aria-live="polite"` 元素中。

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

你应该有一个 `button` 元素。

```js
assert.exists(document.querySelector("button"));
```

你的 `button` 元素应具有 `Switch Theme` 文本。

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

你的 `button` 元素应具有设置为 `"theme-switcher-button"` 的 `id` 属性。

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

你的 `button` 元素应具有设置为 `"true"` 的 `aria-haspopup` 属性。 

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

你的 `button` 元素应具有设置为 `"false"` 的 `aria-expanded` 属性。 

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

你的 `button` 元素应具有设置为 `"theme-dropdown"` 的 `aria-controls` 属性。

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

你应该有一个 `ul` 元素。

```js
assert.exists(document.querySelector("ul"));
```

你的 `ul` 元素应具有设置为 `"theme-dropdown"` 的 `id` 属性。

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

你的 `ul` 元素应具有设置为 `"menu"` 的 `role` 属性。

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

你的 `ul` 元素应具有设置为 `"theme-switcher-button"` 的 `aria-labelledby` 属性。

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

你的 `ul` 元素应具有 `hidden` 属性。

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

你的 `ul` 元素应至少包含两个 `li` 元素。

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

你的每个 `li` 元素都应具有设置为 `"menuitem"` 的 `role` 属性。

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

你的每个 `li` 元素都应包含一段表现主题的文本。

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

你的每个 `li` 元素都应有一个以 `theme-` 开头并以该元素对应主题结尾的 `id` 属性。确保你的 `id` 值全部为小写。

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

你应该有一个 `aria-live` 属性设置为 `"polite"` 的元素。

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

你应该有一个 `themes` 数组。

```js
assert.isArray(themes);
```

你的 `themes` 数组应至少包含两个对象。

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

你的 `themes` 数组应至少包含两个对象，每个对象都有 `name` 和 `message` 属性。 

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

你的 `name` 和 `message` 属性值应为非空的字串。

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

你的每个 `name` 属性值都应匹配你为 `li` 元素提供的某个值。确保这些值全部为小写。

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

当用户点击 `#theme-switcher-button` 以显示主题选项时，应从 `#theme-dropdown` 元素中移除 `hidden` 属性。

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

当显示 `#theme-dropdown` 元素时，`button` 元素的 `aria-expanded` 属性应设置为 `"true"`。

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

当用户点击 `#theme-switcher-button` 且 `#theme-dropdown` 元素显示时，下拉菜单应设置回 `hidden`。

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

当用户点击 `#theme-switcher-button` 且 `#theme-dropdown` 元素显示时，`aria-expanded` 属性应设置为 `"false"`。

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

当用户点击 `#theme-switcher-button` 并选择一个主题时，应将对应的 `theme-<name>` 类添加到 `body` 元素。

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

当用户点击 `#theme-switcher-button` 并选择一个主题时，应在 `aria-live="polite"` 元素中显示与所选 `theme` 来自 `themes` 数组的相关 `message`。


```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
