---
id: 669e81368e52b3a5c35a2dc5
title: 建置一個影片彙整頁面
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `main` 元素作為 `body` 元素的唯一子元素。
1. 你應該有一個包含你頁面主題的 `h1` 元素。
1. 你應該在你的 `h1` 元素下方放置一段介紹頁面主題的段落。
1. 你應該在第一段落下方有三個 `section` 元素。
1. 每個區段應該依序包含一個 `h2` 元素、一段文字，以及一個 `iframe` 元素。
1. 三個 `iframe` 元素的 `src` 屬性應設定為有效的影片。
1. 每個 `iframe` 元素也應該有一個 `title` 屬性來描述嵌入的內容，並且有一個 `height` 屬性和一個 `width` 屬性來設定元素的適當大小。

# --hints--

你應該在你的 `body` 元素內有一個 `main` 元素。

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

你的 `main` 元素應該是 `body` 元素的唯一子元素。

```js
assert.equal(document.querySelector('body').children.length, 1);
```

你應該在 `main` 元素內有一個包含頁面主題的 `h1` 元素。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你應該在你的 `h1` 元素下方放置一段介紹頁面主題的段落。

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

你應該在你的第一個 `p` 元素下方有三個 `section` 元素。

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

每個 `section` 元素應該以一個作為該區段標題的 `h2` 元素開始。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

每個 `section` 元素應包含一個 `p` 元素，作為其第二個子元素來介紹影片內容。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

每個 `section` 元素的第三個子元素應包含一個 `iframe` 元素。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

你的第一個 `iframe` 元素應該有一個設定為有效影片的 `src` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

你的第一個 `iframe` 元素應該有一個 `title` 屬性來描述嵌入的內容。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

你的第一個 `iframe` 元素應該有一個 `height` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

你的第一個 `iframe` 元素應該有一個 `width` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

你的第二個 `iframe` 元素應該有一個設定為有效影片的 `src` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

你的第二個 `iframe` 元素應該有一個 `title` 屬性來描述嵌入的內容。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

你的第二個 `iframe` 元素應該有一個 `height` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

你的第二個 `iframe` 元素應該有一個 `width` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

你的第三個 `iframe` 元素應該有一個 `src` 屬性設定為有效的影片。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

你的第三個 `iframe` 元素應該有一個 `title` 屬性來描述嵌入的內容。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

你的第三個 `iframe` 元素應該有一個 `height` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

你的第三個 `iframe` 元素應該有一個 `width` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Frontend Web Development</h1>
        <p>Frontend Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Frontend
            developers
            ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
