---
id: 669e81368e52b3a5c35a2dc5
title: 創建一個視頻集頁面
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. 你應該有一個 `main` 元素，作爲 `body` 元素的唯一子元素。
1. 你應該有一個包含頁面主題的 `h1` 元素。
1. 你應該在 `h1` 元素下方有一個介紹頁面主題的段落。
1. 你應該在第一個段落下方添加三個 `section` 元素。
1. 每個部分應該依次包含一個 `h2` 元素、一個段落和一個 `iframe` 元素。
1. 三個 `iframe` 都應該有設置爲有效視頻的 `src` 屬性。
1. 每個 `iframe` 元素也都應該有一個用於描述嵌入內容的 `title` 屬性，以及設置爲合適大小的 `height` 屬性和 `width` 屬性。

# --hints--

你應該在 `body` 元素中包含一個 `main` 元素。

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

你的 `main` 元素應該是 `body` 元素的唯一子元素。

```js
assert.equal(document.querySelector('body').children.length, 1);
```

你應該在 `main` 元素中加入一個 `h1` 元素，用於頁面主題。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你應該在 `h1` 元素下方有一個介紹頁面主題的段落。

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

你應該在第一個 `p` 元素下方添加三個 `section` 元素。

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

每個 `section` 元素都應該以一個 `h2` 開頭，作爲該部分的標題。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

每個 `section` 元素都應該包含一個 `p` 元素作爲它的第二個子元素，用於介紹視頻內容。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

每個 `section` 元素都應該包含一個 `iframe` 元素，作爲它的第三個子元素。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

你的第一個 `iframe` 元素應該有一個設置爲有效視頻的 `src` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

你的第一個 `iframe` 元素應該有一個描述嵌入內容的 `title` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

你的第一個 `iframe` 元素應該有一個 `height` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

你的第一個 `iframe` 元素應該有一個 `width` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

你的第二個 `iframe` 元素應該有一個設置爲有效視頻的 `src` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

你的第二個 `iframe` 元素應該有一個描述嵌入內容的 `title` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

你的第二個 `iframe` 元素應該有一個 `height` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

你的第二個 `iframe` 元素應該有一個 `width` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

你的第三個 `iframe` 元素應該有一個設置爲有效視頻的 `src` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

你的第三個 `iframe` 元素應該有一個描述嵌入內容的 `title` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

你的第三個 `iframe` 元素應該有一個 `height` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

你的第三個 `iframe` 元素應該有一個 `width` 屬性。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front End Web Development</h1>
        <p>Front End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front End
            developers
            ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
