---
id: 694bf6eb530e19a9c48c59f7
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

文字列の末尾からスライスするために負の数を使うこともできます。例えば、`-1` は最後の文字を、`-2` は最後から2番目の文字を表します。

文字列の最後の3文字を取得するには、`string[-3:]` を使います。コロンの後に停止パラメータが省略されていることに注目してください。これはスライスが文字列の境界まで続くことを意味します。

`last_three` という名前の変数を作成してください。負のインデックスを使って、ID番号を表す `employee_code` から最後の3文字を抽出します。最後に、`last_three` をターミナルに出力してください。

これで社員カードのワークショップは完了です。

# --hints--

`last_three` という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

負のインデックスを使って `employee_code` の最後の3文字をスライスし、それを `last_three` 変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

`last_three` をターミナルに出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
