---
id: 694bf6eb530e19a9c48c59f4
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

`+` を使って多くの文字列を連結したり、`str()` を使って数字を変換したりすると、コードが複雑で読みづらくなります。

Python 3.6 ではこれを解決するために **f-strings** が導入されました。開き引用符の前に `f` を付けることで、波括弧 `{}` で表される置換フィールドの中に変数や式を入れられます。例えば：

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

変数 `employee_card` を作成し、`Employee:` に続けてスペースと `full_name` 変数の値を表示する f-string を代入してください。

# --hints--

`employee_card` という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

`employee_card` 変数には `f'Employee: {full_name}'` の値が入っているはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
