---
id: 694bf6eb530e19a9c48c59f5
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

文字列を扱うとき、特定の部分を取り出す必要がよくあります。これを**スライス**と呼びます。

構文は `string[start:stop]` で、次のようになっています：

* `start` はスライスが始まるインデックス（**含む**）です。
* `stop` はスライスが終わるインデックス（**含まない**）です。

例えば、`text = 'Python'` の場合、`text[0:2]` は `'Py'` を返します。

`employee_code` を `'DEV-2026-JD-001'` として定義してください。その後、変数 `department` を作成し、`employee_code` のインデックス `0` から `3` までのスライスを代入してください。最後に `department` をターミナルに出力してください。

# --hints--
`employee_code` という名前の変数が必要です。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

`DEV-2026-JD-001` という文字列を `employee_code` 変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

`department` という名前の変数が必要です。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

`employee_code` の最初の3文字をスライスして `department` 変数に代入してください。

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

`department` をターミナルに出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
