---
id: 694bf3af05d10fd89dab05cc
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

Pythonでは、`+`演算子を使って文字列を結合できます。これを連結と呼びます。例を示します。

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

`full_name`と`first_name`を連結して、`last_name`という変数を作成してください。その後、`full_name`を出力してください。

# --hints--

`full_name`という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

`first_name`と`last_name`を連結してできた文字列を、`full_name`変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

`full_name`変数を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
