---
id: 694bf6eb530e19a9c48c59f1
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

住所がまだ不完全なようです。従業員が住んでいるアパートの番号も追加したいので、変数を修正してください。

既存の文字列変数の末尾に内容を追加したい場合は、**複合代入**演算子 `+=` を使うことができます。これは `var = var + 'new text'` と書くよりも短くなります。例えば：

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

文字列は不変なので、この操作は元の文字列を変更しません。代わりに新しい文字列を作成して変数に再代入します。

`+=` 演算子を使って、文字列 `, Apartment 4B` を `address` 変数に追加してください。

# --hints--

文字列 `+=` を `, Apartment 4B` 変数に追加するには、`address` 演算子を使うべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
