---
id: 694bf6eb530e19a9c48c59f6
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

文字列の先頭だけでなく、どの部分からでもスライスできます。これは多くの場合に役立ちます。

変数`year_code`を作成し、`employee_code`のインデックス`4`から`8`までのスライスを代入してください。これにより`2026`が抽出されます。

次に、変数`initials`を作成し、`employee_code`のインデックス`9`から`11`までのスライスを代入してください。これにより`JD`が抽出されます。

最後に、両方の変数をターミナルに出力してください。

# --hints--

`year_code`という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

`employee_code`をインデックス`4`から`8`までスライスし、その結果を`year_code`変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

`initials`という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

`employee_code`をインデックス`9`から`11`までスライスし、その結果を`initials`変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

`year_code`をターミナルに出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

`initials`をターミナルに出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
