---
id: 694bf6eb530e19a9c48c59f3
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

今度は `str()` 関数をもう一度使います。age と同じように、数値の変数は他のテキストと連結する前に文字列に変換する必要があります。

`experience_years` という名前の変数を作成し、整数 `5` を代入してください。

次に、`experience_info` という変数を作成します。`'Experience: '`、`experience_years` 変数（文字列に変換済み）、`' years'` を連結してできた文字列を代入してください。結果をターミナルに出力してください。

# --hints--

`experience_years` という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

`5` 変数に整数 `experience_years` を代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

`experience_info` という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

`'Experience: '` 変数に `str(experience_years)`、`' years'`、`experience_info` を連結してできた文字列を代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

`experience_info` を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
