---
id: 694bf6eb530e19a9c48c59f2
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

`employee_info` の代入を更新して、文字列 ` years old` を末尾に連結し、文を完成させてください。文字列の先頭にスペースを含めることを忘れないでください。

最後に、`employee_info` を出力してください。

# --hints--

`' years old'` 変数の末尾に `employee_info` を連結してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

`employee_info` を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
