---
id: 694bf6ea530e19a9c48c59f0
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

`'John' + 'Doe'` のように2つの文字列を連結すると、結果はスペースなしの `'JohnDoe'` になります。これを直すには、その間にスペースを含む文字列（`' '`）を連結する必要があります。

`full_name` 変数を更新して、`first_name`、スペース、`last_name` を連結するようにしてください。

# --hints--

`' '` と `first_name` の間にスペース（`last_name`）を追加してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
