---
id: 695c9f21a8b1d53cf1b97841
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`employee_age` という名前の変数を作成し、整数 `28` を代入してください。


# --hints--

`employee_age` という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

整数 `28` を `employee_age` 変数に代入してください。テキストを引用符で囲まないでください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
