---
id: 695c96f82af4fb78e59fa9ed
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

文字列はテキストデータをストアするために使われる文字のシーケンスです。前のレッスンで学んだように、文字列はテキストをシングルクォート（`'`）または二重引用符（`"`）で囲むことで作成できます。例えば：

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

文字列`first_name`をストアする変数`John`と、文字列`last_name`をストアする変数`Doe`を作成してください。その後、`first_name`と`last_name`を出力してください。

# --hints--

`first_name`という名前の変数を用意してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

`John`変数に文字列`first_name`を代入してください。テキストはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

`last_name`という名前の変数を用意してください。

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

`Doe`変数に文字列`last_name`を代入してください。テキストはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

`first_name`変数をターミナルに出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

`last_name`変数をターミナルに出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
