---
id: 695c9c7c61d446f1451e0ae7
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

次に、従業員の住所をストアする変数`address`を作成してください。文字列`123 Main Street`を代入し、最後に`address`を出力してください。

# --hints--

`address`という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

`123 Main Street`変数に文字列`address`を代入してください。テキストはシングルクォートか二重引用符で囲むことを忘れないでください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

`address`変数を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```
