---
id: 695df6480de23758af5a29cd
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

次に、従業員の年齢を表示する文字列を作成します。

まず、`employee_info` という変数を作成し、次のものを連結した結果を代入してください。

- `full_name` 変数
- 前後にスペースがある `is` という文字列

# --hints--

`employee_info` という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

`full_name + ' is '` を `employee_info` 変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
