---
id: 695ea1e55c0e619e597fc5c9
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

ご覧のとおり、Pythonは`TypeError: can only concatenate str (not "int") to str`を発生させました。これは、Pythonがテキスト（文字列）と数字（整数）を直接連結することを許可していないためです。

これを修正するには、まず`str()`関数を使って数字を文字列に変換する必要があります。この関数はオブジェクトの文字列バージョンを返します。

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

`employee_info`の代入を更新して、`employee_age`を`str(employee_age)`を使って文字列に変換してください。

# --hints--

`employee_info`変数の代入を修正して、`employee_age`を`str(employee_age)`で文字列に変換してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
