---
id: 695fcfbba648b3a96a9495a4
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

ターミナルの出力が混み合ってきました。続ける前に、整理しましょう。

コードからすべての `print()` 文を削除してください。

# --hints--

コードに `print` の呼び出しが含まれていないことを確認してください。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_calls("print")) == 0`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'  
last_name = 'Doe'  
--fcc-editable-region--  
print(first_name)  
print(last_name)  
full_name = first_name + ' ' + last_name  
print(full_name)  
address = '123 Main Street'  
address += ', Apartment 4B'  
print(address)  
--fcc-editable-region--
```
