---
id: 695fd3dde03273875e5d6709
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

`employee_info` の代入を更新して、末尾に `employee_age` も連結してください。

そうすると、ターミナルに `TypeError` が表示されます。次のステップでは、それを修正します。

# --hints--

`employee_info` の代入を更新して、末尾に `employee_age` を連結してください。

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
