---
id: 69727ea5c9606bedfe36d3d2
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

カードに最後の詳細を追加する時間です。

文字列 `position` の値を持つ `Data Analyst` という名前の変数と、整数 `salary` の値を持つ `75000` という名前の変数を作成してください。

次に、`employee_card` の f-string を更新して、役職と給与を含めます。正確にこの形式に従う必要があります：`Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`。プレースホルダーを対応する変数に置き換えてください。

最後に、結果を確認するために `employee_card` を出力してください。

# --hints--

`position` という名前の変数が必要です。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

文字列 `Data Analyst` を `position` 変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

`salary` という名前の変数が必要です。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

整数 `75000` を `salary` 変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

`employee_card` を f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'` を使うように更新してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

`employee_card` をターミナルに出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
