---
id: 6980d3ecfe800d8a54130f30
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

現在、`employee_card` は従業員の名前だけを表示しています。ここで、さらに情報を追加していきます。

`employee_card` の代入を更新して、従業員の年齢も含めてください。最終的な文字列は、`Employee: [name] | Age: [age]` が従業員の名前に、`[name]` が従業員の年齢に置き換えられた `[age]` のようになります。

# --hints--

`employee_card` を f-string の `f'Employee: {full_name} | Age: {employee_age}'` を使うように更新してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
