---
id: 6823e637c1c0ed56f781b4fc
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Schließlich macht JavaScript Ihre Webseite interaktiv – es ermöglicht Ihnen, der Seite zu sagen, was sie tun soll, wenn jemand auf einen Button klickt, ein Formular absendet oder viele andere Dinge ausführt.

Für den letzten Schritt des Workshops fügen Sie der Seite ein weiteres `h3`- und `p`-Element hinzu, das JavaScript beschreibt.

Fügen Sie zuerst ein `h3`-Element mit dem Text hinzu:

```md
Introduction to JavaScript
```

Dann fügen Sie unterhalb dieses `h3`-Elements ein `p`-Element mit folgendem Text hinzu:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

Sie sollten der Seite ein drittes `h3`-Element hinzufügen.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Der Text Ihres `h3`-Elements sollte `Introduction to JavaScript` sein.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

Sie sollten der Seite ein viertes `p`-Element hinzufügen.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Ihr `p`-Element sollte den Text `JavaScript adds interactivity to a webpage` haben.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
