---
id: 6823d6244511f252c8300eed
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Wenn Sie einen Absatz zu einer Webseite hinzufügen müssen, können Sie das `p`-Element so verwenden:

```html
<p>This is a paragraph element.</p>
```

Verwandeln Sie `Learn the skills to become a full-stack developer` in ein Absatz-Element.

# --hints--

Ihr `p`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

Ihr `p`-Element sollte ein schließendes Tag haben. Schließende Tags haben diese Syntax: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

Der Text Ihres `p`-Elements sollte `Learn the skills to become a full-stack developer` sein.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

Ihr `p`-Element sollte unter Ihrem `h2`-Element stehen.

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
