---
id: 6823ac607bfdbc46331b2559
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML steht für HyperText Markup Language. Es ist der Code, der die Struktur und den Inhalt einer Webseite definiert. Dies ist Ihr Code-Editor, in dem Sie HTML schreiben.

Suchen Sie Zeile 1 im Editor und geben Sie diesen Text ein:

```md
Welcome to freeCodeCamp
```

Wenn Sie fertig sind, klicken Sie auf die Schaltfläche „Überprüfen Sie Ihren Code“, um zu sehen, ob alles korrekt ist.

# --hints--

Sie sollten den Text `Welcome to freeCodeCamp` in Ihrem Editor haben. Überprüfen Sie die Rechtschreibung sorgfältig.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
