---
id: 6823e036ea4b71553558c01b
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Beachten Sie, dass jede Überschrift in der Vorschau etwas anders aussieht. Gleiches gilt für den Absatz.

Fügen Sie als Nächstes ein weiteres `p`-Element hinzu, das den folgenden Text anzeigt:

```md
HTML represents the content and structure of a webpage
```

# --hints--

Sie sollten ein zweites `p`-Element haben.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Ihr `p`-Element sollte den Text `HTML represents the content and structure of a webpage` enthalten. Sie können den Text aus den Anweisungen kopieren.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
