---
id: 6823d9ac8bdc3853df65a1ff
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Es gibt sechs Überschriften-Elemente in HTML: `h1` bis `h6`. Sie werden verwendet, um die Wichtigkeit von Abschnitten auf Ihrer Webseite anzuzeigen, wobei `h1` die wichtigste und `h6` die unwichtigste ist.

Fügen Sie unter Ihrem `p`-Element eine `h3`-Überschrift mit dem Text hinzu:

```md
Introduction to HTML
```

# --hints--

Ihr `h3`-Element sollte ein öffnendes `<h3>`-Tag haben.

```js
assert.exists(document.querySelector("h3"));
```

Ihr `h3`-Element sollte ein schließendes `</h3>`-Tag haben.

```js
assert.match(code, /<\/h3\s*\>/);
```

Der Text Ihres `h3`-Elements sollte `Introduction to HTML` sein.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
