---
id: 682ba2318000b62f179bdf04
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML besteht aus Elementen. Das erste, das Sie verwenden werden, ist das `h1`-Element:

```html
<h1>Welcome to freeCodeCamp</h1>
```

Es beginnt mit einem öffnenden Tag (`<h1>`), endet mit einem schließenden Tag (`</h1>`) und enthält den Text, der zwischen den Tags angezeigt wird.

Verwandeln Sie Ihren `Welcome to freeCodeCamp`-Text in ein `h1`-Element, indem Sie davor ein öffnendes Tag und dahinter ein schließendes Tag hinzufügen.

# --hints--

Ihr `h1`-Element sollte ein öffnendes `<h1>`-Tag haben.

```js
assert.exists(document.querySelector("h1"));
```

Ihr `h1`-Element sollte ein schließendes `</h1>`-Tag haben.

```js
assert.match(code, /<\/h1\s*\>/);
```

Ihr `h1`-Element sollte so aussehen: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
