---
id: 682cd206883fc7b25eb539c5
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Beachten Sie, dass das HTML, das Sie im Editor schreiben, in der Vorschau angezeigt wird. In diesem Workshop schreiben Sie das HTML für eine Teilseite eines Lehrplans.

Geben Sie unter Ihrem `h1`-Element auf der leeren Zeile Folgendes ein:

```md
Full-Stack Curriculum
```

# --hints--

Sie sollten den Text `Full-Stack Curriculum` in Ihrem Editor haben. Überprüfen Sie die Rechtschreibung sorgfältig.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
