---
id: 683921c4769fd23dbadec2fe
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Während HTML die Struktur und den Inhalt einer Webseite definiert, wird CSS verwendet, um Stil hinzuzufügen – Dinge wie Farben, Schriftarten, Abstände und Layout.

Fügen Sie unter dem `h3`, das Sie gerade hinzugefügt haben, ein weiteres Absatz-Element mit dem Text ein:

```md
CSS is used to style a webpage
```

# --hints--

Sie sollten genau zwei `h3`-Elemente haben. Achten Sie darauf, die vorhandenen nicht zu entfernen und keine neuen hinzuzufügen.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Der Text Ihres `h3`-Elements sollte `Introduction to CSS` sein.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

Sie sollten ein drittes `p`-Element zur Seite hinzufügen.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Ihr `p`-Element sollte den Text `CSS is used to style a webpage` enthalten. Sie können den Text aus den Anweisungen kopieren.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
--fcc-editable-region--

--fcc-editable-region--
```
