---
id: 6808fe91a5cb0c3355f23027
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você aprenderá como trabalhar com formulários em React construindo um formulário de aplicativo de super-herói.

Algum boilerplate foi fornecido, você só precisa criar o restante do componente React.

Comece retornando um elemento `div` com um `className` de `form-wrap`. Dentro do `div`, crie um elemento `h2` com o texto `Superhero Application Form`. Abaixo disso, crie um elemento `p` com o texto `Please complete all fields`.

# --hints--

Você deve retornar um elemento `div` com o `className` de `form-wrap`.

```js
assert.exists(document?.querySelector("div.form-wrap"));
```

Você deve criar um elemento `h2` dentro do seu elemento `div`.

```js
assert.exists(document?.querySelector(".form-wrap > h2"));
```

Seu elemento `h2` deve conter o texto `Superhero Application Form`.

```js
const h2El = document?.querySelector("h2")
assert.equal(h2El?.textContent, "Superhero Application Form");
```

Você deve criar um elemento `p` dentro do seu elemento `div`. Certifique-se de que venha depois do `h2`.

```js
assert.exists(document?.querySelector(".form-wrap > h2 + p"));
```

Seu elemento `p` deve conter o texto `Please complete all fields`.

```js
const pEl = document?.querySelector("p")
assert.equal(pEl?.textContent, "Please complete all fields");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  --fcc-editable-region--
  
  --fcc-editable-region--

};
```
