---
id: 680900675ae3d54ee19590c7
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Crie outro `label` com o texto `Real Name` e um `input` do tipo `password` dentro dele. Dê ao `input` um `value` de `realName`. 

Então use o evento `onChange` para atualizar `realName` da mesma forma que você atualizou `heroName`. 

# --hints--

Você deve criar um elemento `label`.

```js
const labels = document.querySelectorAll("label");
assert.lengthOf(labels, 2);
```

Seu elemento `label` deve ter `Real Name` como seu texto.

```js
const labels = document.querySelectorAll("label");
assert.equal(labels[1]?.textContent.trim(), "Real Name");
```

Você deve criar um `input` do tipo `password` dentro do seu elemento `label`.

```js
const inputEl = document.querySelectorAll("label > input");
assert.equal(inputEl[1]?.getAttribute("type"), "password");
```

Seu elemento `input` deve ter seu valor definido para a variável de estado `realName`.

```js
assert.match(code, /value\s*=\s*\{\s*realName\s*\}/)
```

Você deve usar o evento `onChange` para definir o valor da variável de estado `realName` para que você possa digitar no campo `input`.

```js
assert.match(code, /onChange\s*=\s*\{\s*(e|\(\s*e\s*\))\s*=>\s*setRealName\s*\(\s*e\s*\.\s*target\s*\.\s*value\s*\)\s*\}/)
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          <label>
              Hero Name
              <input
                type='text'
                value={heroName}
                onChange={e => setHeroName(e.target.value)}
              />
          </label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </div>
      </form>
    </div>
  )
};
```
