---
id: 60eebd07ea685b0e777b5583
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Willkommen beim Projekt Registrierungsformular! Beginnen Sie damit, die `!DOCTYPE html`-Deklaration oben im Dokument hinzuzufügen, damit der Browser weiß, welchen Dokumenttyp er liest.

# --hints--

Ihr Code sollte die `DOCTYPE`-Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Sie sollten nach der `DOCTYPE`-Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Sie sollten den Dokumenttyp als `html` definieren.

```js
assert(code.match(/html/gi));
```

Sie sollten die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/html\s*>/gi));
```

Ihre `DOCTYPE`-Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
