---
id: 60f02e7361b68405e27b62a5
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Fügen Sie innerhalb des `body` einen Kontext für die Überschrift des Inhalts hinzu, indem Sie ein `h1` mit dem Text `Registration Form` ergänzen.

# --hints--

Sie sollten das `h1` innerhalb des `body` hinzufügen.

```js
assert.exists(document.querySelector('body > h1'));
```

Sie sollten dem `h1` den Text `Registration Form` geben.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
