---
id: 60f027099a15b00485563dd2
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie unter dem `DOCTYPE` ein `html`-Element mit einem `lang`-Attribut hinzu, das auf `en` gesetzt ist, damit Sie einen Platz haben, um mit dem Schreiben von Code zu beginnen.

# --hints--

Ihre `DOCTYPE`-Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ihr `html`-Element sollte ein öffnendes Tag und ein `lang`-Attribut mit dem Wert `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Ihre `html`-Tags sollten in der richtigen Reihenfolge sein.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Sie sollten nur ein `html`-Element haben.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
