---
id: 60f0286404aefb0562a4fdf9
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Fügen Sie ein `title`- und ein `meta`-Element innerhalb des `head`-Elements hinzu. Geben Sie Ihrem Projekt den Titel `Registration Form` und fügen Sie dem `charset`-Element das `utf-8`-Attribut mit dem Wert `meta` hinzu.

# --hints--

Ihr Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ihr `title`-Element sollte sich innerhalb des `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > title'));
```

Ihr Projekt sollte den Titel `Registration Form` haben. Denken Sie daran, dass Groß- und Kleinschreibung sowie Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Sie sollten ein `meta`-Element innerhalb des `head`-Elements erstellen.

```js
assert.exists(document.querySelector('head > meta'));
```

Sie sollten dem `meta`-Element ein `charset`-Attribut mit dem Wert `utf-8` geben.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Ihr `meta`-Element sollte ein leeres Element sein, es hat kein End-Tag `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
