---
id: 60f1a5e2d2c23707a4f9a660
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Entfernen Sie jetzt die horizontale Bildlaufleiste, indem Sie den standardmäßigen `body` `margin`, der von einigen Browsern hinzugefügt wird, auf `0` setzen.

# --hints--

Sie sollten `margin` innerhalb des `body`-Elementselektors hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

Sie sollten dem `margin` den Wert `0` zuweisen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
}
--fcc-editable-region--
```
