---
id: 60f84ec41116b209c280ba91
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Mit `type="password"` können Sie das `pattern`-Attribut verwenden, um einen regulären Ausdruck zu definieren, dem das Passwort entsprechen muss, damit es als gültig gilt.

Fügen Sie dem Passwort-`input`-Element ein `pattern`-Attribut hinzu, das verlangt, dass die Eingabe dem Ausdruck `[a-z0-5]{8,}` entspricht.

Der oben stehende reguläre Ausdruck entspricht acht oder mehr Kleinbuchstaben oder den Ziffern `0` bis `5`. Entfernen Sie dann das `minlength`-Attribut und probieren Sie es aus.

# --hints--

Sie sollten dem Passwort-`input`-Element ein `pattern`-Attribut geben.

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

Sie sollten dem `pattern`-Attribut den Wert `[a-z0-5]{8,}` zuweisen.

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

Sie sollten das `minlength`-Attribut vom Passwort-`input`-Element entfernen.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
