---
id: 60f80e0081e0f2052ae5b505
title: Schritt 19
challengeType: 0
dashedName: step-19
---

# --description--

Das Festlegen des `type`-Attributs eines `input`-Elements ist wichtig, damit der Browser weiß, welche Art von Daten er erwarten soll. Wenn das `type` nicht angegeben ist, verwendet der Browser standardmäßig `text`.

Geben Sie den ersten beiden `input`-Elementen ein `type`-Attribut mit dem Wert `text`, dem dritten ein `type`-Attribut mit dem Wert `email` und dem vierten ein `type`-Attribut mit dem Wert `password`.

Der `email`-Typ erlaubt nur E-Mails mit einem `@` und einem `.` in der Domain.
Der `password`-Typ verschleiert die Eingabe und warnt, wenn die Website kein HTTPS verwendet.

# --hints--

Sie sollten dem ersten `input`-Element ein `type`-Attribut mit dem Wert `text` geben.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Sie sollten dem zweiten `input`-Element ein `type`-Attribut mit dem Wert `text` geben.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Sie sollten dem dritten `input`-Element ein `type`-Attribut mit dem Wert `email` geben.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

Sie sollten dem vierten `input`-Element ein `type`-Attribut mit dem Wert `password` geben.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
