---
id: 62cc5b1779e4d313466f73c5
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Verschachteln Sie ein `link`-Element innerhalb des `head`-Elements. Geben Sie ihm ein `rel`-Attribut mit dem Wert `stylesheet` und ein `href`-Attribut mit dem Wert `styles.css`.

# --hints--

Ihr Code sollte ein `link`-Element enthalten.

```js
assert.isNotNull(document.querySelector('link'));
```

Sie sollten Ihre vorhandenen `head`-Tags nicht ändern. Stellen Sie sicher, dass Sie das schließende Tag nicht gelöscht haben.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

Sie sollten ein `link`-Element haben.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Ihr `link`-Element sollte ein Void-Element sein, es darf kein schließendes Tag `</link>` haben.

```js
assert.notMatch(code, /<\/link>/);
```

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
