---
id: 62ff8e998d3e7eae14d6ae3b
title: Schritt 30
challengeType: 0
dashedName: step-30
---

# --description--

Sie müssen bestätigen, dass der Benutzer die Geschäftsbedingungen gelesen hat.

Fügen Sie ein `label`-Element hinzu. Fügen Sie innerhalb des neu erstellten `label`-Elements ein `input`-Element hinzu und setzen Sie das `type`-Attribut auf `checkbox`. Machen Sie dieses `input`-Element `required`, damit sich Benutzer nicht anmelden können, ohne den Geschäftsbedingungen zuzustimmen.

Fügen Sie den Elementen für die Barrierefreiheit ein `id`- und `for`-Attribut mit dem Wert `terms-and-conditions` hinzu.

# --hints--

Sie sollten ein `label` nach dem dritten `fieldset`-Element hinzufügen.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

Sie sollten ein `input` zum `label`-Element hinzufügen.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

Sie sollten ein `type`-Attribut mit dem Wert `checkbox` zum `input`-Element hinzufügen.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

Sie sollten ein `required`-Attribut zum `input`-Element hinzufügen.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

Das `input`-Element sollte ein `id` von `terms-and-conditions` haben.

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

Das `label`-Element sollte ein `for`-Attribut mit dem Wert `terms-and-conditions` haben.

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
