---
id: 60f02e7361b68405e27b62a5
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتعلم كيف تصمم نماذج من خلال إنشاء نموذج تسجيل.

تم توفير كل قالب HTML الجاهز (`DOCTYPE`، `html`، `head`، و`body`) لك.

داخل `body`، أضف عنوانًا ليعطي سياقًا للنموذج باستخدام عنصر `h1` يحتوي على النص `Registration Form`.

# --hints--

يجب أن تضيف `h1` داخل `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

يجب أن تعطي `h1` النص `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
