---
id: 60f030d388cb74067cf291c3
title: خطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أسفل العنوان، استخدم النص التالي داخل عنصر فقرة لتشجيع المستخدمين على التسجيل:

```md
Please fill out this form with the required information
```

# --hints--

يجب إضافة عنصر `p` داخل عنصر `body`.

```js
assert.exists(document.querySelector('body > p'));
```

يجب إضافة عنصر `p` أسفل عنصر `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

يجب أن يحتوي عنصر `p` على النص `Please fill out this form with the required information`.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  --fcc-editable-region--
  
  --fcc-editable-region--
</body>
</html>
```
