---
id: 60f1a9cbd23023082e149fee
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

هذا أفضل. الآن، اجعل الخلفية مريحة للعينين بتغيير `background-color` في `body` إلى `#1b1b32`. ثم، لرؤية النص، غيّر `color` إلى `#f5f6f7`.

# --hints--

يجب إضافة `background-color` داخل محدِّد عنصر `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

يجب إعطاء `background-color` القيمة `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

يجب إضافة `color` داخل محدِّد عنصر `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

يجب إعطاء `color` القيمة `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
