---
id: 60f5c3e399ff1a05629964e4
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

كما يشير العنوان، أنت تنشئ نموذجًا. لذا، بعد عنصر `p`، أدخل عنصر `form` مع خاصية `action` تستهدف `https://register-demo.freecodecamp.org`.

# --hints--

يجب أن تضيف عنصر `form` بجانب عنصر `p`.

```js
assert.exists(document.querySelector('p + form'));
```

يجب أن تعطي عنصر `form` خاصية `action`.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

يجب أن تعطي خاصية `action` القيمة `https://register-demo.freecodecamp.org`.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

يجب أن يحتوي عنصر `form` على وسم الإغلاق `</form>`.

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
