---
id: 60f5d2776c854e069560fbe6
title: خطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

ستحتوي مجموعة الحقول (`fieldset`) الأولى على حقول الاسم والبريد الإلكتروني وكلمة المرور. ابدأ بإضافة أربعة عناصر مسمى (`label`) إلى مجموعة الحقول (`fieldset`) الأولى.

# --hints--

يجب أن تحتوي الصفحة على ثلاثة عناصر `fieldset` في المجموع.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

يجب أن تكون جميع عناصر `fieldset` أبناء مباشرين لعنصر الـ `form`.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

يجب أن تكون العناصر الأربعة `label` كلها أبناء مباشرين لأول عنصر `fieldset`. تأكد من إغلاق عناصر `label`.

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

يجب أن تكون آخر عنصرين من نوع `fieldset` فارغين.

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

تأكد من إغلاق عناصر `label`.

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
