---
id: 60f805f813eaf2049bc2ceea
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

قم بتسكين عنصر `input` داخل كل `label`. تأكد من إضافة كل `input` بعد نص الـ `label`، وأضف مسافة بعد النقطتين.

# --hints--

يجب إضافة أربعة عناصر `input` إلى عنصر `fieldset`.

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

يجب تسكين عناصر `input` داخل عناصر `label`.

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

يجب إضافة أول `input` بعد نص الـ `label` `Enter Your First Name:`، وأضف مسافة بعد النقطتين.

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

يجب إضافة ثاني `input` بعد نص الـ `label` `Enter Your Last Name:`، وأضف مسافة بعد النقطتين.

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

يجب إضافة ثالث `input` بعد نص الـ `label` `Enter Your Email:`، وأضف مسافة بعد النقطتين.

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

يجب إضافة رابع `input` بعد نص الـ `label` `Create a New Password:`، وأضف مسافة بعد النقطتين.

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
