---
id: 60f84ec41116b209c280ba91
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

باستخدام `type="password"` يمكنك استخدام خاصية `pattern` لتعريف تعبير نمطي يجب أن يطابقه كلمة المرور ليُعتبر صالحًا.

أضف خاصية `pattern` إلى عنصر `input` الخاص بكلمة المرور لتطلب أن يطابق المُدخَل التعبير التالي: `[a-z0-5]{8,}`

ما سبق هو تعبير نمطي يطابق ثمانية أحرف صغيرة أو أكثر أو الأرقام من `0` إلى `5`. بعد ذلك، احذف خاصية `minlength` وجرب النتيجة.

# --hints--

يجب أن تُضيف خاصية `pattern` إلى عنصر `input` الخاص بكلمة المرور.

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

يجب أن تعطي خاصية `pattern` القيمة `[a-z0-5]{8,}`.

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

يجب إزالة الخاصية `minlength` من عنصر `input` الخاص بكلمة المرور.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
--fcc-editable-region--
      <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
