---
id: 60f80e0081e0f2052ae5b505
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

تحديد خاصية `type` لعنصر `input` مهم للمتصفح ليعرف نوع البيانات التي يجب أن يتوقعها. إذا لم تُحدد خاصية `type`، سيستخدم المتصفح القيمة الافتراضية `text`.

أعطِ أول عنصرين من عناصر `input` خاصية `type` بقيمة `text`، والثالث خاصية `type` بقيمة `email`، والرابع خاصية `type` بقيمة `password`.

نوع `email` يسمح فقط بالبريد الإلكتروني الذي يحتوي على `@` و `.` في النطاق.
نوع `password` يخفي المُدخَل، وينبه إذا لم يستخدم الموقع HTTPS.

# --hints--

يجب أن تعطي أول عنصر `input` خاصية `type` بقيمة `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

يجب أن تعطي ثاني عنصر `input` خاصية `type` بقيمة `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

يجب أن تعطي ثالث عنصر `input` خاصية `type` بقيمة `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

يجب أن تعطي رابع عنصر `input` خاصية `type` بقيمة `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
      <label for="email">Enter Your Email: <input id="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
