---
id: 60fab4a123ce4b04526b082b
title: الخطوة 24
challengeType: 0
dashedName: step-24
---

# --description--

يمكن للمستخدمين حاليًا إرسال النموذج دون تحديد مُدخلات الراديو. رغم أنك استخدمت سابقًا الخاصية `required` للدلالة على أن المُدخل مطلوب، إلا أنها لن تعمل في هذه الحالة لأن إضافة `required` إلى كلا المُدخلين ستنقل معلومات خاطئة للمستخدمين.

لحل هذه المشكلة، يمكنك توضيح المطلوب بإضافة عنصر `legend` يحتوي على النص `Account type (required)` قبل عناصر `label` داخل مجموعة الحقول الثانية `fieldset`. بعد ذلك، أضف خاصية `checked` إلى مُدخل `Personal` لضمان إرسال النموذج مع البيانات المطلوبة.

# --hints--

يجب أن تبقى عناصر `input` من نوع `type` بـ `radio`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

يجب إضافة عنصر `legend` داخل العنصر الثاني من نوع `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

يجب إضافة نص `Account type (required)` إلى عنصر `legend`.

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

يجب أن تضيف خاصية `checked` إلى مُدخَل `Personal`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      
      <label><input type="radio" name="account-type" value="personal" /> Personal</label>
--fcc-editable-region--
      <label><input type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
