---
id: 62ff8e998d3e7eae14d6ae3b
title: الخطوة 26
challengeType: 0
dashedName: step-26
---

# --description--

يجب التأكد من أن المستخدم قد قرأ الشروط والأحكام.

أضف عنصر `label`. داخل عنصر `label` الجديد أضف عنصر `input` واضبط الخاصية `type` على `checkbox`. اجعل هذا العنصر `input` بـ `required` حتى لا يتمكن المستخدمون من التسجيل دون الموافقة على الشروط والأحكام.

أضف خاصيتي `id` و `for` بالقيمة `terms-and-conditions` للعناصر من أجل إمكانية الوصول.

# --hints--

يجب إضافة `label` بعد عنصر `fieldset` الثالث.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

يجب إضافة `input` إلى عنصر `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

يجب إضافة خاصية `type` بقيمة `checkbox` إلى عنصر `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

يجب إضافة خاصية `required` إلى عنصر `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

يجب أن يحتوي عنصر `input` على خاصية `id` بقيمة `terms-and-conditions`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

يجب أن يحتوي عنصر `label` على خاصية `for` بقيمة `terms-and-conditions`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
--fcc-editable-region--
    
--fcc-editable-region--
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
