---
id: 65045fa2267ce52da6a73676
title: الخطوة 25
challengeType: 0
dashedName: step-25
---

# --description--

اتبع أفضل الممارسات لإمكانية الوصول بربط عناصر `input` وعناصر `label` في مجموعة الحقول الثانية `fieldset`.

استخدم `personal-account` و `business-account` كقيم لخاصيتي `id` المقابلتين.

# --hints--

يجب أن يحتوي أول عنصر `input` على خاصية `id` بقيمة `personal-account`.

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(2) input')?.[0]?.matches('#personal-account'))
```

يجب أن يحتوي ثاني عنصر `input` على خاصية `id` بقيمة `business-account`.

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(2) input')?.[1]?.matches('#business-account'))
```

يجب أن يحتوي أول عنصر `label` على خاصية `for` بقيمة `personal-account`.

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(2) label')?.[0]?.matches('label[for="personal-account"]'))
```

يجب أن يحتوي ثاني عنصر `label` على خاصية `for` بقيمة `business-account`.

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(2) label')?.[1]?.matches('label[for="business-account"]'))
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
--fcc-editable-region--
      <label><input type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label><input type="radio" name="account-type" value="business" /> Business</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
