---
id: 616968c2c94c8071b349c146
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在 `body` 中，嵌套一个 `h1` 元素，文本为 `CSS Color Markers`。

# --hints--

你的代码应该有一个 `h1` 元素。

```js
const title = document.querySelector('h1');
assert.exists(title);
```

你应该有一个 `h1` 开始标签。

```js
assert.match(code, /<h1\s*>/i);
```

你应该有一个 `h1` 结束标签。

```js
assert.match(code, /<\/h1\s*>/i);
```

你的 `h1` 元素应该在 `body` 元素内。

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

你的 `h1` 元素应该包含文本 `CSS Color Markers`。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

请记住，`h1` 文本的大小写和拼写很重要。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
