---
id: 616d3a67ccf800ad94ec89ae
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

现在，你将添加一些元素，最终将它们设置为彩色马克笔。

首先，在 `body` 元素中，添加一个 `div` 元素，其 `class` 属性设置为 `container`。 确保 `div` 元素位于 `h1` 元素下方。

# --hints--

你应该有一个 `div` 开始标签。

```js
assert.match(code, /<div.*>/i);
```

你应该有一个 `div` 结束标签。

```js
assert.match(code, /<\/div\s*>/i);
```

你的 `div` 元素应该在你的 `body` 元素中。

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

你的 `div` 元素应该在 `h1` 元素之后。

```js
assert.exists(document.querySelector('h1 + div'));
```

你的 `div` 元素应该有一个 `class` 属性，设置为 `container`。

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
