---
id: 617bbb6b97a83f6d1f7d6e38
title: 步骤 26
challengeType: 0
dashedName: step-26
---

# --description--

现在，你已经练习了间色，让我们回顾一下如何创建 <dfn> 复色</dfn>。 复色是由一种原色和一种邻近的间色组合而成。

要创建复色橙色，请更新 `.one` CSS 规则中的 `rgb` 函数，将红色设置为最大值（`255`），将绿色设置为 `127`。

# --hints--

你的 `.one` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(255, 127, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 255, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
