---
id: 617b954d9f4f6217a749380e
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

现在，在马克笔和它们所在的 `container` 元素的边缘之间添加更多的垂直空间。

在 `.container` CSS 规则中，使用简写 `padding` 属性添加上下内边距，各为 `10px`，并将左右内边距设置为 `0`。 这类似于你之前使用的简写 `margin` 属性。

# --hints--

你不应从 `.container` CSS 规则中删除 `background-color` 属性及其值。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

你的 `.container` CSS 规则应该有一个 `padding` 属性，设置为 `10px 0`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.padding, '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
