---
id: 617bb5624a75e86463b7e638
title: 步骤-23
challengeType: 0
dashedName: step-23
---

# --description--

<dfn>间色</dfn>是将原色组合在一起得到的颜色。 在上一步中，由于你改变了红色、绿色和蓝色的值，你可能已经注意到了一些间色。

要创建第一个间色，黄色，请更新 `.one` CSS 规则中的 `rgb` 函数，将纯红和纯绿结合起来。

# --hints--

你的 `.one` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(255, 255, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
