---
id: 617b65579ce424bf5f02ca73
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

As you have learned in the lessons, there are two main color models: the <dfn>additive</dfn> RGB (red, green, blue) model, used in electronic devices, and the <dfn>subtractive</dfn> CMYK (cyan, magenta, yellow, black) model, used in print.

在这个项目中，将使用 RGB 模型。 这意味着颜色从黑色开始，随着红色、绿色和蓝色的不同层次而变化。 一个简单的方法就是使用 CSS `rgb` 函数。

创建一个定位 `container` 类的新 CSS 规则，并使用 `rgb(0, 0, 0)` 将其 `background-color` 设置为黑色。

# --hints--

你应该使用一个类选择器来定位类 `container`。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.container'));
```

你的 `.container` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(0, 0, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
