---
id: 618a00ed1ca871a2b3aca0cb
title: 第 30 步
challengeType: 0
dashedName: step-30
---

# --description--

既然你已经了解了色轮上的所有原色、间色和复色，就更容易理解其他色彩理论概念及其对设计的影响了。

首先，在 `.one`，`.two` 和 `.three` CSS 规则中，调整 `rgb` 函数中的数值，将每个元素的 `background-color` 设置为纯黑色。 请记住，`rgb` 函数使用的是加色模型，即颜色从黑色开始，随着红、绿、蓝三色值的增加而变化。

# --hints--

你的 `.one` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(0, 0, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(0, 0, 0)');
```

你的 `.two` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(0, 0, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

你的 `.three` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(0, 0, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
