---
id: 587d7fb2367417b2b2512bf6
title: Obtén la entrada de parámetros de consulta del cliente
challengeType: 2
forumTopicId: 301512
dashedName: get-query-parameter-input-from-the-client
---

# --description--

Otra forma común de obtener entrada del cliente es codificando los datos después de la ruta, usando una cadena de consulta. La cadena de consulta está delimitada por un signo de interrogación (?), e incluye pares campo=valor. Cada par está separado por un ampersand (&). Express puede interpretar los datos de la cadena de consulta y llenar el objeto `req.query`. Algunos caracteres, como el porcentaje (%), no pueden estar en URLs y deben codificarse en un formato diferente antes de enviarlos. Si usas la API desde JavaScript, puedes usar métodos específicos para codificar/decodificar estos caracteres.

<blockquote>ruta: '/library'<br>URL_de_solicitud_real: '/library?userId=546&#x26;bookId=6754' <br>req.query: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

Construye un endpoint de API, montado en `GET /name`. Responde con un documento JSON, con la estructura `{ name: 'firstname lastname'}`. Los parámetros de nombre y apellido deben codificarse en una cadena de consulta, por ejemplo `?first=firstname&last=lastname`.

**Nota:** En el siguiente ejercicio va a recibir datos de una solicitud POST, a la misma ruta `/name`. Si lo deseas, puedes utilizar el método `app.route(path).get(handler).post(handler)`. Esta sintaxis permite encadenar diferentes manejadores de verbos en la misma ruta. Puedes ahorrar un poco de escritura y tener un código más limpio.

# --hints--

Prueba 1: Tu endpoint de API debería responder con `{ "name": "Mick Jagger" }` cuando se llame al endpoint `/name` con `?first=Mick&last=Jagger`

```js
  const response = await fetch(code + '/name?first=Mick&last=Jagger');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "GET /name" route does not behave as expected'
  );
```

Prueba 2: Tu endpoint de API debe responder con `{ "name": "Keith Richards" }` cuando se llame al endpoint `/name` con `?first=Keith&last=Richards`

```js
  const response = await fetch(code + '/name?last=Richards&first=Keith');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "GET /name" route does not behave as expected'
  );
```

