---
id: 587d7fb2367417b2b2512bf6
title: Obtén la entrada de parámetros de consulta del cliente
challengeType: 2
forumTopicId: 301512
dashedName: get-query-parameter-input-from-the-client
---

# --description--

Otra forma común de obtener la entrada del cliente es codificando los datos después de la ruta, usando una cadena de consulta. La cadena de consulta está delimitada por un signo de interrogación (?), e incluye parejas de campo=valor.  Each couple is separated by an ampersand (&). Express puede analizar los datos de la cadena de consulta, y llenar el objeto `req.query`. Algunos caracteres, como el porcentaje (%), no pueden estar en URLs y tienen que ser codificados en un formato diferente antes de poder enviarlos. Si usas la API desde JavaScript, puedes usar métodos específicos para codificar/decodificar estos caracteres.

<blockquote>route_path: '/library'<br>actual_request_URL: '/library?userId=546&#x26;bookId=6754' <br>req.query: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

Construye un endpoint para el API, montado en `GET /name`. Responde con un documento JSON, tomando la estructura `{ name: 'firstname lastname'}`.  The first and last name parameters should be encoded in a query string e.g. `?first=firstname&last=lastname`.

**Nota:** En el siguiente ejercicio va a recibir datos de una solicitud POST, a la misma ruta `/name`. Si lo deseas, puedes utilizar el método `app.route(path).get(handler).post(handler)`. Esta sintaxis permite encadenar diferentes manejadores de verbos en la misma ruta. Puedes ahorrar un poco de escritura y tener un código más limpio.

# --hints--

Prueba 1: Tu endpoint de API debería responder con `{ "name": "Mick Jagger" }` cuando se llame al endpoint `/name` con `?first=Mick&last=Jagger`

```js
  const response = await fetch(code + '/name?first=Mick&last=Jagger');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "GET /name" route does not behave as expected'
  );
```

Prueba 2: Tu endpoint de API debe responder con `{ "name": "Keith Richards" }` cuando se llame al endpoint `/name` con `?first=Keith&last=Richards`

```js
  const response = await fetch(code + '/name?last=Richards&first=Keith');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "GET /name" route does not behave as expected'
  );
```

