---
id: 669fdc11c9b0263fe0814a7a
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Всередині кожного з елементів `li` потрібно мати елемент посилання.

Для першого елемента посилання текст має бути `About`, а значення атрибута `href` має бути `"#about"`. Символ решітки перед `about` означає ім’я `id`, яке буде додано пізніше в проєкті.

Для другого елемента посилання текст має бути `Posts`, а значення атрибута `href` має бути `"#posts"`.

Для третього елемента посилання текст має бути `Contact`, а значення атрибута `href` має бути `"#contact"`.

# --hints--

Ваш перший елемент посилання має містити текст `About` всередині першого елемента `li`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

Ваш перший елемент посилання має мати атрибут `href`, встановлений у `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Ваш другий елемент посилання має містити текст `Posts` всередині другого елемента `li`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

Ваш другий елемент посилання має мати атрибут `href`, встановлений у `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Ваш третій елемент посилання має містити текст `Contact` всередині третього елемента `li`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

Ваш третій елемент посилання має мати атрибут `href`, встановлений у `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
