---
id: 669fd6fd12918e3de87854d4
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Для вашого блогу має бути можливість користувачам переходити до різних розділів на сторінці.

Елемент `nav` використовується для надання навігаційних посилань на інші розділи в документі або на інші розділи на сайті. Часто ви побачите, що елемент `nav` використовується для меню або змісту.

Ось приклад використання елемента `nav`:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

Під вашим елементом `figure` додайте елемент `nav` з вкладеним всередину елементом `ul`.

Всередині елемента `ul` додайте три елементи `li`.

# --hints--

Ваш елемент `nav` повинен мати початковий тег.

```js
assert.match(code, /<nav\>/);
```

Ваш елемент `nav` повинен мати кінцевий тег.

```js
assert.match(code, /<\/nav\>/);
```

Ваш елемент `nav` повинен бути всередині вашого елемента `header`.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

Ваш елемент `ul` повинен мати початковий тег.

```js
assert.match(code, /<ul\>/);
```

Ваш елемент `ul` повинен мати кінцевий тег.

```js
assert.match(code, /<\/ul\>/);
```

Ваш елемент `ul` повинен бути всередині вашого елемента `nav`.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

Ваш елемент `ul` повинен містити три елементи `li`.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
