---
id: 669fcb5106eeb13ab7afcd09
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з семантичним HTML, створюючи сторінку блогу, присвячену коту на ім’я містер Віскерс.

Перший розділ, який ви створите, — це <dfn>заголовок</dfn> сторінки.

Елемент `header` використовується для представлення вступного вмісту, такого як навігація сторінкою та інша вступна інформація.

Ось приклад із використанням елемента `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Всередині елемента `body` додайте елемент `header`.

# --hints--

У вас має бути відкриваючий тег `header`.

```js
assert.match(code, /<header>/i);
```

У вас має бути закриваючий тег `header`.

```js
assert.match(code, /<\/header>/i);
```

Ваш елемент `header` має бути всередині елемента `body`.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
