---
id: 669fd003cb89ee3c2402e041
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У цьому вступному вмісті ви захочете показати зображення містера Віскерса з підписом.

Під елементом `h1` почніть із додавання елемента `figure`.

Всередині елемента `figure` додайте елемент `img`.

Атрибут `src` елемента `img` повинен мати значення `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`, а текст `alt` повинен мати значення `"a cat in the garden"`.

Під вашим елементом `img` додайте `figcaption` з текстом `Mr. Whiskers in the Garden`.

# --hints--

Ваш елемент `figure` повинен мати початковий тег.

```js
assert.match(code, /<figure\>/);
```

Ваш елемент `figure` повинен мати кінцевий тег.

```js
assert.match(code, /<\/figure\>/);
```

Ваш елемент `figure` повинен бути всередині вашого елемента `header`.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

У вас повинен бути елемент `img`.

```js
assert.exists(document.querySelector("img"));
```

Ваш елемент `img` повинен бути всередині вашого елемента `figure`.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

Ваш елемент `img` повинен мати атрибут `src`, встановлений у `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Ваш елемент `img` повинен мати атрибут `alt`, встановлений у `"a cat in the garden"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

Ваш елемент `figcaption` повинен мати початковий тег.

```js
assert.match(code, /<figcaption\>/);
```

Ваш елемент `figcaption` повинен мати кінцевий тег.

```js
assert.match(code, /<\/figcaption\>/);
```

Ваш елемент `figcaption` повинен бути всередині вашого елемента `figure`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Ваш елемент `figcaption` повинен містити текст `Mr. Whiskers in the Garden`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

Ваш `figcaption` повинен бути під елементом `img`. Ви розташували їх у неправильному порядку.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
