---
id: 669fcb5106eeb13ab7afcd09
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this workshop, you will practice working with semantic HTML by building a blog page dedicated to Mr. Whiskers the cat.

Перший розділ, який потрібно створити, — це <dfn>заголовок</dfn> сторінки.

Елемент `header` використовують, щоб представити вступний вміст (наприклад, навігацію сторінки чи іншу інформацію для ознайомлення).

Ось приклад використання елемента `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Додайте елемент `header` до елемента `body`.

# --hints--

Ви повинні мати початковий тег `header`.

```js
assert.match(code, /<header>/i);
```

Ви повинні мати кінцевий тег `header`.

```js
assert.match(code, /<\/header>/i);
```

Елемент `header` має бути в межах елемента `body`.

```js
assert.match(code, /<\/header>[.\n\s]*<\/body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```
