---
id: 669fdd3965f36f40af9615e5
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Ви побудували заголовок сторінки, і тепер можете перейти до додавання основного вмісту.

Додайте елемент `main` під елементом `header`.

# --hints--

Елемент `main` повинен мати початковий тег.

```js
assert.match(code, /<main\>/);
```

Елемент `main` повинен мати кінцевий тег.

```js
assert.match(code, /<\/main\>/);
```

Елемент `main` має бути під елементом `header`.

```js
assert.match(code, /<\/header>[.\n\s]*<main>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
