---
id: 669fde2081f65141ad703fe4
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Перший розділ на сторінці буде розділом про автора. Цей розділ представить містера Віскера і дасть користувачам уявлення про те, про що цей блог.

Всередині вашого `main` елемента додайте `section` елемент з `id` атрибутом, встановленим у `"about"`.

Всередині `section` елемента додайте `h2` з текстом `About`.

# --hints--

Ваш `section` елемент повинен мати початковий тег.

```js
assert.match(code, /<section.*>/);
```

Ваш `section` елемент повинен мати кінцевий тег.

```js
assert.match(code, /<\/section\>/);
```

Ваш `section` елемент повинен мати `id`, встановлений у `"about"`.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

Ваш `section` елемент повинен бути всередині вашого `main` елемента.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Ваш `h2` елемент повинен мати початковий тег.

```js
assert.match(code, /<h2\>/);
```

Ваш `h2` елемент повинен мати кінцевий тег.

```js
assert.match(code, /<\/h2\>/);
```

Ваш `h2` елемент повинен містити текст `About`. Перевірте правопис.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

Ваш `h2` елемент повинен бути всередині вашого `section` елемента.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
