---
id: 669fdfc9a5e3da42d2376609
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Тепер, коли ви додали розділ about, спробуйте натиснути на `About` посилання, щоб побачити, як сторінка прокрутиться до цього розділу.

Наступним розділом на сторінці блогу буде список публікацій про містера Віскера.

Додайте ще один `section` елемент із `id` множиною, встановленою на `"posts"`.

Всередині `section` елемента додайте `h2` елемент із текстом `Posts`.

# --hints--

Ваш код має містити відкриваючий `section` тег.

```js
assert.match(code, /<section.*>/i);
```

Ваш код має містити закриваючий `section` тег.

```js
assert.match(code, /<\/section>/i);
```

Ваш `section` елемент має мати `id` множину, встановлену на `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

Ваш `section` елемент має бути всередині `main` елемента.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Ваш `h2` елемент має бути всередині `section` елемента.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

Ваш `h2` елемент має містити текст `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
