---
id: 669fdc11c9b0263fe0814a7a
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Всередині кожного з елементів `li` має бути елемент посилання.

Текстом першого посилання має бути `About`, а значенням атрибута `href` має бути `"#about"`. Решітка перед `about` позначає `id`, що буде додано пізніше.

Текстом другого посилання має бути `Posts`, а значенням атрибута `href` має бути `"#posts"`.

Текстом третього посилання має бути `Contact`, а значенням атрибута `href` має бути `"#contact"`.

# --hints--

Перший елемент посилання повинен мати текст `About` в межах першого елемента `li`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

Перший елемент посилання повинен мати атрибут `href` зі значенням `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Другий елемент посилання повинен мати текст `Posts` в межах другого елемента `li`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

Другий елемент посилання повинен мати атрибут `href` зі значенням `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Третій елемент посилання повинен мати текст `Contact` в межах третього елемента `li`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

Третій елемент посилання повинен мати атрибут `href` зі значенням `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
