---
id: 66a334a1a7cca6354999f9bf
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Для першого допису в блозі ви використаєте елемент <dfn>article</dfn>.

Елемент `article` представляє автономний вміст на вебсторінці.

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

Під елементом `h2` додайте елемент `article`.

Всередині елемента `article` додайте елемент `h3` з текстом `Mr. Whiskers' First Day Home`.

Причина використання `h3` тут полягає в тому, що підтримка правильної структурної ієрархії для заголовків є важливою. Оскільки підзаголовок допису є елементом `h2`, наступним рівнем ієрархії буде `h3`.

# --hints--

Ваш код має містити відкриваючий тег `article`.

```js
assert.match(code, /<article>/i);
```

Ваш код має містити закриваючий тег `article`.

```js
assert.match(code, /<\/article>/i);
```

Ваш елемент `article` має бути всередині елемента `section` з пропом `id`, встановленим у `posts`.

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

Ваш код має містити відкриваючий тег `h3`.

```js
assert.match(code, /<h3>/i);
```

Ваш код має містити закриваючий тег `h3`.

```js
assert.match(code, /<\/h3>/i);
```

Ваш елемент `h3` має містити текст `Mr. Whiskers' First Day Home`.

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

Ваш елемент `h3` має бути всередині елемента `article`.

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
