---
id: 66a338be7f1dee383a0e0ecb
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Для другого допису в блозі вам потрібно додати ще один елемент `article`.

Всередині елемента `article` додайте елемент `h3` з текстом `Mr. Whiskers' First Bath`.

Під вашим елементом `h3` додайте два абзаци тексту lorem ipsum.

# --hints--

У вас має бути другий елемент `article` всередині елемента `section` з множиною `id`, встановленою на `posts`.

```js
const articleElement = document.querySelector('#posts article:nth-of-type(2)');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

У вас має бути елемент `h3` всередині елемента `article` з текстом `Mr. Whiskers' First Bath`.

```js
const h3Element = document.querySelector('#posts article:nth-of-type(2) h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Bath");
```

Ваш елемент `h3` має бути всередині вашого елемента `article`.

```js
const h3Element = document.querySelector('#posts article:nth-of-type(2) h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

У вас має бути два абзаци всередині елемента `article`.

```js
const paragraphElements = document.querySelectorAll('#posts article:nth-of-type(2) p');
assert.strictEqual(paragraphElements.length, 2);
```

Ваш перший абзац має містити текст lorem ipsum. Перевірте орфографію та пунктуацію.

```js
const paragraphElement = document.querySelector('#posts article:nth-of-type(2) p:first-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim(), loremIpsum);
```

Ваш другий абзац має містити текст lorem ipsum. Перевірте орфографію та пунктуацію.

```js
const paragraphElement = document.querySelector('#posts article:nth-of-type(2) p:last-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim(), loremIpsum);
```

Ваші абзаци мають йти після елемента `h3`.

```js
let h3 = document.querySelectorAll('#posts article:nth-of-type(2) > p ~ h3');
assert.isEmpty(h3);
const paragraphs = document.querySelectorAll('#posts article:nth-of-type(2) > h3 ~ p');
assert.lengthOf(paragraphs, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>

        <article>
          <h3>Mr. Whiskers' First Day Home</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
