---
id: 66a334a1a7cca6354999f9bf
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Для першого допису використайте елемент <dfn>article</dfn>.

Елемент `article` представляє самостійний вміст на вебсторінці.

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

Додайте елемент `article` під елементом `h2`.

Додайте елемент `h3` з текстом `Mr. Whiskers' First Day Home` до елемента `article`.

Тут використано `h3`, оскільки важливо підтримувати належну структурну ієрархію заголовків. Оскільки підзаголовком дописів є елемент `h2`, наступним рівнем в ієрархії буде `h3`.

# --hints--

Код повинен містити початковий тег `article`.

```js
assert.match(code, /<article>/i);
```

Код повинен містити кінцевий тег `article`.

```js
assert.match(code, /<\/article>/i);
```

Елемент `article` має бути в межах елемента `section` з `id` зі значенням `posts`.

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

Код повинен містити початковий тег `h3`.

```js
assert.match(code, /<h3>/i);
```

Код повинен містити кінцевий тег `h3`.

```js
assert.match(code, /<\/h3>/i);
```

Елемент `h3` повинен містити текст `Mr. Whiskers' First Day Home`.

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

Елемент `h3` має бути в межах елемента `article`.

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
