---
id: 66a49f685961e997e337cab1
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

`header` буде відповідальним за відображення головного заголовка, зображення та навігації сторінок для блогу.

Всередині елемента `header` додайте `h1` з текстом `Welcome to Mr. Whiskers' Blog Page!`.

# --hints--

У вас має бути відкриваючий тег `h1`.

```js
assert.match(code, /<h1>/i);
```

У вас має бути закриваючий тег `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Ваш елемент `h1` має містити текст `Welcome to Mr. Whiskers' Blog Page!`. Перевірте на наявність орфографічних помилок.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

Ваш елемент `h1` має бути вкладений всередину елемента `header`.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
