---
id: 66a49f685961e997e337cab1
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Елемент `header` відповідатиме за відтворення головного заголовка, зображення і навігації сторінки.

Додайте елемент `h1` з текстом `Welcome to Mr. Whiskers' Blog Page!` до елемента `header`.

# --hints--

Ви повинні мати початковий тег `h1`.

```js
assert.match(code, /<h1>/i);
```

Ви повинні мати кінцевий тег `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Елемент `h1` повинен містити текст `Welcome to Mr. Whiskers' Blog Page!`. Двічі перевірте наявність орфографічних помилок.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

Елемент `h1` має бути вкладеним в елементі `header`.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--

--fcc-editable-region--
    </header>
  </body>
</html>
```
