---
id: 66c057041df6394ca796bf33
title: بناء صانع جمل
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

في هذا المختبر، ستنشئ قصتين مختلفتين باستخدام قالب جملة. ستستخدم المتغيرات لتخزين أجزاء مختلفة من القصة ثم تعرض القصص في وحدة التحكم.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب إعلان المتغيرات التالية باستخدام `let`:
   
    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. يجب تعيين قيم سلسلة نصية للمتغيرات السابقة حسب اختيارك.

1. يجب إعلان متغير `firstStory`.
   
1. يجب عليك استخدام قالب القصة التالي لإنشاء القصة الأولى وتعيينها إلى المتغير `firstStory`: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`;

1. يجب عرض قصتك الأولى في وحدة التحكم باستخدام الرسالة `"First story: [firstStory]"`.

1. يجب تعيين قيم جديدة للمتغيرات `adjective`، `noun`، `verb`، `place`، `adjective2`، و `noun2`.

1. يجب إعلان متغير `secondStory`.

1. أنشئ قصة أخرى باستخدام نفس القالب وعيّنها إلى المتغير `secondStory`.

1. يجب عرض قصتك الثانية في وحدة التحكم باستخدام الرسالة `"Second story: [secondStory]"`.


# --hints--

يجب إعلان متغير `adjective`.

```js
assert.isNotNull(adjective);
```

يجب إعلان متغير `noun`.

```js
assert.isNotNull(noun);
```

يجب إعلان متغير `verb`.

```js
assert.isNotNull(verb);
```

يجب إعلان متغير `place`.

```js
assert.isNotNull(place);
```

يجب إعلان متغير `adjective2`.

```js
assert.isNotNull(adjective2);
```

يجب إعلان متغير `noun2`.

```js
assert.isNotNull(noun2);
```

يجب تعيين قيمة سلسلة نصية للمتغير `adjective`.

```js
assert.isString(adjective);
```

يجب تعيين قيمة سلسلة نصية للمتغير `noun`.

```js
assert.isString(noun);
```

يجب تعيين قيمة سلسلة نصية للمتغير `verb`.

```js
assert.isString(verb);
```

يجب تعيين قيمة سلسلة نصية للمتغير `place`.

```js
assert.isString(place);
```

يجب تعيين قيمة سلسلة نصية للمتغير `adjective2`.

```js
assert.isString(adjective2);
```

يجب تعيين قيمة سلسلة نصية للمتغير `noun2`.

```js
assert.isString(noun2);
```

يجب إعلان متغير `firstStory`.

```js
assert.isNotNull(firstStory);
```

يجب أن تستخدم صيغة القصة الصحيحة للقصة الأولى: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. انتبه إلى الفراغات.

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

يجب تجميع قصتك الأولى باستخدام المتغيرات التي أعلنتها بالترتيب الصحيح.

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

يجب تسجيل قصتك الأولى باستخدام الرسالة `"First story: [firstStory]"`.

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

يجب إعلان متغير `secondStory`.

```js
assert.isNotNull(secondStory);
```

يجب إعادة تعيين المتغير `adjective` للقصة الثانية.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

يجب إعادة تعيين المتغير `noun` للقصة الثانية.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

يجب إعادة تعيين المتغير `verb` للقصة الثانية.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

يجب إعادة تعيين المتغير `place` للقصة الثانية.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

يجب إعادة تعيين المتغير `adjective2` للقصة الثانية.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

يجب إعادة تعيين المتغير `noun2` للقصة الثانية.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

يجب أن تستخدم صيغة القصة الصحيحة للقصة الثانية: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. انتبه إلى الفراغات.

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

يجب تجميع قصتك الثانية باستخدام المتغيرات التي أعلنتها بالترتيب الصحيح.

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

يجب تسجيل قصتك الثانية باستخدام الصيغة `"Second story: [secondStory]"`.

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

يجب ألا تكون القصة `firstStory` هي نفسها القصة `secondStory`.

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```

