---
id: 687ad7be1ff45032ccf1d92f
title: Costruisci un selettore di tema
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

In questo laboratorio costruirai un'app che cambia tra temi diversi. Quando cambi tema, il colore di sfondo cambierà e verrà mostrato un messaggio diverso.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `button` con il testo `Switch Theme`.
2. Il tuo elemento `button` dovrebbe avere i seguenti attributi:
   - Un attributo `id` impostato su `"theme-switcher-button"`.
   - Un attributo `aria-haspopup` impostato su `"true"`.
   - Un attributo `aria-expanded` impostato su `"false"`.
   - Un attributo `aria-controls` impostato su `"theme-dropdown"`.
3. Dovresti avere un elemento `ul` con i seguenti attributi:
   - Un attributo `id` impostato su `"theme-dropdown"`.
   - Un attributo `role` impostato su `"menu"`.
   - Un attributo `aria-labelledby` impostato su `"theme-switcher-button"`.
   - Un attributo `hidden`.
4. Il tuo elemento `ul` dovrebbe avere almeno due elementi `li` con un attributo `role` impostato su `"menuitem"` e un testo a tua scelta che rappresenta un tema diverso.
5. Ciascuno dei tuoi elementi `li` dovrebbe avere un attributo `id` che inizia con `theme-` e termina con il tema che hai impostato per il testo dell'elemento `li`. Per esempio, se uno dei tuoi temi è `Light` allora il tuo `id` dovrebbe essere `theme-light`.
6. Dovresti avere un elemento con un attributo `aria-live` impostato su `"polite"`.
7. Dovresti avere un array `themes` con almeno due oggetti che contengono ciascuno una proprietà `name` e `message`. La `name` rappresenterà un tema diverso e la `message` verrà mostrata quando il tema cambia. Sei libero di inventare valori `name` e `message` a tua scelta ma i valori `name` dovrebbero corrispondere al contenuto testuale dei tuoi elementi `li`.
8. Quando un utente clicca sul `#theme-switcher-button`, l'elemento `#theme-dropdown` dovrebbe mostrare i temi disponibili.
9. Quando l'elemento `#theme-dropdown` è mostrato, l'attributo `aria-expanded` dovrebbe essere impostato su `"true"` per l'elemento `button`.
10. Quando l'utente clicca sul `#theme-switcher-button` mentre l'elemento `#theme-dropdown` è mostrato, il menu dovrebbe chiudersi e l'attributo `hidden` dovrebbe essere aggiunto e l'attributo `aria-expanded` dovrebbe essere impostato su `"false"`.
11. Quando un utente clicca sul `#theme-switcher-button` e seleziona un tema, la classe `theme-<name>` corrispondente dovrebbe essere aggiunta all'elemento `body` e il messaggio del tema correlato che hai impostato nell'array `themes` dovrebbe essere mostrato nell'elemento `aria-live="polite"`.

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

Dovresti avere un elemento `button`.

```js
assert.exists(document.querySelector("button"));
```

Il tuo elemento `button` dovrebbe avere il testo `Switch Theme`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

Il tuo elemento `button` dovrebbe avere un attributo `id` impostato su `"theme-switcher-button"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

Il tuo elemento `button` dovrebbe avere un attributo `aria-haspopup` impostato su `"true"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

Il tuo elemento `button` dovrebbe avere un attributo `aria-expanded` impostato su `"false"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

Il tuo elemento `button` dovrebbe avere un attributo `aria-controls` impostato su `"theme-dropdown"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

Dovresti avere un elemento `ul`.

```js
assert.exists(document.querySelector("ul"));
```

Il tuo elemento `ul` dovrebbe avere un attributo `id` impostato su `"theme-dropdown"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

Il tuo elemento `ul` dovrebbe avere un attributo `role` impostato su `"menu"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

Il tuo elemento `ul` dovrebbe avere un attributo `aria-labelledby` impostato su `"theme-switcher-button"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

Il tuo elemento `ul` dovrebbe avere un attributo `hidden`.

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

Il tuo elemento `ul` dovrebbe avere almeno due elementi `li`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

Ciascuno dei tuoi elementi `li` dovrebbe avere un attributo `role` impostato su `"menuitem"`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

Ciascuno dei tuoi elementi `li` dovrebbe avere un testo che rappresenta un tema.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

Ciascuno dei tuoi elementi `li` dovrebbe avere un attributo `id` che inizia con `theme-` e termina con il tema corrispondente per quell'elemento. Assicurati che il valore `id` sia tutto minuscolo.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

Dovresti avere un elemento con un attributo `aria-live` impostato su `"polite"`.

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

Dovresti avere un array `themes`.

```js
assert.isArray(themes);
```

Il tuo array `themes` dovrebbe avere almeno due oggetti.

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

Il tuo array `themes` dovrebbe avere almeno due oggetti ciascuno con una proprietà `name` e `message`.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

I valori delle proprietà `name` e `message` dovrebbero essere stringhe non vuote.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

Ciascuno dei valori della proprietà `name` dovrebbe corrispondere al contenuto testuale di uno dei tuoi elementi `li`. Assicurati che questi valori siano tutti in minuscolo.

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

Quando un utente clicca sul `#theme-switcher-button` per mostrare le opzioni di tema, l'attributo `hidden` dovrebbe essere rimosso dall'elemento `#theme-dropdown`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

Quando l'elemento `#theme-dropdown` è mostrato, l'attributo `aria-expanded` dovrebbe essere impostato su `"true"` per l'elemento `button`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

Quando l'utente clicca sul `#theme-switcher-button` mentre l'elemento `#theme-dropdown` è mostrato, il menu a discesa dovrebbe tornare a `hidden`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

Quando l'utente clicca sul `#theme-switcher-button` mentre l'elemento `#theme-dropdown` è mostrato, l'attributo `aria-expanded` dovrebbe essere impostato su `"false"`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

Quando un utente clicca sul `#theme-switcher-button` e seleziona un tema, la classe `theme-<name>` corrispondente dovrebbe essere aggiunta all'elemento `body`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

Quando un utente clicca sul `#theme-switcher-button` e seleziona un tema, un `message` relativo al `theme` selezionato dall'array `themes` dovrebbe essere mostrato nell'elemento `aria-live="polite"`.


```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
