---
id: 68cabc534d2e741f33a3fe1d
title: Створіть трекер характеристик персонажа гри
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

У цій лабораторній роботі ви створите трекер характеристик персонажа гри. Програма дозволить створити персонажа з певними атрибутами, оновлювати ці атрибути та отримувати поточні характеристики персонажа.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Створіть клас із назвою `GameCharacter`, який представляє персонажа гри та керує його характеристиками.
2. Під час створення нового об’єкта `GameCharacter` він повинен мати такі властивості:
    - множина `_name` зі значенням рядка, переданого під час створення.
    - множина `_health` зі значенням `100`.
    - множина `_mana` зі значенням `50`.
    - множина `_level` зі значенням `1`.
3. Створіть властивість `name` для доступу лише для читання до імені персонажа.
4. Для властивості `health`:
     - Визначте гетер, який повертає поточне значення здоров’я.
     - Визначте сетер, який не дозволяє встановлювати здоров’я нижче `0` і обмежує його значення `100`.
5. Для властивості `mana`:
     - Визначте гетер, який повертає поточне значення мани.
     - Визначте сетер, який не дозволяє встановлювати ману нижче `0` і обмежує її значення `50`.
6. Створіть гетер для `level`, який повертає поточний рівень персонажа.
7. Визначте метод із назвою `level_up`, який:
    - Збільшує рівень персонажа на 1.
    - Скидає здоров’я до `100` і ману до `50` за допомогою відповідних сеттерів.
    - Виводить повідомлення у форматі `<name> leveled up to <level>!` (де `<name>` і `<level>` потрібно замінити на ім’я персонажа та новий рівень відповідно).
8. Визначте метод `__str__`, який повертає відформатований рядок, що містить:
    - Ім’я персонажа.
    - Рівень персонажа.
    - Поточне здоров’я персонажа.
    - Поточну ману персонажа.
   Наприклад, персонаж на ім’я `Kratos` одразу після створення повинен бути представлений так:

   ```md
   Name: Kratos
   Level: 1
   Health: 100
   Mana: 50
   ```

## Приклад використання

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

Ви повинні мати клас `GameCharacter`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

У класі `__init__` має бути метод `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

Метод `__init__` повинен мати два параметри, перший з яких — `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

Ви повинні призначити множину `self._name` рядку, переданому під час створення, всередині методу `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

Ви повинні призначити `100` множині `self._health` всередині методу `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

Ви повинні призначити `50` множині `self._mana` всередині методу `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

Ви повинні призначити `1` множині `self._level` всередині методу `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

У класі `name` має бути метод `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

Метод `name` повинен мати лише один параметр — `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

У методі `self._name` потрібно повернути `name`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

Метод `name` має мати декоратор `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

У класі `health` має бути метод `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

Метод `health` повинен мати лише один параметр — `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

У методі `self._health` потрібно повернути `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

Метод `health` має мати декоратор `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

Ви повинні створити сетер `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

Сетер `health` має мати декоратор `@health.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

Сетер `health` повинен мати два параметри, перший з яких — `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Ви повинні призначити множину `self._health` значенню `0`, якщо значення, передане сетеру `health`, менше за `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

Ви повинні призначити множину `self._health` значенню `100`, якщо значення, передане сетеру `health`, більше за `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

Ви повинні призначити множину `self._health` значенню, переданому сетеру `health`, якщо значення знаходиться між `0` і `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
    char.health = 0
    assert char.health == 0
    char.health = 100
    assert char.health == 100
`)})
```

У класі `mana` має бути метод `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

Метод `mana` повинен мати лише один параметр — `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

У методі `self._mana` потрібно повернути `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

Метод `mana` має мати декоратор `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

Ви повинні створити сетер `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

Сетер `mana` має мати декоратор `@mana.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

Сетер `mana` повинен мати два параметри, перший з яких — `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Ви повинні призначити множину `self._mana` значенню `0`, якщо значення, передане сетеру `mana`, менше за `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

Ви повинні призначити множину `self._mana` значенню `50`, якщо значення, передане сетеру `mana`, більше за `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

Ви повинні призначити множину `self._mana` значенню, переданому сетеру `mana`, якщо значення знаходиться між `0` і `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

У класі `level` має бути метод `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

Метод `level` повинен мати лише один параметр — `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

У методі `self._level` потрібно повернути `level`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

Метод `level` має мати декоратор `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

У класі `level_up` має бути метод `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

Метод `level_up` повинен мати лише один параметр — `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

Ви повинні збільшити `self._level` на `1` у методі `level_up`.

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

Ви повинні призначити множину `self.health` значенню `100` у методі `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

Ви повинні призначити множину `self.mana` значенню `50` у методі `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

Метод `level_up` має вивести `<name> leveled up to <level>!` (де `<name>` і `<level>` потрібно замінити на ім’я персонажа та новий рівень відповідно).

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

У класі `__str__` має бути метод `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

Метод `__str__` повинен мати лише один параметр — `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

Ваш метод `__str__` має повертати рядок із характеристиками персонажа у вказаному форматі.

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
