---
id: 69a92f0b9cc04bb0d5327bb5
title: implementar uma pilha
challengeType: 26
dashedName: implement-a-stack
---

# --description--

Neste laboratﾃｳrio, vocﾃｪ vai implementar uma estrutura de dados pilha usando funﾃｧﾃｵes. Uma pilha ﾃｩ uma estrutura de dados Last-In-First-Out (LIFO), onde os elementos sﾃ｣o adicionados e removidos do topo.

**Objetivo:** Atenda ﾃs user stories abaixo e faﾃｧa todos os testes passarem para concluir o laboratﾃｳrio.

**User stories:**

1. Vocﾃｪ deve ter uma funﾃｧﾃ｣o `initStack` que retorna um objeto com uma propriedade `collection` definida como um array vazio.
1. Vocﾃｪ deve ter uma funﾃｧﾃ｣o `push` que adiciona um elemento ao topo da pilha.
1. Vocﾃｪ deve ter uma funﾃｧﾃ｣o `pop` que remove e retorna o elemento do topo da pilha, ou `undefined` se nﾃ｣o houver nenhum.
1. Vocﾃｪ deve ter uma funﾃｧﾃ｣o `peek` que retorna o elemento do topo da pilha sem removﾃｪ-lo, ou `undefined` se nﾃ｣o houver nenhum.
1. Vocﾃｪ deve ter uma funﾃｧﾃ｣o `isEmpty` que retorna `true` se a pilha nﾃ｣o contiver elementos, e `false` caso contrﾃ｡rio.
1. Vocﾃｪ deve ter uma funﾃｧﾃ｣o `clear` que remove todos os elementos da pilha.

**Nota:** A maioria dos testes depende das funﾃｧﾃｵes `initStack` e `push`. Implemente-as primeiro, pois os testes para `pop`, `peek`, `isEmpty` e `clear` exigem adicionar elementos ﾃ pilha.

# --hints--

Vocﾃｪ deve ter uma funﾃｧﾃ｣o `push`.

```js
assert.isFunction(push);
```

A funﾃｧﾃ｣o `push` deve adicionar um elemento ao topo da pilha.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

Vocﾃｪ deve ter uma funﾃｧﾃ｣o `pop`.

```js
assert.isFunction(pop);
```

A funﾃｧﾃ｣o `pop` deve remover e retornar o elemento do topo da pilha.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

A funﾃｧﾃ｣o `pop` deve retornar `undefined` se a pilha estiver vazia.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

A funﾃｧﾃ｣o `pop` deve retornar valores falsy corretamente.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

Vocﾃｪ deve ter uma funﾃｧﾃ｣o `peek`.

```js
assert.isFunction(peek);
```

A funﾃｧﾃ｣o `peek` deve retornar o elemento do topo da pilha sem removﾃｪ-lo.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

A funﾃｧﾃ｣o `peek` deve retornar `undefined` se a pilha estiver vazia.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

A funﾃｧﾃ｣o `peek` deve retornar valores falsy corretamente.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

Vocﾃｪ deve ter uma funﾃｧﾃ｣o `isEmpty`.

```js
assert.isFunction(isEmpty);
```

A funﾃｧﾃ｣o `isEmpty` deve retornar `true` para uma pilha vazia.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

A funﾃｧﾃ｣o `isEmpty` deve retornar `false` para uma pilha nﾃ｣o vazia.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

A funﾃｧﾃ｣o `isEmpty` deve retornar `false` quando o elemento do topo for falsy.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

Vocﾃｪ deve ter uma funﾃｧﾃ｣o `clear`.

```js
assert.isFunction(clear);
```

A funﾃｧﾃ｣o `clear` deve remover todos os elementos da pilha.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
