---
id: 69a92f0b9cc04bb0d5327bb5
title: Implementar uma pilha
challengeType: 26
dashedName: implement-a-stack
---

# --description--

Neste laboratório, você vai implementar uma estrutura de dados pilha usando funções. Uma pilha é uma estrutura de dados Last-In-First-Out (LIFO), onde os elementos são adicionados e removidos do topo.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função `initStack` que retorna um objeto com uma propriedade `collection` definida como um array vazio. O último elemento de `collection` representa o topo da pilha.
1. Você deve ter uma função `push` que recebe um objeto pilha e um elemento como argumentos, e adiciona o elemento ao topo da pilha colocando-o no final da `collection` da pilha.
1. Você deve ter uma função `pop` que recebe um objeto pilha como argumento, e remove e retorna o elemento do topo da pilha, ou `undefined` se não houver nenhum.
1. Você deve ter uma função `peek` que recebe um objeto pilha como argumento, e retorna o elemento do topo da pilha sem removê-lo, ou `undefined` se não houver nenhum.
1. Você deve ter uma função `isEmpty` que recebe um objeto pilha como argumento, e retorna `true` se a pilha não contiver elementos, e `false` caso contrário.
1. Você deve ter uma função `clear` que recebe um objeto pilha como argumento, e remove todos os elementos da `collection` da pilha.

**Nota:** A maioria dos testes depende das funções `initStack` e `push`. Implemente-as primeiro, pois os testes para `pop`, `peek`, `isEmpty` e `clear` exigem adicionar elementos à pilha.

# --hints--

Você deve ter uma função `push`.

```js
assert.isFunction(push);
```

A função `push` deve adicionar um elemento ao topo da pilha.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

Você deve ter uma função `pop`.

```js
assert.isFunction(pop);
```

A função `pop` deve remover e retornar o elemento do topo da pilha.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.deepEqual(stack.collection, ['first', 'second']);
assert.strictEqual(pop(stack), 'second');
assert.deepEqual(stack.collection, ['first']);
```

A função `pop` deve retornar `undefined` se a pilha estiver vazia.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

A função `pop` deve retornar valores falsy corretamente.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

Você deve ter uma função `peek`.

```js
assert.isFunction(peek);
```

A função `peek` deve retornar o elemento do topo da pilha sem removê-lo.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

A função `peek` deve retornar `undefined` se a pilha estiver vazia.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

A função `peek` deve retornar valores falsy corretamente.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

Você deve ter uma função `isEmpty`.

```js
assert.isFunction(isEmpty);
```

A função `isEmpty` deve retornar `true` para uma pilha vazia.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

A função `isEmpty` deve retornar `false` para uma pilha não vazia.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

A função `isEmpty` deve retornar `false` quando o elemento do topo for falsy.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

Você deve ter uma `clear` função.

```js
assert.isFunction(clear);
```

A função `clear` deve remover todos os elementos da pilha.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
