---
id: 669e81368e52b3a5c35a2dc5
title: 비디오 모음 페이지 만들기
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `main` 요소의 유일한 자식으로 `body` 요소가 있어야 합니다.
1. 페이지 주제를 나타내는 `h1` 요소가 있어야 합니다.
1. `h1` 요소 아래에 페이지 주제를 소개하는 단락이 있어야 합니다.
1. 첫 번째 단락 아래에 세 개의 `section` 요소가 있어야 합니다.
1. 각 섹션에는 순서대로 `h2` 요소, 단락, `iframe` 요소가 포함되어야 합니다.
1. 세 개의 `iframe` 요소에는 유효한 비디오가 설정된 `src` 속성이 있어야 합니다.
1. 각 `iframe` 요소에는 임베드된 콘텐츠를 설명하는 `title` 속성과, 적절한 크기로 설정하는 `height` 속성과 `width` 속성이 있어야 합니다.

# --hints--

`main` 요소 안에 `body` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

`main` 요소의 유일한 자식으로 `body` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector('body').children.length, 1);
```

`h1` 요소 안에 페이지 주제를 나타내는 `main` 요소가 있어야 합니다.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

`h1` 요소 아래에 페이지 주제를 소개하는 단락이 있어야 합니다.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

첫 번째 `section` 요소 아래에 세 개의 `p` 요소가 있어야 합니다.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

각 `section` 요소는 해당 섹션의 제목 역할을 하는 `h2` 요소로 시작해야 합니다.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

각 `section` 요소는 두 번째 자식으로 비디오 콘텐츠를 소개하는 `p` 요소를 포함해야 합니다.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

각 `section` 요소는 세 번째 자식으로 `iframe` 요소를 포함해야 합니다.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

첫 번째 `iframe` 요소에는 유효한 비디오가 설정된 `src` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

첫 번째 `iframe` 요소에는 임베드된 콘텐츠를 설명하는 `title` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

첫 번째 `iframe` 요소에는 `height` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

첫 번째 `iframe` 요소에는 `width` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

두 번째 `iframe` 요소에는 유효한 비디오가 설정된 `src` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

두 번째 `iframe` 요소에는 임베드된 콘텐츠를 설명하는 `title` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

두 번째 `iframe` 요소에는 `height` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

두 번째 `iframe` 요소에는 `width` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

세 번째 `iframe` 요소에는 유효한 비디오가 설정된 `src` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

세 번째 `iframe` 요소에는 임베드된 콘텐츠를 설명하는 `title` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

세 번째 `iframe` 요소에는 `height` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

세 번째 `iframe` 요소에는 `width` 속성이 있어야 합니다.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front-End Web Development</h1>
        <p>Front-End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front-end
            developers ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
