---
id: 68de7360e1832ea180577a8a
title: Construa um Calculador do N-ésimo Número de Fibonacci
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `fibonacci`.
2. Você deve definir uma lista chamada `sequence` dentro da função `fibonacci` e ela deve ser inicializada com os valores `[0, 1]`.
3. A função `fibonacci` deve aceitar um parâmetro, um inteiro positivo `n`.  
4. Chamar `fibonacci(n)` deve usar uma abordagem de programação dinâmica para calcular e retornar o `n`-ésimo número da sequência de Fibonacci, onde cada número é a soma dos dois números anteriores.
5. Cada número calculado na posição `n` na sequência de Fibonacci deve ser armazenado na lista `sequence` no índice `n - 1`.

# --hints--

Você deve ter uma função chamada `fibonacci`.

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

Sua função `fibonacci` deve receber um único parâmetro.

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

Você deve ter uma lista chamada `sequence` dentro da função `fibonacci` inicializada como `[0, 1]`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)` deve retornar `0`.

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)` deve retornar `1`.

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)` deve retornar `1`.

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)` deve retornar `2`.

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)` deve retornar `5`.

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)` deve retornar `55`.

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)` deve retornar `610`.

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

Você não deve usar recursão no seu código.

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```

