---
id: 66e45c8140f9fda5c105ae26
title: 建置一個 Event Flyer 頁面
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該在 body 中有一個 `header` 元素。
1. 你的 `header` 元素中應該包含一張用於你的事件的影像，以及一個帶有你的事件標題的 `h1`，順序如上。如果你願意，可以使用這張影像：`https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. 你應該在 body 中有一個 `main` 元素。
1. 你的 `main` 元素中應至少包含兩個 `section` 元素，以展示事件特性。
1. 你的每個 `section` 元素中都應該包含一個 `h2`。
1. 你的主體應該有上下各 `50px` 的內邊距。
1. 你的 body 應該有 `0` 的上下邊距，以及能自動置中的左右邊距。
1. 你的 body 應該設定相對於 viewport 寬度的寬度。
1. 你的 body 應該使用 `calc` 函式（程式）將其 `min-height` 屬性設定為視窗高度的 100％ 減去套用於 body 頂部和底部的所有內距。
1. 你的傳單中應該至少有一個 `hr`。
1. 你應該將所有 `hr` 和 `section` 元素的 `width` 設定為相對於其父元素的百分比值。

**注意：** 請務必在你的 HTML 中聯結你的樣式表並套用你的 CSS。

# --hints--

你應該在 `body` 元素中有一個 `header` 元素。

```js
assert.exists(document.querySelector('body > header'));
```

你應該在 `header` 中將 `img` 元素作為第一個子元素。

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

你應該在 `header` 中將 `h1` 元素設為第二個子元素。

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

你應該只擁有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `h1` 應該包含你的事件標題。

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

你應該在 `body` 元素中有一個 `main` 元素。

```js
assert.exists(document.querySelector('body > main'));
```

你應該在你的 `main` 元素中至少有兩個 `section` 元素。

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

你的 `section` 元素中每個都應該包含一個 `h2`。

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

你的 `h2` 元素不應該是空的。

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

你的 `body` 元素應該有上下各 `50px` 的內邊距。

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

你的 `body` 元素應該有上下邊距為 `0`。

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

你的 `body` 元素應該有左右 `auto` 邊距。

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

你的 `body` 元素應該使用 `vw` 設定寬度。

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

你的 `body` 應該使用 `calc` 將其 `min-height` 設定為 `100vh - 100px`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

你應該至少有一個 `hr` 元素。

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

你的 `hr` 元素的寬度應該使用百分比值來設定。

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

你的 `section` 元素的寬度應該使用百分比值設定。

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
