---
id: a3566b1109230028080c9345
title: تصميم خوارزمية لجمع كل الأرقام
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

في هذا المختبر، ستحتاج إلى تصميم خوارزمية لجمع كل الأرقام.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك دالة باسم `sumAll` تقبل مصفوفة تحتوي على رقمين.
1. يجب أن تُرجع `sumAll([n, m])` مجموع `n` و `m` بالإضافة إلى مجموع كل الأرقام بينهما. الرقم الأقل لن يكون دائمًا في البداية. على سبيل المثال، يجب أن تُرجع `sumAll([4,1])` القيمة `10` لأن مجموع كل الأرقام بين `1` و `4` (شاملة) هو `10`.

# --hints--

يجب أن يكون لديك دالة باسم `sumAll`.

```js
assert.isFunction(sumAll);
```

يجب أن تُرجع `sumAll([1, 4])` رقمًا.

```js
assert.isNumber(sumAll([1, 4]));
```

يجب أن تُرجع `sumAll([1, 4])` القيمة `10`.

```js
assert.equal(sumAll([1, 4]), 10);
```

يجب أن تُرجع `sumAll([4, 1])` القيمة `10`.

```js
assert.equal(sumAll([4, 1]), 10);
```

يجب أن تُرجع `sumAll([5, 10])` القيمة `45`.

```js
assert.equal(sumAll([5, 10]), 45);
```

يجب أن تُرجع `sumAll([10, 5])` القيمة `45`.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
