---
id: 66c8ba975ee7230e29f6c4af
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

申請者の年収が `minIncomeForCondo` 以上、かつクレジットスコアが `minCreditScoreForCondo` 以上であれば、マンションと自動車ローンの資格があります。

`getLoanMessage` 関数でこれが正しいか確認してください。正しければ、文字列 `"You qualify for a condo and car loan."` を返してください。

# --hints--

`getLoanMessage` 関数は、申請者がマンションと自動車ローンの資格がある場合に文字列を返す必要があります。

```js
assert.isString(getLoanMessage(45000, 680));
```

`getLoanMessage` 関数は文字列 `"You qualify for a condo and car loan."` を返す必要があります。

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

`getLoanMessage` 関数は、申請者の年収とクレジットスコアがマンションローンの条件を満たさない場合に `undefined` を返す必要があります。

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
