---
id: 66c8ba975ee7230e29f6c4b0
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

今度は、申請者が自動車ローンの資格があるかどうかだけを確認してください。資格がある場合は、文字列 `"You qualify for a car loan."` を返してください。

# --hints--

あなたの `getLoanMessage` 関数は、申請者が自動車ローンの資格がある場合に文字列を返す必要があります。

```js
assert.isString(getLoanMessage(30000, 650));
```

あなたの `getLoanMessage` 関数は、文字列 `"You qualify for a car loan."` を返す必要があります。

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

あなたの `getLoanMessage` 関数は、申請者の年収と信用スコアが自動車ローンの要件を満たさない場合に `undefined` を返す必要があります。

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 
  
--fcc-editable-region--
}
```
