---
id: 66c8ba975ee7230e29f6c4ad
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

`annualIncome` と `creditScore` に基づいてユーザーがどのローンに適格かを確認するには、`if/else if` 関数の中で `getLoanMessage` 文または三項演算子を使う必要があります。その後、各条件のブロック内で適切なメッセージを返してください。

まずデュプレックスローンから、`annualIncome` が `minIncomeForDuplex` 以上であり、かつ `creditScore` が `minCreditScoreForDuplex` 以上かどうかを確認してください。

その条件が真であれば、申請者はデュプレックスローンおよび他のローンの適格者です。したがって、そのチェックの中で文字列 `"You qualify for a duplex, condo, and car loan."` を返してください。

# --hints--

`getLoanMessage` 関数は文字列を返す必要があります。

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

`getLoanMessage` 関数は文字列 `"You qualify for a duplex, condo, and car loan."` を返す必要があります。

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

`getLoanMessage` 関数は申請者の年収とクレジットスコアがデュプレックスローンの要件を満たさない場合、`undefined` を返す必要があります。

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
