---
id: 66c8ba41a77db20a93f9d7a1
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、ローン資格チェッカーアプリを作りながらJavaScriptの条件分岐をレビューします。

このアプリは、ユーザーの年収とクレジットスコアに基づいて、デュプレックス、車、コンドミニアムのローン資格があるかどうかをチェックします。クレジットスコアとは、借りたお金をどれだけうまく管理できているかを表す数値です。

まずは、以下の変数と値を作成してください。

| 変数名 | 値   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

`minIncomeForDuplex`という変数を`60000`に設定してください。

```js
assert.equal(minIncomeForDuplex, 60000);
```

`minCreditScoreForDuplex`を`700`に設定してください。

```js
assert.equal(minCreditScoreForDuplex, 700);
```

`minIncomeForCar`という変数を`30000`に設定してください。

```js
assert.equal(minIncomeForCar, 30000);
```

`minCreditScoreForCar`という変数を`650`に設定してください。

```js
assert.equal(minCreditScoreForCar, 650);
```

`minIncomeForCondo`という変数を`45000`に設定してください。

```js
assert.equal(minIncomeForCondo, 45000);
```

`minCreditScoreForCondo`という変数を`680`に設定してください。

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
