---
id: 66c8ba975ee7230e29f6c4b2
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

いよいよ`getLoanMessage`関数をテストする時間です。

以下のテーブルを使って`4`変数とその値を作成してください。

| 変数名       | 値         |
| ----------- | --------- |
| `duplexLoanMsg`       | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)`   |
| `carLoanMsg`   | `getLoanMessage(45000, 660)`   |
| `noLoanMsg`   | `getLoanMessage(25000, 550)`   |

その後、各変数をコンソールにログ出力してメッセージを確認してください。

これでローン資格チェッカープロジェクトは完成です！

# --hints--

`duplexLoanMsg`という名前の変数があるはずです。

```js
assert.isNotNull(duplexLoanMsg);
```

`duplexLoanMsg`変数は`getLoanMessage(85000, 850)`の結果に設定してください。

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

`duplexLoanMsg`変数をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

`condoLoanMsg`という名前の変数があるはずです。

```js
assert.isNotNull(condoLoanMsg);
```

`condoLoanMsg`変数は`getLoanMessage(65000, 690)`の結果に設定してください。

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

`condoLoanMsg`変数をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

`carLoanMsg`という名前の変数があるはずです。

```js
assert.isNotNull(carLoanMsg);
```

`carLoanMsg`変数は`getLoanMessage(45000, 660)`の結果に設定してください。

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

`carLoanMsg`変数をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

`noLoanMsg`という名前の変数があるはずです。

```js
assert.isNotNull(noLoanMsg);
```

`noLoanMsg`変数は`getLoanMessage(25000, 550)`の結果に設定してください。

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

`noLoanMsg`変数をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
