---
id: 66c8ba975ee7230e29f6c4ac
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

ユーザーがローンの対象となる場合、コンソールにメッセージを表示したいです。

そのために、申請者がどのローンの対象かを返すチェックを行う関数を作成します。

`getLoanMessage`という空の関数を、`annualIncome`と`creditScore`のパラメータ付きで作成してください。

# --hints--

`getLoanMessage`関数を作成してください

```js
assert.isFunction(getLoanMessage)
```

`getLoanMessage`関数は`annualIncome`と`creditScore`をパラメータとして持つべきです。

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage`は空の関数であるべきです。

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
