---
id: 6852dd8486c7ee0171322950
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

ستحتاج إلى متغير للاحتفاظ بالشفرة السرية أثناء فك تشفيرها، لذا احذف `pass` وأنشئ متغيرًا اسمه `secret_code` داخل الدالة وعيّن له سلسلة نصية فارغة.

# --hints--

يجب أن يكون لديك متغير `secret_code` داخل الدالة `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

يجب أن تكون قيمة المتغير `secret_code` سلسلة نصية فارغة.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

يجب ألا يكون لديك `pass` في شفرتك بعد الآن.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
