---
id: 6852db755c9c1274d75aae9a
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستنشئ مستخرج رمز PIN. أرقام الـ PIN مخفية في كل سطر من قصيدة.

ابدأ بإنشاء دالة تسمى `pin_extractor` مع معلمة `poem`. الدالة التي لا تحتوي على جسم ترجع خطأ، لذا أضف `pass` داخل الدالة كعنصر نائب مؤقت لجعلها صالحة.

# --hints--

يجب أن يكون لديك دالة باسم `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

يجب أن تحتوي الدالة `pin_extractor` على معلمة باسم `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
