---
id: 68530231c05b93324e071104
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

أثناء عملك على الدالة، سيكون من المفيد وجود استدعاء دالة لأغراض تصحيح الأخطاء، فأنشئ متغيرًا عامًّا باسم `poem` خارج الدالة يحتوي على هذه القصيدة:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

يمكنك استخدام سلسلة نصية متعددة الأسطر لذلك.

ثم، استدعِ الدالة مع المتغير `poem` كمعلمة.

# --hints--

يجب أن يكون لديك متغير عام باسم `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

يجب أن يحتوي المتغير `poem` على القيمة `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

يجب أن تستدعي الدالة `pin_extractor` مع `poem` كمعلمة.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
