---
id: 68531fbb934d735415053274
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

الآن بعد أن أضفت استدعاء دالة يمكنك رؤية أي خطأ يحدث داخل الدالة في الطرفية، ويمكنك دائمًا استخدام `print` لرؤية قيمة الأشياء.

يوجد رقم من رمز التعريف مخفي في كل سطر، لذا داخل الدالة استخدم طريقة `split` لتقسيم السلسلة النصية إلى قائمة من الأسطر. قسم الأسطر عند حرف السطر الجديد (`\n`)، وعيّن القائمة الناتجة إلى متغير اسمه `lines`.

# --hints--

يجب أن يكون لديك متغير اسمه `lines` داخل دالة `pin_extractor` الخاصة بك.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

يجب أن يحتوي المتغير `lines` على القيمة `poem.split('\n')`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
