---
id: 68559f3561d7df287586142e
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

الآن يعمل مستخرج الرمز، لكنه يمكن أن يكون أكثر كفاءة! يمكنه استخراج الرمز من عدة قصائد في نفس الوقت!

أضف قصيدتين إضافيتين في النطاق العام، `poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` و `poem3 = 'There\nonce\nwas\na\ndragon'`، وللآن، أضف تعليقًا على `print` مع الاستدعاء لـ `pin_extractor`.


# --hints--

يجب أن توجد المتغيرات `poem2` و `poem3` وأن تحمل القيمة الصحيحة.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

يجب إضافة تعليق على استدعاء `print(pin_extractor(poem))`.

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

--fcc-editable-region--


print(pin_extractor(poem))
--fcc-editable-region--

```
