---
id: 685578aae65e4106ad2b4300
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

كما تعلمت في درس سابق، تتيح لك الدالة `enumerate` تتبع الفهرس أثناء التكرار على قابل للتكرار:

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

الرقم النوني من الشفرة السرية يأتي من الكلمة النونية من السطر النوني، لذا تحتاج إلى معرفة رقم كل سطر.

غيّر الحلقة لتتكرر على `enumerate(lines)`، وأضف متغير حلقة آخر قبل `line` اسمه `line_index`.

أيضًا، حدّث استدعاء `print` إلى `print(line_index, line)`.

# --hints--

يجب أن تتكرر الحلقة على `enumerate(lines)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

يجب أن يكون لمتغير الحلقة `line_index, line`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

يجب أن يكون لديك `print(line_index, line)` داخل الحلقة.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
