---
id: 68553423d775cc60134ac8bb
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

يجب أن تعمل على كل سطر بشكل مستقل: أنشئ حلقة تكرار على `lines` تستخدم `line` كمتغير للحلقة. داخل الحلقة، اطبع `line`.

# --hints--

يجب أن تحتوي دالة `pin_extractor` على حلقة `for`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

يجب أن تكون الحلقة تكرارًا على `lines`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

يجب أن يكون متغير الحلقة هو `line` (`for line in lines`).

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

يجب أن تحتوي الحلقة على `print(line)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
