---
id: 685535bf527b2c77236b3a9a
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

الرقم النوني من رمز التعريف مخفي كطول الكلمة النونية في السطر النوني.

لإيجاد طول الكلمة النونية، الخطوة التالية هي فصل سطر القصيدة إلى قائمة كلمات: داخل الحلقة، أنشئ متغيرًا `words` وعيّن له قيمة `line` مقسمة إلى كلمات باستخدام طريقة `split`.

ثم أضف استدعاء `print` مع `words` كمعلمته.

# --hints--

يجب أن يكون لديك متغير `words` داخل الحلقة.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

يجب أن يحتوي متغير `words` على القيمة `line.split()`.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

يجب أن يكون لديك `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
