---
id: 6855918e87c33c6979d21d3f
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

الآن الدالة تقوم بعمل جيد في استخراج الشفرة من القصيدة، لكن هناك حالة خاصة يجب مراعاتها: قد يكون سطر القصيدة أقصر من المتوقع، وهذا سيؤدي إلى حدوث خطأ في الدالة (يمكنك تجربة ذلك بحذف كلمة `bright` من القصيدة).

ضع السطر `secret_code += str(len(words[line_index]))` داخل جملة `if` تتحقق من وجود عدد كافٍ من الكلمات في قائمة `words`.

# --hints--

داخل الحلقة يجب أن يكون لديك جملة `if`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

شرط جملة if يجب أن يكون `len(words) > line_index`.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

داخل جملة `if` يجب أن يكون لديك `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
