---
id: 6859a037b3d6721c7fbb8055
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

قبل الحلقة، عرّف متغيرًا جديدًا باسم `secret_codes` بقيمة ابتدائية قائمة فارغة.

بعدها استبدل السطر `return secret_code` بسطر يضيف `secret_code` إلى قائمة `secret_codes`.

# --hints--

يجب أن يكون لديك متغير `secret_codes` داخل الدالة `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_codes')
    `)
})
```

يجب أن تكون قيمة المتغير `secret_codes` قائمة فارغة.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_codes').is_equivalent('secret_codes = []')
    `)
})
```

في نهاية الحلقة يجب أن يكون لديك `secret_codes.append(secret_code)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('secret_codes.append(secret_code)')
    `)
})
```

لا يجب أن تُرجع `secret_code`.


```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_return('secret_code')
    `)
})
```


# --seed--

## --seed-contents--

```py
def pin_extractor(poems):
--fcc-editable-region--

    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        return secret_code
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))

```
