---
id: 68592b42b9c7f27cbe123203
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

السطر الثالث من القصيدة يفتقد كلمة ثالثة، لذلك رمز التعريف أقصر من المتوقع.

أضف بند `else` يربط `'0'` إلى `secret_code` عندما لا تكون هناك كلمات كافية، بحيث تُستخدم كل أسطر القصيدة لاستخراج الأرقام.

# --hints--

يجب أن يكون لديك بند `else`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1] == _Node()
    `)
})
```

يجب أن يحتوي بند `else` على `secret_code += '0'`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_if('len(words) > line_index').is_equivalent('if len(words) > line_index:\\n  secret_code += str(len(words[line_index]))\\nelse:\\n  secret_code += \\'0\\'')
    `)
})
```

يجب أن تُرجع الدالة `pin_extractor(poem)` القيمة `5202`.


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5202'
    `)
})
```

يجب أن تُرجع الدالة `secret_code('There\nonce\nwas\na\ndragon')` القيمة `50000`.

```js
({
    test: () => runPython(`
    assert pin_extractor('There\\nonce\\nwas\\na\\ndragon') == '50000'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
--fcc-editable-region--
        
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))

```
