---
id: 68592730aaf34b61ae279168
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

الطول الذي تُرجعه `len` هو عدد صحيح، تحتاج إلى تحويله إلى سلسلة نصية باستخدام الدالة `str` لتتمكن من إضافته إلى سلسلة `secret_code`.

حدّث استدعاء `print` ليطبع الطول كسلسلة نصية.

# --hints--

يجب أن يكون لديك `print(str(len(words[line_index])))` في نهاية الحلقة.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(len(words[line_index]))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
