---
id: 6989f1ba358acba651d5ee6b
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora configurarás la estructura de datos para almacenar tus resultados. Crea una variable llamada `result` e inicialízala como un arreglo vacío. Este arreglo almacenará todas las combinaciones válidas de paréntesis que generes.

Actualiza tu sentencia return para que devuelva `result` en lugar de un arreglo vacío.

# --hints--

Debes declarar una variable llamada `result` dentro de tu función `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.result);
```

Debes inicializar `result` como un arreglo vacío.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { result } = genParentheses?.variables ?? {};
assert.isFalse(result?.value.isEmpty());
const resultValue = eval(result?.value.toString());
assert.isArray(resultValue);
assert.isEmpty(resultValue);
```

Debes retornar `result` al final de tu función.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.isTrue(genParentheses?.hasReturn('result'));
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
