---
id: 6989f207306332c06b0add39
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Para el enfoque de búsqueda en anchura, usarás una cola para rastrear diferentes estados mientras construyes las combinaciones de paréntesis. Cada estado se representará como un arreglo que contiene tres elementos:

- La cadena actual que se está construyendo
- El número de paréntesis de apertura usados hasta ahora
- El número de paréntesis de cierre usados hasta ahora

Crea una variable llamada `queue` e inicialízala con un arreglo que contenga un arreglo: `['', 0, 0]`. Esto representa el estado inicial con una cadena vacía y cero paréntesis usados.

# --hints--

Debes declarar una variable llamada `queue` en tu función `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.queue);
```

Debes inicializar `queue` con un arreglo que contenga el arreglo `['', 0, 0]`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { queue } = genParentheses?.variables ?? {};
assert.isFalse(queue?.value.isEmpty());
const queueValue = eval(queue?.value.toString());
assert.deepEqual(queueValue, [['', 0, 0]]);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
