---
id: 6989f10a72ad424d4ad6d514
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

A continuación, necesitas validar que el número de pares sea al menos uno, ya que no puedes generar combinaciones de paréntesis con cero o un número negativo de pares.

Agrega otra sentencia `if` para verificar si `pairs` es menor que `1`. Si esta condición es verdadera, devuelve la cadena `The number of pairs should be at least 1`.

# --hints--

Deberías tener una segunda sentencia `if` con la condición `pairs < 1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

Debes devolver la cadena `The number of pairs should be at least 1` cuando `pairs` sea menor que `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
