---
id: 6989f2e35e6941e13e445156
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Ahora implementarás el ciclo principal de BFS. Crea un ciclo `while` que continúe mientras `queue` no esté vacío, es decir, mientras `queue.length` sea mayor que `0`.

Dentro del ciclo, muestra `queue` en la consola.

# --hints--

Debes crear un ciclo `while` con la condición `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

Debes mostrar `queue` en la consola dentro de tu ciclo `while`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
