---
id: 6989f207306332c06b0add39
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Para el enfoque de búsqueda en anchura, usarás una cola para rastrear diferentes estados mientras construyes las combinaciones de paréntesis. Cada estado se representará como un arreglo que contiene tres elementos:

- La cadena actual que se está construyendo
- La cantidad de paréntesis de apertura usados hasta ahora
- La cantidad de paréntesis de cierre usados hasta ahora

Crea una variable llamada `queue` e inicialízala con un arreglo que contenga un arreglo: `['', 0, 0]`. Esto representa el estado inicial con una cadena vacía y cero paréntesis usados.

# --hints--

Debes declarar una variable llamada `queue` en tu función `genParentheses`.

```js
assert.match(genParentheses.toString(), /queue/);
```

Debes inicializar `queue` con un arreglo que contenga el arreglo `['', 0, 0]`.

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  return result;
}
```
