---
id: 6989eb65324d97775c4965fe
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, implementarás una función que genere todas las combinaciones válidas de paréntesis usando un enfoque de búsqueda en anchura (BFS). Por ejemplo, las combinaciones válidas de dos pares de paréntesis son `(())` y `()()`.

Comienza creando una función llamada `genParentheses` con un solo parámetro `pairs`. Por ahora, devuelve un arreglo vacío desde la función.

# --hints--

Debes definir una función llamada `genParentheses`.

```js
assert.isFunction(genParentheses);
```

Tu función `genParentheses` debe tener un solo parámetro llamado `pairs`.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

Tu función debe devolver un arreglo vacío.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
