---
id: 6989ef885db3a3488a13b426
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Antes de implementar el algoritmo principal, necesitas validar la entrada. El parámetro `pairs` debe ser un entero, ya que representa la cantidad de pares de paréntesis a generar.

Agrega una sentencia `if` al inicio de tu función para verificar si `pairs` no es un entero. Usa el método `Number.isInteger()` para eso.

Si la condición es verdadera, devuelve la cadena `The number of pairs should be an integer`.

# --hints--

Debes tener una sentencia if con la condición `!Number.isInteger(pairs)` dentro de la función `genParentheses`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

Debes devolver una cadena cuando la entrada no sea un entero.

```js
let result = genParentheses("i");
assert.isString(result);
```

Debes devolver la cadena `The number of pairs should be an integer` cuando la entrada sea una cadena.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
