---
id: 6989f555728e63ad8b349d84
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Si la cadena actual aún no está completa, necesitas explorar los siguientes estados posibles. Añade una cláusula `else` para la declaración `if` actual.

Dentro del bloque `else`, manejarás la adición de paréntesis de apertura. Añade una declaración `if` para verificar si `opensUsed` es menor que `pairs`. Esto asegura que no uses más paréntesis de apertura de los permitidos.

Si esta condición es `true`, agrega un nuevo arreglo a la `queue`: `[current + '(', opensUsed + 1, closesUsed]`. Esto representa el estado después de añadir un paréntesis de apertura.

# --hints--

Deberías tener una cláusula `else` para tu declaración `if` dentro del ciclo `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

Deberías tener una declaración `if` dentro del bloque `else` con la condición `opensUsed < pairs`.

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

Deberías agregar `[current + '(', opensUsed + 1, closesUsed]` a la cola cuando la condición sea `true`.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
