---
id: 6989f48c3f0292c84f8d8938
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Ahora necesitas identificar cuándo has construido una combinación completa de paréntesis. Una combinación completa tiene una longitud igual al doble del número de pares (ya que cada par aporta un paréntesis de apertura y uno de cierre).

Dentro de tu ciclo `while`, añade una instrucción `if` para verificar si `current.length === 2 * pairs`. Cuando esta condición sea verdadera, agrega `current` al arreglo `result`.

# --hints--

Deberías tener una instrucción `if` con la condición `current.length === 2 * pairs` dentro de tu ciclo `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

Deberías agregar `current` al arreglo `result` dentro de tu nueva instrucción `if`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
