---
id: 6989f3632f8b8bdf168aa321
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Dentro de tu ciclo `while`, usa `queue.shift()` para eliminar y obtener el primer elemento de la cola. Esto implementa el comportamiento primero en entrar, primero en salir (FIFO) característico del BFS.

Desestructura este arreglo en tres variables: `current`, `opensUsed` y `closesUsed`.

# --hints--

Debes usar `queue.shift()` para obtener el primer elemento de la cola.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

Debes desestructurar el resultado en tres variables: `current`, `opensUsed` y `closesUsed`.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return result;
}
```
