---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Ahora necesitas manejar la adición de paréntesis de cierre. La regla clave es que solo puedes agregar un paréntesis de cierre si mantiene el equilibrio, es decir, debe haber más paréntesis de apertura usados que de cierre.

Todavía dentro del bloque `else`, agrega una segunda sentencia `if` para verificar si `closesUsed` es menor que `opensUsed`.

Si esta condición es verdadera, agrega otro nuevo arreglo a `queue`: `[current + ')', opensUsed, closesUsed + 1]`. Esto representa el estado después de agregar un paréntesis de cierre.

# --hints--

Deberías tener una segunda sentencia `if` con la condición `closesUsed < opensUsed` dentro de tu bloque `else`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

Deberías agregar `[current + ')', opensUsed, closesUsed + 1]` a la cola cuando la condición sea verdadera.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
