---
id: 56104e9e514f539506016a5c
title: for 루프로 홀수 반복하기
challengeType: 1
forumTopicId: 18212
dashedName: iterate-odd-numbers-with-a-for-loop
---

# --description--

for 루프는 한 번에 하나씩 반복할 필요가 없습니다. `final-expression`를 변경하면 짝수 단위로 셀 수 있습니다.

`i = 0`에서 시작해서 `i < 10` 동안 반복할 것입니다. `i`를 사용해 각 반복마다 `i += 2`를 2씩 증가시킬 것입니다.

```js
const ourArray = [];

for (let i = 0; i < 10; i += 2) {
  ourArray.push(i);
}
```

`ourArray`에는 이제 `[0, 2, 4, 6, 8]`가 들어 있습니다. 홀수를 세기 위해 `initialization`를 변경해 봅시다.

# --instructions--

`myArray` 루프를 사용해 1부터 9까지 홀수를 `for`에 추가하세요.

# --hints--

이 작업에는 `for` 루프를 사용해야 합니다.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

`myArray`는 `[1, 3, 5, 7, 9]`와 같아야 합니다.

```js
assert.deepEqual(myArray, [1, 3, 5, 7, 9]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 10; i += 2) {
  myArray.push(i);
}
```
