---
id: 56105e7b514f539506016a5e
title: for 루프로 역순으로 세기
challengeType: 1
forumTopicId: 16808
dashedName: count-backwards-with-a-for-loop
---

# --description--

for 루프는 올바른 조건을 정의할 수 있다면 역순으로도 셀 수 있습니다.

각 반복마다 2씩 감소시키려면 초기화, 조건, 마지막 표현식을 변경해야 합니다.

`i = 10`에서 시작해서 `i > 0`인 동안 반복할 것입니다. `i`를 사용해 매 루프마다 `i -= 2`를 2씩 감소시킬 것입니다.

```js
const ourArray = [];

for (let i = 10; i > 0; i -= 2) {
  ourArray.push(i);
}
```

`ourArray`에는 이제 `[10, 8, 6, 4, 2]`가 들어갑니다. 초기화와 마지막 표현식을 바꿔 2씩 역순으로 세어 내림차순 홀수 배열을 만들어 봅시다.

# --instructions--

`myArray` 루프를 사용해 9부터 1까지의 홀수를 `for`에 추가하세요.

# --hints--

이 작업에는 `for` 루프를 사용해야 합니다.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

배열 메서드 `push`를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/myArray.push/));
```

`myArray`은 `[9, 7, 5, 3, 1]`과 같아야 합니다.

```js
assert.deepEqual(myArray, [9, 7, 5, 3, 1]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 9; i > 0; i -= 2) {
  myArray.push(i);
}
```
