---
id: 56533eb9ac21ba0edf2244ad
title: JavaScript로 숫자 1 감소시키기
challengeType: 1
forumTopicId: 17558
dashedName: decrement-a-number-with-javascript
---

# --description--

`--` 연산자를 사용하면 변수를 1만큼 쉽게 <dfn>감소</dfn>시킬 수 있습니다.

```js
i--;
```

는 다음과 같습니다.

```js
i = i - 1;
```

**참고:** 등호가 필요 없도록 전체 줄이 `i--;`가 됩니다.

# --instructions--

코드를 수정하여 `--`에 `myVar` 연산자를 사용하세요.

# --hints--

`myVar`는 `10`와 같아야 합니다.

```js
assert(myVar === 10);
```

`myVar = myVar - 1;`를 변경해야 합니다.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*myVar\s*[-]\s*1.*?;?/));
```

`myVar`에 `10`를 할당하면 안 됩니다.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*10.*?;?/));
```

`--`에 `myVar` 연산자를 사용해야 합니다.

```js
assert(/[-]{2}\s*myVar|myVar\s*[-]{2}/.test(__helpers.removeJSComments(code)));
```

지정된 주석 위의 코드는 변경하지 마세요.

```js
assert(/let myVar = 11;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 11;

// Only change code below this line
myVar = myVar - 1;
```

# --solutions--

```js
let myVar = 11;
myVar--;
```
