---
id: 56533eb9ac21ba0edf2244bd
title: 인수를 사용해 함수에 값 전달하기
challengeType: 1
forumTopicId: 18254
dashedName: passing-values-to-functions-with-arguments
---

# --description--

<dfn>매개변수</dfn>는 함수가 호출될 때 입력될 값들을 위한 자리 표시자 역할을 하는 변수입니다. 함수를 정의할 때 보통 하나 이상의 매개변수를 함께 정의합니다. 함수가 호출될 때 실제로 입력(또는 <dfn>"전달"</dfn>)되는 값들은 <dfn>인수</dfn>라고 합니다.

다음은 두 개의 매개변수 `param1`와 `param2`를 가진 함수입니다:

```js
function testFun(param1, param2) {
  console.log(param1, param2);
}
```

그다음 `testFun`를 이렇게 호출할 수 있습니다: `testFun("Hello", "World");`. 두 개의 문자열 인수 `Hello`와 `World`를 전달했습니다. 함수 내부에서 `param1`는 문자열 `Hello`와 같고 `param2`는 문자열 `World`와 같습니다. `testFun`를 다른 인수로 다시 호출하면 매개변수는 새로운 인수의 값을 갖게 됩니다.

# --instructions--

<ol><li>두 개의 인수를 받아서 그 합을 개발자 콘솔에 출력하는 <code>functionWithArgs</code>라는 함수를 만드세요.</li><li>두 개의 숫자를 인수로 하여 함수를 호출하세요.</li></ol>

# --before-each--

```js
var logOutput;
```

# --hints--

`functionWithArgs`는 함수여야 합니다.

```js
assert(typeof functionWithArgs === 'function');
```

`functionWithArgs(1,2)`는 `3`를 출력해야 합니다.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(1, 2);
assert(logOutput == 3);
```

`functionWithArgs(7,9)`는 `16`를 출력해야 합니다.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(7, 9);
assert(logOutput == 16);
```

함수를 정의한 후 두 개의 숫자를 인수로 하여 `functionWithArgs`를 호출해야 합니다.

```js
assert(
  /functionWithArgs\([-+]?\d*\.?\d*,[-+]?\d*\.?\d*\)/.test(
    __helpers.removeJSComments(code).replace(/\s/g, '')
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function functionWithArgs(a, b) {
  console.log(a + b);
}
functionWithArgs(10, 5);
```
