---
id: 56533eb9ac21ba0edf2244d0
title: 동등 연산자 비교
challengeType: 1
forumTopicId: 16784
dashedName: comparison-with-the-equality-operator
---

# --description--

JavaScript에는 많은 <dfn>비교 연산자</dfn>가 있습니다. 이 연산자들은 모두 불리언 `true` 또는 `false` 값을 반환합니다.

가장 기본적인 연산자는 동등 연산자 `==`입니다. 동등 연산자는 두 값을 비교하여 같으면 `true`를, 다르면 `false`를 반환합니다. 동등은 오른쪽 값을 왼쪽 변수에 할당하는 할당(`=`)과 다르다는 점에 유의하세요.

```js
function equalityTest(myVal) {
  if (myVal == 10) {
    return "Equal";
  }
  return "Not Equal";
}
```

만약 `myVal`가 `10`와 같으면, 동등 연산자는 `true`를 반환하여 중괄호 안의 코드가 실행되고 함수는 `Equal`를 반환합니다. 그렇지 않으면 함수는 `Not Equal`를 반환합니다. JavaScript가 서로 다른 <dfn>데이터 타입</dfn> (예: `numbers`와 `strings`)을 비교하려면 한 타입을 다른 타입으로 변환해야 합니다. 이를 타입 강제 변환이라고 합니다. 변환 후에는 다음과 같이 비교할 수 있습니다:

```js
1   ==  1  // true
1   ==  2  // false
1   == '1' // true
"3" ==  3  // true
```

# --instructions--

지정된 줄에 동등 연산자를 추가하여 `Equal`가 `val`와 같을 때 함수가 문자열 `12`를 반환하도록 하세요.

# --hints--

`testEqual(10)`는 문자열 `Not Equal`를 반환해야 합니다

```js
assert(testEqual(10) === 'Not Equal');
```

`testEqual(12)`는 문자열 `Equal`를 반환해야 합니다

```js
assert(testEqual(12) === 'Equal');
```

`testEqual("12")`는 문자열 `Equal`를 반환해야 합니다

```js
assert(testEqual('12') === 'Equal');
```

`==` 연산자를 사용하세요

```js
assert(__helpers.removeJSComments(code).match(/==/g) && !__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function testEqual(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testEqual(10);
```

# --solutions--

```js
function testEqual(val) {
  if (val == 12) {
    return "Equal";
  }
  return "Not Equal";
}
```
