---
id: 56533eb9ac21ba0edf2244d7
title: 작거나 같다 연산자 비교
challengeType: 1
forumTopicId: 16788
dashedName: comparison-with-the-less-than-or-equal-to-operator
---

# --description--

작거나 같다 연산자(`<=`)는 두 숫자의 값을 비교합니다. 왼쪽 숫자가 오른쪽 숫자보다 작거나 같으면 `true`을 반환합니다. 왼쪽 숫자가 오른쪽 숫자보다 크면 `false`를 반환합니다. 동등 연산자처럼 작거나 같다 연산자도 데이터 타입을 변환합니다.

**예시**

```js
4   <= 5 // true
'7' <= 7 // true
5   <= 5 // true
3   <= 2 // false
'8' <= 4 // false
```

# --instructions--

return 문이 올바르게 작동하도록 표시된 줄에 작거나 같다 연산자를 추가하세요.

# --hints--

`testLessOrEqual(0)`는 문자열 `Smaller Than or Equal to 12`를 반환해야 합니다.

```js
assert(testLessOrEqual(0) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(11)`는 문자열 `Smaller Than or Equal to 12`를 반환해야 합니다.

```js
assert(testLessOrEqual(11) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(12)`는 문자열 `Smaller Than or Equal to 12`를 반환해야 합니다.

```js
assert(testLessOrEqual(12) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(23)`는 문자열 `Smaller Than or Equal to 24`를 반환해야 합니다.

```js
assert(testLessOrEqual(23) === 'Smaller Than or Equal to 24');
```

`testLessOrEqual(24)`는 문자열 `Smaller Than or Equal to 24`를 반환해야 합니다.

```js
assert(testLessOrEqual(24) === 'Smaller Than or Equal to 24');
```

`testLessOrEqual(25)`는 문자열 `More Than 24`를 반환해야 합니다.

```js
assert(testLessOrEqual(25) === 'More Than 24');
```

`testLessOrEqual(55)`는 문자열 `More Than 24`를 반환해야 합니다.

```js
assert(testLessOrEqual(55) === 'More Than 24');
```

작거나 같다 연산자 `<=`를 최소 두 번 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/val\s*<=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessOrEqual(val) {
  if (val) {  // Change this line
    return "Smaller Than or Equal to 12";
  }

  if (val) {  // Change this line
    return "Smaller Than or Equal to 24";
  }

  return "More Than 24";
}

testLessOrEqual(10);
```

# --solutions--

```js
function testLessOrEqual(val) {
  if (val <= 12) {  // Change this line
    return "Smaller Than or Equal to 12";
  }

  if (val <= 24) {  // Change this line
    return "Smaller Than or Equal to 24";
  }

  return "More Than 24";
}
```
