---
id: 56533eb9ac21ba0edf2244cd
title: 중첩 배열 접근하기
challengeType: 1
forumTopicId: 16160
dashedName: accessing-nested-arrays
---

# --description--

앞서 본 예제들처럼, 객체는 중첩 객체와 중첩 배열을 모두 포함할 수 있습니다. 중첩 객체에 접근하는 것과 비슷하게, 배열 대괄호 표기법을 연결해서 중첩 배열에 접근할 수 있습니다.

다음은 중첩 배열에 접근하는 방법의 예입니다:

```js
const ourPets = [
  {
    animalType: "cat",
    names: [
      "Meowzer",
      "Fluffy",
      "Kit-Cat"
    ]
  },
  {
    animalType: "dog",
    names: [
      "Spot",
      "Bowser",
      "Frankie"
    ]
  }
];

ourPets[0].names[1];
ourPets[1].names[0];
```

`ourPets[0].names[1]`는 문자열 `Fluffy`이고, `ourPets[1].names[0]`는 문자열 `Spot`입니다.

# --instructions--

점 표기법과 대괄호 표기법을 사용해서, `secondTree` 배열의 두 번째 객체에 있는 `list` 배열에서 두 번째 요소를 변수 `myPlants`에 할당하세요.

# --hints--

`secondTree`는 문자열 `pine`와 같아야 합니다.

```js
assert(secondTree === 'pine');
```

코드에서 점 표기법과 대괄호 표기법을 사용해 `myPlants`에 접근해야 합니다.

```js
assert(/=\s*myPlants\[1\].list\[1\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = "";
```

# --solutions--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = myPlants[1].list[1];
```
