---
id: 56533eb9ac21ba0edf2244d1
title: 엄격한 동등 연산자와 비교하기
challengeType: 1
forumTopicId: 16790
dashedName: comparison-with-the-strict-equality-operator
---

# --description--

엄격한 동등 연산자(`===`)는 동등 연산자(`==`)의 대응 관계입니다. 하지만 두 값을 공통 타입으로 변환하려고 시도하는 동등 연산자와 달리, 엄격한 동등 연산자는 타입 변환을 수행하지 않습니다.

비교하는 값들의 타입이 다르면, 서로 같지 않은 것으로 간주되어 엄격한 동등 연산자는 false를 반환합니다.

**예시**

```js
3 ===  3  // true
3 === '3' // false
```

두 번째 예시에서 `3`는 `Number` 타입이고 `'3'`는 `String` 타입입니다.

# --instructions--

`if` 문장에 엄격한 동등 연산자를 사용하여 `Equal`이 `val`와 엄격히 같을 때 함수가 문자열 `7`을 반환하도록 하세요.

# --hints--

`testStrict(10)`은 문자열 `Not Equal`을 반환해야 합니다

```js
assert(testStrict(10) === 'Not Equal');
```

`testStrict(7)`는 문자열 `Equal`을 반환해야 합니다

```js
assert(testStrict(7) === 'Equal');
```

`testStrict("7")`는 문자열 `Not Equal`을 반환해야 합니다

```js
assert(testStrict('7') === 'Not Equal');
```

`===` 연산자를 사용해야 합니다

```js
assert(__helpers.removeJSComments(code).match(/(val\s*===\s*\d+)|(\d+\s*===\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrict(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testStrict(10);
```

# --solutions--

```js
function testStrict(val) {
  if (val === 7) {
    return "Equal";
  }
  return "Not Equal";
}
```
