---
id: 56533eb9ac21ba0edf2244cb
title: 복잡한 객체 다루기
challengeType: 1
forumTopicId: 18208
dashedName: manipulating-complex-objects
---

# --description--

때때로 유연한 <dfn>자료 구조</dfn>에 데이터를 저장하고 싶을 때가 있습니다. JavaScript 객체는 유연한 데이터를 다루는 한 가지 방법입니다. 객체는 임의의 <dfn>문자열</dfn>, <dfn>숫자</dfn>, <dfn>불리언</dfn>, <dfn>배열</dfn>, <dfn>함수</dfn>, 그리고 <dfn>객체</dfn> 조합을 허용합니다.

다음은 복잡한 자료 구조의 예입니다:

```js
const ourMusic = [
  {
    "artist": "Daft Punk",
    "title": "Homework",
    "release_year": 1997,
    "formats": [ 
      "CD", 
      "Cassette", 
      "LP"
    ],
    "gold": true
  }
];
```

이것은 하나의 객체를 포함하는 배열입니다. 객체는 앨범에 관한 다양한 <dfn>메타데이터</dfn>를 가지고 있습니다. 또한 중첩된 `formats` 배열도 포함되어 있습니다. 더 많은 앨범 기록을 추가하려면 최상위 배열에 기록을 추가하면 됩니다. 객체는 키-값 형식의 속성에 데이터를 저장합니다. 위 예제에서 `"artist": "Daft Punk"`은 `artist`라는 키와 `Daft Punk`라는 값을 가진 속성입니다.

**참고:** 배열 내 모든 객체 뒤에는 마지막 객체가 아닌 경우 반드시 쉼표를 넣어야 합니다.

# --instructions--

`myMusic` 배열에 새 앨범을 추가하세요. `artist`와 `title` 문자열, `release_year` 숫자, 그리고 문자열로 이루어진 `formats` 배열을 추가하세요.

# --hints--

`myMusic`는 배열이어야 합니다

```js
assert(Array.isArray(myMusic));
```

`myMusic`은 최소 두 개의 요소를 가져야 합니다

```js
assert(myMusic.length > 1);
```

`myMusic` 배열의 요소들은 객체여야 합니다

```js
myMusic.forEach(object => {assert.typeOf(object, 'object')})
```

`myMusic`에 있는 객체는 최소 4개의 속성을 가져야 합니다

```js
myMusic.forEach(object => {assert(Object.keys(object).length > 3); });
```

`myMusic`에 있는 객체는 문자열인 `artist` 속성을 포함해야 합니다

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['artist']);
  assert.typeOf(object.artist, 'string')
})
```

`myMusic`에 있는 객체는 문자열인 `title` 속성을 포함해야 합니다

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['title']);
  assert.typeOf(object.title, 'string')
})
```

`myMusic`에 있는 객체는 숫자인 `release_year` 속성을 포함해야 합니다

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['release_year']);
  assert.typeOf(object.release_year, 'number')
})
```

`myMusic`에 있는 객체는 배열인 `formats` 속성을 포함해야 합니다

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['formats']);
  assert.typeOf(object.formats, 'array')
})
```

`formats`는 최소 두 개의 요소를 가진 문자열 배열이어야 합니다

```js
myMusic.forEach(object => {
  object.formats.forEach(format => {
    assert.typeOf(format, 'string')
  });
  assert.isAtLeast(object.formats.length, 2)
})
```

# --seed--

## --seed-contents--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CD",
      "8T",
      "LP"
    ],
    "gold": true
  }
];
```

# --solutions--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP" ],
    "gold": true
  },
  {
    "artist": "ABBA",
    "title": "Ring Ring",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP",
    "CD",
  ]
  }
];
```
