---
id: 56533eb9ac21ba0edf2244c0
title: 함수에서 전역 범위와 지역 범위
challengeType: 1
forumTopicId: 18194
dashedName: global-vs--local-scope-in-functions
---

# --description--

같은 이름의 <dfn>지역</dfn> 변수와 <dfn>전역</dfn> 변수를 동시에 가질 수 있습니다. 이럴 때는 지역 변수가 전역 변수보다 우선합니다.

이 예제에서:

```js
const someVar = "Hat";

function myFun() {
  const someVar = "Head";
  return someVar;
}
```

함수 `myFun`는 지역 변수 버전이 존재하기 때문에 문자열 `Head`을 반환합니다.

# --instructions--

`myOutfit` 함수에 지역 변수를 추가하여 `outerWear`의 값을 문자열 `sweater`로 재정의하세요.

# --hints--

전역 `outerWear`의 값은 변경하지 마세요.

```js
assert(outerWear === 'T-Shirt');
```

`myOutfit`는 문자열 `sweater`을 반환해야 합니다.

```js
assert(myOutfit() === 'sweater');
```

return 문은 변경하지 마세요.

```js
assert(/return outerWear/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const outerWear = "T-Shirt";

function myOutfit() {
  // Only change code below this line

  // Only change code above this line
  return outerWear;
}

myOutfit();
```

# --solutions--

```js
const outerWear = "T-Shirt";
function myOutfit() {
  const outerWear = "sweater";
  return outerWear;
}
```
