---
id: 56533eb9ac21ba0edf2244c2
title: "Return 키워드로 함수에서 값 반환하기"
challengeType: 1
forumTopicId: 18271
dashedName: return-a-value-from-a-function-with-return
---

# --description--

<dfn>인수</dfn>를 사용해 함수에 값을 전달할 수 있습니다. `return` 문을 사용해 함수 밖으로 값을 다시 보낼 수 있습니다.

**예시**

```js
function plusThree(num) {
  return num + 3;
}

const answer = plusThree(5);
```

`answer`는 `8` 값을 가집니다.

`plusThree`는 `num`에 대한 <dfn>인수</dfn>를 받고 `num + 3`와 같은 값을 반환합니다.

# --instructions--

하나의 인수를 받아 `timesFive`을 곱하고 새로운 값을 반환하는 `5` 함수를 만드세요.

# --hints--

`timesFive`는 함수여야 합니다

```js
assert(typeof timesFive === 'function');
```

`timesFive(5)`는 `25`을 반환해야 합니다

```js
assert(timesFive(5) === 25);
```

`timesFive(2)`는 `10`를 반환해야 합니다

```js
assert(timesFive(2) === 10);
```

`timesFive(0)`는 `0`를 반환해야 합니다

```js
assert(timesFive(0) === 0);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function timesFive(num) {
  return num * 5;
}
timesFive(10);
```
