---
id: 56533eb9ac21ba0edf2244b1
title: 복합 대입과 증강 곱셈
challengeType: 1
forumTopicId: 16662
dashedName: compound-assignment-with-augmented-multiplication
---

# --description--

`*=` 연산자는 변수를 숫자와 곱합니다.

```js
myVar = myVar * 5;
```

는 `myVar`에 `5`를 곱합니다. 이것은 다음과 같이 다시 쓸 수 있습니다:

```js
myVar *= 5;
```

# --instructions--

`a`, `b`, `c`에 대한 할당을 `*=` 연산자를 사용하도록 바꾸세요.

# --hints--

`a`는 `25`와 같아야 합니다.

```js
assert(a === 25);
```

`b`는 `36`와 같아야 합니다.

```js
assert(b === 36);
```

`c`는 `46`와 같아야 합니다.

```js
assert(c === 46);
```

각 변수에 `*=` 연산자를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\*=/g).length === 3);
```

지정된 주석 위의 코드는 수정하지 마세요.

```js
assert(
  /let a = 5;/.test(__helpers.removeJSComments(code)) &&
    /let b = 12;/.test(__helpers.removeJSComments(code)) &&
    /let c = 4\.6;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 12;
let c = 4.6;

// Only change code below this line
a = a * 5;
b = 3 * b;
c = c * 10;
```

# --solutions--

```js
let a = 5;
let b = 12;
let c = 4.6;

a *= 5;
b *= 3;
c *= 10;
```
