---
id: 56533eb9ac21ba0edf2244b8
title: 더하기 할당 연산자로 문자열 연결하기
challengeType: 1
forumTopicId: 16803
dashedName: concatenating-strings-with-the-plus-equals-operator
---

# --description--

`+=` 연산자를 사용해 기존 문자열 변수 끝에 문자열을 <dfn>연결</dfn>할 수도 있습니다. 긴 문자열을 여러 줄로 나누는 데 매우 유용합니다.

**참고:** 공백에 주의하세요. 연결은 문자열 사이에 공백을 자동으로 추가하지 않으므로 직접 넣어야 합니다.

예시:

```js
let ourStr = "I come first. ";
ourStr += "I come second.";
```

`ourStr`은 문자열 `I come first. I come second.` 값을 갖게 됩니다.

# --instructions--

`myStr`을 여러 줄에 걸쳐 `This is the first sentence.`와 `This is the second sentence.` 두 문자열을 `+=` 연산자를 사용해 연결하여 만드세요. 예시처럼 `+=` 연산자를 사용하고 두 문자열 사이에 공백을 반드시 포함하세요. 먼저 첫 번째 문자열을 `myStr`에 할당한 뒤 두 번째 문자열을 더하세요.

# --hints--

`myStr`에는 두 문자열 사이에 공백 문자 하나가 있어야 합니다.

```js
assert(/sentence\. This/.test(myStr));
```

`myStr`은 문자열 `This is the first sentence. This is the second sentence.` 값을 가져야 합니다.

```js
assert(myStr === 'This is the first sentence. This is the second sentence.');
```

`+=`을 만들 때 `myStr` 연산자를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*(["']).*\1/g));
```

# --seed--

## --seed-contents--

```js
let myStr;
```

# --solutions--

```js
let myStr = "This is the first sentence. ";
myStr += "This is the second sentence.";
```
