---
id: 56533eb9ac21ba0edf2244b2
title: 복합 대입과 증강 나누기
challengeType: 1
forumTopicId: 16659
dashedName: compound-assignment-with-augmented-division
---

# --description--

`/=` 연산자는 변수를 다른 숫자로 나눕니다.

```js
myVar = myVar / 5;
```

`myVar`을 `5`로 나눕니다. 이것은 다음과 같이 다시 쓸 수 있습니다:

```js
myVar /= 5;
```

# --instructions--

`a`, `b`, `c`에 대한 할당을 `/=` 연산자를 사용하도록 바꾸세요.

# --hints--

`a`은 `4`과 같아야 합니다.

```js
assert(a === 4);
```

`b`은 `27`과 같아야 합니다.

```js
assert(b === 27);
```

`c`은 `3`과 같아야 합니다.

```js
assert(c === 3);
```

각 변수에 `/=` 연산자를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\/=/g).length === 3);
```

지정된 주석 위의 코드는 수정하지 마세요.

```js
assert(
  /let a = 48;/.test(__helpers.removeJSComments(code)) &&
    /let b = 108;/.test(__helpers.removeJSComments(code)) &&
    /let c = 33;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 48;
let b = 108;
let c = 33;

// Only change code below this line
a = a / 12;
b = b / 4;
c = c / 11;
```

# --solutions--

```js
let a = 48;
let b = 108;
let c = 33;

a /= 12;
b /= 4;
c /= 11;
```
