---
id: 56533eb9ac21ba0edf2244d6
title: 작다 연산자 비교
challengeType: 1
forumTopicId: 16789
dashedName: comparison-with-the-less-than-operator
---

# --description--

작다 연산자(`<`)는 두 숫자의 값을 비교합니다. 왼쪽 숫자가 오른쪽 숫자보다 작으면 `true`을 반환합니다. 그렇지 않으면 `false`를 반환합니다. 동등 연산자처럼 작다 연산자도 비교할 때 데이터 유형을 변환합니다.

**예시**

```js
2   < 5 // true
'3' < 7 // true
5   < 5 // false
3   < 2 // false
'8' < 4 // false
```

# --instructions--

return 문이 의미 있게 작동하도록 표시된 줄에 작다 연산자를 추가하세요.

# --hints--

`testLessThan(0)`는 문자열 `Under 25`를 반환해야 합니다.

```js
assert(testLessThan(0) === 'Under 25');
```

`testLessThan(24)`는 문자열 `Under 25`를 반환해야 합니다.

```js
assert(testLessThan(24) === 'Under 25');
```

`testLessThan(25)`는 문자열 `Under 55`를 반환해야 합니다.

```js
assert(testLessThan(25) === 'Under 55');
```

`testLessThan(54)`는 문자열 `Under 55`를 반환해야 합니다.

```js
assert(testLessThan(54) === 'Under 55');
```

`testLessThan(55)`는 문자열 `55 or Over`를 반환해야 합니다.

```js
assert(testLessThan(55) === '55 or Over');
```

`testLessThan(99)`는 문자열 `55 or Over`를 반환해야 합니다.

```js
assert(testLessThan(99) === '55 or Over');
```

작다 연산자 `<`를 최소 두 번 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/val\s*<\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessThan(val) {
  if (val) {  // Change this line
    return "Under 25";
  }

  if (val) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}

testLessThan(10);
```

# --solutions--

```js
function testLessThan(val) {
  if (val < 25) {  // Change this line
    return "Under 25";
  }

  if (val < 55) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}
```
