---
id: 56533eb9ac21ba0edf2244c8
title: 대괄호 표기법으로 객체 속성 접근하기
challengeType: 1
forumTopicId: 16163
dashedName: accessing-object-properties-with-bracket-notation
---

# --description--

객체 속성에 접근하는 두 번째 방법은 대괄호 표기법(`[]`)입니다. 접근하려는 객체 속성 이름에 공백이 포함되어 있으면 대괄호 표기법을 사용해야 합니다.

하지만 공백이 없는 객체 속성에도 대괄호 표기법을 사용할 수 있습니다.

다음은 대괄호 표기법을 사용해 객체 속성을 읽는 예시입니다:

```js
const myObj = {
  "Space Name": "Kirk",
  "More Space": "Spock",
  "NoSpace": "USS Enterprise"
};

myObj["Space Name"];
myObj['More Space'];
myObj["NoSpace"];
```

`myObj["Space Name"]`은 문자열 `Kirk`이고, `myObj['More Space']`은 문자열 `Spock`이며, `myObj["NoSpace"]`는 문자열 `USS Enterprise`입니다.

속성 이름에 공백이 포함된 경우에는 반드시 따옴표(홑따옴표 또는 겹따옴표)로 감싸야 한다는 점을 유의하세요.

# --instructions--

대괄호 표기법을 사용해 `an entree`의 속성 `the drink`과 `testObj`의 값을 읽고 각각 `entreeValue`과 `drinkValue`에 할당하세요.

# --hints--

`entreeValue`은 문자열이어야 합니다

```js
assert(typeof entreeValue === 'string');
```

`entreeValue`의 값은 문자열 `hamburger`이어야 합니다

```js
assert(entreeValue === 'hamburger');
```

`drinkValue`은 문자열이어야 합니다

```js
assert(typeof drinkValue === 'string');
```

`drinkValue`의 값은 문자열 `water`이어야 합니다

```js
assert(drinkValue === 'water');
```

대괄호 표기법을 두 번 사용해야 합니다

```js
assert(__helpers.removeJSComments(code).match(/testObj\s*?\[('|")[^'"]+\1\]/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};

// Only change code below this line
const entreeValue = testObj;   // Change this line
const drinkValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};
const entreeValue = testObj["an entree"];
const drinkValue = testObj['the drink'];
```
