---
id: 56533eb9ac21ba0edf2244d2
title: 부등 연산자 비교
challengeType: 1
forumTopicId: 16787
dashedName: comparison-with-the-inequality-operator
---

# --description--

부등 연산자(`!=`)는 동등 연산자의 반대입니다. 부등은 같지 않음을 의미합니다. 부등 연산자는 동등 연산자가 `false`를 반환할 때 `true`을 반환하며, 그 반대도 마찬가지입니다. 동등 연산자처럼 부등 연산자도 비교할 때 값의 데이터 유형을 변환합니다.

**예시**

```js
1 !=  2    // true
1 != "1"   // false
1 != '1'   // false
1 != true  // false
0 != false // false
```

# --instructions--

함수가 `!=`이 `if`과 같지 않을 때 문자열 `Not Equal`를 반환하도록 `val` 문장에 부등 연산자 `99`를 추가하세요.

# --hints--

`testNotEqual(99)`는 문자열 `Equal`를 반환해야 합니다

```js
assert(testNotEqual(99) === 'Equal');
```

`testNotEqual("99")`는 문자열 `Equal`를 반환해야 합니다

```js
assert(testNotEqual('99') === 'Equal');
```

`testNotEqual(12)`는 문자열 `Not Equal`를 반환해야 합니다

```js
assert(testNotEqual(12) === 'Not Equal');
```

`testNotEqual("12")`는 문자열 `Not Equal`를 반환해야 합니다

```js
assert(testNotEqual('12') === 'Not Equal');
```

`testNotEqual("bob")`는 문자열 `Not Equal`를 반환해야 합니다

```js
assert(testNotEqual('bob') === 'Not Equal');
```

`!=` 연산자를 사용해야 합니다

```js
assert(__helpers.removeJSComments(code).match(/(?!!==)!=/));
```

# --seed--

## --seed-contents--

```js
// Setup
function testNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testNotEqual(10);
```

# --solutions--

```js
function testNotEqual(val) {
  if (val != 99) {
    return "Not Equal";
  }
  return "Equal";
}
```
