---
id: 56533eb9ac21ba0edf2244c4
title: 함수에서 조기 반환 패턴
challengeType: 1
forumTopicId: 18272
dashedName: return-early-pattern-for-functions
---

# --description--

`return` 문에 도달하면 현재 함수의 실행이 정지되고 제어가 호출한 위치로 돌아갑니다.

**예시**

```js
function myFun() {
  console.log("Hello");
  return "World";
  console.log("byebye")
}
myFun();
```

위 코드는 콘솔에 `Hello` 문자열을 출력하고 `World` 문자열을 반환합니다. `byebye` 문자열은 콘솔에 절대 출력되지 않는데, 함수가 `return` 문에서 종료되기 때문입니다.

# --instructions--

`abTest` 함수를 수정하여 `a` 또는 `b` 이 `0` 보다 작으면 함수가 즉시 `undefined` 값을 반환하고 종료되도록 하세요.

**힌트**  
<a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-javascript/understanding-uninitialized-variables" target="_blank" rel="noopener noreferrer nofollow"><code>undefined</code>는 문자열이 아니라 키워드</a>임을 기억하세요.

# --hints--

`abTest(2, 2)` 은 숫자를 반환해야 합니다

```js
assert(typeof abTest(2, 2) === 'number');
```

`abTest(2, 2)` 은 `8` 를 반환해야 합니다

```js
assert(abTest(2, 2) === 8);
```

`abTest(-2, 2)` 은 `undefined` 를 반환해야 합니다

```js
assert(abTest(-2, 2) === undefined);
```

`abTest(2, -2)` 은 `undefined` 를 반환해야 합니다

```js
assert(abTest(2, -2) === undefined);
```

`abTest(2, 8)` 은 `18` 를 반환해야 합니다

```js
assert(abTest(2, 8) === 18);
```

`abTest(3, 3)` 은 `12` 를 반환해야 합니다

```js
assert(abTest(3, 3) === 12);
```

`abTest(0, 0)` 은 `0` 를 반환해야 합니다

```js
assert(abTest(0, 0) === 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function abTest(a, b) {
  // Only change code below this line



  // Only change code above this line

  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}

abTest(2,2);
```

# --solutions--

```js
function abTest(a, b) {
  if(a < 0 || b < 0) {
    return undefined;
  }
  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}
```
