---
id: 56533eb9ac21ba0edf2244bc
title: 쇼핑 목록
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

변수 `myList`에 쇼핑 목록을 만드세요. 목록은 여러 개의 하위 배열을 포함하는 다차원 배열이어야 합니다.

각 하위 배열의 첫 번째 요소는 항목 이름을 나타내는 문자열이어야 합니다. 두 번째 요소는 수량을 나타내는 숫자여야 합니다. 예를 들어

```js
["Chocolate Bar", 15]
```

목록에는 최소 5개의 하위 배열이 있어야 합니다.

# --hints--

`myList`은 배열이어야 합니다.

```js
assert(Array.isArray(myList));
```

각 하위 배열의 첫 번째 요소는 모두 문자열이어야 합니다.

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

각 하위 배열의 두 번째 요소는 모두 숫자여야 합니다.

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

목록에는 최소 5개의 항목이 있어야 합니다.

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
