---
id: 56533eb9ac21ba0edf2244c6
title: 줄 서기
challengeType: 1
forumTopicId: 18307
dashedName: stand-in-line
---

# --description--

컴퓨터 과학에서 <dfn>queue</dfn>는 항목들이 순서대로 유지되는 추상 <dfn>자료 구조</dfn>입니다. 새 항목은 큐의 뒤에 추가되고, 오래된 항목은 큐의 앞에서 제거됩니다.

# --instructions--

배열(`nextInLine`)과 숫자(`arr`)를 인수로 받는 함수 `item`를 작성하세요.

숫자를 배열의 끝에 추가한 다음, 배열의 첫 번째 요소를 제거하세요.

그 후 `nextInLine` 함수는 제거된 요소를 반환해야 합니다.

# --hints--

`nextInLine([], 5)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(nextInLine([], 5));
```

`nextInLine([], 1)`는 `1`를 반환해야 합니다.

```js
assert(nextInLine([], 1) === 1);
```

`nextInLine([2], 1)`는 `2`를 반환해야 합니다.

```js
assert(nextInLine([2], 1) === 2);
```

`nextInLine([5,6,7,8,9], 1)`는 `5`를 반환해야 합니다.

```js
assert(nextInLine([5, 6, 7, 8, 9], 1) === 5);
```

# --seed--

## --seed-contents--

```js
function nextInLine(arr, item) {
  // Only change code below this line
  
  return item;
  // Only change code above this line
}

// Setup
let testArr = [1, 2, 3, 4, 5];

// Display code
console.log("Before: " + JSON.stringify(testArr));
console.log(nextInLine(testArr, 6));
console.log("After: " + JSON.stringify(testArr));
```

# --solutions--

```js
let testArr = [1, 2, 3, 4, 5];

function nextInLine(arr, item) {
    arr.push(item);
    return arr.shift();
}
```
