---
id: 56533eb9ac21ba0edf2244bb
title: 단어 빈칸 채우기
challengeType: 1
forumTopicId: 18377
dashedName: word-blanks
---

# --description--

명사, 동사, 형용사, 부사 같은 몇몇 단어가 빠진 문장이 제공됩니다. 그런 다음 완성된 문장이 의미가 통하도록 원하는 단어들로 빈칸을 채우세요.

다음 문장을 살펴보세요:

```md
It was really ____, and we ____ ourselves ____.
```

이 문장에는 형용사, 동사, 부사 세 개의 빈칸이 있으며, 원하는 단어를 넣어 완성할 수 있습니다. 완성된 문장은 다음과 같이 변수에 할당할 수 있습니다:

```js
const sentence = "It was really " + "hot" + ", and we " + "laughed" + " ourselves " + "silly" + ".";
```

# --instructions--

이 도전 과제에서는 명사, 동사, 형용사, 부사를 제공합니다. 제공된 단어들과 함께 원하는 단어를 사용해 완전한 문장을 만드세요.

제공된 변수 `+`, `myNoun`, `myAdjective`, `myVerb`를 사용해 문자열 연결 연산자 `myAdverb`로 새 문자열을 만들어야 합니다. 그리고 만들어진 문자열을 `wordBlanks` 변수에 할당하세요. 변수에 할당된 단어들은 변경하지 마세요.

문자열 사이에 공백이 들어가도록 처리해야 하며, 최종 결과는 완전한 문장이 되어야 합니다.

# --before-each--

```js
const removeAssignments = str => str
  .replace(/myNoun\s*=\s*["']dog["']/g, '')
  .replace(/myAdjective\s*=\s*["']big["']/g, '')
  .replace(/myVerb\s*=\s*["']ran["']/g, '')
  .replace(/myAdverb\s*=\s*["']quickly["']/g, '');
```

# --hints--

`wordBlanks`는 문자열이어야 합니다.

```js
assert(typeof wordBlanks === 'string');
```

`myNoun`, `myVerb`, `myAdjective`, `myAdverb`에 할당된 값은 변경하지 마세요.

```js
assert(
  myNoun === 'dog' &&
    myVerb === 'ran' &&
    myAdjective === 'big' &&
    myAdverb === 'quickly'
);
```

`dog`, `ran`, `big`, `quickly` 값을 직접 사용해 `wordBlanks`를 만들면 안 됩니다.

```js
const newCode = removeAssignments(__helpers.removeJSComments(code));
assert(
  !/dog/.test(newCode) &&
    !/ran/.test(newCode) &&
    !/big/.test(newCode) &&
    !/quickly/.test(newCode)
);
```

`wordBlanks`에는 `myNoun`, `myVerb`, `myAdjective`, `myAdverb` 변수에 할당된 모든 단어가 비단어 문자로 구분되어 포함되어야 하며, 원하는 추가 단어도 포함할 수 있습니다.

```js
assert(
  /\bdog\b/.test(wordBlanks) &&
    /\bbig\b/.test(wordBlanks) &&
    /\bran\b/.test(wordBlanks) &&
    /\bquickly\b/.test(wordBlanks)
);
```

# --seed--

## --seed-contents--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

// Only change code below this line
const wordBlanks = ""; // Change this line
// Only change code above this line
```

# --solutions--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

let wordBlanks = "Once there was a " + myNoun + " which was very " + myAdjective + ". ";
wordBlanks += "It " + myVerb + " " + myAdverb + " around the yard.";
```
