---
id: 56533eb9ac21ba0edf2244b4
title: 작은따옴표로 문자열 감싸기
challengeType: 1
forumTopicId: 18260
dashedName: quoting-strings-with-single-quotes
---

# --description--

<dfn>문자열</dfn> 값은 JavaScript에서 작은따옴표나 큰따옴표로 작성할 수 있으며, 시작과 끝에 같은 종류의 따옴표를 사용해야 합니다. 다른 프로그래밍 언어와 달리, JavaScript에서는 작은따옴표와 큰따옴표가 동일하게 작동합니다.

```js
const doubleQuoteStr = "This is a string"; 
const singleQuoteStr = 'This is also a string';
```

한 종류의 따옴표를 다른 종류보다 사용하고 싶은 이유는 문자열 안에 두 가지 따옴표를 모두 사용해야 할 때입니다. 예를 들어, 문자열 안에 대화를 저장하고 그 대화를 따옴표로 감싸고 싶을 때가 있습니다. 또 다른 경우는 문자열 안에 여러 속성이 따옴표로 감싸진 `<a>` 태그를 저장할 때입니다.

```js
const conversation = 'Finn exclaims to Jake, "Algebraic!"';
```

하지만 가장 바깥쪽 따옴표를 문자열 안에서 사용해야 한다면 문제가 생깁니다. 문자열은 시작과 끝에 같은 종류의 따옴표가 있어야 합니다. 그런데 중간에 같은 따옴표가 있으면 문자열이 조기에 끝나고 오류가 발생합니다.

```js
const goodStr = 'Jake asks Finn, "Hey, let\'s go on an adventure?"'; 
const badStr = 'Finn responds, "Let's go!"';
```

여기서 `badStr`는 오류를 발생시킵니다.

위의 <dfn>goodStr</dfn>에서는 역슬래시 `\`를 이스케이프 문자로 사용해 두 종류의 따옴표를 안전하게 사용할 수 있습니다.

**참고:** 역슬래시 `\`는 슬래시 `/`와 혼동하지 마세요. 두 문자는 같은 역할을 하지 않습니다.

# --instructions--

제공된 문자열을 시작과 끝에 작은따옴표를 사용하고 이스케이프 문자가 없는 문자열로 바꾸세요.

현재 문자열 안의 `<a>` 태그는 모든 곳에 큰따옴표를 사용하고 있습니다. 이스케이프 문자를 제거하려면 바깥쪽 따옴표를 작은따옴표로 바꿔야 합니다.

# --hints--

모든 역슬래시(`\`)를 제거하세요.

```js
assert(
  !/\\/g.test(__helpers.removeJSComments(code)) &&
    myStr.match(
      '\\s*<a href\\s*=\\s*"http://www.example.com"\\s*target\\s*=\\s*"_blank">\\s*Link\\s*</a>\\s*'
    )
);
```

작은따옴표 두 개 `'`와 큰따옴표 네 개 `"`가 있어야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/"/g).length === 4 && __helpers.removeJSComments(code).match(/'/g).length === 2);
```

# --seed--

## --seed-contents--

```js
const myStr = "<a href=\"http://www.example.com\" target=\"_blank\">Link</a>";
```

# --solutions--

```js
const myStr = '<a href="http://www.example.com" target="_blank">Link</a>';
```
