---
id: 56533eb9ac21ba0edf2244af
title: 복합 대입과 증강 덧셈
challengeType: 1
forumTopicId: 16661
dashedName: compound-assignment-with-augmented-addition
---

# --description--

프로그래밍에서는 변수의 내용을 수정하기 위해 할당을 자주 사용합니다. 등호 오른쪽에 있는 모든 것이 먼저 평가된다는 점을 기억하세요. 그래서 다음과 같이 말할 수 있습니다:

```js
myVar = myVar + 5;
```

`5`를 `myVar`에 더하기 위해서입니다. 이런 패턴이 매우 흔하기 때문에, 수학 연산과 할당을 한 번에 수행하는 연산자가 있습니다.

그중 하나가 `+=` 연산자입니다.

```js
let myVar = 1;
myVar += 5;
console.log(myVar);
```

콘솔에는 `6`가 출력될 것입니다.

# --instructions--

`a`, `b`, `c`에 대한 할당을 `+=` 연산자를 사용하도록 바꾸세요.

# --hints--

`a`는 `15`와 같아야 합니다.

```js
assert(a === 15);
```

`b`는 `26`와 같아야 합니다.

```js
assert(b === 26);
```

`c`는 `19`와 같아야 합니다.

```js
assert(c === 19);
```

각 변수에 `+=` 연산자를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\+=/g).length === 3);
```

지정된 주석 위의 코드는 수정하지 마세요.

```js
assert(
  /let a = 3;/.test(__helpers.removeJSComments(code)) &&
    /let b = 17;/.test(__helpers.removeJSComments(code)) &&
    /let c = 12;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 3;
let b = 17;
let c = 12;

// Only change code below this line
a = a + 12;
b = 9 + b;
c = c + 7;
```

# --solutions--

```js
let a = 3;
let b = 17;
let c = 12;

a += 12;
b += 9;
c += 7;
```
