---
id: 56533eb9ac21ba0edf2244e1
title: 중첩된 For 루프
challengeType: 1
forumTopicId: 18248
dashedName: nesting-for-loops
---

# --description--

다차원 배열이 있다면, 이전 웨이포인트와 같은 논리로 배열과 모든 하위 배열을 반복할 수 있습니다. 예시는 다음과 같습니다:

```js
const arr = [
  [1, 2], [3, 4], [5, 6]
];

for (let i = 0; i < arr.length; i++) {
  for (let j = 0; j < arr[i].length; j++) {
    console.log(arr[i][j]);
  }
}
```

이 코드는 `arr`의 각 하위 요소를 한 번에 하나씩 출력합니다. 내부 루프에서는 `.length` 자체가 배열이므로 `arr[i]`의 `arr[i]`을 확인하는 점에 주의하세요.

# --instructions--

`multiplyAll` 함수를 수정하여 `arr`의 하위 배열에 있는 모든 숫자의 곱을 반환하게 하세요.

# --hints--

`multiplyAll([[1], [2], [3]])`는 `6`을 반환해야 합니다.

```js
assert(multiplyAll([[1], [2], [3]]) === 6);
```

`multiplyAll([[1, 2], [3, 4], [5, 6, 7]])`는 `5040`을 반환해야 합니다.

```js
assert(
  multiplyAll([
    [1, 2],
    [3, 4],
    [5, 6, 7]
  ]) === 5040
);
```

`multiplyAll([[5, 1], [0.2, 4, 0.5], [3, 9]])`는 `54`을 반환해야 합니다.

```js
assert(
  multiplyAll([
    [5, 1],
    [0.2, 4, 0.5],
    [3, 9]
  ]) === 54
);
```

# --seed--

## --seed-contents--

```js
function multiplyAll(arr) {
  let product = 1;
  // Only change code below this line

  // Only change code above this line
  return product;
}

multiplyAll([[1, 2], [3, 4], [5, 6, 7]]);
```

# --solutions--

```js
function multiplyAll(arr) {
  let product = 1;
  for (let i = 0; i < arr.length; i++) {
    for (let j = 0; j < arr[i].length; j++) {
      product *= arr[i][j];
    }
  }
  return product;
}
```
