---
id: 56533eb9ac21ba0edf2244d5
title: "`>=` 연산자 비교"
challengeType: 1
forumTopicId: 16785
dashedName: comparison-with-the-greater-than-or-equal-to-operator
---

# --description--

크거나 같음 연산자(`>=`)는 두 숫자의 값을 비교합니다. 왼쪽 숫자가 오른쪽 숫자보다 크거나 같으면 `true`을 반환합니다. 그렇지 않으면 `false`를 반환합니다.

동등 연산자처럼 크거나 같음 연산자도 비교할 때 데이터 유형을 변환합니다.

**예시**

```js
6   >=  6  // true
7   >= '3' // true
2   >=  3  // false
'7' >=  9  // false
```

# --instructions--

return 문이 의미가 통하도록 표시된 줄에 크거나 같음 연산자를 추가하세요.

# --hints--

`testGreaterOrEqual(0)`는 문자열 `Less than 10`를 반환해야 합니다.

```js
assert(testGreaterOrEqual(0) === 'Less than 10');
```

`testGreaterOrEqual(9)`는 문자열 `Less than 10`를 반환해야 합니다.

```js
assert(testGreaterOrEqual(9) === 'Less than 10');
```

`testGreaterOrEqual(10)`는 문자열 `10 or Over`를 반환해야 합니다.

```js
assert(testGreaterOrEqual(10) === '10 or Over');
```

`testGreaterOrEqual(11)`는 문자열 `10 or Over`를 반환해야 합니다.

```js
assert(testGreaterOrEqual(11) === '10 or Over');
```

`testGreaterOrEqual(19)`는 문자열 `10 or Over`를 반환해야 합니다.

```js
assert(testGreaterOrEqual(19) === '10 or Over');
```

`testGreaterOrEqual(100)`는 문자열 `20 or Over`를 반환해야 합니다.

```js
assert(testGreaterOrEqual(100) === '20 or Over');
```

`testGreaterOrEqual(21)`는 문자열 `20 or Over`를 반환해야 합니다.

```js
assert(testGreaterOrEqual(21) === '20 or Over');
```

`>=` 연산자를 최소 두 번 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/val\s*>=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterOrEqual(val) {
  if (val) {  // Change this line
    return "20 or Over";
  }

  if (val) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}

testGreaterOrEqual(10);
```

# --solutions--

```js
function testGreaterOrEqual(val) {
  if (val >= 20) {  // Change this line
    return "20 or Over";
  }

  if (val >= 10) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}
```
