---
id: 56533eb9ac21ba0edf2244dd
title: switch 문으로 여러 옵션 중 선택하기
challengeType: 1
forumTopicId: 18277
dashedName: selecting-from-many-options-with-switch-statements
---

# --description--

하나의 값을 여러 옵션과 비교해야 할 때 <dfn>switch</dfn> 문을 사용할 수 있습니다. `switch` 문은 값을 다양한 가능한 값을 정의하는 <dfn>case</dfn> 문과 비교합니다. 유효한 JavaScript 문은 <dfn>case</dfn> 블록 안에서 실행할 수 있으며, 첫 번째로 일치하는 `case` 값부터 `break`를 만날 때까지 실행됩니다.

다음은 `switch` 문의 예입니다:

```js
switch (fruit) {
  case "apple":
    console.log("The fruit is an apple");
    break;
  case "orange":
    console.log("The fruit is an orange");
    break;
}
```

`case` 값은 엄격한 동등성(`===`)으로 테스트됩니다. `break`는 JavaScript에게 문 실행을 중단하라고 지시합니다. `break`가 생략되면 다음 문이 실행됩니다.

# --instructions--

`val`를 테스트하고 다음 조건에 따라 `answer`를 설정하는 switch 문을 작성하세요:  
`1` - `alpha`  
`2` - `beta`  
`3` - `gamma`  
`4` - `delta`

# --hints--

`caseInSwitch(1)`는 문자열 `alpha` 값을 가져야 합니다

```js
assert(caseInSwitch(1) === 'alpha');
```

`caseInSwitch(2)`는 문자열 `beta` 값을 가져야 합니다

```js
assert(caseInSwitch(2) === 'beta');
```

`caseInSwitch(3)`는 문자열 `gamma` 값을 가져야 합니다

```js
assert(caseInSwitch(3) === 'gamma');
```

`caseInSwitch(4)`는 문자열 `delta` 값을 가져야 합니다

```js
assert(caseInSwitch(4) === 'delta');
```

`if`나 `else` 문을 사용하지 마세요

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

최소 3개의 `break` 문이 있어야 합니다

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function caseInSwitch(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

caseInSwitch(1);
```

# --solutions--

```js
function caseInSwitch(val) {
  let answer = "";

  switch (val) {
    case 1:
      answer = "alpha";
      break;
    case 2:
      answer = "beta";
      break;
    case 3:
      answer = "gamma";
      break;
    case 4:
      answer = "delta";
  }
  return answer;
}
```
