---
id: 56533eb9ac21ba0edf2244aa
title: 초기화되지 않은 변수 이해하기
challengeType: 1
forumTopicId: 18335
dashedName: understanding-uninitialized-variables
---

# --description--

JavaScript 변수를 선언하면 초기값은 `undefined`입니다. `undefined` 변수에 수학 연산을 하면 결과는 `NaN`가 되는데, 이는 <dfn>"숫자가 아님"</dfn>을 의미합니다. 문자열과 `undefined` 변수를 연결하면 `undefined`의 <dfn>문자열</dfn>이 생성됩니다.

# --instructions--

세 변수 `a`, `b`, `c`을 각각 `5`, `10`, `"I am a"`으로 초기화하여 `undefined` 상태가 되지 않도록 하세요.

# --hints--

`a`은 정의되어 있어야 하며 최종 값은 `6`이어야 합니다.

```js
assert(typeof a === 'number' && a === 6);
```

`b`은 정의되어 있어야 하며 최종 값은 `15`이어야 합니다.

```js
assert(typeof b === 'number' && b === 15);
```

`c`에는 `undefined`이 포함되어 있지 않아야 하며 최종 값은 문자열 `I am a String!`이어야 합니다.

```js
assert(!/undefined/.test(c) && c === 'I am a String!');
```

지정된 주석 아래 코드는 변경하지 마세요.

```js
assert(
  /a = a \+ 1;/.test(__helpers.removeJSComments(code)) &&
    /b = b \+ 5;/.test(__helpers.removeJSComments(code)) &&
    /c = c \+ " String!";/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
var a;
var b;
var c;
// Only change code above this line

a = a + 1;
b = b + 5;
c = c + " String!";
```

# --solutions--

```js
var a = 5;
var b = 10;
var c = "I am a";
a = a + 1;
b = b + 5;
c = c + " String!";
```
