---
id: 56533eb9ac21ba0edf2244bf
title: 지역 범위와 함수
challengeType: 1
forumTopicId: 18227
dashedName: local-scope-and-functions
---

# --description--

함수 내에서 선언된 변수와 함수 매개변수는 <dfn>지역</dfn> 범위를 가집니다. 즉, 해당 함수 내에서만 보입니다.

다음은 `myTest`라는 지역 변수를 가진 함수입니다 `loc`.

```js
function myTest() {
  const loc = "foo";
  console.log(loc);
}

myTest();
console.log(loc);
```

`myTest()` 함수 호출은 콘솔에 문자열 `foo`을 표시합니다. `console.log(loc)` 줄(`myTest` 함수 외부)은 오류를 발생시키는데, `loc`가 함수 외부에서 정의되지 않았기 때문입니다.

# --instructions--

편집기에는 무슨 일이 일어나는지 볼 수 있도록 두 개의 `console.log`가 있습니다. 코드를 작성하면서 콘솔을 확인해 변화를 보세요. `myVar` 안에 지역 변수 `myLocalScope`를 선언하고 테스트를 실행하세요.

**참고:** 콘솔에는 여전히 `ReferenceError: myVar is not defined`이 표시되지만, 이는 테스트 실패를 일으키지 않습니다.

# --hints--

코드에는 전역 `myVar` 변수가 없어야 합니다.

```js
function declared() {
  myVar;
}

assert.throws(declared, ReferenceError);
```

지역 `myVar` 변수를 추가하세요.

```js
assert(
  /functionmyLocalScope\(\)\{.*(var|let|const)myVar[\s\S]*}/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

# --seed--

## --seed-contents--

```js
function myLocalScope() {
  // Only change code below this line

  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```

# --solutions--

```js
function myLocalScope() {
  // Only change code below this line
  let myVar;
  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```
