---
id: 56533eb9ac21ba0edf2244b9
title: 변수로 문자열 만들기
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

가끔 문자열을 만들어야 할 때가 있습니다. 연결 연산자(`+`)를 사용하면 하나 이상의 변수를 만들고 있는 문자열에 삽입할 수 있습니다.

예시:

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

`ourStr`은 문자열 `Hello, our name is freeCodeCamp, how are you?` 값을 갖게 됩니다.

# --instructions--

`myName`을 자신의 이름과 같은 문자열로 설정하고 `myStr`과 `myName` 문자열 사이에 `My name is `를 넣어 ` and I am well!`를 만드세요.

# --hints--

`myName`은 최소 3자 이상의 문자열로 설정하세요.

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

`+`을 포함하는 `myStr`을 만들려면 `myName` 연결 연산자를 두 번 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
