---
id: 56533eb9ac21ba0edf2244b7
title: 더하기 연산자로 문자열 연결하기
challengeType: 1
forumTopicId: 16802
dashedName: concatenating-strings-with-plus-operator
---

# --description--

JavaScript에서 `+` 연산자가 `String` 값과 함께 사용될 때, 이를 <dfn>연결</dfn> 연산자라고 합니다. 여러 문자열을 <dfn>연결</dfn>하여 새로운 문자열을 만들 수 있습니다.

**예시**

```js
'My name is Alan,' + ' I concatenate.'
```

**참고:** 공백에 주의하세요. 연결 연산자는 문자열 사이에 공백을 자동으로 추가하지 않으므로 직접 공백을 넣어야 합니다.

예시:

```js
const ourStr = "I come first. " + "I come second.";
```

문자열 `I come first. I come second.`가 콘솔에 출력됩니다.
# --instructions--

`myStr`를 문자열 `This is the start.`와 `This is the end.`를 `+` 연산자를 사용해 만드세요. 두 문자열 사이에 공백을 반드시 포함하세요.

# --hints--

`myStr`에는 두 문자열 사이에 공백 문자 하나가 있어야 합니다.

```js
assert(/start\. This/.test(myStr));
```

`myStr`는 문자열 `This is the start. This is the end.` 값을 가져야 합니다.

```js
assert(myStr === 'This is the start. This is the end.');
```

`+`를 만들 때 `myStr` 연산자를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/(["']).*\1\s*\+\s*(["']).*\2/g));
```

`myStr`는 `const` 키워드를 사용해 만들어야 합니다.

```js
assert(/const\s+myStr/.test(__helpers.removeJSComments(code)));
```

결과를 `myStr` 변수에 할당하세요.

```js
assert(/myStr\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "This is the start. " + "This is the end.";
```
