---
id: 56533eb9ac21ba0edf2244ab
title: 변수에서 대소문자 구분 이해하기
challengeType: 1
forumTopicId: 18334
dashedName: understanding-case-sensitivity-in-variables
---

# --description--

JavaScript에서는 모든 변수와 함수 이름이 대소문자를 구분합니다. 즉, 대문자와 소문자가 중요합니다.

`MYVAR`는 `MyVar`도 아니고 `myvar`도 아닙니다. 같은 이름이지만 대소문자가 다른 여러 개의 별도 변수를 가질 수 있습니다. 명확성을 위해 이 언어 기능을 *사용하지 않는 것*을 강력히 권장합니다.

**모범 사례**

JavaScript에서 변수 이름은 <dfn>카멜 케이스</dfn>로 작성하세요. <dfn>카멜 케이스</dfn>에서는 여러 단어로 된 변수 이름의 첫 단어는 소문자이고, 그 다음 단어들의 첫 글자는 대문자입니다.

**예시:**

```js
var someVariable;
var anotherVariableName;
var thisVariableNameIsSoLong;
```

# --instructions--

기존 선언과 할당을 수정하여 이름이 <dfn>카멜 케이스</dfn>를 사용하도록 하세요.

새 변수를 만들지 마세요.

# --hints--

`studlyCapVar`는 선언되어 있어야 하며 `10` 값을 가져야 합니다.

```js
assert(typeof studlyCapVar !== 'undefined' && studlyCapVar === 10);
```

`properCamelCase`는 선언되어 있어야 하며 문자열 `A String` 값을 가져야 합니다.

```js
assert(
  typeof properCamelCase !== 'undefined' && properCamelCase === 'A String'
);
```

`titleCaseOver`는 선언되어 있어야 하며 `9000` 값을 가져야 합니다.

```js
assert(typeof titleCaseOver !== 'undefined' && titleCaseOver === 9000);
```

`studlyCapVar`는 선언과 할당 모두에서 카멜 케이스를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/studlyCapVar/g).length === 2);
```

`properCamelCase`는 선언과 할당 모두에서 카멜 케이스를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/properCamelCase/g).length === 2);
```

`titleCaseOver`는 선언과 할당 모두에서 카멜 케이스를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/titleCaseOver/g).length === 2);
```

# --seed--

## --seed-contents--

```js
// Variable declarations
var StUdLyCapVaR;
var properCamelCase;
var TitleCaseOver;

// Variable assignments
STUDLYCAPVAR = 10;
PRoperCAmelCAse = "A String";
tITLEcASEoVER = 9000;
```

# --solutions--

```js
var studlyCapVar;
var properCamelCase;
var titleCaseOver;

studlyCapVar = 10;
properCamelCase = "A String";
titleCaseOver = 9000;
```
