---
id: 56533eb9ac21ba0edf2244b6
title: 문자열에서의 이스케이프 시퀀스
challengeType: 1
forumTopicId: 17567
dashedName: escape-sequences-in-strings
---

# --description--

문자열 안에서 이스케이프할 수 있는 문자는 따옴표만 있는 것이 아닙니다. 이스케이프 시퀀스를 사용하면 문자열에서 일반적으로 사용할 수 없는 문자를 사용할 수 있습니다.

<table><thead><tr><th>코드</th><th>출력</th></tr></thead><tbody><tr><td><code>\'</code></td><td>작은따옴표</td></tr><tr><td><code>\"</code></td><td>큰따옴표</td></tr><tr><td><code>\\</code></td><td>역슬래시</td></tr><tr><td><code>\n</code></td><td>줄바꿈</td></tr><tr><td><code>\t</code></td><td>탭</td></tr><tr><td><code>\r</code></td><td>캐리지 리턴</td></tr><tr><td><code>\b</code></td><td>백스페이스</td></tr><tr><td><code>\f</code></td><td>폼 피드</td></tr></tbody></table>

*역슬래시 자체를 출력하려면 역슬래시를 이스케이프해야 한다는 점을 유의하세요.*

# --instructions--

다음 세 줄의 텍스트를 이스케이프 시퀀스를 사용해 단일 변수 `myStr`에 할당하세요.

<pre>
FirstLine
    \SecondLine
ThirdLine
</pre>

특수 문자를 올바르게 삽입하려면 이스케이프 시퀀스를 사용해야 합니다. 또한 위에 보이는 대로 공백 없이 이스케이프 시퀀스와 단어 사이에 공백이 없도록 띄어쓰기를 지켜야 합니다.

**참고:** `SecondLine`의 들여쓰기는 공백이 아니라 탭 이스케이프 문자로 구현해야 합니다.

# --hints--

`myStr`에는 공백이 없어야 합니다

```js
assert(!/ /.test(myStr));
```

`myStr`에는 `FirstLine`, `SecondLine` 그리고 `ThirdLine` 문자열이 포함되어야 합니다 (대소문자 구분을 기억하세요)

```js
assert(
  /FirstLine/.test(myStr) && /SecondLine/.test(myStr) && /ThirdLine/.test(myStr)
);
```

`FirstLine` 다음에는 줄바꿈 문자 `\n`가 와야 합니다

```js
assert(/FirstLine\n/.test(myStr));
```

`myStr`에는 줄바꿈 문자 다음에 탭 문자 `\t`가 포함되어야 합니다

```js
assert(/\n\t/.test(myStr));
```

`SecondLine` 앞에는 역슬래시 문자 `\`가 있어야 합니다

```js
assert(/\\SecondLine/.test(myStr));
```

`SecondLine`과 `ThirdLine` 사이에는 줄바꿈 문자가 있어야 합니다

```js
assert(/SecondLine\nThirdLine/.test(myStr));
```

`myStr`에는 지시문에 나온 문자만 포함되어야 합니다

```js
assert(myStr === 'FirstLine\n\t\\SecondLine\nThirdLine');
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "FirstLine\n\t\\SecondLine\nThirdLine";
```
