---
id: 56533eb9ac21ba0edf2244c9
title: 변수로 객체 속성 접근하기
challengeType: 1
forumTopicId: 16165
dashedName: accessing-object-properties-with-variables
---

# --description--

객체에서 대괄호 표기법을 사용하는 또 다른 방법은 변수의 값으로 저장된 속성에 접근하는 것입니다. 이는 객체의 속성을 반복하거나 조회 표를 접근할 때 매우 유용합니다.

다음은 변수를 사용해 속성에 접근하는 예시입니다:

```js
const dogs = {
  Fido: "Mutt",
  Hunter: "Doberman",
  Snoopie: "Beagle"
};

const myDog = "Hunter";
const myBreed = dogs[myDog];
console.log(myBreed);
```

문자열 `Doberman`가 콘솔에 출력됩니다.

변수 이름을 사용할 때는 변수의 *이름*이 아니라 *값*을 사용하기 때문에 변수 이름 주위에 따옴표를 *사용하지 않는 것*에 주의하세요.

# --instructions--

`playerNumber` 변수를 `16`로 설정하세요. 그런 다음, 변수를 사용해 플레이어의 이름을 찾아 `player`에 할당하세요.

# --hints--

`playerNumber`는 숫자여야 합니다

```js
assert(typeof playerNumber === 'number');
```

변수 `player`는 문자열이어야 합니다

```js
assert(typeof player === 'string');
```

`player`의 값은 문자열 `Montana`이어야 합니다

```js
assert(player === 'Montana');
```

`testObj`에 접근할 때 대괄호 표기법을 사용해야 합니다

```js
assert(/testObj\s*?\[.*?\]/.test(__helpers.removeJSComments(code)));
```

변수 `Montana`에 값 `player`를 직접 할당하면 안 됩니다

```js
assert(!__helpers.removeJSComments(code).match(/player\s*=\s*"|\'\s*Montana\s*"|\'\s*;/gi));
```

대괄호 표기법에서 변수 `playerNumber`를 사용해야 합니다

```js
assert(/testObj\s*?\[\s*playerNumber\s*\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};

// Only change code below this line
const playerNumber = 42;  // Change this line
const player = testObj;   // Change this line
```

# --solutions--

```js
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};
const playerNumber = 16;
const player = testObj[playerNumber];
```
