---
id: 56533eb9ac21ba0edf2244ba
title: 문자열 불변성 이해하기
challengeType: 1
forumTopicId: 18331
dashedName: understand-string-immutability
---

# --description--

JavaScript에서 `String` 값은 <dfn>불변</dfn>이어서, 한 번 생성되면 변경할 수 없습니다.

예를 들어, 다음 코드는 문자열 `B` 안의 문자 `Bob`를 문자 `J`로 바꾸려고 해서 오류가 발생합니다:

```js
let myStr = "Bob";
myStr[0] = "J";
```

여기서 `myStr`를 재할당할 수 없다는 뜻은 아닙니다. `myStr`를 바꾸려면, 다음과 같이 새 값으로 할당해야 합니다:

```js
let myStr = "Bob";
myStr = "Job";
```

# --instructions--

예시에서 보여준 방법을 사용하여 `myStr`에 문자열 `Hello World` 값을 할당하도록 수정하세요.

# --hints--

`myStr`는 문자열 `Hello World` 값을 가져야 합니다.

```js
assert(myStr === 'Hello World');
```

지정된 주석 위의 코드는 변경하지 마세요.

```js
assert(/myStr = "Jello World"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let myStr = "Jello World";

// Only change code below this line
myStr[0] = "H"; // Change this line
// Only change code above this line
```

# --solutions--

```js
let myStr = "Jello World";
myStr = "Hello World";
```
