---
id: 56533eb9ac21ba0edf2244be
title: 전역 범위와 함수
challengeType: 1
forumTopicId: 18193
dashedName: global-scope-and-functions
---

# --description--

JavaScript에서 <dfn>범위</dfn>는 변수의 가시성을 의미합니다. 함수 블록 외부에 정의된 변수는 <dfn>전역</dfn> 범위를 가집니다. 즉, JavaScript 코드 어디에서나 볼 수 있습니다.

`let` 또는 `const` 키워드 없이 선언된 변수는 자동으로 `global` 범위에 생성됩니다. 이는 코드의 다른 부분이나 함수를 다시 실행할 때 의도하지 않은 결과를 초래할 수 있습니다. 항상 `let` 또는 `const`로 변수를 선언해야 합니다.

# --instructions--

`let` 또는 `const`를 사용하여 함수 외부에 `myGlobal`라는 전역 변수를 선언하세요. 값은 `10`로 초기화하세요.

함수 `fun1` 내부에서 `5`, `oopsGlobal`, `var` 키워드를 사용하지 않고 `let`에 `const`을 할당하세요.

# --before-each--

```js
var oopsGlobal;
```

# --hints--

`myGlobal`가 정의되어 있어야 합니다

```js
assert(typeof myGlobal != 'undefined');
```

`myGlobal`는 `10` 값을 가져야 합니다

```js
assert(myGlobal === 10);
```

`myGlobal`는 `let` 또는 `const` 키워드로 선언되어야 합니다

```js
assert(/(let|const)\s+myGlobal/.test(__helpers.removeJSComments(code)));
```

`oopsGlobal`는 전역 변수여야 하며 `5` 값을 가져야 합니다

```js
fun1();
assert(typeof oopsGlobal != 'undefined');
```

# --seed--

## --seed-contents--

```js
// Declare the myGlobal variable below this line


function fun1() {
  // Assign 5 to oopsGlobal here

}

// Only change code above this line

function fun2() {
  let output = "";
  if (typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if (typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```

# --solutions--

```js
const myGlobal = 10;

function fun1() {
  oopsGlobal = 5;
}

function fun2() {
  let output = "";
  if(typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if(typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```
