---
id: 56533eb9ac21ba0edf2244d8
title: 논리 and 연산자를 사용한 비교
challengeType: 1
forumTopicId: 16799
dashedName: comparisons-with-the-logical-and-operator
---

# --description--

때때로 한 번에 두 가지 이상을 테스트해야 할 때가 있습니다. <dfn>논리 and</dfn> 연산자(`&&`)는 좌우 <dfn>피연산자</dfn>가 모두 참일 때에만 `true`을 반환합니다.

같은 효과는 한 `if` 안에 `if` 문을 중첩하여도 얻을 수 있습니다.

```js
if (num > 5) {
  if (num < 10) {
    return "Yes";
  }
}
return "No";
```

이 코드는 `num`가 `5`보다 크고 `10`보다 작으면 `Yes`를 반환합니다. 같은 논리를 <dfn>논리 and</dfn> 연산자로도 작성할 수 있습니다.

```js
if (num > 5 && num < 10) {
  return "Yes";
}
return "No";
```

# --instructions--

두 개의 if 문을 `&&` 연산자를 사용한 하나의 문으로 바꾸세요. `val`이 `50`보다 작거나 같고 `25`보다 크거나 같으면 문자열 `Yes`를 반환합니다. 그렇지 않으면 문자열 `No`를 반환하세요.

# --hints--

`&&` 연산자를 한 번 사용하세요

```js
assert(__helpers.removeJSComments(code).match(/&&/g).length === 1);
```

`if` 문은 한 개만 있어야 합니다

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalAnd(0)`는 문자열 `No`를 반환해야 합니다

```js
assert(testLogicalAnd(0) === 'No');
```

`testLogicalAnd(24)`는 문자열 `No`를 반환해야 합니다

```js
assert(testLogicalAnd(24) === 'No');
```

`testLogicalAnd(25)`는 문자열 `Yes`를 반환해야 합니다

```js
assert(testLogicalAnd(25) === 'Yes');
```

`testLogicalAnd(30)`는 문자열 `Yes`를 반환해야 합니다

```js
assert(testLogicalAnd(30) === 'Yes');
```

`testLogicalAnd(50)`는 문자열 `Yes`를 반환해야 합니다

```js
assert(testLogicalAnd(50) === 'Yes');
```

`testLogicalAnd(51)`는 문자열 `No`를 반환해야 합니다

```js
assert(testLogicalAnd(51) === 'No');
```

`testLogicalAnd(75)`는 문자열 `No`를 반환해야 합니다

```js
assert(testLogicalAnd(75) === 'No');
```

`testLogicalAnd(80)`는 문자열 `No`를 반환해야 합니다

```js
assert(testLogicalAnd(80) === 'No');
```

# --seed--

## --seed-contents--

```js
function testLogicalAnd(val) {
  // Only change code below this line

  if (val) {
    if (val) {
      return "Yes";
    }
  }

  // Only change code above this line
  return "No";
}

testLogicalAnd(10);
```

# --solutions--

```js
function testLogicalAnd(val) {
  if (val >= 25 && val <= 50) {
    return "Yes";
  }
  return "No";
}
```
