---
id: 56533eb9ac21ba0edf2244b0
title: 복합 대입과 증강 뺄셈
challengeType: 1
forumTopicId: 16660
dashedName: compound-assignment-with-augmented-subtraction
---

# --description--

`+=` 연산자처럼, `-=`는 변수에서 숫자를 뺍니다.

```js
myVar = myVar - 5;
```

는 `5`에서 `myVar`를 뺍니다. 이것은 다음과 같이 다시 쓸 수 있습니다:

```js
myVar -= 5;
```

# --instructions--

`a`, `b`, 그리고 `c`에 대한 할당을 `-=` 연산자를 사용하도록 바꾸세요.

# --hints--

`a`는 `5`와 같아야 합니다.

```js
assert(a === 5);
```

`b`는 `-6`와 같아야 합니다.

```js
assert(b === -6);
```

`c`는 `2`와 같아야 합니다.

```js
assert(c === 2);
```

각 변수에 `-=` 연산자를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/-=/g).length === 3);
```

지정된 주석 위의 코드는 수정하지 마세요.

```js
assert(
  /let a = 11;/.test(__helpers.removeJSComments(code)) && /let b = 9;/.test(__helpers.removeJSComments(code)) && /let c = 3;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 11;
let b = 9;
let c = 3;

// Only change code below this line
a = a - 6;
b = b - 15;
c = c - 1;
```

# --solutions--

```js
let a = 11;
let b = 9;
let c = 3;

a -= 6;
b -= 15;
c -= 1;
```
