---
id: 56533eb9ac21ba0edf2244a8
title: 할당 연산자(Assignment Operator) 를 통해 값을 저장
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

In JavaScript, you can store a value in a variable with the <dfn>assignment</dfn> operator (`=`).

```js
myVariable = 5;
```

위의 예는 `Number` 값인 `5`를 `myVariable` 에 할당하고 있습니다.

만약 `=` 연산자의 오른쪽에 무언가 계산식이 있을 경우에는, 그 계산이 실행되고 나서, 왼쪽 변수에 값이 할당됩니다.

```js
var myVar;
myVar = 5;
```

우선, 이 코드는 `myVar`라는 이름의 변수를 작성하고 있습니다. 그리고 코드는 `5`를 `myVar`에 대입하고 있습니다. 이제 만약 `myVar`가 코드에서 다시 나오게 된다면, 프로그램은 그것을 `5`라고 판단할 것입니다.

# --instructions--

값 `7`을 변수 `a`에 할당해 주세요.

# --hints--

당신은 지정된 코멘트 위의 코드를 변경하면 안됩니다.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a`의 값은 7이어야 합니다.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --before-user-code--

```js
if (typeof a != 'undefined') {
  a = undefined;
}
```

## --after-user-code--

```js
(function(a){return "a = " + a;})(a);
```

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
