---
id: 56533eb9ac21ba0edf2244df
title: switch 문에서 동일한 옵션 여러 개 사용하기
challengeType: 1
forumTopicId: 18242
dashedName: multiple-identical-options-in-switch-statements
---

# --description--

만약 `break` 문이 `switch` 문의 `case`에서 생략되면, `case`를 만날 때까지 다음 `break` 문들이 계속 실행됩니다. 동일한 출력 값을 갖는 여러 입력이 있다면, `switch` 문에서 이렇게 나타낼 수 있습니다:

```js
let result = "";
switch (val) {
  case 1:
  case 2:
  case 3:
    result = "1, 2, or 3";
    break;
  case 4:
    result = "4 alone";
}
```

1, 2, 3에 대한 케이스는 모두 같은 결과를 만듭니다.

# --instructions--

다음 범위에 대해 `answer`를 설정하는 switch 문을 작성하세요:  
`1-3` - `Low`  
`4-6` - `Mid`  
`7-9` - `High`

**참고:** 범위 내 각 숫자마다 `case` 문이 필요합니다.

# --hints--

`sequentialSizes(1)`는 문자열 `Low`를 반환해야 합니다.

```js
assert(sequentialSizes(1) === 'Low');
```

`sequentialSizes(2)`는 문자열 `Low`를 반환해야 합니다.

```js
assert(sequentialSizes(2) === 'Low');
```

`sequentialSizes(3)`는 문자열 `Low`를 반환해야 합니다.

```js
assert(sequentialSizes(3) === 'Low');
```

`sequentialSizes(4)`는 문자열 `Mid`를 반환해야 합니다.

```js
assert(sequentialSizes(4) === 'Mid');
```

`sequentialSizes(5)`는 문자열 `Mid`를 반환해야 합니다.

```js
assert(sequentialSizes(5) === 'Mid');
```

`sequentialSizes(6)`는 문자열 `Mid`를 반환해야 합니다.

```js
assert(sequentialSizes(6) === 'Mid');
```

`sequentialSizes(7)`는 문자열 `High`를 반환해야 합니다.

```js
assert(sequentialSizes(7) === 'High');
```

`sequentialSizes(8)`는 문자열 `High`를 반환해야 합니다.

```js
assert(sequentialSizes(8) === 'High');
```

`sequentialSizes(9)`는 문자열 `High`를 반환해야 합니다.

```js
assert(sequentialSizes(9) === 'High');
```

`if`나 `else` 문을 사용하지 마세요.

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

아홉 개의 `case` 문이 있어야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/case/g).length === 9);
```

# --seed--

## --seed-contents--

```js
function sequentialSizes(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

sequentialSizes(1);
```

# --solutions--

```js
function sequentialSizes(val) {
  let answer = "";

  switch (val) {
    case 1:
    case 2:
    case 3:
      answer = "Low";
      break;
    case 4:
    case 5:
    case 6:
      answer = "Mid";
      break;
    case 7:
    case 8:
    case 9:
      answer = "High";
  }

  return answer;
}
```
