---
id: 56533eb9ac21ba0edf2244a9
title: 대입 연산자로 변수 초기화하기
challengeType: 1
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

변수를 선언하는 같은 줄에서 초기값으로 <dfn>초기화</dfn>하는 것이 일반적입니다.

```js
var myVar = 0;
```

`myVar`라는 새 변수를 만들고 `0`의 초기값을 할당합니다.

# --instructions--

`a` 변수를 `var`로 선언하고 `9` 값으로 초기화하세요.

# --hints--

`a`를 `9` 값으로 초기화해야 합니다.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
