---
id: 56533eb9ac21ba0edf2244ae
title: JavaScript에서 나머지 구하기
challengeType: 1
forumTopicId: 18184
dashedName: finding-a-remainder-in-javascript
---

# --description--

<dfn>나머지</dfn> 연산자 `%`는 두 수를 나눈 나머지를 구합니다.

**예시**

<pre>
5 % 2 = 1
5 / 2 = 2 나머지 1
2 * 2 = 4
5 - 4 = 1
</pre>

**사용법**  
수학에서 어떤 수가 짝수인지 홀수인지 확인하려면 그 수를 `2`로 나눈 나머지를 확인하면 됩니다. 짝수는 나머지가 `0`이고, 홀수는 나머지가 `1`입니다.

<pre>
17 % 2 = 1
48 % 2 = 0
</pre>

**참고:** <dfn>나머지</dfn> 연산자는 때때로 잘못 modulus 연산자라고 불리기도 합니다. modulus와 매우 비슷하지만 음수에 대해서는 제대로 작동하지 않습니다.

# --instructions--

`remainder`에 `11`를 `3`로 나눈 나머지를 <dfn>나머지</dfn> (`%`) 연산자를 사용해 할당하세요.

# --hints--

변수 `remainder`는 초기화되어야 합니다

```js
assert(/(const|let|var)\s+?remainder/.test(__helpers.removeJSComments(code)));
```

`remainder`의 값은 `2`이어야 합니다

```js
assert(remainder === 2);
```

`%` 연산자를 사용해야 합니다

```js
assert(/\s+?remainder\s*?=\s*?.*%.*;?/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const remainder = 0;
```

# --solutions--

```js
const remainder = 11 % 3;
```
