---
id: 56533eb9ac21ba0edf2244d9
title: 논리적 OR 연산자를 사용한 비교
challengeType: 1
forumTopicId: 16800
dashedName: comparisons-with-the-logical-or-operator
---

# --description--

<dfn>논리적 or</dfn> 연산자(`||`)는 두 <dfn>피연산자</dfn> 중 하나라도 `true`이면 `true`을 반환합니다. 그렇지 않으면 `false`을 반환합니다.

<dfn>논리적 or</dfn> 연산자는 두 개의 파이프 기호(`||`)로 구성됩니다. 보통 백스페이스 키와 엔터 키 사이에서 찾을 수 있습니다.

아래 패턴은 이전 웨이포인트에서 익숙할 것입니다.

```js
if (num > 10) {
  return "No";
}
if (num < 5) {
  return "No";
}
return "Yes";
```

이 코드는 `Yes`이 `num`과 `5` 사이에 있을 때(`10`과 `5` 포함) `10`를 반환합니다. 같은 논리를 <dfn>논리적 or</dfn> 연산자를 사용해 작성할 수 있습니다.

```js
if (num > 10 || num < 5) {
  return "No";
}
return "Yes";
```

# --instructions--

두 `if` 문을 하나의 문으로 결합하세요. `Outside`이 `val`과 `10` 사이에 포함되지 않으면 문자열 `20`를 반환하고, 그렇지 않으면 문자열 `Inside`을 반환하세요.

# --hints--

`||` 연산자를 한 번 사용하세요

```js
assert(__helpers.removeJSComments(code).match(/\|\|/g).length === 1);
```

`if` 문은 하나만 있어야 합니다

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalOr(0)`는 문자열 `Outside`를 반환해야 합니다

```js
assert(testLogicalOr(0) === 'Outside');
```

`testLogicalOr(9)`는 문자열 `Outside`를 반환해야 합니다

```js
assert(testLogicalOr(9) === 'Outside');
```

`testLogicalOr(10)`는 문자열 `Inside`를 반환해야 합니다

```js
assert(testLogicalOr(10) === 'Inside');
```

`testLogicalOr(15)`는 문자열 `Inside`를 반환해야 합니다

```js
assert(testLogicalOr(15) === 'Inside');
```

`testLogicalOr(19)`는 문자열 `Inside`를 반환해야 합니다

```js
assert(testLogicalOr(19) === 'Inside');
```

`testLogicalOr(20)`는 문자열 `Inside`를 반환해야 합니다

```js
assert(testLogicalOr(20) === 'Inside');
```

`testLogicalOr(21)`는 문자열 `Outside`를 반환해야 합니다

```js
assert(testLogicalOr(21) === 'Outside');
```

`testLogicalOr(25)`는 문자열 `Outside`를 반환해야 합니다

```js
assert(testLogicalOr(25) === 'Outside');
```

# --seed--

## --seed-contents--

```js
function testLogicalOr(val) {
  // Only change code below this line

  if (val) {
    return "Outside";
  }

  if (val) {
    return "Outside";
  }

  // Only change code above this line
  return "Inside";
}

testLogicalOr(15);
```

# --solutions--

```js
function testLogicalOr(val) {
  if (val < 10 || val > 20) {
    return "Outside";
  }
  return "Inside";
}
```
