---
id: 56533eb9ac21ba0edf2244ac
title: JavaScript로 숫자 1씩 증가시키기
challengeType: 1
forumTopicId: 18201
dashedName: increment-a-number-with-javascript
---

# --description--

`++` 연산자를 사용하면 변수에 1을 더하거나 쉽게 <dfn>증가</dfn>시킬 수 있습니다.

```js
i++;
```

는 다음과 같습니다

```js
i = i + 1;
```

**참고:** 등호가 필요 없도록 전체 줄이 `i++;`가 됩니다.

# --instructions--

`++`에 `myVar` 연산자를 사용하도록 코드를 변경하세요.

# --hints--

`myVar`는 `88`와 같아야 합니다.

```js
assert(myVar === 88);
```

대입 연산자를 사용하면 안 됩니다.

```js
assert(
  /let\s+myVar\s*=\s*87;\s*\/*.*\s*([+]{2}\s*myVar|myVar\s*[+]{2})/.test(__helpers.removeJSComments(code))
);
```

`++` 연산자를 사용해야 합니다.

```js
assert(/[+]{2}\s*myVar|myVar\s*[+]{2}/.test(__helpers.removeJSComments(code)));
```

지정된 주석 위의 코드는 변경하지 마세요.

```js
assert(/let myVar = 87;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 87;

// Only change code below this line
myVar = myVar + 1;
```

# --solutions--

```js
let myVar = 87;
myVar++;
```
