---
id: 56533eb9ac21ba0edf2244ca
title: 조회용 객체 사용하기
challengeType: 1
forumTopicId: 18373
dashedName: using-objects-for-lookups
---

# --description--

객체는 사전처럼 키/값 저장소로 생각할 수 있습니다. 표 형식 데이터가 있다면 `switch` 문이나 `if/else` 체인 대신 객체를 사용해 값을 조회할 수 있습니다. 입력 데이터가 특정 범위로 제한된 경우에 가장 유용합니다.

다음은 article 객체의 예입니다:

```js
const article = {
  "title": "How to create objects in JavaScript",
  "link": "https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/",
  "author": "Kaashan Hussain",
  "language": "JavaScript",
  "tags": "TECHNOLOGY",
  "createdAt": "NOVEMBER 28, 2018"
};

const articleAuthor = article["author"];
const articleLink = article["link"];

const value = "title";
const valueLookup = article[value];
```

`articleAuthor`는 문자열 `Kaashan Hussain`, `articleLink`는 문자열 `https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/`, `valueLookup`는 문자열 `How to create objects in JavaScript`입니다.

# --instructions--

switch 문을 `lookup`라는 객체로 변환하세요. 이를 사용해 `val`를 조회하고 연결된 문자열을 `result` 변수에 할당하세요.

# --hints--

`phoneticLookup("alpha")`는 문자열 `Adams`와 같아야 합니다

```js
assert(phoneticLookup('alpha') === 'Adams');
```

`phoneticLookup("bravo")`는 문자열 `Boston`와 같아야 합니다

```js
assert(phoneticLookup('bravo') === 'Boston');
```

`phoneticLookup("charlie")`는 문자열 `Chicago`와 같아야 합니다

```js
assert(phoneticLookup('charlie') === 'Chicago');
```

`phoneticLookup("delta")`는 문자열 `Denver`와 같아야 합니다

```js
assert(phoneticLookup('delta') === 'Denver');
```

`phoneticLookup("echo")`는 문자열 `Easy`와 같아야 합니다

```js
assert(phoneticLookup('echo') === 'Easy');
```

`phoneticLookup("foxtrot")`는 문자열 `Frank`와 같아야 합니다

```js
assert(phoneticLookup('foxtrot') === 'Frank');
```

`phoneticLookup("")`는 `undefined`와 같아야 합니다

```js
assert(typeof phoneticLookup('') === 'undefined');
```

`return` 문은 수정하지 마세요

```js
assert(__helpers.removeJSComments(code).match(/return\sresult;/));
```

`case`, `switch`, `if` 문을 사용하지 마세요

```js
assert(
  !/case|switch|if/g.test(__helpers.removeJSComments(code).replace(/([/]{2}.*)|([/][*][^/*]*[*][/])/g, ''))
);
```

# --seed--

## --seed-contents--

```js
// Setup
function phoneticLookup(val) {
  let result = "";

  // Only change code below this line
  switch(val) {
    case "alpha":
      result = "Adams";
      break;
    case "bravo":
      result = "Boston";
      break;
    case "charlie":
      result = "Chicago";
      break;
    case "delta":
      result = "Denver";
      break;
    case "echo":
      result = "Easy";
      break;
    case "foxtrot":
      result = "Frank";
  }

  // Only change code above this line
  return result;
}

phoneticLookup("charlie");
```

# --solutions--

```js
function phoneticLookup(val) {
  let result = "";

  const lookup = {
    alpha: "Adams",
    bravo: "Boston",
    charlie: "Chicago",
    delta: "Denver",
    echo: "Easy",
    foxtrot: "Frank"
  };

  result = lookup[val];

  return result;
}
```
