---
id: 56533eb9ac21ba0edf2244d3
title: 엄격한 부등 연산자와 비교
challengeType: 1
forumTopicId: 16791
dashedName: comparison-with-the-strict-inequality-operator
---

# --description--

엄격한 부등 연산자(`!==`)는 엄격한 동등 연산자의 논리적 반대입니다. 이는 "엄격하게 같지 않음"을 의미하며, 엄격한 동등 연산자가 `false`를 반환할 때 `true`을 반환하고 *그 반대도 마찬가지*입니다. 엄격한 부등 연산자는 데이터 유형을 변환하지 않습니다.

**예시**

```js
3 !==  3  // false
3 !== '3' // true
4 !==  3  // true
```

# --instructions--

`if` 문장에 엄격한 부등 연산자를 추가하여 `Not Equal`가 `val`와 엄격하게 같지 않을 때 함수가 문자열 `17`를 반환하도록 하세요.

# --hints--

`testStrictNotEqual(17)`는 문자열 `Equal`를 반환해야 합니다.

```js
assert(testStrictNotEqual(17) === 'Equal');
```

`testStrictNotEqual("17")`는 문자열 `Not Equal`를 반환해야 합니다.

```js
assert(testStrictNotEqual('17') === 'Not Equal');
```

`testStrictNotEqual(12)`는 문자열 `Not Equal`를 반환해야 합니다.

```js
assert(testStrictNotEqual(12) === 'Not Equal');
```

`testStrictNotEqual("bob")`는 문자열 `Not Equal`를 반환해야 합니다.

```js
assert(testStrictNotEqual('bob') === 'Not Equal');
```

`!==` 연산자를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/(val\s*!==\s*\d+)|(\d+\s*!==\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrictNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testStrictNotEqual(10);
```

# --solutions--

```js
function testStrictNotEqual(val) {
  if (val !== 17) {
    return "Not Equal";
  }
  return "Equal";
}
```
