---
id: 56533eb9ac21ba0edf2244ed
title: 문자열에 변수 추가하기
challengeType: 1
forumTopicId: 16656
dashedName: appending-variables-to-strings
---

# --description--

문자열 <dfn>리터럴</dfn>을 여러 줄에 걸쳐 만들 수 있는 것처럼, 더하기 등호(`+=`) 연산자를 사용해 변수도 문자열에 추가할 수 있습니다.

예시:

```js
const anAdjective = "awesome!";
let ourStr = "freeCodeCamp is ";
ourStr += anAdjective;
```

`ourStr`은 `freeCodeCamp is awesome!` 값을 갖게 됩니다.

# --instructions--

`someAdjective`에 최소 3자 이상의 문자열을 할당하고 `myStr`에 `+=` 연산자를 사용해 추가하세요.

# --hints--

`someAdjective`에는 최소 3자 이상의 문자열을 할당해야 합니다.

```js
assert(typeof someAdjective !== 'undefined' && someAdjective.length > 2);
```

`someAdjective`을 `myStr`에 `+=` 연산자를 사용해 추가해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*someAdjective\s*/).length > 0);
```

# --seed--

## --seed-contents--

```js
// Change code below this line
const someAdjective = "";
let myStr = "Learning to code is ";
```

# --solutions--

```js
const someAdjective = "neat";
let myStr = "Learning to code is ";
myStr += someAdjective;
```
