---
id: 56533eb9ac21ba0edf2244a8
title: 대입 연산자로 값 저장하기
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

JavaScript에서는 <dfn>대입</dfn> 연산자(`=`)를 사용해 변수에 값을 저장할 수 있습니다.

```js
myVariable = 5;
```

이 코드는 `Number` 값을 `5`에 `myVariable` 할당합니다.

만약 `=` 연산자 오른쪽에 계산이 있다면, 그 계산이 먼저 수행된 후 왼쪽 변수에 값이 할당됩니다.

```js
var myVar;
myVar = 5;
```

먼저, 이 코드는 `myVar`라는 변수를 만듭니다. 그다음, `5`를 `myVar`에 할당합니다. 이제 코드에서 `myVar`가 다시 나오면, 프로그램은 그것을 `5`로 처리합니다.

# --instructions--

변수 `7`에 값 `a`를 할당하세요.

# --hints--

지정된 주석 위의 코드는 변경하지 마세요.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a`는 값이 7이어야 합니다.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
