---
id: 56533eb9ac21ba0edf2244c3
title: 반환값을 사용한 할당
challengeType: 1
forumTopicId: 16658
dashedName: assignment-with-a-returned-value
---

# --description--

<a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/storing-values-with-the-assignment-operator" target="_blank" rel="noopener noreferrer nofollow">대입 연산자를 사용한 값 저장</a>에 대해 이야기했던 내용을 기억한다면, 등호 오른쪽에 있는 모든 것은 값이 할당되기 전에 먼저 평가된다는 것을 알 수 있습니다. 이는 함수의 반환값을 변수에 할당할 수 있다는 뜻입니다.

두 수를 더하는 `sum` 함수를 정의했다고 가정해 보겠습니다.

```js
ourSum = sum(5, 12);
```

`sum`와 `5` 인수를 사용해 `12` 함수를 호출하면 `17`라는 반환값이 생성됩니다. 이 반환값은 `ourSum` 변수에 할당됩니다.

# --instructions--

`processArg` 함수를 `7` 인수로 호출하고 그 반환값을 `processed` 변수에 할당하세요.

# --hints--

`processed`는 `2` 값을 가져야 합니다.

```js
assert(processed === 2);
```

`processArg`을 `processed`에 할당해야 합니다.

```js
assert(/processed\s*=\s*processArg\(\s*7\s*\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

// Only change code below this line

```

# --solutions--

```js
var processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

processed = processArg(7);
```
