---
id: 56533eb9ac21ba0edf2244de
title: switch 문에 기본 옵션 추가하기
challengeType: 1
forumTopicId: 16653
dashedName: adding-a-default-option-in-switch-statements
---

# --description--

`switch` 문에서는 모든 가능한 값을 `case` 문으로 지정하지 못할 수 있습니다. 대신, 일치하는 `case` 문이 없을 때 실행되는 `default` 문을 추가할 수 있습니다. 이것은 `if/else` 체인에서 마지막 `else` 문과 같다고 생각하세요.

`default` 문은 마지막 `case` 문이어야 합니다.

```js
switch (num) {
  case value1:
    statement1;
    break;
  case value2:
    statement2;
    break;
...
  default:
    defaultStatement;
    break;
}
```

# --instructions--

다음 조건에 대해 `answer` 값을 설정하는 switch 문을 작성하세요:  
`a` - `apple`  
`b` - `bird`  
`c` - `cat`  
`default` - `stuff`

# --hints--

`switchOfStuff("a")` 는 문자열 `apple` 을 반환해야 합니다

```js
assert(switchOfStuff('a') === 'apple');
```

`switchOfStuff("b")` 는 문자열 `bird` 을 반환해야 합니다

```js
assert(switchOfStuff('b') === 'bird');
```

`switchOfStuff("c")` 는 문자열 `cat` 을 반환해야 합니다

```js
assert(switchOfStuff('c') === 'cat');
```

`switchOfStuff("d")` 는 문자열 `stuff` 을 반환해야 합니다

```js
assert(switchOfStuff('d') === 'stuff');
```

`switchOfStuff(4)` 는 문자열 `stuff` 을 반환해야 합니다

```js
assert(switchOfStuff(4) === 'stuff');
```

`if` 나 `else` 문을 사용하면 안 됩니다

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

`default` 문을 사용해야 합니다

```js
assert(switchOfStuff('string-to-trigger-default-case') === 'stuff');
```

최소 3개의 `break` 문이 있어야 합니다

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function switchOfStuff(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

switchOfStuff(1);
```

# --solutions--

```js
function switchOfStuff(val) {
  let answer = "";

  switch(val) {
    case "a":
      answer = "apple";
      break;
    case "b":
      answer = "bird";
      break;
    case "c":
      answer = "cat";
      break;
    default:
      answer = "stuff";
  }
  return answer;
}
```
