---
id: 56533eb9ac21ba0edf2244e0
title: if else 체인을 switch로 교체하기
challengeType: 1
forumTopicId: 18266
dashedName: replacing-if-else-chains-with-switch
---

# --description--

선택할 옵션이 많을 때, 여러 개의 연결된 `switch`/`if` 문장보다 `else if` 문장이 작성하기 더 쉬울 수 있습니다. 다음 코드는:

```js
if (val === 1) {
  answer = "a";
} else if (val === 2) {
  answer = "b";
} else {
  answer = "c";
}
```

다음 코드로 교체할 수 있습니다:

```js
switch (val) {
  case 1:
    answer = "a";
    break;
  case 2:
    answer = "b";
    break;
  default:
    answer = "c";
}
```

# --instructions--

연결된 `if`/`else if` 문장을 `switch` 문장으로 바꾸세요.

# --hints--

에디터 어디에도 `else` 문장을 사용하지 마세요

```js
assert(!/else/g.test(__helpers.removeJSComments(code)));
```

에디터 어디에도 `if` 문장을 사용하지 마세요

```js
assert(!/if/g.test(__helpers.removeJSComments(code)));
```

최소 네 개의 `break` 문장이 있어야 합니다

```js
assert(__helpers.removeJSComments(code).match(/break/g).length >= 4);
```

`chainToSwitch("bob")`는 문자열 `Marley`을 반환해야 합니다

```js
assert(chainToSwitch('bob') === 'Marley');
```

`chainToSwitch(42)`는 문자열 `The Answer`을 반환해야 합니다

```js
assert(chainToSwitch(42) === 'The Answer');
```

`chainToSwitch(1)`는 문자열 `There is no #1`을 반환해야 합니다

```js
assert(chainToSwitch(1) === 'There is no #1');
```

`chainToSwitch(99)`는 문자열 `Missed me by this much!`을 반환해야 합니다

```js
assert(chainToSwitch(99) === 'Missed me by this much!');
```

`chainToSwitch(7)`는 문자열 `Ate Nine`을 반환해야 합니다

```js
assert(chainToSwitch(7) === 'Ate Nine');
```

`chainToSwitch("John")`는 `""` (빈 문자열)을 반환해야 합니다

```js
assert(chainToSwitch('John') === '');
```

`chainToSwitch(156)`는 `""` (빈 문자열)을 반환해야 합니다

```js
assert(chainToSwitch(156) === '');
```

# --seed--

## --seed-contents--

```js
function chainToSwitch(val) {
  let answer = "";
  // Only change code below this line

  if (val === "bob") {
    answer = "Marley";
  } else if (val === 42) {
    answer = "The Answer";
  } else if (val === 1) {
    answer = "There is no #1";
  } else if (val === 99) {
    answer = "Missed me by this much!";
  } else if (val === 7) {
    answer = "Ate Nine";
  }

  // Only change code above this line
  return answer;
}

chainToSwitch(7);
```

# --solutions--

```js
function chainToSwitch(val) {
  let answer = "";

  switch (val) {
    case "bob":
      answer = "Marley";
      break;
    case 42:
      answer = "The Answer";
      break;
    case 1:
      answer = "There is no #1";
      break;
    case 99:
      answer = "Missed me by this much!";
      break;
    case 7:
      answer = "Ate Nine";
  }
  return answer;
}
```
