---
id: 56533eb9ac21ba0edf2244b5
title: 문자열에서 리터럴 따옴표 이스케이프하기
challengeType: 1
forumTopicId: 17568
dashedName: escaping-literal-quotes-in-strings
---

# --description--

문자열을 정의할 때는 반드시 작은따옴표나 큰따옴표로 시작하고 끝내야 합니다. 문자열 안에 `"` 또는 `'` 같은 리터럴 따옴표가 필요하면 어떻게 해야 할까요?

JavaScript에서는 따옴표 앞에 <dfn>역슬래시</dfn>(`\`)를 넣어 따옴표가 문자열의 끝으로 인식되지 않도록 <dfn>이스케이프</dfn>할 수 있습니다.

```js
const sampleStr = "Alan said, \"Peter is learning JavaScript\".";
```

이렇게 하면 JavaScript는 다음 따옴표가 문자열의 끝이 아니라 문자열 안에 포함되어야 한다고 인식합니다. 따라서 이 문자열을 콘솔에 출력하면 다음과 같이 나옵니다:

```js
Alan said, "Peter is learning JavaScript".
```

# --instructions--

<dfn>역슬래시</dfn>를 사용해 `myStr` 변수에 문자열을 할당하세요. 콘솔에 출력하면 다음과 같이 보여야 합니다:

```js
I am a "double quoted" string inside "double quotes".
```

# --hints--

큰따옴표(`"`) 두 개와 이스케이프된 큰따옴표(`\"`) 네 개를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\\"/g).length === 4 && __helpers.removeJSComments(code).match(/[^\\]"/g).length === 2);
```

변수 `myStr`에는 문자열 `I am a "double quoted" string inside "double quotes".`이 들어 있어야 합니다.

```js
assert(/I am a "double quoted" string inside "double quotes(\."|"\.)$/.test(myStr));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "I am a \"double quoted\" string inside \"double quotes\".";
```
