---
id: 56533eb9ac21ba0edf2244da
title: Else 문 소개
challengeType: 1
forumTopicId: 18207
dashedName: introducing-else-statements
---

# --description--

`if` 문 조건이 참일 때, 그 뒤에 오는 코드 블록이 실행됩니다. 그렇다면 그 조건이 거짓일 때는 어떻게 될까요? 보통은 아무 일도 일어나지 않습니다. `else` 문을 사용하면 대체 코드 블록을 실행할 수 있습니다.

```js
if (num > 10) {
  return "Bigger than 10";
} else {
  return "10 or Less";
}
```

# --instructions--

`if` 문들을 하나의 `if/else` 문으로 결합하세요.

# --hints--

편집기에는 `if` 문이 하나만 있어야 합니다

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`else` 문을 사용해야 합니다

```js
assert(/else/g.test(__helpers.removeJSComments(code)));
```

`testElse(4)` 는 문자열 `5 or Smaller` 을 반환해야 합니다

```js
assert(testElse(4) === '5 or Smaller');
```

`testElse(5)` 는 문자열 `5 or Smaller` 을 반환해야 합니다

```js
assert(testElse(5) === '5 or Smaller');
```

`testElse(6)` 은 문자열 `Bigger than 5` 을 반환해야 합니다

```js
assert(testElse(6) === 'Bigger than 5');
```

`testElse(10)` 는 문자열 `Bigger than 5` 을 반환해야 합니다

```js
assert(testElse(10) === 'Bigger than 5');
```

지정된 주석 위나 아래의 코드는 변경하지 마세요.

```js
assert(/let result = "";/.test(__helpers.removeJSComments(code)) && /return result;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function testElse(val) {
  let result = "";
  // Only change code below this line

  if (val > 5) {
    result = "Bigger than 5";
  }

  if (val <= 5) {
    result = "5 or Smaller";
  }

  // Only change code above this line
  return result;
}

testElse(4);
```

# --solutions--

```js
function testElse(val) {
  let result = "";
  if(val > 5) {
    result = "Bigger than 5";
  } else {
    result = "5 or Smaller";
  }
  return result;
}
```
