---
id: 56533eb9ac21ba0edf2244dc
title: if else 문 체이닝
challengeType: 1
forumTopicId: 16772
dashedName: chaining-if-else-statements
---

# --description--

`if/else` 문은 복잡한 논리를 위해 체이닝할 수 있습니다. 다음은 여러 개의 체이닝된 `if` / `else if` 문의 <dfn>의사 코드</dfn>입니다:

```js
if (condition1) {
  statement1
} else if (condition2) {
  statement2
} else if (condition3) {
  statement3
. . .
} else {
  statementN
}
```

# --instructions--

다음 조건을 충족하는 체이닝된 `if`/`else if` 문을 작성하세요:

`num < 5` - `Tiny` 반환  
`num < 10` - `Small` 반환  
`num < 15` - `Medium` 반환  
`num < 20` - `Large` 반환  
`num >= 20` - `Huge` 반환

# --hints--

최소 네 개의 `else` 문이 있어야 합니다

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 3);
```

최소 네 개의 `if` 문이 있어야 합니다

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 3);
```

최소 한 개의 `return` 문이 있어야 합니다

```js
assert(__helpers.removeJSComments(code).match(/return/g).length >= 1);
```

`testSize(0)` 은 문자열 `Tiny` 를 반환해야 합니다

```js
assert(testSize(0) === 'Tiny');
```

`testSize(4)` 은 문자열 `Tiny` 를 반환해야 합니다

```js
assert(testSize(4) === 'Tiny');
```

`testSize(5)` 은 문자열 `Small` 를 반환해야 합니다

```js
assert(testSize(5) === 'Small');
```

`testSize(8)` 은 문자열 `Small` 를 반환해야 합니다

```js
assert(testSize(8) === 'Small');
```

`testSize(10)` 은 문자열 `Medium` 를 반환해야 합니다

```js
assert(testSize(10) === 'Medium');
```

`testSize(14)` 은 문자열 `Medium` 를 반환해야 합니다

```js
assert(testSize(14) === 'Medium');
```

`testSize(15)` 은 문자열 `Large` 를 반환해야 합니다

```js
assert(testSize(15) === 'Large');
```

`testSize(17)` 은 문자열 `Large` 를 반환해야 합니다

```js
assert(testSize(17) === 'Large');
```

`testSize(20)` 은 문자열 `Huge` 를 반환해야 합니다

```js
assert(testSize(20) === 'Huge');
```

`testSize(25)` 은 문자열 `Huge` 를 반환해야 합니다

```js
assert(testSize(25) === 'Huge');
```

# --seed--

## --seed-contents--

```js
function testSize(num) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

testSize(7);
```

# --solutions--

```js
function testSize(num) {
  if (num < 5) {
    return "Tiny";
  } else if (num < 10) {
    return "Small";
  } else if (num < 15) {
    return "Medium";
  } else if (num < 20) {
    return "Large";
  } else {
    return "Huge";
  }
}
```
