---
id: 56592a60ddddeae28f7aa8e1
title: 인덱스로 다차원 배열 접근하기
challengeType: 1
forumTopicId: 16159
dashedName: access-multi-dimensional-arrays-with-indexes
---

# --description--

<dfn>다차원</dfn> 배열을 생각하는 한 가지 방법은 *배열의 배열*로 보는 것입니다. 배열에 접근할 때 괄호를 사용하면, 첫 번째 괄호 세트는 가장 바깥쪽(첫 번째 수준) 배열의 항목을 가리키고, 추가 괄호 쌍은 내부의 다음 수준 항목을 가리킵니다.

**예시**

```js
const arr = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14]
];

const subarray = arr[3];
const nestedSubarray = arr[3][0];
const element = arr[3][0][1];
```

이 예시에서 `subarray`는 `[[10, 11, 12], 13, 14]` 값을 가지고, `nestedSubarray`는 `[10, 11, 12]` 값을 가지며, `element`는 `11` 값을 가집니다.

**참고:** 배열 이름과 대괄호 사이에 공백이 있으면 안 됩니다. 예를 들어 `array [0][0]`와 심지어 이 `array [0] [0]`도 허용되지 않습니다. JavaScript는 이를 올바르게 처리할 수 있지만, 다른 프로그래머가 코드를 읽을 때 혼란을 줄 수 있습니다.

# --instructions--

대괄호 표기법을 사용해 `myArray`에서 `myData`가 `8`과 같도록 요소를 선택하세요.

# --hints--

`myData`는 `8`와 같아야 합니다.

```js
assert(myData === 8);
```

대괄호 표기법을 사용해 `myArray`에서 올바른 값을 읽어야 합니다.

```js
assert(/myData=myArray\[2\]\[1\]/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --seed-contents--

```js
const myArray = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14],
];

const myData = myArray[0][0];
```

# --solutions--

```js
const myArray = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [[10, 11, 12], 13, 14]];
const myData = myArray[2][1];
```
