---
id: 565bbe00e9cc8ac0725390f4
title: 카드 세기
challengeType: 1
forumTopicId: 16809
dashedName: counting-cards
---

# --description--

카지노 게임 블랙잭에서 플레이어는 덱에 남아 있는 높은 카드와 낮은 카드의 상대적인 수를 추적하여 다음 판에서 하우스보다 유리한지 판단할 수 있습니다. 이것을 카드 세기라고 합니다.

덱에 높은 카드가 더 많이 남아 있으면 플레이어에게 유리합니다. 각 카드에는 아래 표에 따라 값이 할당됩니다. 카운트가 양수일 때는 플레이어가 높은 베팅을 해야 합니다. 카운트가 0이거나 음수일 때는 낮은 베팅을 해야 합니다.

<table><thead><tr><th>카운트 변화</th><th>카드</th></tr></thead><tbody><tr><td>+1</td><td>2, 3, 4, 5, 6</td></tr><tr><td>0</td><td>7, 8, 9</td></tr><tr><td>-1</td><td>10, 'J', 'Q', 'K', 'A'</td></tr></tbody></table>

카드 세기 함수를 작성하세요. 이 함수는 숫자 또는 문자열일 수 있는 `card` 매개변수를 받고, 카드 값에 따라 전역 `count` 변수를 증가시키거나 감소시킵니다(표 참조). 그런 다음 현재 카운트와 카운트가 양수일 때는 `Bet` 문자열, 0이거나 음수일 때는 `Hold` 문자열을 포함하는 문자열을 반환해야 합니다. 현재 카운트와 플레이어의 결정(`Bet` 또는 `Hold`)은 한 칸의 공백으로 구분하세요.

**예시 출력:** `-3 Hold` 또는 `5 Bet`

**힌트**  
값이 7, 8, 9일 때 `count`을 0으로 초기화하지 마세요.  
배열을 반환하지 마세요.  
출력에 따옴표(작은따옴표 또는 큰따옴표)를 포함하지 마세요.

# --hints--

함수는 카운트 값과 텍스트(`Bet` 또는 `Hold`)를 한 칸의 공백으로 구분하여 반환해야 합니다.

```js
assert(//
  (function () {
    count = 0;
    let out = cc(10);
    const hasSpace = /-?\d+ (Bet|Hold)/.test('' + out);
    return hasSpace;
  })()
);
```

카드 수열 2, 3, 4, 5, 6은 문자열 `5 Bet`을 반환해야 합니다.

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(3);
    cc(4);
    cc(5);
    var out = cc(6);
    if (out === '5 Bet') {
      return true;
    }
    return false;
  })()
);
```

카드 수열 7, 8, 9는 문자열 `0 Hold`를 반환해야 합니다.

```js
assert(
  (function () {
    count = 0;
    cc(7);
    cc(8);
    var out = cc(9);
    if (out === '0 Hold') {
      return true;
    }
    return false;
  })()
);
```

카드 수열 10, J, Q, K, A는 문자열 `-5 Hold`을 반환해야 합니다.

```js
assert(
  (function () {
    count = 0;
    cc(10);
    cc('J');
    cc('Q');
    cc('K');
    var out = cc('A');
    if (out === '-5 Hold') {
      return true;
    }
    return false;
  })()
);
```

카드 수열 3, 7, Q, 8, A는 문자열 `-1 Hold`를 반환해야 합니다.

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(7);
    cc('Q');
    cc(8);
    var out = cc('A');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

카드 수열 2, J, 9, 2, 7은 문자열 `1 Bet`를 반환해야 합니다.

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc('J');
    cc(9);
    cc(2);
    var out = cc(7);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

카드 수열 2, 2, 10은 문자열 `1 Bet`를 반환해야 합니다.

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(2);
    var out = cc(10);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

카드 수열 3, 2, A, 10, K는 문자열 `-1 Hold`를 반환해야 합니다.

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(2);
    cc('A');
    cc(10);
    var out = cc('K');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

# --seed--

## --seed-contents--

```js
let count = 0;

function cc(card) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

cc(2); cc(3); cc(7); cc('K'); cc('A');
```

# --solutions--

```js
let count = 0;
function cc(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
