---
id: 5664820f61c48e80c9fa476c
title: 골프 점수 코드
challengeType: 1
forumTopicId: 18195
dashedName: golf-code
---

# --description--

골프 게임에서 각 홀에는 `par`가 있습니다. 이는 골퍼가 공을 홀에 넣어 플레이를 완료하기 위해 예상되는 평균 `strokes` 수를 의미합니다. `par`보다 `strokes`가 얼마나 높거나 낮은지에 따라 다른 별명이 붙습니다.

함수는 `par`와 `strokes` 인수를 받습니다. 아래 표에 따라 우선순위가 높은 순서대로 나열된 획 수에 맞는 올바른 문자열을 반환하세요:

<table><thead><tr><th>획 수</th><th>반환값</th></tr></thead><tbody><tr><td>1</td><td>"Hole-in-one!"</td></tr><tr><td>&#x3C;= par - 2</td><td>"Eagle"</td></tr><tr><td>par - 1</td><td>"Birdie"</td></tr><tr><td>par</td><td>"Par"</td></tr><tr><td>par + 1</td><td>"Bogey"</td></tr><tr><td>par + 2</td><td>"Double Bogey"</td></tr><tr><td>>= par + 3</td><td>"Go Home!"</td></tr></tbody></table>

`par`와 `strokes`는 항상 숫자이며 양수입니다. 편의를 위해 모든 이름을 배열로 추가했습니다.

# --hints--

`golfScore(4, 1)`는 문자열 `Hole-in-one!`를 반환해야 합니다

```js
assert(golfScore(4, 1) === 'Hole-in-one!');
```

`golfScore(4, 2)`는 문자열 `Eagle`를 반환해야 합니다

```js
assert(golfScore(4, 2) === 'Eagle');
```

`golfScore(5, 2)`는 문자열 `Eagle`를 반환해야 합니다

```js
assert(golfScore(5, 2) === 'Eagle');
```

`golfScore(4, 3)`는 문자열 `Birdie`를 반환해야 합니다

```js
assert(golfScore(4, 3) === 'Birdie');
```

`golfScore(4, 4)`는 문자열 `Par`를 반환해야 합니다

```js
assert(golfScore(4, 4) === 'Par');
```

`golfScore(1, 1)`는 문자열 `Hole-in-one!`를 반환해야 합니다

```js
assert(golfScore(1, 1) === 'Hole-in-one!');
```

`golfScore(5, 5)`는 문자열 `Par`를 반환해야 합니다

```js
assert(golfScore(5, 5) === 'Par');
```

`golfScore(4, 5)`는 문자열 `Bogey`를 반환해야 합니다

```js
assert(golfScore(4, 5) === 'Bogey');
```

`golfScore(4, 6)`는 문자열 `Double Bogey`를 반환해야 합니다

```js
assert(golfScore(4, 6) === 'Double Bogey');
```

`golfScore(4, 7)`는 문자열 `Go Home!`를 반환해야 합니다

```js
assert(golfScore(4, 7) === 'Go Home!');
```

`golfScore(5, 9)`는 문자열 `Go Home!`를 반환해야 합니다

```js
assert(golfScore(5, 9) === 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

function golfScore(par, strokes) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

golfScore(5, 4);
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
