---
id: 5675e877dbd60be8ad28edc6
title: 배열을 For 루프로 순차적으로 반복하기
challengeType: 1
forumTopicId: 18216
dashedName: iterate-through-an-array-with-a-for-loop
---

# --description--

JavaScript에서 흔히 하는 작업 중 하나는 배열의 내용을 순차적으로 반복하는 것입니다. 그 방법 중 하나가 `for` 루프를 사용하는 것입니다. 이 코드는 배열 `arr`의 각 요소를 콘솔에 출력합니다:

```js
const arr = [10, 9, 8, 7, 6];

for (let i = 0; i < arr.length; i++) {
   console.log(arr[i]);
}
```

배열은 0부터 시작하는 인덱스를 가진다는 점을 기억하세요. 즉, 배열의 마지막 인덱스는 `length - 1`입니다. 이 루프의 조건은 `i < arr.length`이며, `i`가 `length`와 같아지면 루프가 정지합니다. 이 경우 마지막 반복은 `i === 4` 즉, `i`이 `arr.length - 1`와 같아지고 `6`를 콘솔에 출력할 때입니다. 그 후 `i`이 `5`로 증가하고, `i < arr.length`가 `false`이 되어 루프가 종료됩니다.

# --instructions--

변수 `total`를 선언하고 `0`로 초기화하세요. `for` 루프를 사용해 `myArr` 배열의 각 요소 값을 `total`에 더하세요.

# --hints--

`total`는 선언하고 0으로 초기화해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s*?total\s*=\s*0.*?;?/));
```

`total`는 20과 같아야 합니다.

```js
assert(total === 20);
```

`for`를 순차적으로 반복하려면 `myArr` 루프를 사용해야 합니다.

```js
assert(/for\s*\(/g.test(__helpers.removeJSComments(code)) && /myArr\s*\[/g.test(__helpers.removeJSComments(code)));
```

`total`에 직접 20을 할당하려고 시도하지 마세요.

```js
assert(!__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/total[=+-]0*[1-9]+/gm));
```

# --seed--

## --seed-contents--

```js
// Setup
const myArr = [2, 3, 4, 5, 6];

// Only change code below this line

```

# --solutions--

```js
const myArr = [2, 3, 4, 5, 6];
let total = 0;

for (let i = 0; i < myArr.length; i++) {
  total += myArr[i];
}
```
