---
id: 5679ceb97cbaa8c51670a16b
title: 함수에서 불리언 값을 반환하기
challengeType: 1
forumTopicId: 18273
dashedName: returning-boolean-values-from-functions
---

# --description--

<a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/comparison-with-the-equality-operator" target="_blank" rel="noopener noreferrer nofollow">동등 연산자와 비교</a>에서 모든 비교 연산자가 불리언 `true` 또는 `false` 값을 반환한다는 것을 기억할 수 있습니다.

때때로 사람들은 다음과 같이 비교를 위해 `if/else` 문을 사용합니다:

```js
function isEqual(a, b) {
  if (a === b) {
    return true;
  } else {
    return false;
  }
}
```

하지만 더 나은 방법이 있습니다. `===`이 `true` 또는 `false`를 반환하므로, 비교 결과를 바로 반환할 수 있습니다:

```js
function isEqual(a, b) {
  return a === b;
}
```

# --instructions--

`isLess` 함수에서 `if/else` 문을 제거하세요.

# --hints--

`isLess(10, 15)`은 `true`를 반환해야 합니다.

```js
assert(isLess(10, 15) === true);
```

`isLess(15, 10)`은 `false`를 반환해야 합니다.

```js
assert(isLess(15, 10) === false);
```

어떤 `if`나 `else` 문도 사용하지 마세요.

```js
assert(!/if|else/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function isLess(a, b) {
  // Only change code below this line
  if (a < b) {
    return true;
  } else {
    return false;
  }
  // Only change code above this line
}

isLess(10, 15);
```

# --solutions--

```js
function isLess(a, b) {
  return a < b;
}
```
