---
id: 567af2437cbaa8c51670a16c
title: 속성 존재 여부 확인하기
challengeType: 1
forumTopicId: 18324
dashedName: testing-objects-for-properties
---

# --description--

주어진 객체에 특정 속성이 존재하는지 확인하려면 `.hasOwnProperty()` 메서드를 사용할 수 있습니다. `someObject.hasOwnProperty(someProperty)`은 속성이 객체에 있으면 `true`를, 없으면 `false`를 반환합니다.

**예시**

```js
function checkForProperty(object, property) {
  return object.hasOwnProperty(property);
}

checkForProperty({ top: 'hat', bottom: 'pants' }, 'top'); // true
checkForProperty({ top: 'hat', bottom: 'pants' }, 'middle'); // false
```

첫 번째 `checkForProperty` 함수 호출은 `true`를 반환하고, 두 번째 호출은 `false`를 반환합니다.

# --instructions--

함수 `checkObj`을 수정하여, 함수 매개변수 `obj`로 전달된 객체에 함수 매개변수 `checkProp`로 전달된 특정 속성이 있는지 테스트하세요. `checkProp`로 전달된 속성이 `obj`에 있으면 그 속성의 값을 반환하세요. 그렇지 않으면 `Not Found`를 반환하세요.

# --hints--

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "gift")`은 문자열 `pony`를 반환해야 합니다.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'gift') === 'pony'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "pet")`은 문자열 `kitten`를 반환해야 합니다.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'pet') === 'kitten'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "house")`은 문자열 `Not Found`를 반환해야 합니다.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'house') ===
    'Not Found'
);
```

`checkObj({city: "Seattle"}, "city")`은 문자열 `Seattle`를 반환해야 합니다.

```js
assert(checkObj({ city: 'Seattle' }, 'city') === 'Seattle');
```

`checkObj({city: "Seattle"}, "district")`은 문자열 `Not Found`를 반환해야 합니다.

```js
assert(checkObj({ city: 'Seattle' }, 'district') === 'Not Found');
```

`checkObj({pet: "kitten", bed: "sleigh"}, "gift")`은 문자열 `Not Found`를 반환해야 합니다.

```js
assert(checkObj({ pet: 'kitten', bed: 'sleigh' }, 'gift') === 'Not Found');
```

# --seed--

## --seed-contents--

```js
function checkObj(obj, checkProp) {
  // Only change code below this line
  return "Change Me!";
  // Only change code above this line
}
```

# --solutions--

```js
function checkObj(obj, checkProp) {
  if(obj.hasOwnProperty(checkProp)) {
    return obj[checkProp];
  } else {
    return "Not Found";
  }
}
```
