---
id: 5688e62ea601b2482ff8422b
title: 프로필 조회
challengeType: 1
forumTopicId: 18259
dashedName: profile-lookup
---

# --description--

연락처 목록에 있는 여러 사람을 나타내는 객체들의 배열이 있습니다.

`lookUpProfile` 함수가 `name`과 속성(`prop`)을 인수로 받도록 미리 작성되어 있습니다.

함수는 `name`이 실제 연락처의 `firstName`인지, 그리고 주어진 속성(`prop`)이 그 연락처의 속성인지 확인해야 합니다.

두 조건이 모두 참이면, 그 속성의 "값"을 반환하세요.

만약 `name`이 어떤 연락처와도 일치하지 않으면 문자열 `No such contact`을 반환하세요.

`prop`이 `name`와 일치하는 연락처의 유효한 속성과 일치하지 않으면 문자열 `No such property`을 반환하세요.

# --hints--

`lookUpProfile("Kristian", "lastName")`은 문자열 `Vos`을 반환해야 합니다.

```js
assert(lookUpProfile('Kristian', 'lastName') === 'Vos');
```

`lookUpProfile("Sherlock", "likes")`은 `["Intriguing Cases", "Violin"]`을 반환해야 합니다.

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")`은 배열을 반환해야 합니다.

```js
assert(typeof lookUpProfile('Harry', 'likes') === 'object');
```

`lookUpProfile("Bob", "number")`은 문자열 `No such contact`을 반환해야 합니다.

```js
assert(lookUpProfile('Bob', 'number') === 'No such contact');
```

`lookUpProfile("Bob", "potato")`은 문자열 `No such contact`을 반환해야 합니다.

```js
assert(lookUpProfile('Bob', 'potato') === 'No such contact');
```

`lookUpProfile("Akira", "address")`은 문자열 `No such property`을 반환해야 합니다.

```js
assert(lookUpProfile('Akira', 'address') === 'No such property');
```

# --seed--

## --seed-contents--

```js
// Setup
const contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

function lookUpProfile(name, prop) {
  // Only change code below this line

  // Only change code above this line
}

lookUpProfile("Akira", "likes");
```

# --solutions--

```js
const contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];
function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
