---
id: 5690307fddb111c6084545d7
title: if else 문에서 논리 순서
challengeType: 1
forumTopicId: 18228
dashedName: logical-order-in-if-else-statements
---

# --description--

`if`, `else if` 문에서는 순서가 중요합니다.

함수는 위에서 아래로 실행되므로 어떤 문장이 먼저 오는지 주의해야 합니다.

다음 두 함수를 예로 들어 보겠습니다.

첫 번째 함수는 다음과 같습니다:

```js
function foo(x) {
  if (x < 1) {
    return "Less than one";
  } else if (x < 2) {
    return "Less than two";
  } else {
    return "Greater than or equal to two";
  }
}
```

두 번째 함수는 문장의 순서만 바꿨습니다:

```js
function bar(x) {
  if (x < 2) {
    return "Less than two";
  } else if (x < 1) {
    return "Less than one";
  } else {
    return "Greater than or equal to two";
  }
}
```

이 두 함수는 거의 동일해 보이지만, 둘 다 숫자를 전달하면 서로 다른 결과가 나옵니다.

```js
foo(0)
bar(0)
```

`foo(0)`는 문자열 `Less than one`를 반환하고, `bar(0)`는 문자열 `Less than two`를 반환합니다.

# --instructions--

모든 경우에 올바른 문장을 반환하도록 함수 내 논리 순서를 변경하세요.

# --hints--

`orderMyLogic(4)`는 문자열 `Less than 5`를 반환해야 합니다.

```js
assert(orderMyLogic(4) === 'Less than 5');
```

`orderMyLogic(6)`는 문자열 `Less than 10`를 반환해야 합니다.

```js
assert(orderMyLogic(6) === 'Less than 10');
```

`orderMyLogic(11)`는 문자열 `Greater than or equal to 10`를 반환해야 합니다.

```js
assert(orderMyLogic(11) === 'Greater than or equal to 10');
```

# --seed--

## --seed-contents--

```js
function orderMyLogic(val) {
  if (val < 10) {
    return "Less than 10";
  } else if (val < 5) {
    return "Less than 5";
  } else {
    return "Greater than or equal to 10";
  }
}

orderMyLogic(7);
```

# --solutions--

```js
function orderMyLogic(val) {
  if(val < 5) {
    return "Less than 5";
  } else if (val < 10) {
    return "Less than 10";
  } else {
    return "Greater than or equal to 10";
  }
}
```
