---
id: 56bbb991ad1ed5201cd392d2
title: JavaScript 객체에 새 속성 추가하기
challengeType: 1
forumTopicId: 301169
dashedName: add-new-properties-to-a-javascript-object
---

# --description--

기존 JavaScript 객체에 새 속성을 추가하는 방법은 수정하는 방법과 같습니다.

다음은 `bark`에 `ourDog` 속성을 추가하는 방법입니다:

```js
ourDog.bark = "bow-wow";
```

또는

```js
ourDog["bark"] = "bow-wow";
```

이제 `ourDog.bark`를 평가하면 그의 짖는 소리인 `bow-wow`를 얻을 수 있습니다.

예시:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};

ourDog.bark = "bow-wow";
```

# --instructions--

`bark`에 `myDog` 속성을 추가하고 "woof" 같은 개 짖는 소리로 설정하세요. 점 표기법이나 대괄호 표기법 중 어느 것을 사용해도 됩니다.

# --hints--

`bark`에 `myDog` 속성을 추가해야 합니다.

```js
assert(myDog.bark !== undefined);
```

`bark`의 초기화에 `myDog`를 추가하면 안 됩니다.

```js
assert(!/bark[^\n]:/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};


```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.bark = "Woof Woof";
```
