---
id: 56bbb991ad1ed5201cd392d0
title: JavaScript 객체 만들기
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

`object`라는 용어를 들어본 적이 있을 거예요.

객체는 `arrays`와 비슷하지만, 데이터에 접근하고 수정할 때 인덱스를 사용하는 대신 `properties`라고 하는 것을 통해 객체 내 데이터를 접근합니다.

객체는 데이터를 구조화된 방식으로 저장하는 데 유용하며, 고양이 같은 실제 세계의 객체를 나타낼 수도 있습니다.

다음은 고양이 객체의 예입니다:

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

이 예제에서 모든 속성은 `name`, `legs`, `tails` 같은 문자열로 저장되어 있습니다. 하지만 숫자를 속성으로 사용할 수도 있습니다. 한 단어 문자열 속성의 경우 따옴표를 생략할 수도 있습니다:

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

하지만 객체에 문자열이 아닌 속성이 있으면 JavaScript가 자동으로 형변환하여 문자열로 만듭니다.

# --instructions--

`myDog`라는 이름의 개를 나타내는 객체를 만드세요. 이 객체는 `name` (문자열), `legs`, `tails`, `friends` 속성을 포함해야 합니다.

이 객체 속성들은 원하는 값으로 설정할 수 있지만, `name`은 문자열이어야 하고, `legs`와 `tails`은 숫자여야 하며, `friends`는 배열이어야 합니다.

# --hints--

`myDog`는 `name` 속성을 포함해야 하며, `string`이어야 합니다.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog`는 `legs` 속성을 포함해야 하며, `number`이어야 합니다.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog`는 `tails` 속성을 포함해야 하며, `number`이어야 합니다.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog`는 `friends` 속성을 포함해야 하며, `array`이어야 합니다.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog`는 주어진 모든 속성만 포함해야 합니다.

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
