---
id: 56bbb991ad1ed5201cd392d1
title: 객체 속성 업데이트
challengeType: 1
forumTopicId: 18336
dashedName: updating-object-properties
---

# --description--

JavaScript 객체를 만든 후에는 다른 변수처럼 언제든지 속성을 업데이트할 수 있습니다. 점 표기법이나 대괄호 표기법 중 어느 쪽이든 사용해 업데이트하세요.

예를 들어, `ourDog`을 살펴봅시다.

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```

그가 특히 행복한 강아지이므로 이름을 문자열 `Happy Camper`로 바꿔보겠습니다. 객체의 name 속성을 업데이트하는 방법은 `ourDog.name = "Happy Camper";` 또는 `ourDog["name"] = "Happy Camper";`입니다. 이제 `ourDog.name`를 평가하면 `Camper` 대신 새 이름인 `Happy Camper`가 나옵니다.

# --instructions--

`myDog` 객체의 name 속성을 업데이트하세요. 이름을 `Coder`에서 `Happy Coder`로 바꾸겠습니다. 점 표기법이나 대괄호 표기법 중 어느 쪽이든 사용하세요.

# --hints--

`myDog`의 `name` 속성을 문자열 `Happy Coder`로 업데이트해야 합니다.

```js
assert(/happy coder/gi.test(myDog.name));
```

`myDog` 정의는 수정하지 마세요.

```js
assert(/"name": "Coder"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.name = "Happy Coder";
```
