---
id: 56bbb991ad1ed5201cd392d0
title: 자바스크립트 객체 만들기
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

You may have heard the term `object` before.

객체는 `arrays`(배열)과 유사하지만, 데이터에 접근하고 수정하기 위해 인덱스를 사용하는 대신 `properties`(속성)을 통해 접근합니다.

객체는 데이터를 구조화된 방식으로 저장하는 데 유용하며, 고양이와 같은 실제 객체를 나타낼 수 있습니다.

다음은 샘플 고양이 객체입니다.

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

이 예에서 모든 속성은 `name`, `legs`, `tails`와 같은 문자열로 저장됩니다. 하지만, 숫자도 속성으로 사용할 수 있습니다. 또한, 단어 하나로 이루어진 문자열 속성에는 따옴표를 생략할 수도 있습니다.

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

그러나 객체에 문자열이 아닌 속성이 있는 경우 자바스크립트는 자동으로 그것들을 문자열로 변환합니다.

# --instructions--

`myDog`라는 이름의 개를 나타내는 객체를 만들어보겠습니다. 이 객체는 `name` (문자열), `legs`, `tails`, `friends` 속성을 포함해야 합니다.

이 객체 속성들은 원하는 값으로 설정할 수 있지만, `name`은 문자열, `legs`와 `tails`는 숫자, `friends`는 배열이어야 합니다.

# --hints--

`myDog`는 `name` 속성을 포함해야 하며, 이것은 `string`(문자열)이어야 합니다.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog`는 `legs` 속성을 포함해야 하며, 이것은 `number`(숫자)여야 합니다.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog`는 `tails` 속성을 포함해야 하며, 이것은 `number`(숫자)여야 합니다.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog`는 `friends` 속성을 포함해야 하며, 이것은 `array`(배열)여야 합니다.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog`는 주어진 속성만을 포함해야 합니다.

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --after-user-code--

```js
(function(z){return z;})(myDog);
```

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
