---
id: 56bbb991ad1ed5201cd392ca
title: 인덱스로 배열 데이터에 접근하기
challengeType: 1
forumTopicId: 16158
dashedName: access-array-data-with-indexes
---

# --description--

배열 안의 데이터는 <dfn>인덱스</dfn>를 사용해 접근할 수 있습니다.

배열 인덱스는 문자열이 사용하는 대괄호 표기법과 동일하게 작성하지만, 문자를 지정하는 대신 배열의 항목을 지정합니다. 문자열과 마찬가지로 배열도 <dfn>0부터 시작하는</dfn> 인덱스를 사용하므로, 배열의 첫 번째 요소는 인덱스가 `0`입니다.

<br>

**예시**

```js
const array = [50, 60, 70];
console.log(array[0]);
const data = array[1];
```

`console.log(array[0])`는 `50`를 출력하고, `data`의 값은 `60`입니다.

# --instructions--

`myData`의 첫 번째 값을 대괄호 표기법으로 접근하여 `myArray`라는 변수를 만들고 그 값으로 설정하세요.

# --hints--

변수 `myData`는 `myArray`의 첫 번째 값과 같아야 합니다.

```js
assert(
  (function () {
    if (
      typeof myArray !== 'undefined' &&
      typeof myData !== 'undefined' &&
      myArray[0] === myData
    ) {
      return true;
    } else {
      return false;
    }
  })()
);
```

변수 `myArray`의 데이터는 대괄호 표기법으로 접근해야 합니다.

```js
assert(
  (function () {
    if (__helpers.removeJSComments(code).match(/\s*=\s*myArray\[0\]/g)) {
      return true;
    } else {
      return false;
    }
  })()
);
```

# --seed--

## --seed-contents--

```js
const myArray = [50, 60, 70];


```

# --solutions--

```js
const myArray = [50, 60, 70];
const myData = myArray[0];
```
