---
id: 56bbb991ad1ed5201cd392cd
title: shift 메서드로 배열 조작하기
challengeType: 1
forumTopicId: 18238
dashedName: manipulate-arrays-with-shift
---

# --description--

`pop()`는 항상 배열의 마지막 요소를 제거합니다. 그렇다면 첫 번째 요소를 제거하려면 어떻게 해야 할까요?

바로 `.shift()`가 필요합니다. `.pop()`와 똑같이 작동하지만 마지막 대신 첫 번째 요소를 제거합니다.

예시:

```js
const ourArray = ["Stimpson", "J", ["cat"]];
const removedFromOurArray = ourArray.shift();
```

`removedFromOurArray`는 문자열 `Stimpson` 값을 가지며, `ourArray`는 `["J", ["cat"]]` 값을 가집니다.

# --instructions--

`.shift()` 함수를 사용해 `myArray`에서 첫 번째 항목을 제거하고, "shifted off"된 값을 새로운 변수 `removedFromMyArray`에 할당하세요.

# --hints--

`myArray`는 이제 `[["dog", 3]]`와 같아야 합니다.

```js
assert(
  (function (d) {
    if (d[0][0] == 'dog' && d[0][1] === 3 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`removedFromMyArray`는 `["John", 23]`를 포함해야 합니다.

```js
assert(
  (function (d) {
    if (
      d[0] == 'John' &&
      d[1] === 23 &&
      typeof removedFromMyArray === 'object'
    ) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line
const removedFromMyArray = myArray.shift();
```
