---
id: 56bbb991ad1ed5201cd392ce
title: unshift 메서드로 배열 조작하기
challengeType: 1
forumTopicId: 18239
dashedName: manipulate-arrays-with-unshift
---

# --description--

배열의 시작 부분에서 요소를 `shift`할 수 있을 뿐만 아니라, 배열 앞에 요소를 `unshift`할 수도 있습니다.

`.unshift()`는 `.push()`와 똑같이 작동하지만, 배열 끝에 요소를 추가하는 대신 `unshift()`는 배열의 시작 부분에 요소를 추가합니다.

예시:

```js
const ourArray = ["Stimpson", "J", "cat"];
ourArray.shift();
ourArray.unshift("Happy");
```

`shift` 후에 `ourArray`는 `["J", "cat"]` 값을 갖게 됩니다. `unshift` 후에 `ourArray`는 `["Happy", "J", "cat"]` 값을 갖게 됩니다.

# --instructions--

`["Paul", 35]`를 사용해 `myArray` 변수의 시작 부분에 `unshift()`을 추가하세요.

# --hints--

`myArray`는 이제 `[["Paul", 35], ["dog", 3]]`를 가져야 합니다.

```js
assert(
  (function (d) {
    if (
      typeof d[0] === 'object' &&
      d[0][0] == 'Paul' &&
      d[0][1] === 35 &&
      d[1][0] != undefined &&
      d[1][0] == 'dog' &&
      d[1][1] != undefined &&
      d[1][1] == 3
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();
myArray.unshift(["Paul", 35]);
```
