---
id: 56bbb991ad1ed5201cd392d3
title: JavaScript 객체에서 속성 삭제하기
challengeType: 1
forumTopicId: 17560
dashedName: delete-properties-from-a-javascript-object
---

# --description--

객체에서 속성을 이렇게 삭제할 수도 있습니다:

```js
delete ourDog.bark;
```

예시:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"],
  "bark": "bow-wow"
};

delete ourDog.bark;
```

위에 나온 마지막 줄 이후에 `ourDog`는 다음과 같습니다:

```js
{
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
}
```

# --instructions--

`tails`에서 `myDog` 속성을 삭제하세요. 점 표기법이나 대괄호 표기법 중 어느 것을 사용해도 됩니다.

# --hints--

`tails`에서 `myDog` 속성을 삭제해야 합니다.

```js
assert(typeof myDog === 'object' && myDog.tails === undefined);
```

`myDog` 설정은 수정하지 마세요.

```js
assert(__helpers.removeJSComments(code).match(/"tails": 1/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};
delete myDog.tails;
```
