---
id: 56bbb991ad1ed5201cd392cc
title: pop 메서드로 배열 조작하기
challengeType: 1
forumTopicId: 18236
dashedName: manipulate-arrays-with-pop
---

# --description--

배열의 데이터를 변경하는 또 다른 방법은 `.pop()` 함수입니다.

`.pop()`은 배열의 끝에서 값을 꺼내는 데 사용합니다. 꺼낸 값을 변수에 할당하여 저장할 수 있습니다. 다시 말해, `.pop()`는 배열의 마지막 요소를 제거하고 그 요소를 반환합니다.

숫자, 문자열, 심지어 중첩된 배열 등 어떤 유형의 항목도 배열에서 꺼낼 수 있습니다.

```js
const threeArr = [1, 4, 6];
const oneDown = threeArr.pop();
console.log(oneDown);
console.log(threeArr);
```

첫 번째 `console.log`는 `6` 값을 표시하고, 두 번째는 `[1, 4]` 값을 표시할 것입니다.

# --instructions--

`.pop()` 함수를 사용해 `myArray`에서 마지막 항목을 제거하고, 꺼낸 값을 새 변수 `removedFromMyArray`에 할당하세요.

# --hints--

`myArray`에는 `[["John", 23]]`만 포함되어야 합니다.

```js
assert(
  (function (d) {
    if (d[0][0] == 'John' && d[0][1] === 23 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`pop()`을 `myArray`에 사용하세요.

```js
assert(/removedFromMyArray\s*=\s*myArray\s*.\s*pop\s*(\s*)/.test(__helpers.removeJSComments(code)));
```

`removedFromMyArray`에는 `["cat", 2]`만 포함되어야 합니다.

```js
assert(
  (function (d) {
    if (d[0] == 'cat' && d[1] === 2 && d[2] == undefined) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
const removedFromMyArray = myArray.pop();
```
