---
id: 56bbb991ad1ed5201cd392cb
title: push 메서드로 배열 조작하기
challengeType: 1
forumTopicId: 18237
dashedName: manipulate-arrays-with-push
---

# --description--

배열 끝에 데이터를 추가하는 쉬운 방법은 `push()` 메서드를 사용하는 것입니다.

`push()` 메서드는 하나 이상의 <dfn>인수</dfn>를 받아 배열 끝에 인수들이 나타나는 순서대로 추가합니다. 그리고 배열의 새 길이를 반환합니다.

예시:

```js
const arr1 = [1, 2, 3];
arr1.push(4, 5);

const arr2 = ["Stimpson", "J", "cat"];
arr2.push(["happy", "joy"]);
```

`arr1`는 이제 `[1, 2, 3, 4, 5]` 값을 가지며 `arr2`는 `["Stimpson", "J", "cat", ["happy", "joy"]]` 값을 가집니다.

# --instructions--

`["dog", 3]`를 `myArray` 변수의 끝에 추가하세요.

# --hints--

`myArray`는 이제 `[["John", 23], ["cat", 2], ["dog", 3]]`와 같아야 합니다.

```js
assert(
  (function (d) {
    if (
      d[2] != undefined &&
      d[0][0] == 'John' &&
      d[0][1] === 23 &&
      d[2][0] == 'dog' &&
      d[2][1] === 3 &&
      d[2].length == 2
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
myArray.push(["dog",3]);
```
