---
id: 56bbb991ad1ed5201cd392cf
title: 함수로 재사용 가능한 JavaScript 작성하기
challengeType: 1
forumTopicId: 18378
dashedName: write-reusable-javascript-with-functions
---

# --description--

JavaScript에서는 코드를 재사용 가능한 부분인 <dfn>함수</dfn>로 나눌 수 있습니다.

함수의 예시는 다음과 같습니다:

```js
function functionName() {
  console.log("Hello World");
}
```

이 함수를 호출하거나 <dfn>실행</dfn>하려면 이름 뒤에 소괄호를 붙여 사용하세요. 예를 들어 `functionName();` 함수가 호출될 때마다 개발자 콘솔에 `Hello World` 메시지를 출력합니다. 중괄호 사이의 모든 코드는 함수가 호출될 때마다 실행됩니다.

# --instructions--

<ol>
  <li>
    <code>reusableFunction</code>라는 이름의 함수를 만들어 개발자 콘솔에 문자열 <code>Hi World</code>를 출력하세요.
  </li>
  <li>
    함수를 호출하세요.
  </li>
</ol>

# --before-each--

```js
function testConsole() {
  var originalConsole = console;
  var nativeLog = console.log;
  var hiWorldWasLogged = false;
  console.log = function (message) {
    if (message === 'Hi World') {
      console.warn(message);
      hiWorldWasLogged = true;
    }
    if (nativeLog.apply) {
      nativeLog.apply(originalConsole, arguments);
    } else {
      var nativeMsg = Array.prototype.slice.apply(arguments).join(' ');
      nativeLog(nativeMsg);
    }
  };
  reusableFunction();
  console.log = nativeLog;
  return hiWorldWasLogged;
}
```

# --hints--

`reusableFunction`는 함수여야 합니다.

```js
assert(typeof reusableFunction === 'function');
```

`reusableFunction`가 호출되면 문자열 `Hi World`를 콘솔에 출력해야 합니다.

```js
assert(testConsole());
```

`reusableFunction`가 정의되면 한 번 호출해야 합니다.

```js
const functionStr = reusableFunction && __helpers.removeWhiteSpace(reusableFunction.toString());
const codeWithoutFunction = __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).replace(/reusableFunction\(\)\{/g, '');
assert(/reusableFunction\(\)/.test(codeWithoutFunction));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reusableFunction() {
  console.log("Hi World");
}
reusableFunction();
```
