---
id: 587d7b7e367417b2b2512b24
title: 조건부(삼항) 연산자 사용하기
challengeType: 1
forumTopicId: 301181
dashedName: use-the-conditional-ternary-operator
---

# --description--

<dfn>조건부 연산자</dfn>는 <dfn>삼항 연산자</dfn>라고도 하며, 한 줄짜리 if-else 표현식으로 사용할 수 있습니다.

구문은 `a ? b : c`이며, 여기서 `a`은 조건, `b`는 조건이 `true`일 때 실행할 코드, `c`는 조건이 `false`일 때 실행할 코드입니다.

다음 함수는 `if/else` 문을 사용해 조건을 확인합니다:

```js
function findGreater(a, b) {
  if(a > b) {
    return "a is greater";
  }
  else {
    return "b is greater or equal";
  }
}
```

이것은 조건부 연산자를 사용해 다시 작성할 수 있습니다:

```js
function findGreater(a, b) {
  return a > b ? "a is greater" : "b is greater or equal";
}
```

# --instructions--

`checkEqual` 함수에서 조건부 연산자를 사용해 두 숫자가 같은지 다른지 확인하세요. 함수는 문자열 `Equal` 또는 문자열 `Not Equal`를 반환해야 합니다.

# --hints--

`checkEqual`은 조건부 연산자를 사용해야 합니다

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?/.test(__helpers.removeJSComments(code)));
```

`checkEqual(1, 2)`은 문자열 `Not Equal`를 반환해야 합니다

```js
assert(checkEqual(1, 2) === 'Not Equal');
```

`checkEqual(1, 1)`은 문자열 `Equal`를 반환해야 합니다

```js
assert(checkEqual(1, 1) === 'Equal');
```

`checkEqual(1, -1)`은 문자열 `Not Equal`를 반환해야 합니다

```js
assert(checkEqual(1, -1) === 'Not Equal');
```

# --seed--

## --seed-contents--

```js
function checkEqual(a, b) {

}

checkEqual(1, 2);
```

# --solutions--

```js
function checkEqual(a, b) {
  return a === b ? "Equal" : "Not Equal";
}
```
