---
id: 587d7b7e367417b2b2512b21
title: 여러 개의 조건부(삼항) 연산자 사용하기
challengeType: 1
forumTopicId: 301179
dashedName: use-multiple-conditional-ternary-operators
---

# --description--

이전 도전 과제에서는 단일 조건부 연산자를 사용했습니다. 여러 조건을 확인하기 위해 조건부 연산자를 연결해서 사용할 수도 있습니다.

다음 함수는 `if`, `else if`, 그리고 `else` 문을 사용해 여러 조건을 확인합니다:

```js
function findGreaterOrEqual(a, b) {
  if (a === b) {
    return "a and b are equal";
  }
  else if (a > b) {
    return "a is greater";
  }
  else {
    return "b is greater";
  }
}
```

위 함수는 여러 조건부 연산자를 사용해 다음과 같이 다시 작성할 수 있습니다:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" 
    : (a > b) ? "a is greater" 
    : "b is greater";
}
```

여러 조건부 연산자를 사용할 때는 각 조건을 별도의 줄에 작성하는 것이 모범 사례로 간주됩니다. 위 예시처럼 말이죠. 적절한 들여쓰기 없이 여러 조건부 연산자를 사용하면 코드를 읽기 어렵게 만들 수 있습니다. 예를 들어:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" : (a > b) ? "a is greater" : "b is greater";
}
```

# --instructions--

`checkSign` 함수에서 여러 조건부 연산자를 사용하세요. `findGreaterOrEqual`에서 사용한 권장 형식을 따라 숫자가 양수인지, 음수인지, 아니면 0인지 확인하세요. 함수는 `positive`, `negative` 또는 `zero`을 반환해야 합니다.

# --hints--

`checkSign`은 여러 조건부 연산자를 사용해야 합니다

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?\s*?\?\s*?.+?\s*?:\s*?.+?/gi.test(__helpers.removeJSComments(code)));
```

`checkSign(10)`은 문자열 `positive`을 반환해야 합니다. 대소문자를 구분한다는 점에 유의하세요

```js
assert(checkSign(10) === 'positive');
```

`checkSign(-12)`은 문자열 `negative`를 반환해야 합니다. 대소문자를 구분한다는 점에 유의하세요

```js
assert(checkSign(-12) === 'negative');
```

`checkSign(0)`은 문자열 `zero`를 반환해야 합니다. 대소문자를 구분한다는 점에 유의하세요

```js
assert(checkSign(0) === 'zero');
```

# --seed--

## --seed-contents--

```js
function checkSign(num) {

}

checkSign(10);
```

# --solutions--

```js
function checkSign(num) {
  return (num > 0) ? 'positive' : (num < 0) ? 'negative' : 'zero';
}
```
