---
id: 587d7b87367417b2b2512b3f
title: var와 let 키워드 차이점 살펴보기
challengeType: 1
forumTopicId: 301202
dashedName: explore-differences-between-the-var-and-let-keywords
---

# --description--

`var` 키워드로 변수를 선언할 때 가장 큰 문제 중 하나는 변수 선언을 쉽게 덮어쓸 수 있다는 점입니다.

```js
var camper = "James";
var camper = "David";
console.log(camper);
```

위 코드에서 `camper` 변수는 원래 `James`로 선언되었지만, 이후 `David`로 덮어쓰였습니다. 콘솔에는 문자열 `David`가 출력됩니다.

작은 애플리케이션에서는 이런 문제가 발생하지 않을 수 있습니다. 하지만 코드베이스가 커질수록 의도하지 않은 변수를 실수로 덮어쓸 수 있습니다. 이 동작은 오류를 발생시키지 않기 때문에 버그를 찾고 고치기가 더 어려워집니다.

ES6라는 JavaScript의 주요 업데이트에서 `let` 키워드의 잠재적 문제를 해결하기 위해 `var`라는 키워드가 도입되었습니다. 이후 도전 과제에서 다른 ES6 기능도 배우게 됩니다.

위 코드에서 `var`을 `let`로 바꾸면 오류가 발생합니다:

```js
let camper = "James";
let camper = "David";
```

오류는 브라우저 콘솔에서 확인할 수 있습니다.

따라서 `var`와 달리 `let`을 사용할 때는 같은 이름의 변수를 한 번만 선언할 수 있습니다.

# --instructions--

코드를 `let` 키워드만 사용하도록 수정하세요.

# --hints--

코드에 `var`가 없어야 합니다.

```js
assert.notMatch(code, /var/g);
```

`catName`는 문자열 `Oliver`여야 합니다.

```js
assert.equal(catName, 'Oliver');
```

`catSound`는 문자열 `Meow!`여야 합니다.

```js
assert.equal(catSound, 'Meow!');
```

# --seed--

## --seed-contents--

```js
var catName = "Oliver";
var catSound = "Meow!";
```

# --solutions--

```js
let catName = "Oliver";
let catSound = "Meow!";
```
