---
id: 587d7b7e367417b2b2512b23
title: parseInt 함수 사용하기
challengeType: 1
forumTopicId: 301183
dashedName: use-the-parseint-function
---

# --description--

`parseInt()` 함수는 문자열을 파싱하여 정수를 반환합니다. 예시는 다음과 같습니다:

```js
const a = parseInt("007");
```

위 함수는 문자열 `007`을 정수 `7`로 변환합니다. 만약 문자열의 첫 번째 문자를 숫자로 변환할 수 없으면 `NaN`을 반환합니다.

# --instructions--

`parseInt()`를 `convertToInteger` 함수 안에서 사용하여 입력 문자열 `str`을 정수로 변환하고 반환하세요.

# --hints--

`convertToInteger`는 `parseInt()` 함수를 사용해야 합니다

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("56")`는 숫자를 반환해야 합니다

```js
assert(typeof convertToInteger('56') === 'number');
```

`convertToInteger("56")`는 56을 반환해야 합니다

```js
assert(convertToInteger('56') === 56);
```

`convertToInteger("77")`는 77을 반환해야 합니다

```js
assert(convertToInteger('77') === 77);
```

`convertToInteger("JamesBond")`는 `NaN`을 반환해야 합니다

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("56");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str);
}
```
