---
id: 587d7b87367417b2b2512b41
title: const 키워드로 읽기 전용 변수 선언하기
challengeType: 1
forumTopicId: 301201
dashedName: declare-a-read-only-variable-with-the-const-keyword
---

# --description--

키워드 `let`만이 변수를 선언하는 새로운 방법이 아닙니다. ES6에서는 `const` 키워드를 사용해 변수를 선언할 수도 있습니다.

`const`는 `let`가 가진 모든 멋진 기능을 가지고 있으며, 여기에 `const`로 선언된 변수는 읽기 전용이라는 추가 이점이 있습니다. 이는 상수 값으로, `const`로 변수가 할당되면 다시 할당할 수 없다는 뜻입니다:

```js
const FAV_PET = "Cats";
FAV_PET = "Dogs";
```

값을 다시 할당했기 때문에 콘솔에 오류가 표시됩니다 `FAV_PET`.

다시 할당하지 않을 변수를 선언할 때는 항상 `const` 키워드를 사용하세요. 이렇게 하면 상수로 유지되어야 하는 변수를 실수로 다시 할당하려 할 때 도움이 됩니다.

**참고:** 개발자들은 변경 불가능한 값에는 대문자 변수 식별자를, 변경 가능한 값(객체와 배열)에는 소문자나 카멜 케이스를 사용하는 것이 일반적입니다. 객체, 배열, 변경 불가능한 값과 변경 가능한 값에 대해서는 이후 도전 과제에서 더 배우게 됩니다. 또한 이후 도전 과제에서 대문자, 소문자, 카멜 케이스 변수 식별자 예시도 볼 수 있습니다.

# --instructions--

모든 변수를 `let` 또는 `const`로 선언하도록 코드를 변경하세요. 변수가 변경될 수 있어야 할 때는 `let`을, 변수가 상수로 유지되어야 할 때는 `const`을 사용하세요. 또한 `const`로 선언된 변수 이름을 일반적인 관례에 맞게 바꾸세요. 변수에 할당된 문자열은 변경하지 마세요.

# --hints--

코드에 `var`가 없어야 합니다.

```js
assert.notMatch(code, /var/g);
```

`fCC`를 모두 대문자로 바꾸세요.

```js
assert.match(code, /(FCC)/);
assert.notMatch(code, /(fCC)/);
```

`FCC`는 `const`로 선언된 상수 변수여야 합니다.

```js
assert.match(code, /const\s+FCC/);
```

`FCC`에 할당된 문자열은 변경하지 마세요.

```js
assert.equal(FCC, 'freeCodeCamp');
```

`fact`는 `let`로 선언해야 합니다.

```js
assert.match(code, /(let\s+fact)/g);
```

`console.log`는 `FCC`과 `fact` 변수를 출력하도록 바꾸세요.

```js
assert.match(code, /console\.log\(\s*FCC\s*\,\s*fact\s*\)\s*;?/g);
```

# --seed--

## --seed-contents--

```js
var fCC = "freeCodeCamp"; // Change this line
var fact = "is cool!"; // Change this line
fact = "is awesome!";
console.log(fCC, fact); // Change this line
```

# --solutions--

```js
const FCC = "freeCodeCamp";
let fact = "is cool!";

fact = "is awesome!";
console.log(FCC, fact);
```
