---
id: 587d7b7e367417b2b2512b22
title: 기수를 사용해 parseInt 함수 활용하기
challengeType: 1
forumTopicId: 301182
dashedName: use-the-parseint-function-with-a-radix
---

# --description--

`parseInt()` 함수는 문자열을 파싱하여 정수를 반환합니다. 이 함수는 두 번째 인수로 기수를 받는데, 이는 문자열에 있는 숫자의 진법을 지정합니다. 기수는 2부터 36 사이의 정수일 수 있습니다.

함수 호출은 다음과 같습니다:

```js
parseInt(string, radix);
```

예시는 다음과 같습니다:

```js
const a = parseInt("11", 2);
```

기수 변수는 `11`이 2진법, 즉 기수 2임을 나타냅니다. 이 예시는 문자열 `11`를 정수 `3`로 변환합니다.

# --instructions--

`parseInt()`를 `convertToInteger` 함수 안에서 사용하여 2진수를 정수로 변환하고 반환하세요.

# --hints--

`convertToInteger`는 `parseInt()` 함수를 사용해야 합니다

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("10011")`는 숫자를 반환해야 합니다

```js
assert(typeof convertToInteger('10011') === 'number');
```

`convertToInteger("10011")`는 19를 반환해야 합니다

```js
assert(convertToInteger('10011') === 19);
```

`convertToInteger("111001")`는 57을 반환해야 합니다

```js
assert(convertToInteger('111001') === 57);
```

`convertToInteger("JamesBond")`는 `NaN`를 반환해야 합니다

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("10011");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str, 2);
}
```
