---
id: 598e8944f009e646fc236146
title: 함수에서 반환된 undefined 값 이해하기
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

함수는 `return` 문장을 포함할 수도 있고 포함하지 않을 수도 있습니다. 함수에 `return` 문장이 없으면, 함수를 호출할 때 함수는 내부 코드를 처리하지만 반환 값은 `undefined`입니다.

**예시**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum`는 `return` 문장이 없는 함수입니다. 이 함수는 전역 `sum` 변수를 변경하지만 함수의 반환 값은 `undefined`입니다.

# --instructions--

인수가 없는 `addFive` 함수를 만드세요. 이 함수는 `sum` 변수에 5를 더하지만 반환 값은 `undefined`입니다.

# --hints--

`addFive`는 함수여야 합니다.

```js
assert(typeof addFive === 'function');
```

두 함수가 모두 실행된 후 `sum`는 `8`와 같아야 합니다.

```js
assert(sum === 8);
```

`addFive`의 반환 값은 `undefined`여야 합니다.

```js
assert(addFive() === undefined);
```

`addFive` 함수 내부에서 `5`를 `sum` 변수에 더해야 합니다.

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
