---
id: 598e8944f009e646fc236146
title: 함수로부터 반환된 undefined 값 이해하기
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

A function can include the `return` statement but it does not have to. In the case that the function doesn't have a `return` statement, when you call it, the function processes the inner code but the returned value is `undefined`.

**예)**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum`은 `return`문이 없는 함수입니다. 함수는 전역 `sum` 변수를 변형시킬 것이지만 해당 함수의 반환값은 `undefined`입니다.

# --instructions--

어떤 인자도 받지 않은 `addFive` 함수를 생성하시오. 이 함수는 `sum` 변수에 5를 더하지만 그 반환값은 `undefined`입니다.

# --hints--

`addFive`는 함수이어야 합니다.

```js
assert(typeof addFive === 'function');
```

두 함수가 동작되면 `sum`은 `8`과 같아야 합니다.

```js
assert(sum === 8);
```

`addFive`의 반환값은 `undefined`이어야 합니다.

```js
assert(addFive() === undefined);
```

`addFive` 함수 안에서 `sum` 변수에 `5`가 더해져야 합니다.

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
