---
id: 599a789b454f2bbd91a3ff4d
title: 서로 다른 값 비교 연습
challengeType: 1
forumTopicId: 301174
dashedName: practice-comparing-different-values
---

# --description--

지난 두 도전 과제에서 동등 연산자(`==`)와 엄격한 동등 연산자(`===`)에 대해 배웠습니다. 이제 이 연산자들을 다시 검토하고 더 연습해 보겠습니다.

비교하는 값들의 타입이 같지 않으면 동등 연산자는 타입 변환을 수행한 후 값을 평가합니다. 반면 엄격한 동등 연산자는 타입 변환 없이 데이터 타입과 값을 그대로 비교합니다.

**예시**

`3 == '3'`는 JavaScript가 문자열에서 숫자로 타입 변환을 수행하기 때문에 `true`를 반환합니다. `3 === '3'`는 타입이 다르고 타입 변환이 수행되지 않으므로 `false`를 반환합니다.

**참고:** JavaScript에서는 `typeof` 연산자를 사용해 변수나 값의 타입을 다음과 같이 확인할 수 있습니다:

```js
typeof 3
typeof '3'
```

`typeof 3`는 문자열 `number`를 반환하고, `typeof '3'`는 문자열 `string`를 반환합니다.

# --instructions--

편집기의 `compareEquality` 함수는 동등 연산자를 사용해 두 값을 비교합니다. 값이 엄격히 같을 때만 문자열 `Equal`를 반환하도록 함수를 수정하세요.

# --hints--

`compareEquality(10, "10")`는 문자열 `Not Equal`를 반환해야 합니다.

```js
assert(compareEquality(10, '10') === 'Not Equal');
```

`compareEquality("20", 20)`는 문자열 `Not Equal`를 반환해야 합니다.

```js
assert(compareEquality('20', 20) === 'Not Equal');
```

`===` 연산자를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function compareEquality(a, b) {
  if (a == b) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

compareEquality(10, "10");
```

# --solutions--

```js
function compareEquality(a,b) {
  if (a === b) {
    return "Equal";
  }
  return "Not Equal";
}
```
