---
id: 5a2efd662fb457916e1fe604
title: "JavaScript `do...while` 반복문으로 반복하기"
challengeType: 1
forumTopicId: 301172
dashedName: iterate-with-javascript-do---while-loops
---

# --description--

다음으로 배울 반복문 유형은 `do...while` 반복문입니다. `do...while` 반복문이라고 부르는 이유는, 조건과 상관없이 먼저 반복문 내부 코드를 한 번 `do` 실행하고, 그다음 지정된 조건이 `true`일 때까지 `while` 실행하기 때문입니다.

```js
const ourArray = [];
let i = 0;

do {
  ourArray.push(i);
  i++;
} while (i < 5);
```

위 예시는 다른 반복문과 비슷하게 동작하며, 결과 배열은 `[0, 1, 2, 3, 4]`처럼 보일 것입니다. 하지만 `do...while` 반복문이 다른 반복문과 다른 점은 조건이 처음 검사될 때 실패하면 어떻게 동작하는지입니다. 실제로 살펴보겠습니다. 다음은 조건이 `i < 5` 동안 반복문 내부 코드를 실행하는 일반적인 `while` 반복문입니다:

```js
const ourArray = []; 
let i = 5;

while (i < 5) {
  ourArray.push(i);
  i++;
}
```

이 예시에서는 `ourArray` 값을 빈 배열로 초기화하고 `i` 값을 5로 설정했습니다. `while` 반복문을 실행하면, `i`이 5보다 작지 않으므로 조건이 `false`로 평가되어 반복문 내부 코드를 실행하지 않습니다. 결과적으로 `ourArray`에는 값이 추가되지 않고, 위 예시의 모든 코드 실행이 끝난 후에도 `[]`처럼 남아 있습니다. 이제 `do...while` 반복문을 살펴보겠습니다:

```js
const ourArray = []; 
let i = 5;

do {
  ourArray.push(i);
  i++;
} while (i < 5);
```

이번에는 `i` 값을 5로 초기화했는데, `while` 반복문과 같습니다. 다음 줄로 넘어가면 평가할 조건이 없으므로 중괄호 내부 코드를 실행합니다. 배열에 단일 요소를 추가하고 `i`를 증가시킨 후 조건 검사를 합니다. 마지막 줄에서 조건 `i < 5`을 평가할 때 `i`이 6이 되어 조건 검사를 통과하지 못하므로 반복문을 종료합니다. 위 예시가 끝날 때 `ourArray` 값은 `[5]`입니다. 본질적으로 `do...while` 반복문은 반복문 내부 코드가 최소 한 번은 실행되도록 보장합니다. 이제 배열에 값을 추가하는 `do...while` 반복문을 만들어 보겠습니다.

# --instructions--

코드에서 `while` 반복문을 `do...while` 반복문으로 바꾸어, 반복문이 숫자 `10`만 `myArray`에 추가하고, 코드 실행이 끝났을 때 `i`이 `11`과 같도록 하세요.

# --hints--

이 연습에서는 `do...while` 반복문을 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/do/g));
```

`myArray`은 `[10]`와 같아야 합니다.

```js
assert.deepEqual(myArray, [10]);
```

`i`은 `11`와 같아야 합니다.

```js
assert.equal(i, 11);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];
let i = 10;

// Only change code below this line
while (i < 5) {
  myArray.push(i);
  i++;
}
```

# --solutions--

```js
const myArray = [];
let i = 10;
do {
  myArray.push(i);
  i++;
} while (i < 5)
```
