---
id: 5cc0bd7a49b71cb96132e54c
title: 재귀를 사용하여 숫자의 범위를 만들기
challengeType: 1
forumTopicId: 301180
dashedName: use-recursion-to-create-a-range-of-numbers
---

# --description--

Continuing from the previous challenge, we provide you another opportunity to create a recursive function to solve a problem.

# --instructions--

두 개의 매개변수를 가진 `rangeOfNumbers`라는 함수를 정의했습니다. 이 함수는 `startNum` 매개변수로 표현된 숫자로 시작하여 `endNum` 매개변수로 표현된 숫자로 끝나는 정수 배열을 반환해야 합니다. 시작 숫자는 항상 끝나는 숫자보다 작거나 같을 것입니다. 함수는 자체적으로 재귀 호출을 사용해야 하며, 어떤 종류의 반복문도 사용해서는 안 됩니다. 또한 `startNum`과 `endNum`이 동일한 경우에도 작동해야 합니다.

# --hints--

이 함수는 배열을 반환해야 합니다.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

코드는 반복문 문법(`for` 또는 `while` 또는 `forEach`, `map`, `filter`, `reduce`와 같은 고차 함수)을 사용해서는 안 됩니다.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

`rangeOfNumbers` 함수는 이 과제를 해결하기 위해 재귀(자기 자신을 호출)를 사용해야 합니다.

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

`rangeOfNumbers(1, 5)`은 `[1, 2, 3, 4, 5]`을 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)`은 `[6, 7, 8, 9]`을 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)`은 `[4]`을 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

전역 변수를 사용하여 배열을 캐시하지 않아야 합니다.

```js
rangeOfNumbers(1, 3)
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js
function rangeOfNumbers(startNum, endNum) {
  return [];
};
```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
