---
id: 5cc0bd7a49b71cb96132e54c
title: 재귀를 사용해 숫자 범위 만들기
challengeType: 1
forumTopicId: 301180
dashedName: use-recursion-to-create-a-range-of-numbers
---

# --description--

이전 도전 과제에 이어, 문제를 해결하기 위해 재귀 함수를 만드는 또 다른 기회를 제공합니다.

# --instructions--

두 개의 매개변수를 가진 `rangeOfNumbers`라는 함수를 정의했습니다. 이 함수는 `startNum` 매개변수가 나타내는 숫자로 시작해서 `endNum` 매개변수가 나타내는 숫자로 끝나는 정수 배열을 반환해야 합니다. 시작 숫자는 항상 끝 숫자보다 작거나 같습니다. 함수는 자기 자신을 호출하는 재귀를 사용해야 하며 어떤 종류의 반복문도 사용하지 않아야 합니다. `startNum`와 `endNum`가 같은 경우에도 작동해야 합니다.

# --hints--

함수는 배열을 반환해야 합니다.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

코드에 반복문 구문(`for` 또는 `while` 또는 `forEach`, `map`, `filter`, `reduce` 같은 고차 함수)을 사용하지 않아야 합니다.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

`rangeOfNumbers`는 이 도전 과제를 해결하기 위해 재귀(자기 자신 호출)를 사용해야 합니다.

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

`rangeOfNumbers(1, 5)`는 `[1, 2, 3, 4, 5]`를 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)`는 `[6, 7, 8, 9]`를 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)`는 `[4]`를 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

전역 변수는 배열을 저장하는 데 사용하지 않아야 합니다.

```js
rangeOfNumbers(1, 3)
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js
function rangeOfNumbers(startNum, endNum) {
  return [];
};
```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
