---
id: 5cd9a70215d3c4e65518328f
title: 재귀를 사용해 카운트다운 만들기
challengeType: 1
forumTopicId: 305925
dashedName: use-recursion-to-create-a-countdown
---

# --description--

<a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/replace-loops-using-recursion" target="_blank" rel="noopener noreferrer nofollow">이전 도전 과제</a>에서 재귀를 사용해 `for` 루프를 대체하는 방법을 배웠습니다. 이제 `1`부터 함수에 전달된 숫자까지 연속된 정수를 포함하는 배열을 반환하는 좀 더 복잡한 함수를 살펴보겠습니다.

이전 도전 과제에서 언급했듯이 <dfn>기저 사례</dfn>가 있습니다. 기저 사례는 재귀 함수가 더 이상 자기 자신을 호출하지 않아도 되는 시점을 알려줍니다. 반환값이 이미 알려진 간단한 경우입니다. 또한 <dfn>재귀 호출</dfn>이 있는데, 이는 다른 인수로 원래 함수를 실행하는 것입니다. 함수가 올바르게 작성되면 결국 기저 사례에 도달하게 됩니다.

예를 들어, `1`부터 `n`까지의 숫자를 포함하는 배열을 반환하는 재귀 함수를 작성한다고 가정해 보겠습니다. 이 함수는 최종 숫자를 나타내는 인수 `n`를 받아야 합니다. 그런 다음 `n`의 값이 점점 작아지도록 자기 자신을 호출하다가 `1`에 도달할 때까지 반복해야 합니다. 함수를 다음과 같이 작성할 수 있습니다:

```javascript
function countup(n) {
  if (n < 1) {
    return [];
  } else {
    const countArray = countup(n - 1);
    countArray.push(n);
    return countArray;
  }
}
console.log(countup(5));
```

콘솔에는 `[1, 2, 3, 4, 5]` 값이 출력됩니다.

처음에는 `n` 값이 *감소*하는데 최종 배열의 값은 *증가*한다는 점이 직관에 반하는 것처럼 보입니다. 이는 push가 재귀 호출이 반환된 후 마지막에 실행되기 때문입니다. `n`가 배열에 추가되는 시점에는 `countup(n - 1)`이 이미 평가되어 `[1, 2, ..., n - 1]`을 반환한 상태입니다.

# --instructions--

`countdown`라는 이름의 함수가 매개변수 `n` 하나와 함께 정의되어 있습니다. 이 함수는 재귀를 사용해 `n` 매개변수를 기준으로 `1`부터 `n`까지의 정수를 포함하는 배열을 반환해야 합니다. 만약 함수가 1보다 작은 숫자로 호출되면 빈 배열을 반환해야 합니다. 예를 들어, 이 함수를 `n = 5`로 호출하면 `[5, 4, 3, 2, 1]` 배열을 반환해야 합니다. 함수는 반드시 자기 자신을 호출하는 재귀를 사용해야 하며 어떤 종류의 루프도 사용해서는 안 됩니다.

# --hints--

`countdown(-1)`는 빈 배열을 반환해야 합니다.

```js
assert.isEmpty(countdown(-1));
```

`countdown(10)`는 `[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`을 반환해야 합니다.

```js
assert.deepStrictEqual(countdown(10), [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

`countdown(5)`는 `[5, 4, 3, 2, 1]`을 반환해야 합니다.

```js
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

코드는 어떤 종류의 루프(`for`, `while` 또는 `forEach`, `map`, `filter`, `reduce` 같은 고차 함수)를 사용하지 않아야 합니다.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

이 문제는 재귀를 사용해 해결해야 합니다.

```js
assert(
  countdown.toString().match(/countdown\s*\(.+\)/)
);
```

배열을 캐시하기 위해 전역 변수를 사용해서는 안 됩니다.

```js
countdown(1)
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
function countdown(n){
  return;
}
// Only change code above this line
```

# --solutions--

```js
function countdown(n){
   return n < 1 ? [] : [n].concat(countdown(n - 1));
}
```
