---
id: 5cfa3679138e7d9595b9d9d4
title: 재귀를 사용해 반복문 대체하기
challengeType: 1
videoUrl: >-
  https://www.freecodecamp.org/news/how-recursion-works-explained-with-flowcharts-and-a-video-de61f40cb7f9/
forumTopicId: 301175
dashedName: replace-loops-using-recursion
---

# --description--

재귀는 함수가 자기 자신을 이용해 표현될 수 있다는 개념입니다. 이를 이해하기 쉽게, 다음 과제를 생각해 보세요: 배열의 처음 `n` 요소를 곱해서 그 요소들의 곱을 만드는 것입니다. `for` 반복문을 사용하면 이렇게 할 수 있습니다:

```js
  function multiply(arr, n) {
    let product = 1;
    for (let i = 0; i < n; i++) {
      product *= arr[i];
    }
    return product;
  }
```

하지만 `multiply(arr, n) == multiply(arr, n - 1) * arr[n - 1]`임을 주목하세요. 이것은 `multiply`을 자기 자신을 이용해 다시 쓸 수 있고 반복문을 전혀 사용할 필요가 없다는 뜻입니다.

```js
  function multiply(arr, n) {
    if (n <= 0) {
      return 1;
    } else {
      return multiply(arr, n - 1) * arr[n - 1];
    }
  }
```

`multiply`의 재귀 버전은 이렇게 나눌 수 있습니다. <dfn>기본 사례</dfn>에서는 `n <= 0`일 때 1을 반환합니다. `n`가 더 큰 값일 때는 자기 자신을 호출하지만 `n - 1`를 인자로 넘깁니다. 이 함수 호출은 같은 방식으로 평가되어 `multiply`을 다시 호출하며 `n <= 0`가 될 때까지 계속됩니다. 이 시점에서 모든 함수가 반환할 수 있고 원래 `multiply`이 답을 반환합니다.

**참고:** 재귀 함수는 다시 호출하지 않고 반환하는 기본 사례가 반드시 있어야 합니다(이 예에서는 `n <= 0`일 때), 그렇지 않으면 실행이 끝나지 않습니다.

# --instructions--

배열 `sum(arr, n)`의 처음 `n` 요소의 합을 반환하는 재귀 함수 `arr`를 작성하세요.

# --hints--

`sum([1], 0)`는 0이어야 합니다.

```js
assert.equal(sum([1], 0), 0);
```

`sum([2, 3, 4], 1)`는 2이어야 합니다.

```js
assert.equal(sum([2, 3, 4], 1), 2);
```

`sum([2, 3, 4, 5], 3)`는 9이어야 합니다.

```js
assert.equal(sum([2, 3, 4, 5], 3), 9);
```

코드는 어떤 종류의 반복문(`for`, `while` 등)이나 `forEach`, `map`, `filter`, `reduce` 같은 고차 함수에 의존해서는 안 됩니다.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

이 문제는 재귀를 사용해 해결해야 합니다.

```js
assert(
  sum.toString().match(/sum\(.*\)/g).length > 1
);
```

# --seed--

## --seed-contents--

```js
function sum(arr, n) {
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function sum(arr, n) {
  // Only change code below this line
  if(n <= 0) {
    return 0;
  } else {
    return sum(arr, n - 1) + arr[n - 1];
  }
  // Only change code above this line
}
```
