---
id: 5ee127a03c3b35dd45426493
title: 한 변수의 값을 다른 변수에 할당하기
challengeType: 1
forumTopicId: 418265
dashedName: assigning-the-value-of-one-variable-to-another
---

# --description--

<dfn>할당</dfn> 연산자를 사용해 변수에 값이 할당된 후, 그 변수의 값을 다른 변수에 <dfn>할당</dfn> 연산자를 사용해 할당할 수 있습니다.

```js
var myVar;
myVar = 5;
var myNum;
myNum = myVar;
```

위 코드는 `myVar` 변수를 값 없이 선언한 다음 `5` 값을 할당합니다. 그다음 `myNum`라는 변수를 값 없이 선언합니다. 그리고 `myVar`의 내용(`5`)을 `myNum` 변수에 할당합니다. 이제 `myNum`도 `5` 값을 갖게 됩니다.

# --instructions--

`a`의 내용을 `b` 변수에 할당하세요.

# --hints--

지정된 주석 위의 코드는 변경하지 마세요.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)) && /a = 7;/.test(__helpers.removeJSComments(code)) && /var b;/.test(__helpers.removeJSComments(code)));
```

`b`은 `7` 값을 가져야 합니다.

```js
assert(typeof b === 'number' && b === 7);
```

`a`은 `b`를 사용해 `=`에 할당되어야 합니다.

```js
assert(/b\s*=\s*a\s*/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
var a;
a = 7;
var b;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
var b;
b = a;
```
