---
id: bd7993c9c69feddfaeb7bdef
title: JavaScript로 소수 두 개 곱하기
challengeType: 1
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

JavaScript에서는 정수처럼 소수로도 계산할 수 있습니다.

두 소수를 곱해서 그 곱셈 결과를 구해보세요.

# --instructions--

`0.0`를 변경해서 product가 `5.0`과 같아지게 하세요.

# --hints--

변수 `product`는 `5.0`과 같아야 합니다.

```js
assert(product === 5.0);
```

`*` 연산자를 사용해야 합니다.

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
