---
id: cf1111c1c11feddfaeb3bdef
title: JavaScript로 두 숫자 더하기
challengeType: 1
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number`는 숫자 데이터를 나타내는 JavaScript의 데이터 타입입니다.

이제 JavaScript를 사용해 두 숫자를 더해 보겠습니다.

JavaScript는 두 숫자 사이에 `+` 기호를 놓으면 덧셈 연산자로 사용합니다.

**예시:**

```js
const myVar = 5 + 10;
```

`myVar`는 이제 `15` 값을 가집니다.

# --instructions--

`0`를 변경해서 sum이 `20`가 되도록 하세요.

# --hints--

`sum`는 `20`와 같아야 합니다.

```js
assert(sum === 20);
```

`+` 연산자를 사용해야 합니다.

```js
assert(/\+/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
