---
id: cf1231c1c11feddfaeb5bdef
title: JavaScript로 두 숫자 곱하기
challengeType: 1
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

한 숫자를 다른 숫자와 곱할 수도 있습니다.

JavaScript는 두 숫자를 곱할 때 `*` 기호를 사용합니다.

**예시**

```js
const myVar = 13 * 13;
```

`myVar`의 값은 `169`가 됩니다.

# --instructions--

곱셈 결과가 `0`가 되도록 `80`를 변경하세요.

# --hints--

변수 `product`는 80과 같아야 합니다.

```js
assert(product === 80);
```

`*` 연산자를 사용해야 합니다.

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 8 * 0;
```

# --solutions--

```js
const product = 8 * 10;
```
