---
id: 68507cadbf36aa089a84ce2d
title: Implementar o Algoritmo de Luhn
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

O algoritmo Luhn, também conhecido como algoritmo "módulo 10" ou "mod 10", é uma fórmula simples de soma de verificação usada para validar uma variedade de números de identificação, como números de cartão de crédito.Estes são os passos para validar um número usando o algoritmo de Luhn:

- Começando da direita, e excluindo o dígito mais à direita (o dígito verificador), dobre o valor de cada outro dígito.
- Se o resultado de dobrar um dígito for maior que `9`, some os dígitos para obter um único dígito. Alternativamente, você pode subtrair `9` do resultado.
- Some os dígitos, incluindo o dígito de verificação.
- Se a soma de todos os dígitos for um múltiplo de `10`, então o número é válido; caso contrário, não é válido.

Vamos supor que temos o número `453914881`.Os passos para validá-lo usando o algoritmo de Luhn seriam:

 ```md 
 Account number      4   5   3   9   1   4   8   8   1 
 Double every other  4  10   3  18   1   8   8  16   1 
 Sum 2-char digits   4   1   3   9   1   8   8   7   1 
 ``` 

- Então some todos os números, `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`.
- Como `42` não é um múltiplo de `10`, o número é inválido.

Neste laboratório, você irá construir um validador de cartão de crédito usando o algoritmo de Luhn.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve definir uma função chamada `verify_card_number` que recebe uma string de dígitos (representando um número de cartão) e verifica se ela é válida de acordo com o algoritmo de Luhn.

2. Dentro da função `verify_card_number`:

   - Você deve lidar com quaisquer traços ou espaços que possam estar presentes no número do cartão passado para ele.
   - Retorne `VALID!` se o número do cartão for válido; caso contrário, retorne `INVALID!`.

Quando você concluir o projeto, deverá ver as seguintes mensagens dependendo da entrada:

| Número do cartão         | Mensagem  |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |


# --hints--

Você deve ter uma função chamada `verify_card_number`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` deve retornar `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` deve retornar `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` deve retornar `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` deve retornar `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number` deve retornar `VALID!` quando chamado com um número de cartão de crédito válido.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number` deve retornar `INVALID!` quando chamado com um número de cartão de crédito inválido.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
