---
id: 69baafd6d683dd9921b75db3
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 JavaScript를 사용해 음악 축제 참석자들의 흐름을 시뮬레이션합니다.

아침과 밤 근무 시간 동안 축제 게이트를 설명하는 두 개의 데이터셋을 다룹니다:

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

각 게이트 객체는 다음 프로퍼티를 포함합니다:

- `id`: 게이트를 식별하는 문자열입니다.
- `capacity`: 게이트가 한 틱당 처리할 수 있는 참석자 수입니다.
- `queue`: 특정 틱 동안 게이트에 도착하는 참석자 수를 나타내는 숫자 배열입니다.

먼저, `initializeThroughput`라는 매개변수를 받는 빈 함수 `gates`를 만드세요.

# --hints--

`initializeThroughput`라는 이름의 함수를 만드세요.

```js
assert.isFunction(initializeThroughput);
```

`initializeThroughput` 함수는 `gates` 매개변수를 가져야 합니다.

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
