---
id: 69bac03ae5faf2979a3339ff
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

이제 어느 게이트로 라우트할지 알았으니, 오버플로우 참석자를 이동할 수 있습니다.

`overflowAmount`를 가져와서 다음 게이트의 같은 `tickIndex`에 추가하세요.

구체적으로 다음을 해야 합니다:

- `nextGateIndex`를 사용해 다음 게이트에 접근하세요.
- 그 `queue`에 접근하세요.
- `overflowAmount`를 `tickIndex` 위치의 값에 추가하세요.

# --hints--

`rerouteOverflow` 함수는 다음 게이트의 `overflowAmount`에서 `tickIndex` 위치의 값에 `queue`를 추가해야 합니다.

```js
const testGates1 = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const northGate = testGates1[0];
const originalEastQueue = testGates1[1].queue[0];
rerouteOverflow(testGates1, northGate, 0, 5);
assert.strictEqual(
  testGates1[1].queue[0],
  originalEastQueue + 5,
  "Rerouting from North gate should add overflow to East gate's queue at tickIndex 0"
);

const testGates2 = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const westGate = testGates2[3];
const originalNorthQueue = testGates2[0].queue[2];
rerouteOverflow(testGates2, westGate, 2, 3);
assert.strictEqual(
  testGates2[0].queue[2],
  originalNorthQueue + 3,
  "Rerouting from West gate should wrap around and add overflow to North gate's queue at tickIndex 2"
);

const testGates3 = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const southGate = testGates3[2];
const originalQueues = testGates3.map(gate => [...gate.queue]);
rerouteOverflow(testGates3, southGate, 1, 2);
assert.deepStrictEqual(testGates3[0].queue, originalQueues[0], "North gate queue should not be modified");
assert.deepStrictEqual(testGates3[1].queue, originalQueues[1], "East gate queue should not be modified");
assert.deepStrictEqual(testGates3[2].queue, originalQueues[2], "South gate queue should not be modified");
assert.strictEqual(testGates3[3].queue[1], originalQueues[3][1] + 2, "Only West gate at tickIndex 1 should be modified");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
--fcc-editable-region--

--fcc-editable-region--
}
```
