---
id: 69bac067cd15c4685d07906d
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

이제 개별 눈금을 처리하고 오버플로우를 다루는 함수가 준비되었으니, 하나의 게이트를 제어하는 함수를 만들어 볼 차례입니다. 이 함수는 시뮬레이션의 한 눈금 동안 하나의 게이트에서 일어나는 모든 동작을 처리하게 됩니다.

`handleGateAtTick`이라는 빈 함수를 다음 매개변수와 함께 만드세요:

- `gates`: 모든 게이트 객체가 들어 있는 전체 배열입니다.
- `gate`: 현재 처리 중인 게이트입니다.
- `tickIndex`: `queue` 배열에서 현재 틱의 인덱스입니다.
- `throughputSummary`: 게이트별로 처리된 전체 참가자 수를 추적하는 객체입니다.

# --hints--

`handleGateAtTick`라는 함수를 만드세요.

```js
assert.isFunction(handleGateAtTick);
```

`handleGateAtTick` 함수는 `gates`, `gate`, `tickIndex`, `throughputSummary` 매개변수를 가져야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { handleGateAtTick } = explorer.allFunctions;
const parameters = handleGateAtTick?.parameters;
assert.lengthOf(parameters, 4);
assert.equal(parameters[0].toString(), "gates");
assert.equal(parameters[1].toString(), "gate");
assert.equal(parameters[2].toString(), "tickIndex");
assert.equal(parameters[3].toString(), "throughputSummary");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
