---
id: 69bac027df64d4433c708d44
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

오버플로우를 재경로 지정할 때, 참석자를 `gates` 배열의 다음 게이트로 보내야 합니다. 이를 위해 다음 게이트의 인덱스를 찾아야 합니다. 보통은 현재 인덱스에 1을 더하면 됩니다. 하지만 현재 게이트가 배열의 마지막이라면, 첫 번째 게이트로 다시 돌아가야 합니다.

이럴 때 나머지 연산자(`%`)를 사용할 수 있습니다:

```js
const nextIndex = (currentIndex + 1) % array.length;
```

이 방법이 작동하는 이유는 `currentIndex + 1`가 `array.length`와 같을 때 결과가 0이 되기 때문입니다.

이 방식을 사용해 `nextGateIndex` 배열에서 다음 게이트의 인덱스를 저장하는 `gates`라는 변수를 만드세요.

# --hints--

나머지 연산 기법을 사용해 `currentIndex` 배열에서 `gates`의 다음 인덱스에 접근하고 `nextGateIndex`라는 변수에 할당하세요. 예제와 같은 한 줄 패턴을 따르세요.

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--

--fcc-editable-region--
}
```
