---
id: 69bac18fe3729f6761344365
title: 27단계
challengeType: 1
dashedName: step-27
---

# --description--

시뮬레이션이 몇 눈금 동안 실행되어야 하는지 알아야 합니다. 다음을 기억하세요:

- 각 게이트에는 각 눈금마다 몇 명의 참석자가 도착하는지 보여주는 `queue` 배열이 있습니다.  
- 모든 게이트가 같은 눈금 수를 가지므로 첫 번째 게이트 대기열의 길이를 사용하면 됩니다.  

`maxTicks`라는 변수를 만들고 `gates[0].queue.length` 값을 할당하세요.

이 값은 시뮬레이션 루프를 제어하는 데 사용되어 모든 게이트에 대해 각 눈금이 처리되도록 합니다. 

# --hints--

`maxTicks`라는 변수를 만들고 `gates[0].queue.length` 값을 할당해야 합니다.

```js
const rawFuncStr = __helpers.removeJSComments(simulateFestival.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+maxTicks/,
  "You must have at least one space between the declaration keyword (let/const) and 'maxTicks'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)maxTicks=gates\[0\]\.queue\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `simulateFestival` function, create a variable named `maxTicks` and assign it the value of `gates[0].queue.length`."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
  console.log(
    gate.queue[tickIndex] + " attendees arriving."
  );
  const result = processGateFlow(gate, tickIndex);
  throughputSummary[gate.id] += result.processed;
  if (result.overflow > 0) {
    console.log(
      "Overflow of " + result.overflow +
      " attendees. Rerouting..."
    );
    rerouteOverflow(gates, gate, tickIndex, result.overflow);
  }
}

function printSummary(summary) {
  console.log("\nThroughput Summary");
  for (const gateId in summary) {
    console.log(
      gateId + ": " + summary[gateId] +
      " attendees processed"
    );
  }
}

function simulateFestival(gates, timeBlock) {
  console.log("\n" + timeBlock + " Simulation");
  const throughputSummary = initializeThroughput(gates);
--fcc-editable-region--

--fcc-editable-region--
}
```
