---
id: 69bac1a094d01d84ae4c9e75
title: 28단계
challengeType: 1
dashedName: step-28
---

# --description--

시뮬레이션 중 현재 처리 중인 눈금을 추적하려면 `tickIndex`라는 변수를 만들고 초기값으로 `0`을 할당하세요.  

이 변수는 시뮬레이션의 모든 눈금을 반복하면서 각 게이트가 올바른 순서로 처리되도록 하는 데 사용됩니다.

# --hints--

`tickIndex`라는 변수를 만들고 `0` 값을 할당하세요.

```js
const rawFuncStr = __helpers.removeJSComments(simulateFestival.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+tickIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'tickIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)tickIndex=0;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `simulateFestival` function, create a variable named `tickIndex` and assign it a value of `0`."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
  console.log(
    gate.queue[tickIndex] + " attendees arriving."
  );
  const result = processGateFlow(gate, tickIndex);
  throughputSummary[gate.id] += result.processed;
  if (result.overflow > 0) {
    console.log(
      "Overflow of " + result.overflow +
      " attendees. Rerouting..."
    );
    rerouteOverflow(gates, gate, tickIndex, result.overflow);
  }
}

function printSummary(summary) {
  console.log("\nThroughput Summary");
  for (const gateId in summary) {
    console.log(
      gateId + ": " + summary[gateId] +
      " attendees processed"
    );
  }
}

function simulateFestival(gates, timeBlock) {
  console.log("\n" + timeBlock + " Simulation");
  const throughputSummary = initializeThroughput(gates);
  const maxTicks = gates[0].queue.length;
--fcc-editable-region--

--fcc-editable-region--
}
```
