---
id: 69bab8292750d84cac5b52a0
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이 워크숍에서 처리량은 단일 시뮬레이션(오전 또는 오후) 동안 모든 눈금에서 게이트가 처리한 참석자 총수를 의미합니다. `initializeThroughput` 함수는 각 게이트가 처리한 참석자 수를 추적할 객체를 준비할 것입니다.

먼저 `summary`라는 빈 객체를 만드세요.

# --hints--

`initializeThroughput` 함수는 빈 객체를 생성하고 이를 `summary`라는 변수에 할당해야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { initializeThroughput } = explorer.allFunctions;
const { summary } = initializeThroughput?.variables ?? {};
const validExpressions = ["{}", "new Object()"];
assert.isTrue(
  validExpressions.some((expression) => summary?.value.matches(expression))
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
