---
id: 69babfd65c4c4632a0c76f63
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

게이트가 한 틱 동안 모든 참석자를 처리하지 못하면 일부 참석자가 대기열에 남습니다. 이 오버플로우를 처리하기 위해 오버플로우된 참석자를 다른 게이트로 라우팅하는 함수를 만드세요.

다음 매개변수를 가진 빈 함수를 `rerouteOverflow`라는 이름으로 만드세요:

- `gates`: 게이트 객체 전체 배열입니다.
- `currentGate`: 오버플로우가 발생한 게이트입니다.
- `tickIndex`: 대기열에서 현재 틱 위치입니다.
- `overflowAmount`: 처리하지 못한 참석자 수입니다.

# --hints--

`rerouteOverflow`라는 이름의 함수를 만드세요.

```js
assert.isFunction(rerouteOverflow);
```

`rerouteOverflow` 함수는 `gates`, `currentGate`, `tickIndex`, `overflowAmount` 매개변수를 가져야 합니다.

```js
const regex = __helpers.functionRegex('rerouteOverflow', ['gates', 'currentGate', 'tickIndex', 'overflowAmount']);
assert.match(__helpers.removeJSComments(rerouteOverflow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
