---
id: 69babdf3f896711c11516373
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

각 시뮬레이션 틱 동안:

- 일정 수의 참석자가 게이트에 도착합니다 (해당 게이트의 `queue`에서).
- 게이트는 `capacity`에 따라 참석자를 처리합니다.
- 게이트가 처리할 수 있는 인원보다 더 많은 참석자가 도착하면 일부는 남아 있게 됩니다 (오버플로우).

이제 한 게이트에서 한 틱 동안 이 로직을 처리하는 함수를 만드세요.

`processGateFlow`라는 이름의 빈 함수를 다음 매개변수와 함께 만드세요:

- `gate`: 처리 중인 게이트 객체입니다.
- `tickIndex`: `queue` 배열에서 현재 틱 위치입니다.

# --hints--

`processGateFlow`라는 이름의 함수를 만드세요.

```js
assert.isFunction(processGateFlow);
```

`processGateFlow` 함수는 `gate`와 `tickIndex` 매개변수를 가져야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { processGateFlow } = explorer.allFunctions;
const parameters = processGateFlow?.parameters;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), "gate");
assert.equal(parameters[1].toString(), "tickIndex");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
