---
id: 69babea99725b43a01adb6d1
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

먼저, 현재 틱 동안 도착하는 참석자 수를 구하세요. 입력 게이트의 `queue` 배열에서 인덱스 `tickIndex`에 있는 값을 접근하여 `currentTickQueue`라는 변수에 할당하세요.

# --hints--

`gate.queue[tickIndex]`에 있는 요소를 `currentTickQueue`라는 변수에 할당해야 합니다.

```js
const rawFuncStr = __helpers.removeJSComments(processGateFlow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+currentTickQueue/,
  "You must have at least one space between the declaration keyword (let/const) and 'currentTickQueue'."
);

const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(processGateFlow.toString()));
const solution = /(var|const|let)currentTickQueue=gate\.queue\[tickIndex\];?}$/;

assert(
  solution.test(cleaned),
  "Inside your `processGateFlow` function, assign the element at `gate.queue[tickIndex]` to a variable named `currentTickQueue` as the only statement inside the function."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
--fcc-editable-region--

--fcc-editable-region--
}
```
