---
id: 69bac0e7112d02531998dbf0
title: 20단계
challengeType: 1
dashedName: step-20
---

# --description--

`handleGateAtTick` 함수에서 마지막 작업은 현재 틱에 대한 오버플로우를 처리하는 것입니다.

오버플로우가 존재하면(`result.overflow > 0`), 다음을 수행하세요:

- 다음 메시지를 기록하세요 `"Overflow of " + result.overflow + " attendees. Rerouting..."`.
- `rerouteOverflow()` 함수를 `gates`, `gate`, `tickIndex`, `result.overflow`와 함께 호출하여 오버플로우를 다시 라우트하세요.

# --hints--

오버플로우가 존재하면(`result.overflow > 0`), 먼저 메시지 `"Overflow of " + result.overflow + " attendees. Rerouting..."`를 기록하세요.

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];

const throughputSummary = { North: 0, East: 0, South: 0, West: 0 };
const spy = __helpers.spyOn(console, "log");
try {
  const westGate = testGates[3];
  handleGateAtTick(testGates, westGate, 0, throughputSummary);
  const overflowMessage = spy.calls.find(call => call[0].includes("Overflow of"));
  assert.isTrue(
    typeof overflowMessage !== "undefined",
    "handleGateAtTick should log an overflow message when overflow exists"
  );
  assert.match(
    overflowMessage[0],
    /Overflow of \d+ attendees\. Rerouting\.\.\.$/,
    "Overflow message should match the format 'Overflow of X attendees. Rerouting...'"
  );
} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

오버플로우가 존재하면(`result.overflow > 0`), `rerouteOverflow()`을 `gates`, `gate`, `tickIndex`, `result.overflow`와 함께 호출하여 오버플로우를 다시 라우트하세요.

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const throughputSummary = { North: 0, East: 0, South: 0, West: 0 };
const originalRerouteOverflow = rerouteOverflow;
let called = false;
let callArgs = null;
try {
  rerouteOverflow = function (...args) {
    called = true;
    callArgs = args;
    return originalRerouteOverflow(...args);
  };
  const consoleSpy = __helpers.spyOn(console, "log");
  try {
    const westGate = testGates[3];
    const originalNorthQueue = testGates[0].queue[0];
    handleGateAtTick(testGates, westGate, 0, throughputSummary);
    assert.isTrue(
      called,
      "handleGateAtTick should call rerouteOverflow() when overflow exists"
    );
    assert.deepEqual(
      callArgs,
      [testGates, westGate, 0, 2],
      "rerouteOverflow should be called with gates, gate, tickIndex, and overflow amount"
    );
    assert.strictEqual(
      testGates[0].queue[0],
      originalNorthQueue + 2,
      "Overflow from West should be rerouted to North gate's queue (next gate with wrap-around)"
    );
  } finally {
    consoleSpy.restore();
  }
} catch (err) {
  assert.fail(err);
} finally {
  rerouteOverflow = originalRerouteOverflow;
}
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
  console.log(
    gate.queue[tickIndex] + " attendees arriving."
  );
  const result = processGateFlow(gate, tickIndex);
  throughputSummary[gate.id] += result.processed;
--fcc-editable-region--

--fcc-editable-region--
}
```
