---
id: 69bac067cd15c4685d07906d
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

개별 눈금을 처리하고 오버플로우를 다루는 함수를 만들었으니, 이제 단일 게이트를 위한 제어 함수 하나를 만드세요. 이 함수는 시뮬레이션의 한 눈금 동안 한 게이트에 대한 모든 작업을 처리할 것입니다.

다음 매개변수를 가진 빈 함수 `handleGateAtTick`를 만드세요:

- `gates`: 게이트 객체 전체 배열
- `gate`: 현재 처리 중인 게이트
- `tickIndex`: 시뮬레이션의 현재 눈금 인덱스
- `throughputSummary`: 게이트별 총 처리된 참석자 수를 추적하는 객체

# --hints--

`handleGateAtTick`라는 함수를 만드세요.

```js
assert.isFunction(handleGateAtTick);
```

`handleGateAtTick` 함수는 `gates`, `gate`, `tickIndex`, `throughputSummary` 매개변수를 가져야 합니다.

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
