---
id: 69cfdc80443b715a8aa452a7
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

`initializeThroughput` 함수를 완성하려면:

- 입력 `gates` 배열의 각 게이트를 순차적으로 반복하세요. 각 게이트마다 `summary`에 다음과 같은 속성을 추가하세요:
  - 키는 게이트의 `id`입니다.
  - 값은 `0`입니다.
- 반복이 끝난 후 `summary` 객체를 반환하세요.

`initializeThroughput` 함수는 원본 입력 배열 `gates`을 변경하지 않아야 합니다.

# --hints--

`initializeThroughput()`을 인수로 `[]`을 호출하면 원본 입력을 변경하지 않고 `{}`을 반환해야 합니다.

```js
const input = [];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  {},
  "Calling `initializeThroughput([])` should return an empty object."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given an empty array."
);
```

`initializeThroughput()`을 인수로 `[{ id: "North", capacity: 5, queue: [1, 2, 3] }]`을 호출하면 원본 입력을 변경하지 않고 `{ North: 0 }`을 반환해야 합니다.

```js
const input = [{ id: "North", capacity: 5, queue: [1, 2, 3] }];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0 },
  "Calling `initializeThroughput([{ id: 'North', ... }])` should return `{ North: 0 }`."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given one gate."
);
```

`initializeThroughput()`을 인수로 `[{ id: "North", capacity: 5, queue: [] }, { id: "East", capacity: 3, queue: [] }, { id: "South", capacity: 4, queue: [] }]`를 호출하면 원본 입력을 변경하지 않고 `{ North: 0, East: 0, South: 0}`을 반환해야 합니다.

```js
const input = [
  { id: "North", capacity: 5, queue: [] },
  { id: "East", capacity: 3, queue: [] },
  { id: "South", capacity: 4, queue: [] }
];

const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0, East: 0, South: 0 },
  "Calling `initializeThroughput` with multiple gates should return an object with an entry for each gate id set to 0."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given multiple gates."
);
```

입력 `for` 배열 안의 각 게이트 객체를 순차적으로 반복하려면 `gates` 반복문을 사용하세요.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));
assert.match(cleaned, /for\s*\(/);
```

반복이 끝난 후 `summary`를 반환하세요. 이 문장은 함수의 마지막 줄이어야 합니다.

```js
const funcStr = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));

assert.match(
  funcStr,
  /returnsummary;?}$/,
  "Your `initializeThroughput` function should end with `return summary`."
);
```

`initializeThroughput` 함수는 원본 입력 배열 `gates`을 변경하지 않아야 합니다.

```js
const morningCopy = JSON.parse(JSON.stringify(morningGates));
const nightCopy = JSON.parse(JSON.stringify(nightGates));

initializeThroughput(morningGates);
initializeThroughput(nightGates);

assert.deepEqual(
  morningGates,
  morningCopy,
  "`initializeThroughput` should not mutate the `morningGates` seed array."
);

assert.deepEqual(
  nightGates,
  nightCopy,
  "`initializeThroughput` should not mutate the `nightGates` seed array."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
--fcc-editable-region--

--fcc-editable-region--
}
```
