---
id: 69d048937c7bb4b4c22ebb18
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

이제 참석자를 게이트를 통해 처리할 시간입니다. `while` 반복문 안에서:

- 감소 연산자(`currentTickQueue`)를 사용해 `--`을 1만큼 감소시켜 한 명의 참석자가 통과했음을 표시하세요.
- 증가 연산자(`processed`)를 사용해 `++`을 1만큼 증가시켜 처리된 참석자 수를 추적하세요.

# --hints--

당신의 `while` 루프는 감소 연산자 `--`를 사용해 `currentTickQueue`를 감소시키고 증가 연산자 `++`를 사용해 `processed`를 증가시켜야 합니다.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const viableSolutions = [
  "while(currentTickQueue>0&&processed<gate.capacity){processed++;currentTickQueue--;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){processed++currentTickQueue--;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){processed++;currentTickQueue--}}",
  "while(currentTickQueue>0&&processed<gate.capacity){processed++currentTickQueue--}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--;processed++;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--processed++;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--;processed++}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--processed++}}"
];
const endsCorrectly = viableSolutions.some(sol => cleaned.endsWith(sol));
assert(
  endsCorrectly,
  "Your `while` loop should contain both the `processed++` increment and the `currentTickQueue--` decrement, and nothing else."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
