---
id: 68c341332d98fe38fae733af
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Jetzt ist es Zeit, die Tabelle weiter aufzubauen.

Beginnen Sie damit, unter Ihrem `table`-Element ein `h1`-Element hinzuzufügen. Fügen Sie innerhalb des `table`-Elements ein `caption`-Element mit dem Text `Schedule by Track and Time` hinzu.

# --hints--

Sie sollten ein `table`-Element haben.

```js
assert.exists(document.querySelector("table"));
```

Ihr `table`-Element sollte ein `caption`-Element enthalten.

```js
assert.exists(document.querySelector("table caption"));
```

Ihr `caption`-Element sollte den Text `Schedule by Track and Time` enthalten.

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
