---
id: 68c34f4aeb2b1a4170506751
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Fügen Sie als Nächstes eine fünfte Zeile zur Tabelle hinzu. Beginnen Sie damit, ein weiteres `tr`-Element hinzuzufügen. Fügen Sie innerhalb dieses `tr`-Elements ein `th`-Element mit einem `scope`-Attribut hinzu, das auf `"row"` gesetzt ist, und dem Textinhalt `12:30 PM`. Dann fügen Sie unterhalb dieses `th`-Elements ein `td`-Element mit dem Textinhalt `Lunch Break` hinzu. Ihr `td`-Element sollte außerdem ein `colspan`-Attribut mit dem Wert `3` haben, damit es sich über alle drei Spuren erstreckt.

# --hints--

Sie sollten ein fünftes `tr`-Element innerhalb des `tbody`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

Sie sollten ein `th`-Element innerhalb des fünften `tr`-Elements haben.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

Ihr `th`-Element sollte den Textinhalt `12:30 PM` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

Ihr `th`-Element sollte ein `scope`-Attribut mit dem Wert `"row"` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

Sie sollten ein `td`-Element innerhalb des fünften `tr`-Elements haben.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

Ihr `td`-Element sollte den Textinhalt `Lunch Break` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

Ihr `td`-Element sollte ein `colspan`-Attribut mit dem Wert `3` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
