---
id: 68c34e201b347c40acda8111
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Jetzt ist es Zeit, der Tabelle eine weitere Zeile hinzuzufügen.

Beginnen Sie damit, ein weiteres `tr`-Element hinzuzufügen. Fügen Sie innerhalb dieses `tr`-Elements ein `th`-Element mit einem `scope`-Attribut hinzu, das auf `"row"` gesetzt ist, und dem Textinhalt `11:30 AM`.

Fügen Sie dann drei `td`-Elemente mit folgendem Textinhalt hinzu:

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

Sie sollten ein viertes `tr`-Element innerhalb des `tbody`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

Sie sollten ein `th`-Element innerhalb des vierten `tr`-Elements haben.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

Ihr `th`-Element sollte den Textinhalt `11:30 AM` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

Ihr `th`-Element sollte ein `scope`-Attribut mit dem Wert `"row"` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

Ihr viertes `tr`-Element sollte drei `td`-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

Ihr erstes `td`-Element sollte den Textinhalt `AR/VR in Education` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

Ihr zweites `td`-Element sollte den Textinhalt `JavaScript Fundamentals` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

Ihr drittes `td`-Element sollte den Textinhalt `Design Systems at Scale` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
