---
id: 68c34c608cf9e53f5d36b1d5
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Fügen Sie als Nächstes eine dritte Zeile zur Tabelle hinzu. Beginnen Sie damit, ein weiteres `tr`-Element hinzuzufügen. Fügen Sie innerhalb dieses `tr`-Elements ein `th`-Element mit einem `scope`-Attribut hinzu, das auf `"row"` gesetzt ist, und dem Textinhalt von `11:00 AM`. Dann fügen Sie unter diesem `th`-Element ein `td`-Element mit dem Textinhalt von `Break` hinzu.

# --hints--

Sie sollten ein drittes `tr`-Element innerhalb des `tbody`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

Sie sollten ein `th`-Element innerhalb des dritten `tr`-Elements haben.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

Ihr `th`-Element sollte den Textinhalt von `11:00 AM` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

Ihr `th`-Element sollte ein `scope`-Attribut haben, das auf `"row"` gesetzt ist.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

Sie sollten ein `td`-Element innerhalb des dritten `tr`-Elements haben.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

Ihr `td`-Element sollte den Textinhalt von `Break` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
