---
id: 68c3482278a14e3cf225e897
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Fügen Sie innerhalb Ihres `tr`-Elements ein `th`-Element mit dem Text von `9:00 AM` hinzu. Dann fügen Sie unter diesem `th`-Element drei `td`-Elemente mit dem Text von `Keynote: Tech Future`, `Intro to Web Dev` und `UX for All` hinzu.

# --hints--

Sie sollten ein `th`-Element innerhalb Ihres `tr`-Elements haben.

```js
assert.exists(document.querySelector("tbody tr th"));
```

Ihr `th`-Element sollte den Textinhalt von `9:00 AM` haben.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

Sie sollten insgesamt drei `td`-Elemente unter Ihrem `th`-Element haben.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

Ihr erstes `td`-Element sollte den Textinhalt von `Keynote: Tech Future` haben.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

Ihr zweites `td`-Element sollte den Textinhalt von `Intro to Web Dev` haben.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

Ihr drittes `td`-Element sollte den Textinhalt von `UX for All` haben.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
