---
id: 68c34329d430583a6cd5f7f1
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Fügen Sie innerhalb Ihres `tr`-Elements vier `th`-Elemente hinzu. Das erste `th`-Element sollte den Text von `Time` enthalten, das zweite den Text von `Track A`, das dritte den Text von `Track B` und das vierte den Text von `Track C`.

# --hints--

Ihr `tr`-Element sollte vier `th`-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

Ihr erstes `th`-Element sollte den Text von `Time` enthalten.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

Ihr zweites `th`-Element sollte den Text von `Track A` enthalten.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

Ihr drittes `th`-Element sollte den Text von `Track B` enthalten.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

Ihr viertes `th`-Element sollte den Text von `Track C` enthalten.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
