---
id: 68c34d57ed874e4004f7b5d4
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Im Moment erstreckt sich das `td`-Element mit dem Textinhalt `Break` nur über eine Spalte. Es wäre jedoch schön, wenn es sich über alle drei Spalten erstrecken würde.

Wie Sie aus früheren Workshops und Lektionen wissen, können Sie das `colspan`-Attribut verwenden, um eine Tabellenzelle über mehrere Spalten zu erstrecken.

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

Fügen Sie dem `td`-Element ein `colspan`-Attribut hinzu und setzen Sie dessen Wert auf `3`.

# --hints--

Ihr `td`-Element mit dem Textinhalt `Break` sollte ein `colspan`-Attribut mit dem Wert `3` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
