---
id: 68c343692285523aef0fc6bf
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Das Attribut `scope` wird verwendet, um anzugeben, ob eine Kopfzeile eine Kopfzeile für eine Zeile, Spalte oder Gruppe von Zeilen oder Spalten ist. Hier ist ein Beispiel:

```html
<th scope="col">Example Header</th>
```

Dies hilft Screenreadern, die Beziehung zwischen Kopf- und Datenzellen zu verstehen.

Für alle `th`-Elemente fügen Sie ein `scope`-Attribut mit dem Wert `col` hinzu.

# --hints--

Jedes `th`-Element sollte ein `scope`-Attribut mit dem Wert `"col"` haben.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
