---
id: 68c3420cb0176439a95af1b8
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Der nächste Schritt ist das Hinzufügen des Tabellenkopfs.

Beginnen Sie damit, ein `thead`-Element unter dem `caption`-Element einzufügen. Fügen Sie innerhalb des `thead` ein `tr`-Element hinzu.

# --hints--

Sie sollten ein `thead`-Element haben.

```js
assert.exists(document.querySelector("table thead"));
```

Sie sollten ein `tr`-Element innerhalb des `thead` haben.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
