---
id: 68c33e502cd97f2fd52cb63b
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop lernen Sie, wie Sie zugängliche Tabellenelemente erstellen, indem Sie einen Zeitplan für eine Tech-Konferenz anlegen.

Beginnen Sie damit, ein `h1`-Element mit dem Text `Tech Conference 2025 Schedule` hinzuzufügen.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

Ihr `h1`-Element sollte den Text `Tech Conference 2025 Schedule` enthalten.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
