---
id: 68c34ace730ca43e5fa7b8d6
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Jetzt ist es Zeit, eine weitere Zeile zur Tabelle hinzuzufügen.

Beginnen Sie damit, ein weiteres `tr`-Element hinzuzufügen. Fügen Sie innerhalb dieses `tr`-Elements ein `th`-Element mit einem `scope`-Attribut hinzu, das auf `"row"` gesetzt ist, und dem Textinhalt `10:00 AM`.

Fügen Sie dann drei `td`-Elemente mit folgendem Textinhalt hinzu:

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

Sie sollten ein zweites `tr`-Element innerhalb des `tbody`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

Sie sollten ein `th`-Element innerhalb des zweiten `tr`-Elements haben.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

Ihr `th`-Element sollte den Textinhalt `10:00 AM` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

Ihr `th`-Element sollte ein `scope`-Attribut mit dem Wert `"row"` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

Ihr zweites `tr`-Element sollte drei `td`-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

Ihr erstes `td`-Element sollte den Textinhalt `Accessibility Deep Dive` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

Ihr zweites `td`-Element sollte den Textinhalt `CSS for Beginners` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

Ihr drittes `td`-Element sollte den Textinhalt `Inclusive Design Principles` haben.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
