---
id: 68c346d03709343bff132707
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Für die nächsten Schritte bauen Sie den Körper der Tabelle auf.

Beginnen Sie damit, ein `tbody`-Element unter dem `thead`-Element hinzuzufügen. Fügen Sie dann innerhalb des `tbody`-Elements ein `tr`-Element hinzu.

# --hints--

Sie sollten ein `tbody`-Element unter dem `thead`-Element haben.

```js
assert.exists(document.querySelector("tbody"));
```

Sie sollten ein `tr`-Element innerhalb des `tbody`-Elements haben.

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
