---
id: 60a3e3396c7b40068ad69971
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

حان وقت CSS.

رغم أن `<div>` الخاص بك لا يحتوي على نص، فإنه يُعامل كصندوق يحتوي على محتوى.
اكتب قاعدة CSS تستخدم محدد الصنف `.canvas` واضبط خاصية `width` على 500 بكسل.
إليك قاعدة CSS تضبط عرض الصنف `card` إلى 300 بكسل:

```css
.card {
  width: 300px;
}
```

# --hints--

يجب أن يحتوي شفرتك على محدد `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

يجب ضبط خاصية `width` على `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

يجب أن يكون لعنصر `.canvas` عرض `width` مقداره `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
