---
id: 60a3e3396c7b40068ad6996e
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

الهامش هو المساحة خارج الصندوق، ويمكن استخدامه للتحكم في المسافة بين الصناديق أو العناصر الأخرى.

هنا العنصر السفلي لديه هامش علوي أكبر، مما يدفعه إلى أسفل الصفحة أكثر.

بعد أن راجعت بسرعة نموذج صندوق CSS، حان الوقت للبدء في لوحة روثكو.

احذف العنصر `img`.

# --hints--

يجب ألا يكون لديك عنصر `img` في شفرتك.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
