---
id: 60a3e3396c7b40068ad6996c
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

في الدروس السابقة عن نموذج الصندوق في CSS، تعلمت أن كل عنصر HTML يُعامل كصندوق يحتوي على أربع مساحات.

تخيل أنك تستلم صندوقًا من متجر الإنترنت المفضل لديك — المحتوى هو العنصر داخل الصندوق، وفي حالتك، يكون قسمًا علويًا أو فقرة أو عنصر صورة.

غيّر الخاصية `src` في `<img>` من `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` إلى `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

يجب أن يحتوي عنصر `img` على خاصية `src` بقيمة `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

