---
id: 60a3e3396c7b40068ad6996d
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

يُحيط بالمحتوى مساحة تُسمى الحشو، تشبه الطريقة التي يفصل بها التغليف الفقاعي العنصر عن الصندوق المحيط به.

فكر في الحدود مثل صندوق الكرتون الذي شُحن فيه العنصر الخاص بك.

غيّر الخاصية `src` إلى `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

يجب أن يحتوي عنصر `img` على خاصية `src` بقيمة `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
