---
id: 60a3e3396c7b40068ad6996b
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ابدأ مشروعك بإضافة عنصر `img` مع خاصية `src` بقيمة `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` داخل عنصر `body`.

# --hints--

يجب أن تحتوي شفرتك على عنصر `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

يجب أن يحتوي عنصر `img` على خاصية `src` بقيمة `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
