---
id: 60a3e3396c7b40068ad6997a
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

اكتب قاعدة جديدة تستهدف `.one` وحدد `width` بقيمة 425 بكسل.

# --hints--

يجب أن يكون لديك محدِّد `.one`.

```js
const hasOne = new __helpers.CSSHelp(document).getStyle('.one');
assert(hasOne);
```

يجب أن تحدد خاصية `width` إلى `425px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '425px');
assert(hasWidth);
```

يجب أن يكون لعنصر `.one` قيمة `width` تساوي `425px`.

```js
const oneWidth = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('width');
assert(oneWidth === '425px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
