---
id: 68c34329d430583a6cd5f7f1
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Всередині вашого елемента `tr` додайте чотири елементи `th`. Перший елемент `th` має містити текст `Time`, другий — текст `Track A`, третій — текст `Track B`, а четвертий — текст `Track C`.

# --hints--

Ваш елемент `tr` має містити чотири елементи `th`.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

Перший елемент `th` має містити текст `Time`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

Другий елемент `th` має містити текст `Track A`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

Третій елемент `th` має містити текст `Track B`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

Четвертий елемент `th` має містити текст `Track C`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
