---
id: 68c34d57ed874e4004f7b5d4
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Зараз елемент `td` з текстовим вмістом `Break` займає лише одну колонку. Але було б добре, якби він займав усі три колонки.

Як ви пам’ятаєте з попередніх практичних занять і уроків, ви можете використати атрибут `colspan`, щоб зробити так, щоб клітинка таблиці охоплювала кілька колонок.

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

Додайте атрибут `colspan` до елемента `td` і встановіть його значення `3`.

# --hints--

Ваш елемент `td` з текстовим вмістом `Break` має мати атрибут `colspan`, встановлений у `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
