---
id: 68c33e502cd97f2fd52cb63b
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви навчитеся створювати доступні елементи таблиці, створюючи розклад для технічної конференції.

Почніть із додавання елемента `h1` з текстом `Tech Conference 2025 Schedule`.

# --hints--

У вас має бути елемент `h1`.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

Ваш елемент `h1` повинен містити текст `Tech Conference 2025 Schedule`.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
