---
id: 68c34c608cf9e53f5d36b1d5
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Далі додайте третій рядок до таблиці. Почніть із додавання ще одного `tr` елемента. Всередині цього `tr` елемента додайте `th` елемент з атрибутом `scope`, встановленим у `"row"`, та текстовим вмістом `11:00 AM`. Потім під цим `th` елементом додайте `td` елемент із текстовим вмістом `Break`.

# --hints--

У вас має бути третій `tr` елемент всередині `tbody` елемента.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

У третьому `th` елементі має бути `tr` елемент.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

Ваш `th` елемент має містити текстовий вміст `11:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

Ваш `th` елемент має мати атрибут `scope`, встановлений у `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

У третьому `td` елементі має бути `tr` елемент.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

Ваш `td` елемент має містити текстовий вміст `Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
