---
id: 68c34e201b347c40acda8111
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Тепер настав час додати ще один рядок до таблиці.

Почніть із додавання ще одного `tr` елемента. Всередині цього `tr` елемента додайте `th` елемент з `scope` атрибутом, встановленим у `"row"`, і текстовим вмістом `11:30 AM`.

Потім додайте три `td` елементи з таким текстовим вмістом:

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

У вас має бути четвертий `tr` елемент всередині `tbody` елемента.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

У четвертому `th` елементі має бути `tr` елемент.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

Ваш `th` елемент має містити текстовий вміст `11:30 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

Ваш `th` елемент має мати `scope` атрибут, встановлений у `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

Ваш четвертий `tr` елемент має містити три `td` елементи.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

Перший `td` елемент має містити текстовий вміст `AR/VR in Education`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

Другий `td` елемент має містити текстовий вміст `JavaScript Fundamentals`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

Третій `td` елемент має містити текстовий вміст `Design Systems at Scale`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
