---
id: 68c346d03709343bff132707
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

У наступних кількох кроках ви будете створювати тіло таблиці.

Почніть із додавання елемента `tbody` під елементом `thead`. Потім всередині `tbody` додайте елемент `tr`.

# --hints--

У вас має бути елемент `tbody` під елементом `thead`.

```js
assert.exists(document.querySelector("tbody"));
```

У вас має бути елемент `tr` всередині елемента `tbody`.

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
