---
id: 68c34ace730ca43e5fa7b8d6
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Тепер настав час додати ще один рядок до таблиці.

Почніть із додавання ще одного `tr` елемента. Всередині цього `tr` елемента додайте `th` елемент з `scope` атрибутом, встановленим у `"row"`, та текстовим вмістом `10:00 AM`.

Потім додайте три `td` елементи з таким текстовим вмістом:

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

У вас має бути другий `tr` елемент всередині `tbody` елемента.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

У другому `th` елементі має бути `tr` елемент.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

Ваш `th` елемент має містити текстовий вміст `10:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

Ваш `th` елемент має мати `scope` атрибут, встановлений у `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

Ваш другий `tr` елемент має містити три `td` елементи.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

Ваш перший `td` елемент має містити текстовий вміст `Accessibility Deep Dive`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

Ваш другий `td` елемент має містити текстовий вміст `CSS for Beginners`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

Ваш третій `td` елемент має містити текстовий вміст `Inclusive Design Principles`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
