---
id: 68c343692285523aef0fc6bf
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Атрибут `scope` використовується, щоб вказати, чи є клітинка заголовком для рядка, стовпця або групи рядків чи стовпців. Ось приклад:

```html
<th scope="col">Example Header</th>
```

Це допомагає читачам екрана зрозуміти зв’язок між заголовками та клітинками з даними.

Для всіх елементів `th` додайте атрибут `scope` зі значенням `col`.

# --hints--

Кожен елемент `th` повинен мати атрибут `scope` зі значенням `"col"`.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
