---
id: 68c3420cb0176439a95af1b8
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Наступний крок — додати заголовок таблиці.

Почніть із додавання елемента `thead` під елементом `caption`. Всередині `thead` додайте елемент `tr`.

# --hints--

У вас має бути елемент `thead`.

```js
assert.exists(document.querySelector("table thead"));
```

Усередині `tr` має бути елемент `thead`.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
