---
id: 68c3482278a14e3cf225e897
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Всередині вашого елемента `tr` додайте елемент `th` з текстом `9:00 AM`. Потім під цим елементом `th` додайте три елементи `td` з текстом `Keynote: Tech Future`, `Intro to Web Dev` та `UX for All`. 

# --hints--

У вашому елементі `th` має бути елемент `tr`.

```js
assert.exists(document.querySelector("tbody tr th"));
```

Ваш елемент `th` має містити текст `9:00 AM`.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

Під вашим елементом `td` має бути загалом три елементи `th`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

Перший елемент `td` має містити текст `Keynote: Tech Future`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

Другий елемент `td` має містити текст `Intro to Web Dev`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

Третій елемент `td` має містити текст `UX for All`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
