---
id: 68c34f4aeb2b1a4170506751
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Далі додайте п’ятий рядок до таблиці. Почніть із додавання ще одного `tr` елемента. Всередині цього `tr` елемента додайте `th` елемент з атрибутом `scope`, встановленим у `"row"`, і текстовим вмістом `12:30 PM`. Потім під цим `th` елементом додайте `td` елемент з текстовим вмістом `Lunch Break`. Ваш `td` елемент також повинен мати атрибут `colspan`, встановлений у `3`, щоб він охоплював усі три колонки.

# --hints--

У вас має бути п’ятий `tr` елемент всередині `tbody` елемента.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

У п’ятому `th` елементі має бути `tr` елемент.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

Ваш `th` елемент має містити текст `12:30 PM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

Ваш `th` елемент має мати атрибут `scope`, встановлений у `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

У п’ятому `td` елементі має бути `tr` елемент.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

Ваш `td` елемент має містити текст `Lunch Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

Ваш `td` елемент має мати атрибут `colspan`, встановлений у `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
