---
id: 68c341332d98fe38fae733af
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Тепер настав час почати створювати таблицю.

Почніть із додавання елемента `table` під вашим елементом `h1`. Всередині елемента `table` додайте елемент `caption` з текстом `Schedule by Track and Time`.

# --hints--

У вас має бути елемент `table`.

```js
assert.exists(document.querySelector("table"));
```

Ваш елемент `table` має містити елемент `caption`.

```js
assert.exists(document.querySelector("table caption"));
```

Ваш елемент `caption` має містити текст `Schedule by Track and Time`.

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
