---
id: 68f1f993460f06a64c033061
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Атрибут `poster` — це мініатюрне зображення відео. Уявіть собі відео, які ви дивитесь на YouTube. Воно відображається під час завантаження відео. Якщо атрибут пропущено, під час завантаження показується перший кадр відео.

Тепер додайте атрибут `poster` зі значенням `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` до вашого елемента `video`.

# --hints--

Ваш елемент `video` має містити атрибут `poster` зі значенням `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
