---
id: 68f1fac259a4afabb478b960
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Ви, можливо, помітили, що не направили на справжнє відео. Ви зробите це на наступному кроці. Що стосується типів відеофайлів, існують відмінності у підтримці браузерами. Щоб це врахувати, ви можете використовувати `source` елементи всередині `video` елемента, і браузер вибере перший сумісний `source`.

Ось приклад `source` елемента:

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

`source` елемент є порожнім елементом, тому він не має кінцевого тегу.

Додайте `source` елемент всередину вашого `video` елемента.


# --hints--

У вас має бути `source` елемент всередині вашого `video` елемента.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
