---
id: 68f1f57723debf9383c75947
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Атрибут `loop` перезапустить відео після завершення відтворення. Уявіть інтернет-мем, який повторює відтворення. Якщо опустити атрибут `loop`, відео відтвориться лише один раз.

Атрибут `loop` є булевим і не потребує значення.

Додайте атрибут `loop` до елемента `video`.

# --hints--

Ваш елемент `video` має містити атрибут `loop`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
