---
id: 68f1f8c248306da204655117
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Атрибут `muted` вимкне звук під час початкового відтворення. Якщо у вас увімкнено `controls`, користувач зможе увімкнути звук. Якщо опустити атрибут `muted`, звук відтворюватиметься під час початкового відтворення.

Атрибут `muted` є булевим атрибутом і не потребує значення.

Додайте атрибут `muted` до елемента `video`.

# --hints--

Ваш елемент `video` має містити атрибут `muted`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
