---
id: 68f1f6f97af9749a42260a43
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Атрибут `controls` надає елементові `video` елементи керування відтворенням, включно з відтворенням, перемотуванням і регулюванням гучності.

Атрибут `controls` є булевим атрибутом і не потребує значення.

Додайте атрибут `controls` до елемента `video`.

Тепер ви повинні побачити елемент `video`, відображений на сторінці.

# --hints--

Ваш елемент `video` має мати атрибут `controls`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
