---
id: 68f1eb9ad0a7575e4dbc31d5
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Далі створіть елемент `video` під елементом `h1`. Протягом наступних кількох кроків ви додасте необхідні атрибути, щоб зробити відеопрогравач функціональним.

# --hints--

У вас має бути елемент `video`.

```js
assert.exists(document.querySelector('video'));
```

Новий елемент `video` має бути розташований під елементом `h1`.

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
