---
id: 68f1f371fefd448b4f70b898
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У попередньому уроці ви дізналися про різні атрибути, доступні для елемента `video`. Атрибут `width` визначає ширину відео в пікселях.

Додайте атрибут `width` до елемента `video` зі значенням `640`.

# --hints--

Ваш елемент `video` має містити атрибут `width` зі значенням `640`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
