---
id: 68f27139bddf9f6bf2821acf
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Останній елемент `source`, який ви додасте, буде для MIME-типу `video/quicktime`.

QuickTime — це розширювана мультимедійна архітектура, створена Apple, яка підтримує відтворення, трансляцію, кодування та перекодування різних цифрових форматів медіаконтенту. Хоча вона не така популярна, як формат MP4, вона може знадобитися для підтримки застарілих застосунків.

Під вашим третім елементом `source` додайте четвертий елемент `source` і надайте йому атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` та атрибут `type` зі значенням `video/quicktime`.

Вітаємо! Ви завершили практичне заняття з HTML Video Player.

# --hints--

У вас має бути четвертий елемент `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

Ваш четвертий елемент `source` має мати атрибут `src`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

Ваш четвертий елемент `source` має мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

Ваш четвертий елемент `source` має мати атрибут `type`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

Ваш четвертий елемент `source` має мати атрибут `type` зі значенням `video/quicktime`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
