---
id: 68f26d52b13e0f5b6e2d2e09
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Ще одним загальним MIME-типом є `video/webm` MIME-тип.

WebM — це відкритий формат аудіовізуальних медіафайлів, розроблений Google, який переважно призначений для веб-медіаконтенту. Він підтримує відеокодеки, такі як VP8, VP9 і AV1, а також аудіокодеки Vorbis і Opus, що робить його популярним вибором для елементів відео та аудіо в HTML5.

Під вашим першим елементом `source` додайте ще один елемент `source` і надайте йому атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` та атрибут `type` зі значенням `video/webm`.

# --hints--

У вас має бути другий елемент `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

Ваш другий елемент `source` має мати атрибут `src`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

Ваш другий елемент `source` має мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

Ваш другий елемент `source` має мати атрибут `type`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

Ваш другий елемент `source` має мати атрибут `type` зі значенням `video/webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
