---
id: 68f2701b2160c765f23b60f8
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Ще один загальний MIME-тип — це `video/ogg` MIME-тип.

Ogg — це цифровий мультимедійний контейнер, створений для ефективного потокового передавання та обробки цифрового медіаконтенту. Його підтримує Xiph.Org Foundation, і він є вільним та відкритим, без обмежень програмними патентами. Назва походить від "ogging" — жаргону з комп’ютерної гри Netrek.

Під вашим другим `source` елементом додайте третій `source` елемент і надайте йому атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` та атрибут `type` зі значенням `video/ogg`.

# --hints--

У вас має бути третій `source` елемент.

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

Ваш третій `source` елемент має мати атрибут `src`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

Ваш третій `source` елемент має мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

Ваш третій `source` елемент має мати атрибут `type`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

Ваш третій `source` елемент має мати атрибут `type` зі значенням `video/ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
