---
id: 68f26a5f79e97c5032c10642
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Ви використали відеофайл з розширенням `mp4`, і потрібно повідомити браузеру про це, щоб він знав, як читати цей файл.

Ви використаєте атрибут `type`, щоб вказати MIME-тип `video/mp4`.

MIME (Multipurpose Internet Mail Extensions) — це стандарт для опису документів у формах, відмінних від ASCII-тексту, наприклад, аудіо, відео та зображень.

MP4, офіційно відомий як MPEG-4 Part 14, — це цифровий мультимедійний контейнер. Він широко використовується для зберігання відео та аудіо, але також може містити інші типи даних, як-от субтитри та статичні зображення. Файли MP4 розроблені для потокового передавання через Інтернет і сумісні з багатьма пристроями та платформами.

Тепер додайте атрибут `type` зі значенням `video/mp4`.

# --hints--

Ваш елемент `source` має містити атрибут `type`.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

Ваш елемент `source` має містити атрибут `type` зі значенням `video/mp4`.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
