---
id: 61537a8054753e2f1f2a1574
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie innerhalb Ihres `head`-Elements einen `meta`-Tag mit dem `name`, der auf `viewport` gesetzt ist, und dem `content`, der auf `width=device-width, initial-scale=1` gesetzt ist, hinzu.

Fügen Sie außerdem einen `meta`-Tag mit dem `charset` hinzu, der auf `UTF-8` gesetzt ist.

# --hints--

Sie sollten zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Ein `meta`-Element sollte ein `name` haben, das auf `viewport` gesetzt ist, und ein `content`, das auf `width=device-width, initial-scale=1.0` gesetzt ist.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

Ihr anderes `meta`-Element sollte das `charset`-Attribut haben, das auf `UTF-8` gesetzt ist.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
