---
id: 61537c5f81f0cf325b4a854c
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Fügen Sie ein `header`-Element innerhalb des `body`-Elements hinzu und weisen Sie ihm die Klasse `header` zu.

Erstellen Sie innerhalb des `header` ein `h1` mit `css flexbox photo gallery` als Text.

# --hints--

Sie sollten ein `header`-Element innerhalb Ihres `body`-Elements haben.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Ihr `header`-Element sollte ein `class` mit `header` als Wert haben.

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Ihr `header`-Element sollte ein `h1`-Element darin enthalten.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Ihr `h1`-Element sollte den Text `css flexbox photo gallery` enthalten.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
