---
id: 61537485c4f2a624f18d7794
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Beginnen Sie mit Ihrem Standard-HTML-Boilerplate-Code. Fügen Sie eine `DOCTYPE`-Deklaration, ein `html`-Element, ein `head`-Element und ein `body`-Element hinzu.

Fügen Sie dem öffnenden `lang`-Tag das `<html>`-Attribut mit `en` als Wert hinzu.

# --hints--

Ihr Code sollte die `DOCTYPE`-Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Sie sollten nach der `DOCTYPE`-Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Sie sollten `html` als Dokumenttyp definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Sie sollten die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Ihr `html`-Element sollte ein öffnendes Tag haben.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Ihr öffnendes `<html>`-Tag sollte das `lang`-Attribut mit `en` als Wert haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

Ihre `DOCTYPE`-Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Sie sollten ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Sie sollten ein schließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Sie sollten ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Sie sollten ein schließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Die `head`- und `body`-Elemente sollten Geschwister sein.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Das `head`-Element sollte innerhalb des `html`-Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body`-Element sollte innerhalb des `html`-Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
