---
id: 61537bb9b1a29430ac15ad38
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Fügen Sie innerhalb Ihres `head`-Elements ein `title`-Element mit dem Text `Photo Gallery` hinzu und ein `link`-Element, um Ihre `styles.css`-Datei auf der Seite einzufügen.

# --hints--

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Ihr Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ihr Projekt sollte den Titel `Photo Gallery` haben.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

Denken Sie daran, dass Groß- und Kleinschreibung sowie die Schreibweise für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
