---
id: 61537c9eecea6a335db6da79
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Unter Ihrem `.header`-Element erstellen Sie ein neues `div`-Element und weisen ihm eine `class` von `gallery` zu. Dieses `div` fungiert als Container für die Galerie-Bilder.

Innerhalb dieses `.gallery`-Elements erstellen Sie neun `img`-Elemente.

# --hints--

Sie sollten ein `div`-Element in Ihrem `body`-Element erstellen.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

Ihr neues `div`-Element sollte eine `class` mit `gallery` als Wert haben.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

Ihr neues `div`-Element sollte nach Ihrem `.header`-Element kommen.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Ihr `.gallery`-Element sollte neun `img`-Elemente enthalten.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
