---
id: 6153893900438b4643590367
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Richten Sie Ihren `.header`-Text zentriert aus. Machen Sie den Text mit der `text-transform`-Eigenschaft und `uppercase` als Wert großgeschrieben.

Geben Sie ihm auf allen Seiten einen Innenabstand von `32px`. Setzen Sie den Hintergrund auf `#0a0a23` und die Textfarbe auf `#fff`.

Fügen Sie eine `border-bottom` mit `4px solid #fdb347` als Wert hinzu.

# --hints--

Sie sollten einen `.header`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Ihr `.header`-Selektor sollte eine `text-align`-Eigenschaft mit dem Wert `center` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

Ihr `.header`-Selektor sollte eine `text-transform`-Eigenschaft mit dem Wert `uppercase` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

Ihr `.header`-Selektor sollte eine `padding`-Eigenschaft mit dem Wert `32px` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

Ihr `.header`-Selektor sollte eine `background-color`-Eigenschaft mit dem Wert `#0a0a23` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

Ihr `.header`-Selektor sollte eine `color`-Eigenschaft mit dem Wert `#fff` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

Ihr `.header`-Selektor sollte eine `border-bottom`-Eigenschaft mit dem Wert `4px solid #fdb347` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
