---
id: 6493bc0d99879635209565aa
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Das `border-box`-Sizing-Modell macht das Gegenteil von `content-box`. Die Gesamtbreite des Elements, einschließlich Innenabstand und Rahmen, entspricht der explizit festgelegten Breite. Der Inhalt des Elements wird verkleinert, um Platz für den Innenabstand und den Rahmen zu schaffen.

Ändern Sie die `box-sizing`-Eigenschaft in `border-box`. Beachten Sie, wie Ihre blauen Bildrahmen jetzt innerhalb des roten Galerie-Rahmens passen.

# --hints--


Ihr *-Selektor sollte eine `box-sizing`-Eigenschaft mit dem Wert `border-box` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>


```

```css
--fcc-editable-region--
 * {
  box-sizing: content-box;
}
--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
