---
id: 6494da0daf5df5197963671d
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Beachten Sie, wie der blaue Rahmen des Bildes über den roten Rahmen der Galerie hinausgeht. Das liegt daran, wie Browser die Größe von Container-Elementen berechnen.

Die `box-sizing`-Eigenschaft wird verwendet, um dieses Verhalten einzustellen. Standardmäßig wird das `content-box`-Modell verwendet. Bei diesem Modell wird, wenn ein Element eine bestimmte Breite hat, diese Breite nur auf Basis des Inhalts des Elements berechnet. Innenabstand und Rahmenwerte werden zur Gesamtbreite hinzugefügt, sodass das Element wächst, um diese Werte aufzunehmen.

Versuchen Sie, `box-sizing` explizit auf `content-box` zu setzen, mit dem globalen `*`-Selektor. Zu diesem Zeitpunkt werden Sie keine Änderungen sehen, da Sie den Standardwert verwenden.

# --hints--

Sie sollten einen `*`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Ihr `*`-Selektor sollte eine `box-sizing`-Eigenschaft haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

Ihr `*`-Selektor sollte seine `box-sizing`-Eigenschaft auf `content-box` gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
