---
id: 587d78ac367417b2b2512af6
title: 使用 justify-content 属性对齐元素
challengeType: 0
forumTopicId: 301102
dashedName: align-elements-using-the-justify-content-property
---

# --description--

flex 子元素有时不能充满整个 flex 容器， 所以我们经常需要告诉 CSS 以什么方式排列 flex 子元素，以及调整它们的间距。 幸运的是，我们可以通过 `justify-content` 属性的不同值来实现。 在介绍属性的可选值之前，我们要先理解一些重要术语。

<a href="https://www.freecodecamp.org/news/flexbox-the-ultimate-css-flex-cheatsheet/" target="_blank" rel="noopener noreferrer nofollow">有关 flex-box 属性的更多简介，请点击这里阅读</a>

回忆一下，如果把 flex 容器设为一个行，它的子元素会从左到右逐个排列。 如果把 flex 容器设为一个列，它的子元素会从上到下逐个排列。 子元素排列的方向被称为 **main axis（主轴）**。 对一行来说，主轴水平贯穿每一个对象； 对于列，主轴垂直贯穿每一个对象。

对于如何沿主轴线排放 flex 项目，有几种选择。 常用的一种是 `justify-content: center;`：即 flex 子元素在 flex 容器中居中排列。 其他选择包括：

<ul><li><code>flex-start</code>：将项对齐到弹性容器的起始位置。对于行，这会将项推到容器的左侧。对于列，这会将项推到容器的顶部。如果没有指定 <code>justify-content</code>，这是默认的对齐方式。</li><li><code>flex-end</code>：将项对齐到弹性容器的结束位置。对于行，这会将项推到容器的右侧。对于列，这会将项推到容器的底部。</li><li><code>space-between</code>：将项沿主轴居中对齐，项之间留有额外空间。第一个和最后一个项被推到弹性容器的最边缘。例如，在行中，第一个项靠近容器左侧，最后一个项靠近容器右侧，剩余空间均匀分布在其他项之间。</li><li><code>space-around</code>：类似于 <code>space-between</code>，但第一个和最后一个项不锁定在容器边缘，空间均匀分布在所有项周围，弹性容器两端各有半个空间。</li><li><code>space-evenly</code>：在弹性项之间均匀分布空间，弹性容器两端各有一个完整空间。</li></ul>

# --instructions--

这个例子可以帮助你理解这个属性。 请为 `#box-container` 元素添加 CSS 属性 `justify-content`，并将其属性值设置为 `center`。

**奖励**尝试代码编辑器中 `justify-content` 属性的其他选项以查看它们的区别。但请注意，只有值为 `center` 的选项才能通过此挑战。

# --hints--

`#box-container` 所选择的元素应有 `justify-content` 属性，且其属性值应为 `center`。

```js
const boxContainer = document.querySelector('#box-container');
const justifyDirection =
  window.getComputedStyle(boxContainer)['justify-content'];
assert.strictEqual(justifyDirection, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 100%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 100%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    justify-content: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 100%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 100%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
