---
id: 587d78ad367417b2b2512af8
title: 使用 align-items 属性对齐元素
challengeType: 0
forumTopicId: 301101
dashedName: align-elements-using-the-align-items-property
---

# --description--

`align-items` 属性与 `justify-content` 类似。 回忆一下，`justify-content` 属性使 flex 子元素沿主轴排列。 行的主轴是水平线，列的主轴是垂直线。

Flex 容器中，与主轴垂直的叫做 **cross axis（交叉轴）**。 行的交叉轴是垂直的，列的交叉轴是水平的。

CSS 中的 `align-items` 属性用来定义 flex 子元素沿交叉轴的对齐方式。 对行来说，定义的是元素的上下对齐方式； 对列来说，是定义元素的左右对齐方式。

`align-items` 的可选值包括：

<ul><li><code>flex-start</code>：将项对齐到 flex 容器的起始位置。对于行，这会将项对齐到容器的顶部。对于列，这会将项对齐到容器的左侧。</li><li><code>flex-end</code>：将项对齐到 flex 容器的结束位置。对于行，这会将项对齐到容器的底部。对于列，这会将项对齐到容器的右侧。</li><li><code>center</code>：将项对齐到中心。对于行，这会垂直对齐项（项的上下空间相等）。对于列，这会水平对齐项（项的左右空间相等）。</li><li><code>stretch</code>：拉伸项以填满 flex 容器。例如，行项会被拉伸以填满容器的上下空间。如果未指定 <code>align-items</code> 值，则这是默认值。</li><li><code>baseline</code>：将项对齐到它们的基线。基线是一个文本概念，可以理解为字母所坐落的那条线。</li></ul>

# --instructions--

这个例子可以帮助你理解这个属性。 请在 `#box-container` 里添加 CSS 属性 `align-items` 并将值设为 `center`。

**奖励**尝试代码编辑器中 `align-items` 属性的其他选项以查看它们的区别。但请注意，只有值为 `center` 的选项才能通过此挑战。

# --hints--

`#box-container` 所选择的元素应有 `align-items` 属性，且其属性值应为 `center`。

```js
const boxContainer = document.querySelector('#box-container');
const alignment = window.getComputedStyle(boxContainer)['align-items'];
assert.strictEqual(alignment, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    align-items: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```
