---
id: 616971b3cd990186b66c99a1
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

既然你的外部 CSS 檔案已設定完成，你就可以開始為頁面設計樣式。

提醒你，以下是如何定位段落元素並將其對齊到右側的方法：

```css
p {
  text-align: right;
}
```

創建一個新的 CSS 規則，標的為 `h1` 元素，並將其 `text-align` 屬性設為 `center`。

# --hints--

你應該使用 `h1` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

你的 `h1` CSS 規則應該有一個 `text-align` 屬性設定為 `center`。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
