---
id: 61695c4aad56f95497c19583
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

請記住，`title` 元素會提供搜尋引擎關於頁面的額外資訊。它還會以另外兩種方式顯示該 `title` 元素的內容：

* 當頁面開啟時在標題列
* 當你暫留在瀏覽器分頁上時，即使該分頁不是作用中的，當你暫留在該分頁時，`title` 文字會顯示出來。

在 `head` 元素內，巢狀一個帶有文字 `Colored Markers` 的 `title` 元素。

# --hints--

你應該有一個開啟的 `title` 頁籤。

```js
assert(code.match(/<title\s*>/i));
```

你應該有一個結束的 `title` 頁籤。

```js
assert(code.match(/<\/title\s*>/i));
```

你的專案應該有標題 `Colored Markers`。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

請記住，標題文字的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
