---
id: 616965351e74d4689eb6de30
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

你可以在一個網頁上擁有多個 `meta` 元素。每個 `meta` 元素都會為該頁面添加其他 HTML 元素無法表達的資訊。

在 `head` 中添加另一個 `meta` 元素。為它設定 `name` 屬性為 `"viewport"`，並設定 `content` 屬性為 `"width=device-width, initial-scale=1.0"`，這樣你的頁面在所有裝置上看起來都會相同。

# --hints--

你應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

你的新 `meta` 元素應該是空元素，不應該有結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```


你的新 `meta` 元素應該有一個 `name` 屬性設定為 `"viewport"`，以及一個 `content` 屬性設定為 `"width=device-width, initial-scale=1.0"`。

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

你的兩個 `meta` 元素應該位於 `head` 元素內。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
