---
id: 61695197ac34f0407e339882
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

As you've seen in the previous projects, webpages should start with a `DOCTYPE html` declaration, followed by an `html` element.

在文檔頂部添加 `DOCTYPE html` 聲明，然後添加 `html` 元素。 給 `html` 元素一個值爲 `en` 的 `lang` 屬性。

# --hints--

代碼應包含 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

應該在 `DOCTYPE` 聲明之後包含一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該把文檔類型定義爲 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

應該在 `DOCTYPE` 聲明的類型後面加上一個 `>`。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個值爲 `en` 的 `lang` 屬性的開始標籤。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/gi));
```

你的 `DOCTYPE` 聲明應該位於 HTML 的開始。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
