---
id: 61695d1fbc003856628bf561
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

`charset` 屬性指定文件使用的字元編碼。`utf-8`（Unicode Transformation Format – 8-bit）是一種用於電子通訊的字元編碼標準。

在 `head` 元素內，巢狀一個 `meta` 元素，其屬性 `charset` 設定為 `"utf-8"`。

# --hints--

你應該有一個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

你的 `meta` 元素應該是空元素，不應該有結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 元素應該有一個 `charset` 屬性，並設定為值 `"utf-8"`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

你的 `meta` 元素應該在 `head` 元素內。

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
