---
id: 61696ef7ac756c829f9e4048
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

在這個專案中，你將使用外部 CSS 檔案來設計頁面。我們已經為你創建了一個 `styles.css` 檔案。但在你能使用它之前，你需要將它連結到頁面。

將 `link` 元素巢狀於 `head` 元素內。為它的 `rel` 屬性設定為 `"stylesheet"`，並為 `href` 屬性設定為 `"styles.css"`。

# --hints--

你應該有一個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

你的 `link` 元素應該是空元素，不應該有結束標籤 `</link>`。

```js
assert.notMatch(code, /<\/link>/i);
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `"stylesheet"`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，值為 `"styles.css"`。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

你的 `link` 元素應該在 `head` 元素內。

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
