---
id: 61695197ac34f0407e339882
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

如你在先前的專案中所見，網頁應該以 `DOCTYPE html` 宣告開始，接著是 `html` 元素。

在文件頂部為你的文件添加 `DOCTYPE html` 宣告，並在其後添加一個 `html` 元素。為 `html` 元素添加一個 `lang` 屬性，其值為 `en`。

# --hints--

你的程式碼應該包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應該在 `DOCTYPE` 引用後面加上一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該定義文件型別為 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你應該在 `DOCTYPE` 宣告式的型別後面以 `>` 關閉。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個開啟標籤，並且帶有 `lang` 屬性且值為 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/gi));
```

你的 `DOCTYPE` 宣告應該放在你的 HTML 開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
