---
id: 61695ab9f6ffe951c16d03dd
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

將 `head` 元素巢狀於 `html` 元素內。在 `head` 元素之後，添加一個 `body` 元素。

# --hints--

你應該有一個開啟的 `head` 頁籤。

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個結束的 `head` 頁籤。

```js
assert(code.match(/<\/head\s*>/i));
```

你應該有一個開啟的 `body` 頁籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個結束的 `body` 頁籤。

```js
assert(code.match(/<\/body\s*>/i));
```

只能有一個 `head` 開啟標籤。

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

應該只有一個 `head` 結束標籤。

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

應該只有一個 `body` 開啟標籤。

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

應該只有一個 `body` 結束標籤。

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

你的 `head` 元素應該是空的。

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

你的 `body` 元素應該是空的。

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

你的 `body` 元素應該放在你的 `head` 元素之後。

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

你的 `head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

你的 `body` 元素應該位於 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
