---
id: 616d55bd160a95e22453a081
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

現在你已經將一個標記置中並上色，是時候為其他標記添加顏色了。

在 `container` `div` 中，添加兩個 `div` 元素，並分別給它們一個 `marker` 類別。

# --hints--

你的第一個新的 `div` 元素應該有一個開啟的標籤。

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

你的第一個新的 `div` 元素應該有一個結束標籤。

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

你的第二個新的 `div` 元素應該有一個開啟的標籤。

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

你的第二個新的 `div` 元素應該有一個結束標籤。

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

你的新 `div` 元素應該位於具有 `container` 類別的 `div` 內。

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

你的新 `div` 元素都應該有一個 `marker` 的類別。

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
