---
id: 616d3a67ccf800ad94ec89ae
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

現在你將添加一些元素，最終你會將它們設計成彩色標記。

首先，在 `body` 元素內，添加一個 `div` 元素並將其 `class` 屬性設為 `container`。確保該 `div` 元素位於 `h1` 元素下方。

# --hints--

你應該有一個開啟的 `div` 頁籤。

```js
assert(code.match(/<div.*>/i));
```

你應該有一個結束的 `div` 頁籤。

```js
assert(code.match(/<\/div\s*>/i));
```

你的 `div` 元素應該位於你的 `body` 元素內。

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

你的 `div` 元素應該在 `h1` 元素之後。

```js
assert.exists(document.querySelector('h1 + div'));
```

你的 `div` 元素應該有一個設定為 `container` 的 `class` 屬性。

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
