---
id: 617b954d9f4f6217a749380e
title: 第 25 步
challengeType: 0
dashedName: step-25
---

# --description--

現在在你的標記與它們所在的 `container` 元素邊緣之間添加更多垂直間距。

在 `.container` CSS 規則中，使用縮寫的 `padding` 屬性為上下添加 `10px` 的內距，並將左右內距設為 `0`。這與你之前使用的縮寫 `margin` 屬性類似。

# --hints--

你不應該從 `.container` CSS 規則中移除 `background-color` 屬性及其值。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

你的 `.container` CSS 規則應該有一個 `padding` 屬性設定為 `10px 0`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
