---
id: 617b97abd9f3f61d1590b815
title: 步驟 26
challengeType: 0
dashedName: step-26
---

# --description--

在加法 RGB 色彩模型中，<dfn>原色</dfn>是指當混合時會生成純白色的顏色。但要達成這個效果，每種顏色都需要處於最高強度。

在你混合顏色之前，將你的綠色標記恢復為純綠色。對於 `.two` CSS 規則中的 `rgb` 函式（程式），將綠色恢復到最大值 `255`。

# --hints--

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設定為 `rgb(0, 255, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
