---
id: 618a132676346ac9f7fd59dd
title: 步驟 39
challengeType: 0
dashedName: step-39
---

# --description--

注意你的眼睛如何自然地被中央的紅色所吸引？在設計網站時，你可以利用此效果來吸引注意重要的標題、按鈕或聯結。

除了互補色之外，還有其他幾種重要的顏色組合，但你會在稍後學到那些。

目前，請在 `.two` CSS 規則中使用 `rgb` 函式（程式）將 `background-color` 設為黑色。

# --hints--

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設為 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
