---
id: 619b761916dac02643940022
title: 步驟 46
challengeType: 0
dashedName: step-46
---

# --description--

使用 CSS 為元素套用顏色的一個非常常見的方法是使用 <dfn>十六進位</dfn> 或 hex 值。雖然 hex 值聽起來很複雜，但它們其實只是 RGB 值的另一種形式。

十六進位顏色值以 `#` 字元開頭，並採用 0-9 與 A-F 的六個字元。第一對組字元表述紅色，第二對組表述綠色，第三對組表述藍色。例如，`#4B5320`。

在 `.green` 類別選擇器中，將 `background-color` 屬性設為紅色為 `00`、綠色為 `FF`、藍色為 `00` 的十六進位色碼。

# --hints--

你的 `.green` CSS 規則應該有一個 `background-color` 屬性設定為 `#00FF00`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
