---
id: 61a49d15bdbb5e57cc6fd280
title: 步驟 54
challengeType: 0
dashedName: step-54
---

# --description--

色彩停點允許你微調顏色在漸層線上的位置。它們是一種長度單位，像是 `px` 或百分比，跟隨在 `linear-gradient` 函式中的顏色後面。

例如，在這個紅黑漸層中，從紅色到黑色的過渡發生在漸層線的 90％ 點，因此紅色佔據了大部分可用空間：

```css
linear-gradient(90deg, red 90%, black);
```

在 `linear-gradient` 函式中，為你的第一個紅色引數後添加一個 `75%` 色彩停點。不要為其他顏色引數添加色彩停點。

# --hints--

你的 `.red` CSS 規則應該有一個 `background` 屬性，設定為 `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`。

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
