---
id: 61adc91467b5d59328b9f781
title: 步驟 79
challengeType: 0
dashedName: step-79
---

# --description--

在上一個專案中，你學到了一點關於邊框和 `border-color` 屬性的知識。

所有 HTML 元素都有邊框，儘管它們通常預設為 `none`。使用 CSS，你可以控制元素邊框的所有面向，並設定所有邊框，或一次只設定一個邊框。要讓邊框可見，你需要設定其寬度和樣式。

在 `.sleeve` CSS 規則中，為你的 `.sleeve` 添加 `border-left-width` 屬性，值為 `10px`。

# --hints--

你的 `.sleeve` CSS 規則應該有一個 `border-left-width` 屬性，且其值為 `10px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeftWidth === '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}
--fcc-editable-region--

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
