---
id: 61b0a44a6b865738ba49b9fb
title: 步驟 86
challengeType: 0
dashedName: step-86
---

# --description--

最後你要做的是為每個標記添加一點陰影，使它們看起來更逼真。

`box-shadow` 屬性讓你可以在元素周圍套用一個或多個陰影。以下是基本語法：

```css
box-shadow: offsetX offsetY color;
```

以下是 `offsetX` 和 `offsetY` 值的運作方式：

* `offsetX` 和 `offsetY` 都接受以 `px` 及其他 CSS 單位表示的數值。
* 正的 `offsetX` 值會將陰影向右移動，負值則會向左移動。
* 正的 `offsetY` 值會將陰影向下移動，負值則會向上移動
* 如果你想要 `offsetX` 和 `offsetY` 中任一或兩者的值為零（`0`），你不需要添加單位。每個瀏覽器都理解零表示沒有變化。

陰影的高度和寬度由其套用的元素的高度和寬度決定。你也可以使用選用的 `spreadRadius` 值來擴展陰影的範圍。稍後會有更多說明。

首先為紅色標記添加一個簡單的陰影。

在 `.red` CSS 規則中，為你的 `box-shadow` 屬性添加 `5px` 的 `offsetX`、`5px` 的 `offsetY` 和 `red` 的 `color`。

# --hints--

你的 `.red` CSS 規則應該有一個 `box-shadow` 簡寫屬性，且值為 `5px 5px red`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
