---
id: 5a24c314108439a4d403616b
title: Use Default Props
challengeType: 6
forumTopicId: 301418
dashedName: use-default-props
---

# --description--

У React також є можливість встановлювати пропси за замовчуванням. Ви можете призначити пропси за замовчуванням компоненту як властивість самого компонента, і React призначить цю пропс за замовчуванням, якщо це необхідно. Це дозволяє вказати, яким має бути значення пропса, якщо явно не передано жодного значення. Наприклад, якщо ви оголосите `MyComponent.defaultProps = { location: 'San Francisco' }`, ви визначили проп location, яке встановлене на рядок `San Francisco`, якщо не вказано інакше. React призначає пропси за замовчуванням, якщо пропси є undefined, але якщо ви передасте `null` як значення пропса, воно залишиться `null`.

# --instructions--

У редакторі коду показано компонент `ShoppingCart`. Визначте пропси за замовчуванням для цього компонента, які вказують проп `items` зі значенням `0`.

# --hints--

Компонент `ShoppingCart` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Компонент `ShoppingCart` має мати проп за замовчуванням `{ items: 0 }`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    mockedComponent.setProps({ items: undefined });
    return mockedComponent.find('ShoppingCart').props().items === 0;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};
// Change code below this line
```

# --solutions--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};

// Change code below this line
ShoppingCart.defaultProps = {
  items: 0
}
```
