---
id: 5a24c314108439a4d403616e
title: Access Props Using this.props
challengeType: 6
forumTopicId: 301375
dashedName: access-props-using-this-props
---

# --description--

Останні кілька завдань охоплювали базові способи передачі пропсів дочірнім компонентам. Але що робити, якщо дочірній компонент, якому ви передаєте пропс, є компонентом класу ES6, а не безстанним функціональним компонентом? Компонент класу ES6 використовує трохи іншу конвенцію для доступу до пропсів.

Кожного разу, коли ви посилаєтеся на компонент класу всередині нього самого, ви використовуєте ключове слово `this`. Щоб отримати доступ до пропсів у компоненті класу, ви перед кодом, який використовуєте для доступу, ставите `this`. Наприклад, якщо у компонента класу ES6 є проп з назвою `data`, ви пишете `{this.props.data}` у JSX.

# --instructions--

Відтворіть екземпляр компонента `Welcome` у батьківському компоненті `App`. Тут дайте `Welcome` проп `name` і призначте йому значення рядка. Усередині дочірнього компонента `Welcome` отримайте доступ до пропса `name` у тегах `strong`.

# --hints--

Компонент `App` має повертати один елемент `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Дочірнім елементом `App` має бути компонент `Welcome`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.children().childAt(0).name() === 'Welcome'
    );
  })()
);
```

Компонент `Welcome` має мати проп з назвою `name`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.find('Welcome').props().name;
  })()
);
```

Компонент `Welcome` має відображати рядок, який ви передаєте як проп `name`, у тегах `strong`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.find('strong').text() ===
      mockedComponent.find('Welcome').props().name
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome />
            { /* Change code above this line */ }
        </div>
    );
  }
};

class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong></strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};
```

# --solutions--

```jsx
class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong>{this.props.name}</strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};

class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome name="Quincy"/>
            { /* Change code above this line */ }
        </div>
    );
  }
};
```
