---
id: 5a24c314108439a4d403617b
title: Передайте функцію зворотного виклику як пропси
challengeType: 6
forumTopicId: 301400
dashedName: pass-a-callback-as-props
---

# --description--

Ви можете передавати `state` як пропси дочірнім компонентам, але ви не обмежені лише передачею даних. Ви також можете передавати обробники або будь-який метод, визначений у React-компоненті, дочірньому компоненту. Саме так ви дозволяєте дочірнім компонентам взаємодіяти зі своїми батьківськими компонентами. Ви передаєте методи дочірньому компоненту так само, як звичайний проп. Йому призначається ім’я, і ви маєте доступ до цього імені методу під `this.props` у дочірньому компоненті.

# --instructions--

У редакторі коду наведено три компоненти. Компонент `MyApp` є батьківським, який відтворюватиме дочірні компоненти `GetInput` і `RenderInput`. Додайте компонент `GetInput` до методу відтворення в `MyApp`, потім передайте йому проп з ім’ям `input`, призначений `inputValue` з `MyApp` `state`. Також створіть проп з ім’ям `handleChange` і передайте йому обробник введення `handleChange`.

Далі додайте `RenderInput` до методу відтворення в `MyApp`, потім створіть проп з ім’ям `input` і передайте йому `inputValue` з `state`. Коли ви закінчите, ви зможете вводити текст у поле `input` у компоненті `GetInput`, що викликає метод обробника у його батьківському елементі через пропси. Це оновлює введення у `state` батьківського компонента, яке передається як пропси обом дочірнім. Спостерігайте, як дані рухаються між компонентами і як єдине джерело істини залишається `state` батьківського компонента. Звісно, цей приклад дещо штучний, але він ілюструє, як дані та функції зворотного виклику можна передавати між React-компонентами.

# --hints--

Компонент `MyApp` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('MyApp').length === 1;
  })()
);
```

Компонент `GetInput` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('GetInput').length === 1;
  })()
);
```

Компонент `RenderInput` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('RenderInput').length === 1;
  })()
);
```

Компонент `GetInput` має отримувати стан `MyApp` властивості `inputValue` як пропси і містити елемент `input`, який змінює стан `MyApp`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: '' });
    return waitForIt(() => mockedComponent.state());
  };
  const state_2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => mockedComponent.state());
  };
  const updated_1 = await state_1();
  const updated_2 = await state_2();
  assert(updated_1.inputValue === '' && updated_2.inputValue === 'TestInput');
```

Компонент `RenderInput` має отримувати стан `MyApp` властивості `inputValue` як пропси.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: 'TestName' });
    return waitForIt(() => mockedComponent);
  };
  const updated_1 = await state_1();
  assert(updated_1.find('p').text().includes('TestName'));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyApp />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
  this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
         <GetInput
           input={this.state.inputValue}
           handleChange={this.handleChange}/>
         <RenderInput
           input={this.state.inputValue}/>
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```
