---
id: 5a24c314108439a4d4036160
title: Визначення HTML-класу в JSX
challengeType: 6
forumTopicId: 301393
dashedName: define-an-html-class-in-jsx
---

# --description--

Тепер, коли ви починаєте впевнено писати JSX, можливо, вам цікаво, чим він відрізняється від HTML.

Поки що може здаватися, що HTML і JSX — це одне й те саме.

Одна ключова відмінність у JSX полягає в тому, що ви більше не можете використовувати слово `class` для визначення HTML-класів. Це тому, що `class` є зарезервованим словом у JavaScript. Замість цього в JSX використовується `className`.

Насправді, конвенція іменування для всіх HTML-атрибутів і посилань на події в JSX переходить у верблюжий регістр. Наприклад, подія click у JSX — це `onClick`, замість `onclick`. Аналогічно, `onchange` стає `onChange`. Хоча це й тонка відмінність, її важливо пам’ятати надалі.

# --instructions--

Застосуйте клас `myDiv` до `div`, який наведено в коді JSX.

# --hints--

Константа `JSX` має повертати елемент `div`.

```js
assert.strictEqual(JSX.type, 'div');
```

`div` має мати клас `myDiv`.

```js
assert.strictEqual(JSX.props.className, 'myDiv');
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(JSX, document.getElementById('root'))
```

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h1>Add a class to this div</h1>
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div className = 'myDiv'>
  <h1>Add a class to this div</h1>
</div>);
```
