---
id: 5a24c314108439a4d403616a
title: Передача масиву як пропсів
challengeType: 6
forumTopicId: 301401
dashedName: pass-an-array-as-props
---

# --description--

У попередньому завданні показано, як передавати інформацію від батьківського компонента до дочірнього як `props` або властивості. У цьому завданні розглядається, як масиви можна передавати як `props`. Щоб передати масив у JSX-елемент, його потрібно розглядати як JavaScript і обгорнути у фігурні дужки.

```jsx
<ParentComponent>
  <ChildComponent colors={["green", "blue", "red"]} />
</ParentComponent>
```

Дочірній компонент тоді має доступ до властивості масиву `colors`. При доступі до властивості можна використовувати методи масиву, такі як `join()`.

```jsx
const ChildComponent = (props) => <p>{props.colors.join(', ')}</p>
```

Це об’єднає всі елементи масиву `colors` у рядок, розділений комами, і виведе: `<p>green, blue, red</p>`. Пізніше ми дізнаємося про інші загальні методи для відтворення масивів даних у React.

# --instructions--

У редакторі коду є `List` і `ToDo` компоненти. Під час відтворення кожного `List` із компонента `ToDo` передайте властивість `tasks`, призначену масиву завдань, наприклад `["walk dog", "workout"]`. Потім отримайте доступ до цього масиву `tasks` у компоненті `List`, показуючи його значення всередині елемента `p`. Використайте `join(", ")`, щоб відобразити масив `props.tasks` у елементі `p` як список, розділений комами. Список на сьогодні має містити щонайменше 2 завдання, а на завтра — щонайменше 3 завдання.

# --hints--

Компонент `ToDo` має повертати один зовнішній `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().type() === 'div';
  })()
);
```

Третій дочірній елемент компонента `ToDo` має бути екземпляром компонента `List`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(2).name() === 'List';
  })()
);
```

П’ятий дочірній елемент компонента `ToDo` має бути екземпляром компонента `List`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(4).name() === 'List';
  })()
);
```

Обидва екземпляри компонента `List` мають мати властивість `tasks`, а `tasks` має бути типу масив.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      Array.isArray(mockedComponent.find('List').get(0).props.tasks) &&
      Array.isArray(mockedComponent.find('List').get(1).props.tasks)
    );
  })()
);
```

Перший компонент `List`, що представляє завдання на сьогодні, має містити 2 або більше елементів.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(0).props.tasks.length >= 2;
  })()
);
```

Другий компонент `List`, що представляє завдання на завтра, має містити 3 або більше елементів.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(1).props.tasks.length >= 3;
  })()
);
```

Компонент `List` має відтворювати значення з пропса `tasks` у тегу `p`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      mockedComponent
        .find('p')
        .get(0)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(0)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',') &&
      mockedComponent
        .find('p')
        .get(1)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(1)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',')
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<ToDo />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const List = (props) => {
  { /* Change code below this line */ }
  return <p>{}</p>
  { /* Change code above this line */ }
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        { /* Change code below this line */ }
        <List/>
        <h2>Tomorrow</h2>
        <List/>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const List= (props) => {
  return <p>{props.tasks.join(', ')}</p>
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        <List tasks={['study', 'exercise']} />
        <h2>Tomorrow</h2>
        <List tasks={['call Sam', 'grocery shopping', 'order tickets']} />
      </div>
    );
  }
};
```
