---
id: 5a24c314108439a4d403616f
title: Повторення використання пропсів у безстанових функціональних компонентах
challengeType: 6
forumTopicId: 301411
dashedName: review-using-props-with-stateless-functional-components
---

# --description--

Окрім останнього завдання, ви передавали пропси безстановим функціональним компонентам. Ці компоненти працюють як чисті функції. Вони приймають пропси як вхідні дані і повертають той самий вигляд щоразу, коли їм передають однакові пропси. Можливо, ви цікавитесь, що таке стан, і наступне завдання детальніше це розгляне. А поки що — повторення термінології для компонентів.

*Безстановий функціональний компонент* — це будь-яка функція, яку ви пишете, що приймає пропси і повертає JSX. *Безстановий компонент*, натомість, — це клас, який розширює `React.Component`, але не використовує внутрішній стан (це розглядається в наступному завданні). Нарешті, *компонент зі станом* — це класовий компонент, який підтримує власний внутрішній стан. Ви можете бачити компоненти зі станом просто як компоненти або React-компоненти.

Загальний шаблон — намагатися мінімізувати використання стану і створювати безстанові функціональні компоненти, де це можливо. Це допомагає обмежити управління станом у певній області вашого застосунку. У свою чергу, це покращує розробку та обслуговування вашого застосунку, роблячи простішим відстеження того, як зміни стану впливають на його поведінку.

# --instructions--

У редакторі коду є компонент `CampSite`, який відтворює компонент `Camper` як дочірній елемент. Визначте компонент `Camper` і призначте йому пропси за замовчуванням `{ name: 'CamperBot' }`. Усередині компонента `Camper` відтворіть будь-який код, який хочете, але обов’язково має бути один елемент `p`, що містить лише значення `name`, яке передається як проп `prop`. Нарешті, визначте `propTypes` у компоненті `Camper`, щоб вимагати наявність пропса `name` і перевірити, що він має тип `string`.

# --hints--

Компонент `CampSite` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('CampSite').length === 1;
  })()
);
```

Компонент `Camper` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('Camper').length === 1;
  })()
);
```

Компонент `Camper` має містити пропси за замовчуванням, які призначають рядок `CamperBot` ключу `name`.

```js
assert(
  /Camper.defaultProps={name:(['"`])CamperBot\1,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

Компонент `Camper` має містити типи пропсів, які вимагають, щоб проп `name` був типу `string`.

```js
assert(
  /Camper.propTypes={name:PropTypes.string.isRequired,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

Компонент `Camper` має містити елемент `p` з лише текстом із пропса `name`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return (
      mockedComponent.find('p').text() ===
      mockedComponent.find('Camper').props().name
    );
  })()
);
```

# --seed--

## --before-user-code--

```jsx
var PropTypes = {
   string: { isRequired: true }
};
```

## --after-user-code--

```jsx
ReactDOM.render(<CampSite />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line
```

# --solutions--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line

const Camper = (props) => {
   return (
     <div>
       <p>{props.name}</p>
     </div>
   );
};

Camper.propTypes = {
  name: PropTypes.string.isRequired
};

Camper.defaultProps = {
  name: 'CamperBot'
};
```
