---
id: 5a24c314108439a4d4036179
title: Створення контрольованої форми
challengeType: 6
forumTopicId: 301384
dashedName: create-a-controlled-form
---

# --description--

Останнє завдання показало, що React може контролювати внутрішній стан для певних елементів, таких як `input` і `textarea`, що робить їх контрольованими компонентами. Це також стосується інших елементів форми, включно зі звичайним HTML-елементом `form`.

# --instructions--

Компонент `MyForm` налаштований з порожнім `form` та обробником відправлення. Обробник відправлення буде викликаний, коли форму відправлять.

Ми додали кнопку, яка відправляє форму. Ви можете побачити, що у неї встановлено `type` зі значенням `submit`, що вказує на те, що це кнопка, яка контролює форму. Додайте елемент `input` у `form` і встановіть його атрибути `value` та `onChange()`, як у попередньому завданні. Потім завершіть метод `handleSubmit`, щоб він встановлював властивість стану компонента `submit` у поточне значення вхідних даних у локальному `state`.

**Примітка:** Ви також повинні викликати `event.preventDefault()` в обробнику відправлення, щоб запобігти стандартній поведінці відправлення форми, яка оновить вебсторінку. Для зручності користувача стандартна поведінка тут вимкнена, щоб уникнути скидання коду завдання через оновлення.

Нарешті, створіть тег `h1` після `form`, який відтворюватиме значення `submit` зі стану компонента `state`. Тоді ви зможете вводити текст у форму і натискати кнопку (або клавішу Enter), і побачите, як ваш ввід відтворюється на сторінці.

# --hints--

`MyForm` має повертати елемент `div`, який містить `form` і тег `h1`. Форма має включати `input` і `button`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyForm));
    return (
      mockedComponent.find('div').children().find('form').length === 1 &&
      mockedComponent.find('div').children().find('h1').length === 1 &&
      mockedComponent.find('form').children().find('input').length === 1 &&
      mockedComponent.find('form').children().find('button').length === 1
    );
  })()
);
```

Стан `MyForm` має ініціалізуватися властивостями `input` і `submit`, обидві з яких встановлені у порожні рядки.

```js
assert(
  Enzyme.mount(React.createElement(MyForm)).state('input') === '' &&
    Enzyme.mount(React.createElement(MyForm)).state('submit') === ''
);
```

Введення в елемент `input` має оновлювати властивість `input` у стані компонента.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return mockedComponent.state('input');
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return {
      state: mockedComponent.state('input'),
      inputVal: mockedComponent.find('input').props().value
    };
  };
  const before = _1();
  const after = _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
})();
```

Відправлення форми має запускати `handleSubmit`, який має встановлювати властивість `submit` у стані рівною поточному вводу.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'SubmitInput' } });
    return mockedComponent.state('submit');
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.state('submit');
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'SubmitInput');
})();
```

`handleSubmit` має викликати `event.preventDefault`

```js
assert(
  __helpers.isCalledWithNoArgs(
    'event.preventDefault',
    MyForm.prototype.handleSubmit.toString()
  )
);
```

Елемент заголовка `h1` має відтворювати значення поля `submit` зі стану компонента.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return mockedComponent.find('h1').text();
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.find('h1').text();
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'TestInput');
})();
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyForm />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          {/* Change code below this line */}

          {/* Change code above this line */}
          <button type='submit'>Submit!</button>
        </form>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    event.preventDefault();
    this.setState(state => ({
      submit: state.input
    }));
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          <input value={this.state.input} onChange={this.handleChange} />
          <button type='submit'>Submit!</button>
        </form>
        <h1>{this.state.submit}</h1>
      </div>
    );
  }
}
```
