---
id: 5a24c314108439a4d403617d
title: Use the Lifecycle Method componentDidMount
challengeType: 6
forumTopicId: 301422
dashedName: use-the-lifecycle-method-componentdidmount
---

# --description--

Більшість веб-розробників у певний момент мають викликати кінцеву точку API, щоб отримати дані. Якщо ви працюєте з React, важливо знати, де саме виконувати цю дію.

Найкраща практика в React — розміщувати виклики API або будь-які виклики до вашого сервера в методі життєвого циклу `componentDidMount()`. Цей метод викликається після того, як компонент змонтовано в DOM. Будь-які виклики `setState()` тут спричинять повторне відтворення вашого компонента. Коли ви викликаєте API в цьому методі і оновлюєте стан за допомогою даних, які повертає API, це автоматично спричинить оновлення після отримання даних.

# --instructions--

У `componentDidMount()` є імітація виклику API. Вона оновлює стан через 2,5 секунди, щоб симулювати виклик сервера для отримання даних. У цьому прикладі запитується поточна загальна кількість активних користувачів сайту. У методі відтворення відтворіть значення `activeUsers` у `h1` після тексту `Active Users:`. Спостерігайте, що відбувається у попередньому перегляді, і не соромтеся змінювати час очікування, щоб побачити різні ефекти.

# --hints--

`MyComponent` має відтворити `div` елемент, який обгортає тег `h1`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

Стан компонента слід оновити за допомогою функції таймера в `componentDidMount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return new RegExp('setTimeout(.|\n)+setState(.|\n)+activeUsers').test(
      String(mockedComponent.instance().componentDidMount)
    );
  })()
);
```

Тег `h1` має відтворювати значення `activeUsers` зі стану `MyComponent`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ activeUsers: 1237 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ activeUsers: 1000 });
    return mockedComponent.find('h1').text();
  };
  assert(new RegExp('1237').test(first()) && new RegExp('1000').test(second()));
})();
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        {/* Change code below this line */}
        <h1>Active Users: </h1>
        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        <h1>Active Users: {this.state.activeUsers}</h1>
      </div>
    );
  }
}
```
