---
id: 5a24c314108439a4d4036178
title: Створити контрольований ввід
challengeType: 6
forumTopicId: 301385
dashedName: create-a-controlled-input
---


# --description--

Ваш застосунок може мати більш складні взаємодії між `state` та відтвореним інтерфейсом користувача. Наприклад, елементи керування формою для введення тексту, такі як `input` і `textarea`, підтримують власний стан у DOM під час введення користувачем. За допомогою React ви можете перемістити цей змінний стан у `state` компонента React. Вхідні дані користувача стають частиною `state` застосунку, тому React контролює значення цього поля введення. Зазвичай, якщо у вас є компоненти React із полями введення, у які користувач може вводити текст, це буде керована форма введення.

# --instructions--

Редактор коду містить каркас компонента під назвою `ControlledInput`, щоб створити керований `input` елемент. Стан компонента вже ініціалізовано з `input` властивістю, яка містить порожній рядок. Це значення представляє текст, який користувач вводить у `input` поле.

Спершу створіть метод під назвою `handleChange()`, який має параметр `event`. Коли метод викликається, він отримує об’єкт `event`, що містить рядок тексту з `input` елемента. Ви можете отримати доступ до цього рядка через `event.target.value` всередині методу. Оновіть властивість `input` стану компонента `state` цим новим рядком.

У методі `render` створіть `input` елемент над тегом `h4`. Додайте атрибут `value`, який дорівнює властивості `input` стану компонента `state`. Потім додайте властивість `onChange`, встановлену на обробник події `handleChange()`.

Коли ви вводите текст у поле введення, цей текст обробляється методом `handleChange()`, встановлюється як властивість `input` у локальному стані `state` і відтворюється як значення у `input` полі на сторінці. Стан компонента `state` є єдиним джерелом правди щодо вхідних даних.

І нарешті, не забудьте додати необхідні зв’язки у конструкторі.

# --hints--

`ControlledInput` має повертати `div` елемент, який містить `input` і тег `p`.

```js
assert(
  Enzyme.mount(React.createElement(ControlledInput))
    .find('div')
    .children()
    .find('input').length === 1 &&
    Enzyme.mount(React.createElement(ControlledInput))
      .find('div')
      .children()
      .find('p').length === 1
);
```

Стан `ControlledInput` має ініціалізуватися з властивістю `input`, встановленою в порожній рядок.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(ControlledInput)).state('input'),
  ''
);
```

Введення тексту в елемент введення має оновлювати стан і значення введення, а елемент `p` має відтворювати цей стан під час введення.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(ControlledInput));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return waitForIt(() => mockedComponent.state('input'));
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => ({
      state: mockedComponent.state('input'),
      text: mockedComponent.find('p').text(),
      inputVal: mockedComponent.find('input').props().value
    }));
  };
  const before = await _1();
  const after = await _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.text === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<ControlledInput />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        { /* Change code below this line */}

        { /* Change code above this line */}
        <h4>Controlled Input:</h4>
        <p>{this.state.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  render() {
    return (
      <div>
        <input
          value={this.state.input}
          onChange={this.handleChange} />
        <h4>Controlled Input:</h4>

        <p>{this.state.input}</p>
      </div>
    );
  }
};
```
