---
id: 5a24c314108439a4d403617a
title: Pass State as Props to Child Components
challengeType: 6
forumTopicId: 301403
dashedName: pass-state-as-props-to-child-components
---

# --description--

Ви бачили багато прикладів, де пропси передавалися дочірнім JSX-елементам і дочірнім React-компонентам у попередніх завданнях. Можливо, ви задумуєтеся, звідки беруться ці пропси. Загальний шаблон полягає в тому, що є компонент зі станом, який містить `state` важливий для вашого застосунку, і який потім відтворює дочірні компоненти. Ви хочете, щоб ці компоненти мали доступ до деяких частин цього `state`, які передаються як пропси.

Наприклад, можливо, у вас є `App` компонент, який відтворює `Navbar`, серед інших компонентів. У вашому `App` є `state`, який містить багато інформації про користувача, але `Navbar` потребує доступу лише до імені користувача, щоб його відобразити. Ви передаєте цю частину `state` компоненту `Navbar` як проп.

Цей шаблон ілюструє кілька важливих принципів у React. Перший — це *унідипрекційний потік даних*. Стан рухається в одному напрямку вниз по дереву компонентів вашого застосунку, від батьківського компонента зі станом до дочірніх компонентів. Дочірні компоненти отримують лише ті дані стану, які їм потрібні. Другий — складні застосунки зі станом можна розбити на кілька, або навіть один, компонент зі станом. Решта ваших компонентів просто отримують стан від батьківського елемента як пропси і відтворюють інтерфейс користувача на основі цього стану. Це починає створювати розділення, де керування станом відбувається в одній частині коду, а відтворення інтерфейсу — в іншій. Цей принцип розділення логіки стану від логіки інтерфейсу — один із ключових принципів React. Коли його правильно застосовують, це значно полегшує проєктування складних застосунків зі станом.

# --instructions--

Компонент `MyApp` є зі станом і відтворює компонент `Navbar` як дочірній. Передайте властивість `name` зі стану `state` вниз до дочірнього компонента, а потім покажіть `name` у тегу `h1`, який є частиною методу відтворення `Navbar`. `name` має з’явитися після тексту `Hello, my name is:`.

# --hints--

Компонент `MyApp` має відтворювати компонент `Navbar` всередині.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return (
      mockedComponent.find('MyApp').length === 1 &&
      mockedComponent.find('Navbar').length === 1
    );
  })()
);
```

Компонент `Navbar` має отримати властивість стану `MyApp` `name` як пропси.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').props());
  };
  const navProps = await setState();
  assert(navProps.name === 'TestName');
```

Елемент `h1` у `Navbar` має відтворити проп `name`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const navH1Before = mockedComponent.find('Navbar').find('h1').text();
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').find('h1').text());
  };
  const navH1After = await setState();
  assert(new RegExp('TestName').test(navH1After) && navH1After !== navH1Before);
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         {/* Change code below this line */}
         <Navbar />
         {/* Change code above this line */}
       </div>
    );
  }
};

class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      {/* Change code below this line */}
      <h1>Hello, my name is: </h1>
      {/* Change code above this line */}
    </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         <Navbar name={this.state.name}/>
       </div>
    );
  }
};
class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      <h1>Hello, my name is: {this.props.name}</h1>
    </div>
    );
  }
};
```
