---
id: 5a24c314108439a4d4036187
title: Use a Ternary Expression for Conditional Rendering
challengeType: 6
forumTopicId: 301414
dashedName: use-a-ternary-expression-for-conditional-rendering
---

# --description--

Перед тим як перейти до технік динамічного відтворення, є ще один спосіб використати вбудовані JavaScript умови для відтворення потрібного вмісту: <dfn>тернарний оператор</dfn>. Тернарний оператор часто використовується як скорочення для `if/else` інструкцій у JavaScript. Вони не такі потужні, як традиційні `if/else` інструкції, але дуже популярні серед розробників React. Одна з причин цього полягає в тому, як компілюється JSX: `if/else` інструкції не можна вставляти безпосередньо в код JSX. Ви могли помітити це у кількох попередніх завданнях — коли потрібна була `if/else` інструкція, вона завжди була *поза межами* `return` інструкції. Тернарні вирази можуть бути чудовою альтернативою, якщо ви хочете реалізувати умовну логіку всередині JSX. Нагадаємо, що тернарний оператор має три частини, але ви можете поєднувати кілька тернарних виразів разом. Ось базовий синтаксис:

```jsx
condition ? expressionIfTrue : expressionIfFalse;
```

# --instructions--

У редакторі коду визначено три константи всередині `CheckUserAge` компонента в методі `render()`. Вони називаються `buttonOne`, `buttonTwo` і `buttonThree`. Кожній з них призначено простий JSX вираз, який представляє елемент кнопки. Спочатку ініціалізуйте стан `CheckUserAge` зі значеннями `input` і `userAge`, обидва встановлені в порожній рядок.

Коли компонент відтворює інформацію на сторінці, користувачі повинні мати можливість взаємодіяти з нею. Усередині `return` інструкції компонента налаштуйте тернарний вираз, який реалізує таку логіку: коли сторінка завантажується вперше, відтворіть кнопку відправки `buttonOne` на сторінці. Потім, коли користувач вводить свій вік і натискає кнопку, відтворіть іншу кнопку залежно від віку. Якщо користувач вводить число менше за `18`, відтворіть `buttonThree`. Якщо користувач вводить число, яке дорівнює або більше за `18`, відтворіть `buttonTwo`.

# --hints--

Компонент `CheckUserAge` має відтворювати один елемент `input` і один елемент `button`.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('input')
    .length === 1 &&
    Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('button')
      .length === 1
);
```

Стан компонента `CheckUserAge` має бути ініціалізований властивістю `userAge` і властивістю `input`, обидві встановлені в порожній рядок.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).state().input === '' &&
    Enzyme.mount(React.createElement(CheckUserAge)).state().userAge === ''
);
```

Коли компонент `CheckUserAge` вперше відтворюється в DOM, внутрішній текст `button` має бути Submit.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('button').text() ===
    'Submit'
);
```

Коли у елемент `input` введено число менше за 18 і натиснуто `button`, внутрішній текст `button` має бути `You Shall Not Pass`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter3AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '3' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter17AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '17' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge3 = enter3AndClickButton();
  const userAge17 = enter17AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge3 === 'You Shall Not Pass' &&
      userAge17 === 'You Shall Not Pass'
  );
})();
```

Коли у елемент `input` введено число, яке дорівнює або більше за 18, і натиснуто `button`, внутрішній текст `button` має бути `You May Enter`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter35AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '35' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const userAge35 = enter35AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge18 === 'You May Enter' &&
      userAge35 === 'You May Enter'
  );
})();
```

Після того, як число було відправлено, і значення `input` знову змінено, `button` має повернутися до тексту `Submit`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const changeInputDontClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '5' } });
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter10AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '10' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const changeInput1 = changeInputDontClickButton();
  const userAge10 = enter10AndClickButton();
  const changeInput2 = changeInputDontClickButton();
  assert(
    userAge18 === 'You May Enter' &&
      changeInput1 === 'Submit' &&
      userAge10 === 'You Shall Not Pass' &&
      changeInput2 === 'Submit'
  );
})();
```

Ваш код не повинен містити жодних `if/else` інструкцій.

```js
assert(
  new RegExp(/(\s|;)if(\s|\()/).test(
    code
  ) === false
);
```

# --seed--

## --seed-contents--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userAge: '',
      input: ''
    };
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {this.state.userAge === ''
          ? buttonOne
          : this.state.userAge >= 18
          ? buttonTwo
          : buttonThree}
      </div>
    );
  }
}
```
