---
id: 5a24c314108439a4d4036172
title: Відтворення стану в інтерфейсі користувача іншим способом
challengeType: 6
forumTopicId: 301408
dashedName: render-state-in-the-user-interface-another-way
---

# --description--

Існує інший спосіб отримати доступ до `state` у компоненті. У методі `render()`, перед оператором `return`, ви можете писати JavaScript безпосередньо. Наприклад, ви можете оголосити функції, отримати дані з `state` або `props`, виконувати обчислення над цими даними тощо. Потім ви можете призначити будь-які дані змінним, до яких матимете доступ у операторі `return`.

# --instructions--

У методі render `MyComponent` визначте змінну `const` з ім’ям `name` і призначте їй значення name зі стану компонента `state`. Оскільки ви можете писати JavaScript безпосередньо в цій частині коду, не потрібно брати цей референс у фігурні дужки.

Далі, у операторі return відтворіть це значення в тегу `h1` за допомогою змінної `name`. Пам’ятайте, що у операторі return потрібно використовувати JSX-синтаксис (фігурні дужки для JavaScript).

# --hints--

`MyComponent` має мати ключ `name` зі значенням `freeCodeCamp`, збереженим у своєму стані.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

`MyComponent` має відтворювати заголовок `h1`, укладений у один `div`.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

Відтворений тег `h1` має містити референс до `{name}`.

```js
assert(/<h1>\n*\s*\{\s*name\s*\}\s*\n*<\/h1>/.test(code));
```

Відтворений заголовок `h1` має містити текст, відтворений зі стану компонента.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(firstValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line

    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line
    const name = this.state.name;
    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
