---
id: 5a24c314108439a4d4036189
title: Змінити рядковий CSS умовно залежно від стану компонента
challengeType: 6
forumTopicId: 301380
dashedName: change-inline-css-conditionally-based-on-component-state
---

# --description--

На цьому етапі ви вже бачили кілька застосувань умовного відтворення та використання рядкових стилів. Ось ще один приклад, який поєднує обидві ці теми. Ви також можете умовно відтворювати CSS залежно від стану React-компонента. Для цього ви перевіряєте умову, і якщо ця умова виконується, ви змінюєте об’єкт стилів, який призначений JSX-елементам у методі відтворення.

Ця парадигма важлива для розуміння, бо вона є суттєвим відходом від більш традиційного підходу застосування стилів шляхом безпосередньої зміни DOM-елементів (що дуже поширено, наприклад, у jQuery). У тому підході потрібно відстежувати, коли елементи змінюються, а також безпосередньо керувати маніпуляціями. Це може ускладнити відстеження змін і зробити інтерфейс непередбачуваним. Коли ви встановлюєте об’єкт стилів на основі умови, ви описуєте, як інтерфейс має виглядати як функцію стану застосунку. Існує чіткий потік інформації, який рухається лише в одному напрямку. Це переважний метод при написанні застосунків на React.

# --instructions--

У редакторі коду є простий контрольований компонент введення з стилізованим кордоном. Ви хочете стилізувати цей кордон червоним, якщо користувач введе більше ніж 15 символів тексту у блок введення. Додайте умову для перевірки цього, і якщо умова виконується, встановіть стиль кордону введення в `3px solid red`. Ви можете перевірити це, ввівши текст у поле введення.

# --hints--

Компонент `GateKeeper` має відтворювати елемент `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.find('div').length === 1;
  })()
);
```

Компонент `GateKeeper` має бути ініціалізований зі станом, у якому ключ `input` встановлений у порожній рядок.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.state().input === '';
  })()
);
```

Компонент `GateKeeper` має відтворювати тег `h3` і тег `input`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('h3').length === 1 &&
      mockedComponent.find('input').length === 1
    );
  })()
);
```

Тег `input` спочатку має мати стиль `1px solid black` для властивості `border`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('input').props().style.border === '1px solid black'
    );
  })()
);
```

Тег `input` має бути стилізований кордоном `3px solid red`, якщо значення введення у стані довше за 15 символів.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  let initialStyle = mockedComponent.find('input').props().style.border;
  const state_1 = () => {
    mockedComponent.setState({ input: 'this is 15 char' });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const state_2 = () => {
    mockedComponent.setState({
      input: 'A very long string longer than 15 characters.'
    });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const style_1 = await state_1();
  const style_2 = await state_2();
  assert(
    initialStyle === '1px solid black' &&
      style_1 === '1px solid black' &&
      style_2 === '3px solid red'
  );
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<GateKeeper />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    // Change code below this line

    // Change code above this line
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    if (this.state.input.length > 15) {
      inputStyle.border = '3px solid red';
    };
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```
