---
id: 5a24c314108439a4d4036169
title: Передача пропсів у безстанний функціональний компонент
challengeType: 6
forumTopicId: 301402
dashedName: pass-props-to-a-stateless-functional-component
---

# --description--

Попередні завдання охоплювали багато тем про створення та компонування JSX-елементів, функціональних компонентів і класових компонентів у стилі ES6 у React. З цією базою настав час розглянути ще одну дуже поширену функцію в React: **пропси**. У React ви можете передавати пропси, або властивості, дочірнім компонентам. Припустимо, у вас є компонент `App`, який відтворює дочірній компонент `Welcome`, що є безстанним функціональним компонентом. Ви можете передати `Welcome` властивість `user`, написавши:

```jsx
<App>
  <Welcome user='Mark' />
</App>
```

Ви використовуєте **кастомні HTML-атрибути**, створені вами і підтримувані React, щоб передати їх компоненту. У цьому випадку створена властивість `user` передається компоненту `Welcome`. Оскільки `Welcome` є безстанним функціональним компонентом, він має доступ до цього значення так:

```jsx
const Welcome = (props) => <h1>Hello, {props.user}!</h1>
```

Зазвичай це значення називають `props`, і при роботі з безстанними функціональними компонентами ви фактично розглядаєте його як аргумент функції, яка повертає JSX. Ви можете отримати доступ до значення аргументу в тілі функції. У класових компонентах це трохи інакше.

# --instructions--

У редакторі коду є `Calendar` і `CurrentDate` компоненти. Під час відтворення `CurrentDate` з компонента `Calendar` передайте властивість `date`, призначену поточній даті з об’єкта `Date` JavaScript. Потім отримайте доступ до цього `prop` у компоненті `CurrentDate`, показуючи його значення всередині тегів `p`. Зверніть увагу, що для того, щоб значення `prop` оцінювалися як JavaScript, їх потрібно брати в фігурні дужки, наприклад `date={Date()}`.

# --hints--

Компонент `Calendar` має повертати один елемент `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Другим дочірнім елементом компонента `Calendar` має бути компонент `CurrentDate`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).name() === 'CurrentDate';
  })()
);
```

Компонент `CurrentDate` має мати проп `date`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).props().date;
  })()
);
```

Проп `date` компонента `CurrentDate` має містити рядок тексту.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    const prop = mockedComponent.children().childAt(1).props().date;
    return typeof prop === 'string' && prop.length > 0;
  })()
);
```

Проп `date` має бути створений викликом `Date()`

```js
assert(/<CurrentDatedate={Date\(\)}\/>/.test(__helpers.removeWhiteSpace(code)));
```

Компонент `CurrentDate` має відтворювати значення пропа `date` у тегу `p`.

```js
let date = 'dummy date';
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(CurrentDate, { date })
    );
    return mockedComponent.find('p').html().includes(date);
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<Calendar />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: </p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: {props.date}</p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate date={Date()} />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
