---
id: 5a24c314108439a4d4036173
title: Встановлення стану за допомогою this.setState
challengeType: 6
forumTopicId: 301412
dashedName: set-state-with-this-setstate
---


# --description--

Попередні завдання охоплювали компонент `state` і як ініціалізувати стан у `constructor`. Існує також спосіб змінити властивість `state` компонента. React надає метод для оновлення компонентного `state`, який називається `setState`. Ви викликаєте метод `setState` у своєму класі компонента ось так: `this.setState()`, передаючи об’єкт із парами ключ-значення. Ключі — це ваші властивості `state`, а значення — оновлені дані стану. Наприклад, якщо ми зберігали `username` у `state` і хотіли його оновити, це виглядало б так:

```jsx
this.setState({
  username: 'Lewis'
});
```

React очікує, що ви ніколи не змінюватимете `state` безпосередньо, натомість завжди використовуйте `this.setState()` при зміні стану. Також слід зауважити, що React може об’єднувати кілька оновлень стану для покращення продуктивності. Це означає, що оновлення стану через метод `setState` можуть бути асинхронними. Існує альтернативний синтаксис для методу `setState`, який дозволяє обійти цю проблему. Це рідко потрібно, але добре про це пам’ятати! Будь ласка, ознайомтеся з нашим <a href="https://www.freecodecamp.org/news/what-is-state-in-react-explained-with-examples/" target="_blank" rel="noopener noreferrer nofollow">статтею про React</a> для детальнішої інформації.

# --instructions--

У редакторі коду є елемент `button`, який має обробник `onClick()`. Цей обробник спрацьовує, коли `button` отримує подію кліку в браузері, і запускає метод `handleClick`, визначений у `MyComponent`. Усередині методу `handleClick` оновіть компонент `state`, використовуючи `this.setState()`. Встановіть властивість `name` у `state` рівною рядку `React Rocks!`.

Натисніть кнопку і спостерігайте, як оновлюється відтворений стан. Не хвилюйтеся, якщо ви поки що не повністю розумієте, як працює код обробника кліку. Це буде розглянуто в наступних завданнях.

# --hints--

Стан `MyComponent` має ініціалізуватися парою ключ-значення `{ name: Initial State }`.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'Initial State'
);
```

`MyComponent` має відтворити заголовок `h1`.

```js
assert(Enzyme.mount(React.createElement(MyComponent)).find('h1').length === 1);
```

Відтворений заголовок `h1` має містити текст, відтворений зі стану компонента.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(/<h1>TestName<\/h1>/.test(firstValue));
```

Виклик методу `handleClick` на `MyComponent` має встановити властивість name у стані рівною `React Rocks!`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'Before' });
    return waitForIt(() => mockedComponent.state('name'));
  };
  const second = () => {
    mockedComponent.instance().handleClick();
    return waitForIt(() => mockedComponent.state('name'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Before' && secondValue === 'React Rocks!');
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
     // Change code below this line
    this.setState({
      name: 'React Rocks!'
    });
    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```
