---
id: 5a24c314108439a4d4036163
title: Створення компонента React
challengeType: 6
forumTopicId: 301386
dashedName: create-a-react-component
---

# --description--

Інший спосіб визначити компонент React — це використання синтаксису ES6 `class`. У наведеному нижче прикладі `Kitten` розширює `React.Component`:

```jsx
class Kitten extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <h1>Hi</h1>
    );
  }
}
```

Це створює ES6 клас `Kitten`, який розширює клас `React.Component`. Таким чином, клас `Kitten` тепер має доступ до багатьох корисних функцій React, таких як локальний стан і хуки життєвого циклу. Не хвилюйтеся, якщо ви ще не знайомі з цими термінами — вони будуть детальніше розглянуті в наступних завданнях. Також зверніть увагу, що клас `Kitten` має визначений у ньому `constructor`, який виконує виклик `super()`. Він використовує `super()` для виклику конструктора батьківського класу, у цьому випадку `React.Component`. Конструктор — це спеціальний метод, який використовується під час ініціалізації об’єктів, створених за допомогою ключового слова `class`. Найкраща практика — викликати конструктор компонента з `constructor` за допомогою `super` і передавати `props` обом. Це гарантує правильну ініціалізацію компонента. Поки що знайте, що цей код є стандартним. Незабаром ви побачите інші способи використання конструктора, а також `props`.

# --instructions--

`MyComponent` визначений у редакторі коду за допомогою синтаксису класу. Завершіть написання методу `render`, щоб він повертав елемент `div`, який містить `h1` з текстом `Hello React!`.

# --hints--

Компонент React має повертати елемент `div`.

```js
assert(Enzyme.shallow(React.createElement(MyComponent)).type() === 'div');
```

Повернутий `div` має відтворювати всередині себе заголовок `h1`.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.shallow(React.createElement(MyComponent)).html()
  )
);
```

Заголовок `h1` має містити рядок `Hello React!`.

```js
assert(
  Enzyme.shallow(React.createElement(MyComponent)).html() ===
    '<div><h1>Hello React!</h1></div>'
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line



    // Change code above this line
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line
    return (
      <div>
        <h1>Hello React!</h1>
      </div>
    );
    // Change code above this line
  }
};
```
