---
id: 5a24c314108439a4d4036180
title: Optimize Re-Renders with shouldComponentUpdate
challengeType: 6
forumTopicId: 301398
dashedName: optimize-re-renders-with-shouldcomponentupdate
---

# --description--

До цього моменту, якщо будь-який компонент отримує нові `state` або нові `props`, він повторно відтворює себе та всі свої дочірні елементи. Зазвичай це нормально. Але React надає метод життєвого циклу, який ви можете викликати, коли дочірні компоненти отримують нові `state` або `props`, і конкретно оголосити, чи слід оновлювати компоненти чи ні. Цей метод називається `shouldComponentUpdate()`, і він приймає `nextProps` та `nextState` як параметри.

Цей метод є корисним способом оптимізувати продуктивність. Наприклад, поведінка за замовчуванням полягає в тому, що ваш компонент повторно відтворюється, коли отримує нові `props`, навіть якщо `props` не змінилися. Ви можете використати `shouldComponentUpdate()`, щоб запобігти цьому, порівнюючи `props`. Метод має повертати значення `boolean`, яке повідомляє React, чи оновлювати компонент. Ви можете порівняти поточні пропси (`this.props`) з наступними пропсами (`nextProps`), щоб визначити, чи потрібно оновлювати, і відповідно повертати `true` або `false`.

# --instructions--

Метод `shouldComponentUpdate()` додано в компонент під назвою `OnlyEvens`. Наразі цей метод повертає `true`, тому `OnlyEvens` повторно відтворюється щоразу, коли отримує нові `props`. Змініть метод так, щоб `OnlyEvens` оновлювався лише тоді, коли `value` його нових пропсів є парним числом. Натисніть кнопку `Add` і спостерігайте порядок подій у консолі браузера, коли викликаються хуки життєвого циклу.

# --hints--

Компонент `Controller` має відтворювати компонент `OnlyEvens` як дочірній елемент.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Controller));
    return (
      mockedComponent.find('Controller').length === 1 &&
      mockedComponent.find('OnlyEvens').length === 1
    );
  })()
);
```

Метод `shouldComponentUpdate` має бути визначений у компоненті `OnlyEvens`.

```js
assert(
  (() => {
    const child = React.createElement(OnlyEvens)
      .type.prototype.shouldComponentUpdate.toString()
      .replace(/s/g, '');
    return child !== 'undefined';
  })()
);
```

Компонент `OnlyEvens` має повертати тег `h1`, який відтворює значення `this.props.value`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 1000 });
    return mockedComponent.find('h1').html();
  };
  const second = () => {
    mockedComponent.setState({ value: 10 });
    return mockedComponent.find('h1').html();
  };
  const firstValue = first();
  const secondValue = second();
  assert(firstValue === '<h1>1000</h1>' && secondValue === '<h1>10</h1>');
})();
```

`OnlyEvens` має повторно відтворюватися лише тоді, коли `nextProps.value` є парним числом.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 8 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ value: 7 });
    return mockedComponent.find('h1').text();
  };
  const third = () => {
    mockedComponent.setState({ value: 42 });
    return mockedComponent.find('h1').text();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(firstValue === '8' && secondValue === '8' && thirdValue === '42');
})();
```

# --seed--

## --seed-contents--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return true;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```

# --solutions--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return nextProps.value % 2 === 0;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```
