---
id: 5a24c314108439a4d4036188
title: Відтворення умовно з пропсів
challengeType: 6
forumTopicId: 301405
dashedName: render-conditionally-from-props
---

# --description--

До цього моменту ви бачили, як використовувати `if/else`, `&&` і тернарний оператор (`condition ? expressionIfTrue : expressionIfFalse`) для прийняття умовних рішень про те, що і коли відтворювати. Однак залишилася одна важлива тема, яку варто обговорити — вона дозволяє поєднати будь-які або всі ці концепції з іншою потужною функцією React: пропсами. Використання пропсів для умовного відтворення коду дуже поширене серед розробників React — тобто вони використовують значення певного пропса, щоб автоматично приймати рішення про те, що відтворювати.

У цьому завданні ви налаштуєте дочірній компонент, щоб він приймав рішення про відтворення на основі пропсів. Ви також використаєте тернарний оператор, але побачите, що кілька інших концепцій, розглянутих у попередніх завданнях, можуть бути так само корисними в цьому контексті.

# --instructions--

У редакторі коду є два компоненти, які частково визначені для вас: батьківський під назвою `GameOfChance` і дочірній під назвою `Results`. Вони використовуються для створення простої гри, де користувач натискає кнопку, щоб дізнатися, чи виграв він, чи програв.

Спершу вам потрібен простий вираз, який випадковим чином повертає різне значення щоразу, коли його запускають. Ви можете використати `Math.random()`. Цей метод повертає значення між `0` (включно) і `1` (виключно) щоразу, коли його викликають. Тож для шансів 50/50 використайте `Math.random() >= .5` у вашому виразі. Статистично цей вираз повертатиме `true` 50% часу і `false` інші 50%. У методі render замініть `null` на наведений вище вираз, щоб завершити оголошення змінної.

Тепер у вас є вираз, який можна використати для випадкового прийняття рішення в коді. Далі потрібно це реалізувати. Відтворіть компонент `Results` як дочірній елемент `GameOfChance` і передайте `expression` як проп з ім’ям `fiftyFifty`. У компоненті `Results` напишіть тернарний вираз, щоб відтворити елемент `h1` з текстом `You Win!` або `You Lose!` залежно від пропса `fiftyFifty`, який передається з `GameOfChance`. Нарешті, переконайтеся, що метод `handleClick()` правильно рахує кожен хід, щоб користувач знав, скільки разів він грав. Це також дає знати користувачу, що компонент справді оновився, якщо він виграв або програв двічі підряд.

# --hints--

Компонент `GameOfChance` має існувати і відтворюватися на сторінці.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('GameOfChance').length,
  1
);
```

`GameOfChance` має повертати один елемент `button`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('button').length,
  1
);
```

`GameOfChance` має повертати один екземпляр компонента `Results`, який має проп з ім’ям `fiftyFifty`.

```js
assert(
  Enzyme.mount(React.createElement(GameOfChance)).find('Results').length ===
    1 &&
    Enzyme.mount(React.createElement(GameOfChance))
      .find('Results')
      .props()
      .hasOwnProperty('fiftyFifty') === true
);
```

Стан `GameOfChance` має бути ініціалізований властивістю `counter` зі значенням `1`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).state().counter,
  1
);
```

Коли компонент `GameOfChance` вперше відтворюється в DOM, має повертатися елемент `p` з внутрішнім текстом `Turn: 1`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('p').text(),
  'Turn: 1'
);
```

Кожного разу, коли кнопку натискають, стан лічильника має збільшуватися на 1, і в DOM має відтворюватися один елемент `p` з текстом `Turn: N`, де `N` — це значення стану лічильника.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    count: comp.state('counter'),
    text: comp.find('p').text()
  });
  const _1 = () => {
    simulate();
    return result();
  };
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _1_val = _1();
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  assert(
    _1_val.count === 2 &&
      _1_val.text === 'Turn: 2' &&
      _2_val.count === 3 &&
      _2_val.text === 'Turn: 3' &&
      _3_val.count === 4 &&
      _3_val.text === 'Turn: 4' &&
      _4_val.count === 5 &&
      _4_val.text === 'Turn: 5' &&
      _5_val.count === 6 &&
      _5_val.text === 'Turn: 6'
  );
})();
```

Коли компонент `GameOfChance` вперше монтується в DOM і кожного разу після натискання кнопки, має повертатися один елемент `h1`, який випадковим чином відтворює або `You Win!`, або `You Lose!`. Примітка: це може випадково не спрацювати. Якщо так станеться, спробуйте ще раз.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    h1: comp.find('h1').length,
    text: comp.find('h1').text()
  });
  const _1 = result();
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _6 = () => {
    simulate();
    return result();
  };
  const _7 = () => {
    simulate();
    return result();
  };
  const _8 = () => {
    simulate();
    return result();
  };
  const _9 = () => {
    simulate();
    return result();
  };
  const _10 = () => {
    simulate();
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  const _6_val = _6();
  const _7_val = _7();
  const _8_val = _8();
  const _9_val = _9();
  const _10_val = _10();
  const __text = new Set([
    _1.text,
    _2_val.text,
    _3_val.text,
    _4_val.text,
    _5_val.text,
    _6_val.text,
    _7_val.text,
    _8_val.text,
    _9_val.text,
    _10_val.text
  ]);
  const __h1 = new Set([
    _1.h1,
    _2_val.h1,
    _3_val.h1,
    _4_val.h1,
    _5_val.h1,
    _6_val.h1,
    _7_val.h1,
    _8_val.h1,
    _9_val.h1,
    _10_val.h1
  ]);
  assert(__text.size === 2 && __h1.size === 1);
})();
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<GameOfChance />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    {/* Change code below this line */}
    return <h1></h1>;
    {/* Change code above this line */}
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      // Complete the return statement:
      return {
        counter: prevState
      }
    });
  }
  render() {
    const expression = null; // Change this line
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        {/* Change code below this line */}

        {/* Change code above this line */}
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```

# --solutions--

```jsx
// We want this to be deterministic for testing purposes.
const randomSequence = [true, false, false, true, true, false, false, true, true, false];
let index = 0;
const fiftyFifty = () => randomSequence[index++ % randomSequence.length];

class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h1>{this.props.fiftyFifty ? 'You Win!' : 'You Lose!'}</h1>;
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      return {
        counter: prevState.counter + 1
      }
    });
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        <Results fiftyFifty={fiftyFifty()} />
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```
