---
id: 5a24c314108439a4d403616c
title: Заміщення пропсів за замовчуванням
challengeType: 6
forumTopicId: 301399
dashedName: override-default-props
---

# --description--

Можливість встановлювати пропси за замовчуванням — це корисна функція в React. Щоб замістити пропси за замовчуванням, потрібно явно задати значення пропсів для компонента.

# --instructions--

Компонент `ShoppingCart` тепер відтворює дочірній компонент `Items`. Цей компонент `Items` має проп `quantity` за замовчуванням, встановлений у ціле число `0`. Замініть проп за замовчуванням, передавши значення `10` для `quantity`.

**Примітка:** Пам’ятайте, що синтаксис додавання пропса до компонента схожий на додавання атрибутів HTML. Однак, оскільки значення для `quantity` є цілим числом, його не потрібно брати в лапки, але слід обгорнути у фігурні дужки. Наприклад, `{100}`. Цей синтаксис каже JSX інтерпретувати значення всередині дужок безпосередньо як JavaScript.

# --hints--

Компонент `ShoppingCart` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Компонент `Items` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

Компонент `Items` має отримувати проп `{ quantity: 10 }` від компонента `ShoppingCart`.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
      return (
        mockedComponent.find('Items').props().quantity == 10 &&
        code.replace(/ /g, '')
          .includes('<Itemsquantity={10}/>')
      );
    })()
  );
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<ShoppingCart />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items />
    { /* Change code above this line */ }
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items quantity = {10} />
    { /* Change code above this line */ }
  }
};
```
