---
id: 5a24c314108439a4d403617e
title: Add Event Listeners
challengeType: 6
forumTopicId: 301377
dashedName: add-event-listeners
---

# --description--

Метод `componentDidMount()` також є найкращим місцем, щоб приєднати будь-які слухачі подій, які потрібно додати для конкретної функціональності. React надає синтетичну систему подій, яка обгортає нативну систему подій, присутню в браузерах. Це означає, що синтетична система подій поводиться однаково незалежно від браузера користувача — навіть якщо нативні події можуть поводитися по-різному в різних браузерах.

Ви вже використовували деякі з цих синтетичних обробників подій, таких як `onClick()`. Синтетична система подій React чудово підходить для більшості взаємодій, які ви будете керувати на DOM-елементах. Однак, якщо ви хочете приєднати обробник подій до об’єктів document або window, потрібно робити це безпосередньо.

# --instructions--

Приєднайте слухача подій у методі `componentDidMount()` для подій `keydown` і зробіть так, щоб ці події викликали функцію зворотного виклику `handleKeyPress()`. Ви можете використати `document.addEventListener()`, який приймає подію (у лапках) як перший аргумент і функцію зворотного виклику як другий аргумент.

Потім у `componentWillUnmount()` видаліть цей самий слухач подій. Ви можете передати ті самі аргументи у `document.removeEventListener()`. Доброю практикою є використання цього методу життєвого циклу для будь-якого очищення компонентів React перед їхнім демонтажем і знищенням. Видалення слухачів подій — це приклад такого очищення.

# --hints--

`MyComponent` має відтворити `div` елемент, який обгортає тег `h1`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').children().find('h1').length === 1;
  })()
);
```

Слухач `keydown` має бути приєднаний до document у `componentDidMount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const didMountString = mockedComponent
      .instance()
      .componentDidMount.toString();
    return new RegExp(
      'document.addEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(didMountString);
  })()
);
```

Слухач `keydown` має бути видалений з document у `componentWillUnmount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const willUnmountString = mockedComponent
      .instance()
      .componentWillUnmount.toString();
    return new RegExp(
      'document.removeEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(willUnmountString);
  })()
);
```

Після того, як компонент змонтується, натискання `enter` має оновити його стан і відтворений тег `h1`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const beforeState = mockedComponent.state('message');
  const beforeText = mockedComponent.find('h1').text();
  const pressEnterKey = () => {
    mockedComponent.instance().handleKeyPress({ keyCode: 13 });
    return waitForIt(() => {
      mockedComponent.update();
      return {
        state: mockedComponent.state('message'),
        text: mockedComponent.find('h1').text()
      };
    });
  };
  const afterKeyPress = await pressEnterKey();
  assert(
    beforeState !== afterKeyPress.state && beforeText !== afterKeyPress.text
  );
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleEnter = this.handleEnter.bind(this);
    this.handleKeyPress = this.handleKeyPress.bind(this);
  }
  // Change code below this line
  componentDidMount() {

  }
  componentWillUnmount() {

  }
  // Change code above this line
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleKeyPress = this.handleKeyPress.bind(this);
    this.handleEnter = this.handleEnter.bind(this);  }
  componentDidMount() {
    // Change code below this line
    document.addEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  componentWillUnmount() {
    // Change code below this line
    document.removeEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```
