---
id: 5a24c314108439a4d4036176
title: Використання стану для перемикання елемента
challengeType: 6
forumTopicId: 301421
dashedName: use-state-to-toggle-an-element
---

# --description--

Іноді потрібно знати попередній стан під час оновлення стану. Однак оновлення стану можуть бути асинхронними — це означає, що React може об’єднувати кілька викликів `setState()` в одне оновлення. Це означає, що не можна покладатися на попереднє значення `this.state` або `this.props` при обчисленні наступного значення. Тому не слід використовувати код на кшталт цього:

```jsx
this.setState({
  counter: this.state.counter + this.props.increment
});
```

Натомість слід передати `setState` функцію, яка дозволяє отримати доступ до стану та пропсів. Використання функції з `setState` гарантує, що ви працюєте з найактуальнішими значеннями стану та пропсів. Це означає, що наведене вище слід переписати так:

```jsx
this.setState((state, props) => ({
  counter: state.counter + props.increment
}));
```

Ви також можете використовувати форму без `props`, якщо вам потрібне лише `state`:

```jsx
this.setState(state => ({
  counter: state.counter + 1
}));
```

Зверніть увагу, що об’єктний літерал потрібно обгорнути в круглі дужки, інакше JavaScript вважатиме це блоком коду.

# --instructions--

`MyComponent` має властивість `visibility`, яка ініціалізується значенням `false`. Метод відтворення повертає один вигляд, якщо значення `visibility` істинне, і інший вигляд, якщо воно хибне.

Наразі немає способу оновити властивість `visibility` у стані компонента `state`. Значення має перемикатися між true і false. На кнопці є обробник кліку, який викликає метод класу з назвою `toggleVisibility()`. Передайте функцію в `setState`, щоб визначити цей метод так, щоб стан `state` `visibility` перемикався на протилежне значення під час виклику методу. Якщо `visibility` дорівнює `false`, метод встановлює його в `true`, і навпаки.

Нарешті, натисніть кнопку, щоб побачити умовне відтворення компонента на основі його `state`.

**Підказка:** Не забудьте прив’язати ключове слово `this` до методу в `constructor`!

# --hints--

`MyComponent` має повертати елемент `div`, який містить `button`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('div').find('button')
    .length,
  1
);
```

Стан `MyComponent` має ініціалізуватися властивістю `visibility` зі значенням `false`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).state('visibility'),
  false
);
```

Натискання на елемент кнопки має перемикати властивість `visibility` у стані між `true` і `false` та умовно відтворювати елемент `h1`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ visibility: false });
    return mockedComponent.state('visibility');
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const third = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(!firstValue && secondValue && !thirdValue);
})();
```

Анонімна функція має бути передана в `setState`.

```js
const paramRegex = '[a-zA-Z$_]\\w*(,[a-zA-Z$_]\\w*)?';
assert(
  new RegExp(
    'this\\.setState\\((function\\(' +
      paramRegex +
      '\\){|([a-zA-Z$_]\\w*|\\(' +
      paramRegex +
      '\\))=>)'
  ).test(__helpers.removeWhiteSpace(code))
);
```

`this` не слід використовувати всередині `setState`

```js
assert(!/this\.setState\([^}]*this/.test(code));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    this.toggleVisibility = this.toggleVisibility.bind(this);
  }
  toggleVisibility() {
    this.setState(state => ({
      visibility: !state.visibility
    }));
  }
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```
