---
id: 5a24c314108439a4d4036165
title: Використання React для відтворення вкладених компонентів
challengeType: 6
forumTopicId: 301420
dashedName: use-react-to-render-nested-components
---

# --description--

Останнє завдання показало простий спосіб складання двох компонентів, але існує багато різних способів, як можна складати компоненти за допомогою React.

Складання компонентів — одна з потужних функцій React. Коли ви працюєте з React, важливо починати мислити про ваш інтерфейс користувача у термінах компонентів, як у прикладі застосунку в останньому завданні. Ви розбиваєте свій інтерфейс на базові будівельні блоки, і ці частини стають компонентами. Це допомагає відокремити код, відповідальний за інтерфейс, від коду, який відповідає за логіку вашого застосунку. Це значно спрощує розробку та обслуговування складних проєктів.

# --instructions--

У редакторі коду визначено два функціональні компоненти з іменами `TypesOfFruit` та `Fruits`. Візьміть компонент `TypesOfFruit` і складіть його, або *вкладіть*, у компонент `Fruits`. Потім вкладіть компонент `Fruits` у компонент `TypesOfFood`. Результатом має бути дочірній компонент, вкладений у батьківський компонент, який, у свою чергу, вкладений у власний батьківський компонент!

# --hints--

Компонент `TypesOfFood` має повертати один елемент `div`.

```js
assert(Enzyme.shallow(React.createElement(TypesOfFood)).type() === 'div');
```

Компонент `TypesOfFood` має повертати компонент `Fruits`.

```js
assert(
  Enzyme.shallow(React.createElement(TypesOfFood)).props().children[1].type
    .name === 'Fruits'
);
```

Компонент `Fruits` має повертати компонент `TypesOfFruit`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('h2').html() ===
    '<h2>Fruits:</h2>'
);
```

Компонент `TypesOfFruit` має повертати елементи `h2` та `ul`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('ul').text() ===
    'ApplesBlueberriesStrawberriesBananas'
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<TypesOfFood />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }

      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }
        <TypesOfFruit />
      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }
        <Fruits />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
