---
id: 5a24c314108439a4d4036182
title: Add Inline Styles in React
challengeType: 6
forumTopicId: 301378
dashedName: add-inline-styles-in-react
---

# --description--

Ви могли помітити в останньому завданні, що існувало кілька інших відмінностей у синтаксисі від HTML рядкових стилів, окрім того, що атрибут `style` призначався об’єкту JavaScript. По-перше, назви деяких CSS властивостей стилю використовують camel case. Наприклад, в останньому завданні розмір шрифту задавався через `fontSize` замість `font-size`. Слова, написані через дефіс, як `font-size`, є некоректним синтаксисом для властивостей об’єктів JavaScript, тому React використовує camel case. За правилом, будь-які властивості стилю, написані через дефіс, у JSX записуються camel case.

Усі одиниці довжини значень властивостей (наприклад, `height`, `width` і `fontSize`) за замовчуванням вважаються в `px`, якщо не вказано інше. Якщо ви хочете використати `em`, наприклад, потрібно обгорнути значення та одиниці в лапки, як `{fontSize: "4em"}`. Окрім значень довжини, які за замовчуванням у `px`, усі інші значення властивостей слід обгорнути в лапки.

# --instructions--

Якщо у вас є великий набір стилів, ви можете призначити множину стилів `object` константі, щоб тримати свій код організованим. Оголосіть константу стилів як глобальну змінну на початку файлу. Ініціалізуйте константу `styles` і призначте їй об’єкт `object` з трьома властивостями стилю та їхніми значеннями. Дайте `div` колір `purple`, розмір шрифту `40` і кордон `2px solid purple`. Потім встановіть атрибут `style` рівним константі `styles`.

# --hints--

Змінна `styles` має бути об’єктом `object` з трьома властивостями.

```js
assert(Object.keys(styles).length === 3);
```

Змінна `styles` має мати властивість `color`, встановлену в значення `purple`.

```js
assert(styles.color === 'purple');
```

Змінна `styles` має мати властивість `fontSize`, встановлену в значення `40`.

```js
assert(styles.fontSize == 40);
```

Змінна `styles` має мати властивість `border`, встановлену в значення `2px solid purple`.

```js
assert(styles.border === '2px solid purple');
```

Компонент має відтворити елемент `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return mockedComponent.type() === 'div';
  })()
);
```

Елемент `div` має мати стилі, визначені об’єктом `styles`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return (
      mockedComponent.props().style.color === 'purple' &&
      mockedComponent.props().style.fontSize == 40 &&
      mockedComponent.props().style.border === '2px solid purple'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={{color: "yellow", fontSize: 24}}>Style Me!</div>
    );
    // Change code above this line
  }
};
```

# --solutions--

```jsx
const styles = {
  color: "purple",
  fontSize: 40,
  border: "2px solid purple"
};
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={styles}>Style Me!</div>
    );
    // Change code above this line
  }
};
```
