---
id: 5a24c314108439a4d4036174
title: "Bind 'this' to a Class Method"
challengeType: 6
forumTopicId: 301379
dashedName: bind-this-to-a-class-method
---

# --description--

Окрім встановлення та оновлення `state`, ви також можете визначати методи для свого класу компонента. Метод класу зазвичай потребує використання ключового слова `this`, щоб мати доступ до властивостей класу (наприклад, `state` і `props`) всередині області видимості методу. Існує кілька способів дозволити методам класу отримувати доступ до `this`.

Одним із загальних способів є явне зв’язування `this` у конструкторі, щоб `this` став прив’язаний до методів класу під час ініціалізації компонента. Ви могли помітити, що в останньому завданні використовували `this.handleClick = this.handleClick.bind(this)` для методу `handleClick` у конструкторі. Потім, коли ви викликаєте функцію, як `this.setState()`, всередині методу класу, `this` посилається на клас і не буде `undefined`.

**Примітка:** Ключове слово `this` є одним із найзаплутаніших аспектів JavaScript, але воно відіграє важливу роль у React. Хоча його поведінка тут цілком звичайна, ці уроки не призначені для глибокого повторення `this`, тому, будь ласка, звертайтеся до інших уроків, якщо наведене вище здається заплутаним!

# --instructions--

Редактор коду має компонент із `state`, який відстежує текст. Також є метод, який дозволяє встановити текст у `You clicked!`. Однак метод не працює, бо використовує ключове слово `this`, яке є невизначеним. Виправте це, явно зв’язавши `this` з методом `handleClick()` у конструкторі компонента.

Далі додайте обробник події натискання на елемент `button` у методі відтворення. Він має викликати метод `handleClick()`, коли кнопка отримує подію натискання. Пам’ятайте, що метод, який ви передаєте обробнику `onClick`, потребує фігурних дужок, оскільки його слід інтерпретувати безпосередньо як JavaScript.

Після виконання вищезазначених кроків ви зможете натиснути кнопку і побачити `You clicked!`.

# --hints--

`MyComponent` має повертати елемент `div`, який обгортає два елементи — кнопку та елемент `h1` — у такому порядку.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).find('div').length === 1 &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(0)
      .type() === 'button' &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(1)
      .type() === 'h1'
);
```

Стан `MyComponent` має ініціалізуватися парою ключ-значення `{ text: "Hello" }`.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('text') === 'Hello'
);
```

Натискання на елемент `button` має запускати метод `handleClick` і встановлювати стан `text` у `You clicked!`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ text: 'Hello' });
    return waitForIt(() => mockedComponent.state('text'));
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return waitForIt(() => mockedComponent.state('text'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Hello' && secondValue === 'You clicked!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    // Change code below this line

    // Change code above this line
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }
        <button>Click Me</button>
        { /* Change code above this line */ }
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```
