---
id: 5a24c314108439a4d4036170
title: Create a Stateful Component
challengeType: 6
forumTopicId: 301391
dashedName: create-a-stateful-component
---

# --description--

Одна з найважливіших тем у React — це `state`. Стан складається з будь-яких даних, які ваш застосунок має знати і які можуть змінюватися з часом. Ви хочете, щоб ваші застосунки реагували на зміни стану і відображали оновлений інтерфейс користувача, коли це необхідно. React пропонує зручне рішення для керування станом сучасних веб-застосунків.

Ви створюєте стан у React-компоненті, оголошуючи властивість `state` у класі компонента в його `constructor`. Це ініціалізує компонент `state` під час його створення. Властивість `state` має бути встановлена як JavaScript-об’єкт `object`. Оголошення виглядає так:

```jsx
this.state = {

}
```

Ви маєте доступ до об’єкта `state` протягом усього життя вашого компонента. Ви можете оновлювати його, відтворювати у вашому інтерфейсі та передавати як пропси дочірнім компонентам. Об’єкт `state` може бути таким складним або простим, як вам потрібно. Зверніть увагу, що для створення `React.Component` таким чином потрібно створити класовий компонент, розширюючи `state`.

# --instructions--

У редакторі коду є компонент, який намагається відтворити властивість `firstName` зі свого `state`. Однак стан `state` не визначено. Ініціалізуйте компонент `state` у `constructor` і призначте своє ім’я властивості `firstName`.

# --hints--

`StatefulComponent` має існувати і відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return mockedComponent.find('StatefulComponent').length === 1;
  })()
);
```

`StatefulComponent` має відтворювати елементи `div` та `h1`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

Стан `StatefulComponent` має бути ініціалізований властивістю `firstName`, встановленою як рядок.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return (
      typeof initialState === 'object' && typeof initialState.firstName === 'string'
    );
  })()
);
```

Властивість `firstName` у стані `StatefulComponent` має відтворюватися в елементі `h1`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return mockedComponent.find('h1').text() === initialState.firstName;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    // Only change code below this line

    // Only change code above this line
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      firstName: 'freeCodeCamp!'
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```
