---
id: 5a24c314108439a4d403617c
title: Використання методу життєвого циклу componentWillMount
challengeType: 6
forumTopicId: 301423
dashedName: use-the-lifecycle-method-componentwillmount
---

# --description--

React-компоненти мають кілька спеціальних методів, які дають змогу виконувати дії у певні моменти життєвого циклу компонента. Ці методи називаються методами життєвого циклу або хуками життєвого циклу, і вони дозволяють "ловити" компоненти у визначені моменти часу. Це може бути перед відтворенням, перед оновленням, перед отриманням пропсів, перед демонтажем і так далі. Ось список деяких основних методів життєвого циклу: `componentWillMount()` `componentDidMount()` `shouldComponentUpdate()` `componentDidUpdate()` `componentWillUnmount()` Наступні кілька уроків охоплять деякі базові випадки використання цих методів життєвого циклу.

**Примітка:** Метод життєвого циклу `componentWillMount` буде застарілим у майбутній версії 16.X і видалений у версії 17. Детальніше читайте в цій <a href="https://www.freecodecamp.org/news/how-to-safely-use-reacts-life-cycles-with-fiber-s-async-rendering-fd4469ebbd8f/" target="_blank" rel="noopener noreferrer nofollow">статті</a>

# --instructions--

Метод `componentWillMount()` викликається перед методом `render()`, коли компонент монтується в DOM. Запишіть щось у консоль у методі `componentWillMount()` — можливо, вам буде зручно мати відкриту консоль браузера, щоб побачити вихідні дані.

# --hints--

`MyComponent` має відтворити елемент `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').length === 1;
  })()
);
```

Метод `console.log` має бути викликаний у `componentWillMount`.

```js
assert(
  (function () {
    const lifecycle = React.createElement(MyComponent)
      .type.prototype.componentWillMount.toString()
      .replace(/ /g, '');
    return lifecycle.includes('console.log(');
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line
    console.log('Component is mounting...');
    // Change code above this line
  }
  render() {
    return <div />
  }
};
```
