---
id: 5a24c314108439a4d4036181
title: Вступ до рядкових стилів
challengeType: 6
forumTopicId: 301395
dashedName: introducing-inline-styles
---

# --description--

Існують інші складні концепції, які додають потужні можливості вашому React-коду. Але ви, можливо, замислюєтеся над більш простою проблемою — як стилізувати ті JSX-елементи, які ви створюєте в React. Ви, ймовірно, знаєте, що це не буде точно так само, як працювати з HTML через <a href="/learn/front-end-development-libraries/react/define-an-html-class-in-jsx" target="_blank" rel="noopener noreferrer nofollow">спосіб застосування класів до JSX-елементів</a>.

Якщо ви імпортуєте стилі з таблиці стилів, це майже не відрізняється. Ви застосовуєте клас до свого JSX-елемента за допомогою атрибута `className`, а стилі задаєте для класу у таблиці стилів. Інший варіант — застосувати рядкові стилі, які дуже поширені в розробці ReactJS.

Ви застосовуєте рядкові стилі до JSX-елементів подібно до того, як це робите в HTML, але з кількома відмінностями JSX. Ось приклад рядкового стилю в HTML:

```jsx
<div style="color: yellow; font-size: 16px">Mellow Yellow</div>
```

JSX-елементи використовують атрибут `style`, але через те, як JSX транслюється, ви не можете встановити значення як `string`. Замість цього ви встановлюєте його рівним JavaScript-об’єкту `object`. Ось приклад:

```jsx
<div style={{color: "yellow", fontSize: 16}}>Mellow Yellow</div>
```

Зверніть увагу, як ми використовуємо верблюжий регістр для властивості `fontSize`? Це тому, що React не приймає ключі у форматі kebab-case в об’єкті стилів. React автоматично застосує правильну назву властивості в HTML.

# --instructions--

Додайте атрибут `style` до `div` у редакторі коду, щоб надати тексту червоний колір і розмір шрифту `72px`.

Зверніть увагу, що ви можете необов’язково встановити розмір шрифту як число, опускаючи одиниці `px`, або записати його як `"72px"`.
# --hints--

Компонент має відтворити елемент `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Елемент `div` має мати колір `red`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().props().style.color === 'red';
  })()
);
```

Елемент `div` має мати розмір шрифту `72px`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return (
      mockedComponent.children().props().style.fontSize === 72 ||
      mockedComponent.children().props().style.fontSize === '72' ||
      mockedComponent.children().props().style.fontSize === '72px'
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<Colorful />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div>Big Red</div>
    );
  }
};
```

# --solutions--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div style={{color: "red", fontSize: 72}}>Big Red</div>
    );
  }
};
```
