---
id: 5a24c314108439a4d4036185
title: Use && for a More Concise Conditional
challengeType: 6
forumTopicId: 301413
dashedName: use--for-a-more-concise-conditional
---

# --description--

Інструкції `if/else` працювали в попередньому завданні, але існує більш лаконічний спосіб досягти того самого результату. Уявіть, що ви відстежуєте кілька умов у компоненті і хочете, щоб різні елементи відтворювалися залежно від кожної з цих умов. Якщо ви напишете багато інструкцій `else if` для повернення трохи різних інтерфейсів, ви можете повторювати код, що створює ризик помилок. Замість цього ви можете використовувати логічний оператор `&&` для виконання умовної логіки більш лаконічно. Це можливо, тому що ви хочете перевірити, чи є умова `true`, і якщо так, повернути деяку розмітку. Ось приклад:

```jsx
{condition && <p>markup</p>}
```

Якщо `condition` є `true`, розмітка буде повернена. Якщо умова `false`, операція одразу поверне `false` після оцінки `condition` і нічого не поверне. Ви можете включати ці інструкції безпосередньо у свій JSX і ланцюжити кілька умов, записуючи `&&` після кожної. Це дозволяє обробляти складнішу умовну логіку у вашому методі `render()` без повторення великої кількості коду.

# --instructions--

Розв’яжіть попередній приклад знову, так щоб `h1` відтворювався лише якщо `display` є `true`, але використайте логічний оператор `&&` замість інструкції `if/else`.

# --hints--

`MyComponent` має існувати і відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length;
  })()
);
```

Коли `display` встановлено в `true`, мають відтворюватися `div`, `button` і `h1`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 2 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 1
  );
```

Коли `display` встановлено в `false`, мають відтворюватися лише `div` і `button`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 1 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 0
  );
```

Метод render має використовувати логічний оператор `&&` для перевірки умови `this.state.display`.

```js
assert(code.includes('&&'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         {this.state.display && <h1>Displayed!</h1>}
       </div>
    );
  }
};
```
