---
id: 5a24c314108439a4d4036166
title: Compose React Components
challengeType: 6
forumTopicId: 301381
dashedName: compose-react-components
---

# --description--

Оскільки завдання продовжують використовувати більш складні композиції з компонентами React і JSX, є один важливий момент, який варто врахувати. Відтворення компонентів класу у стилі ES6 всередині інших компонентів нічим не відрізняється від відтворення простих компонентів, які ви використовували в останніх кількох завданнях. Ви можете відтворювати JSX-елементи, безстанні функціональні компоненти та компоненти класу ES6 всередині інших компонентів.

# --instructions--

У редакторі коду компонент `TypesOfFood` вже відтворює компонент під назвою `Vegetables`. Також є компонент `Fruits` з останнього завдання.

Вкладіть два компоненти всередину `Fruits` — спочатку `NonCitrus`, а потім `Citrus`. Обидва ці компоненти надаються вам за лаштунками. Далі вкладіть класовий компонент `Fruits` у компонент `TypesOfFood`, нижче заголовка `h1` і вище `Vegetables`. Результатом має бути серія вкладених компонентів, яка використовує два різні типи компонентів.

# --before-all--

```js
globalThis.NonCitrus = function NonCitrus() {
  return React.createElement('div', null);
};
globalThis.Citrus = function Citrus() {
  return React.createElement('div', null);
};
globalThis.Vegetables = function Vegetables() {
  return React.createElement('div', null);
};
```

# --hints--

Компонент `TypesOfFood` має повертати один елемент `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Компонент `TypesOfFood` має повертати компонент `Fruits`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

Компонент `Fruits` має повертати компонент `NonCitrus` і компонент `Citrus`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return (
      mockedComponent.find('Fruits').children().find('NonCitrus').length ===
        1 &&
      mockedComponent.find('Fruits').children().find('Citrus').length === 1
    );
  })()
);
```

Компонент `TypesOfFood` має повертати компонент `Vegetables` нижче компонента `Fruits`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
        <Vegetables />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }
        <NonCitrus />
        <Citrus />
        { /* Change code above this line */ }
      </div>
    )
  }
}

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
    render() {
      return (
        <div>
        <h1>Types of Food:</h1>
          { /* Change code below this line */ }
          <Fruits />
          { /* Change code above this line */ }
          <Vegetables />
        </div>
      );
    }
};
```
