---
id: 5a24c314108439a4d403618b
title: Give Sibling Elements a Unique Key Attribute
challengeType: 6
forumTopicId: 301394
dashedName: give-sibling-elements-a-unique-key-attribute
---

# --description--

У попередньому завданні показано, як метод `map` використовується для динамічного відтворення кількох елементів на основі вхідних даних користувача. Однак у тому прикладі бракувало важливої деталі. Коли ви створюєте масив елементів, кожен з них повинен мати атрибут `key`, встановлений унікальним значенням. React використовує ці ключі, щоб відстежувати, які елементи додані, змінені або видалені. Це допомагає зробити процес повторного відтворення більш ефективним, коли список змінюється будь-яким чином.

**Примітка:** Ключі повинні бути унікальними лише серед братських елементів, їм не потрібно бути унікальними у всьому застосунку.

# --instructions--

У редакторі коду є масив із кількома фронтенд-фреймворками та безстановий функціональний компонент із назвою `Frameworks()`. `Frameworks()` потрібно спроєктувати масив у невпорядкований список, подібно до попереднього завдання. Завершіть написання функції зворотного виклику `map`, щоб вона повертала елемент `li` для кожного фреймворку в масиві `frontEndFrameworks`. Цього разу переконайтеся, що кожен `li` має атрибут `key`, встановлений унікальним значенням. Елементи `li` також повинні містити текст із `frontEndFrameworks`.

Зазвичай ключ повинен бути чимось, що унікально ідентифікує елемент, який відтворюється. У крайньому випадку можна використовувати індекс масиву, але зазвичай слід намагатися використати унікальний ідентифікатор.

# --hints--

Компонент `Frameworks` має існувати і відтворюватися на сторінці.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('Frameworks').length === 1
);
```

`Frameworks` має відтворювати елемент `h1`.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('h1').length === 1);
```

`Frameworks` має відтворювати елемент `ul`.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('ul').length === 1);
```

Тег `ul` має відтворювати 6 дочірніх елементів `li`.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('ul').children().length ===
    6 &&
    Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .childAt(0)
      .name() === 'li' &&
    Enzyme.mount(React.createElement(Frameworks)).find('li').length === 6
);
```

Кожен елемент списку має мати унікальний атрибут `key`.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(Frameworks)).find('ul');
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key(),
      ul.childAt(4).key(),
      ul.childAt(5).key()
    ]);
    return keys.size === 6;
  })()
);
```

Кожен елемент списку має містити текст із `frontEndFrameworks`.

```js
assert(
  (() => {
    const li = Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .children();
    return [...Array(5)].every((_, i) =>
      frontEndFrameworks.includes(li.at(i).text())
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = null; // Change this line
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```

# --solutions--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = frontEndFrameworks.map((fw, i) => <li key={i}>{fw}</li>);
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```
