---
id: 5a24c314108439a4d4036164
title: Створення компонента за допомогою композиції
challengeType: 6
forumTopicId: 301383
dashedName: create-a-component-with-composition
---

# --description--

Тепер ми розглянемо, як можна скомпонувати кілька React-компонентів разом. Уявіть, що ви створюєте застосунок і маєте три компоненти: `Navbar`, `Dashboard` та `Footer`.

Щоб скомпонувати ці компоненти разом, ви можете створити `App` *батьківський* компонент, який відтворює кожен із цих трьох компонентів як *дочірні*. Щоб відтворити компонент як дочірній у React-компоненті, ви включаєте ім’я компонента, написане як кастомний HTML-тег у JSX. Наприклад, у методі `render` ви можете написати:

```jsx
return (
 <App>
  <Navbar />
  <Dashboard />
  <Footer />
 </App>
)
```

Коли React зустрічає кастомний HTML-тег, який посилається на інший компонент (ім’я компонента, обгорнуте в `< />`, як у цьому прикладі), він відтворює розмітку для цього компонента на місці тегу. Це ілюструє батьківсько-дочірні відносини між `App` компонентом та `Navbar`, `Dashboard` і `Footer`.

# --instructions--

У редакторі коду є простий функціональний компонент під назвою `ChildComponent` та класовий компонент під назвою `ParentComponent`. Скомпонуйте їх разом, відтворивши `ChildComponent` всередині `ParentComponent`. Переконайтеся, що тег `ChildComponent` закритий косою рискою.

**Примітка:** `ChildComponent` визначений за допомогою стрілкової функції ES6, оскільки це дуже поширена практика при використанні React.

# --hints--

React-компонент має повертати один `div` елемент.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.type() === 'div';
  })()
);
```

Компонент має повертати два вкладені елементи.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.children().length === 2;
  })()
);
```

Компонент має повертати `ChildComponent` як свій другий дочірній елемент.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ParentComponent));
    return (
      mockedComponent.find('ParentComponent').find('ChildComponent').length ===
      1
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<ParentComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }


        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }
        <ChildComponent />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
