---
id: 5a24c314108439a4d4036167
title: Відтворення класового компонента в DOM
challengeType: 6
forumTopicId: 301404
dashedName: render-a-class-component-to-the-dom
---

# --description--

Ви, можливо, пам’ятаєте, що в одному з попередніх завдань використовували API ReactDOM для відтворення JSX-елементів у DOM. Процес відтворення React-компонентів буде дуже схожим. Останні кілька завдань були зосереджені на компонентах і композиції, тому відтворення виконувалося за вас у фоновому режимі. Проте жоден React-код, який ви пишете, не відтвориться в DOM без виклику API ReactDOM.

Ось нагадування про синтаксис: `ReactDOM.render(componentToRender, targetNode)`. Перший аргумент — це React-компонент, який ви хочете відтворити. Другий аргумент — це DOM-вузол, у якому ви хочете відтворити цей компонент.

React-компоненти передаються в `ReactDOM.render()` трохи інакше, ніж JSX-елементи. Для JSX-елементів ви передаєте ім’я елемента, який хочете відтворити. Однак для React-компонентів потрібно використовувати той самий синтаксис, ніби ви відтворюєте вкладений компонент, наприклад `ReactDOM.render(<ComponentToRender />, targetNode)`. Цей синтаксис використовується як для класових компонентів ES6, так і для функціональних компонентів.

# --instructions--

Обидва компоненти `Fruits` і `Vegetables` визначені для вас у фоновому режимі. Відтворіть обидва компоненти як дочірні елементи компонента `TypesOfFood`, а потім відтворіть `TypesOfFood` у DOM. Для цього у вас є `div` з `id='challenge-node'`.

# --hints--

Компонент `TypesOfFood` має повертати один елемент `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Компонент `TypesOfFood` має відтворювати компонент `Fruits` після елемента `h1`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

Компонент `TypesOfFood` має відтворювати компонент `Vegetables` після `Fruits`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

Компонент `TypesOfFood` має відтворюватися в DOM у межах `div` з ідентифікатором `challenge-node`.

```js
assert(
  (function () {
    const html = document.getElementById('challenge-node').childNodes[0]
      .innerHTML;
    return (
      html.includes(
        '<div><h2>Fruits:</h2><h4>Non-Citrus:</h4><ul><li>Apples</li><li>Blueberries</li><li>Strawberries</li><li>Bananas</li></ul><h4>Citrus:</h4><ul><li>Lemon</li><li>Lime</li><li>Orange</li><li>Grapefruit</li></ul></div>'
      ) &&
      html.includes(
        '<div><h2>Vegetables:</h2><ul><li>Brussel Sprouts</li><li>Broccoli</li><li>Squash</li></ul></div>'
      )
    );
  })()
);
```

# --seed--

## --before-user-code--

```jsx
const Fruits = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <h4>Non-Citrus:</h4>
        <ul>
          <li>Apples</li>
          <li>Blueberries</li>
          <li>Strawberries</li>
          <li>Bananas</li>
        </ul>
      <h4>Citrus:</h4>
        <ul>
          <li>Lemon</li>
          <li>Lime</li>
          <li>Orange</li>
          <li>Grapefruit</li>
        </ul>
    </div>
  );
};
const Vegetables = () => {
  return (
    <div>
      <h2>Vegetables:</h2>
      <ul>
        <li>Brussel Sprouts</li>
        <li>Broccoli</li>
        <li>Squash</li>
      </ul>
    </div>
  );
};
```

## --seed-contents--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
```

# --solutions--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}
          <Fruits />
           <Vegetables />
         {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
ReactDOM.render(<TypesOfFood />, document.getElementById('challenge-node'));
```
