---
id: 5a24c314108439a4d4036168
title: Напишіть компонент React з нуля
challengeType: 6
forumTopicId: 301424
dashedName: write-a-react-component-from-scratch
---

# --description--

Тепер, коли ви вивчили основи JSX і компонентів React, настав час написати компонент самостійно. Компоненти React — це основні будівельні блоки застосунків React, тому важливо добре навчитися їх створювати. Пам’ятайте, типовий компонент React — це ES6 `class`, який розширює `React.Component`. Він має метод render, який повертає HTML (із JSX) або `null`. Це базова форма компонента React. Коли ви це добре зрозумієте, будете готові почати створювати більш складні проєкти на React.

# --instructions--

Визначте клас `MyComponent`, який розширює `React.Component`. Його метод render має повертати `div`, що містить тег `h1` з текстом: `My First React Component!` всередині. Використайте цей текст точно, регістр і пунктуація мають значення. Обов’язково викличте конструктор для свого компонента.

Відтворіть цей компонент у DOM за допомогою `ReactDOM.render()`. Для вас доступний `div` з `id='challenge-node'`.

# --hints--

Повинен бути компонент React з назвою `MyComponent`.

```js
assert(__helpers.removeWhiteSpace(code).includes('classMyComponentextendsReact.Component{'));
```

`MyComponent` має містити тег `h1` з текстом `My First React Component!`. Регістр і пунктуація мають значення.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('h1').text() === 'My First React Component!';
  })()
);
```

`MyComponent` має відтворюватися у DOM.

```js
assert(document.getElementById('challenge-node').childNodes.length === 1);
```

`MyComponent` має мати конструктор, який викликає `super` з `props`.

```js
assert(
  MyComponent.toString().includes('MyComponent(props)') &&
    MyComponent.toString().includes('_super.call(this, props)')
);
```

# --seed--

## --seed-contents--

```jsx
// Change code below this line
```

# --solutions--

```jsx
// Change code below this line
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>My First React Component!</h1>
      </div>
    );
  }
};

ReactDOM.render(<MyComponent />, document.getElementById('challenge-node'));
```
