---
id: 5a24c314108439a4d403616d
title: Використання PropTypes для визначення очікуваних пропсів
challengeType: 6
forumTopicId: 301419
dashedName: use-proptypes-to-define-the-props-you-expect
---

# --description--

React надає корисні функції перевірки типів, щоб переконатися, що компоненти отримують пропси правильного типу. Наприклад, ваш застосунок робить виклик API, щоб отримати дані, які ви очікуєте у вигляді масиву, а потім передає їх компоненту як пропс. Ви можете встановити `propTypes` у вашому компоненті, щоб вимагати, щоб дані були типу `array`. Це видасть корисне попередження, якщо дані матимуть будь-який інший тип.

Вважається найкращою практикою встановлювати `propTypes`, коли ви заздалегідь знаєте тип пропса. Ви можете визначити властивість `propTypes` для компонента так само, як ви визначали `defaultProps`. Це перевірить, що пропси з певним ключем присутні і мають заданий тип. Ось приклад, який вимагає тип `function` для пропса з назвою `handleClick`:

```js
MyComponent.propTypes = { handleClick: PropTypes.func.isRequired }
```

У наведеному прикладі частина `PropTypes.func` перевіряє, що `handleClick` є функцією. Додавання `isRequired` повідомляє React, що `handleClick` є обов’язковою властивістю для цього компонента. Ви побачите попередження, якщо цей пропс не буде наданий. Також зверніть увагу, що `func` представляє `function`. Серед семи примітивних типів JavaScript лише `function` і `boolean` (записані як `bool`) мають незвичайне написання. Окрім примітивних типів, доступні й інші типи. Наприклад, ви можете перевірити, що пропс є React-елементом. Будь ласка, зверніться до <a href="https://reactjs.org/docs/typechecking-with-proptypes.html#proptypes" target="_blank" rel="noopener noreferrer nofollow">документації</a> для всіх варіантів.

**Примітка:** Починаючи з React v15.5.0, `PropTypes` імпортується окремо від React, ось так: `import PropTypes from 'prop-types';`

# --instructions--

Визначте `propTypes` для компонента `Items`, щоб вимагати `quantity` як пропс і перевірити, що він має тип `number`.

# --hints--

Компонент `ShoppingCart` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Компонент `Items` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

Компонент `Items` має містити перевірку `propTypes`, щоб вимагати значення для `quantity` і переконатися, що його значення є числом.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(noWhiteSpace.includes('quantity:PropTypes.number.isRequired') && noWhiteSpace.includes('Items.propTypes='));
```

# --seed--

## --before-user-code--

```jsx
var PropTypes = {
  number: { isRequired: true }
};
```

## --after-user-code--

```jsx
ReactDOM.render(<ShoppingCart />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line

// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line
Items.propTypes = {
  quantity: PropTypes.number.isRequired
};
// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```
