---
id: 587d824e367417b2b2512c56
title: Jaribu kama Thamani ni Aina Maalum ya Muundo wa Data
challengeType: 2
forumTopicId: 301601
dashedName: test-if-a-value-is-of-a-specific-data-structure-type
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa kuanzia uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`#typeOf` inathibitisha kuwa aina ya thamani ni mfuatano wa herufi uliotolewa, kama inavyoamuliwa na `Object.prototype.toString`.

# --instructions--

Ndani ya `tests/1_unit-tests.js` chini ya jaribio lililo na lebo `#17` katika safu ya `Objects`, badilisha kila `assert` kuwa `assert.typeOf` au `assert.notTypeOf` ili kufanya jaribio lifaulu (linapaswa kutathminiwa kuwa `true`). Usibadilishe hoja zinazopitishwa kwa uthibitisho.

# --hints--

Jaribio zote zinapaswa kupita.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa kwanza - `typeOf` dhidi ya `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'typeOf',
  'myCar is typeOf Object'
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa pili - `typeOf` dhidi ya `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'typeOf',
  'Car.model is a String'
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa tatu - `typeOf` dhidi ya `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'notTypeOf',
  'Plane.wings is a Number (not a String)'
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa nne - `typeOf` dhidi ya `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'typeOf',
  'Plane.engines is an Array'
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa tano - `typeOf` dhidi ya `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[4].method,
  'typeOf',
  'Car.wheels is a Number'
);
```
