---
id: 587d824e367417b2b2512c57
title: Jaribu kama kitu ni mfano wa kijenzi
challengeType: 2
forumTopicId: 301605
dashedName: test-if-an-object-is-an-instance-of-a-constructor
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`#instanceOf` inathibitisha kuwa kitu ni mfano wa kijenzi.

# --instructions--

Ndani ya `tests/1_unit-tests.js` chini ya jaribio lililo na lebo ya `#18` katika suite ya `Objects`, badilisha kila `assert` kuwa `assert.instanceOf` au `assert.notInstanceOf` ili jaribio lipite (linapaswa kutathminiwa kuwa `true`). Usibadilishe hoja zinazopitishwa kwa uthibitisho.

# --hints--

Jaribio zote zinapaswa kupita.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa kwanza - `instanceOf` dhidi ya `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInstanceOf',
  'myCar is not an instance of Plane'
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa pili - `instanceOf` dhidi ya `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'instanceOf',
  'airlinePlane is an instance of Plane'
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa tatu - `instanceOf` dhidi ya `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'instanceOf',
  'everything is an Object in JavaScript...'
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa nne - `instanceOf` dhidi ya `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'notInstanceOf',
  'myCar.wheels is not an instance of String'
);
```
