---
id: 587d824d367417b2b2512c53
title: Jaribu kama mfuatano wa herufi una mfuatano mdogo
challengeType: 2
forumTopicId: 301597
dashedName: test-if-a-string-contains-a-substring
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`include()` na `notInclude()` pia hufanya kazi kwa mfuatano wa herufi! `include()` inathibitisha kuwa mfuatano halisi una mfuatano mdogo unaotarajiwa.

# --instructions--

Ndani ya `tests/1_unit-tests.js` chini ya jaribio lililo na lebo ya `#14` katika safu ya `Strings`, badilisha kila `assert` kuwa `assert.include` au `assert.notInclude` ili jaribio lifaulu (linapaswa kutathminiwa kuwa `true`). Usibadilishe hoja zinazotumwa kwa uthibitisho.

# --hints--

Jaribio zote zinapaswa kupita.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa kwanza - `include` dhidi ya `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'include',
  "'Arrow' contains 'row'..."
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa pili - `include` dhidi ya `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notInclude',
  "... a 'dart' doesn't contain a 'queue'"
);
```
