---
id: 587d824b367417b2b2512c49
title: Jaribio la kuthibitisha ukweli
challengeType: 2
forumTopicId: 301596
dashedName: test-for-truthiness
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`isTrue()` itajaribu thamani ya kweli au siyo kweli `true` na `isNotTrue()` zitapita wakati zitapewa chochote isipokuwa thamani ya kweli au siyo kweli ya `true`.

```js
assert.isTrue(true, 'This will pass with the boolean value true');
assert.isTrue('true', 'This will NOT pass with the string value "true"');
assert.isTrue(1, 'This will NOT pass with the number value 1');
```

`isFalse()` na `isNotFalse()` pia zipo, na zinafanya kazi kama zile za kweli isipokuwa zinatafuta thamani ya kweli au siyo kweli ya `false`.

# --instructions--

Ndani ya `tests/1_unit-tests.js` chini ya jaribio lililopewa lebo `#4` katika suite ya `Basic Assertions`, badilisha kila `assert` kuwa `assert.isTrue` au `assert.isNotTrue` ili kufanya jaribio lipite (linapaswa kutathminiwa kuwa `true`). Usibadilishe hoja zinazotumwa kwa uthibitisho.

# --hints--

Jaribio zote zinapaswa kupita.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa kwanza - `isTrue` dhidi ya `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isTrue', 'True is true');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa pili - `isTrue` dhidi ya `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isTrue',
    'Double negation of a truthy value is true'
  );
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa tatu - `isTrue` dhidi ya `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isNotTrue',
    'A truthy object is not true - neither is a false one'
  );
```
