---
id: 587d824c367417b2b2512c4f
title: Jaribu kama thamani iko ndani ya anuwai maalum
challengeType: 2
forumTopicId: 301598
dashedName: test-if-a-value-falls-within-a-specific-range
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

```javascript
.approximately(actual, expected, delta, [message])
```

Inathibitisha kuwa `actual` ni sawa na `expected`, ndani ya anuwai ya +/- `delta`.

# --instructions--

Ndani ya `tests/1_unit-tests.js` chini ya jaribio lililo na lebo `#10` katika safu ya `Comparisons`, badilisha kila `assert` kuwa `assert.approximately` ili jaribio lifaulu (linapaswa kutathminiwa kuwa `true`).

Chagua anuwai ya kiwango cha chini (kigezo cha 3) ili jaribio lifaulu kila wakati. Inapaswa kuwa chini ya 1.

# --hints--

Jaribio zote zinapaswa kupita.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Unapaswa kuchagua anuwai sahihi kwa uthibitisho wa kwanza - `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'approximately');
assert.equal(
  data.assertions[0].args[2],
  0.5,
  "weirdNumbers(0.5) is in the range (0.5, 1.5]. It's within 1 +/- 0.5"
);
```

Unapaswa kuchagua anuwai sahihi kwa uthibitisho wa pili - `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'approximately');
assert.equal(
  data.assertions[1].args[2],
  0.8,
  "weirdNumbers(0.2) is in the range (0.2, 1.2]. It's within 1 +/- 0.8"
);
```
