---
id: 587d824e367417b2b2512c55
title: Jaribu kama Kitu kina Sifa
challengeType: 2
forumTopicId: 301604
dashedName: test-if-an-object-has-a-property
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`property` inathibitisha kuwa kitu halisi kina sifa fulani.

# --instructions--

Ndani ya `tests/1_unit-tests.js` chini ya jaribio lililo na lebo `#16` katika suite ya `Objects`, badilisha kila `assert` kuwa `assert.property` au `assert.notProperty` ili kufanya jaribio lifaulu (linapaswa kutathminiwa kuwa `true`). Usibadilishe hoja zinazopitishwa kwa uthibitisho.

# --hints--

Jaribio zote zinapaswa kupita.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa kwanza - `property` dhidi ya `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notProperty',
  'A car has not wings'
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa pili - `property` dhidi ya `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'property',
  'planes have engines'
);
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa tatu - `property` dhidi ya `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'property', 'Cars have wheels');
```
