---
id: 587d824a367417b2b2512c46
title: Jifunze jinsi uthibitisho za JavaScript zinavyofanya kazi
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

Kufanya kazi kwenye changamoto hizi kutahusisha kuandika msimbo wako ukitumia mojawapo ya njia zifuatazo:

- Nakili <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">hii repo ya GitHub</a> na ukamilishe changamoto hizi kwa ndani.
- Tumia mtengenezaji wa tovuti unayempenda kukamilisha mradi. Hakikisha unajumuisha mafaili yote kutoka kwenye repo yetu ya GitHub.

# --instructions--

Ndani ya `tests/1_unit-tests.js` chini ya jaribio lililo na lebo `#1` katika suite ya `Basic Assertions`, badilisha kila `assert` kuwa `assert.isNull` au `assert.isNotNull` ili jaribio lifaulu (linapaswa kutathmini kuwa `true`). Usibadilishe hoja zinazopitishwa kwa uthibitisho.

# --hints--

Jaribio zote zinapaswa kupita.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa kwanza - `isNull` dhidi ya `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa pili - `isNull` dhidi ya `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
```
