---
id: 587d824b367417b2b2512c4a
title: Tumia alama mbili za usawa kuthibitisha usawa
challengeType: 2
forumTopicId: 301609
dashedName: use-the-double-equals-to-assert-equality
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`equal()` inalinganisha vitu kwa kutumia `==`.

# --instructions--

Ndani ya `tests/1_unit-tests.js` chini ya jaribio lililo na lebo `#5` katika mkusanyiko wa `Equality`, badilisha kila `assert` kuwa `assert.equal` au `assert.notEqual` ili kufanya jaribio lifaulu (linapaswa kutathminiwa kuwa `true`). Usibadilishe hoja zinazopitishwa kwa uthibitisho.

# --hints--

Jaribio zote zinapaswa kupita.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa kwanza - `equal` dhidi ya `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'equal',
    'Numbers are coerced into strings with == '
  );
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa pili - `equal` dhidi ya `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notEqual',
    ' == compares object references'
  );
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa tatu - `equal` dhidi ya `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'equal',
    "6 * '2' is 12 ! It should be equal to '12'"
  );
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa nne - `equal` dhidi ya `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[3].method, 'notEqual', "6 + '2' is '62'...");
```
