---
id: 587d824b367417b2b2512c4b
title: Tumia alama tatu za sawa kuthibitisha usawa mkali
challengeType: 2
forumTopicId: 301610
dashedName: use-the-triple-equals-to-assert-strict-equality
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa kuanzia uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`strictEqual()` inalinganisha vitu kwa kutumia `===`.

# --instructions--

Ndani ya `tests/1_unit-tests.js` chini ya jaribio lililo na lebo ya `#6` katika suite ya `Equality`, badilisha kila `assert` kuwa `assert.strictEqual` au `assert.notStrictEqual` ili kufanya jaribio lifaulu (linapaswa kutathmini kuwa `true`). Usibadilishe hoja zinazopitishwa kwa uthibitisho.

# --hints--

Jaribio zote zinapaswa kupita.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa kwanza - `strictEqual` dhidi ya `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'notStrictEqual',
    'with strictEqual the type must match'
  );
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa pili - `strictEqual` dhidi ya `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'strictEqual', '3*2 = 6...');
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa tatu - `strictEqual` dhidi ya `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'strictEqual',
    "6 * '2' is 12. Types match !"
  );
```

Unapaswa kuchagua njia sahihi kwa uthibitisho wa nne - `strictEqual` dhidi ya `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[3].method,
    'notStrictEqual',
    'Even if they have the same elements, the Arrays are notStrictEqual'
  );
```
