---
id: 587d824f367417b2b2512c5a
title: Endesha majaribio ya kitendakazi kwenye jibu la API ukitumia Chai-HTTP III - njia ya PUT
challengeType: 2
forumTopicId: 301590
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iii---put-method
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa kuanzia uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Unapojaribu ombi la `PUT`, mara nyingi utatuma data pamoja nalo. Data unayoingiza pamoja na ombi lako la `PUT` huitwa mwili wa ombi.

Ili kutuma ombi la `PUT` na kitu cha JSON kwa kiishio cha `'/travellers'`, unaweza kutumia nyongeza ya `chai-http` na njia za `put` na `send`:

```js
chai
  .request(server)
  .keepOpen()
  .put('/travellers')
  .send({
    "surname": [last name of a traveller of the past]
  })
  ...
```

Na njia ya data inajibu na:

```json
{
  "name": [first name],
  "surname": [last name],
  "dates": [birth - death years]
}
```

Tazama msimbo wa seva kwa majibu tofauti kwa kiishio cha `'/travellers'`.

# --instructions--

Ndani ya `tests/2_functional-tests.js`, badilisha jaribio la `'Send {surname: "Colombo"}'` (`// #3`) na tumia njia za `put` na `send` kujaribu kiishio cha `'/travellers'`.

Tuma kitu cha JSON kinachofuata pamoja na ombi lako la PUT:

```json
{
  "surname": "Colombo"
}
```

Angalia yafuatayo ndani ya wito wa kurudisha wa `request.end`:

1.  `status` inapaswa kuwa `200`
2.  `type` inapaswa kuwa `application/json`
3.  `body.name` inapaswa kuwa `Cristoforo`
4.  `body.surname` inapaswa kuwa `Colombo`

Fuata mpangilio wa uthibitisho ulio juu - tunategemea hilo. Pia, hakikisha unaondoa `assert.fail()` mara umemaliza.

# --hints--

Majaribio yote yanapaswa kupita.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Unapaswa kujaribu `res.status` iwe 200.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Unapaswa kujaribu `res.type` iwe `'application/json'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

Unapaswa kujaribu `res.body.name` iwe `'Cristoforo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Cristoforo\1/);
```

Unapaswa kujaribu `res.body.surname` iwe `'Colombo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")Colombo\1/);
```
