---
id: 5f8884f4c46685731aabfc41
title: Endesha majaribio ya kitendakazi ukitumia kivinjari kisicho na kichwa II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Ndani ya `tests/2_functional-tests.js`, katika jaribio la `'Submit the surname "Vespucci" in the HTML form'` (`// #6`), fanya otomatiki yafuatayo:

1. Jaza fomu na jina la ukoo `Vespucci`
2. Bonyeza kitufe cha kuwasilisha

Na ndani ya wito wa kurudisha `pressButton`:

1. Thibitisha kuwa hali ni OK `200`
2. Thibitisha kuwa maandishi ndani ya kipengele `span#name` ni `'Amerigo'`
3. Thibitisha kuwa maandishi ndani ya kipengele `span#surname` ni `'Vespucci'`
4. Thibitisha kuwa kipengele(kipengele) `span#dates` kipo na idadi yake ni `1`

Usisahau kuondoa wito `assert.fail()`.

# --hints--

Majaribio yote yanapaswa kupita.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Unapaswa kuthibitisha kuwa ombi la kivinjari kisicho na kichwa limefanikiwa.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

Unapaswa kuthibitisha kuwa maandishi ndani ya kipengele `span#name` ni `'Amerigo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

Unapaswa kuthibitisha kuwa maandishi ndani ya kipengele `span#surname` ni `'Vespucci'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

Unapaswa kuthibitisha kuwa kipengele `span#dates` kipo na idadi yake ni 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
