---
id: 669fc7e141e4703748c558bf
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro del elemento `html`, añade un elemento `head`.

# --hints--

Debes tener solo una etiqueta de apertura `head`.

```js
assert.match(code, /<head>/i);
```

Debes tener una etiqueta de cierre `head` tag.

```js
assert.match(code, /<\/head>/i);
```

La etiqueta de apertura `head` debe situarse antes de la etiqueta de cierre `head`.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
