---
id: 669fde2081f65141ad703fe4
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

La primera sección en la página será la sección sobre. La sección presentará al Sr. Whiskers y dará a los usuarios una idea de qué trata este blog.

Dentro de su elemento `main`, añade un elemento `section` con el atributo `id` establecido en `"about"`.

Dentro del elemento `section`, añade un `h2` con el texto de `About`.

# --hints--

Tu elemento `section` debe tener una etiqueta de apertura.

```js
assert.match(code, /<section.*>/);
```

Tu elemento `section` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/section\>/);
```

Tu elemento `section` debería tener un `id` configurado como `"about"`.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

Su elemento `section` debe estar dentro de su elemento `main`.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Tu elemento `h2` debe tener una etiqueta de apertura.

```js
assert.match(code, /<h2\>/);
```

Tu elemento `h2` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/h2\>/);
```

Tu elemento `h2` debería tener el texto de `About`. Revise la ortografía atentamente.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

Tu elemento `h2` debe estar anidado dentro de tu elemento `section`.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>
```
