---
id: 669fd6fd12918e3de87854d4
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Para tu blog, debería haber una manera para que los usuarios naveguen a diferentes secciones en la página.

El elemento `nav` se utiliza para proporcionar enlaces de navegación a otras secciones en el documento u otras secciones en el sitio web. Muchas veces verás el elemento `nav` utilizado para menús o tablas de contenido.

Aquí tienes un ejemplo de cómo usar el elemento `nav`:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

Debajo de tu elemento `figure`, añade un elemento `nav` con un elemento `ul` anidado dentro.

Dentro del elemento `ul`, añade tres elementos `li`.

# --hints--

Tu elemento `nav` debería tener una etiqueta de apertura.

```js
assert.match(code, /<nav\>/);
```

Tu elemento `nav` debería tener un cierre.

```js
assert.match(code, /<\/nav\>/);
```

Tu elemento `nav` debería estar dentro de tu elemento `header`.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

Tu elemento `ul` debe tener una etiqueta de apertura.

```js
assert.match(code, /<ul\>/);
```

Tu elemento `ul` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/ul\>/);
```

Tu elemento `ul` debería estar dentro de tu elemento `nav`.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

Tu elemento `ul` debería contener tres elementos `li`.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--

--fcc-editable-region--
    </header>
  </body>
</html>
```
