---
id: 669fdd3965f36f40af9615e5
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Ahora que has terminado de construir el encabezado de la página, necesitarás comenzar a agregar el contenido principal.

Debajo de tu `header`, añade un elemento `main`.

# --hints--

Tu elemento `main` debe tener una etiqueta de apertura.

```js
assert.match(code, /<main\>/);
```

Tu elemento `main` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/main\>/);
```

Tu elemento `main` debe estar debajo de tu elemento `header`.

```js
assert.match(code, /<\/header>[.\n\s]*<main>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
