---
id: 669fd003cb89ee3c2402e041
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

En este contenido introductorio, querrás mostrar una imagen de Mr. Whiskers con un pie de foto.

Debajo del elemento `h1`, comienza añadiendo un elemento `figure`.

Dentro del elemento `figure`, añade un elemento `img`.

The `src` attribute of the `img` should have a value of `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` and the `alt` text should have a value of `"a cat in the garden"`.

Below your `img` element, add a `figcaption` with the text `Mr. Whiskers in the Garden`.

# --hints--

Tu elemento `figure` debe tener una etiqueta de apertura.

```js
assert.match(code, /<figure\>/);
```

Tu elemento `figure` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/figure\>/);
```

Tu elemento `figure` debería estar dentro de tu elemento `header`.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

Deberías tener un elemento `img`.

```js
assert.exists(document.querySelector("img"));
```

Tu elemento `img` debería estar dentro de tu elemento `figure`.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

Tu elemento `img` debería tener un atributo `src` establecido a `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Your `img` element should have an `alt` attribute set to `"a cat in the garden"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

Tu elemento `figcaption` debe tener una etiqueta de apertura.

```js
assert.match(code, /<figcaption\>/);
```

Tu elemento `figcaption` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/figcaption\>/);
```

Tu elemento `figcaption` debería estar dentro de tu elemento `figure`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Your `figcaption` element should have a the text of `Mr. Whiskers in the Garden`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText, "Mr. Whiskers in the Garden")
```

Tu `figcaption` debería estar debajo del elemento `img`. Los elementos no están en el orden correcto.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>

    </header>
  </body>
--fcc-editable-region--
</html>
```
