---
id: 669fcb06c3034a39f5431a38
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare creating some actual content, add a `body` element below the `head` element.

# --hints--

Debes tener una etiqueta `<body>` de apertura.

```js
assert.match(code, /<body>/i);
```

Debes tener una etiqueta `</body>` de cierre.

```js
assert.match(code, /<\/body>/i);
```

No debes modificar el elemento `head`. Asegúrese de que no eliminaste tu etiqueta de cierre.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Tu elemento `body` debe estar después del elemento `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>

  --fcc-editable-region--
</html>
```
