---
id: 66a334a1a7cca6354999f9bf
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Para la primera entrada del blog, utilizarás un elemento <dfn>article</dfn>.

El elemento `article` representa contenido autocontenido en una página web.

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

Debajo del elemento `h2`, añade un elemento `article`.

Dentro del elemento `article`, añade un elemento `h3` con el texto `Mr. Whiskers' First Day Home`.

La razón por la que se usa un `h3` aquí es que mantener una jerarquía estructural adecuada para los elementos de encabezado es importante. Dado que el subtítulo del post es un elemento `h2`, el siguiente nivel hacia abajo en la jerarquía sería un `h3`.

# --hints--

Tu código debe tener una etiqueta de apertura `article`.

```js
assert.match(code, /<article>/i);
```

Tu código debe tener una etiqueta de cierre `article`.

```js
assert.match(code, /<\/article>/i);
```

Tu elemento `article` debe estar dentro de tu elemento `section` con el `id` establecido en `posts`.

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

Tu código debe tener una etiqueta de apertura `h3`.

```js
assert.match(code, /<h3>/i);
```

Tu código debe tener una etiqueta de cierre `h3`.

```js
assert.match(code, /<\/h3>/i);
```

Tu elemento `h3` debe tener el texto de `Mr. Whiskers' First Day Home`.

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText, "Mr. Whiskers' First Day Home");
```

Tu elemento `h3` debe estar dentro de tu elemento `article`.

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
--fcc-editable-region--
      <section id="posts">
        <h2>Posts</h2>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
