---
id: 66be24cb4144f955b6bcc550
title: Eine Seite mit Spielkarten erstellen
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Webseite erstellen, die mindestens drei Spielkarten anzeigt.
1. Sie sollten ein `main`-Element mit der ID `playing-cards` haben.
1. Innerhalb Ihres `#playing-cards`-Elements sollten Sie mindestens drei `div`-Elemente haben, von denen jedes die Klasse `card` besitzt.
1. Innerhalb jedes `.card`-Elements sollten Sie drei `div`-Elemente haben, das erste mit der Klasse `left`, das zweite mit der Klasse `middle` und das dritte mit der Klasse `right`.
1. Ihr `#playing-cards`-Element sollte Flexbox verwenden, um seine Kinder horizontal zu zentrieren, das Umbruchverhalten zu erlauben und einen Abstand von `20px` zwischen ihnen einzufügen.
1. Jedes Ihrer `.card`-Elemente sollte Flexbox verwenden, um seine Kinder mit `space-between` auszurichten.
1. Jedes Ihrer `.left`-Elemente sollte Flexbox-Item-Eigenschaften verwenden, um sich am Anfang seines übergeordneten Containers auszurichten.
1. Jedes Ihrer `.middle`-Elemente sollte Flexbox-Item-Eigenschaften verwenden, um sich in der Mitte seines übergeordneten Containers auszurichten.
1. Jedes Ihrer `.right`-Elemente sollte Flexbox-Item-Eigenschaften verwenden, um sich am Ende seines übergeordneten Containers auszurichten.
1. Jedes Ihrer `.middle`-Elemente sollte Flexbox verwenden, um seine Kinder in einer Spalte anzuzeigen.

Hier sind einige Zeichen, die Sie kopieren und in Ihren Karten verwenden können, wenn Sie möchten: `♠`, `♣`, `♥` und `♦`.

# --hints--

Sie sollten ein `main`-Element mit einem `id` von `playing-cards` haben.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

Sie sollten mindestens drei `div`-Elemente mit einer Klasse von `card` innerhalb Ihres `#playing-cards`-Elements haben.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

Jedes Ihrer `.card`-Elemente sollte eine Breite und Höhe haben.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

Jedes Ihrer `.card`-Elemente sollte genau drei `div`-Elemente als Kinder haben.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

Sie sollten ein `div` mit einer Klasse von `left` innerhalb jedes `.card`-Elements haben.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

Sie sollten ein `div` mit einer Klasse von `middle` innerhalb jedes `.card`-Elements haben.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

Sie sollten ein `div` mit einer Klasse von `right` innerhalb jedes `.card`-Elements haben.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

Sie sollten einen `#playing-cards`-Selektor haben, der die `display` seiner Elemente auf `flex` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

Sie sollten einen `#playing-cards`-Selektor haben, der die `justify-content` seiner Elemente auf `center` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

Sie sollten einen `#playing-cards`-Selektor haben, der die `flex-wrap` seiner Elemente auf `wrap` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

Sie sollten einen `#playing-cards`-Selektor haben, der die `gap` seiner Elemente auf `20px` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

Sie sollten einen `.card`-Selektor haben, der die `display` seiner Elemente auf `flex` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

Sie sollten einen `.card`-Selektor haben, der die `justify-content` seiner Elemente auf `space-between` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

Sie sollten einen `.left`-Selektor haben, der die `align-self` seiner Elemente auf `flex-start` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

Sie sollten einen `.middle`-Selektor haben, der die `align-self` seiner Elemente auf `center` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

Sie sollten einen `.right`-Selektor haben, der die `align-self` seiner Elemente auf `flex-end` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

Sie sollten einen `.middle`-Selektor haben, der die `display` seiner Elemente auf `flex` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

Sie sollten in Ihrer `flex-direction`-Klasse `column` auf `.middle` setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
