---
id: 5f332b23c2045fb843337579
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

カフェが2020年に創業されたことを訪問者に知らせるために、`p`要素の下に`h1`要素を追加し、テキストに`Est. 2020`を入れてください。

# --hints--

開く`<p>`タグがあるはずです。

```js
assert(code.match(/<p>/i));
```

閉じる`</p>`タグがあるはずです。

```js
assert(code.match(/<\/p>/i));
```

既存の`h1`要素は変更しないでください。終了タグを削除していないことを確認してください。

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

`p`要素は`h1`要素の下にあるべきです。

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

`p`要素にはテキスト`Est. 2020`が入っているべきです。

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```
