---
id: 5f331e55dfab7a896e53c3a1
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` は、ウェブページ上には表示されませんが、検索エンジンやページの表示方法に役立つ追加情報を提供するいくつかの要素のうちの1つです。

`head` 要素の中に、`meta` 要素をネストし、その属性名を `charset`、値を `utf-8` に設定して、ブラウザーにページの文字エンコード方法を伝えてください。

`meta` 要素は終了タグを必要としない空要素であることを覚えておいてください。

# --hints--

`meta` タグがあるはずです。

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta` 要素は空要素なので、終了タグ `</meta>` を付けてはいけません。

```js
assert.notMatch(code, /<\/meta>/i);
```

`meta` タグには `charset` 属性が必要です。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

`charset` 属性の値は `utf-8` であるべきです。

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
