---
id: 5f33294a6af5e9188dbdb8f3
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

メニューの内容を追加する時間です。既存の`main`要素の中に`body`要素を追加してください。ここには最終的にカフェで提供されるコーヒーとデザートの価格情報が入ります。

# --hints--

コードには`<main>`の開始タグが必要です。

```js
assert(code.match(/<main>/i));
```

コードには`</main>`の終了タグが必要です。

```js
assert(code.match(/<\/main>/i));
```

`body`要素は変更しないでください。終了タグを誤って削除しないように注意してください。

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

`main`タグは`body`タグの中にある必要があります。

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
