---
id: 5f33071498eb2472b87ddee4
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このプロジェクトでは、カフェのメニューを作りながらCSS（カスケーディングスタイルシート）の基本を学びます。CSSはHTMLドキュメントのスタイルを指定する言語で、HTML要素が画面にどのように表示されるかを記述します。

Cat Photo Appの最後のいくつかのステップで学んだように、ウェブページを作り始めるには基本的な構造が必要です。すべてのHTMLドキュメントには`DOCTYPE`宣言と`html`要素が必要です。`DOCTYPE`はブラウザーにドキュメントのHTMLバージョンを伝えます。そして`html`要素はすべての他の要素を含むルート要素を表します。

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

`<!DOCTYPE html>`タグと、`html`属性に`lang`を持つ`en`要素を追加してください。

# --hints--

`<!DOCTYPE html>`宣言があるはずです。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

言語が英語に設定された開く`<html lang="en">`があるはずです。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

閉じる`</html>`タグがあるはずです。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
