---
id: 5f332a88dc25a0fd25c7687a
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

カフェの名前は `CAMPER CAFE` です。`h1` 要素の中に `main` 要素を追加してください。カフェの名前を大文字で書いて目立たせましょう。

# --hints--

開く `<h1>` タグがあるはずです。

```js
assert(code.match(/<h1>/i));
```

閉じる `</h1>` タグがあるはずです。

```js
assert(code.match(/<\/h1>/i));
```

既存の `main` 要素は変更しないでください。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

`h1` 要素は `main` 要素の中にネストされているはずです。

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

`h1` 要素にはテキスト `CAMPER CAFE` が含まれているはずです。

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```
