---
id: 5f3313e74582ad9d063e3a38
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`head` 要素の中に `html` 要素を追加して、`title` 要素を追加できるようにしてください。`title` 要素のテキストは `Cafe Menu` にしてください。

# --hints--

開く `<head>` タグが必要です。

```js
assert(code.match(/<head>/i));
```

閉じる `</head>` タグが必要です。

```js
assert(code.match(/<head>/i));
```

開く `<title>` タグが必要です。

```js
assert(code.match(/<title>/i));
```

閉じる `</title>` タグが必要です。

```js
assert(code.match(/<\/title>/i));
```

`title` 要素は `head` 要素の中にネストされている必要があります。

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

`title` 要素のテキストは `Cafe Menu` である必要があります。スペルを確認してください。

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
