---
id: 5f3326b143638ee1a09ff1e3
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

実際のコンテンツを作成する準備として、`body` 要素の下に `head` 要素を追加してください。

# --hints--

`<body>` の開始タグがあるはずです。

```js
assert(code.match(/<body>/i));
```

`</body>` の終了タグがあるはずです。

```js
assert(code.match(/<\/body>/i));
```

`head` 要素は変更しないでください。終了タグを削除していないことを確認してください。

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

`body` 要素は `head` 要素の後に来るはずです。

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
