---
id: 5f344fbc22624a2976425065
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`h2` 要素の中に `section` 要素を作成し、テキストに `Coffee` を入れてください。

# --hints--

開く `<h2>` タグがあるはずです。

```js
assert(code.match(/<h2\s*>/i));
```

閉じる `</h2>` タグがあるはずです。

```js
assert(code.match(/<\/h2\s*>/i));
```

既存の `section` 要素は変更しないでください。終了タグを削除していないか確認してください。

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

`h2` 要素は `section` 要素の中にあるはずです。

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

`h2` 要素にはテキスト `Coffee` が入っているはずです。

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
