---
id: 5f3477cbcb6ba47918c1da92
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

モバイルでもデスクトップやノートパソコンと同じような見た目になるようにページのスタイルを調整するには、特別な`meta`属性を持つ`content`要素を追加する必要があります。

`head`要素の中に次の内容を追加してください。

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

コードには`meta`要素が2つあるはずです。

```js
assert.lengthOf(code.match(/<meta.*?\/?>/g),2);
```

`meta`要素には、値が`name`の`viewport`属性が必要です。

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/name=('|")viewport\1/) || meta[1]?.outerHTML?.match(/name=('|")viewport\1/));
```

`meta`要素には、値が`content`の`width=device-width, initial-scale=1.0`属性が必要です。

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
