---
id: 5f344fad8bf01691e71a30eb
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

これまでは、コンテンツの表示や見た目をあまり自由にコントロールできませんでした。これを変えるために、`style` 要素の中に `head` 要素を追加してください。

# --hints--

コードには開く `<style>` タグが必要です。

```js
assert(code.match(/<style\s*>/i));
```

コードには閉じる `</style>` タグが必要です。

```js
assert(code.match(/<\/style\s*>/));
```

`style` 要素は `head` 要素の中にネストされている必要があります。

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
