---
id: 5f356ed6199b0cdef1d2be8f
title: ステップ 26
challengeType: 0
dashedName: step-26
---

# --description--

これまで、要素のスタイルを指定するのにタイプセレクターやIDセレクターを使ってきました。しかし、要素のスタイル指定には別のセレクターを使うことのほうが一般的です。

<dfn>クラスセレクター</dfn>は、名前の前にドットを直接付けて定義します。例えば次のようになります。

```css
.class-name {
  styles
}
```

既存の`#menu`セレクターを、`#menu`を`.menu`という名前のクラスに置き換えてクラスセレクターに変更してください。

# --hints--

`.menu`のクラスセレクターがあるはずです。

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

`#menu`のセレクターは使わないでください。

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```
