---
id: 5f35e5c4321f818cdc4bed30
title: ステップ 29
challengeType: 0
dashedName: step-29
---

# --description--

見た目が整ったので、メニュー項目を追加していきましょう。

`article` 見出しの下に空の `Coffee` 要素を追加してください。そこに現在提供している各コーヒーのフレーバーと価格を入れます。


# --hints--

開く `<article>` タグがあるはずです。

```js
assert.match(code,/<article>/i);
```

閉じる `</article>` タグがあるはずです。

```js
assert(code.match(/<\/article>/i));
```

既存の `h2` 要素は変更しないでください。

```js
assert.lengthOf(document.querySelectorAll('h2'),1);
```

`article` 要素は `h2` 要素の後に来るはずです。

```js
const article = document.querySelector('article');
assert.equal(article?.previousElementSibling?.tagName, 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
