---
id: 5f356ed69db0a491745e2bb6
title: ステップ 27
challengeType: 0
dashedName: step-27
---

# --description--

`div` 要素にクラスのスタイルを適用するには、`id` 属性を削除し、`class` 要素の開始タグに `div` 属性を追加してください。クラスの値は `menu` に設定してください。

# --hints--

`div` はまだレンダリングされるはずです。`<div>` タグが正しく記述されているか確認してください。

```js
assert.lengthOf(document.querySelectorAll('div'),1);
```

`div` 要素には `menu` クラスが付いているはずです。

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

`div` 要素にはもう `menu` id が付いていないはずです。

```js
assert.lengthOf(document.querySelectorAll('div#menu'),0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
