---
id: 5f356ed63c7807a4f1e6d054
title: ステップ 21
challengeType: 0
dashedName: step-21
---

# --description--

今の目標は、`div` がページ全体の幅を占めないようにすることです。CSS の `width` プロパティがこれに最適です。

`id` セレクターを使って、特定の `id` 属性を持つ要素を指定できます。`id` セレクターは、ハッシュ記号 `#` を要素の `id` 値の直前に置くことで定義されます。例えば、要素が `id` に `cat` を持っている場合、その要素は次のように指定します。

```css 
#cat {
  width: 250px;
}
```

`#menu` セレクターを使って、要素の幅を `300px` に設定してください。

# --hints--

`#menu` セレクターがあるはずです。

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

`width` プロパティを `300px` に設定してください。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

`div` の幅は 300px にしてください。

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```
