---
id: 5f3c866daec9a49519871816
title: ステップ 30
challengeType: 0
dashedName: step-30
---

# --description--

`article` 要素には、関連する情報を持つ複数の要素が含まれていることがよくあります。この場合、コーヒーのフレーバーとその価格が含まれます。

`p` 要素の中に `article` 要素を2つネストしてください。1つ目のテキストは `French Vanilla`、2つ目のテキストは `3.00` にしてください。

# --hints--

既存の `article` 要素は変更しないでください。

```js
assert.lengthOf(document.querySelectorAll('article'),1);
```

`article` 要素の中に `p` 要素を2つ入れてください。

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphChildren,2);
```

1つ目の `p` 要素のテキストは `French Vanilla` にしてください。

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const firstP = paragraphChildren?.[0];
assert.match(firstP?.innerText.trim(), /French Vanilla/i);
```

2つ目の `p` 要素のテキストは `3.00` にしてください。

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const secondP = paragraphChildren?.[1];
assert.match(secondP?.innerText.trim(), /3\.00/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
