---
id: 5f716bee5838c354c728a7c5
title: ステップ 52
challengeType: 0
dashedName: step-52
---

# --description--

先ほど追加したデザートの下に、残りのデザートと価格を3つの`article`要素で追加してください。それぞれに2つのネストされた`p`要素があります。各要素には正しいデザート名と価格のテキストを入れ、すべてに正しいクラスを設定してください。

```md
Cherry Pie 2.75
Cheesecake 3.00
Cinnamon Roll 2.50
```

# --hints--

2番目の`article`の中に4つの`section`要素があるはずです。

```js
const secondSection = document.querySelectorAll('section')[1]
const articles = secondSection.querySelectorAll('article')
assert(articles.length === 4)
```

4つの`.dessert`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('.dessert'),4);
```

4つの新しい`.price`要素があるはずです。

```js
const section = [...document.querySelectorAll('section')].at(-1);
const prices = section.querySelectorAll(`:scope ${'.price'}`);
assert.lengthOf(prices, 4);
```

`section`要素の中に8つの`p`要素があるはずです。

```js
const section = [...document.querySelectorAll('section')].at(-1);
const paragraphs = section.querySelectorAll(`:scope ${'p'}`);

assert.lengthOf(paragraphs,8);
```

`.dessert`要素には`Donut`、`Cherry Pie`、`Cheesecake`、`Cinnamon Roll`のテキストがあるはずです。

```js
const dessert = document.querySelectorAll('.dessert');
assert.match(dessert?.[0]?.innerText.trim(), /donut/i);
assert.match(dessert?.[1]?.innerText.trim(), /cherry pie/i);
assert.match(dessert?.[2]?.innerText.trim(), /cheesecake/i);
assert.match(dessert?.[3]?.innerText.trim(), /cinnamon roll/i);
```

新しい`.price`要素には`1.50`、`2.75`、`3.00`、`2.50`のテキストがあるはずです。

```js
const section = [...document.querySelectorAll('section')].at(-1);
const prices = section?.querySelectorAll(`:scope ${'.price'}`);
assert.match(prices?.[0]?.innerText.trim(), /1\.50/);
assert.match(prices?.[1]?.innerText.trim(), /2\.75/);
assert.match(prices?.[2]?.innerText.trim(), /3\.00/);
assert.match(prices?.[3]?.innerText.trim(), /2\.50/);
```

`p`要素の間にスペースがあってはいけません。

```js
assert(!code.match(/<\/p>\s+<p/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
--fcc-editable-region--
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%;
}
```
