---
id: 68e45a2b4ab9b2f48b1cc6df
title: 料金プランのレイアウトページをデザインする
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー**

1. ページにはテキスト `h1` を持つ `Pricing Plans` 要素があるべきです。
2. ページには `div` 要素の下にクラス `pricing-container` を持つ `h1` 要素があるべきです。
    - `.pricing-container` セレクターは、値 `display` の `flex` プロパティと値 `flex-wrap` の `wrap` プロパティを持つべきです。
3. `.pricing-container` 要素の中に、料金プランを表すクラス `div` を持つ `pricing-card` 要素が3つあるべきです。
    - `.pricing-card` 要素のうち1つは、`basic-plan` クラスに加えて `pricing-card` クラスを持つべきです。
    - `.pricing-card` 要素のうち1つは、`pro-plan` クラスに加えて `pricing-card` クラスを持つべきです。
    - `.pricing-card` 要素のうち1つは、`premium-plan` クラスに加えて `pricing-card` クラスを持つべきです。
4. `.basic-plan` 要素にはテキスト `h2` を持つ `Basic` 要素があるべきです。
5. `.basic-plan` 要素にはテキスト `p` を持つ `$9/month` 要素があるべきです。
6. `.pro-plan` 要素にはテキスト `h2` を持つ `Pro` 要素があるべきです。
7. `.pro-plan` 要素にはテキスト `p` を持つ `$19/month` 要素があるべきです。
8. `.premium-plan` 要素にはテキスト `h2` を持つ `Premium` 要素があるべきです。
9. `.premium-plan` 要素にはテキスト `p` を持つ `$29/month` 要素があるべきです。
10. 各 `.pricing-card` 要素は以下を満たすべきです：
    - Flexbox を使って内容を縦方向に積み重ね、子要素間のスペースを `space-between` で調整してください。
    - レイアウト内での成長や縮小を防ぐために、幅を一定にするため `flex` プロパティを `0 0 200px` に設定してください。
    - 異なるカードがどのようにスペースを取るかを見るために、`border` プロパティを `2px solid black` に設定してください。
11. `.basic-plan` 要素はレイアウトで最初に表示されるべきです。これには `order` プロパティを使ってください。
12. `.pro-plan` 要素はレイアウトで2番目に表示されるべきです。`order` プロパティを使い、`flex-grow` プロパティを `2` に設定して他のプランより多くのスペースを取るようにしてください。
13. `.premium-plan` 要素はレイアウトで最後に表示されるべきです。これには `order` プロパティを使ってください。

**注意:** スタイルシートをHTMLにリンクし、CSSを適用してください。

# --hints--

`h1` 要素があるべきです。

```js
assert.exists(document.querySelector('h1'));
```

ページにはテキスト `h1` を持つ `Pricing Plans` 要素があるべきです。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

`div` 要素の下に `h1` 要素があるべきです。

```js
assert.exists(document.querySelector('h1 + div'));
```

`div` 要素はクラス `pricing-container` を持つべきです。

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

`.pricing-container` セレクターがあるべきです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

`pricing-container` セレクターは `display` プロパティを値 `flex` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

`.pricing-container` セレクターは `flex-wrap` プロパティを値 `wrap` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

`div` 要素の中にクラス `pricing-card` を持つ `.pricing-container` 要素が3つあるべきです。

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

`.pricing-card` 要素のうち1つはクラス `basic-plan` を持つべきです。

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

`.pricing-card` 要素のうち1つはクラス `premium-plan` を持つべきです。

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

`.pricing-card` 要素のうち1つはクラス `pro-plan` を持つべきです。

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

`.pricing-card` セレクターがあるべきです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

`.pricing-card` セレクターは `display` プロパティを値 `flex` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

`.pricing-card` セレクターは `flex-direction` プロパティを値 `column` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

`.pricing-card` セレクターは `justify-content` プロパティを値 `space-between` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

`.pricing-card` セレクターは `flex` プロパティを値 `0 0 200px` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

`.pricing-card` セレクターは `border` プロパティを値 `2px solid black` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

`.basic-plan` セレクターは `order` プロパティを値 `0` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

`.basic-plan` 要素にはテキスト `h2` を持つ `Basic` 要素があるべきです。

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

`.basic-plan` 要素にはテキスト `p` を持つ `$9/month` 要素があるべきです。

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

`.pro-plan` 要素にはテキスト `h2` を持つ `Pro` 要素があるべきです。

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

`.pro-plan` 要素にはテキスト `p` を持つ `$19/month` 要素があるべきです。

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

`.pro-plan` 要素は `order` プロパティを値 `1` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

`.pro-plan` 要素は `flex-grow` プロパティを値 `2` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

`.premium-plan` 要素にはテキスト `h2` を持つ `Premium` 要素があるべきです。

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

`.premium-plan` 要素にはテキスト `p` を持つ `$29/month` 要素があるべきです。

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

`premium-plan` 要素はレイアウトで最後の要素であり、`order` プロパティを値 `2` で持つべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```
