---
id: 66a207974c806a19d6607073
title: Construye una Aplicación de Inventario de Libros
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Instrucciones:**

1. Debes tener un elemento `h1` con el texto `Book Inventory`.
1. Debes tener un elemento `table` con columnas llamadas `Title`, `Author`, `Category`, `Status` y `Rate`.
1. Tu tabla debe tener al menos cuatro filas, la primera para los encabezados de columna y las demás llenas de información del libro.
1. Cada fila de la tabla dentro del cuerpo de la tabla debe tener la clase `read`, `to-read` o `in-progress`.
1. Los elementos `td` de la columna `Status` deben contener un elemento `span` con la `class` de `status`, rodeando el texto `Read`, `To Read` o `In Progress`, dependiendo de la clase de esa fila.
1. Los elementos `td` de la columna `Rate` deben contener un elemento `span` con la `class` de `rate` que envuelve tres elementos `span` vacíos.
1. Los elementos `.rate` colocados dentro de filas `.read` deben tener una clase adicional con el valor de `one`, `two` o `three`, dependiendo de la calificación personal. Este valor debe venir después de `rate`.
1. Debes crear tres selectores de atributo para dirigir los elementos con la clase `read`, `to-read` e `in-progress`, y establecer su propiedad `background-image` para usar un `linear-gradient` de tu elección.
1. Debes establecer la propiedad `display` de cada elemento `span` en `inline-block`.
1. Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `to-read` y establecer sus propiedades `border` y `background-image`.
1. Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `read` y establecer sus propiedades `border` y `background-image`.
1. Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `in-progress` y establecer sus propiedades `border` y `background-image`.
1. Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` y los elementos `span` con el valor de clase que comienza con `rate` y establecer sus propiedades `height`, `width` y `padding`.
1. Debes utilizar un selector de atributo para dirigir los elementos `span` que son hijos directos de elementos `span` que tienen el valor de clase que comienza con `rate` y establecer sus propiedades `border`, `border-radius`, `margin`, `height`, `width` y `background-color`.
1. Debes utilizar un selector de atributo para dirigir el primer descendiente de elementos `span` que tienen la palabra `one` como parte de su valor `class` y establecer su propiedad `background-image` para usar un `linear-gradient`.
1. Debes utilizar un selector de atributo para dirigir los dos primeros descendientes de elementos `span` que tienen la palabra `two` como parte de su valor `class` y establecer su propiedad `background-image` para usar un `linear-gradient`.
1. Debes utilizar un selector de atributo para dirigir los tres elementos `span` que son descendientes de elementos `span` que tienen la palabra `three` como parte de su valor `class` y establecer su propiedad `background-image` para usar un `linear-gradient`.

# --hints--

Debes tener un elemento `h1` con el texto `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Book Inventory');
```

Debes tener sólo un elemento `h1`.

```js
assert.equal(document.querySelectorAll('h1')?.length, 1);
```

You should have a `table` element.

```js
assert.equal(document.querySelectorAll('table')?.length, 1);
```

Tu elemento `table` debe tener cinco columnas.

```js
assert.equal(document.querySelectorAll('th')?.length, 5);
```

Tu primera columna debe tener el texto `Title` como encabezado.

```js
assert.equal(document.querySelectorAll('th')[0]?.innerText, 'Title');
```

Tu segunda columna debe tener el texto `Author` como encabezado.

```js
assert.equal(document.querySelectorAll('th')[1]?.innerText, 'Author');
```

Tu tercera columna debe tener el texto `Category` como encabezado.

```js
assert.equal(document.querySelectorAll('th')[2]?.innerText, 'Category');
```

Tu cuarta columna debe tener el texto `Status` como encabezado.

```js
assert.equal(document.querySelectorAll('th')[3]?.innerText, 'Status');
```

Tu quinta columna debe tener el texto `Rate` como encabezado.

```js
assert.equal(document.querySelectorAll('th')[4]?.innerText, 'Rate');
```

Tu tabla debe tener al menos cuatro filas.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Cada fila de la tabla excepto la fila de encabezado debe tener la clase `read`, `to-read` o `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

Los elementos `td` de la columna `Status` deben contener un elemento `span`.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isAbove(statusColumnData.length, 0);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Cada elemento `span` de la columna `Status` debe tener la clase de `status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isAbove(statusSpans.length, 0);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Cada elemento `.status` debe tener el texto `Read`, `To Read` o `In Progress`, dependiendo de la clase de su fila.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isAbove(statusSpans.length, 0);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

Los elementos `td` de la columna `Rate` deben contener un elemento `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isAbove(rateColumnData.length, 0);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Cada elemento `span` que es un hijo directo de un elemento `td` de la columna `Rate` debe tener la clase de `rate` como la primera clase.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Cada elemento `.rate` debe contener tres elementos `span` vacíos.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

Los elementos `.rate` colocados dentro de filas `.read` deben tener una clase adicional después de la clase `rate` con el valor de `one`, `two` o `three`, dependiendo de la calificación personal.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isAbove(readBooksRates.length, 0);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

Debes tener un selector de atributo para dirigir filas que tengan la clase de `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Debes utilizar un selector de atributo para dirigir filas que tengan la clase de `read` y establecer su propiedad `background-image` para un gradiente lineal de tu elección.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

Debes tener un selector de atributo para dirigir filas que tengan la clase de `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Debes utilizar un selector de atributo para dirigir filas que tengan la clase de `to-read` y establecer su propiedad `background-image` para un gradiente lineal de tu elección.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

Debes tener un selector de atributo para dirigir filas que tengan la clase de `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Debes utilizar un selector de atributo para dirigir filas que tengan la clase de `in-progress` y establecer su propiedad `background-image` para un gradiente lineal de tu elección.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Debes establecer la propiedad `display` de cada elemento `span` en `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

Debes tener un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `to-read` y establecer su propiedad `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `to-read` y establecer su propiedad `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

Debes tener un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `read` y establecer su propiedad `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `read` y establecer su propiedad `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

Debes tener un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `in-progress` y establecer su propiedad `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` que son descendientes de elementos `tr` con la clase `in-progress` y establecer su propiedad `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

Debes tener un selector de atributo para dirigir los elementos `span` con la clase `status` y los elementos `span` con el valor de clase que comienza con `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` y los elementos `span` con el valor de clase que comienza con `rate` y establecer su propiedad `height`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` y los elementos `span` con el valor de clase que comienza con `rate` y establecer su propiedad `width`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Debes utilizar un selector de atributo para dirigir los elementos `span` con la clase `status` y los elementos `span` con el valor de clase que comienza con `rate` y establecer su propiedad `padding`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

Debes tener un selector de atributo para dirigir los elementos `span` que son hijos directos de elementos `span` con el valor de clase que comienza con `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` que son hijos directos de elementos `span` con el valor de clase que comienza con `rate` y establecer su propiedad `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` que son hijos directos de elementos `span` con el valor de clase que comienza con `rate` y establecer su propiedad `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` que son hijos directos de elementos `span` con el valor de clase que comienza con `rate` y establecer su propiedad `margin`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` que son hijos directos de elementos `span` con el valor de clase que comienza con `rate` y establecer su propiedad `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` que son hijos directos de elementos `span` con el valor de clase que comienza con `rate` y establecer su propiedad `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` que son hijos directos de elementos `span` con el valor de clase que comienza con `rate` y establecer su propiedad `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

Debes tener un selector de atributo para dirigir el primer descendiente de elementos `span` que tienen la palabra `one` como parte de su valor `class`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Debes utilizar un selector de atributo para dirigir el primer descendiente de elementos `span` que tienen la palabra `one` como parte de su valor `class` y establecer su propiedad `background-image` para usar un `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

Debes tener un selector de atributo para dirigir los dos primeros descendientes de elementos `span` que tienen la palabra `two` como parte de su valor `class`.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Debes utilizar un selector de atributo para dirigir los dos primeros descendientes de elementos `span` que tienen la palabra `two` como parte de su valor `class` y establecer su propiedad `background-image` para usar un `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

Debes tener un selector de atributo para dirigir los elementos `span` que son descendientes de elementos `span` que tienen la palabra `three` como parte de su valor `class`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Debes utilizar un selector de atributo para dirigir los elementos `span` que son descendientes de elementos `span` que tienen la palabra `three` como parte de su valor `class` y establecer su propiedad `background-image` para usar un `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
