---
id: 6a3e6c0eb5d050ce07781709
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Si la condición `if` evalúa a `False` (lo que significa que *sí* tienes suficientes huevos), la función omitirá el bloque `if` y continuará ejecutando el resto del código.

Ahora necesitas procesar tus datos y devolver un nuevo valor calculado.

Fuera de tu bloque `if` (al mismo nivel de indentación que la palabra clave `if`), usa una llamada a `print()` y una f-string para imprimir `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`.

Luego, en la siguiente línea, `return` el número restante de huevos restando `eggs_to_use` de `available_eggs`.

# --hints--

Deberías tener una sentencia print que imprima `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` en tu función `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

Deberías devolver el número restante de huevos restando `eggs_to_use` de `available_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

Deberías devolver `available_eggs - eggs_to_use` después de la sentencia `print`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
