---
id: 6a3ea8a55317ba7a08e21d5b
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

Ahora finalmente puedes responder la pregunta de antes: ¿por qué el número global de huevos se mantuvo en `1`? Cuando una función devuelve un valor, no sobrescribe automáticamente tus variables existentes. ¡Debes guardar explícitamente ese nuevo valor!

Estás usando la variable `available_eggs` dentro de la función `make_fried_egg`, pero no es la misma variable definida al inicio de tu código. La que está dentro de la función es una variable local que solo existe mientras la función se ejecuta. Cuando la función termina, la variable local se descarta y la variable global `available_eggs` permanece sin cambios.

Para actualizar tu inventario real de cocina, sobrescribe la variable global `available_eggs` con el nuevo número de huevos que devuelve tu función `make_fried_egg`.

Al final de tu código, llama a la función `make_fried_egg`, pasando `available_eggs` como argumento, y asigna el resultado de nuevo a la variable global `available_eggs`.

# --hints--

Debes reasignar la variable `available_eggs` al resultado de `make_fried_egg(available_eggs)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
