---
id: 6a3e6ad6b9554e31dcc30cd6
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Si no tienes suficientes huevos, ¡no deberías intentar cocinar! Necesitas una forma de detener la ejecución de la función.

Puedes usar la palabra clave `return` para esto. Cuando una función encuentra una sentencia `return`, deja de ejecutarse inmediatamente y devuelve el valor retornado al lugar donde fue llamada.

Por ejemplo:

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

Dentro de tu bloque `if`, debajo de la sentencia para imprimir, usa la palabra clave `return` para devolver el valor de `available_eggs`.

# --hints--

Debes devolver `available_eggs` dentro de tu sentencia `if`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

Debes devolver `available_eggs` después de la sentencia `print` dentro de tu bloque `if`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
