---
id: 6a3e6fd08125dfa615cb431b
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Ahora que tu función `use_eggs` está completa, puedes probarla. Al llamar a una función que requiere **parámetros**, pasas los valores específicos dentro de los paréntesis. Estos valores se conocen como **argumentos**.

Aquí tienes un ejemplo de cómo llamar a una función con argumentos:

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

Al final de tu código, llama a tu función `use_eggs` y pásale dos argumentos: la variable global `available_eggs` y el número `1` (que representa la cantidad de huevos que quieres usar).

# --hints--

Debes llamar a la función `use_eggs` y pasar `available_eggs` y `1` como argumentos.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
