---
id: 6a3e71a0a3a44260e675b157
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Para resolver el misterio de por qué tu conteo de huevos no cambió y realmente cocinar el desayuno, crea una nueva función. Piensa en esta función como tu receta, una que llama a otras funciones para hacer el trabajo.

Debajo de tu sentencia `print(available_eggs)`, define una función vacía llamada `make_fried_egg` que tome `available_eggs` como su único parámetro.

Usa la palabra clave `pass` dentro del cuerpo por ahora.

# --hints--

Debes definir una función llamada `make_fried_egg`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

Tu función `make_fried_egg` debe tener `available_eggs` como parámetro.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
