---
id: 6a3e70b94f67d9ee3ecd8486
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Acabas de llamar a `use_eggs(available_eggs, 1)` para usar un huevo. Eso significa que ahora deberías tener `0` huevos en tu cocina, ¿verdad?

¡Lo descubrirás! Al final de tu código, usa la función `print()` para imprimir el valor de la variable `available_eggs`.

Cuando ejecutes el código, notarás algo sorprendente: la salida sigue siendo `1`. ¿Por qué no cambió el número de huevos? Explorarás este misterio en los próximos pasos.

# --hints--

Debes imprimir la variable `available_eggs` al final de tu código.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
