---
id: 6a3e67fc5ef4045201c28219
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Dentro de tu función `use_eggs`, necesitas verificar si tienes suficientes huevos para cocinar. Puedes hacer esto usando una sentencia `if`.

Una sentencia `if` evalúa una condición. Si la condición es verdadera, se ejecuta el código dentro del bloque `if`.

Reemplaza la palabra clave `pass` con una sentencia `if` que verifique si `eggs_to_use` es mayor (`>`) que `available_eggs`. Dentro del bloque `if`, imprime la cadena `The kitchen does not have enough eggs.`.

# --hints--

Debes tener una sentencia `if` que verifique si `eggs_to_use > available_eggs`.

```js
({
  test: () => runPython(`
  cond = _Node(_code).find_function("use_eggs").find_ifs()[0].find_conditions()[0]
  assert cond.is_equivalent("eggs_to_use > available_eggs") or cond.is_equivalent("available_eggs < eggs_to_use")`)
})
```

Debes imprimir `The kitchen does not have enough eggs.` dentro de tu sentencia `if`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_call("print('The kitchen does not have enough eggs.')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
