---
id: 6a3e58891e1ff13d07f4ab48
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Python lee el bloque `def` y recuerda la función, pero no ejecuta el código dentro de ella inmediatamente.

Para ejecutar el código y ver tu salida, necesitas **llamar** (o invocar) la función. Hazlo escribiendo el nombre de la función seguido de paréntesis.

Por ejemplo:

```py
my_function()
```

Al final de tu código, afuera de la función, llama a tu función `check_kitchen_stock`.

# --hints--

Debes llamar a la función `check_kitchen_stock` al final de tu código. Asegúrate de que no esté indentada, para que se ejecute globalmente.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
