---
id: 6a3e7284e89ce3367e709331
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Antes de hacer un huevo frito, primero necesitas verificar si tienes los ingredientes requeridos. Puedes hacer esto comparando la cantidad disponible de huevos con la cantidad necesaria para la receta.

Dentro de tu función `make_fried_egg`, reemplaza la palabra clave `pass` creando una variable llamada `has_enough_eggs` y asignándole el resultado de verificar si `available_eggs` es mayor o igual a `1`.

# --hints--

Debes asignar `has_enough_eggs` al resultado de verificar si `available_eggs` es mayor o igual a `1`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_variables("has_enough_eggs")[0]
assert node.is_equivalent("has_enough_eggs = available_eggs >= 1") or node.is_equivalent("has_enough_eggs = 1 <= available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
