---
id: 6a3ea771b46a239c67628a07
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Tu función `make_fried_egg` ahora verifica si hay un huevo disponible y actualiza su conteo local de huevos cuando se usa uno. Luego, devuelve ese conteo para que quien llama pueda usarlo.

Al final de tu función `make_fried_egg`, usa la palabra clave `return` para devolver tu variable `available_eggs`. Asegúrate de que esta línea no esté indentada dentro de los bloques `if` o `else`.

# --hints--

Debes devolver `available_eggs` al final de tu función `make_fried_egg`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_return("available_eggs")
`)
})
```

Debes devolver `available_eggs` después de la sentencia `if/else`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg")
assert node.is_ordered(str(node.find_ifs()[0]), "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

--fcc-editable-region--
    
--fcc-editable-region--
```
