---
id: 6a3e5ccb99fd6053af9032a1
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Ahora que sabes qué hay en tu cocina, puedes usar esos elementos para cocinar. Para hacerlo, crea una nueva función que acepte parámetros. Piensa en los parámetros como marcadores de posición para los datos que pasas a la función cuando la llamas.

Aquí tienes una función que acepta dos parámetros, `a` y `b`, y imprime su suma en la terminal.

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

Debajo de tu función `check_kitchen_stock`, define una función vacía llamada `use_eggs` que acepte dos parámetros: `available_eggs` y `eggs_to_use`, en ese orden.

Ten en cuenta que si le das a un parámetro el mismo nombre que una variable global (como `available_eggs`), la función usará su propio parámetro local `available_eggs` en lugar del global. Esto se conoce como sombreado de nombres porque la variable local "ensombrece" a la global.

Usa la palabra clave `pass` dentro del cuerpo por ahora.

# --hints--

Debes definir una función llamada `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

Tu función `use_eggs` debe tener como parámetros `available_eggs` y `eggs_to_use`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
