---
id: 6a3e2b9b7eeca084dd4fb857
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

¡Bienvenido a la cocina! Así como un chef se apoya en recetas para preparar una comida, tú te apoyarás en funciones para construir software. En este taller, aprenderás a escribir funciones reutilizables en Python y a dominar el ámbito, entendiendo exactamente qué ingredientes (variables) están disponibles en tu código en cada momento.

Para comenzar, debes definir algunos datos básicos para el inventario de tu cocina. Crea tres variables: `available_eggs`, `available_flour` y `available_sugar`, y asígnales los valores `1`, `2` y `3`, respectivamente.

# --hints--

Debes definir una variable llamada `available_eggs` y asignarle el valor `1`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

Debes definir una variable llamada `available_flour` y asignarle el valor `2`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

Debes definir una variable llamada `available_sugar` y asignarle el valor `3`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
