---
id: 6a3e73eabddd614f91cc235a
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

¡Ahora vas a resolver el misterio de antes! Cuando tu función `use_eggs` devolvió el número restante de huevos, no guardaste ese nuevo valor en ningún lado. ¡Simplemente desapareció en el vacío!

Para actualizar realmente tu inventario, necesitas capturar ese valor devuelto asignándolo a una variable.

Dentro de tu función `make_fried_egg`, usa una sentencia `if/else` para verificar la variable `has_enough_eggs`.

- Si `has_enough_eggs` es `True`, reasigna tu variable `available_eggs` al resultado de `use_eggs(available_eggs, 1)` e imprime el mensaje `Made a fried egg. Yummy!`.

- De lo contrario, dentro de un bloque `else`, imprime el mensaje `Could not make a fried egg. Not enough eggs!`.

# --hints--

Debes tener una sentencia `if` y usar `has_enough_eggs` como condición.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

Dentro de tu bloque `if`, debes reasignar `available_eggs` al resultado de `use_eggs(available_eggs, 1)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

Debes imprimir el mensaje `Made a fried egg. Yummy!` cuando `has_enough_eggs` sea `True`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

Debes tener un bloque `else` que imprima el mensaje `Could not make a fried egg. Not enough eggs!` cuando `has_enough_eggs` sea `False`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
