---
id: 6a3e5475c35dbc138d77c6f4
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora que has calculado tu total, deberías mostrar el inventario. Recuerda de lecciones anteriores que puedes usar la función `print()` con **f-strings** para incluir variables en tu salida.

Dentro de tu función `check_kitchen_stock`, debajo de tu variable `total_items`, agrega estas cuatro sentencias print exactas:

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

Debes tener la sentencia `print(f'The kitchen has {total_items} total items:')` dentro de tu función.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

La sentencia `print(f'The kitchen has {total_items} total items:')` debe ser la primera llamada a `print()` dentro de tu función.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

Debes tener la sentencia `print(f'- {available_eggs} eggs')` dentro de tu función.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

Debes tener la sentencia `print(f'- {available_flour} flour')` dentro de tu función.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

Debes tener la sentencia `print(f'- {available_sugar} sugar')` dentro de tu función.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
