---
id: 6a3e524fc3394ef52ed5b797
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Las funciones tienen su propio ámbito local. Esto significa que las variables creadas dentro de una función son como ingredientes guardados en la despensa de esa cocina: solo se puede acceder a ellas dentro de esa función específica.

Sin embargo, las funciones no están completamente aisladas. Pueden leer variables globales definidas afuera de su ámbito. Pueden usar esas variables para realizar cálculos, tomar decisiones o crear nuevas variables locales.

Dentro de tu función `check_kitchen_stock`, reemplaza la palabra clave `pass` con una nueva variable llamada `total_items` y asígnale la suma de `available_eggs`, `available_flour` y `available_sugar`.

# --hints--

Debes definir una variable llamada `total_items` dentro de tu función `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

Tu variable `total_items` debe ser la suma de `available_eggs`, `available_flour` y `available_sugar`.

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
