---
id: 6a3e4ae26f2b7e4b5b026cf8
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Es hora de crear una función para verificar lo que tienes en tu cocina. En Python, defines funciones usando la palabra clave `def`, seguida del nombre de la función, paréntesis `()` y dos puntos `:`.

Aquí tienes un ejemplo de cómo definir una función simple:

```py
def my_function():
    pass
```

Define una función vacía llamada `check_kitchen_stock()`. Usa la palabra clave `pass` dentro del cuerpo de la función por ahora para que Python no genere un error de sintaxis.

Recuerda que Python depende de la indentación para definir bloques de código. Así que asegúrate de indentar el cuerpo de la función.

# --hints--

Debes definir una función llamada `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
