---
id: 615380dff67172357fcf0425
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

이제 이미지가 너무 큽니다.

이미지를 대상으로 하는 `.gallery img` 선택자를 만드세요. 모든 이미지에 `width` 속성을 `100%`로, `max-width` 속성을 `350px`로 지정하세요.

또한 이미지 크기를 균일하게 유지하려면 `height` 속성을 `300px`로 설정하세요.

# --hints--

`.gallery img` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

`.gallery img` 선택자에 `width` 속성을 값으로 `100%`으로 설정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width, '100%');
```

`.gallery img` 선택자에 `max-width` 속성을 값으로 `350px`으로 설정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth, '350px');
```

`.gallery img` 선택자에 `height` 속성을 값으로 `300px`으로 설정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height, '300px');
```

# --seed--


## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
