---
id: 6153893900438b4643590367
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

`.header` 텍스트를 가운데 정렬하세요. `text-transform` 속성을 사용하고 값으로 `uppercase`를 지정해 텍스트를 대문자로 만드세요.

모든 면에 `32px`만큼 패딩을 주세요. 배경색은 `#0a0a23`로, 텍스트 색상은 `#fff`로 설정하세요.

값이 `border-bottom`인 `4px solid #fdb347`를 추가하세요.

# --hints--

`.header` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

`.header` 선택자에 `text-align` 속성을 값 `center`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

`.header` 선택자에 `text-transform` 속성을 값 `uppercase`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

`.header` 선택자에 `padding` 속성을 값 `32px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

`.header` 선택자에 `background-color` 속성을 값 `#0a0a23`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

`.header` 선택자에 `color` 속성을 값 `#fff`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

`.header` 선택자에 `border-bottom` 속성을 값 `4px solid #fdb347`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
