---
id: 61537c5f81f0cf325b4a854c
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

`header` 요소 안에 `body` 요소를 추가하고, 여기에 `header` 클래스를 할당하세요.

`header` 안에 `h1`를 만들고, 텍스트로 `css flexbox photo gallery`를 넣으세요.

# --hints--

`header` 요소 안에 `body` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

`header` 요소에는 값이 `class`인 `header`가 있어야 합니다.

```js
assert.isTrue(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

`header` 요소 안에 `h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

`h1` 요소 안에 `css flexbox photo gallery` 텍스트가 있어야 합니다.

```js
assert.equal(document?.querySelector('.header')?.querySelector('h1')?.innerText, 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
