---
id: 6153908a366afb4d57185c8d
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

플렉스박스에는 주축과 교차 축이 있습니다. 주축은 `flex-direction` 속성으로 정의되며, 네 가지 가능한 값이 있습니다:

- `row` (기본값): 플렉스 항목이 왼쪽에서 오른쪽으로 배치되는 수평 축
- `row-reverse`: 플렉스 항목이 오른쪽에서 왼쪽으로 배치되는 수평 축
- `column`: 플렉스 항목이 위에서 아래로 배치되는 수직 축
- `column-reverse`: 플렉스 항목이 아래에서 위로 배치되는 수직 축

**참고**: 축과 방향은 텍스트 방향에 따라 달라집니다. 여기서 보여주는 값은 왼쪽에서 오른쪽으로 읽는 텍스트 방향에 해당합니다.

다양한 값을 시도해 보면서 레이아웃에 어떤 영향을 주는지 확인하세요.

완료하면 `flex-direction` 요소에 `row`를 `.gallery`로 명시적으로 설정하세요.

# --hints--

`.gallery` 선택자에 `flex-direction` 속성을 `row` 값으로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection, 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
