---
id: 615392916d83fa4f02f7e2cf
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--


`flex-wrap` 속성은 플렉스 컨테이너가 너무 작을 때 플렉스 항목이 어떻게 동작할지 결정합니다. `wrap`로 설정하면 항목들이 다음 행이나 열로 감싸지도록 허용합니다. 기본값인 `nowrap`는 항목이 감싸지지 않도록 막고 필요하면 크기를 줄입니다.


플렉스 항목이 공간이 부족할 때 다음 행으로 감싸지도록 만드세요.

# --hints--

`.gallery` 선택자에 `flex-wrap` 속성을 `wrap` 값으로 설정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
  flex-direction: row;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
