---
id: 6153938dce8b294ff8f5a4e9
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

`justify-content` 속성은 플렉스 컨테이너 내부 항목들이 주축을 따라 어떻게 배치되는지 결정하며, 항목들의 위치와 주변 공간에 영향을 줍니다.

`.gallery` 선택자에 `justify-content` 속성을 `center` 값으로 지정하세요.

# --hints--

`.gallery` 선택자에 `justify-content` 속성을 `center` 값으로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.justifyContent, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
