---
id: 61537bb9b1a29430ac15ad38
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 프로젝트는 `styles.css` 파일을 페이지에 연결하는 것부터 시작하세요.

# --hints--

`link` 요소 안에 `head` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("head>link"));
```

`link` 요소에는 `rel="stylesheet"`와 `href="styles.css"`가 있어야 합니다.

```js 
  const link = document.querySelector("head>link"); 
  assert.equal(link?.getAttribute("rel"), "stylesheet"); 
  const href = link?.getAttribute("data-href"); 
  assert.oneOf(href, ["./styles.css", "styles.css"]); 
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
--fcc-editable-region--

--fcc-editable-region--    
  </head>
  <body>
  </body>
</html>
```

```css

```
