---
id: 61537c9eecea6a335db6da79
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

`.header` 요소 아래에 새로운 `div` 요소를 만들고 `class` 값을 가진 `gallery`를 할당하세요. 이 `div`는 갤러리 이미지의 컨테이너 역할을 합니다.

그 `.gallery` 요소 안에 아홉 개의 `img` 요소를 만드세요.

# --hints--

`div` 요소 안에 `body` 요소를 만들어야 합니다.

```js
assert.isAtLeast(document.querySelector('body')?.querySelectorAll('div')?.length, 1);
```

새로운 `div` 요소에는 값이 `class`인 `gallery`이 있어야 합니다.

```js
assert.exists(document.querySelector('body')?.querySelector('.gallery'));
```

새로운 `div` 요소는 `.header` 요소 다음에 와야 합니다.

```js
assert.exists(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

`.gallery` 요소에는 아홉 개의 `img` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelector('.gallery')?.querySelectorAll('img'), 9);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
