---
id: 6493bc0d99879635209565aa
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

`border-box` 크기 조절 모델은 `content-box`와 반대입니다. 패딩과 테두리를 포함한 요소의 전체 너비가 명시적으로 설정한 너비가 됩니다. 요소의 내용은 패딩과 테두리를 위한 공간을 만들기 위해 줄어듭니다.

`box-sizing` 속성을 `border-box`로 변경하세요. 파란색 이미지 테두리가 빨간색 갤러리 테두리 안에 맞춰진 것을 확인하세요.

# --hints--


* 선택자에 `box-sizing` 속성을 `border-box` 값으로 설정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>


```

```css
 * {
--fcc-editable-region--
  box-sizing: content-box;
--fcc-editable-region--
}

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
