---
id: 64dd076a1b33c86d84b4232f
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

요소 크기를 더 잘 시각화하려면 테두리를 추가하는 것이 도움이 됩니다.

`.gallery` 요소에 `width`의 `50%`과 `border`로 설정된 `5px solid red`을 지정하세요.

그다음 `img` 요소들에 `width`의 `100%`, `padding`로 설정된 `5px`, 그리고 `border`로 설정된 `5px solid blue`을 지정하세요.

# --hints--

`.gallery` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

`.gallery` 선택자에는 `width` 속성이 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

`.gallery` 선택자에는 `width`로 설정된 `50%` 속성이 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

`.gallery` 선택자에는 `border` 속성이 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

`.gallery` 선택자에는 `border`로 설정된 `5px solid red` 속성이 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

`img` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

`img` 선택자에는 `width` 속성이 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

`img` 선택자에 `width` 속성을 `100%`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

`img` 선택자에는 `padding` 속성이 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

`img` 선택자에는 `padding`로 설정된 `5px` 속성이 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

`img` 선택자에는 `border` 속성이 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

`img` 선택자에는 `border`로 설정된 `5px solid blue` 속성이 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
