---
id: 66c06d618d075c7f7f1b890a
title: Створіть генератор передбачень
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

У цій лабораторній роботі ви визначите п’ять передбачень і випадково оберете одне, яке відтворюватиметься користувачу.

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Надайте п’ятьом змінним `fortune1`, `fortune2`, `fortune3`, `fortune4` та `fortune5` значення рядка на власний вибір. Якщо хочете, можете використати такі значення:

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. Згенеруйте випадкове число від 1 до 5 (включно) та призначте його до змінної `randomNumber`.

3. Створіть змінну `selectedFortune` і призначте до неї передбачення на основі цих правил:

   - якщо `randomNumber` дорівнює 1, то призначте значення змінної `fortune1` до `selectedFortune`;
   - якщо `randomNumber` дорівнює 2, то призначте значення змінної `fortune2` до `selectedFortune`;
   - якщо `randomNumber` дорівнює 3, то призначте значення змінної `fortune3` до `selectedFortune`;
   - якщо `randomNumber` дорівнює 4, то призначте значення змінної `fortune4` до `selectedFortune`;
   - якщо `randomNumber` дорівнює 5, то призначте значення змінної `fortune5` до `selectedFortune`.

4. Виведіть результат `selectedFortune` на консоль.

# --hints--

Надайте значення рядка до `fortune1`.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

Надайте значення рядка до `fortune2`.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

Надайте значення рядка до `fortune3`.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

Надайте значення рядка до `fortune4`.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

Надайте значення рядка до `fortune5`.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

Використайте метод `Math.random()`, щоб згенерувати випадкове число.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

Згенеруйте випадкове число від 1 до 5 (включно) та призначте його до змінної `randomNumber`.

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Ви повинні мати змінну `selectedFortune`, якій призначено значення на основі значення змінної `randomNumber`.

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` має відповідати передбаченню. Наприклад, якщо `randomNumber` дорівнює 1, то `selectedFortune` має дорівнювати `fortune1` і так далі.

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Виведіть результат `selectedFortune` на консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
