---
id: 66c06d618d075c7f7f1b890a
title: Створіть ворожку
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

У цій лабораторній роботі ви визначите п’ять ворожінь і випадковим чином оберете одне з них, щоб показати користувачу.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні ініціалізувати п’ять змінних `fortune1`, `fortune2`, `fortune3`, `fortune4` і `fortune5` рядковим значенням на свій вибір. Якщо хочете, можете використати наведені нижче варіанти:

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. Ви повинні вибрати випадкове число від 1 до 5 включно і призначити його змінній `randomNumber`.

3. Ви повинні створити змінну `selectedFortune` і призначити відповідне ворожіння на основі таких правил:

   - Якщо `randomNumber` дорівнює 1, призначте значення `fortune1` змінній `selectedFortune`.
   - Якщо `randomNumber` дорівнює 2, призначте значення `fortune2` змінній `selectedFortune`.
   - Якщо `randomNumber` дорівнює 3, призначте значення `fortune3` змінній `selectedFortune`.
   - Якщо `randomNumber` дорівнює 4, призначте значення `fortune4` змінній `selectedFortune`.
   - Якщо `randomNumber` дорівнює 5, призначте значення `fortune5` змінній `selectedFortune`.

4. Ви повинні вивести `selectedFortune` у консоль.

# --hints--

Ви повинні ініціалізувати `fortune1` рядковим значенням.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

Ви повинні ініціалізувати `fortune2` рядковим значенням.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

Ви повинні ініціалізувати `fortune3` рядковим значенням.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

Ви повинні ініціалізувати `fortune4` рядковим значенням.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

Ви повинні ініціалізувати `fortune5` рядковим значенням.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

Ви повинні використати метод `Math.random()` для генерації випадкового числа.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

Ви повинні згенерувати випадкове число від 1 до 5 включно і призначити його змінній `randomNumber`.

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Ви повинні мати змінну `selectedFortune`, якій призначене значення на основі значення `randomNumber`.

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` має відповідати своєму ворожінню. Наприклад, якщо `randomNumber` дорівнює 1, `selectedFortune` має бути рівним `fortune1` і так далі.

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Ви повинні вивести `selectedFortune` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
