---
id: 6391d1a4f7ac71efd0621380
title: 建置一個食譜頁面專案
challengeType: 25
dashedName: top-build-a-recipe-project
---

# --description--

該網站將包含一個主要的索引頁面，該頁面將有指向幾個食譜的聯結。當你完成時，該網站看起來不會很漂亮。

**使用者故事：**

1. 你的食譜頁面應該包含一個 `DOCTYPE` 頁籤。
1. 你的食譜頁面應該包含一個 `html` 元素，該元素有 `head` 和 `body` 元素作為子元素。
1. 你應該在 `head` 元素內有一個 `title` 元素，裡面的文字是 `The Odin Recipes`。
1. 你應該會看到一個 `h1` 元素，其文字為 `Creamy Chocolate Fudge`。
1. 你應該會看到帶有 `alt` 屬性的相關影像。
1. 在影像下方應該有一個 `h2` 元素，其文字為 `Description`。
1. 你應該會看到 `Description` 底下有幾段描述這個配方的段落。
1. 應該有一個 `h2` 元素，其文字為 `Ingredients`。
1. 在 `Ingredients` 標題下應該有一個無序列表，列出食譜所需的材料。
1. 在成分列表下方添加另一個標題，稱為 `Steps`。
1. 你應該會看到一個有序列表，裡面有完成此配方所需的幾個步驟。
1. 在步驟下方應該有一個 `h2` 元素，文字為 `More Recipes`。
1. 你應該會看到一個無序列表，裡面有幾個列表項目，且每個列表項目中都有指向其他食譜的聯結。
1. 這些錨點元素應該有一個 `href` 屬性，其值設定為 `#`。

# --hints--

你應該有一個 `DOCTYPE` 頁籤。

```js
assert(code.match(/<!DOCTYPE\s+?html\s*?>/gi));
```

你應該有一個包含 `head` 和 `body` 元素的 `html` 元素。

```js
const html = document.querySelectorAll('html')[0];
const head = document.querySelectorAll('html > head')[0];
const body = document.querySelectorAll('html > body')[0];

assert(html && head && body);
```

你應該在包含文字 `The Odin Recipes` 的 `head` 元素中有一個 `title` 元素。

```js
assert(document.querySelectorAll('HEAD > TITLE')[0]?.innerText == 'The Odin Recipes');
```

你應該在你的 `body` 元素中有一個 `h1` 元素，該元素包含文字 `Creamy Chocolate Fudge`。

```js
assert(document.querySelectorAll('BODY > H1')[0]?.innerText == 'Creamy Chocolate Fudge');
```

你應該有一個帶有 `alt` 屬性的影像。

```js
const img = document.querySelectorAll('IMG')[0];

assert(img && img.alt !='' && img.src != '')
```

你應該有一個 `h2` 元素，其文字為 `Description`。

```js
const h2 = document.querySelectorAll('H2')[0];

assert(h2?.innerText == 'Description');
```

你應該至少有兩個 `p` 元素來描述這個食譜。

```js
const paragraphs = document.querySelectorAll('P');

assert(paragraphs?.length > 1);
```

你應該有一個 `h2` 元素，裡面包含文字 `Ingredients`。

```js
const h2 = document.querySelectorAll('H2')[1];

assert(h2?.innerText == 'Ingredients');
```

你應該有一個 `ul` 元素，裡面巢狀包含含有成分的 `li` 元素。

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "Ingredients") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);
        } else {
            assert(false);
        }
    }
});
```

你應該有一個 `h2` 元素，裡面包含文字 `Steps`。

```js
const h2 = document.querySelectorAll('H2')[2];

assert(h2?.innerText == 'Steps');
```

你應該有一個 `ol` 元素，該元素包含有食譜步驟的 `li` 元素。

```js
const orderedList = document.querySelectorAll('OL')[0];
const listItems = document.querySelectorAll('OL > LI');

assert(orderedList && listItems && listItems.length > 1);
```

你應該有一個 `h2` 元素，其文字為 `More Recipes`。 

```js
const h2 = document.querySelectorAll('H2')[3];

assert(h2?.innerText == 'More Recipes');
```

你應該有一個 `ul` 元素，內嵌包含指向其他食譜的 `a` 元素的 `li` 元素。

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "More Recipes") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);

            listItems.forEach(listItem => {
                const anchor = listItem.querySelector("A");

                assert(anchor);
            });
        } else {
            assert(false);
        }
    }
});
```

你的錨點元素指向食譜時，應該有一個 `href` 屬性，其值設定為 `#`。

```js
const anchorTags = document.querySelectorAll("a");

const allAnchorsHaveHrefHash = [...anchorTags].every(function(anchorTag) {
  return anchorTag.hasAttribute("href") && anchorTag.getAttribute("href") === "#";
});

assert(allAnchorsHaveHrefHash && anchorTags.length > 0); 
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>The Odin Recipes</title>
  </head>
  <body>
    <h1>Creamy Chocolate Fudge</h1>
    <img src="https://cdn.freecodecamp.org/curriculum/odin-project/build-a-recipe-page/build-a-recipe-page-01.jpg" alt="A delicious chocolate fudge dessert">
    <h2>Description</h2>
    <p>This recipe is for a rich and creamy chocolate fudge that is sure to satisfy your sweet tooth. It's perfect for a special occasion or as a tasty treat for any time of the year.</p>
    <p>This recipe is easy to follow and only requires a few simple ingredients. With just a few steps, you'll be able to create a delicious dessert that everyone will love.</p>
    <h2>Ingredients</h2>
    <ul>
      <li>1 cup sugar</li>
      <li>1/2 cup unsalted butter</li>
      <li>1/4 cup milk</li>
      <li>1/4 cup cocoa powder</li>
      <li>1/4 cup chocolate chips</li>
      <li>1/4 tsp salt</li>
      <li>1 tsp vanilla extract</li>
    </ul>
    <h2>Steps</h2>
    <ol>
      <li>In a medium saucepan, melt the butter over medium heat.</li>
      <li>Add the sugar, milk, cocoa powder, and salt to the saucepan and stir until well combined.</li>
      <li>Bring the mixture to a boil, stirring constantly, and then reduce the heat to low and simmer for 5 minutes.</li>
      <li>Remove the saucepan from the heat and stir in the chocolate chips and vanilla extract until the chocolate is melted and the mixture is smooth.</li>
      <li>Pour the fudge into a greased 8-inch square pan and let it cool completely before cutting into squares.</li>
    </ol>
    <h2>More Recipes</h2>
    <ul>
      <li><a href="#">Peanut Butter Cookies</a></li>
      <li><a href="#">Lemon Bars</a></li>
      <li><a href="#">Chocolate Chip Pancakes</a></li>
    </ul>
  </body>
</html>
```

```css

```
