---
id: 68de7360e1832ea180577a8a
title: Costruisci un calcolatore del numero di Fibonacci n-esimo
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione chiamata `fibonacci`.
2. Dovresti definire una lista chiamata `sequence` all'interno della funzione `fibonacci`, inizializzata con i valori `[0, 1]`.
3. La funzione `fibonacci` dovrebbe accettare un parametro, un intero positivo `n`.  
4. Chiamare `fibonacci(n)` dovrebbe usare un approccio di programmazione dinamica per calcolare e restituire il numero `n`-esimo della sequenza di Fibonacci, dove ogni numero è la somma dei due numeri precedenti.
5. Ogni numero calcolato alla posizione `n` nella sequenza di Fibonacci dovrebbe essere archiviato nella lista `sequence` all'indice `n - 1`.

# --hints--

Dovresti avere una funzione chiamata `fibonacci`.

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

La tua funzione `fibonacci` dovrebbe prendere un solo parametro.

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

Dovresti avere una lista chiamata `sequence` all'interno della funzione `fibonacci` inizializzata a `[0, 1]`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)` dovrebbe restituire `0`.

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)` dovrebbe restituire `1`.

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)` dovrebbe restituire `1`.

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)` dovrebbe restituire `2`.

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)` dovrebbe restituire `5`.

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)` dovrebbe restituire `55`.

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)` dovrebbe restituire `610`.

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

Non dovresti usare la ricorsione nel tuo codice.

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```
