---
id: 646c59e942f35541923104bf
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

では HTML に戻って、`main` 要素を作成してください。 `main` 要素の中に、`div` 要素を追加してクラス `cat-head` を設定してください。

# --hints--

`main` 要素が 1 つ必要です。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

`div` 要素が 1 つ必要です。

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

`div` 要素にはクラス `cat-head` が必要です。

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

`div` 要素は `main` タグの内側に置いてください。

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
