---
id: 646c5e727e56e743c9aed4a1
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

CSS の位置指定を使うと、要素をブラウザー内でどのように配置するか指定できます。 It has a `position` property you can set to `static`, `absolute`, `relative`, `sticky` or `fixed`.

要素の `position` プロパティを設定すると、`top`、`right`、`left`、`bottom` のプロパティにピクセルやパーセントの値を指定することで、要素の配置を動かせるようになります。

ただし、すべての要素の位置指定のデフォルト値は `static` です。 この値を設定しても、`top`、`right`、`left`、`bottom` で動かせるようにはなりません。

Give `.cat-head` a `position` property of `static`, then set the `top` and `left` properties to `100px` each.


# --hints--

Your `.cat-head` selector should have a `position` property set to `static`. 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

`.cat-head` セレクターの `top` プロパティを `100px` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

`.cat-head` セレクターの `left` プロパティを `100px` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
