---
id: 6476fb7cbaafa36d65e9cf35
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`fixed` is a `position` property value that lets you make an element fixed to the page no matter where the user scrolls to on the page.

`fixed` 位置指定の効果を確認するには、いくつかマークアップを追加する必要があります。 HTML に、`box` クラスを持つ `div` 要素を追加してください。

# --hints--

`div` 要素が 1 つ必要です。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Your `div` element should have the class `box`.

```js
assert.equal(document.querySelectorAll('div')[1]?.getAttribute('class'), 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
