---
id: 587d8247367417b2b2512c37
title: Potenziell gefährliche Informationen mit helmet.hidePoweredBy() verbergen
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

Zur Erinnerung: Dieses Projekt basiert auf dem folgenden Starter-Projekt, das von <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Hacker können bekannte Sicherheitslücken in Express/Node ausnutzen, wenn sie sehen, dass Ihre Website mit Express betrieben wird. `X-Powered-By: Express` wird standardmäßig bei jeder Anfrage, die von Express kommt, gesendet. Verwenden Sie die `helmet.hidePoweredBy()`-Middleware, um die X-Powered-By-Kopfzeile zu entfernen.

# --hints--

Die helmet.hidePoweredBy()-Middleware sollte korrekt eingebunden sein

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```

