---
id: 587d8248367417b2b2512c3d
title: DNS-Prefetching mit helmet.dnsPrefetchControl() deaktivieren
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Um die Leistung zu verbessern, holen die meisten Browser DNS-Einträge für die Links auf einer Seite vorab ab. So ist die Ziel-IP bereits bekannt, wenn der Nutzer auf einen Link klickt. Dies kann zu einer Überlastung des DNS-Dienstes führen (wenn Sie eine große Website besitzen, die von Millionen Menschen besucht wird…), zu Datenschutzproblemen (ein Lauscher könnte ableiten, dass Sie sich auf einer bestimmten Seite befinden) oder zu verfälschten Seitenstatistiken (einige Links könnten als besucht erscheinen, obwohl sie es nicht sind). Wenn Sie hohe Sicherheitsanforderungen haben, können Sie das DNS-Prefetching deaktivieren – allerdings auf Kosten einer Leistungseinbuße.

# --instructions--

Verwenden Sie die `helmet.dnsPrefetchControl()` Methode auf Ihrem Server.

# --hints--

Die Middleware helmet.dnsPrefetchControl() sollte korrekt eingebunden sein

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```

