---
id: 587d8248367417b2b2512c3b
title: "Verhindern, dass IE nicht vertrauenswürdiges HTML mit helmet.ieNoOpen() öffnet"
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

Zur Erinnerung: Dieses Projekt basiert auf dem folgenden Starter-Projekt, das von <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Einige Webanwendungen stellen nicht vertrauenswürdiges HTML zum Download bereit. Einige Versionen von Internet Explorer öffnen diese HTML-Dateien standardmäßig im Kontext Ihrer Website. Das bedeutet, dass eine nicht vertrauenswürdige HTML-Seite im Kontext Ihrer Seiten schädliche Aktionen ausführen könnte. Diese Middleware setzt den Header X-Download-Options auf noopen. Dadurch wird verhindert, dass IE-Nutzer Downloads im Kontext der vertrauenswürdigen Website ausführen.

# --instructions--

Verwenden Sie die `helmet.ieNoOpen()`-Methode auf Ihrem Server.

# --hints--

Die Middleware helmet.ieNoOpen() sollte korrekt eingebunden sein

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```

