---
id: 587d8249367417b2b2512c40
title: Helmet mit der „parent“ helmet() Middleware konfigurieren
challengeType: 2
forumTopicId: 301575
dashedName: configure-helmet-using-the-parent-helmet-middleware
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

`app.use(helmet())` wird automatisch alle oben eingeführten Middleware einschließen, außer `noCache()` und `contentSecurityPolicy()`, aber diese können bei Bedarf aktiviert werden. Sie können auch jede andere Middleware einzeln deaktivieren oder konfigurieren, indem Sie ein Konfigurationsobjekt verwenden.

**Beispiel:**

```js
app.use(helmet({
  frameguard: {         // configure
    action: 'deny'
  },
  contentSecurityPolicy: {    // enable and configure
    directives: {
      defaultSrc: ["'self'"],
      styleSrc: ['style.com'],
    }
  },
  dnsPrefetchControl: false     // disable
}))
```

Wir haben jede Middleware einzeln eingeführt, um den Lernzweck zu unterstützen und das Testen zu erleichtern. Die Verwendung der „parent“ `helmet()` Middleware ist in einem echten Projekt einfach zu implementieren.

# --hints--

keine Tests – es handelt sich um eine beschreibende Aufgabe

```js
assert(true);
```

