---
id: 587d8249367417b2b2512c3e
title: Client-seitiges Caching mit helmet.noCache() deaktivieren
challengeType: 2
forumTopicId: 301576
dashedName: disable-client-side-caching-with-helmet-nocache
---

# --description--

Zur Erinnerung: Dieses Projekt basiert auf dem folgenden Starter-Projekt, das von <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Wenn Sie ein Update für Ihre Website veröffentlichen und möchten, dass die Nutzer immer die neuere Version herunterladen, können Sie versuchen, das Caching im Browser des Clients zu deaktivieren. Das kann auch in der Entwicklung nützlich sein. Caching bringt Leistungsverbesserungen mit sich, die Sie dadurch verlieren, verwenden Sie diese Option also nur, wenn es wirklich notwendig ist.

# --instructions--

Verwenden Sie die `helmet.noCache()` Methode auf Ihrem Server.

# --hints--

Die Middleware helmet.noCache() sollte korrekt eingebunden sein

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nocache');
assert.equal(
  data.headers['cache-control'],
  'no-store, no-cache, must-revalidate, proxy-revalidate'
);
```

