---
id: 587d8247367417b2b2512c38
title: Risiko von Clickjacking mit helmet.frameguard() abschwächen
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

Zur Erinnerung: Dieses Projekt basiert auf dem folgenden Starter-Projekt, das von <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Ihre Seite könnte ohne Ihre Zustimmung in einem `<frame>` oder `<iframe>` eingebettet werden. Dies kann unter anderem zu Clickjacking-Angriffen führen. Clickjacking ist eine Technik, bei der ein Benutzer dazu gebracht wird, mit einer Seite zu interagieren, die anders ist, als er denkt. Dies wird erreicht, indem Ihre Seite in einem bösartigen Kontext ausgeführt wird, zum Beispiel durch Iframing. In diesem Kontext kann ein Hacker eine versteckte Ebene über Ihre Seite legen. Versteckte Schaltflächen können verwendet werden, um schädliche Skripte auszuführen. Diese Middleware setzt den X-Frame-Options-Kopfzeilenwert. Sie beschränkt, wer Ihre Seite in einem Frame einbetten darf. Es gibt drei Modi: DENY, SAMEORIGIN und ALLOW-FROM.

Wir benötigen nicht, dass unsere App eingerahmt wird.

# --instructions--

Verwenden Sie `helmet.frameguard()` und übergeben Sie das Konfigurationsobjekt `{action: 'deny'}`.

# --hints--

Die Middleware helmet.frameguard() sollte korrekt eingebunden sein

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

Der Wert 'action' von helmet.frameguard() sollte auf 'DENY' gesetzt sein

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```

