---
id: 587d8248367417b2b2512c3c
title: "Browser anweisen, Ihre Website nur über HTTPS mit helmet.hsts() aufzurufen"
challengeType: 2
forumTopicId: 301573
dashedName: ask-browsers-to-access-your-site-via-https-only-with-helmet-hsts
---

# --description--

Zur Erinnerung: Dieses Projekt basiert auf dem folgenden Starter-Projekt, das von <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

HTTP Strict Transport Security (HSTS) ist eine Web-Sicherheitsrichtlinie, die Websites vor Protokoll-Downgrade-Angriffen und Cookie-Hijacking schützt. Wenn Ihre Website über HTTPS aufgerufen werden kann, können Sie die Browser der Nutzer anweisen, unsicheres HTTP zu vermeiden. Indem Sie die Kopfzeile Strict-Transport-Security setzen, teilen Sie den Browsern mit, HTTPS für zukünftige Anfragen innerhalb eines bestimmten Zeitraums zu verwenden. Dies gilt für alle Anfragen, die nach der ersten Anfrage erfolgen.

# --instructions--

Konfigurieren Sie `helmet.hsts()`, damit HTTPS für die nächsten 90 Tage verwendet wird. Übergeben Sie das Konfigurationsobjekt `{maxAge: timeInSeconds, force: true}`. Sie können eine Variable `ninetyDaysInSeconds = 90*24*60*60;` erstellen, die Sie für `timeInSeconds` verwenden.

Hinweis: Die Konfiguration von HTTPS auf einer eigenen Website erfordert den Erwerb einer Domain und eines SSL/TLS-Zertifikats.

# --hints--

Die Middleware helmet.hsts() sollte korrekt eingebunden sein

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hsts');
assert.property(data.headers, 'strict-transport-security');
```

maxAge sollte 7776000 s (90 Tage) entsprechen

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.match(
  data.headers['strict-transport-security'],
  /^max-age=7776000;?/
);
```

