---
id: 587d8247367417b2b2512c36
title: Helmet installieren und einbinden
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

Bei diesen Aufgaben schreiben Sie Ihren Code mit einer der folgenden Methoden:

- Klonen Sie <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">dieses GitHub-Repository</a> und bearbeiten Sie die Aufgaben lokal.
- Verwenden Sie einen Website-Builder Ihrer Wahl, um das Projekt abzuschließen. Stellen Sie sicher, dass Sie alle Dateien aus unserem GitHub-Repository einbinden.

Helmet hilft Ihnen, Ihre Express-Apps abzusichern, indem es verschiedene HTTP-Kopfzeilen setzt.

# --instructions--

Ihr gesamter Code für diese Lektionen kommt in die `myApp.js`-Datei zwischen den Codezeilen, die wir für Sie vorbereitet haben. Ändern oder löschen Sie nicht den Code, den wir hinzugefügt haben.

Die Helmet-Version `3.21.3` ist bereits installiert, binden Sie sie also als `helmet` in `myApp.js` ein.

# --hints--

`helmet` Version `3.21.3` sollte in `package.json` sein

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```

