---
id: 587d8248367417b2b2512c3a
title: Vermeiden Sie das automatische Erkennen des MIME-Typs der Antwort mit helmet.noSniff()
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

Zur Erinnerung: Dieses Projekt basiert auf dem folgenden Starter-Projekt, das von <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Browser können Content- oder MIME-Sniffing verwenden, um die `Content-Type`-Kopfzeile einer Antwort zu überschreiben und die Daten anhand eines impliziten Content-Typs zu erraten und zu verarbeiten. Obwohl dies in manchen Szenarien praktisch sein kann, kann es auch zu gefährlichen Angriffen führen. Diese Middleware setzt die `X-Content-Type-Options`-Kopfzeile auf `nosniff` und weist den Browser an, die angegebene `Content-Type` nicht zu umgehen.

# --instructions--

Verwenden Sie die `helmet.noSniff()`-Methode auf Ihrem Server.

# --hints--

Die helmet.noSniff()-Middleware sollte korrekt eingebunden sein

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```

