---
id: 58a25bcff9fc0f352b528e7d
title: Passwörter asynchron hashen und vergleichen
challengeType: 2
forumTopicId: 301578
dashedName: hash-and-compare-passwords-asynchronously
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Da das Hashen darauf ausgelegt ist, rechenintensiv zu sein, wird empfohlen, dies asynchron auf Ihrem Server zu tun, um eingehende Verbindungen während des Hashens nicht zu blockieren. Alles, was Sie tun müssen, um ein Passwort asynchron zu hashen, ist, `call`

```js
bcrypt.hash(myPlaintextPassword, saltRounds, (err, hash) => {
  /*Store hash in your db*/
});
```

# --instructions--

Fügen Sie diese Hash-Funktion zu Ihrem Server hinzu (wir haben die in der Funktion verwendeten Variablen bereits für Sie definiert) und geben Sie sie in der Konsole aus, damit Sie sie sehen können! An diesem Punkt würden Sie normalerweise den Hash in Ihrer Datenbank speichern.

Wenn Sie nun herausfinden müssen, ob eine neue Eingabe dieselben Daten wie der Hash enthält, verwenden Sie einfach die Vergleichsfunktion.

```js
bcrypt.compare(myPlaintextPassword, hash, (err, res) => {
  /*res == true or false*/
});
```

Fügen Sie dies in Ihre bestehende Hash-Funktion ein (da Sie auf den Abschluss des Hashings warten müssen, bevor Sie die Vergleichsfunktion aufrufen), nachdem Sie den fertigen Hash protokolliert haben, und geben Sie 'res' innerhalb des Vergleichs in der Konsole aus. Sie sollten in der Konsole zuerst den Hash sehen und dann wird 'true' ausgegeben! Wenn Sie in der Vergleichsfunktion 'myPlaintextPassword' in 'someOtherPlaintextPassword' ändern, sollte 'false' ausgegeben werden.

```js
bcrypt.hash('passw0rd!', 13, (err, hash) => {
  console.log(hash);
  //$2a$12$Y.PHPE15wR25qrrtgGkiYe2sXo98cjuMCG1YwSI5rJW1DSJp0gEYS
  bcrypt.compare('passw0rd!', hash, (err, res) => {
    console.log(res); //true
  });
});

```

Senden Sie Ihre Seite ab, wenn Sie denken, dass alles richtig ist.

# --hints--

Der asynchrone Hash sollte erzeugt und korrekt verglichen werden.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash.*myPlaintextPassword( |),( |)saltRounds( |),( |).*err( |),( |)hash[^]*END_ASYNC/gi,
  'You should call bcrypt.hash on myPlaintextPassword and saltRounds and handle err and hash as a result in the callback'
);
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash[^]*bcrypt.compare.*myPlaintextPassword( |),( |)hash( |),( |).*err( |),( |)res[^]*}[^]*}[^]*END_ASYNC/gi,
  'Nested within the hash function should be the compare function comparing myPlaintextPassword to hash'
);
```
