---
id: 58a25bcff9fc0f352b528e7e
title: Passwörter synchron hashen und vergleichen
challengeType: 2
forumTopicId: 301579
dashedName: hash-and-compare-passwords-synchronously
---

# --description--

Zur Erinnerung: Dieses Projekt baut auf dem folgenden Starter-Projekt auf, das von <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> geklont wurde.

Synchrones Hashen ist genauso einfach durchzuführen, kann aber zu Verzögerungen führen, wenn es serverseitig mit hohen Kosten oder sehr häufigem Hashen verwendet wird. Das Hashen mit dieser Methode ist so einfach wie das Aufrufen von

```js
var hash = bcrypt.hashSync(myPlaintextPassword, saltRounds);
```

Fügen Sie diese Methode des Hashens in Ihren Code ein und geben Sie das Ergebnis in der Konsole aus. Die verwendeten Variablen sind bereits im Server definiert, sodass Sie diese nicht anpassen müssen. Sie werden feststellen, dass das Ergebnis in der Konsole unterschiedlich ist, obwohl Sie dasselbe Passwort wie in der asynchronen Funktion hashen – das liegt daran, dass der Salt jedes Mal zufällig generiert wird, was an den ersten 22 Zeichen im dritten String des Hashes zu erkennen ist. Um nun eine Passworteingabe mit dem neuen synchronen Hash zu vergleichen, verwenden Sie die Methode compareSync:

```js
var result = bcrypt.compareSync(myPlaintextPassword, hash);
```

Das Ergebnis ist ein Boolean true oder false.

# --instructions--

Fügen Sie die Funktion ein und geben Sie das Ergebnis in der Konsole aus, um zu sehen, dass es funktioniert.

Senden Sie Ihre Seite ab, wenn Sie glauben, dass alles richtig ist.

# --hints--

Der synchrone Hash sollte erzeugt und korrekt verglichen werden.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_SYNC[^]*hash.*=.*bcrypt.hashSync.*myPlaintextPassword( |),( |)saltRounds[^]*END_SYNC/gi,
  'You should call bcrypt.hashSync on myPlaintextPassword with saltRounds'
);
assert.match(
  data,
  /START_SYNC[^]*result.*=.*bcrypt.compareSync.*myPlaintextPassword( |),( |)hash[^]*END_SYNC/gi,
  'You should call bcrypt.compareSync on myPlaintextPassword with the hash generated in the last line'
);
```

