---
id: 67fcde1677d474ebe61b4f2e
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

在输入元素中，添加 `value` 属性并将其设置为 `query` 状态变量。 这将允许输入字段显示 `query` 的当前值。 另外，请记住，你需要使用花括号（`{}`）在 JSX 中嵌入 JavaScript 表达式，如 `query`。

然后，向输入元素添加 `onChange` 事件处理程序。 向其传递一个以 `e` 作为参数的匿名函数，该参数是事件对象。 在匿名函数内部，调用 `setQuery()` 并传递 `e.target.value`。 每当用户在输入字段中输入内容时，这将使用输入字段的当前值更新 `query` 状态变量。

# --hints--

你应该向输入元素添加 `value` 属性并将其设置为 `query` 状态变量。

```js
assert.match(code, /<input[^>]*value={\s*query\s*}[^>]*\/?>/);
```

你的 `onChange` 事件处理程序应该是一个以 `e` 作为参数的函数。

```js
const input = document.querySelector('input');
const key = Object.keys(input).find(key => key.startsWith("__reactProps"));
const reactInput = input[key];
assert.match(reactInput.onChange.toString(), /onChange\s*\(\s*e\s*\)/);
```

你的 `onChange` 事件处理程序应该使用 `e.target.value` 调用 `setQuery()`。

```js
const input = document.querySelector('input');
const key = Object.keys(input).find(key => key.startsWith("__reactProps"));
const reactInput = input[key];
assert.match(reactInput.onChange.toString(), /setQuery\s*\(\s*e\.target\.value\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Shopping List</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main id="app-container"></main>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ShoppingList } from "./index.jsx";
      const appContainer = document.getElementById("app-container");
      const root = ReactDOM.createRoot(appContainer);
      root.render(<ShoppingList />);
    </script>
  </body>
</html>

```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  box-sizing: border-box;
}

body {
  font-family: Arial, sans-serif;
  line-height: 1.5;
  color: var(--dark-grey);
  background-color: var(--dark-grey);
}

main,
.container {
  display: flex;
  flex-direction: column;
  align-items: center;
}

.container {
  background-color: var(--light-grey);
  width: 90%;
  margin: 20px;
  padding: 10px;
}

h1 {
  color: var(--dark-grey);
}

form {
  text-align: center;
}

button {
  cursor: pointer;
}

button {
  cursor: pointer;
  width: 100px;
  margin: 3px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

input {
  color: var(--dark-grey);
  margin-left: 5px;
  padding: 3px;
}

li {
  text-align: left;
  margin: 10px 0;
}

.selected-item {
  font-weight: bold;
}

@media (min-width: 425px) {
  .container {
    width: 400px;
  }
}

```

```jsx
const { useState } = React;

const items = [
  "Apples",
  "Bananas",
  "Strawberries",
  "Blueberries",
  "Mangoes",
  "Pineapple",
  "Lettuce",
  "Broccoli",
  "Paper Towels",
  "Dish Soap",
];

export const ShoppingList = () => {
  const [query, setQuery] = useState("");

  return (
    <div className="container">
      <h1>Shopping List</h1>
      <form>
        <label htmlFor="search">Search for an item:</label>
        <input
          id="search"
          type="search"
          placeholder="Search..."
          aria-describedby="search-description"
          --fcc-editable-region--

          --fcc-editable-region--
        /> 
        <p id="search-description">Type to filter the list below:</p>
      </form>
    </div>
  );
};

```
