---
id: 66b102ee0bdbad34a9f85ba0
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utaendelea kujifunza kuhusu mfuatano wa herufi kwa kuunda Chatbot ya Mwalimu.

Kuanza, ongeza tamko la `console`, lenye ujumbe wa `"Hi there!"`.

# --hints--

Unapaswa kuwa na tamko la `console`.

```js
assert.match(code, /console\.log(.*)/);
```

Tamko lako la konsoli linapaswa kutoa ujumbe wa `"Hi there!"`.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
