---
id: 66b59b12a745e10011158f7b
title: Hatua ya 4
challengeType: 1
dashedName: step-4
---

# --description--

Hatua inayofuata ni kuunda vigezo vichache zaidi ambavyo vitatumika katika ujumbe za bot zijazo.

Unda kigezo kinachoitwa `subject` na ukaweke thamani ya mfuatano wa herufi `"JavaScript"`.

Kisha unda kigezo kinachoitwa `topic` na ukaweke thamani ya mfuatano wa herufi `"strings"`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `subject`.

```js
assert.isNotNull(subject);
```

Unapaswa kuweka thamani ya mfuatano wa herufi katika kigezo chako cha `subject`.

```js
assert.isString(subject);
```

Unapaswa kuweka mfuatano wa herufi `"JavaScript"` katika kigezo chako cha `subject`.

```js
assert.equal(subject, "JavaScript");
```

Unapaswa kuwa na kigezo kinachoitwa `topic`.

```js
assert.isNotNull(topic);
```

Unapaswa kuweka thamani ya mfuatano wa herufi katika kigezo chako cha `topic`.

```js
assert.isString(topic);
```

Unapaswa kuweka mfuatano wa herufi `"strings"` katika kigezo chako cha `topic`.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
