---
id: 66b5928b86e907fdfbf34e56
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Sasa ni wakati wa kuweka jina la roboti.

Unda kigezo kinachoitwa `botName` na weka thamani yake kuwa mfuatano wa herufi `"teacherBot"`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `botName`.

```js
assert.isNotNull(botName);
```

Kigezo chako cha `botName` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(botName);
```

Kigezo chako cha `botName` kinapaswa kuwa na thamani ya `"teacherBot"`.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
