---
id: 66b59829dba144ff1351220f
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Sasa ni wakati wa kuunda salamu ukitumia variable ya `botName`.

Katika mafundisho ya awali, ulijifunza jinsi ya kuunganisha mfuatano wa herufi kwa kutumia herufi za kiolezo kama ifuatavyo:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

Anza kwa kuunda variable iitwayo `greeting`.

Kisha, ukitumia sintaksia ya herufi za kiolezo, weka thamani ya mfuatano wa herufi unaosema `My name is`, ukifuatiwa na variable ya `botName`, na kumalizika na nukta (`.`).

Mwishowe, andika variable ya `greeting` kwenye konsoli.

# --hints--

Unapaswa kuwa na variable iitwayo `greeting`.

```js
assert.isNotNull(greeting);
```

Unapaswa kuweka thamani ya mfuatano wa herufi kwenye variable yako ya `greeting`.

```js
assert.isString(greeting);
```

Unapaswa kutumia herufi za kiolezo kuunganisha mfuatano wa herufi `My name is` na variable ya `botName` ikifuatiwa na nukta (`.`).

```js
assert.equal(greeting, "My name is teacherBot.");
```

Unapaswa kuwa na tamko la `console` linaloandika variable ya `greeting` kwenye konsoli.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
