---
id: 66b59be6ab830800c4df9146
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Sasa, ni wakati wa kutumia vigezo ulivyoviunda katika hatua iliyopita.

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na herufi za kiolezo kama hizi:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

Anza kwa kuunda kigezo cha `sentence`.

Ukitegemea sintaksia ya herufi za kiolezo, weka thamani ya mfuatano wa herufi `Today, you will learn about [topic variable goes here] in [subject variable goes here].` katika kigezo cha `sentence`.

Utabadilisha viwakilishi vya `[topic variable goes here]` na `[subject variable goes here]` kwa vigezo vya `topic` na `subject` pamoja na sintaksia ya `${}`.

Mwishowe, andika kigezo cha `sentence` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo cha `sentence`.

```js
assert.isNotNull(sentence);
```

Kigezo chako cha `sentence` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(sentence);
```

Unapaswa kutumia sintaksia ya herufi za kiolezo kuweka thamani ya mfuatano wa herufi `Today, you will learn about [topic variable goes here] in [subject variable goes here].` katika kigezo cha `sentence`. Rejelea mfano ikiwa unahitaji msaada.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

Unapaswa kuandika kigezo cha `sentence` kwenye konsoli.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
