---
id: 66b6fdb76441c738719039fa
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kupata herufi ya mwisho katika mfuatano wa herufi kama ifuatavyo:

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` daima itakupa nambari ya kielezo cha mwisho kwa mfuatano wa herufi.

Tengeneza kigezo kipya kinachoitwa `lastCharacter` na ukaweke thamani ya herufi ya mwisho katika kigezo cha `subject`.

Kisha, andika thamani ya kigezo cha `lastCharacter` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `lastCharacter`.

```js
assert.isNotNull(lastCharacter);
```

Unapaswa kuweka thamani ya herufi ya mwisho katika kigezo cha `subject` kwenye kigezo cha `lastCharacter`. Rejelea mfano ikiwa unahitaji msaada.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

Unapaswa kuandika thamani ya kigezo cha `lastCharacter` kwenye `console`.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
