---
id: 66b6e62423e8a031d6c1f03d
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Sasa ni wakati wa kupata urefu wa mfuatano wa herufi `topic`.

Unaweza kutumia herufi za kiolezo ndani ya tamko za `console` kama hii:

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

Anza kwa kutoa ujumbe `Here is an example of using the length property on the word [topic].` kwenye konsoli.

Kumbuka kubadilisha `[topic]` na kigezo `topic`, na tumia sintaksia sahihi ya herufi za kiolezo kama ulivyofanya katika hatua zilizopita.

Kisha, ongeza tamko la pili la `console.log` linalotoa urefu wa mfuatano wa herufi `topic` kwenye konsoli.

# --hints--

Unapaswa kutumia sintaksia ya herufi za kiolezo kutoa ujumbe `Here is an example of using the length property on the word [topic].` kwenye `console`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

Unapaswa kuwa na tamko la pili la `console` kutoa urefu wa kigezo `topic` kwenye `console`.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
