---
id: 66b6d482bbb9e12f2e5ee1ae
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Kwa sehemu hii inayofuata ya mradi, bot itafundisha jinsi ya kufanya kazi na sifa ya mfuatano wa herufi `length`.

Anza kwa kuunda kigezo kipya kinachoitwa `strLengthIntro`.

Kisha ukitumia sintaksia ya template literal, weka thamani ya mfuatano wa herufi `Here is an example of using the length property on the word [subject].` kwa kigezo `strLengthIntro`.

Badilisha `[subject]` na kigezo `subject` kama ulivyofanya awali.

Mwishowe, andika kigezo `strLengthIntro` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `strLengthIntro`.

```js
assert.isNotNull(strLengthIntro);
```

Kigezo chako cha `strLengthIntro` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(strLengthIntro);
```

Unapaswa kutumia sintaksia ya template literal kuweka thamani ya mfuatano wa herufi `Here is an example of using the length property on the word [subject goes here].` kwa kigezo `strLengthIntro`.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

Unapaswa kuandika kigezo `strLengthIntro` kwenye konsoli.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
