---
id: 66b6ffb42b88e33943788abf
title: Hatua 14
challengeType: 1
dashedName: step-14
---

# --description--

Kwa sehemu ya mwisho ya warsha, utakagua jinsi ya kupata kielezo cha sehemu ndogo ya mfuatano wa herufi ndani ya mfuatano wa herufi. Kumbuka kuwa sehemu ndogo ya mfuatano wa herufi ni sehemu ya mfuatano mzima.

Anza kwa kuunda kigezo kinachoitwa `learningIsFunSentence` na ukaweke thamani ya mfuatano wa herufi `"Learning is fun."`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `learningIsFunSentence`.

```js
assert.isNotNull(learningIsFunSentence);
```

`learningIsFunSentence` yako inapaswa kuwa na thamani ya mfuatano wa herufi.

```js
assert.isString(learningIsFunSentence);
```

Unapaswa kuweka thamani ya `"Learning is fun."` kwenye kigezo `learningIsFunSentence`.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
