---
id: 66b6efddeca35833cd6f0b03
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kupata herufi katika mfuatano wa herufi kama ifuatavyo:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

Kumbuka kwamba nambari za kielezo huanza kwa `0`, hivyo herufi ya kwanza katika mfuatano wa herufi itakuwa daima kwenye kielezo `0`.

Anza kwa kuongeza tamko lingine la `console`. 

Ndani ya tamko la `console`, toa herufi ya kwanza ya mabadiliko ya `subject` kwa kutumia uandishi wa mabano na nambari sahihi ya kielezo.

# --hints--

Tamko lako la `console` linapaswa kutoa herufi ya kwanza ya mabadiliko ya `subject` kwa kutumia uandishi wa mabano na kielezo `0`.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
