---
id: 66b6e39031393e30f2c48d0f
title: Hatua ya 7
challengeType: 1
dashedName: step-7
---

# --description--

Ili kupata urefu wa mfuatano wa herufi, unaweza kutumia sifa ya <dfn>length</dfn>. Sifa hii hurudisha idadi ya herufi katika mfuatano wa herufi.

Hapa kuna mfano:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

Anza kwa kutumia sifa ya `length` kupata urefu wa mfuatano wa herufi wa `subject` na andika thamani hiyo kwenye konsoli.

# --hints--


Unapaswa kutumia sifa ya `length` kupata urefu wa mfuatano wa herufi wa `subject`.

```js
assert.match(code, /subject\.length/);
```

Unapaswa kuandika urefu wa mfuatano wa herufi wa `subject` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
