---
id: 66b6f586767a1534f3097353
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Sasa ni wakati wa kupata herufi ya pili ya mabadiliko ya `subject`.

Anza kwa kuongeza tamko la `console` linalotolewa ujumbe `Here is an example of accessing the second letter in the word [subject].` Kumbuka kubadilisha `[subject]` na thamani halisi ya mabadiliko ya `subject` na tumia sintaksia sahihi ya template literal.

Kisha ongeza tamko lingine la `console` linalotolewa herufi ya pili ya mabadiliko ya `subject` kwa kutumia uandishi wa mabano na nambari sahihi ya kielezo.

# --hints--

Tamko lako la `console` linapaswa kutoa ujumbe `Here is an example of accessing the second letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

Unapaswa kuwa na tamko lingine la `console` linalotolewa herufi ya pili ya mabadiliko ya `subject` kwa kutumia uandishi wa mabano na kielezo cha `1`.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
