---
id: 5895f700f9fc0f352b528e63
title: إعداد محرك قالب
challengeType: 2
forumTopicId: 301564
dashedName: set-up-a-template-engine
---

# --description--

ستتطلب منك هذه التحديات كتابة شفرتك باستخدام إحدى الطرق التالية:

- استنسخ <a href="https://github.com/freeCodeCamp/boilerplate-advancednode/" target="_blank" rel="noopener noreferrer nofollow">مستودع GitHub هذا</a> وأكمل هذه التحديات محليًا.
- استخدم منشئ مواقع من اختيارك لإكمال المشروع. تأكد من دمج جميع الملفات من مستودع GitHub الخاص بنا.

يتيح لك محرك القوالب استخدام ملفات قوالب ثابتة (مثل تلك المكتوبة بـ *Pug*) في تطبيقك. عند وقت التشغيل، يستبدل محرك القوالب المتغيرات في ملف القالب بقيم فعلية يمكن أن يوفرها الخادم الخاص بك. ثم يحول القالب إلى ملف HTML ثابت يُرسل إلى العميل. هذه الطريقة تسهل تصميم صفحة HTML وتسمح بعرض المتغيرات على الصفحة دون الحاجة إلى إجراء استدعاء API من العميل.

تم تثبيت `pug@~3.0.0` بالفعل، وهو مدرج كتبعٍ في ملف `package.json` الخاص بك.

يحتاج Express إلى معرفة محرك القوالب الذي تستخدمه. استخدم طريقة `set` لتعيين `pug` كقيمة خاصية `view engine`:

```javascript
app.set('view engine', 'pug');
```

بعد ذلك، أضف طريقة `set` أخرى لتعيين خاصية `views` في `app` للإشارة إلى مجلد `./views/pug` . هذا يُخبر Express بعرض جميع القوالب نسبةً إلى هذا المجلد.

أخيرًا، استخدم `res.render()` في مسار الصفحة الرئيسية، مع تمرير `index` كالمعلمة الأولى. هذا سيعرض قالب `pug`.

إذا سار كل شيء كما هو مخطط، فلن تكون الصفحة الرئيسية لتطبيقك فارغة بعد الآن. بدلاً من ذلك، ستعرض رسالة تشير إلى أنك عرضت قالب Pug بنجاح!

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-a-template-engine-1" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن يكون Pug تبعية.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'pug',
    'Your project should list "pug" as a dependency'
  );
```

يجب أن يكون محرك العرض هو Pug.

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.['view engine'], "pug");
```

يجب تعيين خاصية `views` في التطبيق إلى `./views/pug`.

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.views, "./views/pug");
```

استخدم طريقة ExpressJS الصحيحة لعرض صفحة الفهرس من الاستجابة.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /FCC Advanced Node and Express/gi,
        'You successfully rendered the Pug template!'
      );
```

يجب أن يعمل Pug.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /pug-success-message/gi,
        'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
      );
```

