---
id: 5895f70df9fc0f352b528e6a
title: إنشاء برمجيات وسيطة جديدة
challengeType: 2
forumTopicId: 301551
dashedName: create-new-middleware
---

# --description--

كما هو الحال، يمكن لأي مستخدم الذهاب إلى `/profile` سواء كان مُوثَّقًا أم لا عن طريق كتابة عنوان URL. تريد منع ذلك بفحص ما إذا كان المستخدم مُوثَّقًا أولًا قبل عرض صفحة الملف الشخصي. هذا مثال مثالي على متى تنشئ برمجيات وسيطة.

التحدي هنا هو إنشاء دالة البرمجيات الوسيطة `ensureAuthenticated(req, res, next)`، التي ستتحقق مما إذا كان المستخدم مُوثَّقًا عن طريق استدعاء طريقة `isAuthenticated` الخاصة بـ Passport على الـ `request` والتي تتحقق مما إذا كان `req.user` معرفًا. إذا كان كذلك، يجب استدعاء `next()`. وإلا، يمكنك الرد على الطلب بإعادة توجيه إلى صفحتك الرئيسية لتسجيل الدخول.

تنفيذ هذه البرمجيات الوسيطة هو:

```javascript
function ensureAuthenticated(req, res, next) {
  if (req.isAuthenticated()) {
    return next();
  }
  res.redirect('/');
};
```

أنشئ دالة البرمجيات الوسيطة أعلاه، ثم مرر `ensureAuthenticated` كبرمجيات وسيطة للطلبات الخاصة بصفحة الملف الشخصي قبل المعلمة الخاصة بطلب GET:

```javascript
app
 .route('/profile')
 .get(ensureAuthenticated, (req,res) => {
    res.render('profile');
 });
```

أرسل صفحتك عندما تعتقد أنك أنجزتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#create-new-middleware-8" target="_blank" rel="noopener noreferrer nofollow">الاطلاع على المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب تنفيذ برمجيات الوسيطة `ensureAuthenticated` وربطها بمسار `/profile`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /ensureAuthenticated[^]*req.isAuthenticated/,
    'Your ensureAuthenticated middleware should be defined and utilize the req.isAuthenticated function'
  );
  assert.match(
    data,
    /profile[^]*get[^]*ensureAuthenticated/,
    'Your ensureAuthenticated middleware should be attached to the /profile route'
  );
```

يجب أن يعيد طلب GET غير المُوثَّق إلى `/profile` إعادة توجيه صحيحة إلى `/`.

```js
  const url = new URL("/profile", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/,
    'An attempt to go to the profile at this point should redirect to the homepage since we are not logged in'
  );
```

