---
id: 5895f70cf9fc0f352b528e65
title: إعداد Passport
challengeType: 2
forumTopicId: 301565
dashedName: set-up-passport
---

# --description--

حان الوقت لإعداد *Passport* حتى تبدأ أخيرًا بالسماح للمستخدم بالتسجيل أو تسجيل الدخول إلى حساب. بالإضافة إلى Passport، ستستخدم Express-session لإدارة الجلسات. يحتوي Express-session على العديد من الميزات المتقدمة التي يمكنك استخدامها، لكنك ستستخدم الآن الأساسيات فقط. استخدام هذه البرمجيات الوسيطة يحفظ معرف الجلسة كملف تعريف ارتباط في العميل، ويسمح لنا بالوصول إلى بيانات الجلسة باستخدام هذا المعرف على الخادم. بهذه الطريقة، تحافظ على معلومات الحساب الشخصية خارج ملف تعريف الارتباط الذي يستخدمه العميل لإخبار خادمك بأن العملاء مُوثَّقون، وتحتفظ بـ *المفتاح* للوصول إلى البيانات المخزنة على الخادم.

تم تثبيت `passport@~0.4.1` و `express-session@~1.17.1` بالفعل، وكلاهما مدرج كاعتماديات في ملف `package.json` الخاص بك.

ستحتاج إلى إعداد إعدادات الجلسة وتهيئة Passport. أولًا، أنشئ المتغيرين `session` و `passport` لاستدعاء `express-session` و `passport` على التوالي.

بعد ذلك، أعد تطبيق Express الخاص بك لاستخدام الجلسة بتعريف الخيارات التالية:

```javascript
app.use(session({
  secret: process.env.SESSION_SECRET,
  resave: true,
  saveUninitialized: true,
  cookie: { secure: false }
}));
```

تأكد من إضافة `SESSION_SECRET` إلى ملف `.env` الخاص بك، وأعطه قيمة عشوائية. تُستخدم هذه القيمة لحساب التجزئة التي تُشفّر بها ملف تعريف الارتباط الخاص بك!

بعد إتمام كل ذلك، أخبر تطبيق express الخاص بك **باستخدام** `passport.initialize()` و `passport.session()`.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-passport-3" target="_blank" rel="noopener noreferrer nofollow">الاطلاع على المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن تكون Passport و Express-session من الاعتماديات.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport',
    'Your project should list "passport" as a dependency'
  );
  assert.property(
    packJson.dependencies,
    'express-session',
    'Your project should list "express-session" as a dependency'
  );
```

يجب استدعاء الاعتماديات بشكل صحيح.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport("|')/,
    'You should have required passport'
  );
  assert.match(
    data,
    /require.*("|')express-session("|')/,
    'You should have required express-session'
  );
```

يجب أن يستخدم تطبيق Express الاعتماديات الجديدة.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /passport\.initialize/, 'Your express app should use "passport.initialize()"');
  assert.match(data, /passport\.session/, 'Your express app should use "passport.session()"');
```

يجب إعداد الجلسة والسر الخاص بالجلسة بشكل صحيح.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /secret *:\s*process\.env(\.SESSION_SECRET|\[(?<q>"|')SESSION_SECRET\k<q>\])/,
    'Your express app should have express-session set up with your secret as process.env.SESSION_SECRET'
  );
```

