---
id: 5895f70bf9fc0f352b528e64
title: استخدام قدرات محرك القوالب
challengeType: 2
forumTopicId: 301567
dashedName: use-a-template-engines-powers
---

# --description--

واحدة من أعظم ميزات استخدام محرك القوالب هي القدرة على تمرير المتغيرات من الخادم إلى ملف القالب قبل عرضه كـ HTML.

في ملف Pug الخاص بك، يمكنك استخدام متغير بالإشارة إلى اسم المتغير كـ `#{variable_name}` مضمّنًا مع نص آخر على عنصر، أو باستخدام علامة المساواة على العنصر بدون فراغ مثل `p=variable_name` التي تُعيّن قيمة المتغير إلى نص عنصر p.

يعتمد Pug على استخدام الفراغات وعلامات الجدولة لإظهار العناصر المسكنة وتقليل كمية الشفرة اللازمة لإنشاء موقع جميل.

خذ شفرة Pug التالية كمثال:
 
```pug
head
  script(type='text/javascript').
    if (foo) bar(1 + 5);
body
  if youAreUsingPug
      p You are amazing
    else
      p Get on it!
```

الشكل أعلاه ينتج عنه HTML التالي:

```html
<head>
  <script type="text/javascript">
    if (foo) bar(1 + 5);
  </script>
</head>
<body>
  <p>You are amazing</p>
</body>
```

ملف `index.pug` المضمّن في مشروعك يستخدم المتغيرين `title` و `message`.

مرر تلك المتغيرات من الخادم إلى ملف Pug بإضافة كائن كمعامل ثاني لاستدعاء `res.render` مع المتغيرات وقيمها. أعطِ `title` القيمة `Hello` و `message` القيمة `Please log in`.

يجب أن يبدو كالتالي:

```javascript
res.render('index', { title: 'Hello', message: 'Please log in' });
```

الآن حدّث صفحتك، ويجب أن ترى تلك القيم معروضة في العرض في الموضع الصحيح كما هو محدد في ملف `index.pug` الخاص بك!

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك الاطلاع على <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#use-a-template-engines-power-2" target="_blank" rel="noopener noreferrer nofollow">المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن يعرض Pug المتغيرات بشكل صحيح.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /pug-variable("|')>Please log in/gi,
    'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
  );
```

