---
id: 5895f70ef9fc0f352b528e6b
title: كيفية تجميع ملف تعريفي
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

بما أنك تستطيع التأكد من أن المستخدم الذي يصل إلى `/profile` مُوثَّق، يمكنك استخدام المعلومات الموجودة في `req.user` على صفحتك.

مرر كائنًا يحتوي على الخاصية `username` وقيمة `req.user.username` كالمعلمة الثانية لطريقة `render` في عرض الملف التعريفي.

ثم، اذهب إلى عرض `profile.pug`، وأضف السطر التالي أسفل عنصر `h1` الموجود، وعلى نفس مستوى التراجع:

```pug
h2.center#welcome Welcome, #{username}!
```

هذا ينشئ عنصر `h2` بفئة `center` ومعرف `welcome` يحتوي على النص `Welcome, ` متبوعًا باسم المستخدم.

أيضًا، في `profile.pug`، أضف رابطًا يشير إلى مسار `/logout`، الذي يستضيف منطق إلغاء توثيق المستخدم:

```pug
a(href='/logout') Logout
```

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن تضيف متغير عرض Pug بشكل صحيح إلى `/profile`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );
```

