---
id: 5895f70cf9fc0f352b528e67
title: تنفيذ تسلسل مستخدم Passport
challengeType: 2
forumTopicId: 301556
dashedName: implement-the-serialization-of-a-passport-user
---

# --description--

أنت لا تقوم بتحميل كائن مستخدم فعلي لأن قاعدة البيانات غير مُعدة. اتصل بقاعدة البيانات مرة واحدة عند بدء الخادم، واحتفظ باتصال مستمر طوال دورة حياة التطبيق. للقيام بذلك، أضف سلسلة اتصال قاعدة البيانات الخاصة بك (على سبيل المثال: `mongodb+srv://<username>:<password>@cluster0-jvwxi.mongodb.net/?retryWrites=true&w=majority`) إلى متغير البيئة `MONGO_URI`. يُستخدم هذا في ملف `connection.js`.

*إذا واجهت مشاكل في إعداد قاعدة بيانات مجانية على MongoDB Atlas، اطلع على هذا <a href="https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/" target="_blank" rel="noopener noreferrer nofollow">الدليل التعليمي</a>.*

الآن تريد الاتصال بقاعدة البيانات، ثم تبدأ في الاستماع للطلبات. الهدف من ذلك هو عدم السماح بالطلبات قبل اتصال قاعدة البيانات أو في حال وجود خطأ في قاعدة البيانات. لتحقيق ذلك، ضع تسلسل البيانات ومسارات التطبيق ضمن الشفرة التالية:

```javascript
myDB(async client => {
  const myDataBase = await client.db('database').collection('users');

  // Be sure to change the title
  app.route('/').get((req, res) => {
    // Change the response to render the Pug template
    res.render('index', {
      title: 'Connected to Database',
      message: 'Please login'
    });
  });

  // Serialization and deserialization here...

  // Be sure to add this...
}).catch(e => {
  app.route('/').get((req, res) => {
    res.render('index', { title: e, message: 'Unable to connect to database' });
  });
});
// app.listen out here...
```

تأكد من إزالة التعليق عن شفرة `myDataBase` في `deserializeUser`، وعدل `done(null, null)` لتشمل `doc`.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implement-the-serialization-of-a-passport-user-5" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن يكون اتصال قاعدة البيانات موجودًا.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Connected to Database/gi,
    'You successfully connected to the database!'
  );
```

يجب أن يستخدم إلغاء التسلسل قاعدة البيانات بشكل صحيح الآن، ويجب استدعاء `done(null, null)` مع `doc`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /null,\s*doc/gi,
    'The callback in deserializeUser of (null, null) should be altered to (null, doc)'
  );
```

