---
id: 5895f70cf9fc0f352b528e66
title: تسلسل كائن المستخدم
challengeType: 2
forumTopicId: 301563
dashedName: serialization-of-a-user-object
---

# --description--

التسلسل وإلغاء التسلسل مفهومان مهمان في مجال التوثيق. يعني تسلسل كائن تحويل محتوياته إلى *مفتاح* صغير يمكن بعد ذلك إلغاء تسلسله إلى الكائن الأصلي. هذا يسمح لنا بمعرفة من تواصل مع الخادم دون الحاجة إلى إرسال بيانات التوثيق، مثل اسم المستخدم وكلمة المرور، في كل طلب لصفحة جديدة.

لإعداد هذا بشكل صحيح، تحتاج إلى دالة تسلسل ودالة إلغاء تسلسل. في Passport، يمكن إنشاء هاتين الدالتين باستخدام:

```javascript
passport.serializeUser(cb);
passport.deserializeUser(cb);
```

تُستدعى دالة الاستدعاء الراجع الممررة إلى `serializeUser` بمعاملين: كائن المستخدم الكامل، ودالة استدعاء راجع يستخدمها passport.

تتوقع دالة الاستدعاء الراجع معاملين: خطأ، إن وجد، ومفتاح فريد لتحديد المستخدم يجب إرجاعه في الاستدعاء الراجع. ستستخدم `_id` الخاص بالمستخدم في الكائن. هذا المفتاح مضمون كونه فريدًا لأنه يُولد بواسطة MongoDB.

بنفس الطريقة، تُستدعى `deserializeUser` بمعاملين: المفتاح الفريد، ودالة استدعاء راجع.

تتوقع دالة الاستدعاء الراجع هذه معاملين: خطأ، إن وجد، وكائن المستخدم الكامل. للحصول على كائن المستخدم الكامل، قم بعمل استعلام بحث عن `_id` في Mongo، كما هو موضح أدناه:


```javascript
passport.serializeUser((user, done) => {
  done(null, user._id);
});

passport.deserializeUser((id, done) => {
  myDataBase.findOne({ _id: new ObjectID(id) }, (err, doc) => {
    done(null, null);
  });
});
```

أضف الدالتين أعلاه إلى الخادم الخاص بك. صنف `ObjectID` يأتي من حزمة `mongodb`. تم إضافة `mongodb@~3.6.0` بالفعل كتبعيات. أعلن هذا الصنف باستخدام:

```javascript
const { ObjectID } = require('mongodb');
```

ستُرمي `deserializeUser` خطأً حتى تُعد اتصال قاعدة البيانات. لذا، في الوقت الحالي، أضف تعليقًا على استدعاء `myDatabase.findOne`، واستدعِ فقط `done(null, null)` في دالة الاستدعاء الراجع لـ `deserializeUser`.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#serialization-of-a-user-object-4" target="_blank" rel="noopener noreferrer nofollow">الاطلاع على المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن تسلسل كائن المستخدم بشكل صحيح.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.serializeUser/gi,
    'You should have created your passport.serializeUser function'
  );
  assert.match(
    data,
    /null,\s*user._id/gi,
    'There should be a callback in your serializeUser with (null, user._id)'
  );
```

يجب أن تلغي تسلسل كائن المستخدم بشكل صحيح.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.deserializeUser/gi,
    'You should have created your passport.deserializeUser function'
  );
  assert.match(
    data,
    /null,\s*null/gi,
    'There should be a callback in your deserializeUser with (null, null) for now'
  );
```

يجب أن تكون MongoDB تبعية.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'mongodb',
    'Your project should list "mongodb" as a dependency'
  );
```

يجب استيراد Mongodb بشكل صحيح بما في ذلك ObjectId.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')mongodb\1/gi,
    'You should have required mongodb'
  );
  assert.match(
    data,
    /new ObjectID.*id/gi,
    'Even though the block is commented out, you should use new ObjectID(id) for when we add the database'
  );
```

