---
id: 5895f70df9fc0f352b528e68
title: استراتيجيات التوثيق
challengeType: 2
forumTopicId: 301547
dashedName: authentication-strategies
---

# --description--

الاستراتيجية هي طريقة لتوثيق المستخدم. يمكنك استخدام استراتيجية للسماح للمستخدمين بالتوثيق بناءً على معلومات محفوظة محليًا (إذا قمت بتسجيلهم أولًا) أو من خلال مجموعة متنوعة من المزودين مثل Google أو GitHub. في هذا المشروع، سنستخدم برمجيات وسيطة Passport. توفر Passport مجموعة شاملة من الاستراتيجيات التي تدعم التوثيق باستخدام اسم المستخدم وكلمة المرور، وGitHub، وGoogle، والمزيد.

تمت إضافة `passport-local@~1.0.0` بالفعل كتبعٍ. أضفه إلى الخادم الخاص بك كما يلي:

```javascript
const LocalStrategy = require('passport-local');
```

أخبر passport **باستخدام** كائن `LocalStrategy` مُنشأ مع بعض الإعدادات المحددة. تأكد من أن هذا (وكذلك كل شيء من هذه النقطة فصاعدًا) محاط باتصال قاعدة البيانات لأنه يعتمد عليه!:

```javascript
passport.use(new LocalStrategy((username, password, done) => {
  myDataBase.findOne({ username: username }, (err, user) => {
    console.log(`User ${username} attempted to log in.`);
    if (err) return done(err);
    if (!user) return done(null, false);
    if (password !== user.password) return done(null, false);
    return done(null, user);
  });
}));
```

هذا يحدد العملية التي تُستخدم عند محاولة توثيق شخص محليًا. أولًا، يحاول العثور على مستخدم في قاعدة بياناتك باسم المستخدم المدخل. ثم، يتحقق من تطابق كلمة المرور. أخيرًا، إذا لم تظهر أخطاء قمت بالتحقق منها (مثل كلمة مرور غير صحيحة)، يُعاد كائن `user` ويتم توثيقه.

يتم إعداد العديد من الاستراتيجيات باستخدام إعدادات مختلفة. عمومًا، من السهل إعدادها بناءً على ملف README في مستودع تلك الاستراتيجية. مثال جيد على ذلك هو استراتيجية GitHub حيث لا تحتاج للقلق بشأن اسم المستخدم أو كلمة المرور لأن المستخدم سيُرسل إلى صفحة توثيق GitHub لتوثيقه. طالما أنهم مسجلون دخولهم ووافقوا، تعيد GitHub ملفهم الشخصي لتستخدمه.

في الخطوة التالية، ستُعد كيفية استدعاء استراتيجية التوثيق فعليًا للتحقق من المستخدم بناءً على بيانات النموذج.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#authentication-strategies-6" target="_blank" rel="noopener noreferrer nofollow">الاطلاع على المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن تكون passport-local تبعية.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport-local',
    'Your project should list "passport-local " as a dependency'
  );
```

يجب استدعاء passport-local بشكل صحيح وإعداده.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport-local("|')/,
    'You should have required passport-local'
  );
  assert.match(
    data,
    /new LocalStrategy/,
    'You should have told passport to use a new strategy'
  );
  assert.match(
    data,
    /findOne/,
    'Your new local strategy should use the findOne query to find a username based on the inputs'
  );
```

