---
id: 5895f70df9fc0f352b528e69
title: كيفية استخدام استراتيجيات Passport
challengeType: 2
forumTopicId: 301555
dashedName: how-to-use-passport-strategies
---

# --description--

في ملف `index.pug` المرفق، يوجد نموذج تسجيل دخول. هو مخفي بسبب جافاسكريبت المضمَّن `if showLogin` مع النموذج متراجع تحته.

في `res.render` لتلك الصفحة، أضف متغيرًا جديدًا إلى الكائن، `showLogin: true`. عند تحديث الصفحة، سترى النموذج! هذا النموذج معد لإرسال طلب **POST** على `/login`. إذًا، هنا يجب أن تضبط استقبال طلب POST وتوثيق المستخدم.

لهذا التحدٍّ، يجب إضافة مسار `/login` لاستقبال طلب POST. لتوثيق المستخدم في هذا المسار، تحتاج إلى إضافة برمجيات وسيطة قبل إرسال الاستجابة. يتم ذلك بتمرير معلمة أخرى مع البرمجيات الوسيطة قبل الاستجابة. البرمجيات الوسيطة المستخدمة هي `passport.authenticate('local')`.

يمكن لـ `passport.authenticate` أيضًا أن تأخذ بعض الخيارات كمعلمة مثل `{ failureRedirect: '/' }` وهي مفيدة جدًا، فتأكد من إضافتها أيضًا. أضف استجابة بعد استخدام البرمجيات الوسيطة (والتي تُستدعى فقط إذا نجحت البرمجيات الوسيطة للتوثيق) تعيد توجيه المستخدم إلى `/profile`. أضف هذا المسار أيضًا واجعله يعرض العرض `profile.pug`.

إذا نجح التوثيق، سيتم حفظ كائن المستخدم في `req.user`.

في هذه المرحلة، إذا أدخلت اسم مستخدم وكلمة مرور في النموذج، يجب أن يعيد التوجيه إلى الصفحة الرئيسية `/`، ويجب أن تعرض وحدة التحكم في الخادم الرسالة `'User {USERNAME} attempted to log in.'`، لأننا حاليًا لا نستطيع تسجيل دخول مستخدم غير مسجل.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-use-passport-strategies-7" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب تنفيذ كل الخطوات بشكل صحيح في `server.js`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /showLogin:( |)true/,
    'You should be passing the variable "showLogin" as true to your render function for the homepage'
  );
  assert.match(
    data,
    /failureRedirect:( |)('|")\/('|")/,
    'Your code should include a failureRedirect to the "/" route'
  );
  assert.match(
    data,
    /login[^]*post[^]*local/,
    'You should have a route for login which accepts a POST and passport.authenticates local'
  );
```

يجب أن يعيد طلب POST إلى `/login` التوجيه بشكل صحيح إلى `/`.

```js
  const url = new URL("/login", code);
  const res = await fetch(url, { method: 'POST' });
  const data = await res.text();
  assert.match(
    data,
    /Looks like this page is being rendered from Pug into HTML!/,
    'A login attempt at this point should redirect to the homepage since we do not have any registered users'
  );
```

