---
id: 589690e6f9fc0f352b528e6e
title: نظف مشروعك باستخدام الوحدات
challengeType: 2
forumTopicId: 301549
dashedName: clean-up-your-project-with-modules
---

# --description--

الآن، كل شيء لديك موجود في ملف `server.js`. هذا قد يؤدي إلى شفرة يصعب إدارتها وليست قابلة للتوسع بشكل جيد. أنشئ ملفين جديدين: `routes.js` و `auth.js`

يجب أن يبدأ كلاهما بالشفرة التالية:

```js
module.exports = function (app, myDataBase) {

}
```

الآن، في أعلى ملف الخادم، استورد هذه الملفات هكذا: `const routes = require('./routes.js');` بعد أن تنشئ اتصالًا ناجحًا مع قاعدة البيانات، أنشئ كل واحد منهما هكذا: `routes(app, myDataBase)`

أخيرًا، خذ كل المسارات في الخادم والصقها في ملفاتك الجديدة، واحذفها من ملف الخادم. خذ أيضًا دالة `ensureAuthenticated`، لأنها أنشئت خصيصًا للمسارات. الآن، عليك إضافة التبعيات المستخدمة بشكل صحيح، مثل `const passport = require('passport');`، في الأعلى تمامًا، فوق سطر التصدير في ملف `routes.js`.

استمر في إضافتها حتى لا تبقى أخطاء، ولا يحتوي ملف الخادم على أي مسارات (**باستثناء المسار في كتلة catch**)!

افعل الشيء نفسه في ملف `auth.js` مع كل ما يتعلق بالتوثيق مثل التسلسل وإعداد الاستراتيجية المحلية واحذفها من ملف الخادم. تأكد من إضافة التبعيات واستدعاء `auth(app, myDataBase)` في نفس المكان في الخادم.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#clean-up-your-project-with-modules-2" target="_blank" rel="noopener noreferrer nofollow">الاطلاع على مثال للمشروع المكتمل</a>.

# --hints--

يجب أن تكون الوحدات موجودة.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\s*\(('|")\.\/routes(\.js)?\1\)/gi,
    'You should have required your new files'
  );
  assert.match(
    data,
    /client\s*\.db[^]*routes/gi,
    'Your new modules should be called after your connection to the database'
  );
```

