---
id: 58965611f9fc0f352b528e6c
title: تسجيل خروج المستخدم
challengeType: 2
forumTopicId: 301560
dashedName: logging-a-user-out
---

# --description--

إنشاء منطق تسجيل الخروج سهل. يجب أن يقوم المسار فقط بإلغاء توثيق المستخدم، وإعادة التوجيه إلى الصفحة الرئيسية بدلاً من عرض أي واجهة.

في passport، إلغاء توثيق المستخدم سهل جدًا بمجرد استدعاء `req.logout()` قبل إعادة التوجيه. أضف هذا المسار `/logout` للقيام بذلك:

```js
app.route('/logout')
  .get((req, res) => {
    req.logout();
    res.redirect('/');
});
```

ربما لاحظت أنك لا تتعامل مع الصفحات المفقودة (404). الطريقة الشائعة للتعامل مع هذا في Node هي باستخدام برمجيات وسيطة التالية. أضفها بعد جميع مساراتك الأخرى:

```js
app.use((req, res, next) => {
  res.status(404)
    .type('text')
    .send('Not Found');
});
```

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#logging-a-user-out-10" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب استدعاء `req.logout()` في مسار `/logout` الخاص بك.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /req.logout/gi,
    'You should be calling req.logout() in your /logout route'
  );
```

يجب أن يعيد مسار `/logout` التوجيه إلى الصفحة الرئيسية.

```js
  const url = new URL("/logout", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/gi,
    'When a user logs out they should be redirected to the homepage'
  );
```

