---
id: 58966a17f9fc0f352b528e6d
title: تسجيل المستخدمين الجدد
challengeType: 2
forumTopicId: 301561
dashedName: registration-of-new-users
---

# --description--

الآن يجب أن تسمح لمستخدم جديد على موقعك بتسجيل حساب. في `res.render` لصفحة البداية، أضف متغيرًا جديدًا إلى الكائن المرسل - `showRegistration: true`. عند تحديث الصفحة، سترى نموذج التسجيل الذي تم إنشاؤه مسبقًا في ملف `index.pug`. هذا النموذج مُعد لإرسال بيانات **POST** إلى `/register`، فأنشئ هذا المسار واجعله يضيف كائن المستخدم إلى قاعدة البيانات باتباع المنطق التالي.

يجب أن يكون منطق مسار التسجيل كما يلي:

1. تسجيل المستخدم الجديد
2. توثيق المستخدم الجديد
3. إعادة التوجيه إلى `/profile`

يجب أن يكون منطق الخطوة 1 كما يلي:

1. استعلام قاعدة البيانات باستخدام `findOne`
2. إذا حدث خطأ، استدعِ `next` مع الخطأ
3. إذا تم إرجاع مستخدم، أعد التوجيه إلى الصفحة الرئيسية
4. إذا لم يُعثر على مستخدم ولم تحدث أخطاء، فقم بـ `insertOne` في قاعدة البيانات مع اسم المستخدم وكلمة المرور. طالما لم تحدث أخطاء، استدعِ `next` للانتقال إلى الخطوة 2، توثيق المستخدم الجديد، الذي كتبت منطقَه مسبقًا في مسار `POST /login`.

```js
app.route('/register')
  .post((req, res, next) => {
    myDataBase.findOne({ username: req.body.username }, (err, user) => {
      if (err) {
        next(err);
      } else if (user) {
        res.redirect('/');
      } else {
        myDataBase.insertOne({
          username: req.body.username,
          password: req.body.password
        },
          (err, doc) => {
            if (err) {
              res.redirect('/');
            } else {
              // The inserted document is held within
              // the ops property of the doc
              next(null, doc.ops[0]);
            }
          }
        )
      }
    })
  },
    passport.authenticate('local', { failureRedirect: '/' }),
    (req, res, next) => {
      res.redirect('/profile');
    }
  );
```

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#registration-of-new-users-11" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

**ملاحظة:** من هذه النقطة فصاعدًا، قد تظهر مشاكل تتعلق باستخدام متصفح *picture-in-picture*. إذا كنت تستخدم بيئة تطوير متكاملة على الإنترنت تعرض معاينة للتطبيق داخل المحرر، يُنصح بفتح هذه المعاينة في تبويب جديد.

# --hints--

يجب أن يكون لديك مسار `/register` وأن تعرض نموذج تسجيل على الصفحة الرئيسية.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /showRegistration:( |)true/gi,
    'You should be passing the variable showRegistration as true to your render function for the homepage'
  );
  assert.match(
    data,
    /register[^]*post[^]*findOne[^]*username:( |)req.body.username/gi,
    'You should have a route that accepts a POST request on /register that queries the db with findOne and the query being username: req.body.username'
  );
```

يجب أن يعمل التسجيل بشكل صحيح.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttp = new XMLHttpRequest();
  xhttp.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      test(this);
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttp.open('POST', url + '/register', true);
  xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
  xhttp.send(`username=${user}&password=${user}`);
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Register should work, and redirect successfully to the profile.'
    );
  }
```

يجب أن يعمل تسجيل الدخول بشكل صحيح.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttpReg = new XMLHttpRequest();
  xhttpReg.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      login();
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttpReg.open('POST', url + '/register', true);
  xhttpReg.setRequestHeader(
    'Content-type',
    'application/x-www-form-urlencoded'
  );
  xhttpReg.send(`username=${user}&password=${user}`);
  function login() {
    const xhttp = new XMLHttpRequest();
    xhttp.onreadystatechange = function () {
      if (this.readyState == 4 && this.status == 200) {
        test(this);
      } else {
        throw new Error(`${this.status} ${this.statusText}`);
      }
    };
    xhttp.open('POST', url + '/login', true);
    xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
    xhttp.send(`username=${user}&password=${user}`);
  }
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Login should work if previous test was done successfully and redirect successfully to the profile.'
    );
    assert.match(
      data,
      new RegExp(user, 'g'),
      'The profile should properly display the welcome to the user logged in'
    );
  }
```

يجب أن يعمل تسجيل الخروج بشكل صحيح.

```js
  $.ajax({
    url: code + '/logout',
    type: 'GET',
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(data, /Home/gi, 'Logout should redirect to home');
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

يجب ألا يعمل ملف التعريف بعد تسجيل الخروج.

```js
  $.ajax({
    url: code + '/profile',
    type: 'GET',
    crossDomain: true,
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(
        data,
        /Home/gi,
        'Profile should redirect to home when we are logged out now again'
      );
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

