---
id: 589a69f5f9fc0f352b528e71
title: تنفيذ التوثيق الاجتماعي II
challengeType: 2
forumTopicId: 301557
dashedName: implementation-of-social-authentication-ii
---

# --description--

الجزء الأخير من إعداد توثيق GitHub هو إنشاء الاستراتيجية نفسها. لقد تمت إضافة `passport-github@~1.1.0` كتبعية بالفعل، لذا استدعها في ملف `auth.js` كـ `GitHubStrategy` هكذا: `const GitHubStrategy = require('passport-github').Strategy;`. لا تنسَ استدعاء وتكوين `dotenv` لاستخدام متغيرات البيئة الخاصة بك.

لإعداد استراتيجية GitHub، عليك أن تخبر Passport باستخدام `GitHubStrategy` مُنشأ، والذي يقبل معلمتين: كائن (يحتوي على `clientID` و`clientSecret` و`callbackURL`) ودالة تُستدعى عند توثيق المستخدم بنجاح، لتحديد ما إذا كان المستخدم جديدًا وما الحقول التي تحفظها مبدئيًا في كائن قاعدة بيانات المستخدم. هذا شائع في العديد من الاستراتيجيات، لكن بعضها قد يتطلب معلومات إضافية كما هو موضح في ملف README الخاص بتلك الاستراتيجية على GitHub. على سبيل المثال، تتطلب Google أيضًا *نطاق* يحدد نوع المعلومات التي يطلب طلبك إرجاعها ويطلب من المستخدم الموافقة على هذا الوصول.

الاستراتيجية الحالية التي تنفذها توثق المستخدمين باستخدام حساب GitHub ورموز OAuth 2.0. يتم تزويد معرف العميل والسر الذي تحصل عليه عند إنشاء التطبيق كخيارات عند إنشاء الاستراتيجية. كما تتطلب الاستراتيجية استدعاء راجع `verify`، الذي يستقبل رمز الوصول ورمز التحديث الاختياري، بالإضافة إلى `profile` الذي يحتوي على ملف تعريف GitHub للمستخدم المُوثَّق. يجب أن يستدعي استدعاء `verify` الدالة `cb` مع توفير مستخدم لإكمال التوثيق.

هكذا يجب أن تبدو استراتيجيتك الجديدة في هذه المرحلة:

```js
passport.use(new GitHubStrategy({
  clientID: process.env.GITHUB_CLIENT_ID,
  clientSecret: process.env.GITHUB_CLIENT_SECRET,
  callbackURL: /*INSERT CALLBACK URL ENTERED INTO GITHUB HERE*/
},
  function(accessToken, refreshToken, profile, cb) {
    console.log(profile);
    //Database logic here with callback containing your user object
  }
));
```

لن ينجح التوثيق بعد، وسيُرمى خطأ فعليًا بدون منطق قاعدة البيانات والاستدعاء الراجع، لكن يجب أن يسجل ملف تعريف GitHub الخاص بك في وحدة التحكم إذا جربته!

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-ii-4" target="_blank" rel="noopener noreferrer nofollow">الاطلاع على المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب إضافة تبعية `passport-github`.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport-github',
    'Your project should list "passport-github" as a dependency'
  );
```

يجب استدعاء `passport-github`.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport-github("|')/gi,
    'You should have required passport-github'
  );
```

يجب إعداد استراتيجية GitHub بشكل صحيح حتى الآن.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport\.use.*new GitHubStrategy/gis,
    'Passport should use a new GitHubStrategy'
  );
  assert.match(
    data,
    /callbackURL:\s*("|').*("|')/gi,
    'You should have a callbackURL'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_SECRET|\[(?<q>"|')GITHUB_CLIENT_SECRET\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_SECRET'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_ID|\[(?<q>"|')GITHUB_CLIENT_ID\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_ID'
  );
```
