---
id: 589a8eb3f9fc0f352b528e72
title: تنفيذ التوثيق الاجتماعي III
challengeType: 2
forumTopicId: 301558
dashedName: implementation-of-social-authentication-iii
---

# --description--

الجزء النهائي من الاستراتيجية هو التعامل مع الملف الشخصي الذي يُعاد من GitHub. نحتاج إلى تحميل كائن قاعدة بيانات المستخدم إذا كان موجودًا، أو إنشاء واحد إذا لم يكن موجودًا، وملء الحقول من الملف الشخصي، ثم إرجاع كائن المستخدم. يوفر لنا GitHub معرفًا فريدًا *id* داخل كل ملف شخصي يمكننا استخدامه للبحث به لتسلسل المستخدم (تم تنفيذه بالفعل). فيما يلي مثال على التنفيذ يمكنك استخدامه في مشروعك — يوضع داخل الدالة التي هي المعلمة الثانية للاستراتيجية الجديدة، مباشرة تحت المكان الذي توجد فيه `console.log(profile);` حاليًا:

```js
myDataBase.findOneAndUpdate(
  { id: profile.id },
  {
    $setOnInsert: {
      id: profile.id,
      username: profile.username,
      name: profile.displayName || 'John Doe',
      photo: profile.photos[0].value || '',
      email: Array.isArray(profile.emails)
        ? profile.emails[0].value
        : 'No public email',
      created_on: new Date(),
      provider: profile.provider || ''
    },
    $set: {
      last_login: new Date()
    },
    $inc: {
      login_count: 1
    }
  },
  { upsert: true, new: true },
  (err, doc) => {
    return cb(null, doc.value);
  }
);
```

تسمح لك `findOneAndUpdate` بالبحث عن كائن وتحديثه. إذا لم يكن الكائن موجودًا، فسيتم إدخاله وجعله متاحًا لدالة الاستدعاء الراجع. في هذا المثال، نُعيّن دائمًا `last_login`، ونزيد `login_count` بمقدار `1`، ونملأ غالبية الحقول فقط عند إدخال كائن جديد (مستخدم جديد). لاحظ استخدام القيم الافتراضية. أحيانًا لا يحتوي الملف الشخصي المعاد على كل المعلومات أو يحتفظ المستخدم بها خاصة. في هذه الحالة، تتعامل معها لمنع حدوث خطأ.

يجب أن تكون قادرًا على تسجيل الدخول إلى تطبيقك الآن. جرّب ذلك!

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-iii-5" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن يكون إعداد استراتيجية GitHub مكتملًا.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /GitHubStrategy[^]*myDataBase/gi,
    'Strategy should use now use the database to search for the user'
  );
  assert.match(
    data,
    /GitHubStrategy[^]*cb/gi,
    'Strategy should return the callback function "cb"'
  );
```

