---
id: 589a69f5f9fc0f352b528e70
title: تنفيذ التوثيق الاجتماعي
challengeType: 2
forumTopicId: 301559
dashedName: implementation-of-social-authentication
---

# --description--

المسار الأساسي الذي سيتبعه هذا النوع من التوثيق في تطبيقك هو:

1. ينقر المستخدم على زر أو رابط يرسله إلى مسارك لتوثيق الدخول باستخدام استراتيجية محددة (مثل GitHub).
2. يستدعي مسارك `passport.authenticate('github')` الذي يعيد توجيه المستخدم إلى GitHub.
3. الصفحة التي يصل إليها المستخدم على GitHub تسمح له بتسجيل الدخول إذا لم يكن مسجلاً بالفعل. ثم تطلب منه الموافقة على وصول تطبيقك إلى ملفه الشخصي.
4. يُعاد المستخدم إلى تطبيقك عند عنوان URL لاستدعاء راجع محدد مع ملفه الشخصي إذا تمت الموافقة عليه.
5. يصبح المستخدم مُوثَّقًا الآن، ويجب على تطبيقك التحقق مما إذا كان ملفه الشخصي عائدًا، أو حفظه في قاعدة البيانات إذا لم يكن كذلك.

تتطلب الاستراتيجيات التي تستخدم OAuth أن يكون لديك على الأقل *معرّف عميل* و*سر عميل*، وهما وسيلة للخدمة للتحقق من مصدر طلب التوثيق وصحته. تُحصل هذه القيم من الموقع الذي تحاول تنفيذ التوثيق معه، مثل GitHub، وهي فريدة لتطبيقك--**لا يجب مشاركتها** ولا ينبغي رفعها إلى مستودع عام أو كتابتها مباشرة في شفرتك. من الممارسات الشائعة وضعها في ملف `.env` الخاص بك والإشارة إليها مثل: `process.env.GITHUB_CLIENT_ID`. في هذا التحدي ستستخدم استراتيجية GitHub.

اتبع <a href="https://www.freecodecamp.org/news/how-to-set-up-a-github-oauth-application/" target="_blank" rel="noopener noreferrer nofollow">هذه التعليمات</a> للحصول على *معرّف العميل والسر* من GitHub. عيّن عنوان URL للصفحة الرئيسية إلى صفحتك الرئيسية (**ليس عنوان شفرة المشروع**)، وعيّن عنوان URL لاستدعاء الراجع إلى نفس عنوان الصفحة الرئيسية مع إضافة `/auth/github/callback` في نهايته. احفظ معرّف العميل وسرك في ملف `.env` الخاص بمشروعك كـ `GITHUB_CLIENT_ID` و `GITHUB_CLIENT_SECRET`.

في ملف `routes.js`، أضف `showSocialAuth: true` إلى مسار الصفحة الرئيسية، بعد `showRegistration: true`. الآن، أنشئ مسارين يقبلان طلبات GET: `/auth/github` و `/auth/github/callback`. يجب أن يستدعي الأول فقط passport لتوثيق `'github'`. أما الثاني فيجب أن يستدعي passport لتوثيق `'github'` مع إعادة توجيه للفشل إلى `/`، وإذا نجح التوثيق يعيد التوجيه إلى `/profile` (مماثل لمشروعك السابق).

مثال على كيفية ظهور `/auth/github/callback` مشابه لكيفية تعاملك مع تسجيل الدخول العادي:

```js
app.route('/login')
  .post(passport.authenticate('local', { failureRedirect: '/' }), (req,res) => {
    res.redirect('/profile');
  });
```

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-3" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع حتى هذه النقطة</a>.

# --hints--

يجب أن يكون مسار `/auth/github` صحيحًا.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
          data.replace(/\s/g, ''),
          /passport.authenticate.*?github/g,
          'Route auth/github should only call passport.authenticate with github'
        );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```

يجب أن يكون مسار `/auth/github/callback` صحيحًا.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
        data.replace(/\s/g, ''),
        /failureRedirect:("|')\/\1/g,
        'Route auth/github/callback should accept a get request and call passport.authenticate for github with a failure redirect to home'
      );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github/callback');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github/callback"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```

