---
id: 589fc831f9fc0f352b528e76
title: التعامل مع فصل الاتصال
challengeType: 2
forumTopicId: 301552
dashedName: handle-a-disconnect
---

# --description--

قد تلاحظ أنه حتى الآن كنت تزيد فقط عدد المستخدمين. التعامل مع فصل اتصال المستخدم سهل تمامًا مثل التعامل مع الاتصال الأولي، إلا أنه يجب الاستماع إليه على كل socket بدلاً من الاستماع على الخادم بأكمله.

للقيام بذلك، أضف مستمعًا آخر داخل مستمع `'connect'` الحالي يستمع إلى `'disconnect'` على الـ socket دون تمرير بيانات. يمكنك اختبار هذه الوظيفة بمجرد تسجيل أن المستخدم قد فصل الاتصال في وحدة التحكم.

```js
socket.on('disconnect', () => {
  /*anything you want to do on disconnect*/
});
```

للتأكد من أن العملاء لديهم دائمًا العدد المحدث للمستخدمين الحاليين، يجب عليك تقليل `currentUsers` بمقدار 1 عند حدوث فصل الاتصال ثم إرسال حدث `'user count'` مع العدد المحدث.

**ملاحظة:** تمامًا مثل `'disconnect'`، يجب التعامل مع جميع الأحداث الأخرى التي يمكن أن يرسلها الـ socket إلى الخادم داخل مستمع الاتصال حيث تم تعريف 'socket'.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#handle-a-disconnect-8" target="_blank" rel="noopener noreferrer nofollow">الاطلاع على المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب على الخادم التعامل مع حدث فصل الاتصال من socket.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /socket.on.*('|")disconnect('|")/s, '');
```

يجب أن يستمع العميل إلى حدث `'user count'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```

