---
id: 589fc831f9fc0f352b528e77
title: التوثيق مع Socket.IO
challengeType: 2
forumTopicId: 301548
dashedName: authentication-with-socket-io
---

# --description--

حاليًا، لا يمكنك معرفة من المتصل بـ web socket الخاص بك. بينما يحتوي `req.user` على كائن المستخدم، هذا فقط عندما يتفاعل المستخدم مع خادم الويب، ومع web sockets لا يوجد لديك `req` (طلب) وبالتالي لا توجد بيانات مستخدم. إحدى الطرق لحل مشكلة معرفة من المتصل بـ web socket هي تحليل وفك تشفير الكوكي الذي يحتوي على جلسة passport ثم إلغاء تسلسلها للحصول على كائن المستخدم. لحسن الحظ، هناك حزمة على NPM مخصصة لهذا الأمر تجعل مهمة كانت معقدة بسيطة!

تمت إضافة `passport.socketio@~3.7.0`، و`connect-mongo@~3.2.0`، و`cookie-parser@~1.4.5` كتبعيات. استوردها كـ `passportSocketIo`، و`MongoStore`، و`cookieParser` على التوالي. أيضًا، نحتاج إلى تهيئة مخزن ذاكرة جديد من `express-session` التي استوردناها سابقًا. يجب أن تبدو هكذا:

```js
const MongoStore = require('connect-mongo')(session);
const URI = process.env.MONGO_URI;
const store = new MongoStore({ url: URI });
```

الآن فقط أخبر Socket.IO باستخدامه واضبط الخيارات. تأكد من إضافة هذا قبل شفرة socket الحالية وليس داخل مستمع الاتصال الحالي. يجب أن يبدو خادمك هكذا:

```js
io.use(
  passportSocketIo.authorize({
    cookieParser: cookieParser,
    key: 'express.sid',
    secret: process.env.SESSION_SECRET,
    store: store,
    success: onAuthorizeSuccess,
    fail: onAuthorizeFail
  })
);
```

لاحظ أن تهيئة توثيق Passport لـ Socket.IO مشابهة جدًا للطريقة التي أعددنا بها برمجيات وسيطة `session` لواجهة برمجة التطبيقات. ذلك لأنهما مصممان لاستخدام نفس طريقة التوثيق — الحصول على معرف الجلسة من الكوكي والتحقق منه.

سابقًا، عندما أعددنا برمجيات وسيطة `session`، لم نحدد صراحة اسم الكوكي للجلسة (`key`). ذلك لأن حزمة `session` كانت تستخدم القيمة الافتراضية. الآن بعد أن أضفنا حزمة أخرى تحتاج إلى الوصول لنفس القيمة من الكوكيز، يجب تحديد قيمة `key` صراحة في كلا كائنَي التهيئة.

تأكد من إضافة `key` مع اسم الكوكي إلى برمجيات وسيطة `session` التي تطابق مفتاح Socket.IO. أضف أيضًا مرجع `store` إلى الخيارات، بالقرب من المكان الذي نحدد فيه `saveUninitialized: true`. هذا ضروري لإخبار Socket.IO بأي جلسة يجب ربطها.

<hr>

الآن، عرّف دالتي الاستدعاء الراجعة `success` و`fail`:

```js
function onAuthorizeSuccess(data, accept) {
  console.log('successful connection to socket.io');

  accept(null, true);
}

function onAuthorizeFail(data, message, error, accept) {
  if (error) throw new Error(message);
  console.log('failed connection to socket.io:', message);
  accept(null, false);
}
```

كائن المستخدم متاح الآن على كائن socket الخاص بك كـ `socket.request.user`. على سبيل المثال، يمكنك الآن إضافة التالي:

```js
console.log('user ' + socket.request.user.username + ' connected');
```

سيُسجل في وحدة التحكم الخاصة بالخادم من المتصل!

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#authentication-with-socketio-9" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع حتى هذه النقطة</a>.

# --hints--

يجب أن تكون `passport.socketio` تبعية.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport.socketio',
    'Your project should list "passport.socketio" as a dependency'
  );
```

يجب أن تكون `cookie-parser` تبعية.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'cookie-parser',
    'Your project should list "cookie-parser" as a dependency'
  );
```

يجب استيراد passportSocketIo بشكل صحيح.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\((['"])passport\.socketio\1\)/gi,
    'You should correctly require and instantiate "passport.socketio"'
  );
```

يجب إعداد passportSocketIo بشكل صحيح.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io\.use\(\s*\w+\.authorize\(/,
    'You should register "passport.socketio" as socket.io middleware and provide it correct options'
  );
```

