---
id: 589fc832f9fc0f352b528e79
title: إرسال وعرض رسائل الدردشة
challengeType: 2
forumTopicId: 301562
dashedName: send-and-display-chat-messages
---

# --description--

حان الوقت لتسمح للعملاء بإرسال رسالة دردشة إلى الخادم ليتم بثها إلى جميع العملاء! في ملف `client.js` الخاص بك، يجب أن ترى أن هناك بالفعل جزءًا من الشفرة يتعامل مع إرسال نموذج الرسالة.

```js
$('form').submit(function() {
  /*logic*/
});
```

داخل شفرة إرسال النموذج، يجب أن تبث حدثًا بعد تعريف `messageToSend` ولكن قبل أن تمسح مربع النص `#m`. يجب أن يُسمى الحدث `'chat message'` والبيانات يجب أن تكون فقط `messageToSend`.

```js
socket.emit('chat message', messageToSend);
```

الآن، على الخادم الخاص بك، يجب أن تستمع إلى المقبس للحدث `'chat message'` مع تسمية البيانات بـ `message`. بمجرد استلام الحدث، يجب أن يبث الحدث `'chat message'` إلى جميع المقابس باستخدام `io.emit`، مرسلًا كائن بيانات يحتوي على `username` و `message`.

في `client.js`، يجب أن تستمع الآن للحدث `'chat message'` وعند استلامه، تضيف عنصر قائمة إلى `#messages` يحتوي على اسم المستخدم، ونقطتين، والرسالة!

في هذه المرحلة، يجب أن تكون الدردشة تعمل بشكل كامل وترسل الرسائل عبر جميع العملاء!

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#send-and-display-chat-messages-11" target="_blank" rel="noopener noreferrer nofollow">الاطلاع على المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن يستمع الخادم للحدث `'chat message'` ثم يبثه بشكل صحيح.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*io.emit.*('|")chat message('|").*username.*message/s,
    'Your server should listen to the socket for "chat message" then emit to all users "chat message" with name and message in the data object'
  );
```

يجب أن يتعامل العميل بشكل صحيح مع البيانات الجديدة من الحدث `'chat message'` ويعرضها.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*messages.*li/s,
    'You should append a list item to #messages on your client within the "chat message" event listener to display the new message'
  );
```

