---
id: 589fc831f9fc0f352b528e75
title: التواصل عن طريق الإرسال
challengeType: 2
forumTopicId: 301550
dashedName: communicate-by-emitting
---

# --description--

<dfn>Emit</dfn> هو أكثر طرق التواصل التي ستستخدمها شيوعًا. عندما ترسل شيئًا من الخادم إلى 'io'، ترسل اسم حدث وبياناته إلى جميع المقابس المتصلة. مثال جيد على هذا المفهوم هو إرسال العدد الحالي للمستخدمين المتصلين في كل مرة يتصل فيها مستخدم جديد!

ابدأ بإضافة متغير لتتبع المستخدمين، قبل المكان الذي تستمع فيه حاليًا للاتصالات.

```js
let currentUsers = 0;
```

الآن، عندما يتصل شخص ما، يجب زيادة العدد قبل إرسال العدد. لذلك، ستضيف الزيادة داخل مستمع الاتصال.

```js
++currentUsers;
```

أخيرًا، بعد زيادة العدد، يجب أن ترسل الحدث (لا يزال داخل مستمع الاتصال). يجب أن يُسمى الحدث 'user count'، والبيانات يجب أن تكون فقط `currentUsers`.

```js
io.emit('user count', currentUsers);
```

الآن، يمكنك تنفيذ طريقة لعميلك ليستمع لهذا الحدث! مشابه للاستماع للاتصال على الخادم، ستستخدم الكلمة المفتاحية `on`.

```js
socket.on('user count', function(data) {
  console.log(data);
});
```

الآن، جرب تحميل تطبيقك، وقم بالمصادقة، ويجب أن ترى في وحدة تحكم العميل '1' تمثل العدد الحالي للمستخدمين! جرب تحميل المزيد من العملاء والمصادقة لترى العدد يرتفع.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#communicate-by-emitting-7" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب تعريف `currentUsers`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /currentUsers/s,
    'You should have variable currentUsers defined'
  );
```

يجب أن يرسل الخادم العدد الحالي للمستخدمين عند كل اتصال جديد.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.emit.*('|")user count('|").*currentUsers/s,
    'You should emit "user count" with data currentUsers'
  );
```

يجب أن يستمع عميلك لحدث `'user count'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```

