---
id: 589fc830f9fc0f352b528e74
title: إعداد بيئة العمل
challengeType: 2
forumTopicId: 301566
dashedName: set-up-the-environment
---

# --description--

ستستخدم التحديات التالية ملف `chat.pug`. لذا، في ملف `routes.js` الخاص بك، أضف مسار GET يشير إلى `/chat` يستخدم `ensureAuthenticated`، ويعرض `chat.pug` مع تمرير `{ user: req.user }` كمعلمة إلى الاستجابة. الآن، عدل مسار `/auth/github/callback` الحالي ليعيّن `req.session.user_id = req.user.id`، ويُعيد التوجيه إلى `/chat`.

تمت إضافة `socket.io@~2.3.0` كتبعية بالفعل، لذا استدعِها وأنشئها في الخادم الخاص بك كما يلي مع `http` (مدمج مع Node.js):

```javascript
const http = require('http').createServer(app);
const io = require('socket.io')(http);
```

بعد أن تم تركيب خادم *http* على تطبيق *express*، يجب أن تستمع من خادم *http*. غيّر السطر الذي يحتوي على `app.listen` إلى `http.listen`.

أول شيء يجب التعامل معه هو الاستماع لاتصال جديد من العميل. الكلمة المفتاحية <dfn>on</dfn> تفعل ذلك بالضبط - تستمع لحدث معين. تتطلب معلمتين: سلسلة نصية تحتوي على عنوان الحدث المُرسل، ودالة تمرر البيانات من خلالها. في حالة مستمع الاتصال لدينا، استخدم `socket` لتعريف البيانات في المعلمة الثانية. الـ socket هو عميل فردي متصل.

للاستماع للاتصالات إلى خادمك، أضف التالي داخل اتصال قاعدة البيانات الخاص بك:

```javascript
io.on('connection', socket => {
  console.log('A user has connected');
});
```

الآن لكي يتصل العميل، فقط أضف التالي إلى `client.js` الذي يتم تحميله من الصفحة بعد توثيقك:

```js
/*global io*/
let socket = io();
```

التعليق يمنع ظهور الخطأ الذي عادةً ما تراه لأن 'io' غير معرف في الملف. لقد أضفت بالفعل شبكة CDN موثوقة لمكتبة Socket.IO في الصفحة داخل `chat.pug`.

جرّب الآن تحميل التطبيق وتوثيق الدخول، ويجب أن ترى في وحدة تحكم الخادم رسالة `A user has connected`.

**ملاحظة:** `io()` تعمل فقط عند الاتصال بـ socket مستضاف على نفس عنوان URL/الخادم. للاتصال بـ socket خارجي مستضاف في مكان آخر، تستخدم `io.connect('URL');`.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-the-environment-6" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن تكون `socket.io` تبعية.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'socket.io',
    'Your project should list "socket.io" as a dependency'
  );
```

يجب أن تستدعي وتُنشئ `http` بشكل صحيح كـ `http`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /http.*=.*require.*('|")http\1/s,
    'Your project should list "http" as a dependency'
  );
```

يجب أن تستدعي وتُنشئ `socket.io` بشكل صحيح كـ `io`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.*=.*require.*('|")socket.io\1.*http/s,
    'You should correctly require and instantiate socket.io as io.'
  );
```

يجب أن تستمع Socket.IO للاتصالات.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.on.*('|")connection\1.*socket/s,
    'io should listen for "connection" and socket should be the 2nd arguments variable'
  );
```

يجب أن يتصل عميلك بخادمك.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.*=.*io/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```

