---
id: 589fc832f9fc0f352b528e78
title: إعلان المستخدمين الجدد
challengeType: 2
forumTopicId: 301546
dashedName: announce-new-users
---

# --description--

تستطيع العديد من غرف الدردشة الإعلان عند اتصال أو قطع اتصال مستخدم، ثم عرض ذلك لجميع المستخدمين المتصلين في الدردشة. بما أنك تصدر حدثًا عند الاتصال والقطع، فقط عدّل هذا الحدث لدعم هذه الميزة. أنسب طريقة لذلك هي إرسال 3 بيانات مع الحدث: اسم المستخدم الذي اتصل أو قطع الاتصال، عدد المستخدمين الحالي، وما إذا كان ذلك المستخدم قد اتصل أو قطع الاتصال.

غيّر اسم الحدث إلى `'user'`، ومرر كائنًا يحتوي على الحقول `username` و`currentUsers` و`connected` (لتكون `true` في حالة الاتصال، أو `false` في حالة قطع الاتصال للمستخدم المرسل). تأكد من تعديل كلا الحدثين `'user count'` واجعل حدث قطع الاتصال يرسل `false` للخاصية `connected` بدلًا من `true` كما في الحدث الصادر عند الاتصال.

```js
io.emit('user', {
  username: socket.request.user.username,
  currentUsers,
  connected: true
});
```

الآن سيكون لدى العميل كل المعلومات اللازمة لعرض عدد المستخدمين الحالي بشكل صحيح والإعلان عند اتصال أو قطع اتصال مستخدم! للتعامل مع هذا الحدث على جانب العميل، استمع إلى `'user'`، ثم حدّث عدد المستخدمين الحالي باستخدام jQuery لتغيير نص `#num-users` إلى `'{NUMBER} users online'`، وأضف عنصر `<li>` إلى القائمة غير المرتبة ذات المعرف `messages` بالنص `'{NAME} has {joined/left} the chat.'`.

يمكن أن تبدو طريقة التنفيذ كما يلي:

```js
socket.on('user', data => {
  $('#num-users').text(data.currentUsers + ' users online');
  let message =
    data.username +
    (data.connected ? ' has joined the chat.' : ' has left the chat.');
  $('#messages').append($('<li>').html('<b>' + message + '</b>'));
});
```

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك مراجعة <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135/3#announce-new-users-10" target="_blank" rel="noopener noreferrer nofollow">المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب إصدار حدث `'user'` مع `username` و`currentUsers` و`connected`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  // Regex is lenient to match both `username` and `name` as the key on purpose.
  assert.match(
    data,
    /io.emit.*('|")user\1.*name.*currentUsers.*connected/s,
    'You should have an event emitted named user sending name, currentUsers, and connected'
  );
```

يجب على العميل التعامل بشكل صحيح مع البيانات الجديدة من حدث `'user'` وعرضها.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*num-users/s,
    'You should change the text of "#num-users" within on your client within the "user" event listener to show the current users connected'
  );
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*messages.*li/s,
    'You should append a list item to "#messages" on your client within the "user" event listener to announce a user came or went'
  );
```

