---
id: 58a25c98f9fc0f352b528e7f
title: تجزئة كلمات المرور الخاصة بك
challengeType: 2
forumTopicId: 301553
dashedName: hashing-your-passwords
---

# --description--

بالعودة إلى قسم أمن المعلومات، قد تتذكر أن تخزين كلمات المرور بنصها الصريح *غير مقبول* أبدًا. حان الوقت الآن لتنفيذ BCrypt لحل هذه المشكلة.

تمت إضافة `bcrypt@~5.0.0` بالفعل كتبعيات، لذا استوردها في الخادم الخاص بك. ستحتاج إلى التعامل مع التجزئة في منطقتين رئيسيتين: حيث تسجل/تحفظ حسابًا جديدًا، وعندما تتحقق من صحة كلمة المرور عند تسجيل الدخول.

حاليًا في مسار التسجيل، تدخل كلمة مرور المستخدم بنصها الصريح في قاعدة البيانات بهذا الشكل: `password: req.body.password`. بدلاً من ذلك، جزئ كلمات المرور بإضافة السطر التالي قبل منطق قاعدة البيانات: `const hash = bcrypt.hashSync(req.body.password, 12);`، واستبدل `req.body.password` في الحفظ بقاعدة البيانات بـ `password: hash` فقط.

في استراتيجية التوثيق الخاصة بك، تتحقق من الشرط التالي في الشفرة قبل إكمال العملية: `if (password !== user.password) return done(null, false);`. بعد التغييرات السابقة، أصبح `user.password` عبارة عن تجزئة. قبل تعديل الشفرة الحالية، لاحظ أن الشرط يتحقق مما إذا كانت كلمة المرور **غير** متساوية ثم يعيد عدم التوثيق. بناءً على ذلك، عدل الشفرة لتبدو كما يلي للتحقق بشكل صحيح من كلمة المرور المدخلة مقابل التجزئة:

```js
if (!bcrypt.compareSync(password, user.password)) { 
  return done(null, false);
}
```

هذا كل ما يلزم لتنفيذ واحدة من أهم ميزات الأمان عند الحاجة إلى تخزين كلمات المرور.

أرسل صفحتك عندما تعتقد أنك أتممتها بشكل صحيح. إذا واجهت أخطاء، يمكنك <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#hashing-your-passwords-1" target="_blank" rel="noopener noreferrer nofollow">مراجعة المشروع المكتمل حتى هذه النقطة</a>.

# --hints--

يجب أن تكون BCrypt تبعية.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json()
  assert.property(
    packJson.dependencies,
    'bcrypt',
    'Your project should list "bcrypt" as a dependency'
  );
```

يجب استيراد BCrypt وتنفيذها بشكل صحيح.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')bcrypt\1/gi,
    'You should have required bcrypt'
  );
  assert.match(
    data,
    /bcrypt.hashSync/gi,
    'You should use hash the password in the registration'
  );
  assert.match(
    data,
    /bcrypt.compareSync/gi,
    'You should compare the password to the hash in your strategy'
  );
```

