---
id: 5f332b23c2045fb843337579
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

To let visitors know the cafe was founded in `2020`, add a `p` element below the `h1` element with the text `Est. 2020`.

# --hints--

你應該有一個 `<p>` 開始標籤。

```js
assert.match(code, /<p>/i);
```

你應該有一個 `</p>` 結束標籤。

```js
assert.match(code, /<\/p>/i);
```

你不應該改變你現有的 `h1` 元素。 確認你沒有刪除結束標籤。

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

你的 `p` 元素應該在 `h1` 元素的下面。

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

你的 `p` 元素應該有文本 `Est. 2020`。

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```

