---
id: 5f332b23c2045fb843337579
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

為讓訪客知道咖啡館成立於 `2020`，請在 `h1` 元素下方添加一個 `p` 元素，內容為 `Est. 2020`。

# --hints--

你應該有一個開啟的 `<p>` 頁籤。

```js
assert.match(code, /<p>/i);
```

你應該有一個結束的 `</p>` 頁籤。

```js
assert.match(code, /<\/p>/i);
```

你不應該更改你現有的 `h1` 元素。請確保你沒有刪除結束標籤。

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

你的 `p` 元素應該在你的 `h1` 元素下方。

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

你的 `p` 元素應該有文字 `Est. 2020`。

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

