---
id: 5f34a1fd611d003edeafd681
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

那個棕色背景讓文字難以閱讀。將 `body` 元素的背景顏色改為 `burlywood`，這樣會有一些顏色，但你仍然能閱讀文字。

# --hints--

你應該將 `background-color` 屬性設為 `burlywood`。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'burlywood');
assert.isTrue(hasBackground);
```

你的 `body` 元素應該有 `burlywood` 背景。

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert.equal(bodyBackground, 'burlywood');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  background-color: brown;
--fcc-editable-region--
}
h1, h2, p {
  text-align: center;
}
```
