---
id: 5f344fad8bf01691e71a30eb
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

到目前爲止，你對內容的展示和外觀的控制有限。 爲了改變這種情況，在 `head` 元素中創建一個 `style` 元素。

# --hints--

你應該有一個 `<style>` 開始標籤。

```js
assert.match(code, /<style\s*>/i);
```

你應該有一個 `</style>` 結束標籤。

```js
assert.match(code, /<\/style\s*>/);
```

你的 `style` 元素都應嵌套在 `head` 元素中。

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
