---
id: 5f344fad8bf01691e71a30eb
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

直到現在，你對內容的呈現和外觀只有有限的控制。若要改變這點，請在 `head` 元素內添加一個 `style` 元素。

# --hints--

你的程式碼應該有一個開啟的 `<style>` 頁籤。

```js
assert.match(code, /<style\s*>/i);
```

你的程式碼應該有一個結尾的 `</style>` 頁籤。

```js
assert.match(code, /<\/style\s*>/);
```

你的 `style` 元素應該巢狀在你的 `head` 元素中。

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
