---
id: 5f3477ae9675db8bb7655b30
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

在前一步驟中，你使用了 <dfn>型別選擇器</dfn> 來設定 `h1` 元素的樣式。透過為每個元素新增一個新的型別選擇器，將 `h2` 和 `p` 元素的內容置中，並將其添加到現有的 `style` 元素中。

# --hints--

你不應該更改現有的 `h1` 選擇器。

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasH1);
```

你不應該新增新的 `style` 頁籤。將新的 CSS 規則添加到現有的 `style` 頁籤。

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

你應該為你的 `h2` 選擇器添加一個新的。

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert.exists(hasH2);
```

你應該為你的 `p` 選擇器添加一個新的。

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert.exists(hasP);
```

你的 `h1` 元素應該有 `text-align` 為 `center`。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

你的 `h2` 元素的 `text-align` 應該是 `center`。

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert.equal(h2TextAlign, 'center');
```

你的 `p` 元素應該有 `text-align` 為 `center`。

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert.equal(pTextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1 {
        text-align: center;
      }
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

