---
id: 5f344fbc22624a2976425065
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `section` 元素中創建一個 `h2` 元素，並為它添加文字 `Coffee`。

# --hints--

你應該有一個開啟的 `<h2>` 頁籤。

```js
assert.match(code, /<h2\s*>/i);
```

你應該有一個結束的 `</h2>` 頁籤。

```js
assert.match(code, /<\/h2\s*>/i);
```

你不應該更改你現有的 `section` 元素。請確保你沒有刪除結束標籤。

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

你的 `h2` 元素應該位於你的 `section` 元素內。

```js
const h2 = document.querySelector('h2');
assert.equal(h2.parentElement.tagName, 'SECTION');
```

你的 `h2` 元素應該有文字 `Coffee`。

```js
const h2 = document.querySelector('h2');
assert.equal(h2?.innerText.trim(), 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
