---
id: 5f356ed6199b0cdef1d2be8f
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

到目前為止，你已經使用 `type` 和 `id` 選擇器來設定元素樣式。然而，更常見的是使用不同的選擇器來設定你的元素樣式。 

你已經學會如何使用 `class` 選擇器，像這樣在之前的課程中：

```css
.class-name {
  styles
}
```

將現有的 `#menu` 選擇器改為類別選擇器，方法是將 `#menu` 替換為名為 `.menu` 的類別。

# --hints--

你應該有一個 `.menu` 類別選擇器。

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert.exists(hasMenu);
```

你不應該有 `#menu` 選擇器。

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert.notExists(hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
--fcc-editable-region--
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

