---
id: 5f356ed656a336993abd9f7c
title: 第 21 步
challengeType: 0
dashedName: step-21
---

# --description--

接著，你想要將 `#menu` 水平置中。你可以透過將它的 `margin-left` 和 `margin-right` 屬性設為 `auto` 來達成。將 margin 想像成元素周圍的隱形空間。使用這兩個 margin 屬性，將 `#menu` 元素置中於 `body` 元素內。

# --hints--

你應該將 `margin-left` 屬性設為 `auto`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-left'] === 'auto');
assert.isTrue(hasMargin);
```

你應該將 `margin-right` 屬性設為 `auto`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-right'] === 'auto');
assert.isTrue(hasMargin);
```

你應該將你的 `#menu` 的 `margin-left` 和 `margin-right` 屬性設為 `auto`。

```js
const divMarginRight = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-right');
const divMarginLeft = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-left');
assert.equal(divMarginRight, 'auto');
assert.equal(divMarginLeft, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 80%;
  background-color: burlywood;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
