---
id: 5f3c866daec9a49519871816
title: 步驟 26
challengeType: 0
dashedName: step-26
---

# --description--

`article` 元素通常包含多個具有相關資訊的元素。在這種情況下，它將包含咖啡風味和該風味的價格。

在你的 `article` 元素內巢狀兩個 `p` 元素。第一個的文字應該是 `French Vanilla`，第二個的文字是 `3.00`。

# --hints--

你不應該更改現有的 `article` 元素。

```js
assert.lengthOf(document.querySelectorAll('article'),1);
```

你的 `article` 元素應該有兩個 `p` 元素。

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphChildren,2);
```

你的第一個 `p` 元素應該有文字 `French Vanilla`。

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const firstP = paragraphChildren?.[0];
assert.match(firstP?.innerText.trim(), /French Vanilla/i);
```

你的第二個 `p` 元素應該有文字 `3.00`。

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const secondP = paragraphChildren?.[1];
assert.match(secondP?.innerText.trim(), /3\.00/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
--fcc-editable-region--
            
--fcc-editable-region--
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
