---
id: 5f3c866de7a5b784048f94b1
title: 步驟 32
challengeType: 0
dashedName: step-32
---

# --description--

這大致上是你想要的，但現在如果口味和價格能在同一行會更好。`p` 元素是 <dfn>區塊層級</dfn> 元素，因此它們會佔據父元素的整個寬度。

要讓它們在同一行，你需要對 `p` 元素套用一些樣式，使它們的行為更像 <dfn>行內</dfn> 元素。

為了做到這一點，請先為 `Coffee` 標題下的第一個 `article` 項目添加一個值為 `item` 的 `class` 屬性。

# --hints--

你應該將 `item` 類別套用到你的 `article` 元素。

```js
const el = document.querySelector('article.item');
assert.exists(el);
```

你應該只擁有一個 `item` 類別元素。

```js
const elements = document.querySelectorAll('.item');
assert.lengthOf(elements, 1);
```

你的第一個 `article` 元素應該有 `item` 類別。

```js
assert.equal(document.querySelector('article')?.className, 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
--fcc-editable-region--
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
