---
id: 5f3ef6e0a81099d9a697b550
title: 步驟 61
challengeType: 0
dashedName: step-61
---

# --description--

在 `address` 內，為你的 `address` 添加一個 `p` 元素。然後，在 `p` 中巢狀一個錨點（`a`）元素，該元素的聯結為 `https://www.freecodecamp.org` 且文字為 `Visit our website`。

確保該聯結透過添加值為 `_blank` 的 `target` 屬性，在新分頁中開啟。

# --hints--

你不應該修改現有的 `footer` 元素。

```js
assert.lengthOf(document.querySelectorAll("footer"), 1);
```

你不應該修改現有的 `address` 元素。

```js
assert.lengthOf(document.querySelectorAll("address"), 1);
```

你的新 `p` 元素應該巢狀在你的 `address` 元素內。你應該只有一個 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll("footer > address > p"), 1);
assert.lengthOf(document.querySelectorAll("footer p"), 1);
```

你的新 `a` 元素應該巢狀在你的新 `p` 元素內。你應該只有一個 `a` 元素。

```js
assert.lengthOf(document.querySelectorAll("footer > address > p > a"), 1);
assert.lengthOf(document.querySelectorAll("footer a"), 1);
```

你的新 `a` 元素應該有文字 `Visit our website`。

```js
assert.equal(document.querySelector("footer > address > p > a")?.innerText.trim(), "Visit our website");
```

你的新 `a` 元素應該聯結到 `https://www.freecodecamp.org`。請記得 `a` 元素使用 `href` 屬性來創建聯結。

```js
assert.equal(document.querySelector("footer > address > p > a")?.href, "https://www.freecodecamp.org/");
```

你的新 `a` 元素應該將 `target` 屬性設定為 `_blank`。

```js
assert.equal(document.querySelector("footer > address > p > a")?.target, "_blank");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
      <footer>
        <address>
--fcc-editable-region--
          
--fcc-editable-region--
        </address>
      </footer>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
}

h1 {
  font-size: 40px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

