---
id: 5f7692f7c5b3ce22a57788b6
title: 步驟 40
challengeType: 0
dashedName: step-40
---

# --description--

為你的所有剩餘 `p` 元素添加適用的類別名稱 `flavor` 和 `price`，以完成樣式設定。

# --hints--

你應該有五個 `.flavor` 元素。

```js
assert.lengthOf(document.querySelectorAll('p.flavor'), 5);
```

你應該有五個 `.price` 元素。

```js
assert.lengthOf(document.querySelectorAll('p.price'), 5);
```

你的 `.flavor` 元素應該是包含文字 `French Vanilla`、`Caramel Macchiato`、`Pumpkin Spice`、`Hazelnut` 和 `Mocha` 的 `p` 元素。

```js
const flavor = document.querySelectorAll('p.flavor');
assert.match(flavor?.[0]?.innerText.trim(), /French Vanilla/i);
assert.match(flavor?.[1]?.innerText.trim(), /Caramel Macchiato/i);
assert.match(flavor?.[2]?.innerText.trim(), /Pumpkin Spice/i);
assert.match(flavor?.[3]?.innerText.trim(), /Hazelnut/i);
assert.match(flavor?.[4]?.innerText.trim(), /Mocha/i);
```

你的 `.price` 元素應該是你的 `p` 元素，且文字為 `3.00`、`3.75`、`3.50`、`4.00` 和 `4.50`。

```js
const price = document.querySelectorAll('p.price');
assert.match(price?.[0]?.innerText.trim(), /3\.00/);
assert.match(price?.[1]?.innerText.trim(), /3\.75/);
assert.match(price?.[2]?.innerText.trim(), /3\.50/);
assert.match(price?.[3]?.innerText.trim(), /4\.00/);
assert.match(price?.[4]?.innerText.trim(), /4\.50/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article class="item">
            <p>Caramel Macchiato</p><p>3.75</p>
          </article>
          <article class="item">
            <p>Pumpkin Spice</p><p>3.50</p>
          </article>
          <article class="item">
            <p>Hazelnut</p><p>4.00</p>
          </article>
          <article class="item">
            <p>Mocha</p><p>4.50</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
```

