---
id: 587d7fb2367417b2b2512bf5
title: クライアントからルートパラメータの入力を取得する
challengeType: 2
forumTopicId: 301513
dashedName: get-route-parameter-input-from-the-client
---

# --description--

APIを構築するとき、ユーザーがサービスから何を取得したいかを伝えられるようにする必要があります。たとえば、クライアントがデータベースに保存されているユーザーの情報をリクエストする場合、どのユーザーに関心があるかを知らせる方法が必要です。この目的を達成する一つの方法がルートパラメータの使用です。ルートパラメータはURLの名前付きセグメントで、スラッシュ（/）で区切られています。各セグメントはURLのその位置に一致する部分の値をキャプチャします。キャプチャされた値は`req.params`オブジェクトで見つけられます。

<blockquote>route_path: '/user/:userId/book/:bookId'<br>actual_request_URL: '/user/546/book/6754' <br>req.params: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

`GET /:word/echo`のルートにマウントされたエコーサーバーを構築してください。`{echo: word}`の構造を持つJSONオブジェクトで応答します。繰り返す単語は`req.params.word`で見つけられます。ブラウザのアドレスバーから、例えば`your-app-rootpath/freecodecamp/echo`のような一致するルートにアクセスしてルートをテストできます。

# --hints--

Test 1 : エコーサーバーが単語を正しく繰り返すことを確認してください

```js
  const response = await fetch(code + '/eChOtEsT/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'eChOtEsT',
    'Test 1: the echo server is not working as expected'
  );
```

Test 2 : エコーサーバーが単語を正しく繰り返すことを確認してください

```js
  const response = await fetch(code + '/ech0-t3st/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'ech0-t3st',
    'Test 2: the echo server is not working as expected'
  );
```
