---
id: 587d7fb0367417b2b2512bef
title: HTMLファイルを配信する
challengeType: 2
forumTopicId: 301516
dashedName: serve-an-html-file
---

# --description--

`res.sendFile(path)` メソッドを使うと、ファイルをレスポンスとして返せます。`app.get('/', ...)` ルートハンドラの中に入れて使えます。このメソッドは裏で、送信するファイルの種類に応じてブラウザがどう扱うか指示する適切なヘッダーを設定します。そしてファイルを読み込んで送信します。このメソッドには絶対パスのファイルパスが必要です。Nodeのグローバル変数 `__dirname` を使って次のようにパスを計算することをお勧めします。

```js
absolutePath = __dirname + '/relativePath/file.ext'
```

# --instructions--

GETリクエストで`/views/index.html` パスにアクセスされたとき、`/` ファイルをレスポンスとして返してください。ライブアプリを表示すると、大きなHTMLの見出し（あとで使うフォームもあります…）がスタイルなしで表示されるはずです。

**注意:** 前のチャレンジの解答を編集しても、新しく作成してもかまいません。新しく作る場合は、Expressはルートを上から順に評価し、最初にマッチしたハンドラを実行することを覚えておいてください。前の解答はコメントアウトしないと、サーバーは文字列のレスポンスを返し続けます。

# --hints--

アプリは views/index.html ファイルを配信する必要があります

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /<h1>.*<\/h1>/,
    'Your app does not serve the expected HTML'
  );
```
