---
id: 587d7fb0367417b2b2512bee
title: 動作するExpressサーバーを起動する
challengeType: 2
forumTopicId: 301519
dashedName: start-a-working-express-server
---

# --description--

ファイル`myApp.js`の最初の2行を見ると、Expressアプリオブジェクトを簡単に作成できることがわかります。このオブジェクトにはいくつかのメソッドがあり、これらのチャレンジで多くを学びます。基本的なメソッドの一つが`app.listen(port)`です。これはサーバーに指定したポートで待ち受けるよう指示し、サーバーを起動状態にします。テストのために、アプリをバックグラウンドで動作させる必要があるので、このメソッドを`server.js`ファイルに追加しました。

最初の文字列を配信してみましょう！Expressでは、ルートは次の構造を取ります：`app.METHOD(PATH, HANDLER)`。METHODは小文字のhttpメソッドです。PATHはサーバー上の相対パスで、文字列や正規表現も使えます。HANDLERはルートがマッチしたときにExpressが呼び出す関数です。ハンドラは`function(req, res) {...}`の形を取り、reqはリクエストオブジェクト、resはレスポンスオブジェクトです。例えば、次のハンドラは

```js
function(req, res) {
  res.send('Response String');
}
```

文字列 'Response String' を配信します。

# --instructions--

`app.get()`メソッドを使って、GETリクエストで`/`（ルート）パスにマッチするものに対し、文字列 "Hello Express" を配信してください。ログを見てコードが動作していることを確認しましょう。

**注意：** これらのレッスンのコードは、最初に用意された数行のコードの間に追加してください。

# --hints--

アプリは文字列 'Hello Express' を配信する必要があります

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.equal(
    data,
    'Hello Express',
    'Your app does not serve the text "Hello Express"'
  );
```
