---
id: 587d7fb1367417b2b2512bf1
title: 特定のルートでJSONを提供する
challengeType: 2
forumTopicId: 301517
dashedName: serve-json-on-a-specific-route
---

# --description--

HTMLサーバーがHTMLを提供するのに対し、APIはデータを提供します。<dfn>REST</dfn>（表現状態転送）APIは、クライアントがサーバーの詳細を知らなくても簡単にデータ交換できる仕組みです。クライアントはリソースの場所（URL）と、そのリソースに対して行いたい操作（動詞）だけを知っていれば十分です。GET動詞は、情報を取得する際に使われ、何も変更しません。現在、ウェブ上で情報をやり取りする際に好まれているデータ形式はJSONです。簡単に言うと、JSONはJavaScriptオブジェクトを文字列として表現する便利な方法で、簡単に送信できます。

`/json`のパスでJSONを返すルートを作成して、シンプルなAPIを作りましょう。通常通り、`app.get()`メソッドを使ってください。ルートハンドラの中で、`res.json()`メソッドを使い、引数としてオブジェクトを渡します。このメソッドはリクエスト・レスポンスループを閉じてデータを返します。内部では、有効なJavaScriptオブジェクトを文字列に変換し、ブラウザにJSONを提供していることを伝える適切なヘッダーを設定し、データを送信します。有効なオブジェクトは通常の構造`{key: data}`を持ちます。`data`は数値、文字列、ネストされたオブジェクト、または配列である場合があります。`data`は変数や関数呼び出しの結果である場合もあり、その場合は文字列に変換される前に評価されます。

# --instructions--

`{"message": "Hello json"}`ルートへのGETリクエストに対して、オブジェクト`/json`をJSON形式でレスポンスとして提供してください。その後、ブラウザで`your-app-url/json`にアクセスすると、画面にメッセージが表示されるはずです。

# --hints--

エンドポイント`/json`はJSONオブジェクト`{"message": "Hello json"}`を提供する必要があります。

```js
  const response = await fetch(code + '/json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.message,
    'Hello json',
    "The '/json' endpoint does not serve the right data"
  );
```
