---
id: 587d7fb0367417b2b2512bed
title: Nodeコンソールを使ってみよう
challengeType: 2
forumTopicId: 301515
dashedName: meet-the-node-console
---

# --description--

これらのチャレンジに取り組む際は、以下のいずれかの方法でコードを書いてください。

- <a href="https://github.com/freeCodeCamp/boilerplate-express/" target="_blank" rel="noopener noreferrer nofollow">このGitHubリポジトリ</a>をクローンして、ローカルでチャレンジを完了する。
- お好きなサイトビルダーを使ってプロジェクトを完成させる。このとき、GitHubリポジトリのすべてのファイルを必ず取り込んでください。

開発中は、コードの動作状況を確認できることが重要です。

NodeはJavaScriptの実行環境です。クライアント側のJavaScriptと同様に、コンソールを使って便利なデバッグ情報を表示できます。ローカルマシンでは、コンソールの出力はターミナルに表示されます。

これらのチャレンジに取り組む際は、ターミナルを開いたままにしておくことをお勧めします。ターミナルの出力を読むことで、発生したエラーを確認できます。

ファイルを変更したら、サーバーを再起動する必要があります。

ターミナルから`Ctrl + C`を使ってサーバーを停止し、`node mainEntryFile.js`を使ってNodeを直接起動するか、`package.json`ファイルの実行スクリプトを`npm run`で起動できます。

例えば、`"start": "node server.js"`スクリプトはターミナルで`npm run start`を使って実行します。

ファイル保存時にサーバーを自動再起動するには、Nodeの`--watch`フラグを起動スクリプト`"start": "node --watch server.js"`に追加するか、`nodemon`のようなnpmパッケージをインストールします。これは演習としてお任せします。

# --instructions--

`myApp.js`ファイルを修正して、コンソールに「Hello World」とログを出力してください。

# --hints--

`"Hello World"`がコンソールに表示されるはずです。

```js
  const response = await fetch(code + '/_api/hello-console');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.passed, '"Hello World" is not in the server console');
```
