---
id: 60a3e3396c7b40068ad69982
title: Passaggio 24
challengeType: 0
dashedName: step-24
---

# --description--

Imposta la `height` dell'elemento `.two` a 200 pixel.

# --hints--

Dovresti impostare la proprietà `height` su `200px`.

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '200px');
assert(hasHeight);
```

Il tuo elemento `.two` dovrebbe avere un valore `height` di `200px`.

```js
const twoHeight = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('height');
assert(twoHeight === '200px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
      </div>
    </div>
  </body>
</html>
```
