---
id: 60a3e3396c7b40068ad69970
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Prima di poter iniziare a stilizzare il `div` che hai aggiunto, devi collegare il tuo CSS al tuo HTML.

Aggiungi un elemento `link` per collegare il tuo file `styles.css`. Imposta il `href` su `styles.css` e ricorda di impostare l'attributo `rel` su `stylesheet`.

# --hints--

Il tuo codice dovrebbe avere un elemento `link`.

```js
assert(/<link/.test(code))
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
