---
id: 60a3e3396c7b40068ad6997f
title: Passaggio 21
challengeType: 0
dashedName: step-21
---

# --description--

Aggiungere 1 pixel di spaziatura interna in alto, in basso, a sinistra e a destra del canvas ha modificato le sue dimensioni a 502 pixel x 602 pixel.

Sostituisci la proprietà `padding` con `overflow` impostata su `hidden` — riportando il canvas alle sue dimensioni originali.

# --hints--

Dovresti rimuovere la proprietà `padding` dal selettore `.canvas`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

Dovresti impostare la proprietà `overflow` su `hidden`.

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

Il tuo elemento `.canvas` dovrebbe avere un valore `overflow` di `hidden`.

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
