---
id: 60a3e3396c7b40068ad6996e
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Il margine è l’area esterna al riquadro e può essere usata per controllare lo spazio tra altri riquadri o elementi.

Qui l’elemento in basso ha un margine superiore più grande, che lo spinge più in basso nella pagina.

Ora che hai rapidamente rivisto il modello a riquadro CSS, è il momento di iniziare con il dipinto di Rothko.

Rimuovi l’elemento `img`.

# --hints--

Non dovresti avere un elemento `img` nel tuo codice.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
