---
id: 60a3e3396c7b40068ad6996b
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Inizia il tuo progetto aggiungendo un elemento `img` con un `src` impostato su `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` all'interno dell'elemento `body`.

# --hints--

Il tuo codice dovrebbe avere un elemento `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

L'elemento `img` dovrebbe avere un `src` impostato su `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
