---
id: 60a3e3396c7b40068ad69977
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Usa la spaziatura interna per regolare lo spazio all'interno di un elemento.

In `.frame`, usa la proprietà abbreviata `padding` per aumentare lo spazio tra gli elementi `.frame` e `.canvas` di `50px`. La proprietà abbreviata aumenterà lo spazio in alto, in basso, a sinistra e a destra del bordo e del canvas dell'elemento.

# --hints--

Dovresti impostare la proprietà `padding` su `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

Il tuo elemento `.frame` dovrebbe avere un valore di `padding` pari a `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
