---
id: 60a3e3396c7b40068ad69974
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Ogni dipinto ha bisogno di una cornice.

Avvolgi l'elemento `.canvas` in un altro `div`. Dai a quel `div` la classe `frame`.
# --hints--

Dovresti aggiungere un nuovo elemento `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

Il tuo elemento `.canvas` dovrebbe essere annidato nel nuovo elemento `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

Il tuo nuovo `div` dovrebbe avere un `class` con il valore `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

Il tuo nuovo `div` dovrebbe essere all'interno del tuo elemento `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
  
  <div class="canvas">
  </div>
  
--fcc-editable-region--
  </body>
  ```
