---
id: 60a3e3396c7b40068ad69971
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

È il momento del CSS.

Anche se il tuo `<div>` non ha testo, viene comunque trattato come un riquadro con contenuto.  
Scrivi una regola CSS che usi il selettore di classe `.canvas` e imposta la sua proprietà `width` a 500 pixel.  
Ecco una regola CSS che imposta la larghezza della classe `card` a 300 pixel:

```css
.card {
  width: 300px;
}
```

# --hints--

Il tuo codice dovrebbe avere un selettore `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

Dovresti impostare la proprietà `width` su `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

Il tuo elemento `.canvas` dovrebbe avere una proprietà `width` di `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
