---
id: 60a3e3396c7b40068ad69975
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Scrivi una nuova regola usando il selettore di classe `.frame`.

Usa la dichiarazione abbreviata `border` per dare all'elemento `.frame` un bordo solido nero con una larghezza di `50px`.

# --hints--

Il tuo codice dovrebbe avere un selettore `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Dovresti impostare la proprietà `border` su `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

Il tuo elemento `.frame` dovrebbe avere una proprietà `border` con valore `50px solid black`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
