---
id: 60a3e3396c7b40068ad69975
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Write a new rule using the `.frame` class selector.

Usa la dichiarazione shorthand `border` per assegnare a `.frame` un bordo nero solido con una larghezza di `50px`.

# --hints--

Il codice dovrebbe avere un selettore `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Dovresti assegnare alla proprietà `border` il valore `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

Your `.frame` element should have a `border` property with a value of `50px solid black`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
