---
id: 65577236b056379d5dbc7000
title: 步骤-23
challengeType: 20
dashedName: step-23
---

# --description--

现在，你将要开始开发算法来计算在你的新图中每一个节点之间的最短路径。

声明一个 `shortest_path` 的空函数。 使用 `pass` 关键字填充函数主体。 

# --hints--

你应该定义一个名为 `shortest_path` 的函数。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
