---
id: 65579228c669fcbebffd01d5
title: 第 52 步
challengeType: 20
dashedName: step-52
---

# --description--

另一个错误则比较隐蔽。 当找到一条距离相邻节点比较短的路径时，`paths[current]` 将被分配给相邻节点路径，即 `paths[node]` 。 

这意味着两个变量指向相同的列表。 因为列表是可变的，当你将邻接节点附加到其路径时， `paths[node]` 和 `paths[current]` 都被修改，因为它们是相同的列表。 这导致路径错误，尽管距离是正确的。

你可以通过将一个 `paths[current]` 分配给邻居节点路径来修复这个错误。 为此，你可以使用切片语法：

```py
my_list[:]
```

`my_list` 是你想要复制的列表。 通过切分 `paths[current]` 来修改 `if` 块内的现有 `paths[node] = paths[current]` 赋值。

# --hints--

你应该将 `paths[current][:]` 分配给相邻节点路径。

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
