---
id: 655773b0591c5f9f4045883e
title: 步骤 24
challengeType: 20
dashedName: step-24
---

# --description--

算法将从指定的节点开始。 然后寻找图中起始节点或 *source* 与所有其他节点之间的最短路径。

所以，你的函数需要两个参数：`graph`, and `start` 。 将它们添加到你的函数声明中。

# --hints--

你的函数应该按照顺序接收 `graph` 和 `start` 作为参数。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
