---
id: 655791e6cf5e03be3de73451
title: 步骤 50
challengeType: 20
dashedName: step-50
---

# --description--

如果你尝试取消注释你的函数调用，它将不起作用。 你有几个bug需要修复。 第一个是在嵌套的 `if` 中，你尝试访问一个可能在 `paths[node]` 列表中不存在的元素。 所以，你需要在访问 `paths[node][-1]` 时确保 `paths[node]` 不为空。

在嵌套的 `if` 语句中添加附加条件，来确保你访问 `paths[node][-1]` 时 `paths[node]` 不为空。

# --hints--

你需要在嵌套的 `if` 语句中添加 `paths[node]` 作为第一个条件。 使用 `and` 操作符合并你的条件。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                   
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
