---
id: 65578d4fc3afc3b8f554c882
title: 步骤 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` 也需要一个只使用关键字的参数。 将一个函数作为附加参数传递到 `min()`, 你可以修改列表项的比较方式。

你刚刚在上一步写下的行的结果是按字母顺序排列的节点。 相反，你想选择距离起始节点最小的未访问节点。

将 `key=distances.get` 作为第二个参数传递给你的 `min()` 调用。 这样，比较将根据每个 `unvisited` 列表项目在 `distances` 字典中的值进行。

# --hints--

您应该将`key=distances.get`作为调用`min()`的第二个参数传递。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')

```
