---
id: 65576ff7888f9e96f52a4be1
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

字典存储数据的形式为 *key*-*value* 对。 键与对应值之间用冒号分隔。 每个键值对与以下键值对用逗号分隔：

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

将新的键值对添加到字典中。 使用字符串`'species'`作为键，字符串`'guinea pig'`作为值。

# --hints--

你应该在你的`copper`字典中有一个`'species'`key，其值为`'guinea pig'`。

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

你的`copper`字典应该只有一个键值对。

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
