---
id: 655773f8b8b5db9fc6d0ae76
title: 步骤 25
challengeType: 20
dashedName: step-25
---

# --description--

为了记录访问过的节点，你需要在图中所有节点的列表。 访问过一个节点，就把它从列表中删除。

现在，用 `unvisited` 变量替换 `pass` 关键字，并给它赋值一个空列表。

# --hints--

在 `shortest_path` 函数中，你应该有一个 `unvisited` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

你应该给 `unvisited` 变量赋值一个空列表。 记得使用 `pass`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
