---
id: 655771d889132f9ccd341060
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

当一个图的每一边都有定量比如距离，时间或者其他的定量值，那这个图叫做 *weighted* graph。

在此例子里定量值将是两个点的距离。 您可以把字典改成一串系的元组来代表一个定量值图。

每一个元组里第一个变量是连接的点，第二个变量是两点的距离。

将 `my_graph['A']` 修改为元组列表，考虑 `A-B` 距离为 `3` ，`A-D` 距离为 `1` 。

# --hints--

`my_graph["A"]` 应该是一个包含元组 `('B', 3)` 和 `('D', 1)` 的列表。

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` 应该有4个键值名为 `'A'`、 `'B'`、 `'C'` 和 `'D'`。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
