---
id: 655774955b097ea14897db12
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

与起始节点的距离为零，因为算法从那里开始评估。

你在循环中将 `node` 添加到 `unvisited` 之后，创建一个 `if` 语句，当节点等于起始节点时触发。 然后在 `distances` 字典内将 `0` 分配给该节点。

# --hints--

你应该创建一个当 `node` 等于 `start` 时执行的 `if` 语句。

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

在你的新 `if` 中，你应该将 `0` 分配给 `distances` 字典中的节点。

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
