---
id: 65578d0f6c78a0b868a43b9c
title: 步骤 40
challengeType: 20
dashedName: step-40
---

# --description--

在 `while` 循环内，首先要做的是定义要访问的当前节点。 为此你可以使用 `min()` 函数。 它从作为参数传递的可迭代对象中返回最小的项。

移除 `pass` ，然后创建一个名为 `current` 的变量，并将其赋值为 `min(unvisited)`。

# --hints--

你应该在 `while` 循环内创建一个 `current` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

你应该将 `min(unvisited)` 赋值给你的 `current` 变量。 记得使用 `pass`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
