---
id: 6557709b0aee699a6a00528c
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

例如，图形可用于表示空间中由路径连接的两个点 `A` 和 `B`。 像这样的图将由两个通过边连接的节点组成。

用字符串 `'A'` 替换现有的 `'species'` 键。 然后，将 correspondent 的值替换为字符串 `'B'` 来表示 `'A'` 和 `'B'` 两个节点之间的连接。

# --hints--

你的字典应该有一个 `'A'` 键。

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

你的 `'A'` 键应以 `'B'` 作为值。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

你的字典应该有两个键。

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
