---
id: 6557746aad2844a0cd864e12
title: 步骤 27
challengeType: 20
dashedName: step-27
---

# --description--

虽然算法能搜索图，但它应该记录起始节点和其他节点之间当前已知的最短距离。

在你的 `for` 循环之前，创建一个新的变量 `distances`，并赋值为一个空字典。

# --hints--

你应该有一个变量 `distances` 。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

你的变量 `distances` 应该是一个空字典。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
