---
id: 65577791ad8c26a7705e2919
title: 步骤 34
challengeType: 20
dashedName: step-34
---

# --description--

有了字典理解，你可以从现有词典开始创建一个字典：

```py
{key: val for key in dict}
```

在上面的示例中， `val` 是 `key` 在新字典中的值。 `dict` 是现有字典。

你想要跟踪起始节点和其他节点之间的路径。

在 `distances` 变量后面，创建一个 `paths` 变量并分配一个包含 `graph` 所有键的字典。 将一个空列表分配给每个键，并使用字典理解来构建你的字典。

# --hints--

你应该有一个名为 `paths` 的变量。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

你的 `paths` 变量应该使用字典理解语法为图中的每个节点分配一个空列表。

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
