---
id: 65577739f57ecca6c39bb4e9
title: 步骤 33
challengeType: 20
dashedName: step-33
---

# --description--

`list()` 类型构造函数使您能够从可迭代对象构建列表。

使用 `list()` 修改 `unvisited` 变量的赋值，并传递 `graph` 作为可迭代对象。

# --hints--

你应该使用 `list()` 从 `graph` 字典生成一个列表。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/list\s*\(\s*graph\s*\)/));
  }
})
```

你应该将 `list(graph)` 分配给 `unvisited` 。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\s*=\s*list\s*\(\s*graph\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
