---
id: 6557716aadbd2d9c42c0e69a
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

添加最后一个节点 `'D'` ，它连接 `'A‘` 和 `'C'`。

修改你的字典以代表此结构。 再次使用列表来代表多个连接。

# --hints--

你的字典需要包含四个键，分别叫做 `'A'`、 `'B'`、 `'C'` 和`'D'`。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` 应该是一个列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["A"]) is list
  `))
})
```

`my_graph['A']` 应该是一个包含 `'B'` 和 `'D'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["A"]) == 2 and "B" in my_graph["A"] and "D" in my_graph["A"]
  `))
})
```

`my_graph['B']` 应该是一个列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']` 应该是一个包含 `'A'` and `'C'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']` 应该是一个列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["C"]) is list
  `))
})
```

`my_graph['C']` 应该是一个包含 `'B'` 和 `'D'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["C"]) == 2 and "B" in my_graph["C"] and "D" in my_graph["C"]
  `))
})
```

`my_graph['D']` 应该是一个列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["D"]) is list
  `))
})
```

`my_graph['D']` 应该是一个包含 `'A'` 和 `'C'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["D"]) == 2 and "A" in my_graph["D"] and "C" in my_graph["D"]
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': ['A', 'C'],
    'C': 'B'
}
--fcc-editable-region--
```
