---
id: 655774d01daeeaa1978b99d5
title: 步骤 29
challengeType: 20
dashedName: step-29
---

# --description--

在开始时，因为距离尚未确定，所以图中的所有节点都被认为与源节点有无限距离。

创建一个 `else` 语句，并在 `distances` 字典中为节点分配一个无限的值。 为此，将字符串 `'inf'` 作为参数调用 `float()` 函数，生成一个代表正无穷的浮点数。

# --hints--

你应该有一个 `else` 子句。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

你应该在 `else` 语句中将 `float('inf')` 赋值给 `distances[node]` 。

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
