---
id: 6557712d77ce2d9bd7e63afd
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

在你的图中添加另一个节点连接 `B` 节点，并将其命名为 `C` 。

修改你现有的字典来表示这种排列方式：在 `my_graph` 中添加另一个键 `'C'` 并为其赋值为字符串 `'B'` 。

另外，更改现有的 `'B'` 键的值为列表 `['A', 'C']` 来表示 `'B'` 节点的多个连接。

# --hints--

你的字典应该包含三个键 - `'A'`、 `'B'`、 和 `'C'` 。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` 的值应该是字符串 `'B'` 。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` 应该是一个列表。

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']` 的值应该是包含 `'A'` 和 `'C'` 的列表。

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']` 的值应该是字符串 `'B'` 。

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
