---
id: 65578cb031cd93b77a285db2
title: 步骤 38
challengeType: 20
dashedName: step-38
---

# --description--

你的函数将探索连接到起始节点的所有节点。 它将计算所有路径的最短路径。 然后，它将从未访问过的节点中删除起始节点。 

接下来，将访问最近的邻节点，并在访问所有节点之前重复该进程。

从现在起，你将在探索图中的节点的主循环上工作。 为了避免在算法开发过程中运行无限循环的问题，将函数调用转化为注释。

# --hints--

你应该将函数调用变成注释。

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
