---
id: 65577a17564ce8a8e06c1460
title: 步骤35
challengeType: 20
dashedName: step-35
---

# --description--

字典推导式也支持条件`if`/`else` 语法：

```py
{key: val_1 if condition else val_2 for key in dict}
```

在上面的例子中，`dict`是现有的字典。 当`condition`的计算结果为`True`时，`key`的值将为`val_1`，否则`val_2`。

使用字典理解来基于`graph`创建字典，并将其分配给`distances`变量。 如果节点等于起始节点，将键的值设为零，否则设为无穷大。 使用 `float('inf')` 来实现后者。

# --hints--

你应该使用字典推导式的语法来为你的`distances`变量赋值。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
