---
id: 6566195b0a021bb660b2b4b1
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

现在，将 `my_graph['B']` 修改为元组列表，其中元组中的第一个元素是连接的节点，第二个元素是距离。 `B-C` 的距离为 `4`。

# --hints--

`my_graph['B']` 应该是一个包含元组 `('A', 3)` 和 `('C', 4)` 的列表。

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph` 应该有4个键值名为 `'A'`、 `'B'`、 `'C'` 和 `'D'`。

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
