---
id: 6578b13757611e2825beb8a5
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

键在字典中必须是唯一的，并且只能是不可变的数据类型。 这意味着你不能将列表或其他字典用作键。

在字典中添加另一个键`'age'`，并将整数`2`作为值。

# --hints--

你应该在你的`copper`字典中，在`'species': 'guinea pig'`之后有一个`'age'`key，其值为`2`。 别忘了逗号。

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

你的`copper`字典中应该有两个键值对。

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
