---
id: 65789506b30453080f77470c
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

到目前为止，你已经遇到了不同的数据类型：

- 不可变的数据类型，如整数、字符串、元组和布尔值。
- 可变数据类型，例如列表和字典。

字典由一对花括号`{}` 标识。

首先创建一个变量 `copper` ，并使用一对花括号为其分配一个空字典，就像使用一对方括号创建一个空列表一样。

# --hints--

你应该有一个名为 `copper` 的变量。

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

你的 `copper` 变量应该有一个空字典的值。 为此使用一对花括号。

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
