---
id: 6578b57361f2f132a02e2a18
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

您可以通过字典的键访问存储在字典中的数据：

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

在字典之后，按照上面的示例访问`copper`的`'species'`key并打印结果。

# --hints--

你不应该修改字典的分配。

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

你应该使用`copper['species']`来访问`'species'`key的值。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

你应该调用`print()`，将`copper['species']`作为参数传递。

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
