---
id: 6579cbab9825b8170974c69a
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

你可以使用`del`关键字从字典中删除键值对：

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

就在你的`for`循环之前，使用`del`>关键字从`copper`中删除`'age'`键及其值。

# --hints--

你应该使用`del`关键字删除`for`循环之前的`copper['age']`。

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
