---
id: 6579cd5f6dd62c189e53ddbb
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

现在你已经复习了字典的基本方面，可以继续构建最短路径算法。

删除`copper`字典声明后的每一行代码。

# --hints--

你应该删除字典声明后的所有代码行。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
