---
id: 6579dd49fa8a8e1fd06b85a9
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

图表是代表各单元之间关系的数据结构。 这些被称为 *nodes*的元素可以是真实的物体、实体、空间中的点数或其他元素。 节点之间的连接称为 *edges*。

下面是图形的可视化表现：

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


将`copper`字典重命名为`my_graph`。 这将代表图来测试你的算法。

# --hints--

你应该将你的`copper`字典重命名为`my_graph`。

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

你的`my_graph`变量应该是一个字典。

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
