---
id: 6579717f0920131304286804
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

想要在字典声明之后添加新的键值对，你可以通过和访问已有的键值相同的方式声明键，并使用赋值操作符设置新键的值：

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

删除你的 `print()` 调用。 然后，在 `copper` 声明之后，添加键 `'food'` 到字典，并赋值为 `'hay'` 。

# --hints--

你不应该在你的代码里出现 `print(copper['age'])` 。

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

你应该在字典声明之后添加键 `'food'` 到 `copper` 。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

你应该在字典声明之后设置 `copper['food']` 的值为 `'hay'` 。

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
