---
id: 6579ca0923cfa7162089d2f0
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

相同的语法可用于更改现有键的值。

就在`print()`调用之前，访问`'species'`键，并将其值重新分配给`'Cavia porcellus'`。

# --hints--

在调用`print()`之前，你应该将`copper['species']`重新分配给字符串`'Cavia porcellus'`。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
