---
id: 65797670e0c0d016f17e7660
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

现在，在你的代码底部，打印 `copper` 。

# --hints--

你应该在代码底部打印 `copper`。

```js
({ test: () =>
  {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /print\s*\(\s*copper\s*\)\s*$/)
  }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

--fcc-editable-region--
```
