---
id: 65affe602a2a828b0382a058
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

将 `'age'` 键替换为字符串 `'B'`，并将其值设置为字符串 `'A'`，以表示两个方向上的节点之间的连接。

# --hints--

你的字典应该有一个 `'B'` 键。

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

你的 `'B'` 键应以 `'A'` 作为值。

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

你的字典应该有一个 `'A'` 键。

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

你的 `'A'` 键应以 `'B'` 作为值。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

你的字典应该有两个键。

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
