---
id: 65b7d3319ebec69b983fb91b
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

正如你从输出中看到的，`.items()` 创建了一个数据结构，将每个键值点对保存为一个独立的元组。要迭代这些元组中的元素，你可以添加第二个循环变量：

```py
for i, j in dict.items():
    print(i, j)
```

修改你的 `for` 循环以采用两个循环变量并将它们两个都打印在循环体内。

# --hints--

你应该修改 `for` 循环以使用两个变量。

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

你应该在 `for` 循环内打印两个循环变量。

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
