---
id: 65b7cf140d34058d7ea8935f
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

如果你想迭代字典键的值，一种方法是使用 `.values()` 方法。

修改 `for` 循环以迭代 `copper.values()` 而不是 `copper`，并查看输出。

# --hints--

你应该修改 `for` 循环以迭代 `copper.values()` 而不是 `copper`。

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
