---
id: 68e5293bd00d2fe134f58984
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Ora è il momento di creare un'implementazione concreta della tua classe astratta. Crea una classe chiamata `PercentageDiscount` che eredita da `DiscountStrategy`.

Dalle un metodo `__init__` che prende `self` e `percent` come parametri. Dai un suggerimento di tipo `int` a `percent` e imposta il tipo di ritorno del metodo su `None`. All'interno del metodo `__init__`, archivia `percent` come attributo dell'istanza.

# --hints--

Dovresti creare una classe chiamata `PercentageDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

La tua classe `PercentageDiscount` dovrebbe ereditare da `DiscountStrategy`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

La tua classe `PercentageDiscount` dovrebbe avere un metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

Il tuo metodo `__init__` dovrebbe avere i parametri `self` e `percent: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

Il tuo metodo `__init__` dovrebbe avere un suggerimento di tipo di ritorno `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

Dovresti assegnare `percent` a `self.percent` all'interno del tuo metodo `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
