---
id: 68e5293bd00d2fe134f58983
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Ora aggiungi un secondo metodo astratto chiamato `apply_discount` alla tua classe `DiscountStrategy`. Usa il decoratore `@abstractmethod` e dagli i parametri `self` e `product`.

Aggiungi gli hint di tipo: `product` dovrebbe essere di tipo `Product`, e il metodo dovrebbe restituire un `float`. Usa `pass` nel corpo del metodo.

# --hints--

La tua classe `DiscountStrategy` dovrebbe avere un metodo chiamato `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

Il tuo metodo `apply_discount` dovrebbe avere il decoratore `@abstractmethod`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

Il tuo metodo `apply_discount` dovrebbe avere i parametri `self` e `product: Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

Il tuo metodo `apply_discount` dovrebbe avere un hint di tipo di ritorno `-> float`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
