---
id: 68e5293bd00d2fe134f58981
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Crea una classe chiamata `DiscountStrategy` che eredita da `ABC`. Questa sarà una classe base astratta che definisce l'interfaccia per tutte le strategie di sconto.

Per ora, usa semplicemente `pass` nel corpo della classe per creare una classe vuota.

# --hints--

Dovresti creare una classe chiamata `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

La tua classe `DiscountStrategy` dovrebbe ereditare da `ABC`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
