---
id: 68e52994bc7ea2e3dec530a6
title: Passaggio 29
challengeType: 20
dashedName: step-29
---

# --description--

Ora testiamo il sistema completo di sconti! Aggiungi un'istruzione `if` che verifica se `__name__ == '__main__'` e sposta `product = Product('Wireless Mouse', 50.0)` al suo interno.

Dopodiché, all'interno del blocco `if`, crea una variabile `user_tier` e impostala sulla stringa `Premium`. Poi crea una variabile `strategies` impostata su una lista di strategie che contiene `PercentageDiscount(10)`, `FixedAmountDiscount(5)` e `PremiumUserDiscount()`.

# --hints--

Dovresti avere un'istruzione `if` che verifica se `__name__ == '__main__'`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent('__name__ == "__main__"')`)
})
```

Dovresti spostare `product = Product('Wireless Mouse', 50.0)` all'interno del blocco `if`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt("product = Product('Wireless Mouse', 50.0)")`)
})
```

Dovresti impostare `user_tier` sulla stringa `Premium` all'interno del blocco `if`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('user_tier = "Premium"')`)
})
```

Dovresti creare una lista chiamata `strategies` contenente `PercentageDiscount(10)`, `FixedAmountDiscount(5)` e `PremiumUserDiscount()`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('strategies = [PercentageDiscount(10), FixedAmountDiscount(5), PremiumUserDiscount()]')`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

class DiscountEngine:
    def __init__(self, strategies: list[DiscountStrategy]) -> None:
        self.strategies = strategies

    def calculate_best_price(self, product: Product, user_tier: str) -> float:
        prices = [product.price]

        for strategy in self.strategies:
            if strategy.is_applicable(product, user_tier):
                discounted = strategy.apply_discount(product)
                prices.append(discounted)

        return min(prices)
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
--fcc-editable-region--
```
