---
id: 68e529226b0defdff26c4a8a
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora, crea un'istanza di `Product` con il nome `Wireless Mouse` e un prezzo di `50.0`. Assegna questa istanza a una variabile chiamata `product`. Poi, stampa `product` nella console per vedere com'è.

# --hints--

Dovresti avere una variabile chiamata `product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

Dovresti assegnare un'istanza di `Product` alla variabile `product`.

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

La tua istanza `Product` dovrebbe avere il nome `Wireless Mouse` e il prezzo `50.0`.

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

Dovresti stampare `product` nella console.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

Quando stampi un `Product` con nome `'Wireless Mouse'` e prezzo `50.0`, dovrebbe mostrare `Wireless Mouse - $50.0`.

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
