---
id: 68e5293bd00d2fe134f58987
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Ora creerai un'altra strategia di sconto. Crea una classe chiamata `FixedAmountDiscount` che eredita da `DiscountStrategy`.

Dalle un metodo `__init__` che prende `self` e `amount` come parametri. Dai al parametro `amount` un tipo indicativo `int` e un tipo di ritorno `None` per il metodo. All'interno del metodo `__init__`, archivia `amount` come attributo dell'istanza.

# --hints--

Dovresti creare una classe chiamata `FixedAmountDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

La tua classe `FixedAmountDiscount` dovrebbe ereditare da `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

La tua classe `FixedAmountDiscount` dovrebbe avere un metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

Il tuo metodo `__init__` dovrebbe avere i parametri `self` e `amount: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

Il tuo metodo `__init__` dovrebbe avere un tipo di ritorno `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

Dovresti assegnare `amount` a `self.amount` all'interno del tuo metodo `__init__`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
