---
id: 68e529226b0defdff26c4a8b
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Ora inizierai a buildare l'interfaccia del sistema di sconti. Python offre un modo per definire classi astratte che fungono da modelli per altre classi. Questo si fa usando la classe `ABC` (Abstract Base Class) del modulo `abc`.

All'inizio del tuo codice, aggiungi la seguente istruzione di importazione:

```py
from abc import ABC, abstractmethod
```

Questa importa la classe `ABC` e il decoratore `abstractmethod`, che userai per definire i metodi che devono essere implementati da qualsiasi classe che eredita dalla tua classe astratta.

# --hints--

Dovresti avere l'istruzione di importazione menzionata all'inizio del tuo codice.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
