---
id: 68e5293bd00d2fe134f58988
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

Implementa il metodo `is_applicable` per la tua classe `FixedAmountDiscount`. Uno sconto a importo fisso dovrebbe essere applicato solo se ha senso dal punto di vista economico.

Restituisci `True` se `product.price * 0.9` è maggiore di `self.amount`, altrimenti restituisci `False`. Questo garantisce che lo sconto sia inferiore al 90% del prezzo originale, quindi non è irrealistico.

# --hints--

La tua classe `FixedAmountDiscount` dovrebbe avere un metodo `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("is_applicable")`)
})
```

Il tuo metodo `is_applicable` dovrebbe avere i parametri `self`, `product: Product` e `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_args("self, product:Product, user_tier:str")`)
})
```

Il tuo metodo `is_applicable` dovrebbe avere un tipo di ritorno `-> bool`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

Il tuo metodo `is_applicable` dovrebbe restituire `True` quando `product.price * 0.9` è maggiore di `self.amount`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(50)
    assert discount.is_applicable(test_product, "regular") == True
  `)
})
```

Il tuo metodo `is_applicable` dovrebbe restituire `False` quando `product.price * 0.9` è minore o uguale a `self.amount`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(95)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
