---
id: 68e5293bd00d2fe134f58985
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Poiché `PercentageDiscount` eredita da una classe astratta, devi implementare tutti i metodi astratti. Inizia implementando il metodo `is_applicable`. Il metodo `is_applicable` dovrebbe prendere `self`, `product: Product` e `user_tier: str` come parametri. Dovrebbe avere un tipo di ritorno indicato come `bool`.

Questo metodo deve determinare se uno sconto percentuale è valido. Restituisci `True` se `self.percent` è minore o uguale a `70`, altrimenti restituisci `False`. Questo evita sconti irrealistici come il 90% di sconto.

# --hints--

La tua classe `PercentageDiscount` dovrebbe avere un metodo `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

Il tuo metodo `is_applicable` dovrebbe avere i parametri `self`, `product: Product` e `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

Il tuo metodo `is_applicable` dovrebbe restituire `True` quando `self.percent` è minore o uguale a `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

Il tuo metodo `is_applicable` dovrebbe restituire `False` quando `self.percent` è maggiore di `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

Il tuo metodo `is_applicable` dovrebbe avere un tipo di ritorno indicato come `bool`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
