---
id: 68e529226b0defdff26c4a89
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Al momento, se crei un'istanza di `Product` e la stampi, vedrai una rappresentazione predefinita dell'oggetto che non è molto informativa. Per esempio, potrebbe apparire qualcosa come `<__main__.Product object at 0x7f9b8c2d1d30>`.

Per trasformare la rappresentazione predefinita dell'oggetto in una stringa leggibile, crea un metodo `__str__` nella tua classe `Product` che restituisca una stringa con il formato `name - $price` (dove `name` e `price` devono essere sostituiti con i valori degli attributi corrispondenti).

# --hints--

La tua classe `Product` dovrebbe avere un metodo `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

Il tuo metodo `__str__` dovrebbe restituire una stringa nel formato `name - $price`.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
