---
id: 68e5293bd00d2fe134f58986
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Ora implementa il metodo `apply_discount` nella tua `PercentageDiscount`.

Questo metodo dovrebbe calcolare e restituire il prezzo scontato.

La formula è: `product.price * (1 - self.percent / 100)`

# --hints--

La tua classe `PercentageDiscount` dovrebbe avere un metodo `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("apply_discount")`)
})
```

Il tuo metodo `apply_discount` dovrebbe avere i parametri `self` e `product: Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_args("self, product: Product")`)
})
```

Il tuo metodo `apply_discount` dovrebbe avere un tipo di ritorno `float`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_returns("float")`)
})
```

Il tuo metodo `apply_discount` dovrebbe calcolare il prezzo scontato usando la formula `product.price * (1 - self.percent / 100)`.

```js
({
  test: () => runPython(`
    test_product = Product("Test", 100.0)
    discount = PercentageDiscount(20)
    result = discount.apply_discount(test_product)
    assert result == 80.0
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
