---
id: 68e524c94e2e4ab7b4bfc7e2
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop costruirai un calcolatore di sconti che può applicare diverse strategie di sconto ai prodotti. Il sistema determinerà il prezzo migliore per un cliente basandosi su regole di sconto multiple.

Inizia definendo una classe chiamata `Product`. Dai alla tua classe un metodo `__init__` con i seguenti parametri: `self`, `name` e `price`. All'interno del metodo `__init__`, assegna i parametri `name` e `price` ad attributi dell'istanza con lo stesso nome.

# --hints--

Dovresti creare una classe chiamata `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

La tua classe `Product` dovrebbe avere un metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

Il tuo metodo `__init__` dovrebbe avere i seguenti parametri: `self`, `name` e `price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

Dovresti assegnare `name` a `self.name` all'interno del tuo metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

Dovresti assegnare `price` a `self.price` all'interno del tuo metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
