---
id: 68e5293bd00d2fe134f58982
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Come discusso nella lezione precedente, i metodi astratti sono metodi che devono essere implementati da qualsiasi classe che eredita dalla classe base astratta. Servono come contratto, garantendo che tutte le strategie di sconto abbiano determinati metodi obbligatori.

All'interno della tua classe `DiscountStrategy`, definisci un metodo chiamato `is_applicable`. Usa il decoratore `@abstractmethod` sopra la definizione del metodo. Il metodo deve avere i seguenti parametri: `self`, `product` e `user_tier`. Usa l'istruzione `pass` nel corpo del metodo.

Aggiungi suggerimenti di tipo per rendere chiara la firma del metodo: `product` dovrebbe essere di tipo `Product`, `user_tier` dovrebbe essere di tipo `str`.


# --hints--

La tua classe `DiscountStrategy` dovrebbe avere un metodo chiamato `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("is_applicable")`)
})
```

Il tuo metodo `is_applicable` dovrebbe avere il decoratore `@abstractmethod`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_decorators("abstractmethod")`)
})
```

Il tuo metodo `is_applicable` dovrebbe avere i parametri `self`, `product: Product` e `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
