---
id: 68e79d91daf8fb4808f93c47
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

Dopo aver implementato il metodo `apply_discount`, testa la tua strategia di sconto. Crea un'istanza di `PercentageDiscount` con `10` percento e assegnala a una variabile chiamata `discount`. Poi applica lo sconto usando `discount.apply_discount(product)` e stampa il risultato per vedere lo sconto in azione.

# --hints--

Dovresti creare una variabile chiamata `discount` con un'istanza di `PercentageDiscount`.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

La tua istanza `PercentageDiscount` dovrebbe essere inizializzata con `10` percento.

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

Dovresti chiamare `discount.apply_discount(product)` e stampare il risultato.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
