---
id: 68e9207ddcba053c5cb24413
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

In Python, i type hint sono segnali opzionali che indicano ad altri sviluppatori quale tipo di dato ci si aspetta per una variabile o una funzione. Ecco un esempio che aggiunge il type hint `str` al parametro `name`.

```py
def greet(name: str):
    return 'Hello, ' + name
```

Per i parametri del metodo `__init__`, aggiorna il parametro `name` aggiungendo il type hint `str` e il parametro `price` aggiungendo il type hint `float`.

# --hints--

Dovresti aggiungere il type hint `str` al parametro `name` e il type hint `float` al parametro `price` nel metodo `__init__` della classe `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
