---
id: 68f119ad905b29ff59d858c8
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Proprio come nel passaggio precedente, nel tuo metodo `__str__` aggiungi un'indicazione del tipo di ritorno `str` dato che questo metodo restituisce una rappresentazione stringa dell'oggetto.

# --hints--

Il tuo metodo `__str__` dovrebbe avere un'indicazione del tipo di ritorno `str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
