---
id: 668f08ea07b99b1f4a91acab
title: Створіть сторінку з рецептом
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Історія користувача:**

1. Ви повинні мати оголошення `!DOCTYPE html`.
1. Ви повинні мати елемент `html` з `lang` зі значенням `en`.
1. Ви повинні мати елемент `head`, який містить елемент `title` з назвою рецепта та елемент `meta` з атрибутом `charset` зі значенням `UTF-8`.
1. Ви повинні мати елемент `body`.
1. Ви повинні мати елемент `h1` з назвою рецепта.
1. Під `h1` має бути елемент `p`, який коротко описує рецепт.
1. Ви повинні мати елемент `h2` з текстом `Ingredients` для розділу з інгредієнтами.
1. Під першим елементом `h2` має бути невпорядкований список (елемент `ul`) з принаймні чотирма елементами списку (елементи `li`), який містить інгредієнти.
1. Ви повинні мати ще один елемент `h2` з текстом `Instructions` для розділу з покроковою інструкцією.
1. Під другим елементом `h2` має бути впорядкований список (елемент `ol`) з принаймні чотирма елементами списку, який містить покрокову інструкцію.
1. Ви повинні мати один елемент `img` з атрибутом `src` зі значенням дійсного зображення (можна використати `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg`) та атрибутом `alt`, що описує зображення.

# --hints--

Сторінка з рецептом має містити оголошення `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Ви повинні мати елемент `html` з `lang` зі значенням `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

Ви повинні мати елемент `head` в межах елемента `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Ви повинні мати елемент `title` в межах елемента `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

Елемент `title` має містити назву рецепта.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Ви повинні мати елемент `meta` в межах елемента `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

Елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Ви повинні мати елемент `body` в межах елемента `html`.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

Ви повинні мати елемент `h1` з назвою рецепта.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Ви повинні мати лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Під елементом `h1` має бути елемент `p`.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

Перший елемент `p` має описувати рецепт.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Перший елемент `h2` повинен містити текст `Ingredients`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

Під першим елементом `h2` має бути невпорядкований список.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

Невпорядкований список з інгредієнтами повинен містити принаймні чотири елементи.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Другий елемент `h2` повинен містити текст `Instructions`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

Під другим елементом `h2` має бути впорядкований список.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

Впорядкований список з інструкцією повинен містити принаймні чотири елементи.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Ви повинні мати принаймні один елемент `img`.

```js
assert.exists(document.querySelector('img'));
```

Всі елементи `img` повинні мати дійсний атрибут `src` та його значення.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

Всі елементи `img` повинні мати атрибут `alt` з описом зображення.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
