---
id: 68e6698b29e562d21b1d260c
title: 단계 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 주요 웹 브라우저 목록을 만들 것입니다. HTML 보일러플레이트가 제공되어 있습니다.

`List of Major Web Browsers` 요소 안에 `h1` 요소를 사용하여 `body`라는 제목을 페이지에 추가하는 것부터 시작하세요.

# --hints--

`h1` 요소 안에 `body` 요소가 있어야 합니다.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

`h1` 요소에는 `List of Major Web Browsers` 텍스트가 포함되어야 합니다.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
