---
id: 68ecfa1003d88990bb96a33d
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

설명 목록은 용어와 세부 정보 목록으로 구성됩니다. `dt` 요소는 설명 용어를 나타내고, `dd` 요소는 설명 세부 정보를 나타냅니다.

설명 목록에 처음 추가할 브라우저는 Google Chrome입니다.

`dl` 요소 안에 `dt` 요소를 추가하고, 그 안에 `Google Chrome` 텍스트를 넣으세요.

`dt` 요소 아래에 다음 텍스트가 포함된 `dd` 요소를 추가하세요:

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

`dt` 요소 안에 `dl` 요소가 있어야 합니다.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

`dt` 요소에는 `Google Chrome` 텍스트가 있어야 합니다.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

`dd` 요소 아래에 `dt` 요소가 있어야 합니다.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

`dd` 요소에는 `This is a free web browser developed by Google and first released in 2008.` 텍스트가 포함되어야 합니다.

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```
