---
id: 68ecfa1003d88990bb96a33e
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

이제 두 번째 브라우저를 목록에 추가할 차례입니다.

`Google Chrome` 항목 아래에 `dt` 요소를 하나 더 추가하고 텍스트 `Firefox`를 포함하세요.

`dt` 요소 아래에 다음 텍스트가 들어간 `dd` 요소를 추가하세요:

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

코드에 `dt` 요소가 두 개 있어야 합니다.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

두 번째 `dt` 요소에는 텍스트 `Firefox`가 있어야 합니다.

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

코드에 `dd` 요소가 두 개 있어야 합니다.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

두 번째 `dd` 요소에는 텍스트 `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`가 포함되어야 합니다.

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
