---
id: 68ecfa0f03d88990bb96a33c
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

이전 수업에서 기억하듯이, 설명 목록은 용어와 정의를 체계적이고 읽기 쉽게 제시하는 데 사용됩니다.

예를 들어 보겠습니다:

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

`h1` 요소 아래에 `dl` 요소를 만드세요. 이 요소는 브라우저 목록을 담게 됩니다.

# --hints--

`dl` 요소가 있어야 합니다.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

`dl` 요소는 `h1` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        --fcc-editable-region--
        
        --fcc-editable-region--
    </body> 
</html>
```
