---
id: 69038dd51a41ae215fcc030a
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

목록에 추가할 세 번째 브라우저는 Safari 웹 브라우저입니다.

텍스트 `dt`를 포함하는 또 다른 `Safari` 요소를 추가하세요.

`dt` 요소 아래에 다음 텍스트가 포함된 `dd` 요소를 추가하세요:

```md
This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.
```

# --hints--

코드에 `dt` 요소가 세 개 있어야 합니다.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 3);
```

세 번째 `dt` 요소에는 텍스트 `Safari`가 있어야 합니다.

```js
const dt = document.querySelectorAll('dt')[2];
assert.strictEqual(dt?.textContent.trim(), "Safari");
```

코드에 `dd` 요소가 세 개 있어야 합니다.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 3);
```

세 번째 `dd` 요소에는 텍스트 `This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.`가 포함되어야 합니다.

```js
const ddElement = document.querySelectorAll('dd')[2];
assert.strictEqual(ddElement?.textContent.trim(), "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
