---
id: 69038fe7da59ca23a5eba82d
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

목록에 추가할 마지막 브라우저는 Arc 브라우저입니다.

텍스트 `dt`를 포함하는 또 다른 `Arc` 요소를 추가하세요.

`dt` 요소 아래에 다음 텍스트가 있는 `dd` 요소를 추가하세요:

```md
This is a free Chromium based web browser first released in 2023 by The Browser Company.
```

이 마지막 추가로 브라우저 목록이 완성됩니다!

# --hints--

코드에 `dt` 요소가 다섯 개 있어야 합니다.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 5);
```

다섯 번째 `dt` 요소에는 텍스트 `Arc`가 있어야 합니다.

```js
const dt = document.querySelectorAll('dt')[4];
assert.strictEqual(dt?.textContent.trim(), "Arc");
```

코드에 `dd` 요소가 다섯 개 있어야 합니다.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 5);
```

다섯 번째 `dd` 요소에는 텍스트 `This is a free Chromium based web browser first released in 2023 by The Browser Company.`가 포함되어야 합니다.

```js
const ddElement = document.querySelectorAll('dd')[4];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free Chromium based web browser first released in 2023 by The Browser Company.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            <dt>Brave</dt>
            <dd>This is a free web browser first released in 2016 that is based on the Chromium web browser.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>List of Browsers and Descriptions</title>
</head>
<body>
    <h1>List of Major Web Browsers</h1>
    <dl>
        <dt>Google Chrome</dt>
        <dd>This is a free web browser developed by Google and first released in 2008.</dd>

        <dt>Firefox</dt>
        <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

        <dt>Safari</dt>
        <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

        <dt>Brave</dt>
        <dd>This is a free web browser first released in 2016 that is based on the Chromium web browser.</dd>

        <dt>Arc</dt>
        <dd>This is a free Chromium based web browser first released in 2023 by The Browser Company.</dd>
    </dl>
</body>
</html>
```
