---
id: 61fd75ea7f663457612dba02
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

새로운 `th` 요소 각각 안에 `span`가 `class`로 설정된 `sr-only year` 요소를 중첩하세요. 다음 텍스트를 순서대로 넣으세요: `2019`, `2020`, 그리고 `2021`.

세 번째 `th` 요소에는 `class` 속성을 `current`로 설정하세요.

`td` 요소는 비워 두세요. 이 요소는 표가 4열 레이아웃을 갖도록 하고 헤더가 올바른 열과 연결되도록 존재합니다.

# --hints--

각 `th` 요소에는 `span` 요소가 있어야 합니다.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

새로운 `span` 요소 각각에는 `class` 속성이 `sr-only year`로 설정되어야 합니다.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

첫 번째 `span` 요소에는 텍스트 `2019`가 있어야 합니다.

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

두 번째 `span` 요소에는 텍스트 `2020`가 있어야 합니다.

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

세 번째 `span` 요소에는 텍스트 `2021`가 있어야 합니다.

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

세 번째 `th` 요소에는 `class`가 `current`로 설정되어야 합니다.

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

`td` 요소는 비워 두세요.

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
