---
id: 61fd778081276b59d59abad6
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

첫 번째 `tr` 안에 텍스트 `th`를 가진 `Cash This is the cash we currently have on hand.` 요소를 추가하세요. `Cash `를 제외한 모든 텍스트를 `span` 요소로 감싸고 `class`를 `description`로 설정하세요.

그 다음, 다음 텍스트를 가진 세 개의 `td` 요소를 순서대로 추가하세요: `$25`, `$30`, `$28`. 세 번째 `td` 요소에 `class` 속성을 `current`로 설정하세요.

# --hints--

첫 번째 `tr`에는 `th` 요소가 있어야 합니다.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th'));
```

`th` 요소에는 텍스트 `Cash This is the cash we currently have on hand.`가 있어야 합니다.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th')?.innerText === 'Cash This is the cash we currently have on hand.');
```

텍스트 `This is the cash we currently have on hand.`는 `span` 요소로 감싸야 합니다.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.textContent === 'This is the cash we currently have on hand.');
```

`span` 요소에는 `class` 속성이 `description`로 설정되어야 합니다.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.classList?.contains('description'));
```

세 개의 `td` 요소가 있어야 합니다.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td').length === 3);
```

첫 번째 `td` 요소에는 텍스트 `$25`가 있어야 합니다.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[0]?.textContent === '$25');
```

두 번째 `td` 요소에는 텍스트 `$30`가 있어야 합니다.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[1]?.textContent === '$30');
```

세 번째 `td` 요소에는 텍스트 `$28`가 있어야 합니다.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.textContent === '$28');
```

세 번째 `td` 요소에는 `class`가 `current`로 설정되어야 합니다.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
