---
id: 61fd5a93fd62bb35968adeab
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 문서가 영어임을 나타내는 `DOCTYPE`, `html`와 `head`, `body` 요소를 포함하여 HTML을 설정하세요.

`head` 요소에 `meta`와 `charset`에 적합한 `viewport` 속성을 주고, 적절한 제목을 가진 `title` 요소와 스타일시트용 `link` 요소를 추가하세요.

# --hints--

코드에 `<!DOCTYPE html>` 선언이 있어야 합니다.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

코드에 `html` 요소가 있어야 합니다.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

`html` 요소의 여는 태그에는 `lang` 속성이 `en`로 설정되어야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 요소에는 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/html\s*>/));
```

코드에 `head` 요소가 `html` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

코드에 `body` 요소가 `html` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

`head` 요소가 `body` 요소보다 앞에 와야 합니다.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

`meta` 요소가 두 개 있어야 합니다.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

한 `meta` 요소에는 `name`가 `viewport`로, `content`가 `width=device-width, initial-scale=1.0`로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

다른 `meta` 요소에는 `charset` 속성이 `UTF-8`로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

코드에 `title` 요소가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title`에는 텍스트가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

코드에 `link` 요소가 있어야 합니다.

```js
assert(/<link/.test(code))
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

`link` 요소에는 값이 `rel`인 `stylesheet` 속성이 있어야 합니다.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

`link` 요소에는 값이 `href`인 `styles.css` 속성이 있어야 합니다.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
