---
id: 61fd6ab779390f49148773bb
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`h1` 요소 아래에 `div` 요소를 만드세요. `id` 속성을 `years`으로 설정하세요. 이 특정 요소는 화면 읽기 프로그램에서 숨기고 싶으니 `aria-hidden` 속성을 `true`로 설정하세요.

# --hints--

`div` 요소 다음에 새 `h1` 요소를 만들어야 합니다.

```js
assert(document.querySelector('h1')?.nextElementSibling?.localName === 'div');
```

새 `div` 요소에 `id` 속성을 `years`으로 설정하세요.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'years');
```

새 `div` 요소에 `aria-hidden` 속성을 `true`으로 설정하세요.

```js
assert(document.querySelector('div')?.getAttribute('aria-hidden') === 'true');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
