---
id: 61fd6cc9475a784b7776233e
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

기존 `div` 요소 아래에 `div`로 설정된 `class`를 가진 새로운 `table-wrap` 요소를 추가하세요. 이 요소가 표들을 담는 컨테이너가 됩니다.

그 안에 세 개의 `table` 요소를 추가하세요. CSS를 사용해 이들을 하나의 표처럼 스타일링할 예정이지만, 별도의 표를 사용하는 것이 화면 읽기 프로그램이 문서 흐름을 이해하는 데 도움이 됩니다.

# --hints--

새로운 `div` 요소를 생성하세요.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

새로운 `div` 요소는 `class`이 `table-wrap`로 설정되어야 합니다.

```js
assert(document.querySelector('.table-wrap')?.localName === 'div');
```

기존 `.table-wrap` 요소 다음에 `div` 요소가 와야 합니다.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap'));
```

`.table-wrap` 요소는 세 개의 `table` 요소를 포함해야 합니다.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert(children?.length === 3);
children.forEach(child => assert(child?.localName === 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
