---
id: 61fd6343fb42013d99bcd7f3
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`body` 요소 안에 `section` 요소 안에 `main` 요소를 중첩하세요.

# --hints--

`body` 요소 안의 첫 번째 요소는 `main` 요소여야 합니다.

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

`section` 요소 안에 `main` 요소가 있어야 합니다.

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
