---
id: 6976db7ecfa770bf21307b20
title: Crear un motor de mezcla de listas de reproducción
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

En este laboratorio, construirás un programa que crea una única lista de reproducción remix a partir de múltiples listas enviadas por los oyentes.

Cada oyente proporciona una lista de canciones que quiere escuchar. Algunas canciones pueden aparecer más de una vez, y algunos artistas pueden aparecer demasiadas veces. Tu tarea es procesar estas listas paso a paso: combinarlas en una sola lista, puntuar cada canción, eliminar canciones duplicadas, limitar cuántas veces aparece el mismo artista y luego crear un orden final de reproducción.

**Objetivo:** Cumple las historias de usuario a continuación y haz que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `flattenPlaylists` que acepte un arreglo de listas de reproducción donde cada lista es un arreglo de objetos con las siguientes propiedades: `trackId`, `artist`, `title`, `votes`, `bpm`. Si la entrada no es un arreglo, `flattenPlaylists` debe devolver un arreglo vacío. Se te ha proporcionado una lista de ejemplo para que puedas probar tu función.

2. `flattenPlaylists` debe devolver un arreglo plano de objetos de pista, donde cada objeto incluye todas las propiedades originales de la pista más una propiedad `source` establecida en un arreglo con el índice de la lista y el índice de la pista que indica de dónde proviene la pista.

3. Debes crear una función llamada `scoreTracks` que acepte un arreglo de objetos de pista como los devueltos por `flattenPlaylists` (cada uno con propiedades `trackId`, `artist`, `title`, `votes`, `bpm` y `source`) y devuelva un nuevo arreglo de objetos de pista, cada uno con una propiedad `score` añadida usando la fórmula: `votes * 10 - Math.abs(bpm - 120)`.

4. Debes crear una función llamada `dedupeTracks` que acepte un arreglo de objetos de pista como los devueltos por `scoreTracks` y devuelva un nuevo arreglo con las entradas duplicadas de `trackId` eliminadas, manteniendo solo la primera ocurrencia de cada una.

5. Debes crear una función llamada `enforceArtistQuota` que acepte un arreglo de objetos de pista como los devueltos por `dedupeTracks` y un número que representa el máximo de apariciones permitidas por artista. La función debe devolver un nuevo arreglo donde ningún artista aparezca más veces que el número dado, manteniendo las ocurrencias más tempranas.

6. Debes crear una función llamada `buildSchedule` que acepte un arreglo de objetos de pista como los devueltos por `enforceArtistQuota` y devuelva un nuevo arreglo de objetos `{ slot, trackId }`, donde `slot` es un índice basado en 1 que representa la posición de cada pista en el orden de transmisión.

7. Debes crear una función llamada `remixPlaylist` que acepte un arreglo de listas de reproducción y el número máximo de apariciones permitidas por artista. La función debe devolver el programa final de transmisión como un arreglo de objetos `{ slot, trackId }`, llamando a `flattenPlaylists`, `scoreTracks`, `dedupeTracks`, `enforceArtistQuota` y `buildSchedule` en orden.

# --hints--

Debes tener una función llamada `flattenPlaylists`.

```js
assert.isFunction(flattenPlaylists);
```

Debes devolver un arreglo vacío desde `flattenPlaylists` cuando la entrada no sea un arreglo.

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

Cada pista devuelta por `flattenPlaylists` debe incluir un campo `source` que sea un arreglo que contenga el índice de la lista y el índice de la pista.

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

Debes tener una función llamada `scoreTracks`.

```js
assert.isFunction(scoreTracks)
```

Cada pista devuelta por `scoreTracks` debe incluir un campo numérico `score`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

Debes calcular `score` usando un BPM objetivo de `120` y esta fórmula: `votes * 10 - Math.abs(bpm - 120)`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

Debes tener una función llamada `dedupeTracks`.

```js
assert.isFunction(dedupeTracks)
```

Cuando existan valores duplicados de `trackId`, `dedupeTracks` debe mantener solo la primera ocurrencia de la pista.

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

Debes tener una función llamada `enforceArtistQuota`.

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` debe asegurar que ningún artista aparezca más de `maxPerArtist` veces eliminando pistas extra mientras mantiene las primeras.

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

Debes tener una función llamada `buildSchedule`.

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` debe devolver un arreglo de objetos con la forma `{ slot, trackId }`, donde `slot` comienza en `1`.

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

Debes tener una función llamada `remixPlaylist`.

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` debe llamar a las funciones auxiliares en orden para producir el programa final.

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
