---
id: 60a3e3396c7b40068ad69976
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

影格太寬了。

在 `.frame` 中，將其 `width` 設為 500 像素。

# --hints--

你應該將 `width` 屬性設為 `500px`。

```js
const widthFilter = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.style.width === '500px');
assert(widthFilter.length === 2);
```

你的 `.frame` 元素應該有 `width` 為 `500px`。

```js
const frameWidth = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('width');
assert(frameWidth === '500px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
