---
id: 60a3e3396c7b40068ad6998b
title: 步驟 32
challengeType: 0
dashedName: step-32
---

# --description--

將外邊距推向一個方向會很有幫助。

在這裏，`.one` 元素的下外邊距將 `.two` 向下推 20 個像素。

在 `.two` 選擇器中，使用 `margin` 簡寫屬性將上外邊距設置爲 `0`，水平外邊距設置爲 `auto`，下外邊距設置爲 `20px`。 這將刪除其上外邊距，將其水平居中，並將其下外邊距設置爲 20 像素。

# --hints--

你應該將 `margin` 屬性設置爲 `0 auto 20px`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

你的 `.two` 元素應該有一個 `margin`，值爲 `0 auto 20px`。

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
