---
id: 60a3e3396c7b40068ad6996b
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

開始項目，在 `body` 元素中添加一個 `img` 元素，其 `src` 爲 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` 。

# --hints--

你的代碼應該有一個 `img` 元素。

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 元素的 `src` 應該是 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
