---
id: 60a3e3396c7b40068ad69978
title: 步驟14
challengeType: 0
dashedName: step-14
---

# --description--

用外邊距調整元素外部的空間。

用 `margin` 屬性爲 `.frame` 元素賦予 `20px` 的垂直外邊距和 `auto` 的水平外邊距。 這會將框架向下移動 20 像素並在頁面上水平居中。

# --hints--

你應將 `margin` 屬性設置爲 `20px auto`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

你的 `.frame` 元素應有一個 `margin`，值爲 `20px auto`。

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
