---
id: 60a3e3396c7b40068ad69975
title: 步驟11
challengeType: 0
dashedName: step-11
---

# --description--

用 `.frame` 類選擇器編寫新一個規則。

用 `border` 簡寫聲明爲 `.frame` 元素賦予一個寬度爲 `50px` 的實線黑色邊框。

# --hints--

你的代碼應有一個 `.frame` 選擇器。

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

你應將 `border` 屬性設置爲 `50px solid black`。

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

你的 `.frame` 元素應有一個 `border` 屬性，值爲 `50px solid black`。

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
