---
id: 60a3e3396c7b40068ad69974
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

每幅畫都需要一個框架。

將 `.canvas` 元素包裝在另一個 `div` 中。 爲 `div` 賦予 `frame` 類。
# --hints--

你應添加一個新的 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 2)
```

你的 `.canvas` 元素應嵌套在新的 `div` 元素中。

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

你的新 `div` 應有一個 `class`，值爲 `frame`。

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

你的新 `div` 應在 `body` 元素中。

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
