---
id: 60a3e3396c7b40068ad6997d
title: 步驟 19
challengeType: 0
dashedName: step-19
---

# --description--

使用外邊距在畫布上定位 `.one` 元素。

添加簡寫 `margin` 屬性，垂直外邊距爲 `20px`，水平外邊距爲 `auto`。

# --hints--

你應將 `margin` 屬性設置爲 `20px auto`。

```js
const marginFilter = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.style.margin === '20px auto');
assert(marginFilter.length === 2);
```

你的 `.one` 元素應該有一個 `margin`，值爲 `20px auto`。

```js
const oneMargin = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('margin');
assert(oneMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
