---
id: 60a3e3396c7b40068ad6996c
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在之前關於 CSS 盒模型的課程中，你學到每個 HTML 元素都被視為具有四個區域的盒子。

想像你從你最喜愛的線上零售商收到一個盒子——盒子裡的內容就是該項目，在你的情況下，是標題、段落或影像元素。

將 `<img>` 中的 `src` 屬性從 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` 更改為 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`。

# --hints--

`img` 元素的 `src` 應該是 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

