---
id: 60a3e3396c7b40068ad6997e
title: 第 20 步
challengeType: 0
dashedName: step-20
---

# --description--

現在 `.one` 水平置中，但它的上邊距超出畫布並推到訊框的邊框上，導致整個畫布向下移動 20 像素。

為你的 `.canvas` 元素添加 `1px` 的 `padding`，以便 `.one` 元素有實體可推動。

# --hints--

你應該將 `padding` 屬性設為 `1px`。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

你的 `.canvas` 元素應該有 `padding` 值為 `1px`。

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
