---
id: 60a3e3396c7b40068ad6996f
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `body` 中添加一個 `div` 元素。

將 `class` 屬性設為 `canvas`。

這將作為你繪畫的畫布。

# --hints--

你的程式碼應該有一個 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 1)
```

`div` 元素應該有一個值為 `canvas` 的 `class`。

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
