---
id: 60a3e3396c7b40068ad69977
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

使用 padding 來調整元素內部的間距。

在 `.frame` 中，使用 `padding` 簡寫屬性將 `.frame` 和 `.canvas` 元素之間的間距增加 `50px`。該簡寫會增加元素邊框和畫布內部的上、下、左、右空間。

# --hints--

你應該將 `padding` 屬性設為 `50px`。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

你的 `.frame` 元素應該有 `padding` 值為 `50px`。

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
