---
id: 68da6138d70960c197ec5eef
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Kwa njia ya utafutaji wa upana kwanza, utatumia foleni kufuatilia hali tofauti unapounda mchanganyiko wa mabano ya kawaida. Kila hali itawakilishwa kama tuple yenye vipengele vitatu:

- Mfuatano wa herufi unaojengwa kwa sasa
- Idadi ya mabano ya kawaida ya kufungua yaliyotumika hadi sasa
- Idadi ya mabano ya kawaida ya kufunga yaliyotumika hadi sasa

Tengeneza kigezo kinachoitwa `queue` na uanze na orodha yenye tuple moja: `('', 0, 0)`. Hii inawakilisha hali ya kuanzia yenye mfuatano wa herufi tupu na mabano sifuri yaliyotumika.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `queue` katika kitendakazi chako cha `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

Unapaswa kuanzisha `queue` na orodha yenye tuple `('', 0, 0)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    result = []
    return result
--fcc-editable-region--
```
