---
id: 68da6138d70960c197ec5ef0
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Sasa utaekeleza mzunguko mkuu wa BFS. Unda mzunguko wa `while` unaoendelea mradi `queue` haijajaa (wakati `queue` inatathminiwa kuwa `True`).

Ndani ya mzunguko huo, chapisha `queue`.

# --hints--

Unapaswa kuwa na mzunguko wa `while` katika kitendakazi chako cha `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

Mzunguko wako wa `while` unapaswa kuwa na sharti la `queue`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

Unapaswa kuchapisha `queue` ndani ya mzunguko wako wa `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
