---
id: 68da6138d70960c197ec5eec
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Kabla ya kutekeleza algoriti kuu, unahitaji kuthibitisha ingizo. Kigezo cha `pairs` kinapaswa kuwa nambari kamili, kwani kinawakilisha idadi ya jozi za mabano ya kawaida za kuzalisha.

Ongeza tamko la `if` mwanzoni mwa kitendakazi chako ili kuangalia kama `pairs` si nambari kamili. Tumia kitendakazi cha `isinstance()` kwa ajili hiyo.

Kama hali ni kweli, rudisha mfuatano wa herufi `The number of pairs should be an integer`.

# --hints--

Unapaswa kuwa na tamko la `if` katika kitendakazi chako cha `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

Tamko lako la `if` linapaswa kuangalia kama `pairs` si nambari kamili kwa kutumia `not isinstance(pairs, int)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

Unapaswa kurudisha mfuatano wa herufi `The number of pairs should be an integer` wakati ingizo si nambari kamili.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
--fcc-editable-region--
    
--fcc-editable-region--
    
    return []
```
