---
id: 68da6138d70960c197ec5ef3
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Sasa unahitaji kushughulikia kuongeza mabano ya kawaida ya kufunga. Kanuni kuu ni kwamba unaweza kuongeza mabano ya kawaida ya kufunga tu ikiwa inahakikisha usawa, yaani lazima kuwe na mabano ya kawaida ya kufungua zaidi kuliko mabano ya kawaida ya kufunga.

Bado ndani ya blokki ya `else`, ongeza tamko la pili la `if` ili kukagua kama `closes_used < opens_used`.

Kama sharti hili ni kweli, ongeza tuple nyingine mpya kwenye `queue`: `(current + ')', opens_used, closes_used + 1)`. Hii inawakilisha hali baada ya kuongeza mabano ya kawaida ya kufunga.

# --hints--

Unapaswa kuwa na tamko la pili la `if` ndani ya blokki yako ya `else`.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

Tamko lako la pili lililopangwa ndani la `if` linapaswa kukagua kama `closes_used < opens_used`.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

Unapaswa kuongeza `(current + ')', opens_used, closes_used + 1)` kwenye foleni wakati sharti ni kweli.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
--fcc-editable-region--
            
--fcc-editable-region--

    return result

print(gen_parentheses(1))
```
