---
id: 68da604b8b6012a33e2344fb
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, uta-tekeleza kitendakazi kinachozalisha mchanganyiko wote halali wa mabano ya kawaida kwa kutumia mbinu ya utafutaji wa upana kwanza (BFS). Kwa mfano, mchanganyiko halali wa mabano ya kawaida ya jozi mbili ni `(())` na `()()`.

Anza kwa kuunda kitendakazi kinachoitwa `gen_parentheses` chenye kigezo kimoja `pairs`. Kwa sasa, rudisha orodha tupu kutoka kwa kitendakazi hicho.

# --hints--

Unapaswa kufafanua kitendakazi kinachoitwa `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

Kitendakazi chako cha `gen_parentheses` kinapaswa kuwa na kigezo kimoja kinachoitwa `pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

Kitendakazi chako kinapaswa kurudisha orodha tupu.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
