---
id: 68da6138d70960c197ec5eed
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Ifuatayo, unahitaji kuthibitisha kuwa idadi ya jozi ni angalau moja, kwa sababu huwezi kuunda mchanganyiko wa mabano ya kawaida ukiwa na jozi sifuri au hasi.

Ongeza tamko jingine la `if` ili kuangalia kama `pairs` ni ndogo kuliko `1`. Ikiwa hali hii ni kweli, rudisha mfuatano wa herufi `The number of pairs should be at least 1`.

# --hints--

Unapaswa kuwa na tamko la pili la `if` katika kitendakazi chako cha `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

Tamko lako la pili la `if` linapaswa kuangalia kama `pairs` ni ndogo kuliko `1`.

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

Unapaswa kurudisha mfuatano wa herufi `The number of pairs should be at least 1` wakati `pairs` ni ndogo kuliko `1`.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
--fcc-editable-region--
    
--fcc-editable-region--
    
    return []
```
