---
id: 68e00c5349cbddfef6dbc053
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Ndani ya mzunguko wako wa `while`, tumia `queue.pop(0)` kuondoa na kupata kipengele cha kwanza kutoka kwenye foleni. Hii inatekeleza tabia ya kwanza-kuingia-kwa-kwanza-kutoka (FIFO) inayojulikana katika BFS.

Fungua tuple hii katika vigezo vitatu: `current`, `opens_used`, na `closes_used`.

# --hints--

Unapaswa kutumia `queue.pop(0)` kupata kipengele cha kwanza kutoka kwenye foleni.

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

Unapaswa kufungua matokeo hayo katika vigezo vitatu: `current`, `opens_used`, na `closes_used`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
--fcc-editable-region--
        
--fcc-editable-region--

    return result
```
