---
id: 69708d59833ceded7d00541f
title: 實作鏈結串列操作
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

在本實驗中，你將為鏈結串列資料結構實作額外的操作，並建立在新增和移除節點的基本功能之上。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事**

1. 你應該有一個 `contains` 函式（程式），接受一個鏈結串列和一個元素。如果指定的元素存在於鏈結串列中，則應該傳回 `true`，否則傳回 `false`。
1. 你應該有一個 `getAt` 函式（程式），它接受一個鏈結串列和一個索引。它應該傳回鏈結串列中該索引位置的元素。如果索引超出範圍，它應該傳回 `undefined`。
1. 你應該有一個 `insertAt` 函式（程式），它接受一個鏈結串列、一個索引和一個元素。它應該在鏈結串列中指定的位置插入給定的元素。如果索引超出範圍，則不應修改列表。
1. 你應該有一個 `removeAt` 函式，接受一個鏈結串列和一個索引。它應該移除鏈結串列中指定索引的節點。如果索引超出範圍，則不應修改列表。
1. 你應該有一個接受鏈結串列的 `clear` 函式（程式）。它應該移除鏈結串列中的所有元素，有效地將其重設為空狀態。

**注意：** 有些後續測試依賴於先前的方法。例如，如果 `getAt` 沒有正確實作，像 `insertAt` 和 `removeAt` 這類函式的測試即使接近正確，也可能會失敗。

# --hints--

`initList` 應該傳回一個 `head` 設為 `null` 且 `length` 設為 `0` 的物件。

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

`add` 應該將列表長度增加一。

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

`remove` 應該將列表長度減少一。

```js
const list = initList();  
add(list, 'apple');  
add(list, 'banana');  
remove(list, 'apple');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(list.head.element, 'banana');  
assert.isNull(list.head.next);
```

你應該有一個 `contains` 函式（程式）。

```js
assert.isFunction(contains);
```

`contains` 應該在元素存在於列表中時傳回 `true`。

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains` 應該在元素不存在於列表中時傳回 `false`。

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains` 應該對空的列表傳回 `false`。

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

你應該有一個 `getAt` 函式（程式）。

```js
assert.isFunction(getAt);
```

`getAt` 應該傳回指定索引處的元素。

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt` 應該為無效的索引傳回 `undefined`。

```js
const list = initList();  
add(list, 'apple');  
assert.isUndefined(getAt(list, -1));  
assert.isUndefined(getAt(list, 1));  
assert.isUndefined(getAt(list, 5));
```

`getAt` 應該為空的列表傳回 `undefined`。

```js
const list = initList();  
assert.isUndefined(getAt(list, 0));
```

你應該有一個 `insertAt` 函式（程式）。

```js
assert.isFunction(insertAt);
```

`insertAt` 應該在列表的開頭插入一個元素。

```js
const list = initList();  
add(list, 'second');  
add(list, 'third');  
insertAt(list, 0, 'first');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` 應該在列表的中間插入一個元素。

```js
const list = initList();  
add(list, 'first');  
add(list, 'third');  
insertAt(list, 1, 'second');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` 應該在列表的末尾插入一個元素。

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` 應該在空的列表中於索引 `0` 插入。  

```js  
const list = initList();  
insertAt(list, 0, 'first');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.isNull(list.head.next); 
```

`insertAt` 不應該修改列表中無效索引的內容。

```js
const list = initList();  
add(list, 'apple');  
insertAt(list, -1, 'invalid');  
insertAt(list, 10, 'invalid');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
assert.isUndefined(getAt(list, 1));
```

你應該有一個 `removeAt` 函式（程式）。

```js
assert.isFunction(removeAt);
```

`removeAt` 應該移除指定索引處的元素。

```js
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'third');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` 應該正確移除第一個元素。

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

`removeAt` 應該正確移除最後一個元素。  

```js  
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 2);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` 在移除唯一元素時應該重設列表。

```js
const list = initList();  
add(list, 'only');  
removeAt(list, 0);  
assert.strictEqual(list.length, 0);  
assert.isNull(list.head);  
assert.isTrue(isEmpty(list));
```

`removeAt` 不應該修改列表中無效索引的內容。

```js
const list = initList();  
add(list, 'apple');  
removeAt(list, -1);  
removeAt(list, 10);  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
```

你應該有一個 `clear` 函式（程式）。

```js
assert.isFunction(clear);
```

`clear` 應該從列表中移除所有元素。

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`clear` 應該允許該列表被重複使用。  

```js  
const list = initList();  
add(list, 'first');  
clear(list);  
clear(list);  
add(list, 'second');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'second');  
assert.isNull(list.head.next); 
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  
}

function getAt(list, index) {
 
}

function insertAt(list, index, element) {
  
}

function removeAt(list, index) {
 
}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
