---
id: 683ec95f2f8781355556ff83
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Puoi vedere che la prima chiamata restituisce `[3]`. Questo perché `3` è il valore centrale in `[1, 2, 3, 4, 5]`.

D'altra parte, la seconda chiamata restituisce una lista vuota. Questo succede perché, dopo aver controllato il valore centrale, il loop attualmente non ha la logica per restringere l'intervallo di ricerca aggiornando le variabili `low` o `high`. Funziona solo come un controllo singolo del punto medio e poi termina senza ulteriori ricerche.

Per permettere alla ricerca binaria di continuare a restringere la ricerca, aggiungi un blocco `elif` che controlli se `value` è maggiore di `value_at_middle`. Aggiungi la parola chiave `pass` all'interno del `elif` per ora.

# --hints--

Dovresti aggiungere un blocco `elif` per quando `value` è maggiore di `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--          
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
