---
id: 683ec95f2f8781355556ff7f
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Ora, includi `value_at_middle` nella lista `path_to_target` per tracciare i passaggi effettuati durante la ricerca, indipendentemente dal fatto che sia l'obiettivo.

# --hints--

Dovresti usare il metodo `append()` per aggiungere `value_at_middle` alla lista `path_to_target`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid] \\npath_to_target.append(value_at_middle)")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        
--fcc-editable-region--
```
