---
id: 683ec95f2f8781355556ff7a
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

L'algoritmo di ricerca binaria che stai costruendo non restituirà solo `True` o `False` per indicare se il valore è stato trovato o meno. Restituirà anche il percorso che segue per trovare quel valore. Per questo, avrai una variabile `path_to_target` per tracciare il percorso.

All'interno della tua funzione `binary_search`, sostituisci la parola chiave `pass` con una variabile `path_to_target` e inizializzala a una lista vuota, per cominciare.

# --hints--

Dovresti creare una variabile `path_to_target` impostata su una lista vuota all'interno della tua funzione `binary_search`. Non dimenticare di rimuovere la parola chiave `pass`.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```
