---
id: 683ec95f2f8781355556ff84
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Se la condizione nel `elif` è vera, allora rimuovi la parola chiave `pass` e aggiorna il valore della variabile `low` aggiungendo `1` alla variabile `mid`.

Questo estenderà la ricerca alla metà destra delle aree di ricerca correnti nella lista, perché se `value` è maggiore di `value_at_middle`, significa che `value` deve trovarsi nella metà destra dell’area di ricerca corrente.

# --hints--

Dovresti rimuovere la parola chiave `pass` e aggiornare la variabile `low` a `mid + 1`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid + 1")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        elif value > value_at_middle:
            pass
--fcc-editable-region--
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
```
