---
id: 683ec39565dbc7e3375263b1
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In una lezione precedente, hai imparato la ricerca binaria e come si differenzia dalla ricerca lineare.

In questo workshop, metterai in pratica questo concetto costruendo un algoritmo di ricerca binaria.

Per cominciare, crea una funzione `binary_search` con due parametri `search_list` e `value`.

`search_list` è la lista di elementi all’interno della quale la funzione cercherà, e `value` è l’obiettivo che la funzione proverà a trovare all’interno di `search_list`.

Per ora, aggiungi la parola chiave `pass` all’interno della funzione così il test può passare.

# --hints--

Dovresti creare una funzione `binary_search`.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

La tua funzione `binary_search` dovrebbe avere due parametri `search_list` e `value`.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
