---
id: 683ec95f2f8781355556ff80
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Il passaggio successivo è creare una condizione che verifichi se l'obiettivo `value` è al centro.

All'interno del loop `while`, crea un'istruzione `if` che controlli se l'obiettivo `value` è uguale a `value_at_middle`. Se lo è, restituisci la variabile `path_to_target`.

# --hints--

Dovresti creare un'istruzione `if` che verifichi se `value` è uguale a `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

Dovresti restituire la variabile `path_to_target` all'interno dell'istruzione `if`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
--fcc-editable-region--
        
--fcc-editable-region--
```
