---
id: 683ec95f2f8781355556ff7c
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora che hai definito i limiti della tua ricerca, è il momento di creare il loop che eseguirà l'algoritmo di ricerca binaria.

La ricerca binaria funziona restringendo ripetutamente lo spazio di ricerca. Questo processo continua finché esiste un intervallo valido di elementi da controllare.

Puoi esprimere questo con un loop `while` che continuerà finché il puntatore `low` è minore o uguale al puntatore `high`.

Se `low` diventa mai maggiore di `high`, significa che lo spazio di ricerca è diventato vuoto e `value` non è nella lista.

All'interno della funzione `binary_search`, crea un loop `while` con una condizione che verifica se `low` è minore o uguale a high. All'interno del loop, aggiungi la parola chiave `pass` come segnaposto per ora.

# --hints--

Dovresti creare un loop `while` che gira mentre `low` è minore o uguale a `high`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--
    
--fcc-editable-region--
```
