---
id: 683ec95f2f8781355556ff81
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Devi testare alcune cose per capire il flusso dell'algoritmo in questa fase iniziale.

Per farlo, devi prima uscire dal loop. Questo perché l'implementazione attuale permette solo un'iterazione, quindi se la condizione non è soddisfatta, si creerà un loop infinito.

Subito dopo l'istruzione `if`, usa la parola chiave `break` per uscire dal loop `while`. Poi, dopo il loop `while`, restituisci una lista vuota per indicare che il valore non è stato trovato.

# --hints--

Dovresti usare la parola chiave `break` per uscire dal loop `while` dopo l'istruzione `if`.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

Dovresti restituire una lista vuota dopo il tuo loop `while`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```
