---
id: 683ec95f2f8781355556ff7d
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Per iniziare con il loop, devi trovare l'indice centrale dello spazio di ricerca attuale.

Puoi ottenere l'indice centrale calcolando la media delle variabili `low` e `high`. Dovrai usare la divisione intera per ottenere la media dopo aver sommato i due valori, così la risposta sarà sempre un numero intero arrotondato per difetto, dato che gli indici devono essere numeri interi.

Rimuovi la parola chiave `pass` dal tuo loop `while` e crea una variabile `mid` impostata sulla media di `low` e `high`.

# --hints--

Dovresti rimuovere la parola chiave `pass` dal while loop e creare una variabile `mid` impostata su `(low + high) // 2`.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```
