---
id: 683ec95f2f8781355556ff7e
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora che hai calcolato l'indice `mid`, devi effettivamente recuperare il valore dal `search_list` a quell'indice. Questo valore è quello che confronterai con il parametro `value`, che è l'obiettivo che stai cercando.

Crea una variabile `value_at_middle` e assegnale l'elemento di `search_list` che si trova all'indice `mid`. 

# --hints--

Dovresti creare una variabile `value_at_middle` impostata su `search_list[mid]`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2

        
--fcc-editable-region--
```
