---
id: 683ec95f2f8781355556ff7b
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Nella ricerca binaria, esiste l'indice più basso possibile e l'indice più alto possibile.

L'indice più basso possibile rappresenta il limite sinistro nella porzione corrente dei valori in cui si sta cercando, in questo caso `search_list`. L'indice più alto possibile, invece, rappresenta il limite destro nei valori in cui si sta cercando.

Per tenere conto di questi due, devi considerare l'intera lista in cui si sta cercando. Quindi, all'interno della funzione `binary_search`, definisci una variabile `low` con valore `0` e una variabile `high` con un valore che rappresenta l'ultimo indice nella lista in cui si sta cercando.

# --hints--

Dovresti avere una variabile chiamata `low` nella tua funzione `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

La tua variabile `low` dovrebbe essere impostata su `0`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

Dovresti avere una variabile chiamata `high` nella tua funzione `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

La tua variabile `high` dovrebbe essere impostata su `len(search_list) - 1`.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []
    
--fcc-editable-region--
```
