---
id: 684169b00a073c9b9f0a703c
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Hai notato che la seconda chiamata della funzione ora ha i valori `3, 5, 4` nella lista? Questo indica che la ricerca binaria sta funzionando come previsto.

Ecco come è successo: l'algoritmo ha prima controllato `3` come elemento centrale della lista iniziale. Poiché `4` è maggiore di `3`, la ricerca si è spostata nella metà destra. Poi ha esaminato `5` come nuovo elemento centrale. Poiché `4` è minore di `5`, la ricerca si è spostata a sinistra, identificando infine `4` come elemento centrale dell'intervallo finale.

Per testare di nuovo la funzione, chiamala con `[1, 3, 5, 9, 14, 22], 10` e stampane subito il risultato. Questa è una situazione in cui il valore non verrà trovato.

# --hints--

Dovresti chiamare la funzione con `binary_search([1, 3, 5, 9, 14, 22], 10)` e stamparla subito.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []
    
print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
