---
id: 68416d66013635cb6325eb6d
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Un ultimo piccolo miglioramento che puoi aggiungere è mostrare anche l’indice in cui si trova l’obiettivo `value`.

Quindi, invece di restituire solo `path_to_target` nell’istruzione `if`, aggiungi un `f-string` con il messaggio `Value found at index {mid}`.

Con questo, il tuo workshop sull’algoritmo di ricerca binaria è completo!

# --hints--

Dovresti restituire `path_to_target, f'Value found at index {mid}'` all’interno del blocco `if`.

```js
({
    test: () => 
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target, f'Value found at index {mid}'")`
            )       
        )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            --fcc-editable-region--
            return path_to_target
            --fcc-editable-region--
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return [], "Value not found"

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
print(binary_search([1, 3, 5, 9, 14, 22], 10))
```

# --solutions--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
          return path_to_target, f"Value found at index {mid}"
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return [], "Value not found"

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
print(binary_search([1, 3, 5, 9, 14, 22], 10))
```
