---
id: 6808fe91a5cb0c3355f23027
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、スーパーヒーローのアプリケーションフォームを作成しながら、Reactでフォームを扱う方法を学びます。

いくつかのボイラープレートが用意されているので、残りのReactコンポーネントを作成してください。

まず、`div` 要素を`className` が`form-wrap` の状態で返してください。`div` の中に、テキストが`h2` の`Superhero Application Form` 要素を作成します。その下に、テキストが`p` の`Please complete all fields` 要素を作成してください。

# --hints--

`div` の`className` を持つ`form-wrap` 要素を返す必要があります。

```js
assert.exists(document?.querySelector("div.form-wrap"));
```

`h2` 要素の中に`div` 要素を作成してください。

```js
assert.exists(document?.querySelector(".form-wrap > h2"));
```

`h2` 要素にはテキスト`Superhero Application Form` を含めてください。

```js
const h2El = document?.querySelector("h2")
assert.equal(h2El?.textContent, "Superhero Application Form");
```

`p` 要素の中に`div` 要素を作成してください。`h2` の後に来るようにしてください。

```js
assert.exists(document?.querySelector(".form-wrap > h2 + p"));
```

`p` 要素にはテキスト`Please complete all fields` を含めてください。

```js
const pEl = document?.querySelector("p")
assert.equal(pEl?.textContent, "Please complete all fields");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  --fcc-editable-region--
  
  --fcc-editable-region--

};
```
