---
id: 680900675ae3d54ee19590c9
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`select` 要素の中に、ユーザーが電源として選べるオプションを作成する必要があります。`select` には以下の7つのオプションをリストしてください。

| 電源                        |
|-----------------------------|
| `Select one`                  |
| `Bitten by a strange creature`  |
| `Radioactive exposure`          |
| `Science experiment`           |
| `Alien heritage `               |
| `Ancient artifact discovery`    |
| `Other`                       |

最初のオプション（`Select one`）は、`value` 属性に空の文字列を持つ必要があります（つまり、`value=''`）。これはUIのプレースホルダーとして機能します。

# --hints--

`option` 要素の中にテキスト `Select one` を持つ最初の `select` 要素を作成してください。

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.textContent, "Select one");
```

最初の `option` 要素に空の文字列を持つ `value` 属性を付けてください。

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.getAttribute('value'), '');
```

`option` 要素の中にテキスト `Bitten by a strange creature` を持つ2番目の `select` 要素を作成してください。

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[1]?.textContent, "Bitten by a strange creature");
```

`option` 要素の中にテキスト `Radioactive exposure` を持つ3番目の `select` 要素を作成してください。

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[2]?.textContent, "Radioactive exposure");
```

`option` 要素の中にテキスト `Science experiment` を持つ4番目の `select` 要素を作成してください。

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[3]?.textContent, "Science experiment");
```

`option` 要素の中にテキスト `Alien heritage` を持つ5番目の `select` 要素を作成してください。

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[4]?.textContent, "Alien heritage");
```

`option` 要素の中にテキスト `Ancient artifact discovery` を持つ6番目の `select` 要素を作成してください。

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[5]?.textContent, "Ancient artifact discovery");
```

`option` 要素の中にテキスト `Other` を持つ7番目の `select` 要素を作成してください。

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[6]?.textContent, "Other");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          <label>
            Hero Name
            <input
              type='text'
              value={heroName}
              onChange={e => setHeroName(e.target.value)}
            />
          </label>
          <label>
            Real Name
            <input
              type='password'
              value={realName}
              onChange={e => setRealName(e.target.value)}
            />
          </label>
        </div>
          <label className='section column'>
            How did you get your powers?
            --fcc-editable-region--
            <select>
              
            </select>
            --fcc-editable-region--
          </label>
      </form>
    </div>
  )
};
```
