---
id: 5a24c314108439a4d4036164
title: Create a Component with Composition
challengeType: 6
forumTopicId: 301383
dashedName: create-a-component-with-composition
---

# --description--

Agora vamos ver como podemos compor vûÀrios componentes React juntos. Imagine que vocûˆ estûÀ construindo um aplicativo e criou trûˆs componentes: uma `Navbar`, um `Dashboard`, e um `Footer`.

Para compor esses componentes juntos, vocûˆ pode criar um componente *pai* `App` que renderiza cada um desses trûˆs componentes como *filhos*. Para renderizar um componente como filho em um componente React, vocûˆ inclui o nome do componente escrito como uma tag HTML personalizada no JSX. Por exemplo, no mûˋtodo `render` vocûˆ pode escrever:

```jsx
return (
 <App>
  <Navbar />
  <Dashboard />
  <Footer />
 </App>
)
```

Quando o React encontra uma tag HTML personalizada que referencia outro componente (um nome de componente envolto em `< />` como neste exemplo), ele renderiza a marcaûÏûÈo para esse componente na localizaûÏûÈo da tag. Isso deve ilustrar a relaûÏûÈo pai/filho entre o componente `App` e a `Navbar`, o `Dashboard`, e o `Footer`.

# --instructions--

No editor de cû°digo, hûÀ um simples componente funcional chamado `ChildComponent` e um componente de classe chamado `ParentComponent`. Componha as duas juntas ao renderizar o `ChildComponent` dentro do `ParentComponent`. Certifique-se de fechar a tag `ChildComponent` com uma barra avanûÏada.

**Nota:** `ChildComponent` ûˋ definido com uma funûÏûÈo arrow ES6 porque esta ûˋ uma prûÀtica muito comum ao usar React.

# --hints--

O componente React deve retornar um û¤nico elemento `div`.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.type() === 'div';
  })()
);
```

O componente deve retornar dois elementos aninhados.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.children().length === 2;
  })()
);
```

O componente deve retornar o `ChildComponent` como seu segundo filho.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ParentComponent));
    return (
      mockedComponent.find('ParentComponent').find('ChildComponent').length ===
      1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }


        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }
        <ChildComponent />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
