---
id: 63ee3f71381756f9716727ef
title: CSS 基礎練習 A
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**目標：**  
在這個練習中，你將練習使用三種方法為 HTML 檔案添加 CSS：外部 CSS、內部 CSS 和行內 CSS。當你透過外部和內部方法添加樣式時，應該只使用型別選擇器。你也應該使用顏色關鍵字（例如「blue」）而非使用 `RGB` 或 `HEX` 值。

## 使用者故事

- 你應該會看到一個帶有一些文字的 `div` 元素。
  - 它應該有 `red` 背景、`white` 文字、`32px` 字型大小、文字置中對齊和 `bold`。
  - `div` 元素的 CSS 應該外部添加，並使用型別選擇器。
- 你應該會看到一個帶有一些文字的 `p` 元素。
  - 它應該有 `green` 背景、`white` 文字，和 `18px` 的字型大小。
  - `p` 元素的 CSS 應該以內部方式添加，並使用型別選擇器。
- 你應該會看到一個帶有一些文字的 `button` 元素。
  - `button` 元素應該有 `orange` 背景和 `18px` 的字型大小。
  - `button` 元素的 CSS 應該使用行內樣式添加。

# --hints--

你應該有一個包含一些文字的 `div` 元素。

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

你應該有一個包含 `div` 元素樣式的外部樣式表。

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

你的 `div` 元素不應該使用內部或行內樣式來添加其 CSS。

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

你的 `div` 元素應該有 `background-color` 為 `red` 和 `color` 為 `white`。

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

你的 `div` 元素應該將 `font-weight` 設為 `bold`、`font-size` 設為 `32px`，並將 `text-align` 設為 `center`。

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

你應該有一個 `p` 元素，且它應該包含一些文字。

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

你的 `p` 元素應該使用 `style` 元素在內部添加其樣式。

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

你的 `p` 元素應該有 `18px` 的 `font-size` 並且將 `color` 設定為白色。

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

你應該有一個包含一些文字的 `button` 元素。

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

你的 `button` 元素應該有行內樣式。

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```
 
你的 `button` 元素應該將 `background-color` 設定為 `orange`。

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

你的 `button` 元素的 `font-size` 應該設定為 `18px`。

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
