---
id: 63ee3fe4381756f9716727f0
title: CSS 基礎練習 B
challengeType: 25
dashedName: css-foundations-exercise-b
---

# --description--

**目標：** 在提供的 HTML 檔案中有數個元素，你必須為它們添加類別或 ID 屬性。接著，你必須使用正確的選擇器語法，在提供的 CSS 檔案中添加規則。

## 使用者故事

1. 你應該會看到 `yellow` 背景用於列表中所有奇數編號的元素。

1. 你應該為列表中所有奇數編號的元素使用 `class` 選取器。

1. 你應該會看到列表中的第二個元素有 `blue` 文字和 `36px` 的 `font-size`。

1. 第二個元素的 `font-size` 和文字顏色應該使用 `id` 屬性來設定。

1. 你應該會看到列表中第三個元素的 `font-size` 為 `24px`。

1. 第三個元素上的 `font-size` 應該使用 `class` 屬性來設定。

1. 你應該會看到列表中的第四個元素有 `red` 背景、`24px` 的 `font-size`，以及 `bold` 的 `font-weight`。

1. 第四個元素的 `font-size` 應該使用 `class` 屬性設定。`font-weight` 和顏色應該使用 `id` 屬性設定。

# --hints--

每個奇數元素都應該有一個 `class` 屬性。

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasClass = p?.every((paragraph) => paragraph.classList.length > 0);

assert(everyPHasClass);
```

你的奇數元素應該有 `background-color` 為 `yellow`。

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasBackgroundColor = p?.every((paragraph) => {
  const style = getComputedStyle(paragraph);
  
  return style?.backgroundColor === 'rgb(255, 255, 0)';
})

assert.equal(everyPHasBackgroundColor, true);
```

你的第二個元素應該有 `blue` 文字和 `font-size` 為 `36px`。

```js
const secondElementId = document.querySelectorAll('div')?.[0]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${secondElementId}`);

assert.equal(style?.color, 'blue');
assert.equal(style?.fontSize, '36px');
```

你的第三個元素應該有文字和 `font-size` 為 `24px`。

```js
const thirdElement = document.querySelectorAll('p')?.[1];

assert(thirdElement?.innerText?.length > 0);

const thirdElementClasses = Array.from(thirdElement?.classList?.values());

assert(thirdElementClasses.some(thirdElementClass => {

  const style = new __helpers.CSSHelp(document).getStyle(`.${thirdElementClass}`);

  return style?.fontSize === '24px';

}))

``` 

第四個元素的 `font-size` 應為 `24px`。

```js
const fourthElementClass = document.querySelectorAll('div')?.[1]?.classList[0];

const style = new __helpers.CSSHelp(document).getStyle(`.${fourthElementClass}`);

assert(style?.fontSize === '24px');
```

第四個元素應該有 `red` 的 `background-color`。

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.backgroundColor === 'red');
```

第四個元素的 `font-weight` 應該是 `bold`。

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.fontWeight === 'bold');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p>Number 1 - I'm a class!</p>
    <div>Number 2 - I'm one ID.</div>
    <p>Number 3 - I'm a class, but cooler!</p>
    <div>Number 4 - I'm another ID.</div>
    <p>Number 5 - I'm a class!</p>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p class="odd">Number 1 - I'm a class!</p>
    <div id="two">Number 2 - I'm one ID.</div>
    <p class="odd adjust-font-size">Number 3 - I'm a class, but cooler!</p>
    <div id="four" class="adjust-font-size">Number 4 - I'm another ID.</div>
    <p class="odd">Number 5 - I'm a class!</p>
  </body>
</html>
```

```css
.odd {
  background-color: yellow;
  font-family: Verdana, "DejaVu Sans", sans-serif;
}

.adjust-font-size {
  font-size: 24px;
}

#two {
  color: blue;
  font-size: 36px;
}

#four {
  background-color: red;
  font-weight: bold;
}
```
