---
id: 66c8ba41a77db20a93f9d7a1
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você revisará condicionais JavaScript criando uma aplicação verificadora de qualificação para empréstimo.

A aplicação verificará se o usuário é elegível para um empréstimo de duplex, carro e condomínio com base na renda anual e pontuação de crédito. Uma pontuação de crédito é um número que representa o nível de gerenciamento de dinheiro emprestado.

Para começar, crie as seguintes variáveis e valores.

| Nome da variável          | Valor   |
| ------------------------- | ------- |
| `minIncomeForDuplex`      | `60000` |
| `minCreditScoreForDuplex` | `700`   |
| `minIncomeForCondo`       | `45000` |
| `minCreditScoreForCondo`  | `680`   |
| `minIncomeForCar`         | `30000` |
| `minCreditScoreForCar`    | `650`   |

# --hints--

Você deve criar uma variável `minIncomeForDuplex` definida como `60000`.

```js
assert.equal(minIncomeForDuplex, 60000);
```

Você deve criar uma `minCreditScoreForDuplex` definida como `700`.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

Você deve criar uma variável `minIncomeForCar` definida como `30000`.

```js
assert.equal(minIncomeForCar, 30000);
```

Você deve criar uma variável `minCreditScoreForCar` definida como `650`.

```js
assert.equal(minCreditScoreForCar, 650);
```

Você deve criar uma variável `minIncomeForCondo` definida como `45000`.

```js
assert.equal(minIncomeForCondo, 45000);
```

Você deve criar uma variável `minCreditScoreForCondo` definida como `680`.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
