---
id: 66c8ba975ee7230e29f6c4ad
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Para verificar para qual empréstimo um usuário está qualificado com base na `annualIncome` e na `creditScore`, você deve usar a instrução `if/else if` ou um ternário dentro da função `getLoanMessage`. You'll then return the appropriate message in the block of each condition.

Começando com o empréstimo duplex, verifique se `annualIncome` é maior ou igual a `minIncomeForDuplex` E se `creditScore` é maior que `minCreditScoreForDuplex`.

Se essa condição for verdadeira, então o candidato é elegível para um empréstimo duplex e os outros empréstimos. Então, dentro da verificação, retorne a string `"You qualify for a car, duplex and Condo loan."`

# --hints--

A função `getLoanMessage` deve retornar uma string.

```js
assert.isString(getLoanMessage(65000, 750));
```

A função `getLoanMessage` deve retornar a string `"You qualify for a duplex, condo, and car loan."`.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
```

A função `getLoanMessage` deve retornar `undefined` se a renda anual do candidato e a pontuação de crédito não atenderem aos requisitos para um empréstimo duplex.

```js
assert.strictEqual(getLoanMessage(59000, 700), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--
function getLoanMessage(annualIncome, creditScore) {

}
--fcc-editable-region--
```
