---
id: 66c8ba975ee7230e29f6c4b0
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora, você deve verificar se o solicitante está qualificado apenas para um empréstimo de carro. Se estiver, retorne a string `"You qualify for a car loan."`.

# --hints--

A função `getLoanMessage` deve retornar uma string se o solicitante estiver qualificado para um empréstimo de carro.

```js
assert.isString(getLoanMessage(30000, 650));
```

A função `getLoanMessage` deve retornar a string `"You qualify for a car loan."`.

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

A função `getLoanMessage` deve retornar `undefined` se a renda anual e a pontuação de crédito do solicitante não atenderem aos requisitos para um empréstimo de carro.

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 

--fcc-editable-region--
}
```
