---
id: 66c8ba975ee7230e29f6c4ac
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Quando o usuário for elegível para um empréstimo, você vai querer exibir uma mensagem para ele no console.

Para isso, você vai criar uma função dentro da qual terá algumas verificações que retornarão para qual empréstimo o candidato é elegível.

Crie uma função vazia `getLoanMessage` com parâmetros `annualIncome` e `creditScore`.

# --hints--

Você deve criar uma função `getLoanMessage`

```js
assert.isFunction(getLoanMessage)
```

A função `getLoanMessage` deve ter `annualIncome` e `creditScore` como parâmetros.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` deve ser uma função vazia.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
