---
id: 683ecb010f4c6bafd428802c
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستتحقق من صحة مجموعة من البيانات الطبية لضمان توافقها مع مجموعة من القواعد.

أعلن عن متغير باسم `medical_records` وقم بتعيين قائمة فارغة له. في الخطوات التالية، ستستخدمه لتخزين بياناتك الطبية.

# --hints--

يجب إعلان متغير باسم `medical_records`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

يجب تعيين قائمة فارغة للمتغير `medical_records`.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

