---
id: 6842aee5bcaf4ab8883f8191
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

كما تعلمت في الدروس السابقة، القاموس هو هيكل بيانات يحتفظ بأزواج المفتاح-القيمة، حيث تكون المفاتيح أنواع بيانات غير قابلة للتغيير (مثل السلاسل النصية) وفريدة داخل القاموس:

```py
person = {
    'name': 'John',
    'age': 33
}
```

ستحتوي قائمة `medical_records` على قواميس، كل منها يمثل مريضًا. أضف قاموسًا بمفتاح `patient_id` وقيمة السلسلة النصية `P1001` إلى قائمة `medical_records`.

# --hints--

يجب أن تحتوي قائمة `medical_records` على قاموس واحد فقط.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

يجب أن يحتوي قاموسك على مفتاح `patient_id` بقيمة السلسلة النصية `P1001`.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
