---
id: 6842d659a4ae7973dabab3b0
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

أكمل القاموس بإضافة أزواج المفتاح-القيمة الثلاثة التالية:

- مفتاح `diagnosis` بقيمة `Hypertension`
- مفتاح `medications` بقيمة القائمة `['Lisinopril']`
- مفتاح `last_visit_id` بقيمة `V2301`.

# --hints--

يجب أن يحتوي قاموسك على مفتاح `diagnosis` بقيمة السلسلة النصية `Hypertension`.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

يجب أن يحتوي قاموسك على مفتاح `medications` بقيمة القائمة `['Lisinopril']`.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

يجب أن يحتوي قاموسك على مفتاح `last_visit_id` بقيمة السلسلة النصية `V2301`.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

يجب أن يحتوي قاموسك على مفتاح `patient_id` بقيمة السلسلة النصية `P1001`.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

يجب أن يحتوي قاموسك على مفتاح `age` بقيمة العدد الصحيح `34`.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

يجب أن يحتوي قاموسك على مفتاح `gender` بقيمة السلسلة النصية `Female`.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
