---
id: 6842d5b537c64266420faf62
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

أضف مفتاح `age` بقيمة العدد الصحيح `34` ومفتاح `gender` بقيمة السلسلة النصية `Female` إلى القاموس الخاص بك. لا تنسَ الفاصلة بين أزواج المفتاح-القيمة.

# --hints--

يجب أن تحتوي قائمة `medical_records` على قاموس واحد فقط.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

يجب أن يحتوي قاموسك على مفتاح `age` بقيمة العدد الصحيح `34`.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

يجب أن يحتوي قاموسك على مفتاح `gender` بقيمة السلسلة النصية `Female`.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

يجب أن يحتوي قاموسك على مفتاح `patient_id` بقيمة السلسلة النصية `P1001`.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
