---
id: 6842ba0086201089ea268c92
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

باتباع نفس الهيكل الذي استخدمته في الخطوات السابقة، تم ملء قائمة `medical_records` لك ببيانات مرضى آخرين. يمكنك الاطلاع عليها بحرية.

بعد ذلك، ستبدأ بكتابة الدالة للتحقق من صحة مجموعة البيانات. أنشئ دالة باسم `validate` بمعلمة واحدة فقط هي `data`.

تريد التأكد من أن بياناتك إما قائمة أو صف مرتب. لذلك، داخل دالة `validate`، أعلن عن متغير باسم `is_sequence` وقم بتعيين استدعاء لـ `isinstance` له. مرر `data` كالمعلمة الأولى وصف مرتب يحتوي على `list` و `tuple` كالمعلمة الثانية.

# --hints--

يجب أن يكون لديك دالة باسم `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

يجب أن تحتوي دالة `validate` على معلمة واحدة هي `data`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

يجب أن يكون لديك متغير باسم `is_sequence` داخل دالة `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

يجب أن تعيّن `isinstance(data, (list, tuple))` إلى متغير `is_sequence`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
