---
id: 6846a1f2805ac73ce671a5d8
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

كما تعلمت في درس سابق، تتيح لك الدالة `enumerate` تتبع الفهرس أثناء التكرار على قابل للتكرار:

```py
person = {'name': 'John', 'age': 33}

for index, item in enumerate(person):
    print(index, item)

# 0 name
# 1 age
```

أنشئ حلقة `for` تتكرر على `data`. استخدم الدالة `enumerate` للحصول على كل من الفهرس والعنصر في `data` في كل تكرار. استخدم `index` و`dictionary` كمتغيرات التكرار.

استخدم `pass` لملء جسم الحلقة في الوقت الحالي.

# --hints--

يجب أن تحتوي دالتك `validate` على حلقة `for`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0]
`) })
```

يجب أن تتكرر حلقة `for` على `enumerate(data)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_iter().is_equivalent("enumerate(data)")
`) })
```

يجب أن تستخدم `index` و`dictionary` كمتغيرات التكرار في حلقة `for`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_vars().is_equivalent("index, dictionary")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

--fcc-editable-region--
    
--fcc-editable-region--
```
