---
id: 68517af94b02d6cc2bc08280
title: الخطوة 40
challengeType: 20
dashedName: step-40
---

# --description--

تقبل تراكيب القوائم أيضًا بنود `if` لتصفية العناصر من قابل للتكرار:

```py
nums = [1, 2, 3, 4, 5, 6]
even_nums = [num for num in nums if num % 2 == 0]
print(even_nums) # [2, 4, 6]
```

بما أنك تريد إرجاع قائمة تحتوي فقط على المفاتيح غير الصالحة، أضف بند `if` إلى تركيبك بحيث يُضاف كل `key` إلى القائمة فقط عندما تكون قيمة `value` خاطئة.

# --hints--

يجب أن تُرجع دالة `find_invalid_records` تركيب قائمة يقيم `key` لكل `key, value` في `constraints.items()` إذا كانت `value` خاطئة.

```js
({ test: () => assert(runPython(`
_if_clause = _Node(_code).find_function("find_invalid_records").find_return().find_comp_ifs()[0]
_if_clause.is_equivalent("not value") or _if_clause.is_equivalent("value == False")
`)) })
```

يجب أن تُرجع دالة `find_invalid_records` تركيب قائمة يتكرر على `constraints.items()`.

```js
({ test: () => assert(runPython(`
_iter = _Node(_code).find_function("find_invalid_records").find_return().find_comp_iters()[0]
_iter.is_equivalent("constraints.items()")
`)) })
```

يجب أن تُرجع دالة `find_invalid_records` تركيب قائمة يستخدم `key, value` للتكرار على `constraints.items()`.

```js
({ test: () => assert(runPython(`
_target = _Node(_code).find_function("find_invalid_records").find_return().find_comp_targets()[0]
_target.is_equivalent("key, value")
`)) })
```

يجب أن تُرجع دالة `find_invalid_records` تركيب قائمة يقيم `key` لكل `key, value` في `constraints.items()`.

```js
({ test: () => assert(runPython(`
_expr = _Node(_code).find_function("find_invalid_records").find_return().find_comp_expr()
_expr.is_equivalent("key")
`)) })
```

# --seed--

## --seed-contents--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
        'patient_id': isinstance(patient_id, str)
        and re.fullmatch('p\d+', patient_id, re.IGNORECASE),
        'age': isinstance(age, int) and age >= 18,
        'gender': isinstance(gender, str) and gender.lower() in ('male', 'female'),
        'diagnosis': isinstance(diagnosis, str) or diagnosis is None,
        'medications': isinstance(medications, list)
        and all([isinstance(i, str) for i in medications]),
        'last_visit_id': isinstance(last_visit_id, str)
        and re.fullmatch('v\d+', last_visit_id, re.IGNORECASE)
    }
--fcc-editable-region--
    return [key for key, value in constraints.items()]
--fcc-editable-region--

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
print(find_invalid_records(**medical_records[0]))
```
