---
id: 587d7b7b367417b2b2512b16
title: 複雑な多次元配列を作成する
challengeType: 1
forumTopicId: 301159
dashedName: create-complex-multi-dimensional-arrays
---

# --description--

すごいです！配列についてたくさん学びましたね！これはかなり高度な概要でしたが、配列を扱うことについてはまだまだ学ぶことがあり、後のセクションでさらに多くのことを見ていきます。でも、<dfn>オブジェクト</dfn>を見る前に、もう一度配列を見てみましょう。配列はこれまでのチャレンジで見たものよりも少し複雑になることがあります。

配列をデータ構造として考えるときの最も強力な機能の一つは、配列が他の配列を含むことができる、あるいは完全に他の配列で構成されることもある点です。これまでのチャレンジで配列が配列を含む例を見ましたが、かなり単純なものでした。しかし、配列は無限の深さで他の配列を含むことができ、それぞれが任意の深さを持つことができます。このようにして、配列は非常に複雑なデータ構造、つまり<dfn>多次元</dfn>またはネストされた配列になることがあります。次の例を考えてみましょう。

```js
let nestedArray = [
  ['deep'],
  [
    ['deeper'], ['deeper'] 
  ],
  [
    [
      ['deepest'], ['deepest']
    ],
    [
      [
        ['deepest-est?']
      ]
    ]
  ]
];
```

`deep` 配列は2段階の深さでネストされています。`deeper` 配列は3段階の深さです。`deepest` 配列は4段階、`deepest-est?` は5段階です。

この例は複雑に見えるかもしれませんが、大量のデータを扱う場合、このレベルの複雑さは珍しくもなく、むしろよくあることです。しかし、このように複雑な配列の最も深いレベルにもブラケット記法で簡単にアクセスできます。

```js
console.log(nestedArray[2][1][0][0][0]);
```

これは文字列 `deepest-est?` をログに出力します。そして、そのデータの場所がわかったので、必要に応じてリセットすることもできます。

```js
nestedArray[2][1][0][0][0] = 'deeper still';

console.log(nestedArray[2][1][0][0][0]);
```

今度は `deeper still` をログに出力します。

# --instructions--

`myNestedArray` という変数が配列に等しく定義されています。`myNestedArray` を、データ要素として <dfn>文字列</dfn>、<dfn>数値</dfn>、<dfn>ブール値</dfn> の任意の組み合わせを使って修正し、ちょうど5段階の深さを持つようにしてください（最も外側の配列がレベル1です）。3段階目のどこかに文字列 `deep` を含め、4段階目に文字列 `deeper`、5段階目に文字列 `deepest` を含めてください。

# --hints--

`myNestedArray` にはデータ要素として数値、ブール値、文字列のみを含めるべきです

```js
assert.strictEqual(
  (function (arr) {
    let flattened = (function flatten(arr) {
      const flat = [].concat(...arr);
      return flat.some(Array.isArray) ? flatten(flat) : flat;
    })(arr);
    for (let i = 0; i < flattened.length; i++) {
      if (
        typeof flattened[i] !== 'number' &&
        typeof flattened[i] !== 'string' &&
        typeof flattened[i] !== 'boolean'
      ) {
        return false;
      }
    }
    return true;
  })(myNestedArray),
  true
);
```

`myNestedArray` はちょうど5段階の深さを持つべきです

```js
assert.strictEqual(
  (function (arr) {
    let depth = 0;
    function arrayDepth(array, i, d) {
      if (Array.isArray(array[i])) {
        arrayDepth(array[i], 0, d + 1);
      } else {
        depth = d > depth ? d : depth;
      }
      if (i < array.length) {
        arrayDepth(array, i + 1, d);
      }
    }
    arrayDepth(arr, 0, 0);
    return depth;
  })(myNestedArray),
  4
);
```

`myNestedArray` は3段階深くネストされた配列に文字列 `deep` をちょうど1回含むべきです

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deep').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deep')[0] === 2
);
```

`myNestedArray` は4段階深くネストされた配列に文字列 `deeper` をちょうど1回含むべきです

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deeper').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deeper')[0] === 3
);
```

`myNestedArray` は5段階深くネストされた配列に文字列 `deepest` をちょうど1回含むべきです

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deepest').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deepest')[0] === 4
);
```

# --seed--

## --seed-contents--

```js
let myNestedArray = [
  // Only change code below this line
  ['unshift', false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
  // Only change code above this line
];
```

# --solutions--

```js
let myNestedArray = [
  ['unshift', ['deep', ['deeper', ['deepest']]],false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
];
```
