---
id: 587d7b7d367417b2b2512b1f
title: オブジェクトに格納された配列を変更する
challengeType: 1
forumTopicId: 301163
dashedName: modify-an-array-stored-in-an-object
---

# --description--

これでJavaScriptオブジェクトの基本的な操作はすべて見てきました。キーと値のペアを追加、変更、削除し、キーの存在を確認し、オブジェクト内のすべてのキーを繰り返すことができます。JavaScriptの学習を続けると、オブジェクトのさらに多用途な使い方が見えてきます。さらに、カリキュラムのCoding Interview Prepセクションにあるデータ構造のレッスンでは、ES6の<dfn>Map</dfn>や<dfn>Set</dfn>オブジェクトも扱います。これらは通常のオブジェクトに似ていますが、追加の機能を提供します。配列とオブジェクトの基本を学んだので、JavaScriptを使ってより複雑な問題に取り組む準備が整いました！

# --instructions--

コードエディターに用意したオブジェクトを見てください。`user`オブジェクトには3つのキーがあります。`data`キーには5つのキーがあり、そのうちの1つには`friends`の配列が含まれています。これにより、オブジェクトがどれほど柔軟なデータ構造であるかがわかります。`addFriend`という関数の作成を始めています。`user`オブジェクトを受け取り、`friend`引数の名前を`user.data.friends`に格納された配列に追加し、その配列を返すように関数を完成させてください。

# --hints--

`user`オブジェクトは`name`、`age`、`data`のキーを持つはずです。

```js
assert('name' in user && 'age' in user && 'data' in user);
```

`addFriend`関数は`user`オブジェクトと`friend`文字列を引数として受け取り、`friends`オブジェクトの`user`配列に友達を追加する必要があります。

```js
assert(
  (function () {
    let L1 = user.data.friends.length;
    addFriend(user, 'Sean');
    let L2 = user.data.friends.length;
    return L2 === L1 + 1;
  })()
);
```

`addFriend(user, "Pete")`は`["Sam", "Kira", "Tomo", "Pete"]`を返すべきです。

```js
assert.deepEqual(
  (function () {
    delete user.data.friends;
    user.data.friends = ['Sam', 'Kira', 'Tomo'];
    return addFriend(user, 'Pete');
  })(),
  ['Sam', 'Kira', 'Tomo', 'Pete']
);
```

# --seed--

## --seed-contents--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  // Only change code below this line

  // Only change code above this line
}

console.log(addFriend(user, 'Pete'));
```

# --solutions--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  userObj.data.friends.push(friend);
  return userObj.data.friends;
}
```
