---
id: 587d7b7b367417b2b2512b17
title: スプレッド演算子で配列を結合する
challengeType: 1
forumTopicId: 301156
dashedName: combine-arrays-with-the-spread-operator
---

# --description--

スプレッド演算子のもう一つの大きな利点は、配列を結合したり、ある配列のすべての要素を別の配列の任意のインデックスに挿入したりできることです。従来の構文では配列を連結できますが、これは一方の配列の末尾ともう一方の配列の先頭でのみ結合できます。スプレッド構文を使うと、次の操作が非常に簡単になります。

```js
let thisArray = ['sage', 'rosemary', 'parsley', 'thyme'];

let thatArray = ['basil', 'cilantro', ...thisArray, 'coriander'];
```

`thatArray` は `['basil', 'cilantro', 'sage', 'rosemary', 'parsley', 'thyme', 'coriander']` の値を持ちます。

スプレッド構文を使うことで、従来の方法を使った場合よりも複雑で冗長になったであろう操作を簡単に実現できました。

# --instructions--

`spreadOut` という関数を定義しており、変数 `sentence` を返します。スプレッド演算子を使って関数を修正し、配列 `['learning', 'to', 'code', 'is', 'fun']` を返すようにしてください。

# --hints--

`spreadOut` は `["learning", "to", "code", "is", "fun"]` を返すべきです

```js
assert.deepEqual(spreadOut(), ['learning', 'to', 'code', 'is', 'fun']);
```

`spreadOut` 関数はスプレッド構文を利用するべきです

```js
assert.notStrictEqual(spreadOut.toString().search(/[...]/), -1);
```

# --seed--

## --seed-contents--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence; // Change this line
  return sentence;
}

console.log(spreadOut());
```

# --solutions--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence = ['learning', ...fragment, 'is', 'fun'];
  return sentence;
}
```
