---
id: 587d7b7b367417b2b2512b13
title: スプレッド演算子で配列をコピーする
challengeType: 1
forumTopicId: 301157
dashedName: copy-an-array-with-the-spread-operator
---

# --description--

`slice()` は配列のどの要素をコピーするか選択的に指定できますが、ES6の新しい<dfn>spread operator</dfn>は、シンプルで非常に読みやすい構文で配列の*すべて*の要素を順番に簡単にコピーできます。スプレッド構文は次のように書きます：`...`

実際には、スプレッド演算子を使って配列を次のようにコピーできます：

```js
let thisArray = [true, true, undefined, false, null];
let thatArray = [...thisArray];
```

`thatArray` は `[true, true, undefined, false, null]` と等しくなります。`thisArray` は変更されず、`thatArray` は `thisArray` と同じ要素を含みます。

# --instructions--

`copyMachine` という関数を定義しており、`arr`（配列）と `num`（数値）を引数に取ります。この関数は `num` の `arr` 回コピーでできた新しい配列を返すはずです。ほとんどの処理は済んでいますが、まだ正しく動作しません。スプレッド構文を使って関数を修正し、正しく動作するようにしてください（ヒント：すでに学んだ別のメソッドが役に立つかもしれません！）。

# --hints--

`copyMachine([true, false, true], 2)` は `[[true, false, true], [true, false, true]]` を返すべきです

```js
assert.deepEqual(copyMachine([true, false, true], 2), [
  [true, false, true],
  [true, false, true]
]);
```

`copyMachine([1, 2, 3], 5)` は `[[1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3]]` を返すべきです

```js
assert.deepEqual(copyMachine([1, 2, 3], 5), [
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3]
]);
```

`copyMachine([true, true, null], 1)` は `[[true, true, null]]` を返すべきです

```js
assert.deepEqual(copyMachine([true, true, null], 1), [[true, true, null]]);
```

`copyMachine(["it works"], 3)` は `[["it works"], ["it works"], ["it works"]]` を返すべきです

```js
assert.deepEqual(copyMachine(['it works'], 3), [
  ['it works'],
  ['it works'],
  ['it works']
]);
```

`copyMachine` 関数は配列 `spread operator` に `arr` を利用すべきです

```js
assert(__helpers.removeJSComments(__helpers.removeJSComments(code)).match(/\.\.\.\s*arr/));
```

# --seed--

## --seed-contents--

```js
function copyMachine(arr, num) {
  let newArr = [];
  while (num >= 1) {
    // Only change code below this line

    // Only change code above this line
    num--;
  }
  return newArr;
}

console.log(copyMachine([true, false, true], 2));
```

# --solutions--

```js
function copyMachine(arr,num){
    let newArr=[];
    while(num >=1){
    newArr.push([...arr]);
    num--;
    }
    return newArr;
}
console.log(copyMachine([true, false, true], 2));
```
