---
id: 587d7b7d367417b2b2512b1e
title: Object.keys() でオブジェクトのすべてのキーの配列を生成する
challengeType: 1
forumTopicId: 301160
dashedName: generate-an-array-of-all-object-keys-with-object-keys
---

# --description--

`Object.keys()` メソッドを使って、オブジェクトに格納されているすべてのキーを含む配列を生成することもできます。このメソッドはオブジェクトを引数に取り、そのオブジェクト内の各プロパティを表す文字列の配列を返します。配列内のエントリーには特定の順序はありません。

# --instructions--

`getArrayOfUsers` 関数の作成を完成させて、引数として受け取ったオブジェクト内のすべてのプロパティを含む配列を返すようにしてください。

# --hints--

`users` オブジェクトには `Alan`、`Jeff`、`Sarah`、`Ryan` のキーだけが含まれているはずです

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

`getArrayOfUsers` 関数は `users` オブジェクトのすべてのキーを含む配列を返すべきです

```js
assert(
  (function () {
    users.Sam = {};
    users.Lewis = {};
    let R = getArrayOfUsers(users);
    return (
      R.indexOf('Alan') !== -1 &&
      R.indexOf('Jeff') !== -1 &&
      R.indexOf('Sarah') !== -1 &&
      R.indexOf('Ryan') !== -1 &&
      R.indexOf('Sam') !== -1 &&
      R.indexOf('Lewis') !== -1
    );
  })() === true
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  // Only change code below this line

  // Only change code above this line
}

console.log(getArrayOfUsers(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  return Object.keys(obj);
}

console.log(getArrayOfUsers(users));
```
