---
id: 587d7b7b367417b2b2512b15
title: 配列のすべての要素を For ループで繰り返し処理する
challengeType: 1
forumTopicId: 301161
dashedName: iterate-through-all-an-arrays-items-using-for-loops
---

# --description--

配列を扱うとき、必要な要素を見つけたり、特定の条件を満たすデータ項目に基づいて配列を操作したりするために、各要素を繰り返し処理できることは非常に便利です。JavaScript には、配列を少しずつ異なる方法で繰り返し処理して異なる結果を得るいくつかの組み込みメソッド（`every()`、`forEach()`、`map()` など）がありますが、最も柔軟で制御性が高い手法はシンプルな `for` ループです。

次の例を考えてみましょう。

```js
function greaterThanTen(arr) {
  let newArr = [];
  for (let i = 0; i < arr.length; i++) {
    if (arr[i] > 10) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}

greaterThanTen([2, 12, 8, 14, 80, 0, 1]);
```

`for` ループを使って、この関数は配列の各要素を繰り返し処理しアクセスし、私たちが作成した簡単なテストにかけています。このようにして、`10` より大きいデータ項目を簡単かつプログラム的に判別し、それらの項目を含む新しい配列 `[12, 14, 80]` を返しています。

# --instructions--

`filteredArray` という関数が定義されており、ネストされた配列 `arr` と `elem` を引数に取り、新しい配列を返します。`elem` は `arr` 内の一つ以上の配列に存在するかもしれない要素を表します。`for` ループを使って関数を修正し、`arr` 内のネストされた配列のうち `elem` を含むものが除外されたフィルター済みの配列を返すようにしてください。

# --hints--

`filteredArray([[10, 8, 3], [14, 6, 23], [3, 18, 6]], 18)` は `[[10, 8, 3], [14, 6, 23]]` を返すべきです

```js
assert.deepEqual(
  filteredArray(
    [
      [10, 8, 3],
      [14, 6, 23],
      [3, 18, 6]
    ],
    18
  ),
  [
    [10, 8, 3],
    [14, 6, 23]
  ]
);
```

`filteredArray([["trumpets", 2], ["flutes", 4], ["saxophones", 2]], 2)` は `[["flutes", 4]]` を返すべきです

```js
assert.deepEqual(
  filteredArray(
    [
      ['trumpets', 2],
      ['flutes', 4],
      ['saxophones', 2]
    ],
    2
  ),
  [['flutes', 4]]
);
```

`filteredArray([["amy", "beth", "sam"], ["dave", "sean", "peter"]], "peter")` は `[["amy", "beth", "sam"]]` を返すべきです

```js
assert.deepEqual(
  filteredArray(
    [
      ['amy', 'beth', 'sam'],
      ['dave', 'sean', 'peter']
    ],
    'peter'
  ),
  [['amy', 'beth', 'sam']]
);
```

`filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3)` は `[]` を返すべきです

```js
assert.deepEqual(
  filteredArray(
    [
      [3, 2, 3],
      [1, 6, 3],
      [3, 13, 26],
      [19, 3, 9]
    ],
    3
  ),
  []
);
```

`filteredArray` 関数は `for` ループを利用すべきです

```js
assert.notStrictEqual(filteredArray.toString().search(/for/), -1);
```

# --seed--

## --seed-contents--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  // Only change code below this line

  // Only change code above this line
  return newArr;
}

console.log(filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3));
```

# --solutions--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  for (let i = 0; i<arr.length; i++) {
    if (arr[i].indexOf(elem) < 0) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}
```
