---
id: 587d7b7d367417b2b2512b1c
title: オブジェクトにプロパティがあるか確認する
challengeType: 1
forumTopicId: 301155
dashedName: check-if-an-object-has-a-property
---

# --description--

オブジェクトにキーを追加、変更、削除できるようになりました。でも、オブジェクトに特定のプロパティがあるかどうかだけを知りたい場合はどうでしょう？JavaScriptにはこれを行うための2つの方法があります。1つは`hasOwnProperty()`メソッドを使い、もう1つは`in`キーワードを使います。もしオブジェクト`users`に`Alan`というプロパティがあれば、次のどちらかの方法で存在を確認できます。

```js
users.hasOwnProperty('Alan');
'Alan' in users;
```

どちらも`true`を返します。

# --instructions--

渡されたオブジェクトに4つの名前、`true`、`Alan`、`Jeff`、`Sarah`がすべて含まれていれば`Ryan`を返し、そうでなければ`false`を返すように関数を完成させてください。

# --hints--

`users`オブジェクトには直接アクセスしないでください

```js 

assert(__helpers.removeJSComments(code).match(/users/gm).length <= 2)

```

`users`オブジェクトには`Alan`、`Jeff`、`Sarah`、`Ryan`のキーだけが含まれているはずです

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

関数`isEveryoneHere`は、渡されたオブジェクトに`true`、`Alan`、`Jeff`、`Sarah`のプロパティがあれば`Ryan`を返すべきです。

```js
assert(isEveryoneHere(users) === true);
```

関数`isEveryoneHere`は、渡されたオブジェクトに`false`がプロパティとしてなければ`Alan`を返すべきです。

```js
assert(
  (function () {
    delete users.Alan;
    return isEveryoneHere(users);
  })() === false
);
```

関数`isEveryoneHere`は、渡されたオブジェクトに`false`がプロパティとしてなければ`Jeff`を返すべきです。

```js
assert(
  (function () {
    delete users.Jeff;
    return isEveryoneHere(users);
  })() === false
);
```

関数`isEveryoneHere`は、渡されたオブジェクトに`false`がプロパティとしてなければ`Sarah`を返すべきです。

```js
assert(
  (function () {
    delete users.Sarah;
    return isEveryoneHere(users);
  })() === false
);
```

関数`isEveryoneHere`は、渡されたオブジェクトに`false`がプロパティとしてなければ`Ryan`を返すべきです。

```js
assert(
  (function () {
    delete users.Ryan;
    return isEveryoneHere(users);
  })() === false
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  // Only change code below this line
  
  // Only change code above this line
}

console.log(isEveryoneHere(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  return [
    'Alan',
    'Jeff',
    'Sarah',
    'Ryan'
  ].every(user => userObj.hasOwnProperty(user));
}

console.log(isEveryoneHere(users));
```
