---
id: 587d7b7a367417b2b2512b12
title: slice() を使って配列の要素をコピーする
challengeType: 1
forumTopicId: 301158
dashedName: copy-array-items-using-slice
---

# --description--

次に説明するメソッドは `slice()` です。配列を変更するのではなく、`slice()` は指定された数の要素を新しい配列にコピーまたは*抽出*し、呼び出された元の配列はそのままにします。`slice()` は2つのパラメータだけを取ります — 最初は抽出を開始するインデックス、2つ目は抽出を停止するインデックス（このインデックスの要素は含まれません）です。これを考えてみましょう：

```js
let weatherConditions = ['rain', 'snow', 'sleet', 'hail', 'clear'];

let todaysWeather = weatherConditions.slice(1, 3);
```

`todaysWeather` は `['snow', 'sleet']` の値を持ち、`weatherConditions` はまだ `['rain', 'snow', 'sleet', 'hail', 'clear']` のままです。

つまり、既存の配列から要素を抽出して新しい配列を作成したことになります。

# --instructions--

配列を引数に取る関数 `forecast` を定義しました。`slice()` を使って引数の配列から情報を抽出し、文字列要素 `warm` と `sunny` を含む新しい配列を返すように関数を修正してください。

# --hints--

`forecast` は `["warm", "sunny"]` を返すべきです

```js
assert.deepEqual(
  forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']),
  ['warm', 'sunny']
);
```

`forecast` 関数は `slice()` メソッドを利用するべきです

```js
assert(/\.slice\(/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function forecast(arr) {
  // Only change code below this line

  return arr;
}

// Only change code above this line
console.log(forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']));
```

# --solutions--

```js
function forecast(arr) {
  return arr.slice(2,4);
}
```
