---
id: 65ef198fde24dfb7ff675b42
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

تسمح جملة `raise` بإجبار حدوث استثناء معين. تتكون من الكلمة المفتاحية `raise` تليها نوع الاستثناء، وتمكنك من تقديم رسالة خطأ مخصصة:

```py
raise ValueError("Invalid value")
```

عند تشغيل الشفرة أعلاه، يُثار استثناء `ValueError` وتُعرض الرسالة `"Invalid value"` على المستخدم.

إذا كان `square_target` أقل من `0`، فلا يمكن حساب الجذر التربيعي بقيمة حقيقية. لذلك، ارفع استثناء `ValueError` مع الرسالة `'Square root of negative number is not defined in real numbers'`. لا تنسَ إزالة الكلمة المفتاحية `pass`.

# --hints--

يجب إزالة الكلمة المفتاحية `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

يجب رفع استثناء `ValueError` مع الرسالة `'Square root of negative number is not defined in real numbers'` داخل جسم جملة `if`.

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
