---
id: 65ef19ec318e2fb9066381e3
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

ستُنشئ حالات منفصلة عندما يكون `square_target` يساوي `0` أو `1`.

ابدأ بإنشاء جملة `if` للتحقق مما إذا كان `square_target` يساوي `1`.

# --hints--

يجب إنشاء جملة `if` للتحقق من أن `square_target == 1`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[0].is_equivalent("square_target == 1")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')

--fcc-editable-region--
```
