---
id: 65ef1c34a541afc12dbb849d
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

تعيد الدالة `abs()` القيمة المطلقة لعدد، وهي دائمًا موجبة بغض النظر عن إشارة العدد. ستستخدمها للتحقق من أن الجذر التربيعي المقدر قريب بما يكفي من القيمة الفعلية.

الآن، أنشئ جملة `if` للتحقق مما إذا كانت القيمة المطلقة للفارق بين `square_mid` و `square_target` ضمن `tolerance` المحددة.

# --hints--

يجب أن تحتوي على جملة `if` بشرط `abs(square_mid - square_target) < tolerance:` داخل جسم حلقة `for`.

```js
({
    test: () =>
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("abs(square_mid - square_target) < tolerance")`))
    } 
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2
            
--fcc-editable-region--
```
