---
id: 65ef19425d1b27b6c930bae6
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

إذا كان الرقم الذي تريد إيجاد الجذر التربيعي له سالبًا، يجب أن ترفع الشفرة خطأً لأن الجذر التربيعي لرقم سالب غير معرف في الأعداد الحقيقية.

احذف جملة `pass` وأنشئ جملة `if` للتحقق مما إذا كان `square_target` أقل من `0`.

# --hints--

يجب حذف الكلمة المفتاحية `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

يجب إنشاء جملة `if` للتحقق مما إذا كان `square_target < 0`.

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
