---
id: 65ef181dad3d96b14a21a03e
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذا المشروع، ستجد الجذر التربيعي التقريبي لعدد معين باستخدام طريقة التقسيم النصفي.

طريقة التقسيم النصفي هي تقنية لإيجاد جذور دالة ذات قيم حقيقية. تعمل عن طريق تضييق فترة تحتوي على الجذر التربيعي حتى تصل إلى قيمة ضمن هامش خطأ محدد.

ابدأ بإنشاء دالة باسم `square_root_bisection`. اترك المعلمات فارغة في الوقت الحالي.

# --hints--

يجب أن تعرف الدالة `square_root_bisection` بدون معلمات.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
