---
id: 65ef1a50049cf9bada13266f
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

أنشئ جملة `elif` للتحقق مما إذا كانت قيمة `square_target` تساوي `0`. إذا كانت كذلك، فعين القيمة `0` للمتغير `root`. كما اطبع الرسالة `'The square root of {square_target} is 0'`. تذكر تنسيق الرسالة باستخدام سلسلة نصية f-string.

# --hints--

يجب أن تحتوي على جملة `elif` للتحقق من الشرط `square_target == 0`.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

يجب تعيين القيمة `0` للمتغير `root` وتمرير المعلمة `f'The square root of {square_target} is 0'` إلى استدعاء `print`.


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
