---
id: 65ef1be2b958cabf712296cd
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

الآن ستقوم بتضييق الفاصل بشكل متكرر عن طريق إيجاد منتصف الفاصل الحالي ومقارنة مربع منتصف الفاصل بالقيمة الهدف.

لهذا، داخل كتلة `else`، أنشئ حلقة `for` تعمل حتى `max_iterations` مرة.

لحلقتك، استخدم الدالة `range` التي تولد متتالية من الأعداد يمكنك التكرار عليها. الصيغة النحوية هي `range(start, stop, step)`، حيث `start` هو العدد الصحيح الابتدائي (شامل)، و`stop` هو العدد الصحيح الأخير (غير شامل)، و`step` هو الفرق بين العدد والعدد السابق في المتتالية.

أيضًا، استخدم `_` كمتغير للحلقة. يعمل `_` كعنصر نائب وهو مفيد عندما تحتاج إلى استخدام متغير لكنك لا تحتاج فعليًا إلى قيمته.

# --hints--

يجب أن تنشئ حلقة `for` للتكرار على `range(max_iterations)`. استخدم `_` كمتغير للحلقة.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
