---
id: 65ef190c6b51e9b5a5f7ed29
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

أعطِ الدالة `square_root_bisection` المعلمات التالية:

- `square_target`: الرقم الذي تريد إيجاد الجذر التربيعي له.
- `tolerance` (اختياري): الفرق المقبول بين مربع القيمة التقريبية للجذر والقيمة الهدف الفعلية (الافتراضي هو `1e-7`). تعني قيمة التحمل `1e-7` أن الحل سيكون دقيقًا ضمن `0.0000001` من القيمة الحقيقية، وهو اختيار افتراضي جيد يوازن بين الدقة والأداء.
- `max_iterations` (اختياري): الحد الأقصى لعدد التكرارات التي تُنفذ (الافتراضي هو `100`). إذا لم تتقارب الطريقة ضمن هذا الحد، ستفترض أن الحل غير موجود.
  
# --hints--

يجب أن تحتوي دالتك على هذه المعلمات: `square_target`، و`tolerance = 1e-7`، و`max_iterations = 100`. ترتيبها مهم.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```

