---
id: 65ef1d5e3d2927c5e0f4997b
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

في بايثون، الكلمة المفتاحية `is` تتحقق من هوية الكائن. تُستخدم لتحديد ما إذا كان متغيران يشيران إلى نفس الكائن في الذاكرة. على عكس `is`، معامل المساواة (`==`) يحدد ما إذا كانت قيم كائنين متساوية، بغض النظر عن كونهما نفس الكائن في الذاكرة أم لا.

خارج حلقة for، أنشئ جملة `if` للتحقق مما إذا كان `root is None`. إذا كان كذلك، اطبع الرسالة `'Failed to converge within {max_iterations} iterations.'`. تذكر تنسيق الرسالة باستخدام سلسلة نصية f-string.

# --hints--

يجب أن تنشئ جملة `if` تتحقق مما إذا كان `root is None`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

يجب أن تمرر المعلمة `f'Failed to converge within {max_iterations} iterations.'` إلى دالة `print`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
