---
id: 587d78af367417b2b2512b00
title: استخدام خاصية align-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--

الخاصية الأخيرة لعناصر المرونة هي `align-self`. تتيح لك هذه الخاصية تعديل محاذاة كل عنصر بشكل فردي، بدلاً من تعيينها جميعًا مرة واحدة. هذا مفيد لأن تقنيات التعديل الشائعة الأخرى باستخدام خصائص CSS مثل `float` و`clear` و`vertical-align` لا تعمل على العناصر المرنة.

تقبل خاصية `align-self` نفس القيم التي تقبلها `align-items` وستتجاوز أي قيمة تم تعيينها بواسطة خاصية `align-items`.

# --instructions--

أضف خاصية CSS `align-self` إلى كل من `#box-1` و`#box-2`. أعطِ `#box-1` القيمة `center` وأعطِ `#box-2` القيمة `flex-end`.

# --hints--

يجب أن يحتوي عنصر `#box-1` على خاصية `align-self` بقيمة `center`.

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

يجب أن يحتوي عنصر `#box-2` على خاصية `align-self` بقيمة `flex-end`.

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
