---
id: 587d78ae367417b2b2512afd
title: استخدام خاصية flex-basis لتعيين الحجم الابتدائي للعنصر
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c3d9nCa'
forumTopicId: 301108
dashedName: use-the-flex-basis-property-to-set-the-initial-size-of-an-item
---

# --description--

تحدد خاصية `flex-basis` الحجم الابتدائي للعنصر قبل أن يُجري CSS تعديلات باستخدام `flex-shrink` أو `flex-grow`.

الوحدات المستخدمة في خاصية `flex-basis` هي نفسها وحدات خصائص الحجم الأخرى (`px`، `em`، `%`، إلخ). القيمة `auto` تحدد حجم العناصر بناءً على المحتوى.

# --instructions--

عيّن الحجم الابتدائي للصناديق باستخدام `flex-basis`. أضف خاصية CSS `flex-basis` لكل من `#box-1` و `#box-2`. اجعل قيمة `#box-1` تساوي `10em` وقيمة `#box-2` تساوي `20em`.

# --hints--

يجب أن يحتوي العنصر `#box-1` على خاصية `flex-basis`.

```js
const boxOne = document.querySelector('#box-1');
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

يجب أن تكون قيمة `flex-basis` للعنصر `#box-1` هي `10em`.

```js
assert.match(code, /#box-1\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?10em;/g);
```

يجب أن يحتوي العنصر `#box-2` على خاصية `flex-basis`.

```js
const boxTwo = document.querySelector('#box-2');
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

يجب أن تكون قيمة `flex-basis` للعنصر `#box-2` هي `20em`.

```js
assert.match(code, /#box-2\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?20em;/g);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-basis: 10em;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-basis: 20em;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
