---
id: 587d78ab367417b2b2512af0
title: "استخدام display: flex لوضع صندوقين"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cgz3QS7'
forumTopicId: 301105
dashedName: use-display-flex-to-position-two-boxes
---

# --description--

يستخدم هذا القسم أنماط تحدٍّ متبادلة لشرح كيفية استخدام CSS لوضع العناصر بطريقة مرنة. أولًا، يشرح تحدٍّ النظريات، ثم يطبق تحدٍّ عملي مفهوم flexbox باستخدام مكوِّن تغريدة بسيط.

وضع خاصية CSS `display: flex;` على عنصر يسمح لك باستخدام خصائص flex الأخرى لبناء صفحة متجاوبة.

# --instructions--

أضف خاصية CSS `display` إلى `#box-container` واضبط قيمتها على `flex`.

# --hints--

يجب أن تحتوي `#box-container` على خاصية `display` بقيمة `flex`.

```js
const boxContainer = document.querySelector('#box-container');
const displayStyle = window.getComputedStyle(boxContainer)['display'];
assert.strictEqual(displayStyle, 'flex');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    height: 500px;

  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    height: 500px;
    display: flex;
  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
