---
id: 587d78ae367417b2b2512afe
title: استخدام الخاصية المختصرة flex
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--

هناك اختصار متاح لضبط عدة خصائص flex مرة واحدة. يمكن ضبط خصائص `flex-grow` و `flex-shrink` و `flex-basis` كلها معًا باستخدام الخاصية `flex`.

على سبيل المثال، `flex: 1 0 10px;` يضبط العنصر ليكون `flex-grow: 1;` و `flex-shrink: 0;` و `flex-basis: 10px;`.

الإعدادات الافتراضية للخاصية هي `flex: 0 1 auto;`.

# --instructions--

أضف خاصية CSS `flex` لكل من `#box-1` و `#box-2`. اجعل قيم `#box-1` بحيث تكون `flex-grow` تساوي `2`، و `flex-shrink` تساوي `2`، و `flex-basis` تساوي `150px`. اجعل قيم `#box-2` بحيث تكون `flex-grow` تساوي `1`، و `flex-shrink` تساوي `1`، و `flex-basis` تساوي `150px`.

هذه القيم تجعل `#box-1` ينمو ليملأ المساحة الإضافية بمعدل ضعف `#box-2` عندما تكون الحاوية أكبر من 300px، وينكمش بمعدل ضعف `#box-2` عندما تكون الحاوية أقل من 300px. 300px هو الحجم المجمّع لقيم `flex-basis` للصندوقين.

# --hints--

يجب أن يكون للعنصر `#box-1` الخاصية `flex` مضبوطة على القيمة `2 2 150px`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

يجب أن يكون للعنصر `#box-2` الخاصية `flex` مضبوطة على القيمة `1 1 150px`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

يجب أن يستخدم شفرتك الخاصية `flex` لكل من `#box-1` و `#box-2`.

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
