---
id: 587d78ac367417b2b2512af4
title: استخدام خاصية flex-direction لإنشاء عمود
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cZmWeA4'
forumTopicId: 301109
dashedName: use-the-flex-direction-property-to-make-a-column
---

# --description--

استخدم التحديان الأخيران خاصية `flex-direction` مضبوطة على `row`. يمكن لهذه الخاصية أيضًا إنشاء عمود عن طريق تكديس العناصر الفرعية لحاوية مرنة عموديًا.

# --instructions--

أضف خاصية CSS `flex-direction` إلى العنصر `#box-container`، واجعل قيمتها `column`.

# --hints--

يجب أن يحتوي العنصر `#box-container` على خاصية `flex-direction` مضبوطة على `column`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: column;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
