---
id: 587d78ad367417b2b2512afb
title: استخدم خاصية flex-shrink لتقليص العناصر
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items
---

# --description--

حتى الآن، كل الخصائص في التحديات تنطبق على الحاوية المرنة (الأصل لعناصر المرنة). مع ذلك، هناك عدة خصائص مفيدة لعناصر المرنة.

الأولى هي خاصية `flex-shrink`. عند استخدامها، تسمح للعنصر بأن يتقلص إذا كانت الحاوية المرنة صغيرة جدًا. تتقلص العناصر عندما يكون عرض الحاوية الأصل أصغر من مجموع عرض كل العناصر المرنة بداخلها.

تأخذ خاصية `flex-shrink` أرقامًا كقيم. كلما زاد الرقم، زاد مقدار التقلص مقارنة بالعناصر الأخرى في الحاوية. على سبيل المثال، إذا كان لأحد العناصر قيمة `flex-shrink` تساوي `1` وللآخر قيمة `flex-shrink` تساوي `3`، فسيتقلص العنصر ذو القيمة `3` بمقدار ثلاثة أضعاف الآخر.

# --instructions--

أضف خاصية CSS `flex-shrink` إلى كل من `#box-1` و `#box-2`. أعطِ `#box-1` قيمة `1` و `#box-2` قيمة `2`.

# --hints--

يجب أن يحتوي عنصر `#box-1` على خاصية `flex-shrink` بقيمة `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

يجب أن يحتوي عنصر `#box-2` على خاصية `flex-shrink` بقيمة `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
