---
id: 587d78ad367417b2b2512af8
title: محاذاة العناصر باستخدام خاصية align-items
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c8aggtk'
forumTopicId: 301101
dashedName: align-elements-using-the-align-items-property
---

# --description--

خاصية `align-items` تشبه خاصية `justify-content`. تذكر أن خاصية `justify-content` كانت تُحاذي العناصر المرنة على طول المحور الرئيسي. بالنسبة للصفوف، المحور الرئيسي هو خط أفقي، وللأعمدة هو خط عمودي.

تحتوي الحاويات المرنة أيضًا على **المحور المتقاطع** وهو عكس المحور الرئيسي. بالنسبة للصفوف، المحور المتقاطع عمودي، وللأعمدة المحور المتقاطع أفقي.

يقدم CSS خاصية `align-items` لمحاذاة العناصر المرنة على طول المحور المتقاطع. بالنسبة للصف، تخبر CSS كيف تدفع العناصر في الصف بأكمله لأعلى أو لأسفل داخل الحاوية. وبالنسبة للعمود، كيف تدفع كل العناصر لليسار أو اليمين داخل الحاوية.

القيم المختلفة المتاحة لـ `align-items` تشمل:

<ul><li><code>flex-start</code>: يحاذي العناصر إلى بداية الحاوية المرنة. بالنسبة للصفوف، يحاذي العناصر إلى أعلى الحاوية. بالنسبة للأعمدة، يحاذي العناصر إلى يسار الحاوية.</li><li><code>flex-end</code>: يحاذي العناصر إلى نهاية الحاوية المرنة. بالنسبة للصفوف، يحاذي العناصر إلى أسفل الحاوية. بالنسبة للأعمدة، يحاذي العناصر إلى يمين الحاوية.</li><li><code>center</code>: يحاذي العناصر إلى المنتصف. بالنسبة للصفوف، يحاذي العناصر عموديًا (مساحة متساوية فوق وتحت العناصر). بالنسبة للأعمدة، يحاذيها أفقيًا (مساحة متساوية على اليسار واليمين من العناصر).</li><li><code>stretch</code>: يمدد العناصر لملء الحاوية المرنة. على سبيل المثال، يتمدد عناصر الصفوف لملء الحاوية من الأعلى إلى الأسفل. هذه هي القيمة الافتراضية إذا لم تُحدد قيمة لـ <code>align-items</code>.</li><li><code>baseline</code>: يحاذي العناصر إلى خطوط الأساس الخاصة بها. خط الأساس هو مفهوم نصي، فكر فيه كأنه الخط الذي تجلس عليه الحروف.</li></ul>

# --instructions--

مثال يساعد في توضيح هذه الخاصية أثناء العمل. أضف خاصية CSS `align-items` إلى العنصر `#box-container`، وأعطه القيمة `center`.

**مكافأة**  
جرّب الخيارات الأخرى لخاصية `align-items` في محرر الشفرة لترى الفروقات بينها. لكن لاحظ أن القيمة `center` هي الوحيدة التي ستجتاز هذا التحدٍّ.

# --hints--

يجب أن يحتوي العنصر `#box-container` على خاصية `align-items` بقيمة `center`.

```js
const boxContainer = document.querySelector('#box-container');
const alignment = window.getComputedStyle(boxContainer)['align-items'];
assert.strictEqual(alignment, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    align-items: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```
