---
id: 587d78ad367417b2b2512afa
title: استخدام خاصية flex-wrap لتغليف صف أو عمود
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cQv9ZtG'
forumTopicId: 301114
dashedName: use-the-flex-wrap-property-to-wrap-a-row-or-column
---

# --description--

تحتوي خاصية flexbox في CSS على ميزة لتقسيم الحاوية المرنة إلى عدة صفوف (أو أعمدة). افتراضيًا، تتسع الحاوية المرنة لجميع العناصر المرنة معًا. على سبيل المثال، يكون الصف كله في سطر واحد.

لكن استخدام الخاصية `flex-wrap` يخبر CSS بتغليف العناصر. هذا يعني أن العناصر الزائدة تنتقل إلى صف أو عمود جديد. تعتمد نقطة التوقف التي يحدث عندها التغليف على حجم العناصر وحجم الحاوية.

تحتوي CSS أيضًا على خيارات لاتجاه التغليف:

<ul><li><code>nowrap</code>: هذا هو الإعداد الافتراضي، ولا يغلف العناصر.</li><li><code>wrap</code>: يغلف العناصر على عدة أسطر من الأعلى إلى الأسفل إذا كانت في صفوف، ومن اليسار إلى اليمين إذا كانت في أعمدة.</li><li><code>wrap-reverse</code>: يغلف العناصر على عدة أسطر من الأسفل إلى الأعلى إذا كانت في صفوف، ومن اليمين إلى اليسار إذا كانت في أعمدة.</li></ul>

# --instructions--

التخطيط الحالي يحتوي على صناديق كثيرة جدًا لصف واحد. أضف خاصية CSS `flex-wrap` إلى عنصر `#box-container`، وأعطه القيمة `wrap`.

# --hints--

يجب أن يحتوي عنصر `#box-container` على الخاصية `flex-wrap` بقيمة `wrap`.

```js
const boxContainer = document.querySelector('#box-container');
const flexWrap = window.getComputedStyle(boxContainer)['flex-wrap'];
assert.strictEqual(flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;
    flex-wrap: wrap;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```
