---
id: 587d78ae367417b2b2512aff
title: استخدام خاصية order لإعادة ترتيب العناصر
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--

تُستخدم خاصية `order` لإخبار CSS بترتيب ظهور العناصر المرنة في الحاوية المرنة. افتراضيًا، تظهر العناصر بنفس الترتيب الذي تأتي به في شفرة HTML المصدرية. تأخذ الخاصية أرقامًا كقيم، ويمكن استخدام الأرقام السالبة.

# --instructions--

أضف خاصية CSS `order` لكل من `#box-1` و `#box-2`. أعطِ `#box-1` القيمة `2` وأعطِ `#box-2` القيمة `1`.

# --hints--

يجب أن يكون للعنصر `#box-1` خاصية `order` بقيمة `2`.

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

يجب أن يكون للعنصر `#box-2` خاصية `order` بقيمة `1`.

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
