---
id: 587d78ae367417b2b2512afc
title: استخدم خاصية flex-grow لتوسيع العناصر
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c2p78cg'
forumTopicId: 301111
dashedName: use-the-flex-grow-property-to-expand-items
---

# --description--

العكس من `flex-shrink` هو خاصية `flex-grow`. تذكر أن `flex-shrink` تتحكم في حجم العناصر عندما تتقلص الحاوية. خاصية `flex-grow` تتحكم في حجم العناصر عندما تتوسع الحاوية الأم.

باستخدام مثال مشابه من التحدٍّ السابق، إذا كان لأحد العناصر قيمة `flex-grow` تساوي `1` وللآخر قيمة `flex-grow` تساوي `3`، فإن العنصر الذي له القيمة `3` سينمو بمقدار ثلاثة أضعاف العنصر الآخر.

# --instructions--

أضف خاصية CSS `flex-grow` لكل من `#box-1` و `#box-2`. أعطِ `#box-1` القيمة `1` و `#box-2` القيمة `2`.

# --hints--

يجب أن يكون للعنصر `#box-1` خاصية `flex-grow` بقيمة `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
assert.equal(flexGrow, '1');
```

يجب أن يكون للعنصر `#box-2` خاصية `flex-grow` بقيمة `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
assert.equal(flexGrow, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-grow: 1;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-grow: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
