---
id: 587d78ab367417b2b2512af2
title: استخدام خاصية flex-direction لإنشاء صف
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row
---

# --description--

إضافة `display: flex` إلى عنصر يحوله إلى حاوية مرنة. هذا يجعل من الممكن محاذاة أي عناصر فرعية لذلك العنصر في صفوف أو أعمدة. تفعل ذلك بإضافة خاصية `flex-direction` إلى العنصر الأب وتعيينها إلى row أو column. إنشاء صف سيُحاذي العناصر الفرعية أفقيًا، وإنشاء عمود سيُحاذيها عموديًا.

خيارات أخرى لـ `flex-direction` هي `row-reverse` و `column-reverse`.

**ملاحظة:** القيمة الافتراضية لخاصية `flex-direction` هي `row`.

# --instructions--

أضف خاصية CSS `flex-direction` إلى عنصر `#box-container`، وعيّن قيمتها إلى `row-reverse`.

# --hints--

يجب أن يحتوي عنصر `#box-container` على خاصية `flex-direction` مضبوطة على `row-reverse`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
