---
id: 69008059d43d6751385883c1
title: Debuggen von Fehlern bei Inkrement- und Dekrementoperatoren in einer fehlerhaften App
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

Ihr Freund hat Sie gebeten, ihm beim Debuggen seines Codes zu helfen. Er übt gerade Inkrement- und Dekrementoperatoren, erhält aber unerwartete Ergebnisse in der Konsole.

Ihre Aufgabe ist es, ihm zu helfen, die Fehler zu erkennen und den Code so zu aktualisieren, dass die Ausgaben in der Konsole den erwarteten Ergebnissen entsprechen.

Sie können auf die vorherigen Lektionen zu Inkrement- und Dekrementoperatoren sowie zur Präfix- und Postfixnotation zurückgreifen.

**User Stories:**

1. Sie sollten eine Variable namens `updatedCounter` haben, die durch Verwendung des Inkrementoperators korrekt den Wert `11` erzeugt. Aktualisieren Sie den Ausdruck, um die richtige Präfix- oder Postfixnotation zu verwenden, anstatt den Wert fest zu codieren.
2. Sie sollten eine Variable namens `finalScore` haben, die durch Verwendung des Inkrementoperators korrekt den Wert `8` erzeugt. Aktualisieren Sie den Ausdruck, um die richtige Präfix- oder Postfixnotation zu verwenden, anstatt den Wert fest zu codieren.
3. Sie sollten eine Variable namens `updatedCoins` haben, die durch Verwendung des Dekrementoperators korrekt den Wert `2` erzeugt. Aktualisieren Sie den Ausdruck, um die richtige Präfix- oder Postfixnotation zu verwenden, anstatt den Wert fest zu codieren.
4. Sie sollten eine Variable namens `newHealth` haben, die durch Verwendung des Dekrementoperators korrekt den Wert `7` erzeugt. Aktualisieren Sie den Ausdruck, um die richtige Präfix- oder Postfixnotation zu verwenden, anstatt den Wert fest zu codieren.

# --hints--

Sie sollten eine Variable namens `updatedCounter` haben, die die Präfix-Inkrementnotation (`++counter`) verwendet.

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` sollte gleich `11` sein.

```js
assert.strictEqual(updatedCounter, 11);
```

Sie sollten eine Variable namens `finalScore` haben, die die Postfix-Inkrementnotation (`score++`) verwendet.

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` sollte gleich `8` sein.

```js
assert.strictEqual(finalScore, 8);
```

Sie sollten eine Variable namens `updatedCoins` haben, die die Präfix-Dekrementnotation (`--coins`) verwendet.

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` sollte gleich `2` sein.

```js
assert.strictEqual(updatedCoins, 2);
```

Sie sollten eine Variable namens `newHealth` haben, die die Postfix-Dekrementnotation (`health--`) verwendet.

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` sollte gleich `7` sein.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
