---
id: 68e6698b29e562d21b1d260c
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستبني قائمة بأهم متصفحات الويب. تم توفير قالب جاهز لـ HTML لك.

ابدأ بإضافة عنوان إلى صفحتك يقرأ `List of Major Web Browsers` باستخدام عنصر `h1` داخل عنصر `body`.

# --hints--

يجب أن يكون لديك عنصر `h1` داخل عنصر `body`.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

يجب أن يحتوي عنصر `h1` الخاص بك على النص `List of Major Web Browsers`.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
