---
id: 68ecfa1003d88990bb96a33e
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

حان الوقت الآن لإضافة المتصفح الثاني إلى القائمة.

أسفل إدخالك لـ `Google Chrome`، أضف عنصر `dt` آخر يحتوي على النص `Firefox`.

أسفل عنصر `dt` الخاص بك، أضف عنصر `dd` بالنص التالي:

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

يجب أن يكون لديك عنصران `dt` في شفرتك.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

يجب أن يحتوي عنصر `dt` الثاني على النص `Firefox`.

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

يجب أن يكون لديك عنصران `dd` في شفرتك.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

يجب أن يحتوي عنصر `dd` الثاني على النص `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
  <head> 
    <meta charset="UTF-8"> 
    <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
    <title>List of Browsers and Descriptions</title> 
  </head> 
  <body> 
    <h1>List of Major Web Browsers</h1> 
    <dl>
      <dt>Google Chrome</dt>
      <dd>This is a free web browser developed by Google and first released in 2008.</dd>

        --fcc-editable-region--
      
        --fcc-editable-region--
    </dl>
  </body> 
</html>
```

