---
id: 68ecfa0f03d88990bb96a33c
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

كما تتذكر من درس سابق، تُستخدم قوائم الوصف لعرض المصطلحات والتعريفات بطريقة منظمة وسهلة القراءة.

إليك مثال:

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

أسفل عنصر `h1`، أنشئ عنصر `dl`. سيحتوي هذا على قائمة المتصفحات.

# --hints--

يجب أن يكون لديك عنصر `dl`.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

يجب أن يكون عنصر `dl` الخاص بك أسفل عنصر `h1`.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
  <head> 
    <meta charset="UTF-8"> 
    <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
    <title>List of Browsers and Descriptions</title> 
  </head> 
  <body> 
    <h1>List of Major Web Browsers</h1> 
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body> 
</html>
```

