---
id: 68ecfa1003d88990bb96a33d
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

قوائم الوصف تتكون من قائمة من المصطلحات والتفاصيل. عنصر `dt` يُمثل مصطلح الوصف، بينما عنصر `dd` يُمثل تفاصيل الوصف.

أول متصفح ستضيفه إلى قائمة الوصف هو Google Chrome.

داخل عنصر `dl` الخاص بك، أضف عنصر `dt` يحتوي على النص `Google Chrome`.

أسفل عنصر `dt` الخاص بك، أضف عنصر `dd` بالنص التالي:

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

يجب أن يكون لديك عنصر `dt` داخل عنصر `dl` الخاص بك.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

يجب أن يحتوي عنصر `dt` الخاص بك على النص `Google Chrome`.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

يجب أن يكون لديك عنصر `dd` أسفل عنصر `dt` الخاص بك.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

يجب أن يحتوي عنصر `dd` الخاص بك على النص `This is a free web browser developed by Google and first released in 2008.`

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```

