---
id: 6853e58d52ec64349a5760ca
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستنفذ خوارزمية أقصر مسار. ستكتب دالة Python تحسب أقصر مسار بين العقد في رسم بياني، وتُرجع أيضًا المسار المتبع.

على سبيل المثال، إذا كان هناك رسم بياني حيث المدن متصلة بطرق بمسافات مختلفة، ستجد الخوارزمية أقصر مسار من مدينة إلى أخرى. إذا أردت السفر من المدينة A إلى المدينة D، قد تجد الخوارزمية أن المسار A ⇨ B ⇨ C ⇨ D (الإجمالي: 15 كم) أقصر من الذهاب مباشرة A ⇨ D (20 كم).

لتبدأ، عرّف متغيرًا اسمه `INF` وعيّن له القيمة `float('inf')`، التي تمثل اللانهاية الموجبة. ستستخدمه لاحقًا للدلالة على مسافة لا نهائية بين عقدتين.

# --hints--

يجب تعريف المتغير `INF` في أعلى ملفك.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

يجب تعيين `float('inf')` للمتغير `INF`.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
