---
id: 686251fb49ebc489ba60ccb5
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

ستحتاج إلى إنشاء قائمة ثنائية الأبعاد لتمثيل المصفوفة المجاورة للرسم البياني. ستُستخدم هذه المصفوفة لتمثيل أوزان الحواف بين العقد في الرسم البياني.

أنشئ متغيرًا باسم `adj_matrix` وقم بتعيين قائمة ثنائية الأبعاد تمثل الرسم البياني بالأوزان التالية:

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

يجب إنشاء متغير باسم `adj_matrix`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

يجب أن تكون `adj_matrix` قائمة ثنائية الأبعاد تحتوي على القيم المقدمة.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
