---
id: 686253f1b051998ad4904e3e
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

الآن ستنشئ الدالة الرئيسية التي تستقبل ثلاث معلمات: المصفوفة المجاورة، العقدة الابتدائية، وعقدة هدف اختيارية.

أنشئ دالة باسم `shortest_path` تستقبل ثلاث معلمات: `matrix`، `start_node`، و`target_node`. عيّن القيمة الافتراضية `None` للمعلمة `target_node`.

المعلمة `target_node` مُعيّنة إلى `None` افتراضيًا، مما يعني أنه إذا لم تُحدد عقدة هدف، يجب على الدالة حساب أقصر المسارات من العقدة الابتدائية إلى جميع العقد الأخرى في الرسم البياني.

أضف جملة `pass` داخل جسم الدالة في الوقت الحالي.

# --hints--

يجب أن تنشئ دالة باسم `shortest_path`.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

يجب أن تستقبل الدالة ثلاث معلمات: `matrix`، `start_node`، و`target_node` مع تعيين القيمة الافتراضية `None` للمعلمة `target_node`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
