---
id: 686b82af323098e7a5203bf4
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

أثناء تنفيذ الخوارزمية، يجب أن تتابع العقد التي زرتها بالفعل حتى لا تعالجها أكثر من مرة.

للقيام بذلك، أنشئ قائمة باسم `visited` وابدأها بالقيمة `False` لكل عقدة.

# --hints--

يجب أن تحتوي دالة `shortest_path` على قائمة باسم `visited`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("visited")`
        )
      );
    }
});
```

ابدأ قائمة `visited` بالقيمة `False` لكل عقدة باستخدام التعبير `[False] * n`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("visited").is_equivalent("visited = [False] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
--fcc-editable-region--
    
--fcc-editable-region--
```
