---
id: 686b7848fbaa41e0c2a7efef
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

الآن بعد أن أعددت قائمة `distances`، عليك تحديث المسافة للعقدة البداية.

بما أن المسافة من العقدة البداية إلى نفسها دائمًا `0`، عيّن القيمة عند فهرس `start_node` في قائمة `distances` إلى `0`.

# --hints--

يجب أن تعيّن `distances[start_node]` إلى `0`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_stmt("distances[start_node] = 0")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
--fcc-editable-region--
    
--fcc-editable-region--
```
