---
id: 686b78bc13d518e1267448e4
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

بالإضافة إلى تتبع المسافات، تحتاج أيضًا إلى تتبع المسارات الفعلية التي تم اتخاذها للوصول إلى كل عقدة.

ستنشئ قائمة يخزن كل عنصر فيها المسار الذي تم اتخاذه للوصول إلى تلك العقدة. في البداية، سيحتوي مسار كل عقدة على نفسها فقط.

توفر تراكيب القوائم طريقة مختصرة لإنشاء القوائم. على سبيل المثال:

```py
[x * 2 for x in range(3)]
```

أنشئ متغيرًا باسم `paths` وابدأ تهيئته باستخدام تركيب قائمة ينشئ قائمة تحتوي على `[node_no]` لكل `node_no` في `range(n)`.

# --hints--

يجب أن يكون لديك متغير يسمى `paths`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

يجب أن تهيئ `paths` باستخدام تركيب قائمة.

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

يجب أن يستخدم تركيب القائمة الخاص بك `node_no` للتكرار على `range(n)`.

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

يجب أن يقيم تركيب القائمة الخاص بك `[node_no]` لكل `node_no` في `range(n)`.

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
