---
id: 6876089f7a500b7c782e5b83
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

يجب أن تقرر ما إذا كانت العقدة الحالية خيارًا أفضل من التي وجدتها بالفعل (إن وجدت). للقيام بذلك، أضف جملة شرطية داخل الحلقة.

يجب أن تقوم الحالة بأمرين:

- التحقق مما إذا لم تُزر العقدة بعد.
- مقارنة المسافة المعروفة مع `min_distance` الحالية.

داخل حلقة `for` الداخلية، أضف جملة `if` تتحقق مما إذا لم تُزر `node_no` **و** ما إذا كانت `distances[node_no]` أقل من `min_distance`.

أضف `pass` كعنصر نائب داخل كتلة الشرط.

# --hints--

يجب أن تضيف جملة `if` داخل حلقة `for` الداخلية.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

يجب أن تتحقق جملة `if` مما إذا كانت العقدة غير مزارة (`visited[node_no]` قيمته خاطئة) وأن `distances[node_no]` أقل من `min_distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
    
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
--fcc-editable-region--
            pass
--fcc-editable-region--
```
