---
id: 68760464eb9c2e79cc913fac
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

في هذه الخطوة، ستضيف حلقة تُنفذ مرة واحدة لكل عقدة في الرسم البياني. تتيح هذه الحلقة للخوارزمية تحديث المسافات والمسارات عبر عدة تمريرات.

أنشئ حلقة `for` تُنفذ `n` مرات. استخدم `_` كمتغير للحلقة لأنك لا تحتاج إلى استخدام قيمة التكرار.

داخل الحلقة، تحتاج إلى التحضير لاختيار العقدة التالية للمعالجة بإنشاء متغيرين:

- واحد للاحتفاظ بأصغر مسافة تم العثور عليها حتى الآن في التكرار الحالي  
- وآخر لتخزين فهرس العقدة التي تملك هذه المسافة الأصغر.

أنشئ المتغيرين `min_distance` و `current`، واضبطهما على `INF` و `-1` على التوالي.

# --hints--

يجب أن تحتوي دالة `shortest_path` على حلقة `for`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

يجب أن تتكرر حلقة `for` الخاصة بك على `range(n)`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

يجب أن تستخدم حلقة for المتغير `_` كمتغير التكرار.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

يجب أن تنشئ متغيرًا باسم `min_distance` وتضبطه على `INF` داخل حلقة for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

يجب أن تنشئ متغيرًا باسم `current` وتضبطه على `-1` داخل حلقة for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
    
--fcc-editable-region--
    
--fcc-editable-region--
```
