---
id: 68760671d950767b055bf2f9
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

الآن عليك فحص كل عقدة لإيجاد العقدة التي تملك أصغر مسافة معروفة ولم تُزر بعد.

للقيام بذلك، أضف حلقة `for` داخل الحلقة الرئيسية.

يجب أن تتكرر الحلقة على كل `node_no` في `range(n)`، حيث تمثل `n` عدد العقد في الرسم البياني. أضف `pass` كجسم للحلقة مؤقتًا.

# --hints--

يجب أن يكون لديك حلقة `for` ثانية داخل دالة `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 1
      `)
      );
    }
});
```

يجب أن تتكرر حلقة `for` الخاصة بك على `range(n)` مع `node_no` كمتغير التكرار.

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
inner_loop = outer_loop.find_bodies()[0].find_for_loops()[0]
assert inner_loop.find_for_vars().is_equivalent("node_no")
assert inner_loop.find_for_iter().is_equivalent("range(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
    
    for _ in range(n):  
        min_distance = INF  
        current = -1
--fcc-editable-region--
        
--fcc-editable-region--
```
