---
id: 68777816e95d4936f129819f
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

قبل محاولة تحديث المسافة إلى الجار، تحقق أولًا من أن الجار يمكن الوصول إليه ولم تتم زيارته بعد. بعد ذلك، احسب المسافة الإجمالية للوصول إلى ذلك الجار عبر العقدة الحالية.

داخل حلقة `for`، أضف عبارة `if` تتحقق من:

- أن `distance` لا تساوي `INF` (مما يعني وجود حافة بين العقد)
- أن الجار `node_no` لم تتم زيارته بعد

داخل كتلة `if`، أنشئ متغيرًا باسم `new_distance` وعيّنه مجموع `distances[current]` (أقصر مسافة إلى العقدة الحالية) و`distance` (المسافة من العقدة الحالية إلى الجار).

# --hints--

يجب أن تحتوي على عبارة `if` تتحقق مما إذا كانت `distance` لا تساوي `INF` وإذا لم تتم زيارة عقدة الجار.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

يجب أن تنشئ متغيرًا باسم `new_distance` داخل كتلة `if`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

يجب أن تعيّن لـ `new_distance` مجموع `distances[current]` و`distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
    
    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no
        
        if current == -1:
            break
        visited[current] = True
        
        for node_no in range(n):  
            distance = matrix[current][node_no]
--fcc-editable-region--
            
--fcc-editable-region--
```
