---
id: 687b3926419dec576850b814
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

بعد أن حسبت المسافة الجديدة الممكنة إلى الجار، تحقق مما إذا كانت أفضل من المسافة المخزنة حاليًا في قائمة `distances`. إذا كانت كذلك، حدّث المسافة.

داخل كتلة `if` الموجودة، أضف عبارة `if` تتحقق مما إذا كانت `new_distance` أقل من `distances[node_no]`.

داخل هذه الكتلة الشرطية الجديدة، حدّث `distances[node_no]` لتخزين `new_distance`.

# --hints--

يجب أن يكون لديك عبارة `if` مسكنة تتحقق مما إذا كانت `new_distance` أقل من `distances[node_no]`.

```js
({
    test: () => {
      runPython(`  
      cond = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]  
      assert cond.is_equivalent("new_distance < distances[node_no]") or cond.is_equivalent("distances[node_no] > new_distance")  
      `);
    }
});
```

يجب أن تحدّث `distances[node_no]` لتخزين `new_distance`.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("distances[node_no] = new_distance")
      `)
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
    
    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no
        
        if current == -1:  
            break  
        visited[current] = True  
        
        for node_no in range(n):
            distance = matrix[current][node_no]  
            if distance != INF and not visited[node_no]:  
                new_distance = distances[current] + distance
--fcc-editable-region--
                
--fcc-editable-region--
```
