---
id: 68e4c0ff9edcc92771d7549a
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

في الوقت الحالي، تحتوي قائمة `distances` على عنصر واحد فقط. مع ذلك، تحتاج إلى قيمة مسافة واحدة لكل عقدة في الرسم البياني.

في بايثون، يمكنك ضرب قائمة بعدد صحيح لتكرارها. على سبيل المثال، `[0] * 3` تُنشئ `[0, 0, 0]`.

حدّث قائمة `distances` لتحتوي على `n` نسخ من `INF` بضرب `[INF]` في `n`.

# --hints--

يجب تحديث `distances` لتحتوي على `n` نسخ من `INF`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
--fcc-editable-region--
    distances = [INF]
--fcc-editable-region--
```
