---
id: 6998b99a8e195cd0c1f211df
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop bauen Sie einen einfachen Wildlife Tracker mit JavaScript-Objekten.

Hier ist ein Beispiel für ein Objekt:

```js
const animal = {}; // empty object
```

Dies erstellt ein leeres Objekt.

Erstellen Sie nun eine Variable mit dem Namen `tiger` und weisen Sie ihr ein leeres Objekt zu.

# --hints--

Sie sollten eine Variable mit dem Namen `tiger` erstellen.

```js
assert.exists(tiger);
```

Der Variablen `tiger` sollte ein leeres Objekt zugewiesen werden.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
