---
id: 6999a0c1698b93d5d41156ed
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Fügen Sie dem `tiger`-Objekt jetzt eine weitere Eigenschaft namens `isEndangered` hinzu.

Setzen Sie die `isEndangered`-Eigenschaft auf `true`.

# --hints--

`tiger` sollte eine `isEndangered`-Eigenschaft haben.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` sollte gleich `true` sein.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
