---
id: 6999b0d79ea938edcb032237
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

In diesem letzten Schritt des Wildlife Tracker Workshops erstellen Sie eine Funktion, die auf eine Eigenschaft eines Objekts mit Klammernotation zugreift.

Hier ist ein Beispiel, wie man mit Klammernotation auf eine Eigenschaft zugreift:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

In diesem Beispiel ermöglicht die Klammernotation den Zugriff auf eine Eigenschaft mithilfe einer Variablen.

Erstellen Sie nun eine Funktion namens `getProperty`.

Die Funktion soll zwei Parameter annehmen: `animal` und `propertyName`.

Geben Sie den Wert der Eigenschaft mit Klammernotation zurück.

Nachdem Sie die Funktion erstellt haben, verwenden Sie `console.log`, um `getProperty(tiger, "species")` und `getProperty(elephant, "age")` aufzurufen, damit Sie die zurückgegebenen Werte in der Konsole sehen können.

Damit ist der Wildlife Tracker Workshop abgeschlossen!

# --hints--

Sie sollten eine Funktion mit dem Namen `getProperty` erstellen.

```js
assert.isFunction(getProperty);
```

Die Funktion `getProperty` sollte zwei Parameter haben: `animal` und `propertyName`.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` sollte Klammernotation verwenden, um auf die Eigenschaft zuzugreifen.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

Sie sollten `getProperty(tiger, "species")` in der Konsole ausgeben.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

Sie sollten `getProperty(elephant, "age")` in der Konsole ausgeben.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

Das Aufrufen von `getProperty(tiger, "species")` sollte `"Tiger"` zurückgeben.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

Das Aufrufen von `getProperty(elephant, "age")` sollte `12` zurückgeben.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` sollte die Funktionsparameter verwenden und mit jedem Objekt funktionieren.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
