---
id: 6999adc72c02812a28c01e31
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

In diesem Schritt erstellen Sie eine Funktion, die eine Eigenschaft eines Objekts aktualisiert.

Hier ist ein Beispiel, wie man eine Eigenschaft innerhalb einer Funktion aktualisiert:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
};

console.log(changeAge(dog, 6)); // { age: 6 }
```

In diesem Beispiel wird die `age`-Eigenschaft auf einen neuen Wert aktualisiert.

Erstellen Sie nun eine Funktion namens `updateAge`. Die Funktion soll zwei Parameter entgegennehmen: `animal` und `newAge`.

Aktualisieren Sie innerhalb der Funktion die `age`-Eigenschaft des `animal`-Objekts auf `newAge`. Geben Sie das aktualisierte `animal`-Objekt zurück.

Nachdem Sie die Funktion erstellt haben, verwenden Sie `console.log`, um `updateAge(elephant, 12)` aufzurufen, damit Sie das aktualisierte `elephant`-Objekt in der Konsole sehen können.

# --hints--

Sie sollten eine Funktion mit dem Namen `updateAge` erstellen.

```js
assert.isFunction(updateAge);
```

Die `updateAge`-Funktion sollte zwei Parameter haben: `animal` und `newAge`.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

Die `updateAge`-Funktion sollte das aktualisierte `animal`-Objekt zurückgeben.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

Sie sollten `updateAge(elephant, 12)` in der Konsole ausgeben.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

Das Aufrufen von `updateAge(elephant, 12)` sollte die Eigenschaft age aktualisieren.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` sollte die Funktionsparameter verwenden und mit jedem Objekt arbeiten.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
