---
id: 6999ac171928136168654774
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

In diesem Schritt erstellen Sie eine Funktion, um auf die `species` Eigenschaft eines Objekts zuzugreifen.

Hier ist ein kleines Beispiel, wie man auf eine Eigenschaft eines Objekts zugreift:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

In diesem Beispiel nimmt die Funktion ein Objekt als Parameter und gibt dessen `species` Eigenschaft zurück.

Erstellen Sie nun eine Funktion namens `getSpecies`.

Die Funktion soll einen Parameter namens `animal` erhalten.

Geben Sie die `species` Eigenschaft von `animal` mit Punktnotation zurück.

Nachdem Sie die Funktion erstellt haben, verwenden Sie `console.log`, um `getSpecies(tiger)` aufzurufen, damit Sie das Ergebnis in der Konsole sehen können.

# --hints--

Sie sollten eine Funktion namens `getSpecies` erstellen.

```js
assert.isFunction(getSpecies);
```

Die `getSpecies` Funktion sollte einen einzelnen Parameter namens `animal` haben.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Sie sollten `getSpecies(tiger)` in der Konsole ausgeben.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` sollte mit Punktnotation auf die Eigenschaft species zugreifen.

```js
assert.match(code, /animal\.species/);
```

Das Aufrufen von `getSpecies(tiger)` sollte `"Tiger"` zurückgeben.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` sollte den Funktionsparameter verwenden und mit jedem Objekt funktionieren.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
