---
id: 6999b07c34c2cde21ee6e1d8
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

In diesem Schritt erstellen Sie eine Funktion, die überprüft, ob ein Objekt eine bestimmte Eigenschaft besitzt.

Hier ist ein Beispiel für die Verwendung von `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

In diesem Beispiel verwendet die `hasColor`-Funktion `hasOwnProperty`, um zu prüfen, ob die `color`-Eigenschaft in einem Objekt vorhanden ist.

Erstellen Sie nun eine Funktion namens `hasHabitat`. Die Funktion soll einen Parameter namens `animal` entgegennehmen.

Verwenden Sie die `hasOwnProperty`-Methode, um `true` zurückzugeben, wenn das `animal`-Objekt eine Eigenschaft namens `"habitat"` besitzt, und `false` andernfalls.

Nachdem Sie die Funktion erstellt haben, verwenden Sie `console.log`, um `hasHabitat(tiger)` und `hasHabitat(elephant)` aufzurufen, damit Sie beide Ergebnisse in der Konsole sehen können.

# --hints--

Sie sollten eine Funktion namens `hasHabitat` erstellen.

```js
assert.isFunction(hasHabitat);
```

Die `hasHabitat`-Funktion sollte einen einzelnen Parameter namens `animal` haben.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` sollte die `hasOwnProperty`-Methode verwenden, um die `habitat`-Eigenschaft zu überprüfen.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

Sie sollten `hasHabitat(tiger)` in der Konsole ausgeben.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

Sie sollten `hasHabitat(elephant)` in der Konsole ausgeben.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

Das Aufrufen von `hasHabitat(tiger)` sollte `true` zurückgeben, nachdem ein Lebensraum hinzugefügt wurde.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Das Aufrufen von `hasHabitat(elephant)` sollte `false` zurückgeben, wenn kein Lebensraum existiert.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` sollte die Funktionsparameter verwenden und mit jedem Objekt funktionieren.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
