---
id: 69999ff7f7862fdfb885ea5d
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Aktualisieren Sie nun Ihr `tiger`-Objekt, indem Sie eine neue Eigenschaft namens `age` hinzufügen.

Setzen Sie die `age`-Eigenschaft auf den Wert `5`.

# --hints--

`tiger` sollte eine `age`-Eigenschaft haben.

```js
assert.property(tiger, "age");
```

`tiger.age` sollte gleich `5` sein.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
