---
id: 6999b0198d18d402c36cc5fe
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

In diesem Schritt erstellen Sie eine Funktion, die eine Eigenschaft aus einem Objekt entfernt.

Hier ist ein Beispiel, wie man eine Eigenschaft mit `delete` entfernt:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

In diesem Beispiel wird die Eigenschaft `canFly` aus dem Objekt `bird` entfernt.

Erstellen Sie nun eine Funktion namens `removeEndangeredStatus`. Die Funktion soll einen Parameter namens `animal` erhalten.

Entfernen Sie innerhalb der Funktion die Eigenschaft `isEndangered` aus dem Objekt `animal` mit dem Schlﾃｼsselwort `delete`. Geben Sie das aktualisierte Objekt `animal` zurﾃｼck.

Nachdem Sie die Funktion erstellt haben, verwenden Sie `console.log`, um `removeEndangeredStatus(tiger)` aufzurufen, damit Sie das aktualisierte Objekt in der Konsole sehen kﾃｶnnen.

# --hints--

Sie sollten eine Funktion namens `removeEndangeredStatus` erstellen.

```js
assert.isFunction(removeEndangeredStatus);
```

Die Funktion `removeEndangeredStatus` sollte einen einzelnen Parameter `animal` haben.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Die Funktion `removeEndangeredStatus` sollte das aktualisierte Objekt `animal` zurﾃｼckgeben.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` sollte das Schlﾃｼsselwort `delete` verwenden, um die Eigenschaft `isEndangered` zu entfernen.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

Sie sollten `removeEndangeredStatus(tiger)` in der Konsole ausgeben.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

Das Aufrufen von `removeEndangeredStatus(tiger)` sollte die Eigenschaft `isEndangered` entfernen.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` sollte die Funktionsparameter verwenden und mit jedem Objekt arbeiten.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
