---
id: 6999faf886634b72be7a819b
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Für diesen Wildtier-Tracker aktualisieren Sie ein Objekt mit neuen Eigenschaften.

Hier ist ein Beispiel, wie man ein Objekt mit einer Eigenschaft erstellt:

```js
const animal = {
  name: "Lion" // property: value
};
```

Aktualisieren Sie nun das `tiger`-Objekt, sodass es eine Eigenschaft namens `species` enthält.

Setzen Sie deren Wert auf `"Tiger"`.

# --hints--

`tiger` sollte ein Objekt sein.

```js
assert.isObject(tiger);
```

`tiger` sollte eine `species`-Eigenschaft haben.

```js
assert.property(tiger, "species");
```

`tiger.species` sollte gleich `"Tiger"` sein.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
