---
id: 699a00de9f564bb0effaa14d
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

In diesem Schritt erstellen Sie eine Funktion, um auf die `age` Eigenschaft eines Objekts zuzugreifen.

Hier ist ein Beispiel:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

In diesem Beispiel nimmt die Funktion ein Objekt als Parameter und gibt dessen `age` Eigenschaft zurück.

Erstellen Sie nun eine Funktion namens `getAge`.

Die Funktion soll einen Parameter namens `animal` erhalten.

Geben Sie die `age` Eigenschaft von `animal` mit Punktnotation zurück.

Nachdem Sie die Funktion erstellt haben, verwenden Sie `console.log`, um `getAge(tiger)` aufzurufen, damit Sie das Ergebnis in der Konsole sehen können.

# --hints--

Sie sollten eine Funktion namens `getAge` erstellen.

```js
assert.isFunction(getAge);
```

Die `getAge` Funktion sollte einen einzelnen Parameter haben, `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` sollte auf die `age` Eigenschaft mit Punktnotation zugreifen.

```js
assert.match(code, /animal\.age/);
```

Sie sollten `getAge(tiger)` in der Konsole ausgeben.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

Das Aufrufen von `getAge(tiger)` sollte `5` zurückgeben.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` sollte den Funktionsparameter verwenden und mit jedem Objekt funktionieren.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
