---
id: 69ba8da6b296391d913096a0
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Erstellen Sie nun ein zweites Objekt namens `elephant`.

FĂĽgen Sie die folgenden Eigenschaften hinzu:

- `species` mit dem Wert `"Elephant"`
- `age` mit dem Wert `10`
- `isEndangered` mit dem Wert `true`

# --hints--

Sie sollten eine Variable namens `elephant` erstellen.

```js
assert.exists(elephant);
```

`elephant` sollte eine `species`-Eigenschaft mit dem Wert `"Elephant"` haben.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` sollte eine `age`-Eigenschaft mit dem Wert `10` haben.

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` sollte eine `isEndangered`-Eigenschaft mit dem Wert `true` haben.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
