---
id: 66cc21d23238dc8240a8a182
title: Passo 16
challengeType: 1
dashedName: step-16
---

# --description--

Seria bom que a calculadora pudesse calcular o quadrado de um número. O quadrado de um número é o número multiplicado por ele mesmo.

Para calcular o quadrado de um número em JavaScript, você pode usar o método <dfn>Math.pow()</dfn>. O método `Math.pow()` aceita dois argumentos: o número base e o expoente. 

Você também pode usar o operador de exponenciação (`**`) para calcular o quadrado de um número.

Aqui está um exemplo:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

Declare uma função chamada `calculateSquare` que aceita um parâmetro `num` e retorna o quadrado de `num`.

# --hints--

Você deve ter uma função chamada `calculateSquare`.

```js
assert.isFunction(calculateSquare);
```

A função `calculateSquare` deve aceitar um parâmetro `num`.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

A função `calculateSquare` deve retornar o quadrado de `num`.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
