---
id: 66cc172af7159a7c67804544
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Agora é hora de adicionar a funcionalidade de multiplicação à calculadora.

Declare uma função chamada `calculateProduct` que receba dois parâmetros, `num1` e `num2`.

Dentro da função, retorne o produto dos dois parâmetros.

Em seguida, adicione um `console.log` que chame a função `calculateProduct` com os argumentos `13` e `5`.

# --hints--

Você deve declarar uma função chamada `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

A função `calculateProduct` deve ter dois parâmetros chamados `num1` e `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateProduct } = explorer.allFunctions;
assert.lengthOf(calculateProduct?.parameters, 2);
assert.equal(calculateProduct?.parameters[0].toString(), "num1");
assert.equal(calculateProduct?.parameters[1].toString(), "num2");
```

A função `calculateProduct` deve retornar o produto dos dois parâmetros.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

Você deve ter um `console.log` que chame a função `calculateProduct` com os argumentos `13` e `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateProduct(13,5))", "console.log(calculateProduct(5,13))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
