---
id: 66cc12fa504b0479dac479a0
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

A vantagem das funções é que podem ser chamadas com diferentes argumentos, permitindo reutilizar o mesmo código com vários valores.

Aqui estão alguns exemplos de chamadas de funções com diferentes argumentos de string:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

Adicione um `console.log` que chame a função `calculateSum` com os argumentos `2` e `5`.

# --hints--

Você deve ter um `console.log` que chame a função `calculateSum` com os argumentos `2` e `5`.

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
