---
id: 66cc16abc30b367ba1d1f9ef
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Agora é hora de testar sua função `calculateDifference`.

Comece adicionando um `console.log` que chama a função `calculateDifference` com os argumentos `22` e `5`.

Então, adicione um segundo `console.log` que chama a função `calculateDifference` com os argumentos `12` e `1`.

Finalmente, adicione um terceiro `console.log` que chama a função `calculateDifference` com os argumentos `17` e `9`.

# --hints--

Você deve ter um `console.log` que chama a função `calculateDifference` com os argumentos `22` e `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(22,5))";
assert.include(strippedCode, expectedStatement);
```

Você deve ter um `console.log` que chama a função `calculateDifference` com os argumentos `12` e `1`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(12,1))";
assert.include(strippedCode, expectedStatement);
```

Você deve ter um `console.log` que chama a função `calculateDifference` com os argumentos `17` e `9`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(17,9))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
