---
id: 66cc2739f687e484d50bb6f1
title: Passo 18
challengeType: 1
dashedName: step-18
---

# --description--

A última funcionalidade da calculadora será o cálculo de uma raiz quadrada.

Um quadrado é um número multiplicado por si mesmo.Assim, a raiz quadrada seria o número que, quando multiplicado por si mesmo, dá o número original.

Por exemplo, a raiz quadrada de `9` é `3` porque `3 * 3 = 9`.

Para obter a raiz quadrada de um número em JavaScript, você pode usar o método <dfn>Math.sqrt()</dfn>.

Aqui está um exemplo:

```js
// result: 3
Math.sqrt(9);
```

Declare uma função chamada `calculateSquareRoot` que recebe um parâmetro `num` e retorna a raiz quadrada de `num`.

# --hints--

Você deve ter uma função chamada `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

A função `calculateSquareRoot` deve receber um parâmetro `num`.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

A função `calculateSquareRoot` deve retornar a raiz quadrada de `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
