---
id: 66cc1a3a39aef47d6473cb2f
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

O próximo passo é adicionar a funcionalidade de divisão à calculadora.

Declare uma função chamada `calculateQuotient` que receba dois parâmetros, `num1` e `num2`.

Dentro da função, retorne o quociente dos dois parâmetros.

Em seguida, adicione um `console.log` que chame a função `calculateQuotient` com os argumentos `7` e `11`.

# --hints--

Você deve declarar uma função chamada `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

A função `calculateQuotient` deve ter dois parâmetros chamados `num1` e `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateQuotient } = explorer.allFunctions;
assert.lengthOf(calculateQuotient?.parameters, 2);
assert.equal(calculateQuotient?.parameters[0].toString(), "num1");
assert.equal(calculateQuotient?.parameters[1].toString(), "num2");
```

A função `calculateQuotient` deve retornar o quociente dos dois parâmetros.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

Você deve ter um `console.log` que chame a função `calculateQuotient` com os argumentos `7` e `11`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateQuotient(7,11))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
