---
id: 66cc0f1ae40802781b2ea972
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Nas lições anteriores, você aprendeu como trabalhar com parâmetros de função.

Um parâmetro de função é uma variável definida na declaração da função e que atua como um placeholder.

Aqui está um exemplo de uma função que tem um parâmetro:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Declare uma função chamada `calculateSum`, que receba dois parâmetros, `num1` e `num2`.

# --hints--

Você deve declarar uma função chamada `calculateSum`.

```js
assert.isFunction(calculateSum);
```

A função `calculateSum` deve receber dois parâmetros, `num1` e `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateSum } = explorer.allFunctions;
assert.lengthOf(calculateSum?.parameters, 2);
assert.equal(calculateSum?.parameters[0].toString(), "num1");
assert.equal(calculateSum?.parameters[1].toString(), "num2");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
