---
id: 66cc0ba5881acb7692cfc4de
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora é hora de adicionar outra função.

Declare uma função chamada `addThreeAndFour` que retorne a soma de `3` e `4`.

Em seguida, chame a função `addThreeAndFour` dentro de um `console.log` para ver o resultado.

# --hints--

Você deve declarar uma função chamada `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

A função `addThreeAndFour` deve retornar a soma de `3` e `4`.

```js
assert.strictEqual(addThreeAndFour(), 7);
```

Sua função `addThreeAndFour` deve retornar a soma usando o operador `+`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

Você deve chamar a função `addThreeAndFour` dentro de um `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
