---
id: 66cc095dfe1682753d2ab030
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Um dos conceitos que você aprendeu em lições anteriores foi como retornar valores de uma função.

Aqui está um lembrete de como retornar um valor de uma função:

```js
function myFunction() {
  return "Hello World";
}
```

Dentro da função `addTwoAndSeven`, retorne a soma de `2` e `7`.

# --hints--

A função `addTwoAndSeven` deve retornar a soma de `2` e `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Sua função `addTwoAndSeven` deve retornar a soma usando o operador `+`.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
const validExpressions = ["2 + 7", "7 + 2", "(2 + 7)", "(7 + 2)"];
assert.isTrue(validExpressions.some(expr => addTwoAndSeven?.hasReturn(expr)));
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
