---
id: 66cc0ba5881acb7692cfc4de
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora, é hora de adicionar outra função.

Declare uma função chamada `addThreeAndFour` que retorne a soma de `3` e `4`.

Em seguida, chame a função `addThreeAndFour` dentro de um `console.log` para ver o resultado.

# --hints--

Você deve declarar uma função chamada `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

A função `addThreeAndFour` deve retornar a soma de `3` e `4`.

```js
assert.equal(addThreeAndFour(), 7);
```

Your `addThreeAndFour` function should return the sum using the `+` operator.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

You should call the `addThreeAndFour` function inside a `console.log`.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());
--fcc-editable-region--

--fcc-editable-region--
```
