---
id: 66cc0f1ae40802781b2ea972
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

In previous lessons, you learned how to work with function parameters.

Um parâmetro de função é uma variável definida na declaração da função e que atua como um placeholder.

Aqui está um exemplo de uma função que tem um parâmetro:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Declare uma função chamada `calculateSum`, que receba dois parâmetros, `num1` e `num2`.

# --hints--

Você deve declarar uma função chamada `calculateSum`.

```js
assert.isFunction(calculateSum);
```

A função `calculateSum` deve receber dois parâmetros, `num1` e `num2`.

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
