---
id: 66cc0a9e06e00b75d6782be9
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Nas lições anteriores, você aprendeu como chamar (invocar) uma função. Chamar uma função significa executar o código dentro da função.

Aqui está um lembrete de como chamar uma função:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Adicione uma instrução `console.log` e, dentro dela, chame a função `addTwoAndSeven`.

Agora você deve ver a soma de `2` e `7` registrada no console.

# --hints--

Você deve chamar a função `addTwoAndSeven` dentro de um `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addTwoAndSeven())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
