---
id: 66cc10958d72e57907f898b2
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Dentro da função `calculateSum`, você precisará retornar a soma dos dois parâmetros.

# --hints--

A função `calculateSum` deve retornar a soma dos dois parâmetros.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function calculateSum(num1, num2) {
  
}
--fcc-editable-region--
```
