---
id: 66cdf76685e4cb5a8726e27b
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Para testar ainda mais a função `calculateSum`, você precisará chamá-la com argumentos diferentes.

Adicione outro `console.log` que chame a função `calculateSum` com os argumentos `10` e `10`.

Abaixo desse `console.log`, adicione outro `console.log` que chame a função `calculateSum` com os argumentos `5` e `5`.

# --hints--

Você deve ter um `console.log` que chama a função `calculateSum` com os argumentos `10` e `10`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateSum(10,10))";
assert.include(strippedCode, expectedStatement);
```

Você deve ter um `console.log` que chama a função `calculateSum` com os argumentos `5` e `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateSum(5,5))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
--fcc-editable-region--

--fcc-editable-region--
```
