---
id: 0095d0567b6817114989a1dd
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви реалізуєте алгоритм пошуку найкоротшого шляху на JavaScript. Ви напишете функцію, яка обчислює найкоротший шлях між вузлами у зваженому графі та повертає як відстані, так і пройдені маршрути.

Наприклад, якщо є граф, де міста з’єднані дорогами з різною довжиною, алгоритм знайде найкоротший маршрут від одного міста до іншого. Якщо ви хочете поїхати з міста A до міста D, алгоритм може визначити, що шлях A ⇨ B ⇨ C ⇨ D (усього: 15 км) коротший за прямий A ⇨ D (20 км).

Щоб почати, визначте змінну з іменем `INF` і призначте їй значення `Infinity`. Пізніше ви використаєте її, щоб позначити, що між двома вузлами немає прямого з’єднання.

# --hints--

Ви повинні визначити змінну з іменем `INF`.

```js
assert.isDefined(INF);
```

Значення вашої змінної `INF` має бути `Infinity`.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
