---
id: 69c918841987bf16c1d1babb
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Тепер ви створите головну функцію, яка виконує всю роботу.

Створіть функцію з іменем `shortestPath`, яка приймає три параметри: `matrix`, `startNode` та `targetNode`.

Протягом наступних кількох кроків ви розробите логіку цієї функції.

# --hints--

Ви повинні створити функцію з іменем `shortestPath`.

```js
assert.isFunction(shortestPath);
```

Ваша функція `shortestPath` має мати параметри `matrix`, `startNode` та `targetNode`.

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
