---
id: 69c918841987bf16c1d1babd
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Вам потрібно відстежувати найкоротшу відому відстань від початкового вузла до кожного іншого вузла. Спочатку припустіть, що кожен вузол знаходиться на нескінченно великій відстані.

У JavaScript можна створити масив, попередньо заповнений певним значенням, використовуючи `new Array(n).fill(value)`. Наприклад, `new Array(3).fill(0)` створює `[0, 0, 0]`.

Усередині функції `shortestPath` створіть змінну з ім’ям `distances` і призначте їй `new Array(n).fill(INF)`.

# --hints--

Ви повинні створити змінну з ім’ям `distances` всередині `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

Ваша змінна `distances` має бути призначена `new Array(n).fill(INF)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
