---
id: 69c918841987bf16c1d1babe
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Відстань від початкового вузла до самого себе завжди дорівнює `0`. Ви маєте оновити масив `distances`, щоб це відобразити.

Після масиву `distances` встановіть `distances[startNode]` рівним `0`.

# --hints--

Ви повинні встановити `distances[startNode]` рівним `0`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*startNode\s*\]\s*=\s*0/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);

}
--fcc-editable-region--
```
