---
id: 69c918841987bf16c1d1babc
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Вам потрібно зберегти число вузлів у графі.

Всередині функції `shortestPath` створіть змінну з ім’ям `n` і призначте їй `matrix.length`.

# --hints--

Ви повинні створити змінну `n` всередині функції `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=/
);
```

Значення вашої змінної `n` має бути `matrix.length`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=\s*matrix\.length/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {

}
--fcc-editable-region--
```
