---
id: 69c918841987bf16c1d1bac5
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

Після сканування всіх вузлів, щоб знайти найближчий невідвіданий, потрібно обробити випадок, коли жодного дійсного вузла не знайдено. Це трапляється, коли всі залишкові вузли недосяжні.

Після внутрішнього `for` циклу, але ще всередині зовнішнього циклу, додайте `if` інструкцію, яка перевіряє, чи `current` строго дорівнює `-1`. Якщо це правда, використайте `break` всередині блоку `if`, щоб раніше зупинити зовнішній цикл.

# --hints--

Ви повинні додати `if` інструкцію, яка перевіряє, чи `current === -1`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

Ви повинні використати `break` всередині блоку `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
