---
id: 69c918841987bf16c1d1babf
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Вам також потрібно відстежувати фактичний шлях, який пройшов до кожного вузла. Ви збережете список шляхів, де кожен запис — це масив індексів вузлів, що представляє пройдений маршрут.

Спочатку шлях кожного вузла міститиме лише його самого. Ви можете створити цю структуру, використовуючи `Array.from()`:

```js
Array.from({ length: n }, (_, i) => [i]);
```

Це створює масив з `n` елементів, де елемент з індексом `i` є `[i]`.

Все ще в межах функції `shortestPath`, створіть змінну з ім’ям `paths` і призначте їй `Array.from({ length: n }, (_, i) => [i])`.

# --hints--

Ви повинні створити змінну `paths` всередині функції `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

Ваша змінна `paths` має бути ініціалізована за допомогою `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

Ваша змінна `paths` має використовувати `{ length: n }` як перший аргумент для `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

Другим аргументом для `Array.from` має бути функція зворотного виклику `(_, i) => [i]`, де `_` — це заповнювач для невикористаного елемента, а `i` — індекс. Вона має повертати `[i]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
