---
id: 69c918841987bf16c1d1bac8
title: Крок 17
challengeType: 1
dashedName: step-17
---

# --description--

Перед оновленням відстаней потрібно перевірити, чи варто враховувати сусіда. Має існувати реальне ребро до нього (`distance !== INF`) і він не повинен бути ще відвіданим.

Всередині циклу `for` по сусідах додайте інструкцію `if`, яка перевіряє, чи `distance` не строго дорівнює `INF` **і** чи `nodeNo` не відвіданий.

Всередині блоку `if` оголосіть змінну `newDistance`, призначену `distances[current] + distance`. Це загальна вартість досягнення сусіда через поточний вузол.

# --hints--

Ви повинні додати інструкцію `if` з умовою `distance !== INF && !visited[nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

Ви повинні оголосити змінну `newDistance` зі значенням `distances[current] + distance` всередині блоку `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
