---
id: 69c918841987bf16c1d1bac7
title: Крок 16
challengeType: 1
dashedName: step-16
---

# --description--

Тепер, коли поточний вузол позначено як відвіданий, потрібно переглянути всіх його сусідів, щоб перевірити, чи можна досягти їх ефективніше через поточний вузол.

Після `visited[current] = true` додайте цикл сусідів `for`, який ітерує кожен `nodeNo` від `0` до `n`. Зверніть увагу, що це ще один внутрішній цикл на тому ж рівні, що й перший.

Всередині циклу оголосіть змінну з ім’ям `distance` і призначте їй `matrix[current][nodeNo]`. Це дасть вам вагу ребра між поточним вузлом і його сусідом.

# --hints--

Ви повинні додати цикл сусідів `for` після позначення поточного вузла як відвіданого.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

Цикл сусідів `for` має ітерувати від `0` до `n`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

Усередині циклу сусідів має бути змінна `distance`, призначена `matrix[current][nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
