---
id: 69c918841987bf16c1d1bac1
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Тепер ви додасте головний цикл, який керує алгоритмом. Він запускається один раз для кожного вузла в графі, щоразу вибираючи найближчий невідвіданий вузол.

Всередині функції `shortestPath` додайте цикл `for`, який виконується `n` разів, використовуючи `i` як змінну циклу.

На початку тіла циклу, використовуючи `let`, оголосіть дві змінні `minDistance` і `current`, встановлені відповідно в `INF` і `-1`. Вони відстежуватимуть найближчий невідвіданий вузол, знайдений у кожній ітерації.

# --hints--

У вас має бути цикл `for` всередині `shortestPath`, який виконується `n` разів. Не забудьте встановити `i` зі значенням `let` всередині нього.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

У вас має бути змінна `minDistance`, встановлена в `INF` всередині циклу.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

У вас має бути змінна `current`, встановлена в `-1` всередині циклу.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
