---
id: 69c918841987bf16c1d1baba
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Вам знадобиться двовимірний масив, щоб представити матрицю суміжності графа. Кожне значення представляє вагу (відстань) ребра між двома вузлами. Значення `INF` означає, що прямого ребра немає.

Створіть іншу змінну з назвою `adjMatrix` і призначте їй двовимірний масив із такими значеннями:

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Ви повинні створити змінну з назвою `adjMatrix`.

```js
assert.isDefined(adjMatrix);
```

`adjMatrix` має бути двовимірним масивом з 6 рядків і 6 стовпців.

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix` має містити правильні значення.

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
