---
id: 69c918841987bf16c1d1bac3
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Всередині внутрішнього циклу `for` потрібно перевірити, чи поточний вузол не відвідувався і чи він ближчий, ніж найкращий, який ви знайшли досі.

Додайте інструкцію `if`, яка перевіряє, чи `nodeNo` не відвідувався і чи `distances[nodeNo]` менший за `minDistance`.

Поки що залиште тіло порожнім.

# --hints--

У вас має бути інструкція `if` з умовою `!visited[nodeNo] && distances[nodeNo] < minDistance` всередині внутрішнього циклу `for`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*!visited\s*\[\s*nodeNo\s*\]\s*&&\s*distances\s*\[\s*nodeNo\s*\]\s*<\s*minDistance\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
