---
id: 66c6491fe4c8e0f16845425f
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora es el momento de añadir otra fruta a la lista.

Usando el mismo método de array que antes, añade la cadena `"Grapes"` al final del array `shoppingList`.

Luego, añade un `console.log` y llama a la función `getShoppingListMsg` dentro del `console.log` para ver la lista actualizada en la consola.

# --hints--

Tu array `shoppingList` ahora debería contener dos elementos.

```js
assert.lengthOf(shoppingList, 2);
```

Tu array `shoppingList` debería contener la cadena `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

Debes llamar a la función `getShoppingListMsg(shoppingList)` dentro del `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
