---
id: 66c63bfa4685e4e3843fa0dc
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, continuarás aprendiendo sobre arrays construyendo una lista de compras de supermercado.

Comienza agregando un `console.log` que registre la cadena `"Grocery shopping list"` en la consola.

# --hints--

Debes tener un `console.log` en tu código.

```js
assert.match(code, /console\.log(.*)/);
```

Tu `console.log` debe mostrar el mensaje `"Grocery shopping list"`. Verifica nuevamente errores de ortografía y capitalización.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
