---
id: 66c644c08ebcd2ef90c750bd
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

En los vídeos de la clase anterior, aprendiste a agregar elementos al final de un array usando el método `push` así:

```js
fruits.push('pear');
```

Usando el método `push`, agrega la cadena `"Apples"` al array `shoppingList`.

# --hints--

Tu array `shoppingList` debe tener un elemento.

```js
assert.lengthOf(shoppingList, 1);
```

Tu array `shoppingList` debe tener la cadena `"Apples"`. Revisa dos veces si hay errores de ortografía y capitalización.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
