---
id: 66c730ee6ae076281721d0b9
title: Paso 14
challengeType: 1
dashedName: step-14
---

# --description--

En los videos de conferencia anteriores, aprendiste cómo eliminar elementos del final de un array usando el método `pop`.

Aquí tienes un recordatorio de cómo usar el método `pop`:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

Usa el método `pop` para eliminar el último elemento del array `shoppingList`.

# --hints--

El array `shoppingList` ahora debería tener `5` elementos en él.

```js
assert.lengthOf(shoppingList, 5);
```

El array `shoppingList` ya no debería tener la cadena `"Potato Chips"` en él.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
