---
id: 66c742d045c9fc2e09fa64b1
title: Paso 19
challengeType: 1
dashedName: step-19
---

# --description--

El último cambio por hacer en la lista de compras es actualizar el primer elemento de la lista.

En los videos de la conferencia anterior, aprendiste cómo actualizar un elemento utilizando la notación de corchetes y el índice del elemento que deseas actualizar.

Aquí tienes un recordatorio de cómo actualizar un elemento en un array:

```js
const array = [1, 2, 3, 4, 5];

array[0] = 10;
// Result: [10, 2, 3, 4, 5]
console.log(array); 
```

Actualiza el primer elemento en el array `shoppingList` a `"Canola Oil"`.

# --hints--

Debes actualizar el primer elemento en el array `shoppingList` a `"Canola Oil"`.

```js
assert.equal(shoppingList[0], "Canola Oil");
assert.notEqual(shoppingList[1], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

shoppingList.shift();

--fcc-editable-region--

--fcc-editable-region--
```
