---
id: 66c73a7798f6f62b2ae58f22
title: Paso 16
challengeType: 1
dashedName: step-16
---

# --description--

Ahora es tiempo de añadir más artículos al principio de la lista de compras.

Inicia añadiendo una sentencia `console.log` que registre el mensaje `"Debería comprar un postre."`

Debajo de esa declaración de `console`, usa el método correcto de array para añadir la cadena `"Chocolate Cake"` al principio del array `shoppingList`.

Finalmente, añade un `console.log` y llama a la función `getShoppingListMsg` dentro del `console.log` para ver la lista actualizada registrada en la consola.

# --hints--

Deberías tener un `console.log` que registre el mensaje `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

Tu array `shoppingList` debería tener un total de `6` elementos.

```js
assert.lengthOf(shoppingList, 6);
```

Deberías usar el método `unshift` para añadir la cadena `"Chocolate Cake"` al principio del array `shoppingList`.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

Deberías llamar a `getShoppingListMsg(shoppingList)` dentro del `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
