---
id: 66c726c34ecf1e238aa9d7d5
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

En los videos de la conferencia, aprendiste cómo agregar elementos al principio de un arreglo usando el método `unshift()`.

Aquí tienes un recordatorio de cómo utilizar el método `unshift()`:

```js
array.unshift(item1, item2, ..., itemX);
```

Utiliza el método `unshift()` para agregar la cadena `"Vegetable Oil"` al principio del arreglo `shoppingList`.

# --hints--

Debe tener tres elementos en tu arreglo `shoppingList`.

```js
assert.lengthOf(shoppingList, 3);
```

Debe utilizar el método `unshift()` para agregar `"Vegetable Oil"` al principio del arreglo `shoppingList`.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
