---
id: 66cbe2319d3845545a293a0b
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Ahora es momento de ver el mensaje registrado en la consola.

Añade un `console.log` y llama a la función `getShoppingListMsg` dentro del `console.log` para ver el mensaje registrado en la consola.

# --hints--

Debes llamar a la función `getShoppingListMsg(shoppingList)` dentro del `console.log`.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
