---
id: 5d822fd413a79914d39e98d6
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

이제 건물처럼 보이는 무언가가 생겼습니다. 첫 변수를 만들 준비가 되었습니다. 이전 수업에서 변수 선언은 두 개의 대시(`-`)로 시작하며, 이름과 값을 이렇게 지정한다는 것을 배웠습니다: `--variable-name: value;`. `bb1` 클래스 규칙 안에서 `--building-color1`라는 이름의 변수를 만들고 `#999` 값을 할당하세요.

# --hints--

`--building-color1`라는 새 변수를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

`--building-color1` 변수를 `.bb1` 안에 정의해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

`--building-color1`에 `#999` 값을 할당해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(), '#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
