---
id: 5d822fd413a79914d39e98cd
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

CSS에서는 별표(*)를 사용해 모든 요소를 선택할 수 있습니다. `*` 선택자를 사용하고 `border` 속성에 `1px solid black` 값을 주어 모든 요소에 테두리를 추가하세요. 이 방법은 요소가 어디에 있고 크기가 어떤지 시각화하는 데 도움이 됩니다. 나중에 이 부분은 제거할 것입니다.

# --hints--

`*` 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

모든 요소를 스타일링하려면 `border` 속성을 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

모든 요소에 `1px solid black` 테두리가 있어야 합니다.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
