---
id: 5d822fd413a79914d39e9917
title: 75단계
challengeType: 0
dashedName: step-75
---

# --description--

쉼표(`,`)로 구분하여 여러 개의 그라디언트를 요소에 추가할 수 있습니다. 예를 들어 다음과 같습니다.

```css
gradient1(
  colors
),
gradient2(
  colors
);
```

아래에 있는 `repeating-linear-gradient`에 `.fb1c`를 추가하세요; `--building-color4`에서 `0%`까지 `10%`를, 그리고 `--window-color4`과 `10%` 사이에 `90%`를 사용하세요. 이렇게 하면 마지막에 추가한 그라디언트 뒤를 채울 수 있습니다.

# --hints--

첫 번째 `repeating-linear-gradient`는 변경하지 마세요.

```js
const fb1c = __helpers.removeCssComments(code).match(/\.fb1c\s*\{([\s\S]*?)\}/);
assert.exists(fb1c);

const background = fb1c[1].match(/background\s*:\s*([^;]*)/);
assert.exists(background);

assert.match(background[1], /repeating-linear-gradient\(\s*90deg\s*,\s*var\(--building-color4\)\s*(?:0%)?\s*,\s*var\(--building-color4\)\s*10%\s*,\s*transparent\s*10%\s*,\s*transparent\s*15%\s*\)/);
```

첫 번째 `,` 뒤에는 쉼표(`repeating-linear-gradient`)가 있어야 합니다.

```js
const fb1c = __helpers.removeCssComments(code).match(/\.fb1c\s*\{([\s\S]*?)\}/);
assert.exists(fb1c);

const background = fb1c[1].match(/background\s*:\s*([^;]*)/);
assert.exists(background);

assert.match(background[1], /repeating-linear-gradient\(\s*90deg\s*,\s*var\(--building-color4\)\s*(?:0%)?\s*,\s*var\(--building-color4\)\s*10%\s*,\s*transparent\s*10%\s*,\s*transparent\s*15%\s*\)\s*,/);
```

`repeating-linear-gradient`에서 `--building-color4`까지 `0%` 색상을 가진 `10%`를 추가하세요.

```js
const fb1c = __helpers.removeCssComments(code).match(/\.fb1c\s*\{([\s\S]*?)\}/);
assert.exists(fb1c);

const background = fb1c[1].match(/background\s*:\s*([^;]*)/);
assert.exists(background);

assert.match(background[1], /repeating-linear-gradient\([\s\S]*?\)\s*,\s*repeating-linear-gradient\(\s*var\(--building-color4\)\s*(?:0%)?\s*,\s*var\(--building-color4\)\s*10%/);
```

`--window-color4`에서 `10%`까지 `90%` 두 번째 색상을 사용하세요.

```js
const fb1c = __helpers.removeCssComments(code).match(/\.fb1c\s*\{([\s\S]*?)\}/);
assert.exists(fb1c);

const background = fb1c[1].match(/background\s*:\s*([^;]*)/);
assert.exists(background);

assert.match(background[1], /repeating-linear-gradient\([\s\S]*?\)\s*,\s*repeating-linear-gradient\([\s\S]*?var\(--window-color4\)\s*10%\s*,\s*var\(--window-color4\)\s*90%/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1 building-wrap">
        <div class="bb1a bb1-window"></div>
        <div class="bb1b bb1-window"></div>
        <div class="bb1c bb1-window"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2">
        <div class="bb2a"></div>
        <div class="bb2b"></div>
      </div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4 building-wrap">
        <div class="bb4a"></div>
        <div class="bb4b"></div>
        <div class="bb4c window-wrap">
          <div class="bb4-window"></div>
          <div class="bb4-window"></div>
          <div class="bb4-window"></div>
          <div class="bb4-window"></div>
        </div>
      </div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div></div>
      <div></div>
      <div class="fb1 building-wrap">
        <div class="fb1a"></div>
        <div class="fb1b"></div>
        <div class="fb1c"></div>
      </div>
      <div class="fb2"></div>
      <div></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
  --window-color1: black;
  --window-color2: #8cd9b3;
  --window-color3: #d98cb3;
  --window-color4: #8cb3d9;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings, .foreground-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
  position: absolute;
  top: 0;
}

.building-wrap {
  display: flex;
  flex-direction: column;
  align-items: center;
}

.window-wrap {
  display: flex;
  align-items: center;
  justify-content: space-evenly;
}

/* BACKGROUND BUILDINGS - "bb" stands for "background building" */
.bb1 {
  width: 10%;
  height: 70%;
}

.bb1a {
  width: 70%;
}
  
.bb1b {
  width: 80%;
}
  
.bb1c {
  width: 90%;
}

.bb1d {
  width: 100%;
  height: 70%;
  background: linear-gradient(
      var(--building-color1) 50%,
      var(--window-color1)
    );
}

.bb1-window {
  height: 10%;
  background: linear-gradient(
      var(--building-color1),
      var(--window-color1)
    );
}

.bb2 {
  width: 10%;
  height: 50%;
}

.bb2a {
  border-bottom: 5vh solid var(--building-color2);
  border-left: 5vw solid transparent;
  border-right: 5vw solid transparent;
}

.bb2b {
  width: 100%;
  height: 100%;
  background: repeating-linear-gradient(
      var(--building-color2),
      var(--building-color2) 6%,
      var(--window-color2) 6%,
      var(--window-color2) 9%
    );
}

.bb3 {
  width: 10%;
  height: 55%;
  background: repeating-linear-gradient(
      90deg,
      var(--building-color3),
      var(--building-color3),
      var(--window-color3) 15%
    );
}

.bb4 {
  width: 11%;
  height: 58%;
}

.bb4a {
  width: 3%;
  height: 10%;
  background-color: var(--building-color4);
}

.bb4b {
  width: 80%;
  height: 5%;
  background-color: var(--building-color4);
}
  
.bb4c {
  width: 100%;
  height: 85%;
  background-color: var(--building-color4);
}

.bb4-window {
  width: 18%;
  height: 90%;
  background-color: var(--window-color4);
}

/* FOREGROUND BUILDINGS - "fb" stands for "foreground building" */
.fb1 {
  width: 10%;
  height: 60%;
}

.fb1b {
  width: 60%;
  height: 10%;
  background-color: var(--building-color4);
}
--fcc-editable-region--
.fb1c {
  width: 100%;
  height: 80%;
  background: repeating-linear-gradient(
      90deg,
      var(--building-color4),
      var(--building-color4) 10%,
      transparent 10%,
      transparent 15%
    )
}
--fcc-editable-region--
.fb2 {
  width: 10%;
  height: 40%;
  background-color: var(--building-color3);
}

.fb3 {
  width: 10%;
  height: 35%;
  background-color: var(--building-color1);
}
  
.fb4 {
  width: 8%;
  height: 45%;
  background-color: var(--building-color1);
  position: relative;
  left: 10%;
}

.fb5 {
  width: 10%;
  height: 33%;
  background-color: var(--building-color2);
  position: relative;
  right: 10%;
}

.fb6 {
  width: 9%;
  height: 38%;
  background-color: var(--building-color3);
}
    
```
