---
id: 5d822fd413a79914d39e98cc
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

CSS 변수 스카이라인 프로젝트에 오신 것을 환영합니다! `link` 요소 안에 `styles.css` 파일을 연결하는 `head` 요소를 추가하는 것부터 시작하세요.

# --hints--

코드에 `link` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('link'));
```

기존 `head` 태그를 변경하지 마세요. 닫는 태그를 삭제하지 않았는지 확인하세요.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

하나의 빈 `link` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 요소는 값이 `rel`인 `stylesheet` 속성을 가져야 합니다.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 요소는 값이 `href`인 `styles.css` 속성을 가져야 합니다.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
