---
id: 5d822fd413a79914d39e98d7
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

변수를 사용하려면 변수 이름을 소괄호 안에 넣고 그 앞에 `var`를 붙이세요. 예를 들어 `var(--variable-name)`처럼요. 변수에 지정한 값이 해당 변수를 사용하는 속성에 적용됩니다.

이전 단계에서 만든 변수 `--building-color1`를 `background-color` 클래스의 `.bb1a` 속성 값으로 추가하세요.

# --hints--

`background-color` 요소의 `.bb1a`가 설정되어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a')?.backgroundColor)
```

`var(--building-color1)`를 사용해 `background-color` 요소의 `.bb1a`를 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
