---
id: 5d822fd413a79914d39e98e7
title: Step 28
challengeType: 0
dashedName: step-28
---

# --description--

여섯 개의 새 요소에 다음 `width`와 `height` 값을 지정하세요: `10%`와 `60%`는 `.fb1`, `10%`와 `40%`는 `.fb2`, `10%`와 `35%`는 `.fb3`, `8%`와 `45%`는 `.fb4`, `10%`와 `33%`는 `.fb5`, `9%`와 `38%`는 `.fb6`입니다.

# --hints--

`.fb1` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb1'));
```

`.fb1` 선택자에 `width`를 `10%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.width, '10%');
```

`.fb1` 선택자에 `height`를 `60%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.height, '60%');
```

`.fb2` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb2'));
```

`.fb2` 선택자에 `width`를 `10%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.width, '10%');
```

`.fb2` 선택자에 `height`를 `40%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.height, '40%');
```

`.fb3` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb3'));
```

`.fb3` 선택자에 `width`를 `10%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.width, '10%');
```

`.fb3` 선택자에 `height`를 `35%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.height, '35%');
```

`.fb4` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb4'));
```

`.fb4` 선택자에 `width`를 `8%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.width, '8%');
```

`.fb4` 선택자에 `height`를 `45%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.height, '45%');
```

`.fb5` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb5'));
```

`.fb5` 선택자에 `width`를 `10%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.width, '10%');
```

`.fb5` 선택자에 `height`를 `33%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.height, '33%');
```

`.fb6` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb6'));
```

`.fb6` 선택자에 `width`를 `9%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.width, '9%');
```

`.fb6` 선택자에 `height`를 `38%`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.height, '38%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
}
--fcc-editable-region--

--fcc-editable-region--
    
```

