---
id: 5d822fd413a79914d39e98e6
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--

`div` 안에 여섯 개의 `.foreground-buildings` 요소를 중첩하고, 순서대로 `fb1`부터 `fb6`까지의 클래스를 부여하세요. "fb"는 "foreground building"의 약자로, 전경에 위치할 여섯 개의 건물을 의미합니다.

# --hints--

클래스가 `div`인 새로운 `fb1`를 생성하세요.

```js
assert.exists(document.querySelector('div.fb1'));
```

클래스가 `div`인 새로운 `fb2`를 생성하세요.

```js
assert.exists(document.querySelector('div.fb2'));
```

클래스가 `div`인 새로운 `fb3`를 생성하세요.

```js
assert.exists(document.querySelector('div.fb3'));
```

클래스가 `div`인 새로운 `fb4`를 생성하세요.

```js
assert.exists(document.querySelector('div.fb4'));
```

클래스가 `div`인 새로운 `fb5`를 생성하세요.

```js
assert.exists(document.querySelector('div.fb5'));
```

클래스가 `div`인 새로운 `fb6`를 생성하세요.

```js
assert.exists(document.querySelector('div.fb6'));
```

이 새로운 `div` 요소들을 `.foreground-buildings` 요소 안에 배치하세요.

```js
assert.exists(document.querySelector('div.foreground-buildings > div.fb1'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb2'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb3'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb4'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb5'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb6'));
```

새로운 `div` 요소들을 올바른 순서로 배치하세요.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.fb1','fb2') && __t('div.fb2','fb3') && __t('div.fb3','fb4') && __t('div.fb4', 'fb5') && __t('div.fb5', 'fb6'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
--fcc-editable-region--
    <div class="foreground-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
}

```

