---
id: 5d822fd413a79914d39e98d1
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

`.background-buildings` 요소에 `width` 값을 가진 `height`과 `100%`를 지정하여 부모인 `body`의 전체 너비와 높이를 차지하게 만드세요.

# --hints--

올바른 요소를 선택하려면 `background-buildings` 클래스를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

`.background-buildings` 요소는 `width` 값을 가진 `100%`을 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

`.background-buildings` 요소는 `height` 값을 가진 `100%`을 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--
    
```

