---
id: 66c057041df6394ca796bf33
title: 創建一個句子生成器
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

在本實驗中，你將使用句子模板創建兩個不同的故事。 你將使用變量來存儲故事的不同部分，然後將故事輸出到控制檯。

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶故事：**

1. 你應使用 `let` 聲明以下變量：

    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. 你應該將上述變量賦值爲你所選擇的字符串值。

1. 你應該聲明一個 `firstStory` 變量。

1. 你應使用以下故事模板創建第一個故事，並將其分配給 `firstStory` 變量：`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`；

1. 你應該將第一個故事輸出到控制檯，使用的信息是 `"First story: [firstStory]"`。

1. 你應該爲 `adjective`、`noun`、`verb`、`place`、`adjective2` 和 `noun2` 變量重新賦值。

1. 你應該聲明一個 `secondStory` 變量。

1. 使用相同的模板創建另一個故事，並將其賦值給 `secondStory` 變量。

1. 你應該將第二個故事輸出到控制檯，使用的信息是 `"Second story: [secondStory]"`。


# --hints--

你應該聲明一個 `adjective` 變量。

```js
assert.isNotNull(adjective);
```

你應該聲明一個 `noun` 變量。

```js
assert.isNotNull(noun);
```

你應該聲明一個 `verb` 變量。

```js
assert.isNotNull(verb);
```

你應該聲明一個 `place` 變量。

```js
assert.isNotNull(place);
```

你應該聲明一個 `adjective2` 變量。

```js
assert.isNotNull(adjective2);
```

你應該聲明一個 `noun2` 變量。

```js
assert.isNotNull(noun2);
```

你應該爲 `adjective` 變量賦值一個字符串。

```js
assert.isString(adjective);
```

你應該爲 `noun` 變量賦值一個字符串。

```js
assert.isString(noun);
```

你應該爲 `verb` 變量賦值一個字符串。

```js
assert.isString(verb);
```

你應該爲 `place` 變量賦值一個字符串。

```js
assert.isString(place);
```

你應該爲 `adjective2` 變量賦值一個字符串。

```js
assert.isString(adjective2);
```

你應該爲 `noun2` 變量賦值一個字符串。

```js
assert.isString(noun2);
```

你應該聲明一個 `firstStory` 變量。

```js
assert.isNotNull(firstStory);
```

你應該爲第一個故事使用正確的故事格式：`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`。 注意空格。

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

You should assemble your first story using the variables you declared in the correct order.

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

你應該將第一個故事輸出到控制檯，使用的信息是 `"First story: [firstStory]"`。

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

你應該聲明一個 `secondStory` 變量。

```js
assert.isNotNull(secondStory);
```

你應該爲第二個故事的 `adjective` 變量重新賦值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

你應該爲第二個故事的 `noun` 變量重新賦值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

你應該爲第二個故事的 `verb` 變量重新賦值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

你應該爲第二個故事的 `place` 變量重新賦值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

你應該爲第二個故事的 `adjective2` 變量重新賦值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

你應該爲第二個故事的 `noun2` 變量重新賦值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

你應該爲第二個故事使用正確的故事格式：`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`。 注意空格。

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

You should assemble your second story using the variables you declared in the correct order.

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

你應該將第二個故事輸出到控制檯，使用的格式是 `"Second story: [secondStory]"`。

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`firstStory` 不能和 `secondStory` 相同。

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```

