---
id: 66c057041df6394ca796bf33
title: 建置一個句子製造器
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

在本實驗中，你將使用句子模板創建兩個不同的故事。你將使用變數來儲存故事的不同零件，然後將故事輸出到主控台。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該使用 `let` 宣告以下變數：
   
    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. 你應該為上述變數指定一些你選擇的字串值。

1. 你應該宣告一個 `firstStory` 變數。
   
1. 你應該使用以下故事模板來創建第一個故事，並將它指定給 `firstStory` 變數：`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`；

1. 你應該使用訊息 `"First story: [firstStory]"` 將你的第一個故事輸出到主控台。

1. 你應該為你的 `adjective`、`noun`、`verb`、`place`、`adjective2` 和 `noun2` 變數指定新值。

1. 你應該宣告一個 `secondStory` 變數。

1. 使用相同的模板創建另一個故事，並將它指定給 `secondStory` 變數。

1. 你應該使用訊息 `"Second story: [secondStory]"` 將你的第二個故事輸出到主控台。


# --hints--

你應該宣告一個 `adjective` 變數。

```js
assert.isNotNull(adjective);
```

你應該宣告一個 `noun` 變數。

```js
assert.isNotNull(noun);
```

你應該宣告一個 `verb` 變數。

```js
assert.isNotNull(verb);
```

你應該宣告一個 `place` 變數。

```js
assert.isNotNull(place);
```

你應該宣告一個 `adjective2` 變數。

```js
assert.isNotNull(adjective2);
```

你應該宣告一個 `noun2` 變數。

```js
assert.isNotNull(noun2);
```

你應該為你的 `adjective` 變數指定字串值。

```js
assert.isString(adjective);
```

你應該為你的 `noun` 變數指定字串值。

```js
assert.isString(noun);
```

你應該為你的 `verb` 變數指定一個字串值。

```js
assert.isString(verb);
```

你應該為 `place` 變數指定一個字串值。

```js
assert.isString(place);
```

你應該為 `adjective2` 變數指定字串值。

```js
assert.isString(adjective2);
```

你應該為你的 `noun2` 變數指定字串值。

```js
assert.isString(noun2);
```

你應該宣告一個 `firstStory` 變數。

```js
assert.isNotNull(firstStory);
```

你應該為第一個故事使用正確的故事格式：`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`。請注意空格。

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

你應該使用你宣告的變數，按照正確的順序組合你的第一個故事。

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

你應該使用訊息 `"First story: [firstStory]"` 記錄你的第一個故事。

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

你應該宣告一個 `secondStory` 變數。

```js
assert.isNotNull(secondStory);
```

你應該為第二個故事重新指定 `adjective` 變數。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

你應該為第二個故事重新指定 `noun` 變數。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

你應該為第二個故事重新指定 `verb` 變數。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

你應該為第二個故事重新指定 `place` 變數。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

你應該為第二個故事重新指定 `adjective2` 變數。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

你應該為第二個故事重新指定 `noun2` 變數。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

你應該為第二個故事使用正確的故事格式：`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`。請注意空格。

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

你應該使用你宣告的變數，依正確順序組合你的第二個故事。

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

你應該使用格式 `"Second story: [secondStory]"` 記錄你的第二個故事。

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`firstStory` 不應該與 `secondStory` 相同。

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```

