---
id: 66eaddd04a9e533fba689001
title: створіть картку допису в блозі
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

У цій лабораторній роботі ви попрактикуєтеся у стилізації фонових зображень і кордонів, створюючи картку допису в блозі.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути `div` елемент із класом `blog-post-card`, який міститиме всі елементи вашої картки.
2. Усередині `.blog-post-card` елемента має бути зображення з дійсним `alt` атрибутом і текстом, а також класом `post-img`. Ви можете використати `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` для `src` атрибута зображення.
3. У `div` елементі має бути `post-content` з класом `.blog-post-card`, який містить:
    - `h2` елемент із класом `post-title` і текстом для заголовка вашого допису в блозі.
    - `p` елемент із класом `post-excerpt` і текстом для короткого опису вашого допису.
    - `a` елемент із класом `read-more` із текстом `Read More`.
4. Ви маєте встановити фоновий колір `body` елемента, відмінний від білого.
5. Застосуйте такі стилі до `.blog-post-card` елемента:
    - Білий фон.
    - Закруглені кути.
    - Ширина на ваш вибір.
    - Вирівнювання тексту на ваш вибір.
6. `.post-img` елемент має бути стилізований так, щоб зображення заповнювало всю ширину картки і мало нижній кордон.
7. `.post-content` елемент має бути стилізований із відступами всередині картки.
8. `.post-title` і `.post-excerpt` елементи мають мати колір тексту, відмінний від за замовчуванням, і поля з усіх боків.
9. `.read-more` елемент має бути стилізований як кнопка і мати:
    - Колір тексту, відмінний від за замовчуванням.
    - Фоновий колір.
    - Поля з усіх боків.
    - Відступи з усіх боків.
    - Закруглені кути.
    - Властивість `display` зі значенням `inline-block`.
    - Ефект наведення, який змінює фоновий колір.

**Примітка:** Обов’язково підключіть свій файл стилів у HTML і застосуйте CSS.

# --hints--

У вас має бути `div` із класом `blog-post-card`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

У вас має бути `img` елемент із `class` рівним `post-img`. Переконайтеся, що у зображення є атрибут `alt` із текстом і атрибут `src` зі значенням.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

У вас має бути `div` із класом `post-content`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

У вас має бути `h2` елемент із класом `post-title`. Переконайтеся, що він містить текст для заголовка вашого допису в блозі.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

У вас має бути `p` із класом `post-excerpt`. Переконайтеся, що він містить текст для короткого опису вашого допису.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

У вас має бути `a` елемент із `class` рівним `read-more`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

Ваш `.read-more` елемент має містити текст `Read More`.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

Ваш `.blog-post-card` елемент має мати властивість `border-radius` зі значенням (не має бути `0` або від’ємним значенням).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

Ваш `.blog-post-card` елемент має мати білий фон.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

Ваш `body` елемент має мати фоновий колір, відмінний від білого.

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

Ви маєте націлитися на `.blog-post-card` і встановити його властивість `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

Ви маєте націлитися на `.blog-post-card` і встановити його властивість `text-align`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

Ви маєте націлитися на `.post-content` і встановити його властивість `padding`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

Ваш `.read-more` елемент має мати ефект наведення.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

Ви маєте націлитися на `.read-more` і встановити його властивість `color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

Ви маєте націлитися на `.read-more` і встановити його властивість `background-color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

Ви маєте націлитися на `.read-more` і встановити його властивість `margin`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

Ви маєте націлитися на `.read-more` і встановити його властивість `display`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

Ви маєте націлитися на `.read-more` і встановити його властивість `border-radius`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

Ви маєте націлитися на `.read-more` і встановити його властивість `padding`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

Ваш `.post-img` елемент має заповнювати `width` картки.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

Ваш `.post-img` елемент має мати значення `border-bottom`.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

Ваші `.post-title` і `.post-excerpt` елементи мають мати поля і кольори тексту, відмінні від за замовчуванням.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
