---
id: 587d7b8f367417b2b2512b60
title: إعادة هيكلة المتغيرات العامة خارج الدوال
challengeType: 1
forumTopicId: 301235
dashedName: refactor-global-variables-out-of-functions
---

# --description--

حتى الآن، رأينا مبدأين مميزين للبرمجة الوظيفية:

1) لا تغير متغيرًا أو كائنًا - أنشئ متغيرات وكائنات جديدة وأعدها إذا لزم الأمر من دالة. تلميح: استخدام شيء مثل `const newArr = arrVar`، حيث `arrVar` مصفوفة، سيُنشئ فقط مرجعًا للمتغير الموجود وليس نسخة. لذا تغيير قيمة في `newArr` سيغير القيمة في `arrVar`.

2) أعلن معلمات الدالة - أي عملية حسابية داخل دالة تعتمد فقط على المعلمات الممررة إليها، وليس على أي كائن أو متغير عام.

إضافة واحد إلى رقم ليست مثيرة جدًا، لكن يمكننا تطبيق هذه المبادئ عند العمل مع المصفوفات أو الكائنات الأكثر تعقيدًا.

# --instructions--

أعد كتابة الشفرة بحيث لا تتغير المصفوفة العامة `bookList` داخل أي من الدالتين. يجب أن تضيف دالة `add` اسم الكتاب المعطى `bookName` إلى نهاية المصفوفة الممررة إليها وتُعيد مصفوفة جديدة (قائمة). يجب أن تزيل دالة `remove` اسم الكتاب المعطى `bookName` من المصفوفة الممررة إليها.

**ملاحظة:** يجب أن تُعيد كلتا الدالتين مصفوفة، ويجب إضافة أي معلمات جديدة قبل معلمة `bookName`.

# --hints--

يجب ألا تتغير `bookList` وتظل تساوي `["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"]`.

```js
add(bookList, "Test");
remove(bookList, "The Hound of the Baskervilles");
assert(
  JSON.stringify(bookList) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'On The Electrodynamics of Moving Bodies',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae'
    ])
);
```

`add(bookList, "A Brief History of Time")` يجب أن تُعيد `["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae", "A Brief History of Time"]`.

```js
assert(
  JSON.stringify(add(bookList, "A Brief History of Time")) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'On The Electrodynamics of Moving Bodies',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae',
      'A Brief History of Time'
    ])
);
```

`remove(bookList, "On The Electrodynamics of Moving Bodies")` يجب أن تُعيد `["The Hound of the Baskervilles", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"]`.

```js
assert(
  JSON.stringify(remove(bookList, 'On The Electrodynamics of Moving Bodies')) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae'
    ])
);
```

`remove(add(bookList, "A Brief History of Time"), "On The Electrodynamics of Moving Bodies");` يجب أن تساوي `["The Hound of the Baskervilles", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae", "A Brief History of Time"]`.

```js
assert(
  JSON.stringify(remove(add(bookList, 'A Brief History of Time'), 'On The Electrodynamics of Moving Bodies')) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae',
      'A Brief History of Time'
    ])
);
```

# --seed--

## --seed-contents--

```js
// The global variable
const bookList = ["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"];

// Change code below this line
function add(bookName) {

  bookList.push(bookName);
  return bookList;
  
  // Change code above this line
}

// Change code below this line
function remove(bookName) {
  const book_index = bookList.indexOf(bookName);
  if (book_index >= 0) {

    bookList.splice(book_index, 1);
    return bookList;

    // Change code above this line
    }
}
```

# --solutions--

```js
// The global variable
const bookList = ["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"];

function add(bookList, bookName) {
  return [...bookList, bookName];
}

function remove(bookList, bookName) {
  const bookListCopy = [...bookList];
  const bookNameIndex = bookList.indexOf(bookName);
  if (bookNameIndex >= 0) {
    bookListCopy.splice(bookNameIndex, 1);
  }
  return bookListCopy;
}
```
