---
id: 587d7b8e367417b2b2512b5e
title: تجنب التحوُّلات والآثار الجانبية باستخدام البرمجة الوظيفية
challengeType: 1
forumTopicId: 301228
dashedName: avoid-mutations-and-side-effects-using-functional-programming
---

# --description--

إذا لم تكتشف ذلك بعد، كانت المشكلة في التحدي السابق مع استدعاء `splice` في دالة `tabClose()`. للأسف، `splice` يغيّر المصفوفة الأصلية التي يُستدعى عليها، لذا استُخدم الاستدعاء الثاني لها على مصفوفة معدّلة، وأعطى نتائج غير متوقعة.

هذا مثال صغير على نمط أكبر بكثير - تستدعي دالة على متغير أو مصفوفة أو كائن، وتغيّر الدالة المتغير أو شيئًا في الكائن.

أحد المبادئ الأساسية في البرمجة الوظيفية هو عدم تغيير الأشياء. التغييرات تؤدي إلى أخطاء برمجية. من الأسهل منع الأخطاء عندما تعرف أن دوالك لا تغيّر أي شيء، بما في ذلك معلمات الدالة أو أي متغير عام.

المثال السابق لم يحتوي على عمليات معقدة، لكن طريقة `splice` غيّرت المصفوفة الأصلية، وأدت إلى خطأ برمجي.

تذكّر أنه في البرمجة الوظيفية، تغيير أو تعديل الأشياء يُسمى <dfn>تحوُّل</dfn>، والنتيجة تُسمى <dfn>أثرًا جانبيًا</dfn>. الدالة، في أفضل الأحوال، يجب أن تكون <dfn>دالة نقية</dfn>، أي لا تسبب أي آثار جانبية.

لنحاول إتقان هذا المبدأ وألا نغيّر أي متغير أو كائن في شفرتنا.

# --instructions--

املأ شفرة دالة `incrementer` بحيث تُرجع قيمة المتغير العام `fixedValue` بعد زيادتها بواحد.

# --hints--

يجب ألا تغيّر دالتك `incrementer` قيمة `fixedValue` (التي تساوي `4`).

```js
incrementer();
assert(fixedValue === 4);
```

يجب أن تُرجع دالة `incrementer` قيمة أكبر بواحد من قيمة `fixedValue`.

```js
const __newValue = incrementer();
assert(__newValue === 5);
```

يجب أن تُرجع دالة `incrementer` قيمة تعتمد على قيمة المتغير العام `fixedValue`.

```js
(function () {
  fixedValue = 10;
  const newValue = incrementer();
  assert(fixedValue === 10 && newValue === 11);
  fixedValue = 4;
})();
```

# --seed--

## --seed-contents--

```js
// The global variable
let fixedValue = 4;

function incrementer() {
  // Only change code below this line


  // Only change code above this line
}
```

# --solutions--

```js
let fixedValue = 4

function incrementer() {
  return fixedValue + 1
}
```
