---
id: 587d7da9367417b2b2512b6a
title: إرجاع مصفوفة مرتبة دون تغيير المصفوفة الأصلية
challengeType: 1
forumTopicId: 301237
dashedName: return-a-sorted-array-without-changing-the-original-array
---

# --description--

أحد الآثار الجانبية لطريقة `sort` هو أنها تغير ترتيب العناصر في المصفوفة الأصلية. بعبارة أخرى، تُغير المصفوفة في مكانها. إحدى الطرق لتجنب ذلك هي أولاً دمج مصفوفة فارغة مع المصفوفة التي تُرتب (تذكر أن `slice` و`concat` تُرجع مصفوفة جديدة)، ثم تشغيل طريقة `sort`.

# --instructions--

استخدم طريقة `sort` في الدالة `nonMutatingSort` لترتيب عناصر المصفوفة تصاعديًا. يجب أن تُرجع الدالة مصفوفة جديدة، ولا تُغير المتغير `globalArray`.

# --hints--

يجب أن يستخدم شفرتك طريقة `sort`.

```js
assert(nonMutatingSort.toString().match(/\.sort/g));
```

يجب ألا يتغير المتغير `globalArray`.

```js
assert(JSON.stringify(globalArray) === JSON.stringify([5, 6, 3, 2, 9]));
```

يجب أن تُرجع `nonMutatingSort(globalArray)` القيمة `[2, 3, 5, 6, 9]`.

```js
assert(
  JSON.stringify(nonMutatingSort(globalArray)) ===
    JSON.stringify([2, 3, 5, 6, 9])
);
```

يجب ألا تكون `nonMutatingSort(globalArray)` مشفرة بشكل ثابت.

```js
assert(!nonMutatingSort.toString().match(/\[.*?[23569].*?\]/gs));
```

يجب أن تُرجع الدالة مصفوفة جديدة، لا المصفوفة التي تم تمريرها إليها.

```js
assert(nonMutatingSort(globalArray) !== globalArray);
```

يجب أن تُرجع `nonMutatingSort([1, 30, 4, 21, 100000])` القيمة `[1, 4, 21, 30, 100000]`.

```js
assert(JSON.stringify(nonMutatingSort([1, 30, 4, 21, 100000])) ===
    JSON.stringify([1, 4, 21, 30, 100000]))
```

يجب أن تُرجع `nonMutatingSort([140000, 104, 99])` القيمة `[99, 104, 140000]`.

```js
assert(JSON.stringify(nonMutatingSort([140000, 104, 99])) ===
    JSON.stringify([99, 104, 140000]))
```

# --seed--

## --seed-contents--

```js
const globalArray = [5, 6, 3, 2, 9];

function nonMutatingSort(arr) {
  // Only change code below this line


  // Only change code above this line
}

nonMutatingSort(globalArray);
```

# --solutions--

```js
const globalArray = [5, 6, 3, 2, 9];
function nonMutatingSort(arr) {
  return [].concat(arr).sort((a,b) => a-b);
}
```
