---
id: 587d7dab367417b2b2512b70
title: مقدمة إلى التطبيق الجزئي والتطبيق الجزئي للدوال
challengeType: 1
forumTopicId: 301232
dashedName: introduction-to-currying-and-partial-application
---

# --description--

الـ <dfn>arity</dfn> لدالة هو عدد المعلمات التي تحتاجها. <dfn>التطبيق الجزئي</dfn> لدالة يعني تحويل دالة ذات arity N إلى N دوال ذات arity 1.

بمعنى آخر، يعيد هيكلة الدالة لتأخذ معلمة واحدة، ثم تُرجع دالة أخرى تأخذ المعلمة التالية، وهكذا.

إليك مثالًا:

```js
function unCurried(x, y) {
  return x + y;
}

function curried(x) {
  return function(y) {
    return x + y;
  }
}

const curried = x => y => x + y

curried(1)(2)
```

`curried(1)(2)` سترجع `3`.

هذا مفيد في برنامجك إذا لم تستطع توفير كل المعلمات للدالة دفعة واحدة. يمكنك حفظ كل استدعاء دالة في متغير، الذي سيحتفظ بمرجع الدالة المُرجعة التي تأخذ المعلمة التالية عندما تتوفر. إليك مثالًا باستخدام الدالة التطبيقية في المثال أعلاه:

```js
const funcForY = curried(1);
console.log(funcForY(2)); // 3
```

بنفس الطريقة، يمكن وصف <dfn>التطبيق الجزئي</dfn> بأنه تطبيق بعض المعلمات على دالة في كل مرة، ثم إرجاع دالة أخرى تُطبق على المزيد من المعلمات. إليك مثالًا:

```js
function impartial(x, y, z) {
  return x + y + z;
}

const partialFn = impartial.bind(this, 1, 2);
partialFn(10); // 13
```

# --instructions--

املأ جسم الدالة `add` بحيث تستخدم التطبيق الجزئي لجمع المعلمات `x` و `y` و `z`.

# --hints--

`add(10)(20)(30)` يجب أن تُرجع `60`.

```js
assert(add(10)(20)(30) === 60);
```

`add(1)(2)(3)` يجب أن تُرجع `6`.

```js
assert(add(1)(2)(3) === 6);
```

`add(11)(22)(33)` يجب أن تُرجع `66`.

```js
assert(add(11)(22)(33) === 66);
```

يجب أن يتضمن شفرتك بيانًا نهائيًا يُرجع `x + y + z`.

```js
assert(__helpers.removeJSComments(code).match(/[xyz]\s*?\+\s*?[xyz]\s*?\+\s*?[xyz]/g));
```

# --seed--

## --seed-contents--

```js
function add(x) {
  // Only change code below this line


  // Only change code above this line
}

add(10)(20)(30);
```

# --solutions--

```js
const add = x => y => z => x + y + z
```
