---
id: 587d7b88367417b2b2512b45
title: "استخدام الدوال ذات الرتبة العليا map و filter و reduce لحل تحدٍّ معقد"
challengeType: 1
forumTopicId: 301311
dashedName: use-higher-order-functions-map-filter-or-reduce-to-solve-a-complex-problem
---

# --description--

بعد أن عملت على عدة تحديات باستخدام الدوال ذات الرتبة العليا مثل `map()` و `filter()` و `reduce()`، يمكنك الآن تطبيقها لحل تحدٍّ أكثر تعقيدًا.

# --instructions--

أكمل شفرة الدالة `squareList` باستخدام أي تركيبة من `map()` و `filter()` و `reduce()`. يجب أن تُرجع الدالة مصفوفة جديدة تحتوي على مربعات *فقط* الأعداد الصحيحة الموجبة (الأعداد العشرية ليست أعدادًا صحيحة) عند تمرير مصفوفة من الأعداد الحقيقية إليها. مثال على مصفوفة أعداد حقيقية هو `[-3, 4.8, 5, 3, -3.2]`.

**ملاحظة:** يجب ألا تستخدم دالتك أي نوع من حلقات `for` أو `while` أو دالة `forEach()`.

# --hints--

يجب أن تكون `squareList` عبارة عن `function`.

```js
assert.typeOf(squareList, 'function'),
  '<code>squareList</code> should be a <code>function</code>';
```

يجب عدم استخدام `for` أو `while` أو `forEach`.

```js
assert(!__helpers.removeJSComments(code).match(/for|while|forEach/g));
```

يجب استخدام `map` أو `filter` أو `reduce`.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/\.(map|filter|reduce)\(/g)
);
```

يجب أن تُرجع الدالة `array`.

```js
assert(Array.isArray(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])));
```

يجب أن تُرجع `squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])` القيمة `[16, 1764, 36]`.

```js
assert.deepStrictEqual(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2]), [
  16,
  1764,
  36
]);
```

يجب أن تُرجع `squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3])` القيمة `[9, 100, 49]`.

```js
assert.deepStrictEqual(squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3]), [
  9,
  100,
  49
]);
```

# --seed--

## --seed-contents--

```js
const squareList = arr => {
  // Only change code below this line
  return arr;
  // Only change code above this line
};

const squaredIntegers = squareList([-3, 4.8, 5, 3, -3.2]);
console.log(squaredIntegers);
```

# --solutions--

```js
const squareList = arr => {
  const positiveIntegers = arr.filter(num => {
    return num >= 0 && Number.isInteger(num);
  });
  const squaredIntegers = positiveIntegers.map(num => {
    return num ** 2;
  });
  return squaredIntegers;
};
```
