---
id: 587d7b8f367417b2b2512b64
title: تنفيذ طريقة filter على نموذج أولي
challengeType: 1
forumTopicId: 301231
dashedName: implement-the-filter-method-on-a-prototype
---

# --description--

قد تتعلم الكثير عن طريقة `filter` إذا نفذت نسختك الخاصة منها. يُنصح باستخدام حلقة `for` أو `Array.prototype.forEach()`.

# --instructions--

اكتب `Array.prototype.myFilter()` الخاصة بك، والتي يجب أن تتصرف تمامًا مثل `Array.prototype.filter()`. لا يجب عليك استخدام طريقة `filter` المدمجة. يمكن الوصول إلى مثيل `Array` في طريقة `myFilter` باستخدام `this`.

# --hints--

`[23, 65, 98, 5, 13].myFilter(item => item % 2)` يجب أن تساوي `[23, 65, 5, 13]`.

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item % 2;
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

`["naomi", "quincy", "camperbot"].myFilter(element => element === "naomi")` يجب أن تُرجع `["naomi"]`.

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element === "naomi";
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

`[1, 1, 2, 5, 2].myFilter((element, index, array) => array.indexOf(element) === index)` يجب أن تُرجع `[1, 2, 5]`.

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array.indexOf(element) === index;
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

يجب ألا يستخدم شفرتك طريقة `filter`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?filter/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myFilter = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myFilter = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    if (callback(this[i], i, this)) newArray.push(this[i]);
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const odd_s = s.myFilter(item => item % 2 === 1);
```
