---
id: 587d7b8d367417b2b2512b5b
title: تعرّف على البرمجة الوظيفية
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

البرمجة الوظيفية هي أسلوب برمجة تعتمد فيه الحلول على دوال بسيطة ومعزولة، دون أي تأثيرات جانبية خارج نطاق الدالة: `INPUT -> PROCESS -> OUTPUT`

البرمجة الوظيفية تعتمد على:

1) دوال معزولة - لا تعتمد على حالة البرنامج، والتي تشمل المتغيرات العامة التي قد تتغير

2) دوال نقية - نفس المُدخَل يعطي دائمًا نفس الناتج

3) دوال ذات تأثيرات جانبية محدودة - أي تغييرات أو تحوُّلات في حالة البرنامج خارج الدالة تُراقب بعناية

# --instructions--

أعضاء freeCodeCamp يحبون الشاي.

في محرر الشفرة، دالتي `prepareTea` و `getTea` معرفتان مسبقًا. استدعِ الدالة `getTea` لتحصل على 40 كوب شاي للفريق، وخزنها في المتغير `tea4TeamFCC`.

# --hints--

يجب أن يحتوي المتغير `tea4TeamFCC` على 40 كوب شاي للفريق.

```js
assert(tea4TeamFCC.length === 40);
```

يجب أن يحتوي المتغير `tea4TeamFCC` على أكواب شاي أخضر.

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40); 
```
