---
id: 587d7da9367417b2b2512b67
title: إضافة عناصر إلى نهاية مصفوفة باستخدام concat بدلاً من push
challengeType: 1
forumTopicId: 301226
dashedName: add-elements-to-the-end-of-an-array-using-concat-instead-of-push
---

# --description--

البرمجة الوظيفية تعتمد على إنشاء واستخدام دوال لا تُغيّر البيانات الأصلية.

التحدٍّ الأخير قدّم طريقة `concat` كوسيلة لدمج المصفوفات في مصفوفة جديدة دون تغيير المصفوفات الأصلية. قارن بين `concat` وطريقة `push`. تضيف `push` عناصر إلى نهاية نفس المصفوفة التي تُستدعى عليها، مما يغيّر تلك المصفوفة. إليك مثالاً:

```js
const arr = [1, 2, 3];
arr.push(4, 5, 6);
```

ستكون قيمة `arr` معدّلة إلى `[1, 2, 3, 4, 5, 6]`، وهذا ليس أسلوب البرمجة الوظيفية.

تقدم `concat` طريقة لدمج عناصر جديدة إلى نهاية مصفوفة دون أي تأثيرات جانبية تغيّر البيانات.

# --instructions--

غيّر دالة `nonMutatingPush` لتستخدم `concat` لدمج `newItem` في نهاية `original` دون تغيير مصفوفات `original` أو `newItem`. يجب أن تُرجع الدالة مصفوفة.

# --hints--

يجب أن يستخدم شفرتك طريقة `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

يجب ألا تستخدم شفرتك طريقة `push`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?push/g));
```

يجب ألا تتغير المصفوفة `first`.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

يجب ألا تتغير المصفوفة `second`.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

يجب أن تُرجع `nonMutatingPush([1, 2, 3], [4, 5])` القيمة `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingPush([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingPush(original, newItem) {
  // Only change code below this line
  return original.push(newItem);

  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingPush(first, second);
```

# --solutions--

```js
function nonMutatingPush(original, newItem) {
  return original.concat(newItem);
}
const first = [1, 2, 3];
const second = [4, 5];
```
