---
id: 587d7daa367417b2b2512b6c
title: دمج مصفوفة في سلسلة نصية باستخدام طريقة join
challengeType: 1
forumTopicId: 18221
dashedName: combine-an-array-into-a-string-using-the-join-method
---

# --description--

تُستخدم طريقة `join` لدمج عناصر المصفوفة معًا لإنشاء سلسلة نصية. تأخذ معلمة للفاصل الذي يُستخدم لفصل عناصر المصفوفة في السلسلة النصية.

إليك مثالًا:

```js
const arr = ["Hello", "World"];
const str = arr.join(" ");
```

سيكون لـ `str` قيمة السلسلة النصية `Hello World`.
# --instructions--

استخدم طريقة `join` (من بين طرق أخرى) داخل الدالة `sentensify` لصنع جملة من الكلمات في السلسلة النصية `str`. يجب أن تُرجع الدالة سلسلة نصية. على سبيل المثال، تتحول `I-like-Star-Wars` إلى `I like Star Wars`. لا تستخدم طريقة `replace` في هذا التحدٍّ.

# --hints--

يجب أن يستخدم شفرتك طريقة `join`.

```js
assert(__helpers.removeJSComments(code).match(/\.join/g));
```

يجب ألا تستخدم شفرتك طريقة `replace`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?replace/g));
```

يجب أن تُرجع `sentensify("May-the-force-be-with-you")` سلسلة نصية.

```js
assert(typeof sentensify('May-the-force-be-with-you') === 'string');
```

يجب أن تُرجع `sentensify("May-the-force-be-with-you")` السلسلة النصية `May the force be with you`.

```js
assert(sentensify('May-the-force-be-with-you') === 'May the force be with you');
```

يجب أن تُرجع `sentensify("The.force.is.strong.with.this.one")` السلسلة النصية `The force is strong with this one`.

```js
assert(
  sentensify('The.force.is.strong.with.this.one') ===
    'The force is strong with this one'
);
```

يجب أن تُرجع `sentensify("There,has,been,an,awakening")` السلسلة النصية `There has been an awakening`.

```js
assert(
  sentensify('There,has,been,an,awakening') === 'There has been an awakening'
);
```

# --seed--

## --seed-contents--

```js
function sentensify(str) {
  // Only change code below this line


  // Only change code above this line
}

sentensify("May-the-force-be-with-you");
```

# --solutions--

```js
function sentensify(str) {
  return str.split(/\W/).join(' ');
}
```
