---
id: 587d7dab367417b2b2512b6e
title: استخدم طريقة every للتحقق من أن كل عنصر في مصفوفة يفي بمعيار
challengeType: 1
forumTopicId: 301312
dashedName: use-the-every-method-to-check-that-every-element-in-an-array-meets-a-criteria
---

# --description--

تعمل طريقة `every` مع المصفوفات للتحقق مما إذا كان *كل* عنصر يمر باختبار معين. تُرجع قيمة منطقية - `true` إذا استوفت كل القيم المعيار، و`false` إذا لم تستوفِ.

على سبيل المثال، الشفرة التالية تتحقق مما إذا كان كل عنصر في مصفوفة `numbers` أقل من 10:

```js
const numbers = [1, 5, 8, 0, 10, 11];

numbers.every(function(currentValue) {
  return currentValue < 10;
});
```

ستُرجع طريقة `every` هنا القيمة `false`.

# --instructions--

استخدم طريقة `every` داخل الدالة `checkPositive` للتحقق مما إذا كان كل عنصر في `arr` موجبًا. يجب أن تُرجع الدالة قيمة منطقية.

# --hints--

يجب أن تستخدم شفرتك طريقة `every`.

```js
assert(__helpers.removeJSComments(code).match(/\.every/g));
```

`checkPositive([1, 2, 3, -4, 5])` يجب أن تُرجع `false`.

```js
assert.isFalse(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` يجب أن تُرجع `true`.

```js
assert.isTrue(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([1, -2, 3, -4, 5])` يجب أن تُرجع `false`.

```js
assert.isFalse(checkPositive([1, -2, 3, -4, 5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.every(num => num > 0);
}
```
