---
id: 587d7da9367417b2b2512b69
title: فرز مصفوفة أبجديًا باستخدام طريقة sort
challengeType: 1
forumTopicId: 18303
dashedName: sort-an-array-alphabetically-using-the-sort-method
---

# --description--

طريقة `sort` تفرز عناصر المصفوفة وفقًا لدالة الاستدعاء الراجع.

على سبيل المثال:

```js
function ascendingOrder(arr) {
  return arr.sort(function(a, b) {
    return a - b;
  });
}

ascendingOrder([1, 5, 2, 3, 4]);
```

سيُرجع هذا القيمة `[1, 2, 3, 4, 5]`.

```js
function reverseAlpha(arr) {
  return arr.sort(function(a, b) {
    return a === b ? 0 : a < b ? 1 : -1;
  });
}

reverseAlpha(['l', 'h', 'z', 'b', 's']);
```

سيُرجع هذا القيمة `['z', 's', 'l', 'h', 'b']`.

طريقة الفرز الافتراضية في JavaScript تعتمد على قيمة نقطة يونيكود للسلسلة النصية، وهذا قد يُنتج نتائج غير متوقعة. لذلك يُفضل توفير دالة استدعاء راجع لتحديد كيفية فرز عناصر المصفوفة. عندما تُزوَّد دالة استدعاء راجع، عادةً ما تُسمى `compareFunction`، تُفرز عناصر المصفوفة بناءً على قيمة الإرجاع من `compareFunction`: إذا أعادت `compareFunction(a,b)` قيمة أقل من 0 لعنصرين `a` و `b`، فإن `a` يأتي قبل `b`. إذا أعادت `compareFunction(a,b)` قيمة أكبر من 0 لعنصرين `a` و `b`، فإن `b` يأتي قبل `a`. إذا أعادت `compareFunction(a,b)` قيمة تساوي 0 لعنصرين `a` و `b`، فإن `a` و `b` يبقيان كما هما دون تغيير.

# --instructions--

استخدم طريقة `sort` في الدالة `alphabeticalOrder` لفرز عناصر `arr` ترتيبًا أبجديًا. يجب أن تُرجع الدالة المصفوفة المُفرزة.

# --hints--

يجب أن يستخدم شفرتك طريقة `sort`.

```js
assert(__helpers.removeJSComments(code).match(/\.sort/g));
```

`alphabeticalOrder(["a", "d", "c", "a", "z", "g"])` يجب أن تُرجع `["a", "a", "c", "d", "g", "z"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'd', 'c', 'a', 'z', 'g'])) ===
    JSON.stringify(['a', 'a', 'c', 'd', 'g', 'z'])
);
```

`alphabeticalOrder(["x", "h", "a", "m", "n", "m"])` يجب أن تُرجع `["a", "h", "m", "m", "n", "x"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['x', 'h', 'a', 'm', 'n', 'm'])) ===
    JSON.stringify(['a', 'h', 'm', 'm', 'n', 'x'])
);
```

`alphabeticalOrder(["a", "a", "a", "a", "x", "t"])` يجب أن تُرجع `["a", "a", "a", "a", "t", "x"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'a', 'a', 'a', 'x', 't'])) ===
    JSON.stringify(['a', 'a', 'a', 'a', 't', 'x'])
);
```

# --seed--

## --seed-contents--

```js
function alphabeticalOrder(arr) {
  // Only change code below this line

  return arr
  // Only change code above this line
}

alphabeticalOrder(["a", "d", "c", "a", "z", "g"]);
```

# --solutions--

```js
function alphabeticalOrder(arr) {
  return arr.sort();
}
```
