---
id: 587d7b8e367417b2b2512b5c
title: فهم مصطلحات البرمجة الوظيفية
challengeType: 1
forumTopicId: 301240
dashedName: understand-functional-programming-terminology
---

# --description--

مر فريق FCC بتقلب مزاجي ويريد الآن نوعين من الشاي: الشاي الأخضر والشاي الأسود. حقيقة عامة: تقلبات مزاج العميل شائعة جدًا.

مع هذه المعلومات، سنحتاج إلى مراجعة دالة `getTea` من التحدي السابق للتعامل مع طلبات الشاي المختلفة. يمكننا تعديل `getTea` لتقبل دالة كمعلمة لتتمكن من تغيير نوع الشاي الذي تُحضره. هذا يجعل `getTea` أكثر مرونة، ويمنح المبرمج تحكمًا أكبر عندما تتغير طلبات العميل.

لكن أولًا، دعنا نغطي بعض المصطلحات الوظيفية:

<dfn>استدعاءات راجعة</dfn> هي الدوال التي تُمرر أو تُسلَّم إلى دالة أخرى لتحديد استدعاء تلك الدالة. ربما رأيتها تُمرر إلى طرق أخرى، مثلًا في `filter`، حيث تخبر دالة الاستدعاء الراجع جافاسكريبت بمعايير كيفية تصفية المصفوفة.

الدوال التي يمكن تعيينها إلى متغير، أو تمريرها إلى دالة أخرى، أو إرجاعها من دالة أخرى مثل أي قيمة عادية تُسمى دوال <dfn>الدرجة الأولى</dfn>. في جافاسكريبت، كل الدوال هي دوال درجة أولى.

الدوال التي تأخذ دالة كمعلمة، أو تُرجع دالة كقيمة إرجاع، تُسمى دوال <dfn>الترتيب الأعلى</dfn>.

عندما تُمرر دوال إلى دالة أخرى أو تُرجع منها، يمكن تسمية تلك الدوال التي تم تمريرها أو إرجاعها <dfn>لامبدا</dfn>.

# --instructions--

حضّر 27 كوبًا من الشاي الأخضر و13 كوبًا من الشاي الأسود وخزنها في المتغيرين `tea4GreenTeamFCC` و`tea4BlackTeamFCC` على التوالي. لاحظ أن دالة `getTea` تم تعديلها لتأخذ دالة كمعلمة أولى الآن.

ملاحظة: البيانات (عدد أكواب الشاي) تُزوَّد كآخر معلمة. سنناقش هذا أكثر في الدروس القادمة.

# --hints--

يجب أن يحتوي المتغير `tea4GreenTeamFCC` على 27 كوبًا من الشاي الأخضر للفريق.

```js
assert(tea4GreenTeamFCC.length === 27);
```

يجب أن يحتوي المتغير `tea4GreenTeamFCC` على أكواب من الشاي الأخضر.

```js
assert(tea4GreenTeamFCC[0] === 'greenTea');
```

يجب أن يحتوي المتغير `tea4BlackTeamFCC` على 13 كوبًا من الشاي الأسود.

```js
assert(tea4BlackTeamFCC.length === 13);
```

يجب أن يحتوي المتغير `tea4BlackTeamFCC` على أكواب من الشاي الأسود.

```js
assert(tea4BlackTeamFCC[0] === 'blackTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareGreenTea = () => 'greenTea';

// Function that returns a string representing a cup of black tea
const prepareBlackTea = () => 'blackTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (prepareTea, numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4GreenTeamFCC = null;
const tea4BlackTeamFCC = null;
// Only change code above this line

console.log(
  tea4GreenTeamFCC,
  tea4BlackTeamFCC
);
```

# --solutions--

```js
const prepareGreenTea = () => 'greenTea';
const prepareBlackTea = () => 'blackTea';

const getTea = (prepareTea, numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

const tea4BlackTeamFCC = getTea(prepareBlackTea, 13);
const tea4GreenTeamFCC = getTea(prepareGreenTea, 27);
```
