---
id: 587d7dab367417b2b2512b6f
title: استخدم طريقة some للتحقق مما إذا كان أي عنصر في مصفوفة يفي بمعيار معين
challengeType: 1
forumTopicId: 301314
dashedName: use-the-some-method-to-check-that-any-elements-in-an-array-meet-a-criteria
---

# --description--

تعمل طريقة `some` مع المصفوفات للتحقق مما إذا كان *أي* عنصر يجتاز اختبارًا معينًا. تُرجع قيمة منطقية — `true` إذا استوفى أي من القيم المعيار، و`false` إذا لم يستوفِ.

على سبيل المثال، الشفرة التالية تتحقق مما إذا كان أي عنصر في مصفوفة `numbers` أقل من 10:

```js
const numbers = [10, 50, 8, 220, 110, 11];

numbers.some(function(currentValue) {
  return currentValue < 10;
});
```

ستُرجع طريقة `some` القيمة `true`.

# --instructions--

استخدم طريقة `some` داخل الدالة `checkPositive` للتحقق مما إذا كان أي عنصر في `arr` موجبًا. يجب أن تُرجع الدالة قيمة منطقية.

# --hints--

يجب أن تستخدم شفرتك طريقة `some`.

```js
assert(__helpers.removeJSComments(code).match(/\.some/g));
```

`checkPositive([1, 2, 3, -4, 5])` يجب أن تُرجع `true`.

```js
assert(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` يجب أن تُرجع `true`.

```js
assert(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([-1, -2, -3, -4, -5])` يجب أن تُرجع `false`.

```js
assert(!checkPositive([-1, -2, -3, -4, -5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.some(elem => elem > 0);
}
```
