---
id: 587d7b8f367417b2b2512b62
title: تنفيذ map على نموذج أولي
challengeType: 1
forumTopicId: 301230
dashedName: implement-map-on-a-prototype
---

# --description--

كما رأيت عند تطبيق `Array.prototype.map()`، أو ببساطة `map()` سابقًا، تُرجع طريقة `map` مصفوفة بنفس طول المصفوفة التي استُدعي عليها. كما أنها لا تُغير المصفوفة الأصلية، طالما أن دالة الاستدعاء الراجع الخاصة بها لا تفعل ذلك.

بعبارة أخرى، `map` دالة نقيّة، ومخرجاتها تعتمد فقط على مُدخلاتها. بالإضافة إلى ذلك، تأخذ دالة أخرى كمعلمتها.

يمكنك أن تتعلم الكثير عن طريقة `map` إذا نفذت نسختك الخاصة منها. يُنصح باستخدام حلقة `for` أو `Array.prototype.forEach()`.

# --instructions--

اكتب دالتك الخاصة `Array.prototype.myMap()`، التي يجب أن تتصرف تمامًا مثل `Array.prototype.map()`. لا يجب عليك استخدام طريقة `map` المدمجة. يمكن الوصول إلى مثيل `Array` في طريقة `myMap` باستخدام `this`.

# --hints--

`[23, 65, 98, 5, 13].myMap(item => item * 2)` يجب أن تساوي `[46, 130, 196, 10, 26]`.

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item * 2;
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`["naomi", "quincy", "camperbot"].myMap(element => element.toUpperCase())` يجب أن تُرجع `["NAOMI", "QUINCY", "CAMPERBOT"]`.

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element.toUpperCase();
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`[1, 1, 2, 5, 2].myMap((element, index, array) => array[index + 1] || array[0])` يجب أن تُرجع `[1, 2, 5, 2, 1]`.

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array[index + 1] || array[0];
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

يجب ألا يستخدم شفرتك طريقة `map`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?map/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    newArray.push(callback(this[i], i, this));
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const doubled_s = s.myMap(item => item * 2);
```
