---
id: 587d7b8e367417b2b2512b5d
title: فهم مخاطر استخدام الشفرة الأمرية
challengeType: 1
forumTopicId: 301241
dashedName: understand-the-hazards-of-using-imperative-code
---

# --description--

البرمجة الوظيفية عادة جيدة. تحافظ على شفرتك سهلة الإدارة، وتجنبك الأخطاء البرمجية الخفية. لكن قبل أن نصل إلى هناك، دعنا نلقي نظرة على أسلوب برمجة أمري لتوضيح الأماكن التي قد تواجه فيها مشاكل.

في اللغة الإنجليزية (وفي العديد من اللغات الأخرى)، يُستخدم زمن الأمر لإعطاء الأوامر. وبالمثل، الأسلوب الأمري في البرمجة هو الذي يعطي الكمبيوتر مجموعة من التعليمات لأداء مهمة.

غالبًا ما تغير التعليمات حالة البرنامج، مثل تحديث المتغيرات العامة. مثال كلاسيكي هو كتابة حلقة `for` تعطي تعليمات دقيقة للتكرار على مؤشرات مصفوفة.

على النقيض من ذلك، البرمجة الوظيفية هي شكل من أشكال البرمجة التصريحية. تخبر الكمبيوتر بما تريد إنجازه عن طريق استدعاء طريقة أو دالة.

تقدم JavaScript العديد من الطرق المعرفة مسبقًا التي تتعامل مع المهام الشائعة حتى لا تحتاج إلى كتابة كيفية تنفيذ الكمبيوتر لها. على سبيل المثال، بدلًا من استخدام حلقة `for` المذكورة أعلاه، يمكنك استدعاء طريقة `map` التي تتولى تفاصيل التكرار على المصفوفة. هذا يساعد على تجنب الأخطاء الدلالية، مثل "أخطاء الانحراف بمقدار واحد" التي تم تناولها في قسم تصحيح الأخطاء.

تخيل السيناريو التالي: تتصفح الويب في متصفحك، وتريد تتبع علامات التبويب التي فتحتها. دعنا نحاول نمذجة هذا باستخدام بعض الشفرة البسيطة الموجهة للكائنات.

كائن Window يتكون من علامات تبويب، وعادة ما يكون لديك أكثر من نافذة مفتوحة. عناوين كل موقع مفتوح في كل كائن Window محفوظة في مصفوفة. بعد العمل في المتصفح (فتح علامات تبويب جديدة، دمج النوافذ، وإغلاق علامات التبويب)، تريد طباعة علامات التبويب التي لا تزال مفتوحة. تُزال علامات التبويب المغلقة من المصفوفة وتُضاف علامات التبويب الجديدة (للتبسيط) إلى نهايتها.

يعرض محرر الشفرة تنفيذ هذه الوظيفة باستخدام دوال `tabOpen()`، و`tabClose()`، و`join()`. المصفوفة `tabs` هي جزء من كائن Window الذي يخزن أسماء الصفحات المفتوحة.

# --instructions--

افحص الشفرة في المحرر. إنها تستخدم طريقة لها تأثيرات جانبية في البرنامج، مما يسبب سلوكًا غير صحيح. القائمة النهائية لعلامات التبويب المفتوحة، المخزنة في `finalTabs.tabs`، يجب أن تكون `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']` لكن القائمة الناتجة عن الشفرة مختلفة قليلاً.

غيّر `Window.prototype.tabClose` بحيث تزيل علامة التبويب الصحيحة.

# --hints--

`finalTabs.tabs` يجب أن تكون `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']`

```js
assert.deepEqual(finalTabs.tabs, [
  'FB',
  'Gitter',
  'Reddit',
  'Twitter',
  'Medium',
  'new tab',
  'Netflix',
  'YouTube',
  'Vine',
  'GMail',
  'Work mail',
  'Docs',
  'freeCodeCamp',
  'new tab'
]);
```

# --seed--

## --seed-contents--

```js
// tabs is an array of titles of each site open within the window
const Window = function(tabs) {
  this.tabs = tabs; // We keep a record of the array inside the object
};

// When you join two windows into one window
Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

// When you open a new tab at the end
Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab'); // Let's open a new tab for now
  return this;
};

// When you close a tab
Window.prototype.tabClose = function(index) {

  // Only change code below this line

  const tabsBeforeIndex = this.tabs.splice(0, index); // Get the tabs before the tab
  const tabsAfterIndex = this.tabs.splice(index + 1); // Get the tabs after the tab

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex); // Join them together

  // Only change code above this line

  return this;
 };

// Let's create three browser windows
const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']); // Your mailbox, drive, and other work sites
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']); // Social sites
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']); // Entertainment sites

// Now perform the tab opening, closing, and other operations
const finalTabs = socialWindow
  .tabOpen() // Open a new tab for cat memes
  .join(videoWindow.tabClose(2)) // Close third tab in video window, and join
  .join(workWindow.tabClose(1).tabOpen());
console.log(finalTabs.tabs);
```

# --solutions--

```js
const Window = function(tabs) {
  this.tabs = tabs;
};

Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab');
  return this;
};

Window.prototype.tabClose = function(index) {
  const tabsBeforeIndex = this.tabs.slice(0, index);
  const tabsAfterIndex = this.tabs.slice(index + 1);

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex);
  return this;
 };

const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']);
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']);
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']);

const finalTabs = socialWindow
  .tabOpen()
  .join(videoWindow.tabClose(2))
  .join(workWindow.tabClose(1).tabOpen());
```
