---
id: 587d7daa367417b2b2512b6b
title: تقسيم سلسلة نصية إلى مصفوفة باستخدام طريقة split
challengeType: 1
forumTopicId: 18305
dashedName: split-a-string-into-an-array-using-the-split-method
---

# --description--

طريقة `split` تُقسّم سلسلة نصية إلى مصفوفة من السلاسل النصية. تأخذ معلمة للفاصل، التي يمكن أن تكون حرفًا يُستخدم لتقسيم السلسلة أو تعبيرًا نمطيًا. على سبيل المثال، إذا كان الفاصل هو مسافة، تحصل على مصفوفة من الكلمات، وإذا كان الفاصل سلسلة فارغة، تحصل على مصفوفة من كل حرف في السلسلة.

إليك مثالين يقسمان سلسلة نصية واحدة حسب المسافات، ثم أخرى حسب الأرقام باستخدام تعبير نمطي:

```js
const str = "Hello World";
const bySpace = str.split(" ");

const otherString = "How9are7you2today";
const byDigits = otherString.split(/\d/);
```

`bySpace` ستكون قيمتها `["Hello", "World"]` و`byDigits` ستكون قيمتها `["How", "are", "you", "today"]`.

بما أن السلاسل النصية غير قابلة للتغيير، تجعل طريقة `split` التعامل معها أسهل.

# --instructions--

استخدم طريقة `split` داخل الدالة `splitify` لتقسيم `str` إلى مصفوفة من الكلمات. يجب أن تُرجع الدالة المصفوفة. لاحظ أن الكلمات ليست دائمًا مفصولة بمسافات، ويجب ألا تحتوي المصفوفة على علامات ترقيم.

# --hints--

يجب أن يستخدم شفرتك طريقة `split`.

```js
assert(__helpers.removeJSComments(code).match(/\.split/g));
```

`splitify("Hello World,I-am code")` يجب أن تُرجع `["Hello", "World", "I", "am", "code"]`.

```js
assert(
  JSON.stringify(splitify('Hello World,I-am code')) ===
    JSON.stringify(['Hello', 'World', 'I', 'am', 'code'])
);
```

`splitify("Earth-is-our home")` يجب أن تُرجع `["Earth", "is", "our", "home"]`.

```js
assert(
  JSON.stringify(splitify('Earth-is-our home')) ===
    JSON.stringify(['Earth', 'is', 'our', 'home'])
);
```

`splitify("This.is.a-sentence")` يجب أن تُرجع `["This", "is", "a", "sentence"]`.

```js
assert(
  JSON.stringify(splitify('This.is.a-sentence')) ===
    JSON.stringify(['This', 'is', 'a', 'sentence'])
);
```

# --seed--

## --seed-contents--

```js
function splitify(str) {
  // Only change code below this line


  // Only change code above this line
}

splitify("Hello World,I-am code");
```

# --solutions--

```js
function splitify(str) {
  return str.split(/\W/);
}
```
