---
id: 587d7b8e367417b2b2512b5f
title: تمرير المعاملات لتجنب الاعتماد الخارجي في دالة
challengeType: 1
forumTopicId: 301234
dashedName: pass-arguments-to-avoid-external-dependence-in-a-function
---

# --description--

كان التحدٍّ الأخير خطوة أقرب إلى مبادئ البرمجة الوظيفية، لكن لا يزال هناك شيء مفقود.

لم نُغير قيمة المتغير العام، لكن الدالة `incrementer` لن تعمل بدون وجود المتغير العام `fixedValue`.

مبدأ آخر من مبادئ البرمجة الوظيفية هو إعلان تبعياتك دائمًا بشكل صريح. هذا يعني أنه إذا كانت الدالة تعتمد على وجود متغير أو كائن، فمرر ذلك المتغير أو الكائن مباشرة إلى الدالة كمعلمة.

هناك عدة نتائج جيدة من هذا المبدأ. تصبح الدالة أسهل في الاختبار، وتعرف بالضبط ما هو المُدخَل الذي تأخذه، ولن تعتمد على أي شيء آخر في برنامجك.

هذا يمنحك ثقة أكبر عند تعديل أو حذف أو إضافة شفرة جديدة. ستعرف ما يمكنك تغييره وما لا يمكنك، وسترى أين تكمن الفخاخ المحتملة.

أخيرًا، ستُنتج الدالة دائمًا نفس المخرجات لنفس مجموعة المُدخَلات، بغض النظر عن أي جزء من الشفرة ينفذها.

# --instructions--

لنحدّث دالة `incrementer` لإعلان تبعياتها بوضوح.

اكتب دالة `incrementer` بحيث تأخذ معلمة، ثم تُرجع نتيجة بعد زيادة القيمة بمقدار واحد.

# --hints--

يجب ألا تغيّر دالتك `incrementer` قيمة `fixedValue` (التي هي `4`).

```js
assert(fixedValue === 4);
```

يجب أن تأخذ دالة `incrementer` معلمة.

```js
assert(incrementer.length === 1);
```

يجب أن تُرجع دالة `incrementer` قيمة أكبر بمقدار واحد من قيمة `fixedValue`.

```js
const __newValue = incrementer(fixedValue);
assert(__newValue === 5);
```

# --seed--

## --seed-contents--

```js
// The global variable
let fixedValue = 4;

// Only change code below this line
function incrementer() {


  // Only change code above this line
}
```

# --solutions--

```js
let fixedValue = 4;

function incrementer(fixedValue) {
  return fixedValue + 1;
}
```
