---
id: 587d7da9367417b2b2512b66
title: دمج مصفوفتين باستخدام طريقة concat
challengeType: 1
forumTopicId: 301229
dashedName: combine-two-arrays-using-the-concat-method
---

# --description--

<dfn>الدمج</dfn> يعني ربط العناصر نهايةً بنهاية. تقدم JavaScript طريقة `concat` لكل من السلاسل النصية والمصفوفات التي تعمل بنفس الطريقة. بالنسبة للمصفوفات، تُستدعى الطريقة على مصفوفة واحدة، ثم تُمرر مصفوفة أخرى كمعلمة إلى `concat`، فتُضاف إلى نهاية المصفوفة الأولى. تُرجع مصفوفة جديدة ولا تُغير أيًا من المصفوفتين الأصليتين. إليك مثالًا:

```js
[1, 2, 3].concat([4, 5, 6]);
```

المصفوفة المرجعة ستكون `[1, 2, 3, 4, 5, 6]`.

# --instructions--

استخدم طريقة `concat` في الدالة `nonMutatingConcat` لدمج `attach` في نهاية `original`. يجب أن تُرجع الدالة المصفوفة المدمجة.

# --hints--

يجب أن يستخدم شفرتك طريقة `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

يجب ألا تتغير المصفوفة `first`.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

يجب ألا تتغير المصفوفة `second`.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingConcat([1, 2, 3], [4, 5])` يجب أن تُرجع `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingConcat([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingConcat(original, attach) {
  // Only change code below this line


  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingConcat(first, second);
```

# --solutions--

```js
function nonMutatingConcat(original, attach) {
  return original.concat(attach);
}
const first = [1, 2, 3];
const second = [4, 5];
```
