---
id: 587d7b90367417b2b2512b65
title: إرجاع جزء من مصفوفة باستخدام طريقة slice
challengeType: 1
forumTopicId: 301239
dashedName: return-part-of-an-array-using-the-slice-method
---

# --description--

ترجع طريقة `slice` نسخة من عناصر معينة من المصفوفة. يمكنها أخذ معلمتين، الأولى تحدد فهرس بداية الجزء المقطوع، والثانية تحدد فهرس نهاية الجزء المقطوع (والتي لا تُشمل في النتيجة). إذا لم تُعطَ المعلمتان، فالافتراضي هو البدء من بداية المصفوفة حتى نهايتها، وهي طريقة سهلة لنسخ المصفوفة كاملة. طريقة `slice` لا تُغير المصفوفة الأصلية، بل تُرجع مصفوفة جديدة.

إليك مثالاً:

```js
const arr = ["Cat", "Dog", "Tiger", "Zebra"];
const newArray = arr.slice(1, 3);
```

`newArray` ستكون قيمتها `["Dog", "Tiger"]`.

# --instructions--

استخدم طريقة `slice` في الدالة `sliceArray` لإرجاع جزء من مصفوفة `anim` بناءً على فهارسي `beginSlice` و `endSlice` المعطاة. يجب أن تُرجع الدالة مصفوفة.

# --hints--

يجب أن يستخدم شفرتك طريقة `slice`.

```js
assert(__helpers.removeJSComments(code).match(/\.slice/g));
```

يجب ألا تتغير قيمة المتغير `inputAnim`.

```js
assert(
  JSON.stringify(inputAnim) ===
    JSON.stringify(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 3)` يجب أن تُرجع `["Dog", "Tiger"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 3)) ===
    JSON.stringify(['Dog', 'Tiger'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 0, 1)` يجب أن تُرجع `["Cat"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 0, 1)) ===
    JSON.stringify(['Cat'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 4)` يجب أن تُرجع `["Dog", "Tiger", "Zebra"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 4)) ===
    JSON.stringify(['Dog', 'Tiger', 'Zebra'])
);
```

# --seed--

## --seed-contents--

```js
function sliceArray(anim, beginSlice, endSlice) {
  // Only change code below this line


  // Only change code above this line
}

const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
sliceArray(inputAnim, 1, 3);
```

# --solutions--

```js
function sliceArray(anim, beginSlice, endSlice) {
  return anim.slice(beginSlice, endSlice);
}
const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
```
