---
id: 5d822fd413a79914d39e98c9
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the CSS Variables Skyline project! Start by adding the `!DOCTYPE html` declaration at the top of the document so the browser knows what type of document it's reading.

# --hints--

Su código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debe incluir un espacio después de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debe definir el tipo de documento para que sea `html`.

```js
assert(code.match(/html/gi));
```

Debe cerrar la declaración `DOCTYPE` con `>` después del tipo.

```js
assert(code.match(/html\s*>/gi));
```

Su declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
