---
id: 5d822fd413a79914d39e98cc
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Dentro del `head`, anide un elemento `meta` con un `charset` de `UTF-8`, un elemento `title` con un título de `City Skyline` y un elemento `link` que vincule su archivo `styles.css`.

# --hints--

Puedes crear una etiqueta `meta` dentro de la etiqueta `head` element.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('meta'));
```

Debe asignar a la etiqueta `meta` un `charset` de `UTF-8`.

```js
const linkElement = document.querySelector('meta');
const charset = linkElement?.getAttribute("charset").toLowerCase();
assert.strictEqual(charset, 'utf-8');
```

Tu código debe tener una etiqueta `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

La etiqueta `title` debe estar dentro de la etiqueta `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

Tu proyecto debe tener estar titulado como `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Recuerda, las mayúsculas, las minúsculas y la ortografía son importantes para el título.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

Tu código debe tener una etiqueta `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

No debes cambiar las etiquetas existentes `head`. Asegúrate de no cerrar la etiqueta de cierre.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Debes tener un elemento vacío `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Su etiqueta `link` debe estar dentro de su etiqueta `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Tu etiqueta `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Tu etiqueta `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
