---
id: a39963a4c10bc8b4d4f06d7e
title: Eine Funktion zum Entfernen von Werten implementieren
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die als erstes Argument ein Anfangs-Array entgegennimmt, gefolgt von einem oder mehreren zusätzlichen Argumenten, die die zu entfernenden Werte darstellen.

**Ziel**: Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `destroyer`-Funktion haben, die ein Array und ein oder mehrere zusätzliche Argumente akzeptiert.
1. Die `destroyer`-Funktion sollte ein neues Array zurückgeben, das alle Elemente aus dem ersten Argument ausschließt, die mit einem der nachfolgenden Argumente übereinstimmen.
1. Die Funktion muss eine unbestimmte Anzahl von Argumenten akzeptieren.

# --hints--

Sie sollten eine `destroyer`-Funktion haben.

```js
assert.isFunction(destroyer);
```

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` sollte `[1, 1]` zurückgeben.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` sollte `[1, 5, 1]` zurückgeben.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` sollte `[1]` zurückgeben.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` sollte `[]` zurückgeben.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` sollte `["hamburger"]` zurückgeben.

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

`destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )` sollte `[12, 92, 65]` zurückgeben.

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```

