---
id: a103376db3ba46b2d50db289
title: 實作一個 Spinal Case Converter
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

Spinal case 是一種字串格式，所有字、單字皆為小寫並以連字號分隔。`"this-is-spinal-tap"` 是 spinal case 字串的範例。

在本實驗中，你將創建一個函式（程式），將給定的 `string` 轉換為 spinal case。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `spinalCase` 的函式（程式）。
2. `spinalCase` 函式（程式）應該接受一個單一引數，一個字串。
3. `spinalCase` 函式應該傳回以 spinal case 格式表示的字串。例如，如果引數是 `ProductLanding page`，該函式應該傳回 `product-landing-page`。

# --hints--

你應該創建一個名為 `spinalCase` 的函式（程式）。

```js
assert.isFunction(spinalCase);
```

`spinalCase` 應該接受單一引數。

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` 應該傳回字串 `this-is-spinal-tap`。

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` 應該傳回字串 `this-is-spinal-tap`。

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` 應該傳回字串 `the-andy-griffith-show`。

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` 應該傳回字串 `teletubbies-say-eh-oh`。

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` 應該傳回字串 `all-the-small-things`。

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
