---
id: a105e963526e7de52b219be9
title: Tekeleza muungano wa kipekee uliopangwa
challengeType: 26
dashedName: implement-a-unique-sorted-union
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite mitihani yote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi kinachoitwa `uniteUnique`.
1. Kitendakazi cha `uniteUnique` kinapaswa kupokea safu za data mbili au zaidi kama hoja.
1. Kitendakazi kinapaswa kurudisha safu mpya ya data inayojumuisha thamani za kipekee kutoka kwa safu za data za hoja, kwa mpangilio wa mara ya kwanza wanapopatikana katika hoja. Kwa mfano, ingizo kama `[1, 2, 4], [2, 3, 5]` kitakuwa na matokeo ya `[1, 2, 4, 3, 5]`.

# --hints--

Unapaswa kuwa na kitendakazi cha `uniteUnique`.

```js
assert.isFunction(uniteUnique);
```

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` inapaswa kurudisha `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` inapaswa kurudisha `[1, 2, 3, 5]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` inapaswa kurudisha `[1, 2, 3, 5, 4, 6, 7, 8]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` inapaswa kurudisha `[1, 3, 2, 5, 4, 6]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` inapaswa kurudisha `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const uniteUnique = (...arrays) => {
  return arrays.reduce((a, b) => {
    return [
      ...a,
      ...b.filter((e, currentIndex) => b.indexOf(e) === currentIndex && a.indexOf(e) === -1)
    ];
  }, []);
};
```
