---
id: 663255f28c59315db74d137b
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

CamperBot이 `getRandomIndex` 함수를 완성했고 예상대로 작동하고 있습니다.

하지만 이제 DOM에서 `body` 요소에 참조를 만들려고 할 때 다음 문제가 발생하고 있습니다:

```js
Uncaught TypeError: document.queryselector is not a function
```

`TypeError`는 코드가 예상한 유형이 아닌 값에 대해 작업을 수행하려고 한다는 뜻입니다.

DOM에서 요소를 선택하는 올바른 메서드 이름으로 `TypeError` 메서드를 업데이트하여 `document.queryselector`를 고치세요.

# --hints--

`queryselector`를 카멜 케이스로 업데이트해야 합니다.

```js
assert.match(code, /querySelector/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```


```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
  console.log(Math.floor(darkColorsArr.length * Math.random()));
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

--fcc-editable-region--
const body = document.queryselector("body");
--fcc-editable-region--
```
