---
id: 663260de72634166b0800fe9
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

CamperBot이 이제 `changeBackgroundColor`라는 함수를 만들었으며, 이 함수는 페이지의 배경색을 `darkColorsArr` 배열에서 무작위 색상으로 변경합니다. 함수는 또한 그 새 색상의 16진수 코드를 표시합니다.

이 함수를 테스트하려고 할 때, 배경색이 바뀌지 않고 텍스트에 다음과 같은 내용이 나타나는 것을 확인합니다:

```js
Hex Code: undefined
```

`undefined`가 여기 나타나는 이유는 `color` 변수가 올바르게 설정되지 않았기 때문입니다.

`darkColorsArr[getRandomIndex]` 줄의 오류를 고쳐서 `color` 변수가 `darkColorsArr` 배열에서 무작위 색상으로 설정되도록 하세요.

# --hints--

`getRandomIndex` 줄 안에서 `darkColorsArr[getRandomIndex]` 함수를 호출해야 합니다.

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

--fcc-editable-region--
function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
changeBackgroundColor();
--fcc-editable-region--
```
