---
id: 66323433f931ca32305a11f5
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

이제 CamperBot이 `darkColorsArr`에서 무작위 인덱스를 반환하는 함수를 만들려고 합니다. 하지만 다음과 같은 오류 메시지가 발생했습니다:

```js
Uncaught ReferenceError: math is not defined
```

존재하지 않는 변수를 참조할 때 `ReferenceError`가 발생합니다. 이번 경우에는 CamperBot이 `math`를 사용하려고 하는데 JavaScript에는 `math` 객체가 없습니다.

`math.random()` 줄에서 CamperBot의 오류를 고치고 콘솔을 다시 열어보세요.

# --hints--

`math.random()` 줄의 대소문자 오류를 고쳐야 합니다.

```js
assert.match(getRandomIndex.toString(), /console\.log\(\s*darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
--fcc-editable-region--
  console.log(darkColorsArr.length * math.random())
--fcc-editable-region--
}
getRandomIndex();
```
