---
id: 66325a250690a3612c1db0f6
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

CamperBot이 `bgHexCodeSpanElement`라는 새 변수를 만들어 `span`라는 `id`를 가진 `bg-hex-code` 요소에 대한 참조를 저장했습니다. 하지만 그 변수를 콘솔에 출력하려고 하면 `null`가 나타납니다.

`null`는 JavaScript에서 값이 없음을 나타내는 특별한 값입니다. 이는 존재하지 않는 객체의 속성에 접근하려 할 때 발생할 수 있습니다.

이번 경우에는 CamperBot이 `document.querySelector` 메서드에 올바른 선택자를 전달하지 않았습니다.

`document.querySelector("bg-hex-code")` 줄을 고쳐 `id`라는 `bg-hex-code`를 가진 요소를 올바르게 선택하도록 하세요.

# --hints--

`"bg-hex-code"`을 고쳐 `id`의 `bg-hex-code`를 올바르게 참조하도록 하세요. CSS에서 `id` 선택자는 해시 기호(`#`)로 시작한다는 점을 기억하세요.

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

--fcc-editable-region--
const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("bg-hex-code");

console.log(bgHexCodeSpanElement);
--fcc-editable-region--
```
