---
id: 6650c9a94d6e13d14a043a69
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

CamperBot이 무작위 배경색 변경기를 만들려고 합니다. 하지만 계속 문제가 발생해서 코드를 디버그하는 데 도움이 필요합니다.

CamperBot은 이미 프로젝트에 필요한 HTML과 CSS를 추가했지만, 페이지에 스타일과 내용이 전혀 나타나지 않는 이유를 잘 모르고 있습니다.

콘솔을 열어보면 다음 메시지가 보입니다:

```js
SyntaxError: unknown: Unexpected token, expected "," (5:2)
```

JavaScript 엔진이 해석할 수 없는 내용을 만나면 구문 오류가 발생합니다. 이번 경우에는 CamperBot이 `darkColorsArr` 배열에 구문 오류가 있는 것 같습니다.

`darkColorsArr` 배열의 구문 오류를 고치면 페이지에 내용과 스타일이 나타날 것입니다.

# --hints--

`darkColorsArr` 배열의 구문 오류를 해결하세요. 배열 요소는 쉼표로 구분해야 한다는 점을 기억하세요.

```js
assert.isArray(darkColorsArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
--fcc-editable-region--
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C"
  "#616A6B"
  "#4A235A"
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
--fcc-editable-region--
```
