---
id: 67fccc2463253b213a000142
title: بناء مشغل وسائط متعددة
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

في الدروس السابقة، تعرفت على كيفية التعامل مع عناصر `audio` و`video`. في هذا المختبر، ستبني مشغل وسائط متعددة يعرض مسار `audio` و`video` مع نص مكتوب.

لعنصر `audio`، تحتاج إلى تضمين عنصر `source` الذي يُستخدم لتحديد الوسائط المستخدمة.

إليك مثالاً:

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

يمكن أيضًا استخدام عنصر `source` داخل عنصر `video` بهذه الطريقة:

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تجتاز لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر `h1` للعنوان الرئيسي للصفحة.
2. يجب أن يكون لديك ثلاثة عناصر `section`.
3. داخل أول عنصر `section`، يجب أن يكون لديك عنصر `h2` لعنوان الأغنية التي تُشغل.
4. أسفل عنصر `h2`، يجب أن يكون لديك عنصر `audio` مع خاصية `controls` وخاصية `aria-label`.
5. داخل عنصر `audio`، يجب أن يكون لديك عنصر `source` مع خاصية `src` تشير إلى ملف صوتي وخاصية `type`. يمكنك استخدام هذا الرابط الصوتي إذا أردت: `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. داخل ثاني عنصر `section`، يجب أن يكون لديك عنصر `h2` لعنوان الفيديو الذي يُشغل.
7. أسفل عنصر `h2`، يجب أن يكون لديك عنصر `video` مع خصائص `controls` و `width` وخاصية `aria-label`.
8. داخل عنصر `video`، يجب أن يكون لديك عنصر `source` مع خاصية `src` تشير إلى ملف فيديو وخاصية `type`. يمكنك استخدام هذا الرابط للفيديو إذا أردت: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. أسفل عنصر `source`، يجب أن يكون لديك عنصر `track` مع خاصية `src` تشير إلى ملف ترجمات وخاصية `kind`، وخاصية `srclang`، وخاصية `label`. يمكنك استخدام هذا الرابط للترجمات إذا أردت: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. داخل ثالث عنصر `section`، يجب أن يكون لديك عنصر `h2` لعنوان القسم مثل "النص المكتوب".
11. أسفل عنصر `h2`، يجب أن يكون لديك عنصر `p` يحتوي على النص المكتوب للفيديو.

# --hints--

يجب أن يكون لديك عنصر `h1`.

```js
assert.isNotNull(document.querySelector("h1"));
```

يجب أن يحتوي عنصر `h1` على نص يمثل العنوان الرئيسي للصفحة.

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```  

يجب أن يكون لديك ثلاثة عناصر `section`.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

يجب أن يحتوي أول عنصر `section` على عنصر `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

يجب أن يحتوي عنصر `h2` على نص يمثل عنوان الأغنية التي تُشغل.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

يجب أن يحتوي أول عنصر `section` على عنصر `audio`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

يجب أن يحتوي عنصر `audio` على خاصية `controls`.

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

يجب أن يحتوي عنصر `audio` على خاصية `aria-label` مع نص يصف الصوت.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

يجب أن يحتوي عنصر `audio` على عنصر `source`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

يجب أن يحتوي عنصر `source` على خاصية `src` تشير إلى ملف صوتي.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

يجب أن يحتوي عنصر `source` على خاصية `type` بقيمة تحدد نوع الصوت.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

يجب أن يحتوي ثاني عنصر `section` على عنصر `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

يجب أن يحتوي عنصر `h2` على نص يمثل عنوان الفيديو الذي يُشغل.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

يجب أن يحتوي ثاني عنصر `section` على عنصر `video`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

يجب أن يحتوي عنصر `video` على خاصية `controls`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

يجب أن يحتوي عنصر `video` على خاصية `width` بقيمة تحدد عرض الفيديو.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

يجب أن يحتوي عنصر `video` على عنصر `source`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

يجب أن يحتوي عنصر `source` على خاصية `src` تشير إلى ملف فيديو.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

يجب أن يحتوي عنصر `source` على خاصية `type` بقيمة تحدد نوع الفيديو.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

يجب أن يحتوي عنصر `video` على عنصر `track`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

يجب أن يحتوي عنصر `track` على خاصية `src` تشير إلى ملف ترجمات.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

يجب أن يحتوي عنصر `track` على خاصية `kind` بنص يصف نوع الملف.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

يجب أن يحتوي عنصر `track` على خاصية `srclang` بنص يصف لغة الترجمات.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

يجب أن يحتوي عنصر `track` على خاصية `label` بنص يصف لغة الترجمات.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

يجب أن يحتوي ثالث عنصر `section` على عنصر `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

يجب أن يحتوي عنصر `h2` على نص يمثل عنوان القسم.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

يجب أن يحتوي ثالث عنصر `section` على عنصر `p` للنص المكتوب.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
