---
id: 694a9ff65d4472dd451d8255
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utazoea jinsi ya kutumia thamani za kweli au siyo kweli na kauli za masharti katika Python kwa kujenga kalkuleta ya kuhifadhi tiketi za sinema.

Anza kwa kutangaza vigezo vinavyohifadhi taarifa kuhusu mtumizi na onyesho la sinema.

Tengeneza kigezo kinachoitwa `base_price` kuhifadhi bei msingi ya tiketi ya sinema na weka thamani yake kuwa `15`. Tengeneza kigezo kingine kinachoitwa `age` kuhifadhi umri wa mtumizi na weka thamani yake kuwa `21`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `base_price`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

Unapaswa kuweka thamani kamili `15` kwa `base_price`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

Unapaswa kutangaza kigezo kinachoitwa `age`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

Unapaswa kuweka thamani kamili `21` kwa `age`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
