---
id: 694cc1b515976889afd4cded
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Sasa unahitaji kuhifadhi baadhi ya thamani za mfuatano wa herufi kuhusu tiketi ya sinema.

Tangaza kigezo kinachoitwa `seat_type` kuhifadhi aina ya kiti ambacho mtumizi ameichagua na weka thamani yake kuwa mfuatano wa herufi `Gold`. Tangaza kigezo kingine kinachoitwa `show_time` kuhifadhi muda wa kuonyeshwa kwa sinema na weka thamani yake kuwa mfuatano wa herufi `Evening`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `seat_type`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `Gold` kwa `seat_type`. Kumbuka kuzunguka thamani hiyo kwa alama za nukuu moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

Unapaswa kutangaza kigezo kinachoitwa `show_time`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `Evening` kwa `show_time`. Kumbuka kuzunguka thamani hiyo kwa alama za nukuu moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
