---
id: 6958dd85587e10c99162dd81
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Kama ulivyojifunza katika mafundisho yaliyopita, katika Python, tamko la `if` linaweza kutumika kuendesha msimbo kulingana na kama hali ni kweli.

Tamko la `if` linajumuisha neno kuu la `if`, linalofuatiwa na hali na nukta. Msimbo unaotakiwa kuendeshwa wakati hali ni kweli, ambao lazima uwe umeingizwa, unaitwa mwili wa tamko la `if`.

```py
if condition:
    # Code to execute if condition is True
```

Katika hatua hii, utahakikisha kama mtumizi anastahili kuagiza tiketi ya sinema kulingana na umri wake.

Tengeneza tamko la `if` ili kuangalia kama `age` ni kubwa kuliko `17`. Ndani ya mwili wa tamko la `if`, chapisha `User is eligible to book a ticket`. Hii itachapisha ujumbe tu wakati umri wa mtumizi ni mkubwa kuliko `17`.

Kumbuka kuingiza mwili wa tamko la `if` na kuzunguka ujumbe na nukuu moja au alama za nukuu maradufu ndani ya wito wa `print()`.

# --hints--

Unapaswa kuwa na tamko la `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

Hali ya tamko la `if` inapaswa kuwa `age > 17`.

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

Unapaswa kuwa na wito wa `print()` ndani ya tamko lako la `if`. Kumbuka kuingiza mwili wa tamko la `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

Unapaswa kuchapisha `User is eligible to book a ticket` ndani ya tamko lako la `if`. Kumbuka kuzunguka ujumbe na nukuu moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
