---
id: 6958dfa0c74b237e1b56ce04
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Mtumizi anastahili punguzo la uanachama ikiwa ni mwanachama.

Tangaza kigezo kinachoitwa `discount` na weka thamani yake kuwa `0`. Hii itahifadhi punguzo ambalo mtumizi anapata kwenye tiketi ya sinema.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `discount`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

Unapaswa kuweka thamani ya nambari kamili `0` kwa `discount`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
