---
id: 6958e2377cacbc46a7f7383a
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Sasa, ongeza kifungu cha `else` kwenye tamko lako la `if` na chapisha `No extra charges will be applied` ndani ya mwili wa `else`. Hii itachapisha ujumbe tu wakati hali ya ada za ziada ni `false`.

Kisha, chini ya kifungu cha `else`, tumia wito wa `print()` kuonyesha ujumbe unaoonyesha `Extra charges:` ikifuatiwa na thamani iliyosasishwa ya `extra_charges` na hakiki matokeo kwenye terminali.

# --hints--

Unapaswa kuwa na kifungu cha `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

Unapaswa kuchapisha `No extra charges will be applied` ndani ya kifungu chako cha `else`. Kumbuka kuzunguka ujumbe huo na nukuu moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

Unapaswa kuwa na `print('Extra charges:', extra_charges)` chini ya tamko lako la `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
