---
id: 6958e59845777391b86d7331
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Katika Python, tamko la `if` pia linaweza kuwekwa ndani ya mwili wa tamko lingine la `if`. Hii huitwa tamko la `if` lililopangwa ndani.

Tamko la `if` lililopangwa ndani linakuwezesha kukagua hali ya ziada tu baada ya hali ya kwanza kutimizwa. Tamko la `if` lililo ndani litaendeshwa tu ikiwa hali ya `if` ya nje ni kweli.

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

Sasa utahesabu ada za huduma kulingana na aina ya kiti ambacho mtumizi ameichagua.

Ndani ya mwili wa tamko la mwisho la `if`, chini ya mstari wa `print('Ticket booking condition satisfied')`, tengeneza kigezo kinachoitwa `service_charges` na uweke thamani yake kuwa `0`. Hakikisha msimbo wako umeingizwa kwa nafasi nne ili kubaki ndani ya mwili wa tamko la nje la `if`.

Kisha, tengeneza tamko la `if` lililopangwa ndani ili ukague kama `seat_type` ni sawa na `Premium`. Ndani ya mwili wa tamko la `if` lililopangwa ndani, sasisha thamani ya `service_charges` kuwa `5`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `service_charges` ndani ya mwili wa tamko lako la mwisho la `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

Unapaswa kuweka thamani ya nambari kamili `0` kwa `service_charges`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

Unapaswa kuwa na tamko la `if` ndani ya mwili wa tamko lako la mwisho la `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

Hali ya tamko jipya la `if` inapaswa kuwa `seat_type == 'Premium'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

Unapaswa kuweka thamani ya nambari kamili `5` kwa kigezo cha `service_charges` ndani ya tamko lako jipya la `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
