---
id: 6958e1875e592b58e8852757
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Sasa tengeneza kigezo kinachoitwa `extra_charges` na weka thamani yake kuwa `0`. Hii itawakilisha ada za ziada za kuongezwa tiketi ya sinema wakati wa wikendi.

Kisha, tengeneza usemi wa `if` kuangalia kama `is_weekend` ni kweli. Ndani ya mwili wa usemi wa `if`, sasisha thamani ya `extra_charges` kuwa `2` na chapisha `Extra charges will be applied` kwenye terminali.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `extra_charges`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

Unapaswa kuweka thamani ya nambari kamili `0` kwa `extra_charges`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

Unapaswa kuwa na usemi wa nne wa `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

Sharti la usemi wako wa nne wa `if` linapaswa kuwa `is_weekend`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

Unapaswa kuweka thamani ya nambari kamili `2` kwa kigezo cha `extra_charges` ndani ya usemi wako mpya wa `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

Unapaswa kuchapisha `Extra charges will be applied` ndani ya usemi wako mpya wa `if`. Kumbuka kuzunguka ujumbe huo na nukuu za moja kwa moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
