---
id: 6958de0d773289a89515a6b8
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Sasa utakagua kama mtumizi ameruhusiwa kuhifadhi tiketi ya onyesho la jioni kulingana na umri wake.

Tengeneza tamko la `if` ili ukague kama `age` ni kubwa au sawa na `21`. Ndani ya mwili wa tamko la `if`, chapisha `User is eligible for Evening shows`.

# --hints--

Unapaswa kuwa na tamko mbili za `if` katika msimbo wako.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

Sharti la tamko lako la pili la `if` linapaswa kuwa `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Unapaswa kuchapisha `User is eligible for Evening shows` ndani ya tamko lako jipya la `if`. Kumbuka kuzunguka ujumbe huo na nukuu moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
