---
id: 6958e0ecc121b0ae3dd9c5e0
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Unapaswa pia kushughulikia hali ambapo mtumizi hapati punguzo la uanachama.

Ongeza kifungu cha `else` kwenye tamko lako la `if` na chapisha `User does not qualify for membership discount` ndani ya mwili wa `else`.

Unaweza kuchapisha ujumbe ukifuatiwa na kigezo kama hii:

```py
print('Total:', total)
```

Pia unataka kuonyesha thamani iliyosasishwa ya `discount`. Chini ya tamko la `if...else`, tumia wito wa `print()` kuonyesha ujumbe unaoonyesha `Discount:` ukifuatiwa na thamani iliyosasishwa ya `discount`. Kisha, hakiki matokeo kwenye terminali.

# --hints--

Unapaswa kuwa na kifungu cha `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

Unapaswa kuchapisha `User does not qualify for membership discount` ndani ya kifungu chako cha `else`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

Unapaswa kuwa na `print('Discount:', discount)` chini ya tamko lako la `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
