---
id: 6960ba6e6c95ed4ad896c6f2
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Kuna kivunjiko kimoja zaidi kwa sheria za uhifadhi. Watumizi kati ya `18` na `21` wanaweza kuhifadhi maonyesho ya jioni ikiwa ni wanachama.

Wakati opereta za kimantiki nyingi zinapotumika katika usemi wa `if`, masharti yaliounganishwa na `and` yanapimwa kabla ya masharti yaliounganishwa na `or`. Mabano ya kawaida `()` hutumika katika Python kuunganisha masharti na kudhibiti mpangilio wa jinsi yanavyopimwa.

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

Sahihisha masharti ya mstari wa `if age >= 21 or age >= 18 and show_time != 'Evening':` kuongeza sharti lingine kwa kutumia opereta wa `or` kuangalia kama `is_member` ni kweli. Tumia mabano ya kawaida `()` kuunganisha masharti ya `show_time != 'Evening'` na `is_member` pamoja kama ilivyoonyeshwa katika mfano hapo juu.

# --hints--

Sharti lako la mwisho la usemi wa `if` linapaswa kuwa `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
