---
id: 6960a5165f7efe52a0c0aa39
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Katika Python, opereta `and` hutumika kuangalia kama masharti mengi ni ya kweli. Hapa ni jinsi unavyoweza kuitumia kuunganisha masharti mawili katika usemi wa `if`:

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

Punguzo la uanachama linapaswa kutumika kwa wanachama tu ikiwa `age` yao ni kubwa au sawa na `21`.

Sahihisha sharti la mstari wa `if is_member:` kwa kutumia opereta `and` kuunganisha sharti lililopo na sharti jingine linaloangalia kama `age` ni kubwa au sawa na `21`.

# --hints--

Sharti lako la tatu la usemi wa `if` linapaswa kuwa `is_member and age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
