---
id: 696099abb6acc539b27055a1
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Katika hatua iliyopita, uligundua hali kwa kutumia usemi wa `if`. Kifungu cha `else` hutumika kushughulikia hali wakati hali hiyo ni potofu. Huu ni sintaksia ya usemi wa `if...else`:

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

Sasa, ongeza kifungu cha `else` kwenye usemi wako wa `if` na chapisha `User is not eligible for Evening shows` ndani ya mwili wa `else`. Hii itachapisha tu wakati hali katika usemi wa `if` ni potofu.

# --hints--

Unapaswa kuwa na kifungu cha `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

Unapaswa kuchapisha `User is not eligible for Evening shows` ndani ya kifungu chako cha `else`. Kumbuka kuzunguka ujumbe huo na nukuu moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
