---
id: 6960b6a6ea06e27dcc7eef37
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Sasa utaangalia kama mtumizi anakidhi masharti ya kuhifadhi tiketi ya sinema. Watumizi wenye umri wa `21` au zaidi wanaweza kuhifadhi tiketi bila vikwazo vyovyote.

Tengeneza usemi wa `if` kuangalia kama `age` ni kubwa au sawa na `21`. Ndani ya mwili wa usemi wa `if`, chapisha `Ticket booking condition satisfied` kwenye terminali.

Kisha, ongeza kifungu cha `else` kwenye usemi wako wa `if` na chapisha `Ticket booking failed due to restrictions` ndani ya mwili wa `else`.

# --hints--

Unapaswa kuwa na usemi wa tano wa `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

Sharti la usemi wako wa tano wa `if` linapaswa kuwa `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Unapaswa kuchapisha `Ticket booking condition satisfied` ndani ya usemi wako mpya wa `if`. Kumbuka kuzunguka ujumbe huo na nukuu moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

Usemi wako mpya wa `if` unapaswa kuwa na kifungu cha `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

Unapaswa kuchapisha `Ticket booking failed due to restrictions` ndani ya kifungu chako cha `else`. Kumbuka kuzunguka ujumbe huo na nukuu moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
