---
id: 696099c1ac72a4e585de1d5b
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Baadhi ya taarifa zinaweza kuwa kweli au si kweli tu. Kama ulivyojifunza katika mafundisho yaliyopita, hii inaweza kuwakilishwa kwa kutumia thamani za kweli au siyo kweli.

Tengeneza kigezo kinachoitwa `is_member` kuonyesha kama mtumizi ni mwanachama na weka thamani yake kuwa `True`.

Chini ya kigezo `is_member` tengeneza kigezo kingine kinachoitwa `is_weekend` kuonyesha kama onesho la sinema ni mwishoni mwa wiki na weka thamani yake kuwa `False`. Usizungushe thamani hiyo kwa nukuu.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `is_member`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

Unapaswa kuweka thamani ya kweli au siyo kweli `True` kwa `is_member`. Usizungushe thamani hiyo kwa nukuu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

Unapaswa kutangaza kigezo kinachoitwa `is_weekend`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

Unapaswa kuweka thamani ya kweli au siyo kweli `False` kwa `is_weekend`. Usizungushe thamani hiyo kwa nukuu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
