---
id: 696cb889fc4fdeb337e52864
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Katika Python, masharti ya `if` hayahitaji kulinganisha thamani moja kwa moja na `True` au `False`. Badala yake, hutegemea uhalali wa thamani. Thamani za kweli huangaliwa kuwa `True` katika muktadha wa thamani ya kweli au siyo kweli, kama vile sharti la `if`. Hizi ni pamoja na nambari zisizo sifuri, mfuatano wa herufi usio tupu, na thamani ya kweli au siyo kweli `True`.

```py
if value:
   print('value is truthy')
```

Tengeneza tamko la `if` ili ukague kama `is_member` ni thamani ya kweli. Ndani ya mwili wa tamko la `if`, sasisha thamani ya `discount` kuwa `3` na chapisha `User qualifies for membership discount` kwenye terminali.

# --hints--

Unapaswa kuwa na tamko tatu za `if` katika msimbo wako.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

Sharti la tamko lako la tatu la `if` linapaswa kuwa `is_member`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

Unapaswa kuweka thamani ya nambari kamili `3` kwa mabadiliko ya `discount` ndani ya tamko lako jipya la `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

Unapaswa kuchapisha `User qualifies for membership discount` ndani ya tamko lako jipya la `if`. Kumbuka kuzunguka ujumbe huo na nukuu moja au alama za nukuu maradufu.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
