---
id: 673339d96c56d36e963e888e
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 JavaScript의 Map과 Set 객체를 연습하면서 식물 묘목장 카탈로그를 만들어 보세요.

카탈로그에 있는 각 식물을 나타내기 위해, 세 가지 속성을 가진 객체를 사용합니다:

- `commonName`, 식물의 일반 이름
- `scientificName`, 식물의 학명(또는 라틴어 이름)
- `cultivar`, 식물의 품종

`Spanish lavender`라는 품종의 `Lavandula stoechas` (학명 `Ballerina`)에 대해 위 세 가지 속성을 가진 객체를 생성하고 `ballerina`라는 변수에 할당하는 것부터 시작하세요.

# --hints--

`ballerina`라는 변수가 있어야 합니다.

```js
assert.exists(ballerina);
```

`ballerina` 변수에 객체를 할당해야 합니다.

```js
assert.isObject(ballerina);
```

`ballerina` 객체에 `commonName` 속성이 `"Spanish lavender"`로 설정되어야 합니다.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

`ballerina` 객체에 `scientificName` 속성이 `"Lavandula stoechas"`로 설정되어야 합니다.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

`ballerina` 객체에 `cultivar` 속성이 `"Ballerina"`로 설정되어야 합니다.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

`ballerina` 객체에는 세 개의 키가 있어야 합니다.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
