---
id: 6733b5e6f210c1bd1afdf8b3
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

나머지 식물 객체들이 추가되었습니다. 이제 카탈로그를 만들어 보세요.

일반 JavaScript 객체처럼 `Map` 객체는 키-값 쌍 모음을 포함합니다. `Map` 객체를 만들려면 `Map` 생성자를 사용할 수 있습니다:

```js
const myMap = new Map();
```

`catalog` 변수를 선언하고 `Map` 객체를 할당하세요.

# --hints--

`catalog`라는 변수가 있어야 합니다.

```js
assert.exists(catalog);
```

`Map` 변수에 `catalog` 객체를 할당해야 합니다.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
