---
id: 6733b5632fb854bc6c58dc0d
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이제 세 가지 속성이 같은 또 다른 객체를 만들어 `Pretty Polly`라는 스페인 라벤더 품종을 나타내고, `prettyPolly`라는 변수에 할당하세요.

# --hints--

`prettyPolly`라는 변수가 있어야 합니다.

```js
assert.exists(prettyPolly);
```

`prettyPolly` 변수에 객체를 할당해야 합니다.

```js
assert.isObject(prettyPolly);
```

`prettyPolly` 객체에 `commonName` 속성이 `"Spanish lavender"`로 설정되어야 합니다.

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

`prettyPolly` 객체에 `scientificName` 속성이 `"Lavandula stoechas"`로 설정되어야 합니다.

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

`prettyPolly` 객체에 `cultivar` 속성이 `"Pretty Polly"`로 설정되어야 합니다.

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

`prettyPolly` 객체에 세 개의 키가 있어야 합니다.

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
