---
id: 6734a2385c52c54d916a8b86
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

카탈로그에 또 다른 키-값 쌍을 추가하세요. `prettyPolly` 변수를 키로 사용하고 값은 다시 정수인 `small`, `medium`, `large` 속성을 가진 객체로 만드세요.

# --hints--

`prettyPolly` 객체를 `catalog` 맵의 키로 설정해야 합니다.

```js
assert.isTrue(catalog.has(prettyPolly));
```

`prettyPolly` 맵에서 `catalog` 키의 값은 `small`, `medium`, `large` 속성을 가진 객체여야 합니다.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`small` 맵에서 `prettyPolly` 키와 연결된 객체의 `catalog` 속성은 양의 정수여야 합니다.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`medium` 맵에서 `prettyPolly` 키와 연결된 객체의 `catalog` 속성은 양의 정수여야 합니다.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`large` 맵에서 `prettyPolly` 키와 연결된 객체의 `catalog` 속성은 양의 정수여야 합니다.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

여전히 `ballerina` 키가 `catalog` 맵에 있어야 합니다.

```js
assert.isTrue(catalog.has(ballerina));
```

`catalog` 맵에는 두 개의 요소가 있어야 합니다.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
