---
id: 6734c10a327760a665b7d5b9
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

다음 몇 단계에서는 식물을 판매할 때 카탈로그를 업데이트할 수 있는 함수를 작성할 것입니다.

식물 객체, 화분 크기, 판매할 화분 수의 세 가지 인수를 순서대로 받는 `sellPlants`라는 이름의 함수를 만드세요.

지정한 화분 수가 사용 가능한 화분 수보다 많으면 함수가 `Not enough <pot-size> size pots for <item-name>. Only <pot-number> left.`을 반환하도록 하세요. `<pot-size>`는 화분 크기로, `<pot-number>`은 해당 크기의 남은 화분 수로, `<item-name>`는 식물의 학명과 한 칸 띄운 뒤 작은따옴표로 감싼 품종 이름으로 바꾸세요.

모든 것이 제대로 작동하는지 확인하려면 `sellPlants` 객체, `ballerina`, `"small"`을 인수로 하여 `25` 함수를 호출한 결과를 콘솔에 출력하세요.

# --hints--

`catalog` 맵은 수정하지 마세요.

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

`sellPlants`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(sellPlants);
```

`sellPlants` 함수는 세 개의 인수를 받아야 합니다.

```js
assert.lengthOf(sellPlants, 3);
```

`sellPlants(ballerina, "small", 25)`은 `Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.`을 반환해야 합니다.

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

`sellPlants(ballerina, "small", 25)`을 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)`은 `undefined`을 반환해야 합니다.

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog)
```
