---
id: 6734a772196245601e4b80e0
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

맵에서 요소를 가져와야 할 때는 `get` 메서드를 사용할 수 있습니다. 이 메서드는 키를 인수로 받아 연결된 값을 반환합니다.

`get` 메서드를 사용해 `ballerina` 키의 값을 카탈로그에서 가져와 `getBallerina`라는 변수에 할당하세요. 그런 다음 새 변수를 콘솔에 출력하세요.

# --hints--

`getBallerina`라는 변수가 있어야 합니다.

```js
assert.exists(getBallerina);
```

`getBallerina` 변수는 `catalog.get(ballerina)` 값을 가져야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+getBallerina\s*=\s*catalog\s*\.\s*get\s*\(\s*ballerina\s*\)/);
```

`getBallerina` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*getBallerina\s*\)/);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
```
