---
id: 6734f0ae9d1cf4dfd43a7805
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

맵 객체에서 요소를 삭제하려면 키를 인수로 받아 해당 키와 연결된 항목을 삭제하는 `delete` 메서드를 사용할 수 있습니다. 이 메서드는 삭제가 성공했는지 여부를 나타내는 불리언 값을 반환합니다.

`removePlant`이라는 이름의 함수를 만들어 식물 객체를 인수로 받아 카탈로그에서 삭제하세요. 지정한 식물이 카탈로그에서 삭제되었는지에 따라 `true` 또는 `false`을 반환하게 만드세요.

그다음, `"ballerina"` 문자열을 인수로 하여 새로 만든 함수를 호출한 결과를 콘솔에 기록하세요.

# --hints--

`removePlant`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(removePlant);
```

`removePlant` 함수는 하나의 인수를 받아야 합니다.

```js
assert.lengthOf(removePlant, 1);
```

`removePlant` 함수는 `delete`에서 `catalog` 메서드를 호출해야 합니다.

```js
let flag = false;
const temp = catalog.delete;
try {
    catalog.delete = (p) => {flag = true}
    removePlant("");
    assert.isTrue(flag);
} finally {
    catalog.delete = temp;
}
```

`removePlant(item)`은 카탈로그에서 `item`을 삭제해야 합니다.

```js
const testArr = [];
const temp = catalog.delete;
try {
    catalog.delete = (p) => {testArr.push(p)}
    removePlant("test");
    assert.deepEqual(testArr, ["test"]);
} finally {
    catalog.delete = temp;
}
```

`removePlant` 함수는 삭제가 성공했는지 여부를 나타내는 불리언 값을 반환해야 합니다.

```js
try {
    assert.isFalse(removePlant("test"));
    assert.isTrue(removePlant(royalCrown));
} finally {
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 })
}
```

`removePlant("ballerina")`을 콘솔에 기록해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*removePlant\s*\(\s*("|')ballerina\1\s*\)\s*\)/);
```



# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

--fcc-editable-region--

--fcc-editable-region--

const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```
