---
id: 67349d6e726b2c3c7f7301e8
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

Map 객체의 `size` 속성은 맵에 있는 요소 개수를 반환합니다. `size`의 `catalog`을 콘솔에 기록하세요.

# --hints--

`size` 맵의 `catalog` 속성을 콘솔에 기록해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\.\s*size\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
console.log(catalog);

--fcc-editable-region--
```
