---
id: 6734a606e5968a5a19b77171
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

이전 단계에서 했던 것처럼 나머지 식물들을 카탈로그에 계속 추가하세요.

# --hints--

`willowVale` 객체를 `catalog` 맵의 키로 설정하세요.

```js
assert.isTrue(catalog.has(willowVale));
```

`willowVale` 맵의 `catalog` 키 값은 `small`, `medium`, `large` 속성을 가진 객체여야 하며, 각각 양의 정수 값을 가져야 합니다.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`hidcote` 객체를 `catalog` 맵의 키로 설정하세요.

```js
assert.isTrue(catalog.has(hidcote));
```

`hidcote` 맵의 `catalog` 키 값은 `small`, `medium`, `large` 속성을 가진 객체여야 하며, 각각 양의 정수 값을 가져야 합니다.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`imperialGem` 객체를 `catalog` 맵의 키로 설정하세요.

```js
assert.isTrue(catalog.has(imperialGem));
```

`imperialGem` 맵의 `catalog` 키 값은 `small`, `medium`, `large` 속성을 가진 객체여야 하며, 각각 양의 정수 값을 가져야 합니다.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`royalCrown` 객체를 `catalog` 맵의 키로 설정하세요.

```js
assert.isTrue(catalog.has(royalCrown));
```

`royalCrown` 맵의 `catalog` 키 값은 `small`, `medium`, `large` 속성을 가진 객체여야 하며, 각각 양의 정수 값을 가져야 합니다.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`ballerina` 맵에 `catalog` 키가 여전히 있어야 합니다.

```js
assert.isTrue(catalog.has(ballerina));
```

`prettyPolly` 맵에 `catalog` 키가 여전히 있어야 합니다.

```js
assert.isTrue(catalog.has(prettyPolly));
```

`catalog` 맵에는 여섯 개의 요소가 있어야 합니다.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
