---
id: 6734a13a29acdf47f2b8e161
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

일반 객체는 키가 항상 문자열이어야 하지만, map 객체는 어떤 데이터 타입이든 키로 사용할 수 있습니다. 식물 목록에서는 각 식물 객체를 키로 사용하고, 다양한 크기의 화분 수를 담은 객체를 값으로 연결하는 객체-객체 키-값 쌍을 사용할 것입니다.

`set` 호출을 수정하여 `ballerina` 객체를 키로 사용하세요. 값으로는 `small`, `medium`, `large` 세 가지 속성을 가진 객체를 사용하며, 각각은 양의 정수 값을 가져야 합니다.

# --hints--

`ballerina` 객체를 `catalog` map의 키로 설정해야 합니다.

```js
assert.isTrue(catalog.has(ballerina));
```

`ballerina` map에서 `catalog` 키의 값은 `small`, `medium`, `large` 속성을 가진 객체여야 합니다.

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`small` map에서 `ballerina` 키에 연결된 객체의 `catalog` 속성은 양의 정수여야 합니다.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`medium` map에서 `ballerina` 키에 연결된 객체의 `catalog` 속성은 양의 정수여야 합니다.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`large` map에서 `ballerina` 키에 연결된 객체의 `catalog` 속성은 양의 정수여야 합니다.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

`catalog` map에는 단 하나의 키-값 쌍만 포함되어야 합니다.

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
