---
id: 67349ba230cff237c0f2bdf2
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

Map 객체는 키가 삽입된 순서를 기억합니다. map 객체에 키-값 쌍을 설정하려면 `set` 메서드를 사용할 수 있습니다. 다음은 `"name"` map 객체에 키 `"Jerry"`를 값 `myMap`로 설정하는 예입니다.

```js
myMap.set("name", "Jerry");
```

`"ballerina"` map에 `ballerina` 키를 `catalog` 객체의 값으로 설정하세요. 그런 다음, 콘솔에 카탈로그를 출력하세요.

map 객체를 제대로 보려면 브라우저 콘솔을 열어야 한다는 점을 참고하세요.

# --hints--

`"ballerina"` map에 `ballerina` 객체의 값을 `catalog` 키로 설정해야 합니다.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

`catalog` map에는 하나의 키-값 쌍이 포함되어야 합니다.

```js
assert.lengthOf(catalog, 1);
```

`catalog` map을 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
