---
id: 67350acfd2006f2777295b96
title: 19단계
challengeType: 1
dashedName: step-19
---

# --description--

`keys` 메서드는 `Map` 객체에서 직접 사용할 수 있으며, 맵의 키를 순회할 수 있는 `MapIterator`를 반환합니다.

`displayCatalog`라는 이름의 함수를 만드세요. 함수 안에서 `for...of` 반복문을 만들어 `catalog`의 키를 순회하며 각각을 콘솔에 출력하세요.

그다음 `displayCatalog`를 호출하고 콘솔에서 결과를 확인하세요.

# --hints--

`displayCatalog`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(displayCatalog);
```

`displayCatalog` 함수를 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /displayCatalog\s*\(\s*\)/);
```

`displayCatalog` 함수는 `keys`에서 `catalog` 메서드를 호출해야 합니다.

```js
const catalogCopy = new Map(catalog);
let flag = false;
const temp = catalog.keys;
try {
    catalog.keys = () => {
        flag = true;
        return catalogCopy.keys()
    };
    displayCatalog();    
    assert.isTrue(flag);
} finally {
    catalog.keys = temp;
}
```

`displayCatalog` 함수는 `catalog`에 포함된 각 키를 콘솔에 출력해야 합니다.

```js
const testArr = [];
const temp = console.log;
try {
    console.log = obj => {testArr.push(obj)};
    displayCatalog();    
    assert.isNotEmpty(testArr);
    const k = catalog.keys();
    for (let i = 0; i < catalog.size; i++) {
        assert.strictEqual(testArr[i], k.next().value)
    }
} finally {
    console.log = temp;
}
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const removePlant = plant => catalog.delete(plant);

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

--fcc-editable-region--

--fcc-editable-region--
```
