---
id: 6737a95d5b431860f96ca7bb
title: 27단계
challengeType: 1
dashedName: step-27
---

# --description--

이제 함수 안에서 `ballerina`에 `catalogSet` 객체를 두 번째로 추가해 보세요. 이미 set에 객체가 저장되어 있고 set에는 중복된 요소가 허용되지 않기 때문에 아무 일도 일어나지 않는 것을 볼 수 있습니다.

원한다면 이전에 map에서 했던 것처럼 `size` 속성을 사용해 set에 있는 요소 개수를 확인할 수 있습니다.

# --hints--

함수 `add` 안에서 `ballerina` 객체를 두 번째로 set에 `displayPlantsSet` 해야 합니다.

```js
const testArr = [];
const temp = Set.prototype.add;
try {
    Set.prototype.add = obj => {testArr.push(obj)};
    displayPlantsSet()
    assert.lengthOf(testArr, 2);
    assert.strictEqual(testArr[0], ballerina);
    assert.strictEqual(testArr[1], ballerina);
} finally {
    Set.prototype.add = temp;
}
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

const removePlant = plant => catalog.delete(plant);

const displayCatalog = () => {
    let catalogString = "";
    catalog.forEach((val, key) => {
        catalogString += `${key.scientificName} '${key.cultivar}': ${val.small} S, ${val.medium} M, ${val.large} L
`
    })
    return catalogString
}
--fcc-editable-region--
const displayPlantsSet = () => {
    const catalogSet = new Set();
    catalogSet.add(ballerina);

    return catalogSet
};

const plantsSet = displayPlantsSet();
console.log(plantsSet);
--fcc-editable-region--
```
