---
id: 6737a1a073a0d14c6544d301
title: 25단계
challengeType: 1
dashedName: step-25
---

# --description--

`plantsSet`라는 변수를 선언하고 `displayPlantsSet` 함수를 호출한 결과를 할당하세요.

그다음 `plantsSet`를 콘솔에 출력하세요. 출력 결과를 보려면 브라우저 콘솔을 열어야 합니다.


# --hints--

`plantsSet`라는 변수가 있어야 합니다.

```js
assert.exists(plantsSet);
```

`displayPlantsSet`를 호출하고 `plantsSet`에 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+plantsSet\s*=\s*displayPlantsSet\s*\(\s*\)/)
```

`plantsSet`를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*plantsSet\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

const removePlant = plant => catalog.delete(plant);

const displayCatalog = () => {
    let catalogString = "";
    catalog.forEach((val, key) => {
        catalogString += `${key.scientificName} '${key.cultivar}': ${val.small} S, ${val.medium} M, ${val.large} L
`
    })
    return catalogString
}
--fcc-editable-region--
const displayPlantsSet = () => new Set();

--fcc-editable-region--
```
