---
id: 5ddb965c65d27e1512d44d9a
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذا المشروع، ستتعلم إنشاء نموذج عداد السعرات الحرارية الذي يمكّن المستخدمين من إدخال ميزانيتهم اليومية من السعرات الحرارية وعدد السعرات الحرارية في وجبات مختلفة. بعد ذلك، يحسب النموذج ويعرض ما إذا كان المستخدم في عجز أو فائض من السعرات الحرارية.

تم تزويدك بقالب جاهز من CSS وHTML. مع ذلك، عليك بناء نموذج عداد السعرات الحرارية الخاص بك.

استكشف HTML وCSS بحرية، ثم أضف عنصر `form` وعيّن له خاصية `id` بقيمة `calorie-counter`.

# --hints--

يجب أن يحتوي المستند على عنصر `form`.

```js
assert.exists(document.querySelector('form'));
```

يجب أن يحتوي عنصر `form` على خاصية `id` بقيمة `calorie-counter`.

```js
assert.equal(document.querySelector('form')?.id, 'calorie-counter');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
