---
id: 63b608ebf40c871d960fc004
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

تحتاج إلى اثنين إضافيين من هذه كتل شفرة `fieldset` – واحدة لـ `snacks` وواحدة لـ `exercise`.

# --hints--

يجب إنشاء عنصر `fieldset` رابع في نموذجك `form`.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]);
```

يجب أن يأتي عنصر `fieldset` الرابع بعد عنصر `fieldset` الثالث.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.previousElementSibling?.tagName, "FIELDSET");
```

يجب أن يحتوي عنصر `fieldset` الرابع على معرف `id` مضبوط على `snacks`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.id, "snacks");
```

يجب أن يحتوي عنصر `fieldset` الرابع على عنصر `legend`.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('legend'));
```

يجب أن يحتوي عنصر `legend` الرابع على النص `Snacks`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('legend')?.innerText, "Snacks");
```

يجب أن يحتوي عنصر `fieldset` الرابع على عنصر `div`.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('div'));
```

يجب أن يحتوي عنصر `div` الرابع على `class` مضبوط على `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('div')?.className, "input-container");
```

يجب إنشاء عنصر `fieldset` خامس في نموذجك `form`.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]);
```

يجب أن يأتي عنصر `fieldset` الخامس بعد عنصر `fieldset` الرابع.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.previousElementSibling?.tagName, "FIELDSET");
```

يجب أن يحتوي عنصر `fieldset` الخامس على معرف `id` مضبوط على `exercise`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.id, "exercise");
```

يجب أن يحتوي عنصر `fieldset` الخامس على عنصر `legend`.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('legend'));
```

يجب أن يحتوي عنصر `legend` الخامس على النص `Exercise`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('legend')?.innerText, "Exercise");
```

يجب أن يحتوي عنصر `fieldset` الخامس على عنصر `div`.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('div'));
```

يجب أن يحتوي عنصر `div` الخامس على `class` مضبوط على `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
