---
id: 63b607af6fcdb119aae9b16a
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

في النموذج الخاص بك، يجب أن يتمكن المستخدمون من إضافة أنواع وجبات مختلفة مع عدد السعرات الحرارية الخاصة بها.

أنشئ عنصر `fieldset` مع تعيين `id` إلى `breakfast`.

داخل ذلك العنصر، أنشئ عنصر `legend` بالنص `Breakfast`، وعنصر `div` فارغ مع تعيين `class` إلى `input-container`.

# --hints--

يجب أن تنشئ عنصر `fieldset` في `form` الخاص بك.

```js
assert.exists(document.querySelector('form fieldset'));
```

يجب أن يأتي عنصر `fieldset` بعد عنصر `input` الخاص بك.

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

يجب أن يحتوي عنصر `fieldset` على `id` معين إلى `breakfast`.

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

يجب أن تنشئ عنصر `legend` داخل عنصر `fieldset` الخاص بك.

```js
assert.exists(document.querySelector('form fieldset legend'));
```

يجب أن يحتوي عنصر `legend` على النص `Breakfast`.

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

يجب أن تنشئ عنصر `div` داخل عنصر `fieldset` الخاص بك.

```js
assert.exists(document.querySelector('form fieldset div'));
```

يجب أن يأتي عنصر `div` بعد عنصر `legend` الخاص بك.

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

يجب أن يحتوي عنصر `div` على `class` معين إلى `input-container`.

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
