---
id: 5d822fd413a79914d39e98cb
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

بعد ذلك، أضف وسمَي `head` و`body` الافتتاحيين والختاميين داخل عنصر `html`.

# --hints--

يجب أن يكون لديك وسم `head` افتتاحي.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك وسم `head` ختامي.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك وسم `body` افتتاحي.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك وسم `body` ختامي.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون عنصرَا `head` و`body` أخوين.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
  
</html>
--fcc-editable-region--

```

