---
id: 5d822fd413a79914d39e98d2
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

قم بتسكين `div` بصنف `bb1` داخل حاوية المباني الخلفية. افتح ملف `styles.css`، وأعطِ `.bb1` عرضًا (`width`) بنسبة `10%` وارتفاعًا (`height`) بنسبة `70%`. "bb" تعني "background building"، وهذا سيكون المبنى الأول لديك.

# --hints--

أنشئ عنصر `div` جديد.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

أعطِ عنصر الـ `div` الجديد صنفًا `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

استخدم محدِّد صنف `.bb1` لتنسيق العنصر.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

يجب أن تُعطي العنصر `.bb1` قيمة `width` تساوي `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

يجب أن تُعطي العنصر `.bb1` قيمة `height` تساوي `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
    
```

