---
id: 5d822fd413a79914d39e98d4
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

أعطِ أجزاء المبنى خاصيتي `width` و `height` بالقيم التالية: `70%` و `10%` لـ `.bb1a`، و `80%` و `10%` لـ `.bb1b`، و `90%` و `10%` لـ `.bb1c`، و `100%` و `70%` لـ `.bb1d`. تذكّر أن هذه النسب مئوية بالنسبة للحاوية الأصلية، ولاحظ أن الارتفاعات ستجمع لتصل إلى 100% - مما يملأ الحاوية عموديًا.

# --hints--

يجب استخدام محدِّد صنف لتنسيق `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

يجب إعطاء `.bb1a` خاصية `width` بقيمة `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

يجب إعطاء `.bb1a` خاصية `height` بقيمة `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

يجب استخدام محدِّد صنف لتنسيق `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

يجب إعطاء `.bb1b` خاصية `width` بقيمة `80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

يجب إعطاء `.bb1b` خاصية `height` بقيمة `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

يجب استخدام محدِّد صنف لتنسيق `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

يجب إعطاء `.bb1c` خاصية `width` بقيمة `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

يجب إعطاء `.bb1c` خاصية `height` بقيمة `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

يجب استخدام محدِّد صنف لتنسيق `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

يجب إعطاء `.bb1d` خاصية `width` بقيمة `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

يجب إعطاء `.bb1d` خاصية `height` بقيمة `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

