---
id: 5d822fd413a79914d39e98cf
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

يمكنك رؤية `body` (إنه الصندوق الداخلي الأعمق في صفحتك)؛ الصندوق المحيط به هو عنصر `html`. اجعل `body` يملأ كامل إطار العرض بإعطائه `height` بقيمة `100vh`. أزل الـ `margin` الافتراضي من `body` بتعيين `margin` إلى `0`. أخيرًا، عيّن خاصية `overflow` إلى `hidden` لإخفاء أي أشرطة تمرير تظهر عندما يمتد شيء ما خارج إطار العرض.

# --hints--

يجب أن تستخدم محدِّد `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

يجب أن يكون لـ `body` ارتفاع `height` بقيمة `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

يجب أن يكون لـ `body` هامش `margin` بقيمة `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

يجب أن تكون خاصية `overflow` في `body` مضبوطة على `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```

