---
id: 5d822fd413a79914d39e98cc
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

داخل `head`، قم بتسكين عنصر `meta` مع خاصية `charset` بقيمة `UTF-8`، وعنصر `title` بعنوان `City Skyline`، وعنصر `link` يربط ملف `styles.css` الخاص بك.

# --hints--

يجب إنشاء عنصر `meta` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('meta'));
```

يجب إعطاء وسم `meta` خاصية `charset` بقيمة `UTF-8`.

```js
const linkElement = document.querySelector('meta');
const charset = linkElement?.getAttribute("charset").toLowerCase();
assert.strictEqual(charset, 'utf-8');
```

يجب أن يحتوي شفرتك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون عنصر `title` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

يجب أن يحتوي مشروعك على عنوان `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

تذكر، أن حالة الأحرف والإملاء مهمة للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

يجب أن يحتوي شفرتك على عنصر `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

لا تغير وسوم `head` الموجودة لديك. تأكد من عدم حذف وسم الإغلاق.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

يجب أن يكون لديك عنصر `link` فارغ واحد.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
