---
id: 5d822fd413a79914d39e98c9
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

مرحبًا بك في مشروع CSS Variables Skyline! ابدأ بإضافة الإعلان `!DOCTYPE html` في أعلى الوثيقة ليعرف المتصفح نوع الوثيقة التي يقرأها.

# --hints--

يجب أن تحتوي شفرتك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب أن تضع مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب أن تحدد نوع الوثيقة ليكون `html`.

```js
assert(code.match(/html/gi));
```

يجب أن تغلق إعلان `DOCTYPE` بعلامة `>` بعد النوع.

```js
assert(code.match(/html\s*>/gi));
```

يجب أن يكون إعلان `DOCTYPE` في بداية شفرة HTML الخاصة بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
