---
id: 5d822fd413a79914d39e98ca
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف وسمَي `html` الافتتاحي والختامي أسفل `DOCTYPE` لتجد مكانًا تبدأ فيه كتابة بعض الكود. تأكد من تعيين اللغة إلى الإنجليزية.

# --hints--

يجب أن يكون إعلان `DOCTYPE` في بداية ملف HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يحتوي عنصر `html` على وسم بداية، وخاصية `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم إغلاق.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن تكون وسوم `html` بالترتيب الصحيح.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

يجب أن يكون لديك عنصر `html` واحد فقط.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
