---
id: 5d822fd413a79914d39e98d1
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

امنح عنصر `.background-buildings` الخاص بك `width` و `height` بقيمة `100%` ليشغل العرض والارتفاع الكاملين لعنصر الأصل، وهو `body`.

# --hints--

يجب استخدام صنف `background-buildings` لاختيار العنصر الصحيح.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

يجب أن يكون لعنصر `.background-buildings` الخاص بك `width` بقيمة `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

يجب أن يكون لعنصر `.background-buildings` الخاص بك `height` بقيمة `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--
    
```

