---
id: 658212ba98182f3e855e85f9
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

في هذا المشروع، ستبني متتبع نفقات بسيط وعملي باستخدام Python.

ابدأ بإعلان دالة باسم `add_expense` تأخذ ثلاث معلمات: `expenses` و `amount` و `category`. استخدم الكلمة المفتاحية `pass` لملء جسم الدالة.

# --hints--

يجب أن تستخدم الكلمة المفتاحية `def` لإعلان دالة باسم `add_expense`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

يجب أن تأخذ دالة `add_expense` ثلاث معلمات: `expenses` و `amount` و `category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
