---
id: 65822773f4c60a479ec89f66
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

داخل الدالة `print_expenses`، أنشئ حلقة `for` تقوم بتكرار كل عنصر في قائمة `expenses`. استخدم `expense` كمتغير للحلقة وانقل `pass` داخل جسم الحلقة.

# --hints--

يجب أن تنشئ حلقة `for` تقوم بتكرار `expenses` داخل دالتك `print_expenses`. تذكر استخدام `pass` داخل جسم الحلقة.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

يجب أن تستخدم `expense` كمتغير للحلقة.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

يجب أن يكون `pass` داخل جسم الحلقة فقط.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
