---
id: 6582352cbdbcde4d0c4f7b0b
title: الخطوة 20
challengeType: 20
dashedName: step-20
---

# --description--

لاستدعاء دالة lambda يمكنك استخدام صيغة الدالة المعتادة مع زوج من الأقواس العادية بعد اسم المتغير.

استدعِ دالة lambda الخاصة بك `test` ومرر `3` كمعلمة. ثم اطبع النتيجة.

# --hints--

يجب أن تحتوي شفرتك على `print(test(3))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
