---
id: 65824111a09164518320088d
title: الخطوة 30
challengeType: 20
dashedName: step-30
---

# --description--

تسمح لك الدالة `filter()` باختيار عناصر من قابل للتكرار، مثل قائمة، بناءً على ناتج دالة:

```py
filter(my_function, my_list)
```

تأخذ `filter()` دالة كمعلمتها الأولى وقابل للتكرار كمعلمتها الثانية. تُرجع مُكرّرًا، وهو كائن خاص يمكّنك من التكرار على عناصر مجموعة، مثل القائمة.

نتيجة المثال أعلاه هي مُكرّر يحتوي على عناصر `my_list` التي تُرجع لها `my_function` القيمة `True`.

داخل الدالة `filter_expenses_by_category`، استدعِ `filter()` مع تمرير دالة `lambda` التي كتبتها في الخطوة السابقة كمعلمتها الأولى وقائمة `expenses` كمعلمتها الثانية.

# --hints--

يجب استدعاء `filter()` داخل الدالة `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

يجب تمرير `lambda expense: expense['category'] == category` كمعلمتها الأولى لاستدعاء `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

يجب تمرير `expenses` كمعلمتها الثانية لاستدعاء `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
