---
id: 65824561f3478e5371a33ae5
title: الخطوة 32
challengeType: 20
dashedName: step-32
---

# --description--

الخطوة التالية هي تعريف الدالة الرئيسية، التي ستكون نقطة الدخول لبرنامج تتبع النفقات التفاعلي.

عرّف دالة باسم `main` بدون معلمات. املأ جسم الدالة بقائمة `expenses` التي أنشأتها في بداية هذا المشروع. ستستخدم هذه القائمة لتخزين سجلات النفقات.

# --hints--

يجب أن تعرف دالة باسم `main()` بدون معلمات.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

يجب أن تنقل قائمة `expenses` داخل جسم الدالة `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
