---
id: 6582507654b3ed5712341382
title: الخطوة 38
challengeType: 20
dashedName: step-38
---

# --description--

ستستخدم عبارات شرطية لفحص اختيار المستخدم. إذا كان الاختيار `'1'`، فهذا يعني أن المستخدم يريد إضافة مصروف.

لا تزال داخل حلقة `while`، تحت المتغير `choice`، اكتب عبارة `if` لفحص ما إذا كان `choice` يساوي السلسلة النصية `'1'`. إذا كان صحيحًا، فسيكون ذلك نقطة البداية لإضافة مصروف جديد.

داخل جسم عبارة `if`، أعلن عن متغير `amount` وقم بتعيين استدعاء `input()` فارغ له.

# --hints--

يجب أن تنشئ عبارة `if` تتحقق مما إذا كان `choice` يساوي `'1'` داخل حلقة `while` الخاصة بك.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
cond.is_equivalent("choice == '1'") or cond.is_equivalent("'1' == choice")
`)) })
```

يجب أن تعلن عن متغير باسم `amount` داخل جسم عبارة `if` الخاصة بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("amount")
`)) })
```

يجب أن تعيّن استدعاء `input()` فارغًا إلى متغير `amount` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input()")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

--fcc-editable-region--
```
