---
id: 658246c28575d653d1f89b59
title: الخطوة 33
challengeType: 20
dashedName: step-33
---

# --description--

حلقة `while` هي نوع آخر من الحلقات التي تُنفذ جزءًا من الشفرة طالما أن شرطًا محددًا هو `True`. تنتهي الحلقة عندما يصبح الشرط `False`:

```py
while condition:
    <code>
```

أسفل قائمة `expenses`، أنشئ حلقة `while`. استخدم `True` كشرط، واطبع السلسلة النصية `'\nExpense Tracker'` داخل جسم الحلقة لعرض عنوان البرنامج.

# --hints--

يجب أن تنشئ حلقة `while` باستخدام `True` كشرط داخل دالة `main()` الخاصة بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

يجب أن تطبع `'\nExpense Tracker'` داخل حلقة `while` الخاصة بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
