---
id: 65822bd82d708c4895080c35
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

في بايثون، من المهم أن تعرف أن نفس نوع الاقتباس المستخدم لتعريف سلسلة نصية لا يمكن استخدامه داخلها. على سبيل المثال، السلسلة النصية `'I'm a string!'` غير صحيحة. لاستخدام الاقتباس المفرد داخل تلك السلسلة يجب عليك إما:

- الهروب من الاقتباس بإضافة شرطة مائلة عكسية قبله: `'I\'m a string!'`
- أو استخدام علامات اقتباس مزدوجة لتعريف السلسلة النصية: `"I'm a string!"` (الأفضل).

يمكنك الوصول إلى القيم في القاموس عبر مفاتيحه. تحتاج إلى استخدام ترميز الأقواس ووضع المفتاح بين الأقواس المربعة:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

أنت حاليًا تدرج قاموس `expense` في سلسلة f-string. عدل التعبير في سلسلة f-string للوصول إلى قيمة المفتاح `'amount'` والمفتاح `'category'` في قاموس `expense`.

# --hints--

يجب أن تمرر `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` إلى استدعاء `print()`. تذكر استخدام علامات اقتباس مزدوجة داخل سلسلة f-string المحاطة بعلامات اقتباس مفردة والعكس صحيح.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
