---
id: 65822934c8a0454803e3ef12
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

بعد ذلك، ستعرض تفاصيل كل مصروف.

داخل حلقة `for`، استبدل `pass` باستدعاء `print()` ومرر له سلسلة نصية بصيغة f التالية: `f'Amount: {expense}, Category: {expense}'`.

# --hints--

يجب أن تطبع `f'Amount: {expense}, Category: {expense}'` داخل حلقة `for` الخاصة بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

يجب ألا يكون هناك `pass` داخل جسم الحلقة.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
