---
id: 658254db6e12485a48130f57
title: الخطوة 40
challengeType: 20
dashedName: step-40
---

# --description--

يجب تحويل مبلغ المصروف قبل إجراء أي حساب. تأخذ الدالة `float()` سلسلة نصية أو عددًا صحيحًا كمعلمة وتُرجع عددًا عشريًا.

مرر `input('Enter amount: ')` إلى الدالة `float()`.

# --hints--

يجب أن تمرر `input('Enter amount: ')` إلى الدالة `float()` في تعيين `amount`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = float(input('Enter amount: '))")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
        
        if choice == '1':
            amount = input('Enter amount: ')
--fcc-editable-region--
```
