---
id: 658238f7604f154ea9a23e1e
title: الخطوة 23
challengeType: 20
dashedName: step-23
---

# --description--

تعيد الدالة `sum()` مجموع العناصر في القابل للتكرار الذي يُمرر إليها كمعلمتها. ستستخدم `sum()` مع دالتي `map()` و `lambda` للحصول على المبلغ الإجمالي للنفقات.

جرّب الآن تعديل استدعاء `print()` الحالي لديك، واستبدل استدعاء `list()` باستدعاء الدالة `sum()` مع تمرير استدعاء `map()` الحالي كمعلمة لها.

# --hints--

يجب أن تطبع `sum(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
