---
id: 65822308ded1ba4632f66aa3
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

القاموس هو نوع بيانات مدمج آخر في Python. القاموس هو مجموعة من البيانات على شكل أزواج المفتاح-القيمة. تُعرّف القواميس باستخدام الأقواس المتعرجة (`{}`) وتحتوي على أزواج المفتاح-القيمة مفصولة بفواصل. يتبع كل مفتاح نقطتان (`:`) ثم القيمة:

```py
{'amount': 50.0, 'category': 'Food'}
```

في المثال أعلاه، `'amount'` و `'category'` هما مفتاحان، و `50.0` و `'Food'` هما القيم المقابلة لهما.

أنشئ قاموسًا بمفتاح `'amount'` وقيمة المعلمة `amount` ومرر القاموس الجديد إلى استدعاء `.append()`.

# --hints--

يجب أن تمرر قاموسًا كمعلمة إلى طريقة `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

يجب أن تمرر `{'amount': amount}` كمعلمة إلى طريقة `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
