---
id: 6582575b8089f85b8b92d7c8
title: الخطوة 42
challengeType: 20
dashedName: step-42
---

# --description--

بمجرد حصولك على تفاصيل المصروف، يجب استدعاء الدالة `add_expense` لإضافة المصروف الجديد إلى قائمة `expenses`.

بعد الحصول على `amount` و `category` باستخدام `input()`، استدعِ الدالة `add_expense` مع تمرير ثلاث معلمات: `expenses`، `amount` و `category`.

- `expenses` هي القائمة الفارغة التي أنشئت في الدالة الرئيسية في وقت سابق من هذا المشروع.
- `amount` هو مقدار المصروف.
- `category` هي فئة المصروف.

# --hints--

يجب استدعاء الدالة `add_expense()` داخل جملة `if` الخاصة بك.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

يجب تمرير ثلاث معلمات إلى استدعاء `add_expense()` الخاص بك: `expenses`، `amount`، و `category`. الترتيب مهم.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

يجب أن يأتي استدعاء `add_expense()` بعد تعيين `amount` و `category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
