---
id: 65823634a308cb4d6bcc5fa6
title: الخطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

يمكن دمج دوال lambda بشكل مفيد مع الدالة `map()`، التي تنفذ دالة محددة لكل عنصر في مجموعة من الكائنات، مثل القائمة:

```py
map(lambda x: x * 2, [1, 2, 3])
```

تُمرر الدالة التي ستُنفذ كالمعلمة الأولى، ويُمرر القابل للتكرار كالمعلمة الثانية.

النتيجة في المثال أعلاه ستكون `[2, 4, 6]`، حيث تضاعف كل عنصر في القائمة الممررة إلى `map()` بفعل دالة lambda.

عدّل استدعاء `print()` ليطبع نتيجة استدعاء `map()` مع `test` كالمعلمة الأولى، و`[2, 3, 5, 8]` كالمعلمة الثانية. لن تتمكن بعد من رؤية ناتج مقروء.

# --hints--

يجب أن يحتوي شفرتك على `print(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
