---
id: 65823040bdb32949d4c1a96f
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

ستحتاج إلى دالة لحساب المبلغ الإجمالي للنفقات.

عرّف دالة باسم `total_expenses` تأخذ معلمة واحدة `expenses`. املأ جسم الدالة بعبارة `pass` في الوقت الحالي.

# --hints--

يجب أن تعرف دالة باسم `total_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

يجب أن تأخذ دالة `total_expenses` معلمة واحدة فقط وهي `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
