---
id: 6582324b5d37fe4b622bc3c0
title: الخطوة 19
challengeType: 20
dashedName: step-19
---

# --description--

الدوال المجهولة (Lambda) هي دوال قصيرة ومجهولة في Python، مثالية للمهام البسيطة التي تُنفذ مرة واحدة فقط. تُعرف باستخدام الكلمة المفتاحية `lambda`، وتستخدم الصيغة النحوية التالية:

```py
lambda x: expr
```

في المثال أعلاه، تمثل `x` معلمة تُستخدم في التعبير `expr`، وتعمل مثل أي معلمة في دالة تقليدية. التعبير `expr` هو التعبير الذي يُحلل ويُرجع عندما يتم استدعاء الدالة المجهولة.

أنشئ متغيرًا باسم `test` وقم بتعيين دالة مجهولة تأخذ معلمة `x` وتُرجع `x * 2`.

# --hints--

يجب إعلان متغير باسم `test`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

يجب تعيين `lambda x: x * 2` إلى متغير `test`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
