---
id: 65821fcc010c3245718f2a06
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

معلمة `expenses` في دالة `add_expense` ستكون قائمة نفقات. تريد أن تضيف عناصر في نهاية القائمة. لهذا ستستخدم طريقة القائمة `.append()`.

أضف استدعاءً لطريقة `.append()` على قائمة `expenses`. لا تمرر أي معلمات إلى `.append()` في الوقت الحالي.

# --hints--

يجب أن تحتوي دالتك على `expenses.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

يجب ألا تحتوي دالتك على `pass`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
