---
id: 65821c1c40ccac44adf38e6b
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

أنشئ قائمة فارغة باسم `expenses`. ستستخدمها لتخزين كل من مصاريفك.

# --hints--

يجب إعلان متغير جديد باسم `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

يجب تعيين قائمة فارغة إلى المتغير `expenses`.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
