---
id: 658225d560369446ece5518b
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

أضف زوج مفتاح-قيمة آخر إلى القاموس الذي تُلحقه بقائمة `expense`. استخدم السلسلة النصية `'category'` كمفتاح، والمعلمة `category` كقيمة.

# --hints--

يجب أن تضيف زوج مفتاح-قيمة ثاني إلى القاموس الذي تمرره كمعلمة إلى الطريقة `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

يجب أن تمرر `{'amount': amount, 'category': category}` كمعلمة إلى الطريقة `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
