---
id: 65823ff0d4b991510fade1a8
title: الخطوة 29
challengeType: 20
dashedName: step-29
---

# --description--

داخل الدالة `filter_expenses_by_category`، استبدل `pass` بدالة `lambda`. استخدم `expense` كمعلمة وقم بتقييم المقارنة بين قيمة المفتاح `'category'` في قاموس `expense` و`category` باستخدام معامل المساواة.

# --hints--

يجب أن تنشئ دالة `lambda` تستخدم معلمة باسم `expense` وتُقيّم التعبير `expense['category'] == category` داخل الدالة `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

يجب ألا يحتوي الدالة `filter_expenses_by_category` على `pass`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
