---
id: 65823cfc74aa564ffc460489
title: الخطوة 26
challengeType: 20
dashedName: step-26
---

# --description--

الآن، نفذ استدعاء `map()` مع تمرير دالة `lambda` كالمعلمة الأولى وقائمة `expenses` كالمعلمة الثانية.

# --hints--

يجب استدعاء الدالة `map()` داخل الدالة `total_expenses`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

يجب تمرير دالة `lambda` كالمعلمة الأولى لاستدعاء `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

يجب تمرير قائمة `expenses` كالمعلمة الثانية لاستدعاء `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
