---
id: 6582267ef39f5047411c1743
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

بعد ذلك، عرّف دالة باسم `print_expenses` تأخذ معلمة واحدة هي `expenses`. ستُستخدم هذه الدالة لاحقًا لعرض كل مصروف في قائمتك.

املأ جسم دالتك الجديدة بعبارة `pass`.

# --hints--

يجب أن تعرّف دالة جديدة باسم `print_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

يجب أن تأخذ دالة `print_expenses` معلمة واحدة فقط هي `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
