---
id: 66314f2add0a8b57da5d6d91
title: الخطوة 22
challengeType: 20
dashedName: step-22
---

# --description--

يجب أن ترى شيئًا مثل `<map object at 0xd273a8>` مطبوعًا في وحدة التحكم، وهو التمثيل النصي لكائن map الذي تُرجعُه الدالة `map()`.

للحصول على مخرجات قابلة للقراءة، تحتاج إلى تحويل كائن map إلى قائمة. افعل ذلك بتمرير استدعاء `map()` كمعلمة إلى الدالة `list()`.

# --hints--

يجب أن تمرر استدعاء `map()` كمعلمة إلى الدالة `list()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
