---
id: 66694fc4bba24f33ca01fa5b
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

يحتوي Python على عدد من طرق القوائم. مثل الطرق لإضافة أو إزالة عناصر من القائمة.

يمكنك إضافة عنصر إلى نهاية القائمة باستخدام الطريقة `append()`. على سبيل المثال:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

الآن `example_list` هي `[4, 5, 6, 7]`.

حاول استخدام الطريقة `append()` لإضافة `3` إلى `my_list`. ثم اطبع القائمة.

# --hints--

يجب أن يكون العنصر الأخير في `my_list` هو `3`.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

يجب أن يكون طول `my_list` هو 3.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

يجب ألا تغير إعلان القائمة.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

يجب أن تطبع القائمة.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

يجب أن تطبع القائمة بعد استخدام `append()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
