---
id: 6669539c1379793f9cb8917c
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

قوائم Python قابلة للتغيير، ما يعني إمكانية تعديل قيمة عناصر القائمة. يمكنك تغيير قيمة عنصر باستخدام ترميز الأقواس.

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

سيجعل هذا قيمة `example_list` تساوي `[4, 'oh', 6, 7]`.

غيّر العنصر الأول في `my_list` إلى `0`، ثم اطبع القائمة للتحقق من القيمة.

# --hints--

لا تغيّر إعلان القائمة.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

يجب أن يكون العنصر الأول في `my_list` هو 0.

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

يجب أن تغيّر العنصر الأول باستخدام `my_list[0] = 0`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

يجب أن تطبع القائمة.

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

يجب أن تطبع القائمة بعد تغيير العنصر الأول.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
