---
id: 666951d2e4b27e3a3c72ca5f
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

يمكنك أيضًا الوصول إلى عنصر واحد للحصول على قيمته.

القوائم تبدأ فهارسها من الصفر مثل السلاسل النصية. هذا يعني أن العنصر الأول في الفهرس `0`، والعنصر الثاني في الفهرس `1`، وهكذا.

للوصول إلى عنصر تستخدم ترميز الأقواس. على سبيل المثال، `example_list[1]` يعني الوصول إلى العنصر في الفهرس `1`، وهو العنصر الثاني في `example_list`.

اطبع العنصر الأول من `my_list`.

# --hints--

يجب أن تطبع العنصر الأول من القائمة باستخدام `print(my_list[0])`.

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
