---
id: 66694eaaeb7d253049fa7883
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

لإضافة عناصر إلى قائمة، يمكنك كتابتها بين الأقواس المربعة، مع فصل القيم بفواصل.

أضف الرقمين `1` و `2` داخل القائمة.

# --hints--

العنصر الأول في `my_list` يجب أن يكون `1`.

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

العنصر الثاني في `my_list` يجب أن يكون `2`.

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

يجب أن يحتوي `my_list` على عنصرين بداخله.

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

يجب تعيين `[1, 2]` مباشرة إلى `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
