---
id: 66694bfa435cef241f6f9a68
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

قبل البدء في المشروع، ستتعلم عن القوائم.

على عكس الأعداد والسلاسل النصية، <dfn>القائمة</dfn> هي نوع بيانات يعمل كحاوية لقيم أخرى:

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

تتميز القائمة بالأقواس المربعة حول كل القيم، وفاصلة بين القيم، مثل: `["A", "happy", "list"]`. إذا لم تحتوي القائمة على أي قيم، فهي قائمة فارغة: `[]`.

يمكن أن تحتوي القائمة على أنواع بيانات مختلفة: `[1, "Up", ["Down", "Twice"]]`. هذا يشمل جميع أنواع البيانات الممكنة. يمكنها أيضًا أن تحتوي على قائمة أخرى!

أنشئ متغيرًا اسمه `my_list` وقم بتعيين قائمة فارغة له.

# --hints--

يجب أن يكون لديك متغير اسمه `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

يجب أن تكون قيمة المتغير `my_list` قائمة فارغة.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
