---
id: 666ab82c874add3c95145213
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

يمكن استخدام الطريقة `pop()` لإزالة عنصر من القائمة. افتراضيًا، تزيل العنصر الأخير من القائمة. يمكنك تمرير فهرس كمعلمة للطريقة، وستزيل العنصر عند الفهرس المحدد.

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

في هذه الحالة، `fruits_list.pop(2)` تزيل العنصر عند الفهرس `2` من القائمة.

استخدم `pop()` لإزالة العنصر الأخير من `my_list`، ثم اطبع `my_list`.

# --hints--

يجب أن تستخدم `pop()`.

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

يجب أن تكون قيمة `my_list` هي `[0, 1, 2]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

يجب أن تطبع `my_list` بعد استخدام `pop()`.

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
