---
id: 666ab9013cd5563f2b39947a
title: خطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

الآن بعد أن استكشفت القوائم، احذف كل الشفرة التي كتبتها لمراجعتك لتتمكن من بدء المشروع.

# --hints--

يجب حذف كل شيء.

```js
({
    test: () => runPython(`
    assert not _Node(_code)
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)

my_list.pop()
print(my_list)

--fcc-editable-region--
```
