---
id: 666aae6a5d30a71f1fd7749f
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

يمكن للطريقة `insert` إضافة عنصر في أي موقع داخل القائمة. المعلمة الأولى هي الموقع الذي يجب إضافة العنصر فيه، والمعلمة الثانية هي العنصر المراد إضافته. على سبيل المثال، إليك كيفية إضافة عنصر جديد في الموقع الثالث من `example_list`:

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

باستخدام `insert()`، أضف `1` إلى `my_list` في الموقع الصحيح بحيث يكون العد تصاعديًا، ثم اطبع القائمة.

# --hints--

يجب إضافة `1` إلى `my_list` عند الفهرس `1` باستخدام `insert()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

يجب أن تحتوي `my_list` على 4 عناصر.

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

يجب أن تساوي `my_list` القيمة `[0, 1, 2, 3]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

يجب طباعة القائمة بعد إضافة العنصر الجديد.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
