---
id: 67eaa957114d373deb3a9149
title: Crear un Convertidor de Moneda
challengeType: 25
dashedName: build-a-currency-converter
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Tu componente `CurrencyConverter` debería renderizar un elemento `input` para aceptar el monto a convertir.
2. Tu elemento `input` debería aceptar números.
3. Tu componente `CurrencyConverter` debería renderizar dos elementos `select` para elegir la moneda a convertir **de** y **a**.
4. Tu elemento `select` debe incluir opciones para **al menos** `USD`, `EUR`, `GBP` y `JPY`. Puedes usar cualquier tasa de cambio, siempre que no haya una correspondencia uno a uno entre las monedas. Aquí tienes algunos ejemplos de asignaciones buenas y malas:

5. Tu componente `CurrencyConverter` debería memorizar el cálculo de los montos convertidos para la moneda **de** manera que un cambio en la opción **a** `select` no recalculará los montos convertidos.
6. Tu componente `CurrencyConverter` debería renderizar un elemento mostrando el monto convertido en el formato `XX.XX CCC`, donde `XX.XX` es el monto convertido y `CCC` es el código de la moneda.
7. El monto convertido debería redondearse a dos decimales.

# --hints--

Deberías exportar un componente `CurrencyConverter`.

```js
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const a = eval(script);

  assert.property(exports, "CurrencyConverter");
```

Deberías tener un elemento `input[type="number"]` para aceptar el monto a convertir.

```js
const inp = document.querySelectorAll('input[type="number"]');
assert.equal(inp.length, 1);
```

Deberías tener dos elementos `select`.

```js
const selects = document.querySelectorAll('select');
assert.equal(selects.length, 2);
```

Los elementos `select` deberían tener opciones para al menos `USD`, `EUR`, `GBP` y `JPY`.

```js
const selects = [...document.querySelectorAll('select')];

assert.equal(selects.length, 2);

for (const select of selects) {
  const options = [...select.options].map((o) => o.value);
  assert.includeMembers(options, ["USD", "EUR", "GBP", "JPY"]);
}
```

Cambiar el valor del primer elemento `select` debería causar que los montos convertidos sean recalculados.

```js
  function spyOn(
    obj,
    method
  ) {
    const original = obj[method];
    const calls = [];
  
    const fn = (cb, deps) => {
      const result = original(() => {calls.push(1); return cb();}, deps);
      return result;
    };
  
    obj[method] = fn;
  
    fn.calls = calls;
    return fn;
  }
  const abuseMemo = spyOn(React, "useMemo");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const a = eval(script);
  const s = await __helpers.prepTestComponent(exports.CurrencyConverter);
  const [first, _second] = s.querySelectorAll('select');
  assert.exists(first);
  await changeSelectValue(first);
  assert.equal(abuseMemo.calls.length, 2);
```

Cambiar el valor del segundo elemento `select` no debería causar recalculación de los montos convertidos.

```js
  function spyOn(
    obj,
    method
  ) {
    const original = obj[method];
    const calls = [];
  
    const fn = (cb, deps) => {
      const result = original(() => {calls.push(1); return cb();}, deps);
      return result;
    };
  
    obj[method] = fn;
  
    fn.calls = calls;
    return fn;
  }
  const abuseMemo = spyOn(React, "useMemo");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const a = eval(script);
  const s = await __helpers.prepTestComponent(exports.CurrencyConverter);
  const [_first, second] = s.querySelectorAll('select');
  assert.exists(second);
  await changeSelectValue(second);
  assert.equal(abuseMemo.calls.length, 1);
```

Cambiar el valor del primer elemento `select` debería mostrar la nueva cantidad convertida.

```js
  const s = await __helpers.prepTestComponent(window.index.CurrencyConverter);

  const nonFormContentBefore = getInnerTextExcept(s, "input,select");
  const resultBefore = nonFormContentBefore.match(/(\d+\.\d{2})\s*([A-Z]{3})/);

  const [first, _second] = s.querySelectorAll('select');
  assert.exists(first);
  await changeSelectValue(first);

  const nonFormContentAfter = getInnerTextExcept(s, "input,select");
  const resultAfter = nonFormContentAfter.match(/(\d+\.\d{2})\s*([A-Z]{3})/);

  try {
    assert.notEqual(nonFormContentBefore, nonFormContentAfter);
    assert.notEqual(resultBefore[1], resultAfter[1]);
    assert.strictEqual(resultBefore[2], resultAfter[2]);
  } catch (e) {
    console.error(e);
    throw e;
  }
```

Cambiar el valor del segundo elemento `select` debería mostrar la nueva cantidad convertida y la moneda.

```js
  const s = await __helpers.prepTestComponent(window.index.CurrencyConverter);

  const nonFormContentBefore = getInnerTextExcept(s, "input,select");
  const resultBefore = nonFormContentBefore.match(/(\d+\.\d{2})\s*([A-Z]{3})/);
  
  const [_first, second] = s.querySelectorAll('select');
  assert.exists(second);
  await changeSelectValue(second);

  const nonFormContentAfter = getInnerTextExcept(s, "input,select");
  const resultAfter = nonFormContentAfter.match(/(\d+\.\d{2})\s*([A-Z]{3})/);

  try {
    assert.notEqual(nonFormContentBefore, nonFormContentAfter);
    assert.notEqual(resultBefore[1], resultAfter[1]);
    assert.notEqual(resultBefore[2], resultAfter[2]);
    assert.strictEqual(resultAfter[2], second.value);
    } catch (e) {
      console.error(e);
      throw e;
    }
```

El monto convertido debería mostrarse en el formato `XX.XX CCC`, donde `XX.XX` es el monto convertido redondeado a dos decimales y `CCC` es el código de la moneda.

```js
  const s = await __helpers.prepTestComponent(window.index.CurrencyConverter);

  const inp = s.querySelector('input[type="number"]');
  assert.exists(inp);
  const [_first, second] = s.querySelectorAll('select');
  assert.exists(second);
  await changeSelectValue(second);
  await React.act(async () => {
    inp.value = 10;
    const ev2 = new Event("change", { bubbles: true, cancelable: false });
    inp[Object.keys(inp).find((k) => k.startsWith("__reactProps"))].onChange({...ev2, target: inp});
  });

  const nonFormContent = getInnerTextExcept(s, "input,select");

  try {
    const currencyCode = second.value;
    const reg = new RegExp(`\\d+\\.\\d{2} ${currencyCode}`);
    assert.match(nonFormContent, reg);
  } catch (e) {
    console.error(e);
    throw e;
  }
```

El monto convertido debería ser diferente del monto de entrada.

```js
  const s = await __helpers.prepTestComponent(window.index.CurrencyConverter);

  const inp = s.querySelector('input[type="number"]');
  assert.exists(inp);
  const [first, second] = s.querySelectorAll('select');
  assert.exists(first);
  assert.exists(second);
  for (let i = 0; i < first.options.length; i++) {
    for (let j = 0; j < second.options.length; j++) {
      if (first.options[i].value === second.options[j].value) {
        continue;
      }

      await React.act(async () => {
        first.value = first.options[i].value;
        const ev = new Event("change", { bubbles: true, cancelable: false });
        first[Object.keys(first).find((k) => k.startsWith("__reactProps"))].onChange({...ev, target: first});
        second.value = second.options[j].value;
        const ev2 = new Event("change", { bubbles: true, cancelable: false });
        second[Object.keys(second).find((k) => k.startsWith("__reactProps"))].onChange({...ev2, target: second});

        inp.value = 10;
        const ev3 = new Event("change", { bubbles: true, cancelable: false });
        inp[Object.keys(inp).find((k) => k.startsWith("__reactProps"))].onChange({...ev3, target: inp});
      });

      const nonFormContent = getInnerTextExcept(s, "input,select");
      const { amount } = nonFormContent.match(/(?<amount>\d+\.\d{2}) [A-Z]{3}/).groups;

      try {
        assert.notEqual(Number(amount), Number(inp.value));
      } catch (e) {
        console.error(e);
        throw e;
      }
    }
  }
```

# --before-all--

```js
async function delay(time) {
  return new Promise((resolve) => setTimeout(resolve, time));
}

function getInnerTextExcept(doc, removingSelector) {
  const body = doc.cloneNode(true);

  const squareElements = body.querySelectorAll(removingSelector);
  squareElements.forEach(element => {
    element.parentNode.removeChild(element)
  });

  return body.innerText;
}

const changeSelectValue = async (selectElement) => {
  await React.act(async () => {
    const notSelected = [...selectElement.options].find((o) => !o.selected);
    selectElement.value = notSelected.value;
    const ev = new Event("change", { bubbles: true, cancelable: false });
    selectElement[Object.keys(selectElement).find((k) => k.startsWith("__reactProps"))].onChange({...ev, target: selectElement});
  });
};

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Currency Converter</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { CurrencyConverter } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<CurrencyConverter />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useMemo } = React;

export function CurrencyConverter() {

}
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Currency Converter</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { CurrencyConverter } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<CurrencyConverter />);
    </script>
</body>

</html>
```

```css
body {
    background-color: #0a0a23;
    font-family: Lato, sans-serif;
}

main {
    background-color: #1b1b32;
    padding: 20px;
    border-radius: 8px;
    box-shadow: 0px 4px 6px rgba(0, 0, 0, 0.1);
    text-align: center;
    width: 400px;
    margin: 50px auto;
}

h1 {
    margin-bottom: 16px;
    color: #f5f6f7;
}

label {
    color: #f5f6f7;
}

input,
select {
    width: 100%;
    display: block;
    box-sizing: border-box;
    padding: 8px;
    margin-bottom: 12px;
    margin-top: 4px;
    border: 1px solid #99c9ff;
    border-radius: 4px;
}

.conversion-display {
    font-size: 16px;
    color: #f5f6f7;
    margin-bottom: 10px;
}

p {
    font-size: 18px;
    font-weight: bold;
    color: #acd157;
}
```

```jsx
export function CurrencyConverter() {
  const [amount, setAmount] = React.useState(1);
  const [startCurrency, setStartCurrency] = React.useState("USD");
  const [endCurrency, setEndCurrency] = React.useState("EUR");
  
  const exchangeRates = {
    USD: 1,
    EUR: 0.85,
    GBP: 0.75,
    JPY: 110
  }

  const convertedAmounts = React.useMemo(() => {
    const converted = {};
    Object.keys(exchangeRates).forEach((curr) => {
      converted[curr] = ((amount / exchangeRates[startCurrency]) * exchangeRates[curr]).toFixed(2);
    });
    return converted;
  }, [amount, startCurrency]);
  
  return (
    <main>
      <h1>Currency Converter</h1>
      <p className="conversion-display">{startCurrency} to {endCurrency} Conversion</p>
      <input
        type="number"
        value={amount}
        onChange={(e) => setAmount(Number(e.target.value))}
      />
      <label>
        Start Currency:
        <select value={startCurrency} onChange={(e) => setStartCurrency(e.target.value)}>
          {Object.keys(exchangeRates).map((curr) => (
            <option key={curr} value={curr}>
              {curr}
            </option>
          ))}
        </select>
      </label>
      <label>
        Target Currency:
        <select value={endCurrency} onChange={(e) => setEndCurrency(e.target.value)}>
          {Object.keys(exchangeRates).map((curr) => (
            <option key={curr} value={curr}>
              {curr}
            </option>
          ))}
        </select>
      </label>
      <p>Converted Amount: <span>{convertedAmounts[endCurrency]}</span> {endCurrency}</p>
    </main>
  );
}
```
