---
id: 66ba762af611169359d9d369
title: Створіть статтю в газеті
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Історія користувача:**

1. Встановіть основне значення властивості `font-size` HTML-документа на `24px`.
1. Ви повинні мати елемент з класом `newspaper`, який містить всі інші елементи.
1. Елемент `.newspaper` повинен мати властивість `font-size` зі значенням `16px` та шрифт `Open Sans` (запасний шрифт — `sans-serif`).
1. В межах елемента `.newspaper` має бути ще принаймні 7 елементів: елемент з класом `name` для назви газети, елемент з класом `date` для дати публікації, елемент з класом `headline` для заголовка, елемент з класом `sub-headline` для підзаголовка, елемент з класом `author` для автора та два абзаци (кожен з класом `text`). Всі елементи мають містити відповідну інформацію.
1. Елемент `.name` повинен мати властивість `font-size`, значення якої вдвічі більше за основне значення `font-size`, і використовувати шрифт `Times New Roman` (запасний шрифт — `serif`).
1. Елементи `.name` та `.author` мають використовувати CSS, щоб їхні символи були у верхньому регістрі.
1. Елемент `.headline` повинен мати властивість `font-size`, значення якої вдвічі більше за значення `font-size` батьківського елемента, і бути жирним.
1. Елемент `.sub-headline` повинен мати властивість `font-weight` зі значенням `100`, властивість `font-size`, значення якої більше в `1.5` разів за значення `font-size` батьківського елемента, і бути курсивом.
1. Елемент `.author` має використовувати CSS, щоб шрифт був жирним.
1. Елементи `.text` повинні мати властивість `text-indent` зі значенням `20px`.
1. Елементи `.text` повинні мати властивість `line-height`, значення якої вдвічі більше за значення `font-size` батьківського елемента.
1. Перша літера елементів `.text` має бути жирною та вдвічі більшою за значення властивості `font-size` батьківського елемента. Для цього використайте селектор `::first-letter`.

# --hints--

Ви повинні мати селектор `html`, який встановлює значення властивості `font-size` на `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

Ви повинні мати елемент з класом `newspaper`.

```js
assert.exists(document.querySelector('.newspaper'));
```

Ви повинні мати селектор `.newspaper`, який встановлює значення властивості `font-size` своїх елементів на `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

Ви повинні мати селектор `.newspaper`, який встановлює значення властивості `font-family` своїх елементів на `'Open Sans', sans-serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

В межах елемента `.newspaper` має бути елемент з класом `name`.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

Елемент `.name` має містити назву газети.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

В межах елемента `.newspaper` має бути елемент з класом `date`.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

Елемент `.date` має містити дату.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

В межах елемента `.newspaper` має бути елемент з класом `headline`.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

Елемент `.headline` має містити заголовок статті.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

В межах елемента `.newspaper` має бути елемент з класом `sub-headline`.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

Елемент `.sub-headline` має містити підзаголовок статті.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

В межах елемента `.newspaper` має бути елемент з класом `author`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

Елемент `.author` має містити автора статті.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

В межах елемента `.newspaper` має бути принаймні 2 абзаци (кожен з класом `text`).

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

Елементи `.text` мають містити текст статті.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

Ви повинні мати селектор `.name`, який встановлює значення властивості `font-size` своїх елементів на `2rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

Ви повинні мати селектор `.name`, який встановлює значення властивості `font-family` своїх елементів на `'Times New Roman', serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

Елемент `.name` повинен мати властивість `text-transform` зі значенням `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Елемент `.author` повинен мати властивість `text-transform` зі значенням `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Ви повинні мати селектор `.headline`, який встановлює значення властивості `font-size` своїх елементів на `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

Ви повинні мати селектор `.headline`, який встановлює значення властивості `font-weight` своїх елементів на `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

Ви повинні мати селектор `.sub-headline`, який встановлює значення властивості `font-weight` своїх елементів на `100`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

Ви повинні мати селектор `.sub-headline`, який встановлює значення властивості `font-size` своїх елементів на `1.5em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

Ви повинні мати селектор `.sub-headline`, який встановлює значення властивості `font-style` своїх елементів на `italic`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

Ви повинні мати селектор `.author`, який встановлює значення властивості `font-weight` своїх елементів на `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

Ви повинні мати селектор `.text`, який встановлює значення властивості `text-indent` своїх елементів на `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

Ви повинні мати селектор `.text`, який встановлює значення властивості `line-height` своїх елементів на `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

Ви повинні мати селектор `.text::first-letter`, який встановлює значення властивості `font-weight` своїх елементів на `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

Ви повинні мати селектор `.text::first-letter`, який встановлює значення властивості `font-size` своїх елементів на `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
