---
id: 683ec95f2f8781355556ff7c
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Nachdem Sie die Grenzen Ihrer Suche definiert haben, ist es Zeit, die Schleife zu erstellen, die den binären Suchalgorithmus ausführt.

Die binäre Suche funktioniert, indem sie den Suchbereich wiederholt eingrenzt. Dieser Prozess läuft so lange weiter, wie ein gültiger Bereich von Elementen zum Überprüfen vorhanden ist.

Sie können dies mit einer `while`-Schleife ausdrücken, die so lange läuft, wie Ihr `low`-Zeiger kleiner oder gleich Ihrem `high`-Zeiger ist.

Wenn `low` jemals größer als `high` wird, bedeutet das, dass der Suchbereich leer geworden ist und das `value` nicht in der Liste enthalten ist.

Erstellen Sie innerhalb der `binary_search`-Funktion eine `while`-Schleife mit einer Bedingung, die überprüft, ob `low` kleiner oder gleich high ist. Fügen Sie innerhalb der Schleife vorerst das Schlüsselwort `pass` als Platzhalter hinzu.

# --hints--

Sie sollten eine `while`-Schleife erstellen, die läuft, solange `low` kleiner oder gleich `high` ist.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
