---
id: 683ec95f2f8781355556ff83
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Sie sehen, dass der erste Aufruf `[3]` zurückgibt. Das liegt daran, dass `3` der mittlere Wert in `[1, 2, 3, 4, 5]` ist.

Andererseits gibt der zweite Aufruf eine leere Liste zurück. Das passiert, weil die Schleife nach der Überprüfung des mittleren Werts derzeit nicht die Logik enthält, um den Suchbereich durch Aktualisieren der Variablen `low` oder `high` einzugrenzen. Sie führt nur eine einzelne Mittelpunktsprüfung durch und beendet sich dann ohne weitere Suche.

Damit die binäre Suche ihren Suchbereich weiter eingrenzen kann, fügen Sie einen `elif`-Block hinzu, der überprüft, ob `value` größer als `value_at_middle` ist. Fügen Sie das Schlüsselwort `pass` vorerst innerhalb von `elif` ein.

# --hints--

Sie sollten einen `elif`-Block hinzufügen, wenn `value` größer als `value_at_middle` ist.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--          
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
