---
id: 683ec95f2f8781355556ff82
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Denken Sie daran, dass die Funktion eine Liste namens `search_list` und das `value` entgegennimmt, das Sie suchen.

Rufen Sie nun die Funktion mit `binary_search([1, 2, 3, 4, 5], 3)` und `binary_search([1, 2, 3, 4, 5, 9], 4)` auf und geben Sie die Aufrufe sofort aus.

# --hints--

Sie sollten die Funktion mit `binary_search([1, 2, 3, 4, 5], 3)` aufrufen und sie sofort ausgeben.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

Sie sollten die Funktion mit `binary_search([1, 2, 3, 4, 5, 9], 4)` aufrufen und sie sofort ausgeben.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        break
    return []
          
--fcc-editable-region--

--fcc-editable-region--
```
