---
id: 683ec39565dbc7e3375263b1
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In einer vorherigen Lektion haben Sie die binäre Suche kennengelernt und erfahren, wie sie sich von der linearen Suche unterscheidet.

In diesem Workshop setzen Sie das in die Praxis um, indem Sie einen Algorithmus für die binäre Suche erstellen.

Beginnen Sie damit, eine `binary_search`-Funktion mit `search_list`- und `value`-Parametern zu erstellen.

`search_list` ist die Liste der Elemente, in der die Funktion suchen wird, und `value` ist das Ziel, das die Funktion innerhalb der `search_list` zu finden versucht.

Fügen Sie vorerst das Schlüsselwort `pass` innerhalb der Funktion hinzu, damit der Test bestanden wird.

# --hints--

Sie sollten eine `binary_search`-Funktion erstellen.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

Ihre `binary_search`-Funktion sollte `search_list`- und `value`-Parameter haben.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
