---
id: 683ec95f2f8781355556ff81
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Sie müssen Dinge testen, damit Sie den Ablauf des Algorithmus in diesem Anfangsstadium verstehen.

Dazu müssen Sie zuerst aus der Schleife ausbrechen. Das liegt daran, dass die aktuelle Implementierung nur eine Iteration zulässt, sodass es bei Nichterfüllung der Bedingung zu einer Endlosschleife kommt.

Direkt nach der `if` Anweisung verwenden Sie das Schlüsselwort `break`, um aus der `while` Schleife auszubrechen. Danach geben Sie nach der `while` Schleife eine leere Liste zurück, um anzuzeigen, dass der Wert nicht gefunden wurde.

# --hints--

Sie sollten das Schlüsselwort `break` verwenden, um nach der `while` Anweisung aus der `if` Schleife auszubrechen.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

Sie sollten nach Ihrer `while` Schleife eine leere Liste zurückgeben.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```
