---
id: 683ec95f2f8781355556ff7b
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Bei der binären Suche gibt es den niedrigstmöglichen Index und den höchstmöglichen Index.

Der niedrigstmögliche Index stellt die linke Grenze im aktuellen Bereich der Werte dar, in dem gesucht wird, in diesem Fall `search_list`. Der höchstmögliche Index hingegen stellt die rechte Grenze der Werte dar, in denen gesucht wird.

Um diese beiden zu berücksichtigen, müssen Sie die gesamte Liste, die durchsucht wird, in Betracht ziehen. Definieren Sie also innerhalb der `binary_search`-Funktion eine `low`-Variable mit dem Wert `0` und eine `high`-Variable mit einem Wert, der den letzten Index in der zu durchsuchenden Liste berücksichtigt.

# --hints--

Sie sollten eine Variable namens `low` in Ihrer `binary_search`-Funktion haben.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

Ihre `low`-Variable sollte auf `0` gesetzt sein.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

Sie sollten eine Variable namens `high` in Ihrer `binary_search`-Funktion haben.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

Ihre `high`-Variable sollte auf `len(search_list) - 1` gesetzt sein.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []
    
--fcc-editable-region--
```
