---
id: 683ec95f2f8781355556ff7e
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Nachdem Sie den `mid`-Index berechnet haben, müssen Sie den Wert aus dem `search_list` an diesem Index tatsächlich abrufen. Dieser Wert ist derjenige, den Sie mit Ihrem `value`-Parameter vergleichen, der das Ziel Ihrer Suche ist.

Erstellen Sie eine `value_at_middle`-Variable und weisen Sie ihr das Element aus `search_list` zu, das sich am `mid`-Index befindet.

# --hints--

Sie sollten eine `value_at_middle`-Variable erstellen, die auf `search_list[mid]` gesetzt ist.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2

        
--fcc-editable-region--
```
