---
id: 683ec95f2f8781355556ff7a
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Der binäre Suchalgorithmus, den Sie erstellen, wird nicht nur `True` oder `False` zurückgeben, um anzuzeigen, ob der Wert gefunden wurde oder nicht. Er wird auch den Pfad zurückgeben, den er nimmt, um diesen Wert zu finden. Dafür haben Sie eine `path_to_target`-Variable, um den Pfad zu verfolgen.

Ersetzen Sie innerhalb Ihrer `binary_search`-Funktion das Schlüsselwort `pass` durch eine `path_to_target`-Variable und initialisieren Sie sie zunächst mit einer leeren Liste.

# --hints--

Sie sollten innerhalb Ihrer `path_to_target`-Funktion eine `binary_search`-Variable erstellen, die auf eine leere Liste gesetzt ist. Vergessen Sie nicht, das Schlüsselwort `pass` zu entfernen.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```
