---
id: 683ec95f2f8781355556ff7d
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Um mit der Schleife zu beginnen, müssen Sie den mittleren Index des aktuellen Suchbereichs finden.

Sie können den mittleren Index berechnen, indem Sie den Durchschnitt der Variablen `low` und `high` ermitteln. Sie müssen eine Ganzzahldivision verwenden, um den Durchschnitt nach dem Addieren der beiden Werte zu erhalten, sodass das Ergebnis immer eine nach unten gerundete ganze Zahl ist, da die Indizes ganze Zahlen sein müssen.

Entfernen Sie das Schlüsselwort `pass` aus Ihrer `while`-Schleife und erstellen Sie eine `mid`-Variable, die auf den Durchschnitt von `low` und `high` gesetzt ist.

# --hints--

Sie sollten das Schlüsselwort `pass` aus der while-Schleife entfernen und eine `mid`-Variable erstellen, die auf `(low + high) // 2` gesetzt ist.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
--fcc-editable-region--
        pass
--fcc-editable-region--
```
