---
id: 683ec95f2f8781355556ff84
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Wenn die Bedingung im `elif` wahr ist, entfernen Sie das Schlüsselwort `pass` und aktualisieren Sie den Wert der Variablen `low`, indem Sie `1` zur Variablen `mid` hinzufügen.

Dies erweitert die Suche auf die rechte Hälfte der aktuellen Suchbereiche in der Liste, denn wenn `value` größer als `value_at_middle` ist, bedeutet das, dass sich `value` in der rechten Hälfte des aktuellen Suchbereichs befinden muss.

# --hints--

Sie sollten das Schlüsselwort `pass` entfernen und die Variable `low` auf `mid + 1` aktualisieren.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid + 1")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        elif value > value_at_middle:
            pass
--fcc-editable-region--
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
```
