---
id: 683ec95f2f8781355556ff80
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Als Nächstes erstellen Sie eine Bedingung, die überprüft, ob das Ziel `value` in der Mitte ist.

Erstellen Sie innerhalb der `while`-Schleife eine `if`-Anweisung, die prüft, ob das Ziel `value` gleich `value_at_middle` ist. Wenn ja, geben Sie die Variable `path_to_target` zurück.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, die überprüft, ob `value` gleich `value_at_middle` ist.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

Sie sollten die Variable `path_to_target` innerhalb der `if`-Anweisung zurückgeben.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
--fcc-editable-region--
        
--fcc-editable-region--
```
