---
id: 684169b00a073c9b9f0a703c
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Ist Ihnen aufgefallen, dass der zweite Funktionsaufruf jetzt die Werte `3, 5, 4` in der Liste enthält? Das zeigt, dass die binäre Suche wie vorgesehen funktioniert.

So ist es passiert: Der Algorithmus hat zuerst `3` als Mitte der Anfangsliste geprüft. Da `4` größer ist als `3`, verlagerte sich die Suche auf die rechte Hälfte. Dann wurde `5` als neue Mitte untersucht. Weil `4` kleiner ist als `5`, ging die Suche nach links und identifizierte schließlich `4` als Mitte des letzten Bereichs.

Um die Funktion erneut zu testen, rufen Sie sie mit `[1, 3, 5, 9, 14, 22], 10` auf und geben Sie den Aufruf sofort aus. Dies ist eine Situation, in der der Wert nicht gefunden wird.

# --hints--

Sie sollten die Funktion mit `binary_search([1, 3, 5, 9, 14, 22], 10)` aufrufen und den Aufruf sofort ausgeben.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1
    
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
