---
id: 6998b99a8e195cd0c1f211df
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda Kifuatilia Wanyamapori rahisi kwa kutumia vitu vya JavaScript.

Hapa kuna mfano wa kitu:

```js
const animal = {}; // empty object
```

Hii inaunda kitu tupu.

Sasa unda kigezo kinachoitwa `tiger` na uweke thamani yake kuwa kitu tupu.

# --hints--

Unapaswa kuunda kigezo kinachoitwa `tiger`.

```js
assert.exists(tiger);
```

Kigezo cha `tiger` kinapaswa kupewa thamani ya kitu tupu.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
