---
id: 6999a0c1698b93d5d41156ed
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa ongeza sifa nyingine kwa kitu cha `tiger` kinachoitwa `isEndangered`.

Weka sifa ya `isEndangered` kuwa `true`.

# --hints--

`tiger` inapaswa kuwa na sifa ya `isEndangered`.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` inapaswa kuwa sawa na `true`.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
