---
id: 6999b07c34c2cde21ee6e1d8
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Katika hatua hii, utaunda kitendakazi kinachokagua kama kitu kina sifa fulani.

Hapa kuna mfano wa kutumia `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

Katika mfano huu, kitendakazi cha `hasColor` kinatumia `hasOwnProperty` kukagua kama sifa ya `color` ipo kwenye kitu.

Sasa unda kitendakazi kinachoitwa `hasHabitat`. Kitendakazi hicho kinapaswa kuchukua kigezo kimoja kinachoitwa `animal`.

Tumia njia ya `hasOwnProperty` kurudisha `true` ikiwa kitu cha `animal` kina sifa inayoitwa `"habitat"` na `false` vinginevyo.

Baada ya kuunda kitendakazi, tumia `console.log` kuita `hasHabitat(tiger)` na `hasHabitat(elephant)` ili uone matokeo yote mawili kwenye konsoli.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `hasHabitat`.

```js
assert.isFunction(hasHabitat);
```

Kitendakazi cha `hasHabitat` kinapaswa kuwa na kigezo kimoja, `animal`.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` inapaswa kutumia njia ya `hasOwnProperty` kukagua sifa ya `habitat`.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

Unapaswa kuandika `hasHabitat(tiger)` kwenye konsoli.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

Unapaswa kuandika `hasHabitat(elephant)` kwenye konsoli.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

Kuita `hasHabitat(tiger)` kunapaswa kurudisha `true` baada ya kuongeza makazi.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Kuita `hasHabitat(elephant)` kunapaswa kurudisha `false` ikiwa hakuna makazi.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` inapaswa kutumia vigezo vya kitendakazi na kufanya kazi na kitu chochote.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
