---
id: 69999ff7f7862fdfb885ea5d
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Sasa sahihisha kitu chako cha `tiger` kwa kuongeza sifa mpya iitwayo `age`.

Weka sifa ya `age` kuwa na thamani ya `5`.

# --hints--

`tiger` inapaswa kuwa na sifa ya `age`.

```js
assert.property(tiger, "age");
```

`tiger.age` inapaswa kuwa sawa na `5`.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
