---
id: 6999b0198d18d402c36cc5fe
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Katika hatua hii, utaunda kitendakazi kinachotoa sifa kutoka kwenye kitu.

Hapa kuna mfano wa kuondoa sifa kwa kutumia `delete`:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

Katika mfano huu, sifa ya `canFly` imeondolewa kutoka kwenye kitu cha `bird`.

Sasa unda kitendakazi kinachoitwa `removeEndangeredStatus`. Kitendakazi hicho kinapaswa kuchukua kigezo kimoja kinachoitwa `animal`.

Ndani ya kitendakazi, ondoa sifa ya `isEndangered` kutoka kwenye kitu cha `animal` kwa kutumia neno kuu la `delete`. Rudisha kitu cha `animal` kilichosasishwa.

Baada ya kuunda kitendakazi, tumia `console.log` kuita `removeEndangeredStatus(tiger)` ili uone kitu kilichosasishwa kwenye konsoli.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `removeEndangeredStatus`.

```js
assert.isFunction(removeEndangeredStatus);
```

Kitendakazi cha `removeEndangeredStatus` kinapaswa kuwa na kigezo kimoja, `animal`.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Kitendakazi cha `removeEndangeredStatus` kinapaswa kurudisha kitu cha `animal` kilichosasishwa.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` inapaswa kutumia neno kuu la `delete` kuondoa sifa ya `isEndangered`.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

Unapaswa kuandika `removeEndangeredStatus(tiger)` kwenye konsoli.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

Kuita `removeEndangeredStatus(tiger)` kunapaswa kuondoa sifa ya `isEndangered`.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` inapaswa kutumia vigezo vya kitendakazi na kufanya kazi na kitu chochote.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
