---
id: 6999adc72c02812a28c01e31
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Katika hatua hii, utaunda kitendakazi kinachosahihisha sifa ya kitu.

Hapa kuna mfano wa kusahihisha sifa ndani ya kitendakazi:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
};

console.log(changeAge(dog, 6)); // { age: 6 }
```

Katika mfano huu, sifa ya `age` imesasishwa kuwa thamani mpya.

Sasa unda kitendakazi kinachoitwa `updateAge`. Kitendakazi hicho kinapaswa kuchukua vigezo viwili: `animal` na `newAge`.

Ndani ya kitendakazi, sasisha sifa ya `age` ya kitu cha `animal` kuwa `newAge`. Rudisha kitu cha `animal` kilichosasishwa.

Baada ya kuunda kitendakazi, tumia `console.log` kuitisha `updateAge(elephant, 12)` ili uone kitu cha `elephant` kilichosasishwa kwenye konsoli.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `updateAge`.

```js
assert.isFunction(updateAge);
```

Kitendakazi cha `updateAge` kinapaswa kuwa na vigezo viwili: `animal` na `newAge`.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

Kitendakazi cha `updateAge` kinapaswa kurudisha kitu cha `animal` kilichosasishwa.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

Unapaswa kuandika `updateAge(elephant, 12)` kwenye konsoli.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

Kuitisha `updateAge(elephant, 12)` kunapaswa kusahihisha sifa ya age.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` inapaswa kutumia vigezo vya kitendakazi na kufanya kazi na kitu chochote.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
