---
id: 6999faf886634b72be7a819b
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Kwa ajili ya mfuatiliaji huyu wa wanyamapori, utasasisha kitu na sifa mpya.

Hapa kuna mfano wa kuunda kitu chenye sifa:

```js
const animal = {
  name: "Lion" // property: value
};
```

Sasa sasisha kitu cha `tiger` ili kiwe na sifa iitwayo `species`.

Weka thamani yake kuwa `"Tiger"`.

# --hints--

`tiger` inapaswa kuwa kitu.

```js
assert.isObject(tiger);
```

`tiger` inapaswa kuwa na sifa ya `species`.

```js
assert.property(tiger, "species");
```

`tiger.species` inapaswa kuwa sawa na `"Tiger"`.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
