---
id: 6999ac171928136168654774
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Katika hatua hii, utaunda kitendakazi cha kupata sifa ya `species` ya kitu.

Hapa kuna mfano mdogo wa kupata sifa kutoka kwa kitu:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

Katika mfano huu, kitendakazi kinapokea kitu kama kigezo na kurudisha sifa yake ya `species`.

Sasa unda kitendakazi kinachoitwa `getSpecies`.

Kitendakazi kinapaswa kupokea kigezo kimoja kinachoitwa `animal`.

Rudisha sifa ya `species` ya `animal` kwa kutumia uandishi wa nukta.

Baada ya kuunda kitendakazi, tumia `console.log` kuita `getSpecies(tiger)` ili uone matokeo kwenye konsoli.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `getSpecies`.

```js
assert.isFunction(getSpecies);
```

Kitendakazi cha `getSpecies` kinapaswa kuwa na kigezo kimoja, `animal`.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Unapaswa kuandika `getSpecies(tiger)` kwenye konsoli.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` inapaswa kupata sifa species kwa kutumia uandishi wa nukta.

```js
assert.match(code, /animal\.species/);
```

Kuita `getSpecies(tiger)` kunapaswa kurudisha `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` inapaswa kutumia kigezo cha kitendakazi na kufanya kazi na kitu chochote.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
