---
id: 699a00de9f564bb0effaa14d
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Katika hatua hii, utaunda kitendakazi cha kupata sifa ya `age` ya kitu.

Hapa kuna mfano:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
};

console.log(getAnimalAge(cat)); // 3
```

Katika mfano huu, kitendakazi kinapokea kitu kama kigezo na kurudisha sifa yake ya `age`.

Sasa unda kitendakazi kinachoitwa `getAge`.

Kitendakazi kinapaswa kupokea kigezo kimoja kinachoitwa `animal`.

Rudisha sifa ya `age` ya `animal` kwa kutumia uandishi wa nukta.

Baada ya kuunda kitendakazi, tumia `console.log` kuita `getAge(tiger)` ili uone matokeo kwenye konsoli.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `getAge`.

```js
assert.isFunction(getAge);
```

Kitendakazi cha `getAge` kinapaswa kuwa na kigezo kimoja, `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` inapaswa kupata sifa ya `age` kwa kutumia uandishi wa nukta.

```js
assert.match(code, /animal\.age/);
```

Unapaswa kuandika `getAge(tiger)` kwenye konsoli.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

Kuita `getAge(tiger)` kunapaswa kurudisha `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` inapaswa kutumia kigezo cha kitendakazi na kufanya kazi na kitu chochote.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
