---
id: 69ba8da6b296391d913096a0
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Sasa tengeneza kitu cha pili kinachoitwa `elephant`.

Ongeza sifa zifuatazo:

- `species` yenye thamani `"Elephant"`
- `age` yenye thamani `10`
- `isEndangered` yenye thamani `true`

# --hints--

Unapaswa kutengeneza kigezo kinachoitwa `elephant`.

```js
assert.exists(elephant);
```

`elephant` inapaswa kuwa na sifa ya `species` yenye thamani sawa na `"Elephant"`.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` inapaswa kuwa na sifa ya `age` yenye thamani sawa na `10`.

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` inapaswa kuwa na sifa ya `isEndangered` yenye thamani sawa na `true`.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
