---
id: 61fd70336ebb3e4f62ee81ba
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML tables use the `caption` element to describe what the table is about. The `caption` element should always be the first child of a `table`, but can be positioned with the `caption-side` CSS property.

1 つ目の `table` に `caption` 要素を追加して、`Assets` というテキストを設定してください。

# --hints--

1 つ目の `table` 要素内に `caption` 要素が 1 つ必要です。

```js
assert(document.querySelector('table')?.children?.[0]?.localName === 'caption');
```

`caption` 要素のテキストは `Assets` にする必要があります。

```js
assert(document.querySelector('caption')?.textContent === 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
