---
id: 587d8259367417b2b2512c86
title: Implementar o Insertion Sort
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

Neste laboratório, você vai implementar o algoritmo de ordenação por inserção. Esse método funciona construindo um array ordenado no início da lista. Ele começa o array ordenado com o primeiro elemento. Depois, inspeciona o próximo elemento e o troca para trás dentro do array ordenado até que ele esteja na posição correta. Continua iterando pela lista e trocando novos itens para trás na parte ordenada até chegar ao fim. Esse algoritmo tem complexidade de tempo quadrática nos casos médio e pior.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma `insertionSort` função.
1. A `insertionSort` função deve receber um array de inteiros e retornar um array com os mesmos inteiros ordenados do menor para o maior.

# --hints--

`insertionSort` deve ser uma função.

```js
assert.isFunction(insertionSort);
```

`insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` deve retornar um array que não foi alterado exceto pela ordem.

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort` deve retornar um array ordenado (do menor para o maior elemento).

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`insertionSort([5, 4, 33, 2, 8])` deve retornar `[2, 4, 5, 8, 33]`.

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` não deve usar o método `.sort()` integrado.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
