---
id: 686ccc2c8b967e17ab18d593
title: 實作二分法
challengeType: 27
dashedName: implement-the-bisection-method
---

# --description--

二分法，也稱為二分搜尋法，使用二分搜尋來尋找實值函式的根。它透過縮小平方根所在的區間，直到收斂到指定容差範圍內的值。

例如，如果容差是 `0.01`，二分法會持續將區間對半分割，直到上下界之間的差異小於或相等於 `0.01`。

在本實驗中，你將實作一個使用二分法來尋找數字平方根的函式。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該定義一個名為 `square_root_bisection` 的函式（程式），並帶有三個參數：
    - 你想要尋找平方根的數字。
    - 容差是結果可接受的誤差範圍。你應該設定一個預設容差值。
    - 要執行的最大迭代次數。你應該設定一個預設的迭代次數。

1. `square_root_bisection` 函式（程式）應該：
    - 如果傳遞給函式的數字為負數，則引發帶有訊息 `Square root of negative number is not defined in real numbers` 的 `ValueError`。
    - 對於數字 `0` 和 `1`，列印訊息：`The square root of [number] is [number]` 並傳回該數字本身作為平方根。
    - 對於任何其他正數，列印帶有訊息的近似平方根：`The square root of [square_target] is approximately [root]`，並傳回計算出的根值。
    - 如果沒有值符合容差條件，列印失敗訊息：`Failed to converge within [maximum] iterations` 並傳回 `None`。

**注意**：你不能為此實驗匯入任何模組。

# --hints--

你不應該匯入任何模組。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_imports()) == 0`)) })
```

你應該有一個名為 `square_root_bisection` 的函式（程式）。

```js
({ test: () => assert(runPython(`_Node(_code).has_function("square_root_bisection")`)) })
```

你的 `square_root_bisection` 函式（程式）應該有三個參數。

```js
({ test: () => runPython(`
    import inspect 
    sig = inspect.signature(square_root_bisection)
    assert len(sig.parameters) == 3
`) })
```

你應該為容差和最大迭代次數設定預設值。

```js
({ test: () => runPython(`
try:
  import inspect 
  sig = inspect.signature(square_root_bisection)
  assert len(sig.parameters) == 3
  square_root_bisection(4)
except TypeError:
  assert False
`) })
```

當傳遞給函式的數字為負數時，你的 `square_root_bisection` 函式應該引發帶有訊息 `Square root of negative number is not defined in real numbers` 的 `ValueError`。

```js
({ test: () => runPython(`
try:
  square_root_bisection(-6)
except ValueError as e:
  assert str(e) == "Square root of negative number is not defined in real numbers"
else:
  assert False
`) })
```

`square_root_bisection(0)` 應該傳回 `0`。

```js
({ test: () => runPython(`assert square_root_bisection(0) == 0`) })
```

`square_root_bisection(0)` 應該列印 `The square root of 0 is 0`。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(0)
assert "The square root of 0 is 0" in _out
`) })
```

`square_root_bisection(0.001, 1e-7, 50)` 應該傳回介於 `0.03162267660168379` 和 `0.031622876601683794` 之間的數字。

```js
({ test: () => runPython(`assert 0.03162267660168379 <= square_root_bisection(0.001, 1e-7, 50) <= 0.031622876601683794`) })
```

`square_root_bisection(0.001, 1e-7, 50)` 應該列印 `The square root of 0.001 is approximately X`，其中 `X` 是介於 `0.03162267660168379` 和 `0.031622876601683794` 之間的數字。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.001, 1e-7, 50)

assert 0.03162267660168379 <= _root <= 0.031622876601683794
assert f"The square root of 0.001 is approximately {_root}" in _out
`) })
```

`square_root_bisection(0.25, 1e-7, 50)` 應該傳回介於 `0.4999999` 和 `0.5000001` 之間的數字。

```js
({ test: () => runPython(`assert 0.4999999 <= square_root_bisection(0.25, 1e-7, 50) <= 0.5000001`) })
```

`square_root_bisection(0.25, 1e-7, 50)` 應該列印 `The square root of 0.25 is approximately X`，其中 `X` 是介於 `0.4999999` 和 `0.5000001` 之間的數字。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.25, 1e-7, 50)

assert 0.4999999 <= _root <= 0.5000001
assert f"The square root of 0.25 is approximately {_root}" in _out
`) })
```

`square_root_bisection(1)` 應該傳回 `1`。

```js
({ test: () => runPython(`assert square_root_bisection(1) == 1`) })
```

`square_root_bisection(1)` 應該列印 `The square root of 1 is 1`。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(1)
assert "The square root of 1 is 1" in _out
`) })
```

`square_root_bisection(81, 1e-3, 50)` 應該傳回介於 `8.999` 和 `9.001` 之間的數字。

```js
({ test: () => runPython(`assert 8.999 <= square_root_bisection(81, 1e-3, 50) <= 9.001`) })
```

`square_root_bisection(81, 1e-3, 50)` 應該列印 `The square root of 81 is approximately X`，其中 `X` 是介於 `8.999` 和 `9.001` 之間的數字。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(81, 1e-3, 50)

assert 8.999 <= _root <= 9.001
assert f"The square root of 81 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-3, 100)` 應該傳回介於 `14.999` 和 `15.001` 之間的數字。

```js
({ test: () => runPython(`assert 14.999 <= square_root_bisection(225, 1e-3, 100) <= 15.001`) })
```

`square_root_bisection(225, 1e-3, 100)` 應該列印 `The square root of 225 is approximately X`，其中 `X` 是介於 `14.999` 和 `15.001` 之間的數字。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-3, 100)

assert 14.999 <= _root <= 15.001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-5, 100)` 應該傳回介於 `14.99999` 和 `15.00001` 之間的數字。

```js
({ test: () => runPython(`assert 14.99999 <= square_root_bisection(225, 1e-5, 100) <= 15.00001`) })
```

`square_root_bisection(225, 1e-5, 100)` 應該列印 `The square root of 225 is approximately X`，其中 `X` 是介於 `14.99999` 和 `15.00001` 之間的數字。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-5, 100)

assert 14.99999 <= _root <= 15.00001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 100)` 應該傳回介於 `14.9999999` 和 `15.0000001` 之間的數字。

```js
({ test: () => runPython(`assert 14.9999999 <= square_root_bisection(225, 1e-7, 100) <= 15.0000001`) })
```

`square_root_bisection(225, 1e-7, 100)` 應該列印 `The square root of 225 is approximately X`，其中 `X` 是介於 `14.9999999` 和 `15.0000001` 之間的數字。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-7, 100)

assert 14.9999999 <= _root <= 15.0000001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 10)` 應該傳回 `None`。

```js
({ test: () => runPython(`assert square_root_bisection(225, 1e-7, 10) is None`) })
```

`square_root_bisection(225, 1e-7, 10)` 應該列印 `Failed to converge within 10 iterations`。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(225, 1e-7, 10)
assert "Failed to converge within 10 iterations" in _out
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = square_target if square_target < 1 else 1
        high = 1 if square_target < 1 else square_target
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if high - low <= tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations")

        else:   
            print(f'The square root of {square_target} is approximately {root}')

    return root
```
