---
id: 68f890b73bbfc9efa7379e7b
title: Debuggen von Typumwandlungsfehlern in einer fehlerhaften App
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Sie sind gerade bei einer lokalen Web-Entwicklungsfirma eingestiegen, und Ihre erste Aufgabe ist es, fehlerhaften Code zu bereinigen, der von den vorherigen Entwicklern hinterlassen wurde.

Sie versuchten einfache arithmetische Operationen durchzuführen, aber etwas ist schiefgelaufen, die Ergebnisse ergeben keinen Sinn. Ihre Aufgabe ist es, den Code zu überprüfen, zu debuggen und zu korrigieren, damit er wie erwartet funktioniert und leichter zu lesen ist.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Variable namens `firstResult` haben, die die Zahlen `5` und `10` korrekt addiert, um den Wert `15` zu erzeugen.
2. Sie sollten eine Variable namens `secondResult` haben, die `5` korrekt von `8` subtrahiert, um den Wert `3` zu erzeugen.
3. Sie sollten eine Variable namens `thirdResult` haben, die den Wert `6` durch Addition von zwei Zahlen erzeugt. Ersetzen Sie den aktuell im Ausdruck verwendeten Boolean durch eine Zahl.
4. Sie sollten eine Variable namens `fourthResult` haben, die den Wert `8` durch Addition von zwei Zahlen erzeugt. Ersetzen Sie den aktuell im Ausdruck verwendeten Boolean durch eine Zahl.
5. Sie sollten eine Variable namens `fifthResult` haben, die zwei Zahlen korrekt multipliziert, um das Produkt `20` zu erhalten. Ersetzen Sie den aktuell im Ausdruck verwendeten String durch eine Zahl.
6. Sie sollten eine Variable namens `sixthResult` haben, die zwei Zahlen korrekt addiert, um den Wert `22` zu erzeugen. Ersetzen Sie den aktuell im Ausdruck verwendeten `null` durch eine Zahl.

# --hints--

Sie sollten eine Variable namens `firstResult` haben.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` sollte den numerischen Wert `15` haben.

```js
assert.strictEqual(firstResult, 15);
```

Sie sollten eine Variable namens `secondResult` haben.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` sollte den numerischen Wert `3` haben.

```js
assert.strictEqual(secondResult, 3);
```

Sie sollten eine Variable namens `thirdResult` haben.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` sollte den numerischen Wert `6` haben.

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` sollte keinen Boolean-Wert (`true` oder `false`) enthalten.

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Sie sollten eine Variable namens `fourthResult` haben.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` sollte den numerischen Wert `8` haben.

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` sollte keinen Boolean-Wert (`true` oder `false`) enthalten.

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Sie sollten eine Variable namens `fifthResult` haben.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` sollte den numerischen Wert `20` haben.

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` sollte keine String-Multiplikation enthalten (z. B. "10" * 2).

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

Sie sollten eine Variable namens `sixthResult` haben.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` sollte den numerischen Wert `22` haben.

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` sollte `null` nicht in seinem Ausdruck enthalten.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
