---
id: 67116d7584d0b469b14579bf
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

`library` 配列の中に、次のプロパティと値を持つオブジェクトを作成してください。

| プロパティ | 値   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

`library` 配列の中にオブジェクトがあるはずです。

```js
assert.isObject(library[0])
```

オブジェクトには `title` プロパティが必要です。

```js
assert.property(library[0], "title")
```

`title` プロパティの値は `"Your Next Five Moves: Master the Art of Business Strategy"` であるべきです。

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

オブジェクトには `author` プロパティが必要です。

```js
assert.property(library[0], "author")
```

`author` プロパティの値は `"Patrick Bet-David and Greg Dinkin"` であるべきです。

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

オブジェクトには `about` プロパティが必要です。

```js
assert.property(library[0], "about")
```

`about` プロパティの値は `"A book on how to plan ahead"` であるべきです。

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

オブジェクトには `pages` プロパティが必要です。

```js
assert.property(library[0], "pages")
```

`pages` プロパティの値は数値であるべきです。

```js
assert.isNumber(library[0]?.pages)
```

`pages` プロパティは `320` に設定してください。

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
