---
id: 67116d7584d0b469b14579c0
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

`library` 配列の中に、次のプロパティと値を持つ別のオブジェクトを作成してください。

| プロパティ | 値   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


`library` 配列の2番目の要素としてオブジェクトがあるはずです。

```js
assert.isObject(library[1]);
```

オブジェクトには `title` プロパティが必要です。

```js
assert.property(library[1], "title")
```

`title` プロパティの値は `"Atomic Habits"` であるべきです。

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

オブジェクトには `author` プロパティが必要です。

```js
assert.property(library[1], "author")
```

`author` プロパティの値は `"James Clear"` であるべきです。

```js
assert.propertyVal(library[1], "author", "James Clear");
```

オブジェクトには `about` プロパティが必要です。

```js
assert.property(library[1], "about")
```

`about` プロパティの値は `"A practical book about discarding bad habits and building good ones"` であるべきです。

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

オブジェクトには `pages` プロパティが必要です。

```js
assert.property(library[1], "pages")
```

`pages` プロパティの値は数値であるべきです。

```js
assert.isNumber(library[1]?.pages)
```

`pages` プロパティは `320` に設定してください。

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
