---
id: 67116c175a77cb64c5c10d49
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

前のレッスンでは、`map`、`filter`、`reduce` のような高階関数の使い方を学びました。

このワークショップでは、さまざまな高階関数を使う練習ができる図書館管理アプリを作成します。

まずは、`library` という変数を作成し、空の配列を代入してください。

# --hints--

`library` 配列を作成してください。

```js
assert.isArray(library)
```

`library` 配列は空である必要があります。

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
