---
id: 681dbe48ee5df6842385d735
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

本を表す残りのオブジェクトはすでに用意されています。

ここでコンソールに本の情報を表示する作業を始めるのに良いタイミングです。次のいくつかのステップで、`map()` メソッドを使ってこの目標を達成する方法を学びます。

まず、メッセージ `"Books in the Library:\n"` をコンソールにログ出力してください。ここで改行文字が追加されているのは、このメッセージと本のリストの間に改行を入れるためです。

# --hints--

メッセージ `"Books in the Library:\n"` をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(("|')Books\s+in\s+the\s+Library:\\n\1\);?/);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

--fcc-editable-region--

--fcc-editable-region--
```
