---
id: 681dc623b2b18887266777b1
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

これまでのレッスンでは、元の配列の各要素に指定した関数を適用して新しい配列を作成する`map()`メソッドについて学びました。

例を示します：

```js
const developers = [
  { name: "Alice", city: "New York", title: "Front-End Developer" },
  { name: "Bob", city: "San Francisco", title: "Back-End Developer" }
];

console.log(developers.map(dev => dev.name));
// ["Alice", "Bob"]
```

この例の`dev`は`developers`配列内の各オブジェクトを表します。次に、ドット記法を使ってオブジェクトから`name`を取得します。最後に、結果は名前の新しい配列になります。

`getBookInformation`関数の中で、`map()`パラメータに対して`catalog`メソッドを使い、本のタイトルだけの新しい配列を返してください。助けが必要な場合は例を参照してください。

# --hints--

`getBookInformation`関数は配列を返す必要があります。

```js
assert.isArray(getBookInformation(library));
```

`getBookInformation`関数は本のタイトルを表す文字列の配列を返す必要があります。

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = ["Title A", "Title B", "Title C"];

assert.deepEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

--fcc-editable-region--
function getBookInformation(catalog) {
  
}
--fcc-editable-region--
```
