---
id: a5deed1811a43193f9f1c841
title: Implementieren Sie einen Element-Skipper
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die Elemente in einem Array überspringt, bis sie ein akzeptables Element basierend auf einer bestimmten Testfunktion findet.

Zum Beispiel, für ein Array wie `[1, 1, 1, 2, 1, 1, 1]` und eine Testfunktion `function(n) {return n === 2}`, ist das erste akzeptable Element dasjenige am Index 3. Daher müssen alle Elemente davor verworfen werden, und die Ausgabe sollte die verbleibenden Elemente `[2, 1, 1, 1]` sein.

**Ziel**: Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `dropElements`-Funktion haben, die ein Array (`arr`) und eine Funktion (`func`) als Argumente akzeptiert.
1. Die `dropElements`-Funktion sollte das Array durchlaufen und Elemente beginnend mit dem ersten entfernen, bis `func` für ein Element `true` zurückgibt.
1. Die `dropElements`-Funktion sollte die verbleibenden Elemente im Array zurückgeben, wenn die Bedingung erfüllt ist.
1. Wenn die Bedingung nie erfüllt wird, sollte sie ein leeres Array zurückgeben.

# --hints--

Sie sollten eine `dropElements`-Funktion haben.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` sollte `[3, 4]` zurückgeben.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` sollte `[1, 0, 1]` zurückgeben.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` sollte `[1, 2, 3]` zurückgeben.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` sollte `[]` zurückgeben.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` sollte `[7, 4]` zurückgeben.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` sollte `[3, 9, 2]` zurückgeben.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
