---
id: 680a7e200400089967b8b1f4
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

`str.maketrans()` 方法接受兩個長度相等的字串，並傳回一個翻譯表，該表將第一個字串的每個字元對應到第二個字串的相應字元。翻譯表中的每個字元皆以 Unicode 序數儲存，該序數是一個唯一識別該字元的數字。

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

你將在課程後續學習更多關於 `str.maketrans()` 傳回的結構型別。現在，為你的 `alphabet` 中的字元對應到 `shifted_alphabet` 中的字元創建一個翻譯表，並將它指定給名為 `translation_table` 的變數。

# --hints--

你應該宣告一個名為 `translation_table` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

你應該呼叫 `str.maketrans()`，並將 `alphabet` 和 `shifted_alphabet` 作為引數傳入，然後將結果指定給 `translation_table`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
