---
id: 680a5e7f3d47715196b12ca0
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

現在你已經有整個位移字母表，移除 `print(shifted_alphabet)` 呼叫。

# --hints--

你應該從你的程式碼中移除 `print(shifted_alphabet)`。

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
