---
id: 680a431692a5f863f726c438
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

在這個工作坊中，你將建置一個 *Caesar cipher*。這是加密文字最簡單的技術之一，其方法是將明文中的每個字母替換成字母表中固定位置數量之後的字母。舉例來說，當位移為 5 時，`a` 會被替換成 `f`，`b` 會被替換成 `g`，以此類推。

要實作這個密碼，你需要為你的字母表創建一個從 `shift` 指示的位置開始的新版本。正如你在之前的課程中學到的，你可以使用字串切片來擷取字串的零件：

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

創建一個名為 `shifted_alphabet` 的變數，並使用切片語法將 `alphabet` 中從 `shift` 索引開始的部分指定給它。然後，呼叫內建的 `print()` 函式（程式）在終端機上列印 `shifted_alphabet`，並查看結果。

# --hints--

你應該宣告一個名為 `shifted_alphabet` 的新變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

你應該將 `alphabet[shift:]` 指定給 `shifted_alphabet`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

你應該呼叫 `print()` 函式，並傳入 `shifted_alphabet` 作為引數。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
