---
id: 680a3d4c065557260de2ff33
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將創建一個名為凱薩密碼的程式，這是一種透過將字母在字母表中位移來編碼訊息的加密方法。

首先創建一個名為 `shift` 的變數，並將值 `5` 指定給你的新變數。

# --hints--

你應該宣告一個名為 `shift` 的變數。請注意將變數名稱放在行首。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

你應該將值 `5` 指定到你的 `shift` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
