---
id: 680a41ef5bfdc059cfd9d4b5
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

在之前的課程中，你學習了可以儲存在變數中的不同資料型別。你剛剛指定了一個整數值。現在你需要指定一個字串，字串是一連串由單引號或雙引號包圍的字元：

```py
string_1 = 'I am a string'
string_2 = "I am also a string"
```

宣告另一個名為 `alphabet` 的變數，並將字串 `abcdefghijklmnopqrstuvwxyz` 指定給此變數。

# --hints--

你應該宣告一個名為 `alphabet` 的變數。請注意將變數名稱放在行首。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

你應該將字串 `abcdefghijklmnopqrstuvwxyz` 指定到你的 `alphabet` 變數。請記得使用單引號或雙引號將字串括起來。

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
