---
id: 680a836235461fce1ff53b3f
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

`translate()` 方法接受由 `str.maketrans()` 產生的翻譯表作為引數。它在字串上被呼叫，並傳回原始字串的複本，其中的字元已根據翻譯表被替換：

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

對 `text` 呼叫 `translate()` 方法，並將 `translation_table` 作為引數傳入，將結果指定到名為 `encrypted_text` 的變數。

# --hints--

你應該有一個名為 `encrypted_text` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

你應該將 `text.translate(translation_table)` 指定給 `encrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
