---
id: 680a82806c4c6bc3567f2708
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

宣告一個名為 `text` 的新變數並指定其字串為 `hello world`。這將是要加密的訊息。

# --hints--

你應該宣告一個名為 `text` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

你應該將字串 `hello world` 指定給 `text`。請記得將字串用單引號或雙引號括起來，並注意字母大小寫。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
