---
id: 680a549f60f5a3fedfc0edd9
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

如你從輸出中所見，位移後的字母表從字母 `f` 開始，因為 `shift` 的值是 `5`。但現在字母表的前五個字母 – `a`、`b`、`c`、`d` 和 `e` – 在位移後的字母表中遺失了，所以你需要將它們添加到位移後字母表的末尾。

`+` 運算子用於將兩個或多個字串串接在一起，這個過程稱為串接，如下所示：

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

修改 `shifted_alphabet` 變數的現有指派：使用切片語法擷取 `alphabet` 缺少的第一部分，並將其串接到 `alphabet[shift:]`。

提醒你，`sentence[start:stop]` 傳回 `sentence` 中從位置 `start`（包含）到 `stop`（不包含）的字元。

# --hints--

你不應該直接修改 `alphabet` 變數。

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

你應該在宣告 `shifted_alphabet` 的同一行將 `alphabet[shift:]` 與 `alphabet` 缺少的第一部分串接後指定給 `shifted_alphabet` 變數。使用 `shift` 來指定切片停止的位置。

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
