---
id: 680a3d4c065557260de2ff33
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

如你可能還記得先前的課程，在 Python 中，你透過在指定運算子（`=`）的左側寫變數名稱，並在右側寫要指定的值來宣告變數：

```py
variable_name = value
```

創建一個名為 `shift` 的變數，並將值 `5` 指定給你的新變數。

# --hints--

你應該宣告一個名為 `shift` 的變數。請注意將變數名稱放在行首。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

你應該將值 `5` 指定到你的 `shift` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
