---
id: 6810891bd6de8a87a833df42
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

如你從輸出中所見，該訊息已被加密。接下來的步驟將是讓你的程式碼可重複使用，以防你想加密不同的訊息。

為此，你需要創建一個函式（程式）。提醒你，以下是如何創建一個名為 `spam` 的函式（程式），該函式（程式）會在終端機上列印 `Spam!`：

```py
def spam():
    print('Spam!')
```

創建一個名為 `caesar` 的函式（程式）。將你所有現有的程式碼放入該函式（程式）主體內。注意保持函式（程式）主體內所有行的相同縮排層次。

# --hints--

你應該有一個名為 `caesar` 的函式（程式）。

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

你應該將目前為止撰寫的所有程式碼移動到 `caesar` 函式（程式）主體內。請確保你保持函式（程式）主體內所有行的相同縮排層級……

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
