---
id: 6810869355fa2468f9ebaf9c
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

現在列印 `encrypted_text` 以在終端機中查看輸出。

# --hints--

你應該列印 `encrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
--fcc-editable-region--

--fcc-editable-region--
```
