---
id: 6818ed1adc249bb91d373729
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

在之前的課程中，你學到 `return` 述語用於從函式傳回一個值，讓你可以在程式碼的其他地方使用它：

```py
def spam():
    return 'Spam!'
```

從你的函式中移除 `print(encrypted_text)` 呼叫。然後，刪除 `encrypted_text` 變數的宣告，並直接從你的函式傳回 `text.translate(translation_table)`。

# --hints--

你不應該在 `caesar` 函式中使用 `print(encrypted_text)`。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

你不應該在 `caesar` 函式中有 `encrypted_text` 變數。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

你的 `caesar` 函式應該傳回 `text.translate(translation_table)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
