---
id: 6818b80519e6c6c47a8bb6e8
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

如你在先前的課程中所學，函式（程式）可以有參數，這些參數是在函式（程式）內可以取用的變數。這裡有一個帶有兩個參數 `num1` 和 `num2` 的 `sum` 函式（程式）。

```py
def sum(num1, num2):
    print(num1 + num2)
```

要加密的訊息和位移仍然硬編碼在你的函式中。為你的函式添加兩個名為 `text` 和 `shift` 的參數，並從 `caesar` 函式主體中刪除 `text` 和 `shift` 兩個變數的宣告。

# --hints--

你的 `caesar` 函式應該有兩個參數，`text` 和 `shift`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

你不應該在 `caesar` 函式（程式）主體內宣告 `text` 變數。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

你不應該在 `caesar` 函式（程式）主體內宣告 `shift` 變數。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
