---
id: 6818fcb250f34e62cab32c39
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

你可能已經注意到，雖然訊息已被加密，但大寫字元仍保持不變。這是因為翻譯表中不包含大寫字元。

要修正它，你需要使用 `upper()` 方法修改傳遞給你的 `str.maketrans()` 呼叫的字串：

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

更新你的 `str.maketrans()` 呼叫，將每個引數與該引數本身的大寫版本串接。

# --hints--

你應該透過將每個引數本身的大寫版本串接到該引數，來修改你的 `str.maketrans()` 呼叫的引數。

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
