---
id: 6818e9c7c8dc5694e8f64fc3
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

現在，透過使用字串 `freeCodeCamp` 和數字 `3` 作為引數來呼叫 `caesar` 函式，測試該函式。將函式呼叫指定給名為 `encrypted_text` 的變數。

# --hints--

你應該在全域的範圍中有一個名為 `encrypted_text` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

你應該將 `caesar('freeCodeCamp', 3)` 指定給 `encrypted_text` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
