---
id: 6818eb3f98c0bca5af7c8d03
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

現在你的程式碼是可重複使用的。然而，`caesar` 函式會在終端機列印一則訊息，並且預設回傳 `None`。為了證明這點，請在你的程式碼結尾列印 `encrypted_text`。 

# --hints--

你應該在函式本體外列印 `encrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
