---
id: 6819ca464db219ab86f37dc4
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

如你已經驗證過，移位用於加密文字應為正數。不過，它不能超過 `25`（`alphabet` 的最後索引）。

為你的 `if` 述語添加第二個條件，以驗證 `shift` 是否大於 25。請記住，Python 中的邏輯 OR 操作是透過 `or` 運算子實作的。

此外，更新傳回的訊息為 `Shift must be an integer between 1 and 25.`

# --hints--

你應該使用 `or` 運算子，在 `shift < 1` 之後添加第二個條件，檢查 `shift` 是否大於 `25`。你不應該修改現有的條件。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

你的第二個 `if` 述語應該傳回 `Shift must be an integer between 1 and 25.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
