---
id: 6819be9055828d65c813eec5
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

如果 `isinstance()` 函式的第一個引數是第二個引數的實體，則傳回 `True`，否則傳回 `False`：

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

將你 `if` 述語中目前的條件取代為 `isinstance()` 呼叫。將 `shift` 作為第一個引數，`int` 作為第二個引數傳入。 

# --hints--

你的 `if` 述語應該使用 `isinstance(shift, int)` 作為條件。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
