---
id: 6819bc0637b80256a33adccc
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

既然你已經實作了加密的基本功能，現在是時候添加一些驗證。為此，你將需要一個 `if` 述語。以下是 `if` 述語的語法提醒：

```py
if condition:
    # code to run when condition is true
```

在你的函式（程式）主體開始處，創建一個 `if` 述語。暫時使用 `True` 作為條件，並在該 `if` 述語主體內傳回字串 `Shift must be an integer value.`

# --hints--

你應該在你的 `caesar` 函式中有一個 `if` 述語。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

你的 `if` 述語應該使用 `True` 作為其條件。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

你應該從你的 `if` 述語傳回字串 `Shift must be an integer value.`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
