---
id: 6819cf38880d8cc9b9f6af7a
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

Python 允許你為函式的參數指定預設值，從而創建一個可以用較少或不帶引數呼叫的函式。以下是如何創建一個具有預設值的 `name` 參數的函式：

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

為你的函式（程式）添加第三個名為 `encrypt` 的參數，並將其預設值設為 `True`。 

# --hints--

你應該為你的函式（程式）添加第三個名為 `encrypt` 的參數，預設值為 `True`。不要修改現有的參數。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
