---
id: 6819c5c9faf2548e556163d5
title: 第 18 步
challengeType: 20
dashedName: step-18
---

# --description--

`not` 運算子會反轉布林值。它會將 `True` 變成 `False`，將 `False` 變成 `True`：

```py
print(not True) # False
print(not False) # True
```

使用 `not` 運算子修正 `if` 述語的條件，使你的函式在 `shift` *不是* 整數時傳回 `Shift must be an integer value.`。

# --hints--

你應該在你的 `if` 述語中使用 `not` 運算子來否定 `isinstance(shift, int)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
