---
id: 6819c7beb4a8289c2b4f24d5
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

你的函式不應接受負數或零的 `shift`。因此，在你的第一個 `if` 述語之後，創建另一個 `if` 述語來檢查 `shift` 是否小於 `1`，並傳回字串 `Shift must be a positive integer.`。

# --hints--

你應該有第二個 `if` 述語來檢查 `shift` 是否小於 `1`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift < 1", "1 > shift", "shift <= 0", "0 => shift"]
assert any(_cond.is_equivalent(option) for option in _options)
`) })
```

你的第二個 `if` 述語應該傳回 `Shift must be a positive integer.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be a positive integer.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
