---
id: 681a2a6cf95e82ca30866364
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

宣告兩個名為 `encrypt` 和 `decrypt` 的函式（程式），兩者皆帶有 `text` 和 `shift` 參數。

你會使用 `encrypt` 進行加密處理程序，並使用 `decrypt` 進行解密，為這兩個動作標註描述性名稱。

從你新的兩個函式（程式）中傳回一個 `caesar` 呼叫，並傳入 `text` 和 `shift`，但請確保對於 `decrypt` 函式（程式）也傳入 `False` 作為第三個引數。

# --hints--

你應該有一個名為 `encrypt` 的函式（程式），並帶有兩個參數，`text` 和 `shift`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

你的 `encrypt` 函式應該傳回 `caesar(text, shift)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

你應該有一個名為 `decrypt` 的函式（程式），並帶有兩個參數，`text` 和 `shift`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

你的 `decrypt` 函式應該傳回 `caesar(text, shift, False)`。

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
