---
id: 681a44f4c3235f7d8f428545
title: 第 25 步
challengeType: 20
dashedName: step-25
---

# --description--

現在你將要測試 `decrypt` 函式（程式）。將指定給 `encrypted_text` 的值取代為以下字串，該字串表述一則要解密的訊息：`Pbhentr vf sbhaq va hayvxryl cynprf.`

然後，宣告一個名為 `decrypted_text` 的變數，並將 `decrypt` 的呼叫指定給它，該呼叫以 `encrypted_text` 作為第一個引數，並以位移量 `13` 作為第二個引數。

最後，在終端機上列印 `decrypted_text`。這樣，凱薩密碼就完成了。

# --hints--

你應該將 `Pbhentr vf sbhaq va hayvxryl cynprf.` 指定到 `encrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'")`)) })
```

你應該有一個名為 `decrypted_text` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("decrypted_text")`)) })
```

你應該將 `decrypt(encrypted_text, 13)` 指定到 `decrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("decrypted_text").is_equivalent("decrypted_text = decrypt(encrypted_text, 13)")`)) })
```

你應該列印 `decrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(decrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = encrypt('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```

# --solutions--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'
decrypted_text = decrypt(encrypted_text, 13)
print(decrypted_text)
```
