---
id: 681a1d60e037127e7c8e0230
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

你將使用前一步驟中添加的參數來判斷該函式（程式）是否應該加密傳入的文字（預設行為，`encrypt=True`），或是否應該解密加密過的訊息。

創建一個 `if` 述語來檢查 `encrypt` 是否不為真。在新的 `if` 述語中，將 `shift` 設定為 `- shift`。這是為了啟用位移相對於加密處理程序以相反方向進行。

# --hints--

當 `encrypt` 不為真時，你應該將 `shift` 設定為 `- shift`。

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
