---
id: 6551f1e8737dbe6e0cb08ea4
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Ahora, imprime tu nueva variable.

# --hints--

No debes eliminar el llamado a `print(type(text))` ya existente.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Debes llamar a la función `print()` pasándole la variable `shift`. Preste atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
