---
id: 6551f5756c20146fc94f8675
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Puedes acceder al número de caracteres en una cadena con la función incorporada `len()`.

Modifica tu llamada existente`print()` pasando`len(text)`en lugar de`text[-1]`.

# --hints--


Debes llamar a la función `len()`.

```js
assert.match(code, /len\s*\(.*\)/)
```

Debes pasar la variable`text` a tu función `len()` incluyendo el nombre de la variable entre los paréntesis.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

Debes pasar `len(text)` a tu función `print()` incluyendo el nombre de la variable entre los paréntesis. Preste atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

No Debes tener `print(text[-1])` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
