---
id: 6551f1e8737dbe6e0cb08ea4
title: '--Paso 10--'
challengeType: 20
dashedName: step-10
---

# --description--

And now print your new variable.

# --hints--

No debes eliminar el llamado a `print(type(text))` ya existente.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

You should call the `print()` function passing in the `shift` variable. Preste atención para realizar la llamada al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

El código contiene una sintaxis no válida y/o sangría inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
