---
id: 6551ef504d91936d2d4e54f8
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Un *argument* es un objeto o una expresión que se pasa a una función - agregado entre paréntesis de apertura y cierre- cuando esta es llamada:

```py
greet = 'Hello!'
print(greet)
```

El código en el ejemplo anterior imprimiría la cadena `'Hello!'`, que es el valor de la variable `greet` pasada a `print()` como argumento.

Imprime tu variable `text` en la pantalla pasando la variable `text` como argumento a la función `print()`.

# --hints--

Debes pasar la variable `text` a tu función `print()` incluyendo el nombre de la variable entre los paréntesis. Preste atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
