---
id: 6551f6f9ec6db270eb83f3db
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Puedes ver el `11` impreso en la terminal, porque `'Hello World'` contiene 11 caracteres.

Otra función integrada útil es `type()`, que devuelve el tipo de dato de una variable. Modifica tu llamada a `print()` para imprimir el tipo de dato de `text`.

# --hints--

Debes llamar a la función `type()`.

```js
assert.match(code, /type\s*\(.*\)/)
```

Debes pasar `text` a tu función `type()` incluyendo el nombre de la variable entre los paréntesis.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

Debes pasar `type(text)` a tu función `print()` incluyendo el nombre de la variable entre los paréntesis. Preste atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

No Debes tener `print(len(text))` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
