---
id: 6551fb531ce7f773071a0b84
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Como puedes ver, el resultado de imprimir `type(text)` es `<class 'str'>`, lo que significa que tu variable es una cadena, indicado por el `str`.

Now go to a new line and create another variable called `shift` and assign the value `3` to this variable.

# --hints--

Debes declarar una variable llamada `shift`. Presta atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^shift\s*=/m)
```

You should assign the value `3` to your `shift` variable.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
