---
id: 6551ecf597038c36637f3759
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

Las variables son una parte esencial de Python, así como de cualquier lenguaje de programación. Una variable es un nombre que hace referencia o apunta a un objeto. Puede declarar una variable escribiendo el nombre de la variable en el lado izquierdo del operador de asignación `=` y especificando el valor para asignar a esa variable en el lado derecho del operador de la asignación:

```py
variable_name = value
```

Crea una variable llamada `number` y asignale el valor `5` a tu nueva variable.

# --hints--

Debes declarar una variable llamada `number`. Presta atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^number\s*=/m)
```

Debes asignale el valor `5` a tu variable `number`.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
