---
id: 6551f22653e6ce6e4f1d1370
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Modifica tu `print(shift)` llamada para imprimir el tipo de datos de tu variable `shift`.

# --hints--

No debes eliminar el llamado a `print()` ya existente.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Debes pasar `shift` a la función `type()`.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

Debes pasar `type(shift)` a tu función `print()`. Preste atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

No Debes tener `print(shift)` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
