---
id: 6551f3bfed58796ebb1268dc
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Cada caracter de una cadena puede ser referenciado por un índice numérico. El recuento de índices comienza en cero. Por lo que el primer caracter de una cadena tiene el índice `0`. Por ejemplo, en la cadena `'Hello World'`, la ` 'H'` está en el índice `0`, la` 'e'` está en el índice `1`, y así sucesivamente. 

Cada carácter de una cadena puede ser accedido usando la notación de corchetes. Necesitas escribir el nombre de la variable seguido de corchetes cuadrados y agregar el índice del carácter entre los corchetes:

```py
text = 'Hello World'
r = text[8]
```

Ahora, en lugar de imprimir `text`, imprime solo el carácter en el índice `6`.

# --hints--

Todavía debes llamar a la función `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Debes pasar `text[6]` a tu función `print()` incluyendo el nombre de la variable entre los paréntesis. Preste atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
