---
id: 6551f4259d66026ee28cb205
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

También puedes acceder a los caracteres de una cadena empezando desde el final de la misma. El último carácter tiene un índice de `-1`, el penúltimo `-2` y así sucesivamente.

Ahora modifica tu llamada existente a `print()` para imprimir el último carácter en tu cadena.

# --hints--

Todavía debes llamar a la función `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Debes pasar `text[-1]` a tu función `print()` incluyendo el nombre de la variable entre los paréntesis. Preste atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
