---
id: 65520f8282faf57a0db4f7fe
title: Paso 21
challengeType: 20
dashedName: step-21
---

# --description--

Como puedes ver en la salida, `'h'` está en el índice `7` en la cadena `alphabet`. Ahora necesitas encontrar la letra en el índice `7` más el valor de `shift`. Para eso, puedes usar el operador de suma, `+`, de la misma manera que lo usarías para una suma matemática.

Modifica tu variable `shifted` para que almacene el valor de `alphabet` en el índice `index + shift`.

# --hints--

Debes asignar `alphabet[index + shift]` a tu variable `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
