---
id: 65522255d5b9cd80f335c6f2
title: Paso 29
challengeType: 20
dashedName: step-29
---

# --description--

Al final del cuerpo de tu ciclo, declara una variable llamada `new_index` y asigna el valor de `index + shift` a esta variable.

# --hints--

Debes declarar una variable llamada `new_index` dentro de tu ciclo `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

Debes asignar `index + shift` a tu nueva variable al final del cuerpo de tu ciclo `for`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
