---
id: 655220953ba90d80514d7ee2
title: Paso 27
challengeType: 20
dashedName: step-27
---

# --description--

Actualmente, la función `print()` está tomando un solo argumento `char`, pero puede tomar múltiples argumentos, separados por una coma.

Agrega un segundo argumento a `print(char)` para que imprima el carácter y su índice dentro del alfabeto.

# --hints--

Debes agregar `index` como el segundo argumento a tu llamada existente `print(char)`. No olvides separar los argumentos con una coma.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
