---
id: 65521fc818947e800bffe48a
title: Paso 26
challengeType: 20
dashedName: step-26
---

# --description--

Dentro del ciclo `for`, antes de imprimir el carácter actual, declara una variable llamada `index` y asigna a esta variable el valor devuelto por `alphabet.find(char)`.

# --hints--

Debes declarar una nueva variable llamada `index` al inicio de tu ciclo `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

Debes asignar `alphabet.find(char)` a tu nueva variable `index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
