---
id: 655261b2e1f2c197093f3993
title: Paso 41
challengeType: 20
dashedName: step-41
---

# --description--

Actualmente, los espacios se encriptan como `'c'`. Para mantener el espaciado original en el mensaje plano, necesitarás una estructura condicional `if`. Esta está compuesta por la palabra clave `if`, una condición y dos puntos `:`.

```py
if x != 0:
    print(x)
```

En el ejemplo anterior, la condición de la sentencia `if` es `x != 0`. El código `print(x)`, dentro del cuerpo de la sentencia `if`, se ejecuta solo cuando la condición evalúa a `True` (en este ejemplo, significa que `x` es diferente de cero).

Al inicio de tu ciclo for, reemplaza `print(char == ' ')` con una estructura `if`. La condición de esta estructura `if` debe evaluar a `True` si `char` es un espacio vacío y a `False` en caso contrario. Dentro del cuerpo del `if`, imprime la cadena `'space!'`. Recuerda indentar esta línea.

# --hints--

No deberías tener `print(char == ' ')` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

Debes reemplazar `print(char == ' ')` con una estructura `if` que se active cuando `char == ' '`. No uses paréntesis para encerrar la condición del `if` y recuerda incluir los dos puntos al final.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

Debes imprimir la cadena `'space!'` dentro de tu nueva estructura `if`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
