---
id: 65520e6f2b9678799977f24d
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

Elimina la última llamada a `print()`. Luego, en lugar de `text[0]`, pasa `text[0].lower()` como argumento a tu llamada `.find()` y observa el resultado.

# --hints--

No deberías tener `print(text.lower())` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

Aún debes imprimir tu variable `index`.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

Debes actualizar tu llamada `alphabet.find(text[0])` para usar `text[0].lower()` como argumento. Presta atención a colocar la llamada al método al inicio de la línea.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
