---
id: 6552008c0d9d9075cbec9772
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Vas a utilizar el método `.find()` para encontrar la posición en el alfabeto de cada letra en tu mensaje. Un método es similar a una función, pero pertenece a un objeto.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

Ve, la `.find()`método es *called on* `sentence` (la cadena a buscar), y `'r'` (el carácter a local) se pasa como el argumento. La llamada `sentence.find('r')` devolverá `4`, que es el índice de la primera aparición de `'r'` en `sentence`.

Al final de su código, llame a `.find()` en `alphabet` y pase `'z'` como argumento al método.

# --hints--

Debes llamar al método `.find()`.

```js
assert.match(code, /find\s*\(.*\)/)
```

Debería llamar al método `.find()` en la variable `alphabet`.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


Debería llamar al método `.find()` en la variable `alphabet` y pasar `'z'` al método. Preste atención para realizar la llamada al principio de la línea.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
