---
id: 65525e359ca28d938baa82c5
title: Paso 39
challengeType: 20
dashedName: step-39
---

# --description--

Puedes obtener el mismo efecto de `a = a + b` usando el operador de asignación de suma:

```py
a += b
```

El operador de asignación de suma te permite agregar un valor a una variable y luego asignar el resultado a esa variable.

Usa el operador `+=` para sumar un valor y asignarlo al mismo tiempo a `encrypted_text`.

# --hints--

Debes usar el operador de asignación de suma para agregar `alphabet[new_index]` al valor actual de `encrypted_text` dentro del cuerpo de tu ciclo.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
