---
id: 655251308f31958d06cdf267
title: Paso 36
challengeType: 20
dashedName: step-36
---

# --description--

Por el momento, el carácter encriptado se actualiza en cada iteración. Sería mejor almacenar la cadena encriptada en una nueva variable. Antes de tu `for` loop, declara una variable llamada `encrypted_text` y asigna una cadena vacía (`''`) a esta variable.

# --hints--

Debes declarar una variable llamada `encrypted_text` antes de tu ciclo `for`.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

Debes asignar una cadena vacía a tu variable `encrypted_text`.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
