---
id: 655209a4a27dd37873c4cac3
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` devuelve el índice del carácter coincidente dentro de la cadena. Si el carácter no se encuentra, devuelve `-1`. Como puedes ver, el primer carácter en `text`, mayúscula `'H'`, no se encuentra, ya que `alphabet` contiene sólo letras minúsculas.

Puedes transformar una cadena en su equivalente en minúsculas con el método `.lower()`. Añade otra llamada `print()` para imprimir `text.lower()` y ver la salida.

# --hints--

Debes tener `print(index)` en tu código. Preste atención para realizar la llamada al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Debería imprimir `text.lower()`. Preste atención para colocar el nombre de la variable al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
