---
id: 655208d59b131e7816f18c96
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

La función `print()` solo te da una salida en la consola, pero las funciones y métodos pueden tener un valor de retorno que puedes usar en tu código.

Ahora asigna `alphabet.find(text[0])` a una variable llamada `index`. De esta manera, `index` almacenará el valor devuelto por `alphabet.find(text[0])`.

# --hints--

Debes declarar una variable llamada `index`. Presta atención a colocar el nombre de la variable al inicio de la línea.

```js
assert.match(code, /^index\s*=/m)
```

Debes asignar el valor devuelto por `alphabet.find(text[0])` a tu variable `index`.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
