---
id: 65524d2a1a253b8bb5197ae2
title: Paso 28
challengeType: 20
dashedName: step-28
---

# --description--

`find` nuevamente está devolviendo `-1` para letras mayúsculas, y también para el carácter de espacio. Te encargarás del espacio más adelante.

Por ahora, en lugar de iterar sobre `text`, cambia el ciclo `for` para iterar sobre `text.lower()`.

# --hints--

Debes modificar tu estructura `for` para iterar sobre `text.lower()` en lugar de `text`.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
