---
id: 65521203d7165c7b84b22ad4
title: Pasos 23
challengeType: 20
dashedName: step-23
---

# --description--

Un ciclo te permite recorrer sistemáticamente una secuencia de elementos y ejecutar acciones en cada uno.

En este caso, usarás un ciclo `for`. Así es como puedes iterar sobre `text`:

```py
for i in text:
```

`for` es la palabra clave que denota el tipo de bucle. `i` es una variable que toma secuencialmente el valor de los elementos en `text`. La instrucción termina con dos puntos, `:`.

Debajo de la línea donde declaraste `alphabet`, escribe un ciclo `for` para iterar sobre `text`. Usa `i` como la variable del ciclo.

Al hacerlo, hay un error en la terminal. Aprenderás sobre ello en el siguiente paso.

# --hints--

Debes usar la palabra clave `for` para crear un bucle. Asegúrate de colocar la palabra clave `for` al principio de la línea y dejar un espacio en blanco después de la palabra clave.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

Debes escribir la variable `i` después de la palabra clave `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

Debes escribir la palabra clave `in` después de `for i `. Asegúrate de dejar un espacio alrededor de la palabra clave `in`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

Debes escribir `text` después de `for i in `. No olvides agregar el `:` final.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

Tu ciclo `for` debe colocarse debajo de la línea de código `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
