---
id: 6553ed69ece88d29594748aa
title: Paso 52
challengeType: 20
dashedName: step-52
---

# --description--

Para ejecutar, una función necesita ser llamada (o invocada) añadiendo un par de paréntesis después de su nombre, así:

```py
function_name()
```

Al final de tu código, llama a tu función `caesar`. Presta atención a la identación.

# --hints--

Debes llamar a tu función `caesar`. Asegúrate de escribir la llamada a la función al inicio de la línea.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
