---
id: 6553f8c570f9982e013a8886
title: Paso 57
challengeType: 20
dashedName: step-57
---

# --description--

Actualmente, cada letra es cifrada siempre con la misma letra, dependiendo del desplazamiento especificado. ¿Y si el desplazamiento fuera diferente para cada letra? Esto sería mucho más difícil de descifrar. Este algoritmo se conoce como el cifrado de Vigenère, donde el desplazamiento para cada letra se determina mediante otro texto, llamado *clave*.

Comienza a transformar tu cifrado César en un cifrado de Vigenère eliminando las dos llamadas a funciones.

# --hints--

Debes eliminar tus dos llamadas a la función `caesar()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
