---
id: 6553980e0527fa115c705646
title: Paso 42
challengeType: 20
dashedName: step-42
---

# --description--

Ahora, en lugar de imprimir `'space!'`, usa el operador de asignación de suma para agregar el espacio (actualmente almacenado en `char`) al valor actual de `encrypted_text`.

# --hints--

No deberías tener `print('space!')` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

Debes usar el operador `+=` para agregar `char` al valor actual de `encrypted_text` dentro de tu nueva estructura `if`.

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
