---
id: 6553a755879b131a445e664c
title: Paso 44
challengeType: 20
dashedName: step-44
---

# --description--

Intenta asignar la cadena `'Hello Zaira'` a tu variable `text` y ve lo que sucede en la consola.

Verás una excepción `string index out of range`. No te preocupes, ¡averiguarás cómo arreglarlo pronto!

# --hints--

Debes tener una variable `text`.

```js
assert.match(code, /^text\s*=/m)
```

Debes asignar la cadena `'Hello Zaira'` a tu variable `text`.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
