---
id: 6553995f412dd8122ed38e4a
title: Paso 43
challengeType: 20
dashedName: step-43
---

# --description--

Una sentencia condicional también puede tener una cláusula `else`. Esta cláusula puede añadirse al final de una sentencia `if` para ejecutar código alternativo si la condición de la sentencia `if` es falsa:

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

Como puedes ver en tu salida, cuando las iteraciones del ciclo llegan al espacio, se añade un espacio a la cadena encriptada. Luego, el código fuera del bloque `if` se ejecuta y se añade una `c` a la cadena encriptada.

Para solucionarlo, agrega una cláusula `else` después de `encrypted_text += char` y sangra todas las líneas de código siguientes excepto la llamada a `print()`.

# --hints--

Debes crear una cláusula `else`. Recuerda incluir los dos puntos finales.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

Debes indentar las líneas de código después de tu cláusula `else` excepto la llamada a `print()`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
