---
id: 6553f6086add4b2cbb99fd78
title: Paso 56
challengeType: 20
dashedName: step-56
---

# --description--

En la parte inferior de tu código, después de tu existente `caesar(text, shift)` llamada, llama a tu función de nuevo.

Esta vez, pase la variable `text` y el entero `13` como argumentos.

# --hints--

Debes volver a llamar a tu función, esta vez pasando como argumentos `text` y `13`.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

Debe mantener la llamada a la función existente.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
