---
id: 6553f4f66099802c6ae94613
title: Paso 55
challengeType: 20
dashedName: step-55
---

# --description--

Actualmente, tu código genera un `TypeError`, porque la función `caesar` se define con dos parámetros (`message` and `offset`), por lo tanto espera ser llamado con dos argumentos **.

Llamar `caesar()` sin los argumentos requeridos detiene la ejecución del código.

Dale valores a `message` y `offset`, pasándole `text` y `shift` como argumentos al llamado a la función `caesar`.

# --hints--

Debes pasar `text` y `shift` como argumentos a tu llamada de función, incluyéndolos dentro de los paréntesis. No olvide separar los argumentos con una coma.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
