---
id: 6553a572f7a65718f1e42e18
title: Paso 49
challengeType: 20
dashedName: step-49
---

# --description--

Una función es esencialmente un bloque reutilizable de código. Ya has conocido algunas funciones integradas, como `print()`, `find()` y `len()`. Pero también puedes definir funciones personalizadas de esta forma:

```py
def function_name():
    <code>
```

Una declaración de función comienza con la palabra clave `def` seguida por el nombre de la función — un nombre de variable válido — y un par de paréntesis. La declaración termina con dos puntos.

Justo después de tu variable `shift`, declara una función llamada `caesar` y sangra todas las líneas siguientes para darle un cuerpo a tu nueva función.

# --hints--

Debes usar la palabra clave `def` para declarar una nueva función.

```js
assert.match(code, /^def\s+/m)
```

Debes escribir `caesar` como el nombre de la función después de la palabra clave `def`. Recuerda agregar un espacio después de `def`.

```js
assert.match(code, /^def\s+caesar/m)
```

Debes agregar un par de paréntesis después del nombre de la función. No olvides el dos puntos final.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

Debes indentar todas las líneas después de `shift = 3` para que se conviertan en tu nuevo cuerpo de función.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
