---
id: 6553a7d8d05cbb1ae335a665
title: Paso 45
challengeType: 20
dashedName: step-45
---

# --description--

Cuando el bucle alcanza la letra `Z`, la suma `index + shift` excede la longitud de la cadena `alphabet`. Por lo tanto,`alphabet[new_index]` está intentando usar un ídice no valido, lo que provoca un `IndexError`.

Puedes notar que la salida de la terminal se detiene en el espacio inmediatamente antes de la `Z`, la última `print` antes de lanzar el error.

En este caso, el operador de módulo, `%`, puede utilizarse para devolver el resto de la división entre dos números. Por ejemplo: `5 % 2` es igual a `1`, porque 5 dividido 2 tiene un cociente de 2 y un resto de 1.

Envuelve `index + shift` entre paréntesis, y realiza el módulo entre dicha expresión y `26`, que es la longitud de `alphabet`.

# --hints--

Debes tener `new_index = (index + shift) % 26` en tu declaración `else`.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
