---
id: 6553efd6ada3f42aa2d75448
title: Paso 53
challengeType: 20
dashedName: step-53
---

# --description--

Ahora deberías ver la salida de tu código nuevamente. Aunque este método funcione, tampoco mejora significativamente la reusabilidad del código. Llamar repetidamente a tu función siempre dará el mismo resultado. Sin embargo, las funciones pueden declararse con *parámetros* para introducir versatilidad y personalización:

```py
def function_name(param_1, param_2):
    <code>
```

Los parámetros son variables que puedes usar dentro de tu función. Una función puede ser declarada con diferentes cantidades de parámetros. En el ejemplo anterior, `param_1` y `param_2` son parámetros.

Modifica la declaración de tu función para que tome dos parámetros llamados `message` y `offset`.

Después de eso, verás que aparece un error en el terminal. Verás cómo solucionarlo en los siguientes pasos.

# --hints--

Tu función `caesar` debe tomar `message` y `offset` como parámetros. Recuerda separar los parámetros con una coma.

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
