---
id: 655491bd5b98b813fa5bedca
title: Preparación 58
challengeType: 20
dashedName: step-58
---

# --description--

Ahora modifica la declaración de tu función: cambia el nombre de la función a `vigenere` y el segundo parámetro a `key`.

# --hints--

Debes modificar el nombre de tu función a `vigenere`.

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

Tu función `vigenere` debe tomar `message` y `key` como parámetros.

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
