---
id: 65549f90cf78131c96ebcf28
title: Paso 61
challengeType: 20
dashedName: step-61
---

# --description--

Al programar, la legibilidad es clave. Los comentarios sirven como notas efectivas, explicando la lógica detrás de tu código. Son valiosos cuando regresas a un proyecto después de un tiempo y también ayudan a tus colegas a entender el código.

En Python, puedes escribir un comentario usando un `#`. Todo lo que venga después del `#` no se ejecutará.

Antes de tu sentencia `if`, añade un comentario que diga `Append space to the message`.

# --hints--

Debes agregar un comentario que diga `Append space to the message`. No olvides el carácter `#` al principio.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
