---
id: 6554930320d70414e7b6acc6
title: Paso 59
challengeType: 20
dashedName: step-59
---

# --description--

Elimina tu variable `shift`. Luego, declara una nueva variable llamada `custom_key` y asigna el valor de la cadena `'python'` a esta variable.

# --hints--

Debes eliminar la variable `shift` y su valor.

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

Debes declarar una variable llamada `custom_key`.

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

Debes asignar la cadena `'python'` a tu variable `custom_key`.

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
