---
id: 6554a88d5af937226f4a9121
title: Paso 65
challengeType: 20
dashedName: step-65
---

# --description--

El método `.index()` es idéntico al método `.find()`, pero lanza una excepción `ValueError` si no puede encontrar la subcadena.

Un `ValueError` es una excepción incorporada que se lanza cuando se pasa un argumento con el tipo correcto pero con un valor inapropiado a una función.

Después de incrementar `key_index`, declara una variable llamada `offset`. Encuentra el índice que tiene `key_char` en `alphabet` y asígnalo a `offset`. Usa `.index()` para encontrar el índice.

# --hints--

Debes declarar una variable llamada `offset`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

Tu variable `offset` debe almacenar el valor de `alphabet.index(key_char)`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
