---
id: 6555d6e12c95701172b55709
title: Paso 77
challengeType: 20
dashedName: step-77
---

# --description--

Las funciones pueden ser llamadas con argumentos por defecto. Un argumento por defecto indica el valor que la función debe tomar si no se pasa el argumento. Por ejemplo, la función a continuación acepta tres argumentos pero puedes llamarla pasando dos argumentos. El tercero asumirá el valor especificado por defecto:

```py
def foo(a, b, c=value):
    <code>
```

Modifica la función `vigenere` para que su parámetro `direction` tenga un valor predeterminado de `1`.

# --hints--

El parámetro `direction` de tu función `vigenere` debe tener un valor predeterminado de `1`.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

--fcc-editable-region--
def vigenere(message, key, direction):
--fcc-editable-region--
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
