---
id: 6555e43e783ed31a0532b1b2
title: Paso 93
challengeType: 20
dashedName: step-93
---

# --description--

El carácter de nueva línea `\n` es una secuencia especial usada para representar una nueva línea. Puedes escribir una barra invertida `\` seguida de una `n` como una secuencia normal de caracteres en una cadena y será reemplazada por una nueva línea en la salida cuando el programa se ejecute.

Pon un carácter de nueva línea al comienzo de tu primera llamada a `print` y observa la salida.

# --hints--

Debes modificar tu primera llamada a print en `print(f'\nEncrypted text: {text}')`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*f("|')\\nEncrypted\stext:\s\{\s*text\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print(f'Encrypted text: {text}')
print(f'Key: {custom_key}')
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
