---
id: 6555e1bfcd374e18c6be8e58
title: Paso 91
challengeType: 20
dashedName: step-91
---

# --description--

En Python, hay una manera de formatear cadenas f치cilmente. *f-strings* te permite interpolar valores en tus cadenas.

Interpolaci칩n significa escribir marcadores de posici칩n que ser치n reemplazados por los valores especificados cuando el programa se ejecute. Por ejemplo, puedes obtener el mismo resultado de `'Encrypted text: ' + text` con `f'Encrypted text: {text}'`. Necesitas poner una `f` antes de las comillas para crear la f-string y escribir las variables o expresiones que quieres interpolar entre llaves.

Modifique la primera llamada `print()` para usar una f-string.

# --hints--

Debe modificar la llamada `print('Encrypted text: ' + text)` para que imprima el mismo resultado usando una f-string.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]

    return final_message

def encrypt(message, key):
    return vigenere(message, key)

def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
