---
id: 655520c20cb1e6177b0641d6
title: Paso 81
challengeType: 20
dashedName: step-81
---

# --description--

El operador `not` se usa para negar una expresión. Cuando se coloca antes de un valor verdadero — un valor que evalúa a `True` — devuelve `False` y viceversa.

Agrega el operador `not` al inicio de la condición `if` para verificar si el carácter no es alfabético.

# --hints--

Debes usar el operador `not` en la condición de tu sentencia `if`.

```js
({ test: () => assert.match(code, /if\s+not\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char.isalpha():
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
