---
id: 6555231eac4d9f19bd3d44b5
title: Paso 83
challengeType: 20
dashedName: step-83
---

# --description--

La palabra clave `pass` puede usarse como un marcador de posición para código futuro. No tiene ningún efecto en tu código, pero puede evitar que recibas errores en caso de código incompleto:

```py
def foo():
    pass
```

Llamar a `vigenere` con `1` para encriptar y `-1` para desencriptar está bien, pero puede ser un poco críptico. Crea una nueva función llamada `encrypt` que tome los parámetros `message` y `key`, y usa `pass` para completar el cuerpo de la función.


# --hints--

Debes declarar una nueva función llamada `encrypt`. Recuerda usar la palabra clave `pass` dentro del cuerpo.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

Tu función `encrypt` debe tomar `message` y `key` como parámetros. El orden importa.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
