---
id: 655619327c7b364166f8dd6f
title: Paso 46
challengeType: 20
dashedName: step-46
---

# --description--

Si deseas incorporar caracteres adicionales en la cadena `alphabet`, como dígitos o caracteres especiales, verás que es necesario modificar manualmente el operando derecho de la operación módulo.

Reemplaza `26` con `len(alphabet)` para evitar este problema.

# --hints--

Debes modificar el valor de `new_index` reemplazando `26` con `len(alphabet)`.

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
  assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/));
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % 26
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
