---
id: 65688b5a1655a7a6caede847
title: Paso 31
challengeType: 20
dashedName: step-31
---

# --description--

Cuando intenta cambiar los caracteres individuales de una cadena como lo hizo en el paso anterior, obtienes un `TypeError`, que ocurre cuando se utiliza un objeto de tipo inapropiado en tu código.

Como puede ver en el mensaje de error, las cadenas no soportan la asignación de elementos, porque son inmutables. Sin embargo, a una variable puede reasignarsele otra cadena:

```py
message = 'Hello World'
message = 'Hello there!'
```

Elimina la linea `texto[0]` y reasigna `text` a la cadena `'Albatross'`.

# --hints--

No Debes tener `text[0] = 'C'` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

Debes reasignar `text` la cadena `'Albatross'`. No modifiques `text = 'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
