---
id: 65688a50e6c998a21d8e41d3
title: Paso 30
challengeType: 20
dashedName: step-30
---

# --description--

Las cadenas son inmutables, lo que significa que no pueden cambiarse una vez creadas. Por ejemplo, podrías pensar que el siguiente código cambia el valor de `my_string` a la cadena `'train'`, pero esto no es válido:

```py
my_string = 'brain'
my_string[0] = 't'
```

Confirma que al usar la notación de corchetes para acceder a la primera letra en `text` e intentar cambiarla por un carácter de tu elección, verás que la salida desaparece y aparece un error.

# --hints--

Debes tener `text = 'Hello World'` en tu código.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

Debes acceder a la primera letra en `text` con `text[0]`.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

Debes usar el operador `=` para asignar un carácter de tu elección a `text[0]`. No olvides encerrar el carácter entre comillas simples o dobles.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
