---
id: 656886ec4e4d9b96ded3c565
title: Paso 25
challengeType: 20
dashedName: step-25
---

# --description--

La variable de iteración puede tener cualquier nombre válido, pero es conveniente darle un nombre significativo.

Renombra tu variable `i` a `char`.

# --hints--

Debes reemplazar `i` con `char`. 

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
