---
id: 65687d2f8c7ee27b0446cef3
title: Pasos 24
challengeType: 20
dashedName: step-24
---

# --description--

El código que se ejecuta en cada iteración — colocado después de los `:` — constituye el cuerpo del bucle. Este código debe estar indentado. En Python, se recomienda usar 4 espacios por nivel de indentación. Este nivel indentado es un bloque de código.

```py
for i in text:
    print(i)
```

Python se basa en la indentación para indicar bloques de código. Dos puntos al final de una línea son una señal de que seguirá un nuevo bloque de código indentado.

Entonces, cuando no se encuentra un bloque indentado después de los dos puntos finales, la ejecución del código se detiene y se lanza un `IndentationError`. Este código no mostrará la salida y en su lugar generará un `IndentationError`:

```py
for i in text:
print(i)
```

Dale un cuerpo a tu `for` loop agregando una llamada a `print(i)`. Recuerda indentar el cuerpo del loop.

# --hints--

Debes agregar `print(i)` al cuerpo de tu `for` loop. Presta atención a la indentación. 

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
