---
id: 659530ca823d6c47f61daa8a
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Imprime tu variable `shifted`.

# --hints--

Debes imprimir tu variable `shifted`. Presta atención a colocar la llamada a la función al inicio de la línea.

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
