---
id: 65951b637ddec63611c482c0
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Puedes utilizar la función incorporada `print()` para mostrar la salida de tu código en la terminal.

Functions are reusable code blocks that you can call, or invoke, to run their code when you need them. Para llamar una función, basta con escribir un par de paréntesis junto a su nombre. Pronto aprenderás más sobre las funciones.

Por ahora, ve a una nueva línea y agrega una llamada vacía a la función `print()`. Aún no deberías ver ninguna salida.

# --hints--

Debes tener `print()` en tu código. Presta atención a colocar la llamada a la función al principio de la línea.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
