---
id: 65952b70924edb449faf1950
title: Paso 19
challengeType: 20
dashedName: step-19
---

# --description--

Declara una nueva variable llamada `shifted`. Usa la notación de corchetes para acceder al valor de `alphabet` en el índice `index` y asígnalo a tu nueva variable.

# --hints--

Debes declarar una variable llamada `shifted`. Presta atención a colocar el nombre de la variable al comienzo de la línea.

```js
assert.match(code, /^shifted\s*=/m)
```

Debes asignar `alphabet[index]` a tu variable `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
