---
id: 659526f3e5424b3eac4b6aa7
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

A continuación, imprime la variable `index` en la consola.

# --hints--

Debes imprimir tu variable `index`. Presta atención a colocar la llamada a la función al inicio de la línea.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
