---
id: 659db3e3670d3e712be82593
title: Paso 22
challengeType: 20
dashedName: step-22
---

# --description--

Repetir el proceso de localizar la letra dentro del alfabeto y determinar la letra desplazada para cada carácter en `text` puede ser tedioso. Afortunadamente, puede simplificarlo utilizando un bucle.

Por ahora, elimina todas las líneas de código debajo de la declaración de la variable `alphabet`.

# --hints--

Debes tener `text = 'Hello World'` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

Debes seguir usando `shift = 3` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

Debes tener `alphabet = 'abcdefghijklmnopqrstuvwxyz'` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

Debes eliminar la variable `index` y su valor.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

Debes tener `print(index)` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

Debes eliminar la variable `shifted` y su valor.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

No Debes tener `print(shifted)` en tu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

El código contiene una sintaxis no válida y/o indentación inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
