---
id: 6624e0ec0bf47017eed84b1f
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

El primer tipo de cifrado que vas a construir se llama cifrado *Caesar*. Específicamente, tomarás cada letra en tu mensaje, encontrarás su posición en el alfabeto, tomarás la letra ubicada 3 posiciones después en el alfabeto y reemplazarás la letra original con la nueva letra.

Para implementar esto, usarás el método `.find()` discutido en el paso anterior. Modifica tu llamada existente a `.find()` pasando `text[0]` como argumento en lugar de `'z'`.

# --hints--

Debes modificar tu llamada existente `alphabet.find('z')` pasando `text[0]` al método.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
