---
id: 66605a64a9822f01ec5ae8ce
title: Pasos 50
challengeType: 20
dashedName: step-50
---

# --description--

En Python, el *alcance* de una variable determina dónde se puede acceder a esa variable:

- Las variables definidas fuera de una función tienen un alcance *global* y pueden ser accedidas desde cualquier parte de tu código.

- Las variables definidas dentro de una función son *locales* a esa función y no pueden ser accedidas desde fuera de ella.

Para ver esto en acción, intenta imprimir la variable `alphabet` al final de tu código. Esto generará una excepción `NameError`.

Deberías ver un mensaje de error que indica que `alphabet` no está definido. Esto se debe a que `alphabet` está definido dentro de la función `caesar` y no es accesible fuera de ella.

# --hints--

Deberías intentar imprimir la variable `alphabet` fuera de la función caesar.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
