---
id: 60a3e3396c7b40068ad6996a
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

지금쯤이면 HTML 페이지가 가져야 할 기본 요소에 익숙해졌을 것입니다.

`DOCTYPE` 선언, 언어가 영어로 설정된 `html` 요소, `head` 요소, 그리고 `body` 요소를 포함하도록 코드를 설정하세요.

# --hints--

코드에 `<!DOCTYPE html>` 선언이 있어야 합니다.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

`<html>`인 `lang` 속성을 가진 여는 `en` 태그가 있어야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

코드에 `head` 요소 안에 `html` 요소가 있어야 합니다.

```js
assert(document.querySelectorAll('head').length === 1);
```

코드에 `body` 요소 안에 `html` 요소가 있어야 합니다.

```js
assert(document.querySelectorAll('body').length === 1);
```

`head` 요소가 `body` 요소보다 앞에 와야 합니다.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
