---
id: 60a3e3396c7b40068ad6996b
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`head` 요소 안에 `meta`를 `charset`로 설정하는 `UTF-8` 태그와 값이 `title`인 `Rothko Painting` 요소를 추가하세요.

`body` 요소 안에 `img`인 `src`을 가진 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` 요소를 추가하세요.
# --hints--

코드에 `meta` 태그가 있어야 합니다.

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` 태그는 `charset`를 `UTF-8`로 설정해야 합니다.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

코드에 `title` 요소가 있어야 합니다.

```js
assert(document.querySelectorAll('title').length === 1);
```

`title`는 `Rothko Painting`이어야 합니다.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

코드에 `img` 요소가 있어야 합니다.

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 요소는 `src`인 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`를 가져야 합니다.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
