---
id: 60a3e3396c7b40068ad69972
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

`height` 규칙에 `600px` 속성을 값 `.canvas`으로 추가하세요.

# --hints--

`height` 속성을 `600px`로 설정해야 합니다.

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

`.canvas` 요소는 `height`의 `600px`을 가져야 합니다.

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
