---
id: 60a3e3396c7b40068ad6998b
title: 33단계
challengeType: 0
dashedName: step-33
---

# --description--

여백이 한 방향으로 밀리는 것이 도움이 됩니다.

이 경우 `.one` 요소의 아래쪽 여백이 `.two`을 20픽셀 아래로 밀어냅니다.

`.two` 선택자에서 `margin` 단축 속성을 사용해 위쪽 여백을 `0`, 수평 여백을 `auto`, 아래쪽 여백을 `20px`로 설정하세요. 이렇게 하면 위쪽 여백이 제거되고, 수평으로 가운데 정렬되며, 아래쪽 여백이 20픽셀로 설정됩니다.

# --hints--

`margin` 속성을 `0 auto 20px`로 설정해야 합니다.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

`.two` 요소는 `margin` 값이 `0 auto 20px`이어야 합니다.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
