---
id: 60a3e3396c7b40068ad6997d
title: 20단계
challengeType: 0
dashedName: step-20
---

# --description--

여백을 사용해 `.one` 요소를 캔버스 위에 배치하세요.

수직 여백은 `margin`, 수평 여백은 `20px`인 축약형 `auto` 속성을 추가하세요.

# --hints--

`margin` 속성을 `20px auto`로 설정해야 합니다.

```js
const marginFilter = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.style.margin === '20px auto');
assert(marginFilter.length === 2);
```

`.one` 요소는 `margin` 값이 `20px auto`이어야 합니다.

```js
const oneMargin = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('margin');
assert(oneMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
