---
id: 60a3e3396c7b40068ad69986
title: 28단계
challengeType: 0
dashedName: step-28
---

# --description--

`.two` 요소 바로 아래에 `class` 값이 `three`인 새로운 `div`를 만드세요. 이것이 세 번째 사각형이 될 것입니다.

# --hints--

기존 `.one`와 `.two` 요소는 변경하지 마세요.

```js
assert(document.querySelectorAll('.one').length === 1);
assert(document.querySelectorAll('.two').length === 1);
```

새로운 `div`는 `.canvas` 요소 안에 중첩되어야 합니다.

```js
assert(document.querySelector('.canvas').children[2].tagName === 'DIV');
```

새로운 `div`는 `.two` 요소 다음에 와야 합니다.

```js
assert(document.querySelector('.two').nextElementSibling.tagName === 'DIV');
```

새로운 `div` 요소는 `three` 값을 가진 `class`를 가져야 합니다.

```js
assert(document.querySelector('.canvas').children[2].className.split(' ').includes('three'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
