---
id: 60a3e3396c7b40068ad6996e
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

여백은 박스 바깥의 영역이며, 다른 박스나 요소 사이의 간격을 조절하는 데 사용할 수 있습니다.

여기서 아래쪽 요소는 더 큰 위쪽 여백을 가지고 있어 페이지 아래로 더 밀려 있습니다.

CSS 박스 모델을 이해했으니, 이제 Rothko 그림 작업을 시작해 보세요.

`<img>` 요소를 제거하세요.

# --hints--

코드에 `img` 요소가 없어야 합니다.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
