---
id: 60a3e3396c7b40068ad69974
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

모든 그림에는 액자가 필요합니다.

`.canvas` 요소를 다른 `div` 요소로 감싸세요. 그 `div`에 `frame` 클래스를 주세요.
# --hints--

새로운 `div` 요소를 추가하세요.

```js
assert(document.querySelectorAll('div').length === 2)
```

`.canvas` 요소는 새 `div` 요소 안에 중첩되어야 합니다.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

새 `div`에는 값이 `class`인 `frame`이 있어야 합니다.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

새 `div`은 `body` 요소 안에 있어야 합니다.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
