---
id: 60a3e3396c7b40068ad6997c
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

`background-color`의 `.one`를 `#efb762`로 설정하세요.

# --hints--

`background-color` 속성을 `#efb762`로 설정해야 합니다.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(239, 183, 98)');
assert(hasBackground)
```

`.one` 요소는 `background-color` 값을 가진 `#efb762`를 가져야 합니다.

```js
const oneBackground = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('background-color');
assert(oneBackground === 'rgb(239, 183, 98)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
