---
id: 60a3e3396c7b40068ad69975
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

`.frame` 클래스 선택자를 사용해 새 규칙을 작성하세요.

`border` 축약 선언을 사용해 `.frame` 요소에 너비가 `50px`인 검은색 실선 테두리를 지정하세요.

# --hints--

코드에 `.frame` 선택자가 있어야 합니다.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

`border` 속성을 `50px solid black`로 설정해야 합니다.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

`.frame` 요소에 `50px solid black` `border`가 있어야 합니다.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
