---
id: 60a3e3396c7b40068ad69973
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

캔버스의 `background-color`를 `#4d0f00`로 변경하세요.

# --hints--

`background-color` 속성을 `#4d0f00`로 설정해야 합니다.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(77, 15, 0)');
assert(hasBackground);
```

`.canvas` 요소는 `background-color`인 `#4d0f00`를 가져야 합니다.

```js
const canvasBackground = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('background-color');
assert(canvasBackground === 'rgb(77, 15, 0)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
