---
id: 60a3e3396c7b40068ad6997e
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

이제 `.one`가 수평으로 가운데 정렬되었지만, 위쪽 여백이 캔버스를 넘어 프레임의 테두리까지 밀어내면서 전체 캔버스가 20픽셀 아래로 이동했습니다.

`padding` 요소가 밀어낼 수 있는 단단한 무언가를 제공하기 위해 `1px` 요소에 `.canvas`의 `.one`을 추가하세요.

# --hints--

`padding` 속성을 `1px`로 설정해야 합니다.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

`.canvas` 요소는 `padding` 속성 값이 `1px`여야 합니다.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
