---
id: 673b567e3ba535dda140d278
title: 投票システムを作成する
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

このラボでは、`Map` を使って投票を作成し、`Set` で重複投票を防ぐ投票システムを作成します。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `poll` 変数を新しい `Map` オブジェクトに初期化してください。

2. `addOption` をパラメータとして受け取る `option` 関数を用意してください。

3. `addOption` 関数内で:

   - `option` がまだ投票に存在しない場合、空の `Set` を値として投票に追加し、有権者を追跡できるようにしてください。また、メッセージ `Option "<option>" added to the poll.` を返してください。

   - `option` がすでに存在する場合は、メッセージ `Option "<option>" already exists.` を返してください。

   - 空のオプションを追加しようとした場合は、メッセージ `Option cannot be empty.` を返してください。

4. `vote`（投票するオプション）と `option`（有権者の一意のID）の2つのパラメータを受け取る `voterId` 関数を用意してください。

5. `vote` 関数内で:

   - `option` が投票に存在しない場合、メッセージ `Option "<option>" does not exist.` を返してください。

   - `option` が存在する場合、`voterId` がすでにこの `option` に投票しているか確認してください。

   - 有権者がすでに投票している場合は、メッセージ `Voter <voterId> has already voted for "<option>".` を返してください。

   - 有権者がまだ投票していない場合は、`voterId` をこのオプションの `Set` の `voters` に追加してください。関数はメッセージ `Voter <voterId> voted for "<option>".` を返してください。

6. `poll` に少なくとも3つのオプションがあること。

7. `poll` に少なくとも3票があること。

8. 投票結果を以下の形式で返す `displayResults` 関数を用意してください。

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

`poll` 変数を新しい `Map` オブジェクトに初期化してください。

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

`addOption` をパラメータとして受け取る `option` 関数を定義してください。

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

`vote` と `option` の2つのパラメータを受け取る `voterId` 関数を定義してください。

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

投票結果を表示する `displayResults` 関数を定義してください。

```js
assert.isFunction(displayResults);
```

`poll` に少なくとも3つのオプションがあること。

```js
assert.isAtLeast(poll.size, 3);
```

`poll` に少なくとも3票があること。

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

各投票オプションが `Set` オブジェクトにマッピングされていることを確認してください。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` は `Option "Egypt" added to the poll.` を返すべきです。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

空のオプションを追加すると `"Option cannot be empty."` を返すべきです。

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

`Turkey` がすでに追加されている場合、`addOption("Turkey")` は `Option "Turkey" already exists.` を返すべきです。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`Malaysia` が投票オプションに存在する場合、`vote("Malaysia", "traveler1")` は `Voter traveler1 voted for "Malaysia".` を返すべきです。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` はオプションの有権者の `Set` を更新すべきです。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`traveler1` が `Algeria` に再度投票しようとすると、`vote("Algeria", "traveler1")` は `Voter traveler1 has already voted for "Algeria".` を返すべきです。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

重複投票は `Set` のサイズを増やすべきではありません。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`Nigeria` が投票オプションにない場合、`vote("Nigeria", "traveler2")` は `Option "Nigeria" does not exist.` を返すべきです。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

一意のオプションは複数の票を受け取ることができます。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

オプション `Turkey` が2票、オプション `Morocco` が1票、オプション `Spain` が票を受け取らなかった場合、`displayResults()` は `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"` を返すべきです。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` は正しい形式で結果を返すべきです。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
