---
id: 66a83601cd819e37f0dccd14
title: Hatua 1
challengeType: 0
dashedName: step-1
---

# --description--

Katika warsha hii, utafanya mazoezi ya kutumia fomu za HTML kwa kujenga Fomu ya Maoni ya Hoteli.

Kwa maandishi ya utangulizi, utataka kuonyesha kichwa kikuu kifuatiwa na maelezo mafupi kuhusu kuacha maoni.

Ndani ya kipengele chako cha `body`, ongeza kipengele cha `header`.

Ndani ya kipengele cha `header`, ongeza kipengele cha `h1`. Maandishi ya kipengele cha `h1` yanapaswa kuwa `Hotel Feedback Form`.

Chini ya kipengele chako cha `h1`, ongeza kipengele cha `p`. Maandishi ya kipengele cha `p` yanapaswa kuwa `Thank you for staying with us. Please provide feedback on your recent stay.`.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `header`.

```js
assert.match(code, /<header>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `header`.

```js
assert.match(code, /<\/header>/i);
```

Unapaswa kuwa na lebo ya ufunguzi ya `h1`.

```js
assert.match(code, /<h1>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

Kipengele chako cha `h1` kinapaswa kuwa ndani ya kipengele cha `header`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

Unapaswa kuwa na lebo ya ufunguzi ya `p`.

```js
assert.match(code, /<p>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `p`.

```js
assert.match(code, /<\/p>/i);
```

Kipengele chako cha aya kinapaswa kuwa na maandishi `Thank you for staying with us. Please provide feedback on your recent stay.`.

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Kipengele chako cha aya kinapaswa kuwa ndani ya `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

Kipengele chako cha aya kinapaswa kuja baada ya `h1`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
