---
id: 66a83e5e491625454b6f62c3
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Fomu zinajumuisha `inputs` ambapo watumizi wanaweza kuingiza data zao. Unaweza kuunganisha ingizo zinazohusiana pamoja kwa kutumia kipengele cha `fieldset`.

Hapa kuna mfano wa kutumia kipengele cha `fieldset`:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

Ndani ya kipengele chako cha `form`, ongeza kipengele cha `fieldset`.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `fieldset`.

```js
assert.match(code, /<fieldset>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `fieldset`.

```js
assert.match(code, /<\/fieldset>/i);
```

Kipengele chako cha `fieldset` kinapaswa kuwa ndani ya kipengele cha `form`.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
