---
id: 66a8380d911e3f4270d5cadc
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Sasa, ni wakati wa kuongeza kipengele cha `main` kinachowakilisha maudhui kuu ya ukurasa.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `main`.

```js
assert.match(code, /<main>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `main`.

```js
assert.match(code, /<\/main>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
