---
id: 66a83fec026a7a4631e084d2
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Wakati unafanya kazi na vipengele vya `fieldset`, ni kawaida kutumia maelezo ya picha kuelezea kundi la vingozi. Unaweza kutumia kipengele cha `legend` kwa hili.

Hapa kuna mfano wa kutumia kipengele cha `legend`:

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

Ndani ya kipengele chako cha `fieldset`, ongeza kipengele cha `legend` chenye maandishi `Personal Information`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `legend`.

```js
assert.match(code, /<legend>/);
```

Unapaswa kuwa na lebo ya kufunga ya `legend`.

```js
assert.match(code, /<\/legend>/);
```

Kipengele chako cha `legend` kinapaswa kuwa ndani ya kipengele cha `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

Kipengele chako cha `legend` kinapaswa kuwa na maandishi `Personal Information`.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
