---
id: 66a84111965a0c46df6bbd0a
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kuhusisha kipengele cha `label` na `input` kama ifuatavyo:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

Sifa ya kipengele ya `for` kwenye kipengele cha `label` inapaswa kufanana na sifa ya kipengele ya `id` kwenye kipengele cha `input`. Hii inajulikana kama uhusiano wa wazi.

Chini ya kipengele chako cha `legend`, ongeza kipengele cha `label` chenye maandishi ya `Name (required):`. Weka sifa ya kipengele ya `for` kwa thamani ya `"full-name"`.

Kisha chini ya kipengele chako cha `label`, ongeza kipengele cha `input` kisicho na sifa za kipengele. Katika hatua zinazofuata, utaongeza sifa za kipengele zinazohitajika.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `label`.

```js
assert.match(code, /<label.*>/);
```

Unapaswa kuwa na lebo ya kufunga ya `label`.

```js
assert.match(code, /<\/label>/);
```

Kipengele chako cha `label` kinapaswa kuwekwa chini ya `legend` ndani ya kipengele cha `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

Kipengele chako cha `label` kinapaswa kuwa na sifa ya kipengele ya `for` yenye thamani ya `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

Kipengele chako cha `label` kinapaswa kuwa na maandishi ya `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

Unapaswa kuwa na kipengele cha `input` chini ya kipengele chako cha `label`.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
