---
id: 66a83bdcf425e7446900b7c4
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na kipengele cha `form` kama hii:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

Sifa ya kipengele ya `action` hutumika kubainisha wapi data ya fomu inapaswa kutumwa wakati fomu inapotumwa.

Sifa ya kipengele ya `method` hutumika kubainisha njia ya HTTP itakayotumika wakati wa kutuma data ya fomu. Njia zinazotumika zaidi ni `GET` na `POST`.

**KUMBUKA**: Utajifunza jinsi njia za HTTP zinavyofanya kazi katika moduli zijazo.

Ndani ya kipengele chako cha `main`, ongeza kipengele cha `form` chenye sifa ya kipengele ya `action` iliyowekwa kuwa `"https://hotel-feedback.freecodecamp.org"` na sifa ya kipengele ya `method` iliyowekwa kuwa `"POST"`.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `form`.

```js
assert.match(code, /<form.*>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `form`.

```js
assert.match(code, /<\/form>/i);
```

Kipengele chako cha `form` kinapaswa kuwa na sifa ya kipengele ya `action` iliyowekwa kuwa `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Kipengele chako cha `form` kinapaswa kuwa na sifa ya kipengele ya `method` iliyowekwa kuwa `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
