---
id: 66a93bbe65a26169dbf3bc39
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Fomu yako ya maoni ya hoteli inapaswa pia kukusanya anwani ya barua pepe kutoka kwa mtumizi.

Anza kwa kuongeza kipengele kipya cha `label` chenye maandishi ya `Email address (required):` kwenye fomu. Kipengele chako cha `label` kinapaswa kuwa na sifa ya kipengele ya `for` iliyowekwa kwa thamani ya `"email"`.

# --hints--

Unapaswa kuwa na kipengele cha `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

Kipengele chako cha `label` kinapaswa kuwa na maandishi ya `Email address (required):`.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

Kipengele chako cha `label` kinapaswa kuwa na sifa ya kipengele ya `for` iliyowekwa kwa thamani ya `"email"`.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
