---
id: 66a9689b1bf24b7750898a1b
title: Hatua 23
challengeType: 0
dashedName: step-23
---

# --description--

Kwa `input` na `label` za mwisho ndani ya `fieldset`, utaongeza kisanduku cha kuchagua `input` chenye sifa za kipengele `id` na `value` zilizowekwa kuwa `"price"`, na sifa ya kipengele `name` iliyowekwa kuwa `"choice"`.

Kisha, kipengele cha `label` chenye maandishi `Price` na sifa ya kipengele `for` iliyowekwa kuwa `"price"`.

Sasa unaweza kujaribu `form` yako kwa kuchagua visanduku mbalimbali vya kuchagua.

# --hints--

Unapaswa kuwa na kisanduku cha kuchagua cha tano `input`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"]:nth-of-type(5)'));
```

Kisanduku chako cha kuchagua cha tano `input` kinapaswa kuwa na sifa ya kipengele `id` iliyowekwa kuwa `"price"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('id'), 'price');
```

Kisanduku chako cha kuchagua cha tano `input` kinapaswa kuwa na sifa ya kipengele `name` iliyowekwa kuwa `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('name'), 'choice');
```

Kisanduku chako cha kuchagua cha tano `input` kinapaswa kuwa na sifa ya kipengele `value` iliyowekwa kuwa `"price"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('value'), 'price');
```

Unapaswa kuwa na kipengele cha `label` baada ya kisanduku chako cha kuchagua `input` chenye maandishi `Price`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.textContent?.trim(), 'Price');
```

`label` yako inapaswa kuwa na sifa ya kipengele `for` iliyowekwa kuwa `"price"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.getAttribute('for'), 'price');
```   

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
