---
id: 66a95d0eff8168747805f1f3
title: Hatua 19
challengeType: 0
dashedName: step-19
---

# --description--

Unapotaka mtumizi achague chaguzi nyingi kutoka kwenye orodha, unaweza kutumia visanduku vya kuchagua.

Hapa kuna mfano wa jinsi ya kufanya kazi na visanduku vya kuchagua vinavyohusiana na chaguzi za chakula:

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

Sifa ya kipengele `value` hutumika kubainisha thamani itakayotumwa kwa seva wakati fomu inapotumwa.

Chini ya kipengele chako cha `legend`, ongeza kisanduku cha kuchagua `input` chenye sifa za `id` na `value` zikiwa zimewekwa kuwa `"ads"`, na sifa ya `name` ikiwa imewekwa kuwa `"choice"`.

Chini ya kisanduku chako cha kuchagua cha `input`, ongeza kipengele cha `label` chenye maandishi ya `Social Media Ads`. Sifa ya `for` inapaswa kuwekwa kuwa `"ads"`.

# --hints--

Unapaswa kuwa na kisanduku cha kuchagua cha `input` chini ya kipengele chako cha `legend`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

Kisanduku chako cha kuchagua cha `input` kinapaswa kuwa na sifa ya `id` iliyowekwa kuwa `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

Kisanduku chako cha kuchagua cha `input` kinapaswa kuwa na sifa ya `name` iliyowekwa kuwa `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "choice");
```

Kisanduku chako cha kuchagua cha `input` kinapaswa kuwa na sifa ya `value` iliyowekwa kuwa `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

Unapaswa kuwa na kipengele cha `label` chini ya kisanduku chako cha kuchagua cha `input`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

Kipengele chako cha `label` kinapaswa kuwa na maandishi ya `Social Media Ads`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

Kipengele chako cha `label` kinapaswa kuwa na sifa ya `for` iliyowekwa kuwa `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
