---
id: 66a9521bc70162712caf118d
title: Hatua 16
challengeType: 0
dashedName: step-16
---

# --description--

Ikiwa unataka watumizi wachague chaguo moja kutoka kwenye orodha ya chaguzi, unaweza kutumia seti ya vitufe vya radio.

Hapa kuna mfano wa vitufe viwili vya radio:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

Katika mfano huu, vitufe vya radio vimeunganishwa pamoja kwa kutumia thamani ile ile ya sifa ya kipengele `name`. Hii ina maana kwamba kitufe cha radio kimoja tu ndicho kinaweza kuchaguliwa kwa wakati mmoja.

Chini ya kipengele chako cha `legend`, ongeza kitufe cha `radio` chenye sifa ya `id` kuwekwa kuwa `"yes-option"`, sifa ya `name` kuwekwa kuwa `"hotel-stay"`, na sifa ya `value` kuwekwa kuwa `"yes"`.

Chini ya kitufe chako cha `radio`, ongeza kipengele cha `label` chenye maandishi ya `Yes` na sifa ya `for` kuwekwa kuwa `"yes-option"`.

# --hints--

Unapaswa kuwa na kitufe cha `radio` chini ya kipengele chako cha `legend`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

Kitufe chako cha `radio` kinapaswa kuwa na sifa ya `id` iliyowekwa kuwa `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

Kitufe chako cha `radio` kinapaswa kuwa na sifa ya `name` iliyowekwa kuwa `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

Kitufe chako cha radio kinapaswa kuwa na sifa ya `value` iliyowekwa kuwa `"yes"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

Unapaswa kuwa na kipengele cha `label` chini ya kitufe chako cha `radio`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

Kipengele chako cha `label` kinapaswa kuwa na sifa ya `for` iliyowekwa kuwa `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

Kipengele chako cha `label` kinapaswa kuwa na maandishi ya `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
