---
id: 66a937e74920ba68ebe5e86d
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na sifa za kipengele za `placeholder` na `required` kama ifuatavyo:

```html
<input type="text" placeholder="e.g., John Doe" required>
```

Kwa kipengele chako cha `input` kilichopo, ongeza sifa ya kipengele ya `placeholder` yenye thamani ya `"e.g., John Doe"`.

Pia, ongeza sifa ya kipengele ya `required` kwa kipengele cha `input`.

# --hints--

Haupaswi kuondoa kipengele cha `input` kilichopo.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

`input` yako inapaswa kuwa na sifa ya kipengele ya `placeholder` yenye thamani ya `"e.g., John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

`input` yako inapaswa kuwa na sifa ya kipengele ya `required`.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
