---
id: 66a962954f4e0d76223b37ed
title: Hatua 21
challengeType: 0
dashedName: step-21
---

# --description--

Ifuatayo, ongeza kisanduku cha kuchagua kingine `input` chenye sifa za kipengele `id` na `value` zikiwa zimewekwa kuwa `"location"`, na sifa ya kipengele `name` ikiwa `"choice"`.

Kwa kipengele cha `label`, maandishi ya `Location` na sifa ya kipengele `for` zinapaswa kuwekwa kuwa `"location"`.

Chini ya kipengele cha `label`, ongeza kisanduku cha kuchagua kingine `input` chenye sifa za kipengele `id` na `value` zikiwa zimewekwa kuwa `"reputation"`, na sifa ya kipengele `name` ikiwa `"choice"`.

Kwa kipengele cha `label`, maandishi ya `Reputation` na sifa ya kipengele `for` zinapaswa kuwekwa kuwa `"reputation"`.

# --hints--

Unapaswa kuwa na kisanduku cha kuchagua cha tatu `input`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

Kisanduku chako cha kuchagua cha tatu `input` kinapaswa kuwa na sifa ya kipengele `id` iliyowekwa kuwa `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

Kisanduku chako cha kuchagua cha tatu `input` kinapaswa kuwa na sifa ya kipengele `name` iliyowekwa kuwa `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'choice');
```

Kisanduku chako cha kuchagua cha tatu `input` kinapaswa kuwa na sifa ya kipengele `value` iliyowekwa kuwa `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

Unapaswa kuwa na kipengele cha `label` baada ya kisanduku chako cha kuchagua `input` chenye maandishi ya `Location`.

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

`label` yako inapaswa kuwa na sifa ya kipengele `for` iliyowekwa kuwa `"location"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

Unapaswa kuwa na kisanduku cha kuchagua cha nne `input` chini ya `label` yako.

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

Unapaswa kuwa na kisanduku cha kuchagua cha nne `input` chenye sifa ya kipengele `id` iliyowekwa kuwa `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

Kisanduku chako cha kuchagua cha nne `input` kinapaswa kuwa na sifa ya kipengele `name` iliyowekwa kuwa `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'choice');
```

Kisanduku chako cha kuchagua cha nne `input` kinapaswa kuwa na sifa ya kipengele `value` iliyowekwa kuwa `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

Unapaswa kuwa na kipengele cha `label` chenye maandishi ya `Reputation` chini ya kisanduku cha kuchagua cha sifa ya hadhi.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

`label` yako inapaswa kuwa na sifa ya kipengele `for` iliyowekwa kuwa `"reputation"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
