---
id: 66a9401c9d660d6bb15993e2
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Fomu yako ya maoni ya hoteli inapaswa kuwa na chaguo kwa watumizi kuongeza umri wao.

Anza kwa kuongeza kipengele cha `label` chenye maandishi ya `Age (optional):` kwenye fomu.

Sifa ya kipengele ya `for` inapaswa kuwekwa kuwa `"age"`.

# --hints--

Unapaswa kuwa na kipengele cha `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

Kipengele chako cha `label` kinapaswa kuwa na sifa ya kipengele ya `for` iliyowekwa kuwa `"age"`.

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

Kipengele chako cha `label` kinapaswa kuwa na maandishi ya `Age (optional):`.

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
