---
id: 66a954b2bcddba72076c1857
title: Hatua 17
challengeType: 0
dashedName: step-17
---

# --description--

Chini ya kipengele chako cha `label`, ongeza kitufe cha `radio` chenye sifa ya `id` kuwekwa kuwa `"no-option"`, sifa ya kipengele `name` kuwekwa kuwa `"hotel-stay"`, na sifa ya kipengele `value` kuwekwa kuwa `"no"`. 

Chini ya kitufe chako kipya cha `radio`, ongeza kipengele kingine cha `label` chenye sifa ya kipengele `for` kuwekwa kuwa `"no-option"`. Maandishi ya `label` yanapaswa kuwa `No`.

Ukiisha, sasa unaweza kujaribu vitufe vya radio kwa kuchagua chaguo moja kwa wakati.

# --hints--

Unapaswa kuwa na kipengele cha `input` chini ya kipengele chako cha kwanza cha `label`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele `type` yenye thamani ya `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele `id` yenye thamani ya `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele `name` yenye thamani ya `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

Kipengele chako cha ingizo kinapaswa kuwa na sifa ya kipengele `value` yenye thamani ya `"no"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) label + input')?.value,
  'no'
);
```

Unapaswa kuwa na kipengele cha pili cha `label` chini ya kipengele chako cha pili cha `input`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

Kipengele chako cha `label` kinapaswa kuwa na sifa ya kipengele `for` yenye thamani ya `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

Kipengele chako cha `label` kinapaswa kuwa na maandishi ya `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
