---
id: 66a93730719e1f68410cce54
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Sifa ya kipengele `name` hutumika kutambua data ya fomu baada ya kuwasilishwa kwa seva.

Hapa kuna mfano wa jinsi ya kutumia sifa ya kipengele `name`:

```html
<input type="email" name="email">
```

Ongeza sifa ya kipengele `name` kwenye kipengele cha `input` kwa thamani ya `"name"`.

# --hints--

Haupaswi kuondoa kipengele cha `input` kilichopo.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Haupaswi kuunda kipengele kipya cha `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele `name` yenye thamani ya `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
