---
id: 66a930b20f589b6664c51cb0
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Wakati mtumizi anatoa jina lake kamili, `input` itakubali maandishi ya kawaida.

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na sifa ya kipengele ya `type` kama hii:

```html
<input type="text">
```

Kwa kipengele chako cha `input` kilichopo, ongeza sifa ya kipengele ya `type` iliyowekwa kuwa `"text"`.

Pia mpe kipengele cha `input` sifa ya kipengele ya `id` yenye thamani ya `"full-name"`.

# --hints--

Haupaswi kuondoa kipengele cha `input` kilichopo.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Haupaswi kuunda kipengele kipya cha `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `type` yenye thamani ya `"text"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `id` yenye thamani ya `"full-name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
