---
id: 66a96127422411756204bc1b
title: Hatua 20
challengeType: 0
dashedName: step-20
---

# --description--

Ongeza kisanduku cha kuchagua kingine `input` chenye sifa za kipengele `id` na `value` zikiwa zimewekwa kuwa `"recommendation"`, na sifa ya kipengele `name` ikiwa `"choice"`.

Chini ya kisanduku cha kuchagua `input`, ongeza `label` kingine chenye maandishi `Personal Recommendation`. Sifa ya kipengele `for` inapaswa kuwekwa kuwa `"recommendation"`.

# --hints--

Unapaswa kuwa na kisanduku cha kuchagua `input` chini ya kipengele chako cha `label`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(3) label + input[type="checkbox"]'));
```

Kisanduku chako cha kuchagua `input` kinapaswa kuwa na sifa ya kipengele `id` iliyowekwa kuwa `"recommendation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.id, "recommendation");
```

Kisanduku chako cha kuchagua `input` kinapaswa kuwa na sifa ya kipengele `name` iliyowekwa kuwa `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.name, "choice");
```

Kisanduku chako cha kuchagua `input` kinapaswa kuwa na sifa ya kipengele `value` iliyowekwa kuwa `"recommendation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.value, "recommendation");
```

Unapaswa kuwa na kipengele cha `label` chini ya kisanduku chako cha kuchagua `input`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label:nth-of-type(2)"));
```

Kipengele chako cha `label` kinapaswa kuwa na maandishi `Personal Recommendation`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"] + label:nth-of-type(2)')?.textContent?.trim(), "Personal Recommendation");
```

Kipengele chako cha `label` kinapaswa kuwa na sifa ya kipengele `for` iliyowekwa kuwa `"recommendation"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"] + label:nth-of-type(2)')?.getAttribute("for"), "recommendation");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
          
          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
