---
id: 66a9419e2d18476c645ce693
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Nambari `input` hutumika kuunda sehemu ya ingizo la nambari.

Hapa kuna mfano wa sehemu ya ingizo la nambari:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

Sifa za kipengele `min` na `max` hutumika kuweka kiwango cha chini na kiwango cha juu ambacho kinaweza kuingizwa katika sehemu ya ingizo.

Chini ya kipengele chako cha `label`, ongeza ingizo lenye sifa ya kipengele `type` iliyowekwa kuwa `"number"` na `id` ya `"age"`.

Sifa ya kipengele `name` inapaswa kuwekwa kuwa `"age"`, sifa ya kipengele `min` inapaswa kuwekwa kuwa `"3"` na sifa ya kipengele `max` inapaswa kuwekwa kuwa `"100"`.

# --hints--

Unapaswa kuwa na sehemu ya ingizo la nambari `input` chini ya kipengele chako cha `label`.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

Sehemu yako ya ingizo la nambari `input` inapaswa kuwa na `id` ya `age`.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

Sehemu yako ya ingizo la nambari `input` inapaswa kuwa na sifa ya kipengele `name` iliyowekwa kuwa `age`.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

Sehemu yako ya ingizo la nambari `input` inapaswa kuwa na sifa ya kipengele `min` iliyowekwa kuwa `"3"`.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

Sehemu yako ya ingizo la nambari `input` inapaswa kuwa na sifa ya kipengele `max` iliyowekwa kuwa `"100"`.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
