---
id: 66a93c95bc58e26a8fe95818
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Ifuatayo, ongeza `input` yenye aina ya `"email"` chini ya `label` yako ya barua pepe. `input` hii inapaswa kuwa na sifa ya kipengele ya `id` iliyowekwa kwa thamani ya `"email"`. Pia, weka sifa ya kipengele ya `name` kwa thamani ya `"email"`. 

`input` hii pia inahitajika, kwa hivyo hakikisha unaongeza sifa ya kipengele ya `required`.

Mwishowe, ongeza sifa ya kipengele ya `placeholder` iliyowekwa kwa `"example@email.com"`.

# --hints--

Unapaswa kuwa na kipengele cha `input` chini ya `label` yako ya barua pepe.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `type` yenye thamani ya `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `id` iliyowekwa kwa `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `name` iliyowekwa kwa `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `required`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `placeholder` iliyowekwa kwa `"example@email.com"`.
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
