---
id: 67a51d3a8a6fe123b77b0c6e
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Vipengele vya `input` vinaweza kuwa na sifa ya kipengele ya `size`. Sifa hii ya kipengele inaelezea idadi ya herufi zinazopaswa kuonekana wakati mtumizi anapoandika katika ingizo. Thamani ya `size` inapaswa kuwa nambari kamili chanya isiyokuwa sifuri. Ikiwa `size` haijatolewa, au imetolewa na thamani isiyo halali, ingizo litakuwa na upana wa kawaida uliowekwa na kivinjari.

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

Toa sifa ya kipengele ya `size` kwa ingizo la jina na barua pepe na thamani ya `"20"`.

# --hints--

Unapaswa kutoa sifa ya kipengele ya `input` kwa `size` ya jina na thamani ya `"20"`.

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

Unapaswa kutoa sifa ya kipengele ya `input` kwa `size` ya barua pepe na thamani ya `"20"`.

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
