---
id: 66a8347f97a33e36ffc81b9f
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

To prepare to create some actual content, add a `body` element below the `head` element.

# --hints--

You should have an opening `body` tag.

```js
assert.match(code, /<body>/i);
```

You should have a closing `body` tag.

```js
assert.match(code, /<\/body>/i);
```

Hupaswi kubadilisha kipengele chako cha `head` kilichopo. Hakikisha kuwa hukufuta tagi ya kufunga.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Kipengele chako cha `body` kinapaswa kuja baada ya kipengele chako cha `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>

  --fcc-editable-region--
</html>
```
