---
id: 68551b677a5f6898af364b85
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتعلم المزيد عن Tailwind CSS باستخدام فئات الأدوات flexbox لبناء مكوِّن دعوة لاتخاذ إجراء (CTA) متجاوب.

تم توفير قالب جاهز أساسي لـ HTML لك، لذا عليك بناء عناصر HTML أولاً.

ابدأ بإنشاء عنصر `div`. داخل هذا العنصر `div`، أنشئ عنصرين `div` منفصلين آخرين.

# --hints--

يجب أن تنشئ عنصر `div`.

```js
assert.exists(document.querySelector("div"));
```

يجب أن يكون لديك عنصران `div` آخران داخل عنصر الـ `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
