---
id: 685523c05fb41d0848c4b52f
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

داخل العنصر الأول `div`، أنشئ عنصر `span` بالنص `Soundflow`، وعنصر `h1` بالنص `Discover New Music`، وعنصر `p` بالنص `Stream your favorite tracks and discover new artists.`

# --hints--

يجب أن تنشئ عنصر `span` بالنص `Soundflow` داخل العنصر الأول `div`.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

يجب أن تنشئ عنصر `h1` بالنص `Discover New Music` داخل العنصر الأول `div`.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

يجب أن تنشئ عنصر `p` بالنص `Stream your favorite tracks and discover new artists.` داخل العنصر الأول `div`.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
