---
id: 68552464ef378e1051a55c28
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

داخل العنصر الثاني `div`، أنشئ عنصرين وسم الرابط. الأول يجب أن يحتوي على النص `Learn more`، والثاني يجب أن يحتوي على النص `Start listening`.

بالنسبة لخاصية `href` في العنصرين وسم الرابط، اجعلهما كلاهما يساوي `#`.

# --hints--

يجب أن تنشئ عنصرين `a` داخل العنصر الثاني `div`

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

يجب أن يحتوي أول عنصر `a` على النص `Learn more` وخاصية `href` بقيمة `#` داخل العنصر الثاني `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

يجب أن يحتوي ثاني عنصر `a` على النص `Start listening` وخاصية `href` بقيمة `#` داخل العنصر الثاني `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
