---
id: 6858ebf2552046ab3de9f17c
title: خطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

في هذه المرحلة، ستلاحظ وجود مساحة كبيرة جدًا على يمين المحتوى داخل الـ `div`. يمكنك تقليل هذه المساحة بإضافة عرض.

أضف الصنف `md:w-1/2` إلى الـ `div` لتقليل العرض إلى 50% على الشاشات المتوسطة والأكبر، أي الشاشات التي عرضها 768 بكسل فما فوق.

لتشاهد ذلك عمليًا، وسّع تبويب المعاينة.

# --hints--

يجب أن يحتوي عنصر الـ `div` الخارجي على الصنف المساعد `md:w-1/2`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
