---
id: 6858f3a67aa8580ff37016c7
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

إذا عدلت عرض تبويب المعاينة ليطابق عرض جهاز محمول مرة أخرى، ستلاحظ أن المحتوى دائمًا موضوع جنبًا إلى جنب. هذا ليس ما تريده على الأجهزة المحمولة.

لإصلاح ذلك، أضف الصنف `flex-col` إلى عنصر `div`. سيغير هذا اتجاه التوزيع إلى عمودي، فتتكدس العناصر فوق بعضها.

إليك كيفية استخدام صنف المساعدة `flex-col`:

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

يجب أن يحتوي عنصر `div` الافتتاحي على صنف المساعدة `flex-col`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
