---
id: 6858e7e23bf9e475e14cb931
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

الآن بعد أن بُني هيكل HTML، حان الوقت للانتقال إلى التنسيق الصحيح باستخدام Tailwind.

ابدأ بعنصر `div` الذي يلف المحتوى بأكمله. عيّن لون الخلفية له بإعطائه `class` بقيمة `bg-indigo-600`.

# --hints--

يجب أن يحتوي `div` الخارجي على خاصية `class` مضبوطة على `bg-indigo-600`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("bg-indigo-600"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
