---
id: 6858ef974c8ad5dcd7abb9f1
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

يمكن تحسين تباعد وترتيب المحتوى داخل `div`. لبقية التباعد والترتيب، تحتاج إلى استخدام flexbox.

أضف الصنف `flex` إلى `div` ليجلس عنصرَا `div` الفرعيان جنبًا إلى جنب.

أيضًا، أضف `justify-around` لتوزيع المساحة حول وبين المحتوى بشكل متساوٍ، و`items-center` لتمركز المحتوى بالكامل أخيرًا داخل `div`.

# --hints--

يجب أن يحتوي عنصر `div` الخارجي على الصنف المساعد `flex`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

يجب أن يحتوي عنصر `div` الخارجي على الصنف المساعد `justify-around`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

يجب أن يحتوي عنصر `div` الخارجي على الصنف المساعد `items-center`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
