---
id: 6858eadd3144c89d45e02bc7
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

إذا لاحظت، فإن `div` الخارجي ملتصق جدًا بالأعلى والمحتوى بداخله مكدس جدًا إلى اليسار. لذا، يجب أن تعمل على بعض التباعد الآن.

أضف الصنف `mt-8` إلى `div` لدفعه للأسفل قليلاً و`p-4` لتحريك النص بعيدًا عن اليسار.

# --hints--

يجب أن يحتوي عنصر `div` الخارجي على الصنف المساعد `mt-8`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

يجب أن يحتوي عنصر `div` الخارجي على الصنف المساعد `p-4`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
