---
id: 6858f90c54647a48ddb8822d
title: خطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

في هذه المرحلة، المحتوى، أي عنصرَي `div` داخل العنصر الخارجي `div`، سيُرتبان فوق بعضهما بغض النظر عن حجم الشاشة.

ذلك لأن لديك الصنف `flex-col` في عنصر `div` المغلف. لإصلاح ذلك، استخدم البادئة `lg:` لإضافة الصنف `flex-row` إلى العنصر `div`.

هذا يعني أن اتجاه التوزيع المرن سيكون `row` على الأجهزة ذات الشاشات الأكبر. لرؤية ذلك مجددًا، قم بتعديل عرض المتصفح إلى حجم أكبر.

# --hints--

يجب أن يحتوي عنصر `div` الخارجي على الصنف المساعد `lg:flex-row`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
