---
id: 68597114730d7c6b8838d204
title: خطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

الفقرة تحت `h1` تبدو جيدة، لذا لا تحتاج إلى المزيد من التنسيق. هذا يعني أنه يمكنك الانتقال إلى عناصر وسم الرابط داخل العنصر الثاني `div`.

أعطِ عنصر وسم الرابط الأول الأصناف `bg-white`، و`hover:bg-gray-100`، و`text-indigo-600`. هذا يمنحه خلفية بيضاء، ولون نص أزرق نيلي، وخلفية فاتحة قليلاً عند التمرير فوقه.

بالنسبة لعنصر وسم الرابط الثاني، أعطه الأصناف `bg-pink-500`، و`hover:bg-pink-600`، و`text-white`.

# --hints--

يجب أن يحتوي عنصر وسم الرابط الأول على صنف الأدوات `bg-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

يجب أن يحتوي عنصر وسم الرابط الأول على صنف الأدوات `hover:bg-gray-100`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

يجب أن يحتوي عنصر وسم الرابط الأول على صنف الأدوات `text-indigo-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

يجب أن يحتوي عنصر وسم الرابط الثاني على صنف الأدوات `bg-pink-500`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

يجب أن يحتوي عنصر وسم الرابط الثاني على صنف الأدوات `hover:bg-pink-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

يجب أن يحتوي عنصر وسم الرابط الثاني على صنف الأدوات `text-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
        --fcc-editable-region--
          Start listening
        </a>
      </div>
    </div>
  </body>
</html>
```
