---
id: 69708d59833ceded7d00541f
title: 連結リストの操作を実装する
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

このラボでは、連結リストの基本的なノードの追加と削除の機能に加えて、追加の操作を実装します。

**目的**：以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー**

1. `contains` 関数を持っている必要があります。この関数は連結リストと要素を受け取り、指定された要素が連結リストに存在する場合は `true` を返し、存在しない場合は `false` を返します。
1. `getAt` 関数を持っている必要があります。この関数は連結リストとインデックスを受け取り、連結リストの指定されたインデックスの要素を返します。インデックスが範囲外の場合は `undefined` を返します。
1. `insertAt` 関数を持っている必要があります。この関数は連結リスト、インデックス、要素を受け取り、指定された位置に要素を挿入します。インデックスが範囲外の場合はリストを変更しません。
1. `removeAt` 関数を持っている必要があります。この関数は連結リストとインデックスを受け取り、指定されたインデックスのノードを削除します。インデックスが範囲外の場合はリストを変更しません。
1. `clear` 関数を持っている必要があります。この関数は連結リストを受け取り、すべての要素を削除してリストを空の状態にリセットします。

**注意:** 後のテストは前のメソッドに依存する場合があります。例えば、`getAt` が正しく実装されていないと、`insertAt` や `removeAt` のテストが正しくても失敗することがあります。

# --hints--

`initList` は `head` が `null` に設定され、`length` が `0` に設定されたオブジェクトを返す必要があります。

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

`add` はリストの長さを1増やす必要があります。

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

`remove` はリストの長さを1減らす必要があります。

```js
const list = initList();  
add(list, 'apple');  
add(list, 'banana');  
remove(list, 'apple');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(list.head.element, 'banana');  
assert.isNull(list.head.next);
```

`contains` 関数を持っている必要があります。

```js
assert.isFunction(contains);
```

`contains` は要素がリストに存在する場合 `true` を返す必要があります。

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains` は要素がリストに存在しない場合 `false` を返す必要があります。

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains` は空のリストに対して `false` を返す必要があります。

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

`getAt` 関数を持っている必要があります。

```js
assert.isFunction(getAt);
```

`getAt` は指定されたインデックスの要素を返す必要があります。

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt` は無効なインデックスに対して `undefined` を返す必要があります。

```js
const list = initList();  
add(list, 'apple');  
assert.isUndefined(getAt(list, -1));  
assert.isUndefined(getAt(list, 1));  
assert.isUndefined(getAt(list, 5));
```

`getAt` は空のリストに対して `undefined` を返す必要があります。

```js
const list = initList();  
assert.isUndefined(getAt(list, 0));
```

`insertAt` 関数を持っている必要があります。

```js
assert.isFunction(insertAt);
```

`insertAt` はリストの先頭に要素を挿入する必要があります。

```js
const list = initList();  
add(list, 'second');  
add(list, 'third');  
insertAt(list, 0, 'first');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` はリストの中央に要素を挿入する必要があります。

```js
const list = initList();  
add(list, 'first');  
add(list, 'third');  
insertAt(list, 1, 'second');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` はリストの末尾に要素を挿入する必要があります。

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` は空のリストにインデックス `0` で挿入する必要があります。

```js  
const list = initList();  
insertAt(list, 0, 'first');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.isNull(list.head.next); 
```

`insertAt` は無効なインデックスに対してリストを変更しない必要があります。

```js
const list = initList();  
add(list, 'apple');  
insertAt(list, -1, 'invalid');  
insertAt(list, 10, 'invalid');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
assert.isUndefined(getAt(list, 1));
```

`removeAt` 関数を持っている必要があります。

```js
assert.isFunction(removeAt);
```

`removeAt` は指定されたインデックスの要素を削除する必要があります。

```js
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'third');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` は最初の要素を正しく削除する必要があります。

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

`removeAt` は最後の要素を正しく削除する必要があります。

```js  
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 2);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` は唯一の要素を削除するときにリストをリセットする必要があります。

```js
const list = initList();  
add(list, 'only');  
removeAt(list, 0);  
assert.strictEqual(list.length, 0);  
assert.isNull(list.head);  
assert.isTrue(isEmpty(list));
```

`removeAt` は無効なインデックスに対してリストを変更しない必要があります。

```js
const list = initList();  
add(list, 'apple');  
removeAt(list, -1);  
removeAt(list, 10);  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
```

`clear` 関数を持っている必要があります。

```js
assert.isFunction(clear);
```

`clear` はリストのすべての要素を削除する必要があります。

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`clear` はリストを再利用可能にする必要があります。

```js  
const list = initList();  
add(list, 'first');  
clear(list);  
clear(list);  
add(list, 'second');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'second');  
assert.isNull(list.head.next); 
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  
}

function getAt(list, index) {
 
}

function insertAt(list, index, element) {
  
}

function removeAt(list, index) {
 
}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
