---
id: 6842a6cd9836f0114a5b7a8a
title: 숫자 패턴 생성기 만들기
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

이 실습에서는 숫자 패턴을 만드는 작은 앱을 만들어 Python 기본기를 연습합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 양의 정수를 나타내는 단일 매개변수 `number_pattern`를 받는 `n`라는 이름의 함수를 정의해야 합니다.
1. `number_pattern`는 `for` 반복문을 사용해야 합니다.
1. `number_pattern(n)`는 1부터 `n`까지(포함) 모든 정수를 공백으로 구분한 문자열을 반환해야 합니다. 예를 들어, `number_pattern(4)`는 문자열 `1 2 3 4`을 반환해야 합니다.
1. 함수에 전달된 인수가 정수 값이 아니면 함수는 `Argument must be an integer value.`을 반환해야 합니다.
1. 함수에 전달된 인수가 1보다 작으면 함수는 `Argument must be an integer greater than 0.`를 반환해야 합니다.

# --hints--

`number_pattern` 함수를 만들어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

`number_pattern` 함수는 `n`라는 이름의 매개변수를 하나 가져야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)`는 `1 2 3 4`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)`는 `1 2 3 4 5 6 7 8 9 10 11 12`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern`는 양의 정수에 대해 공백으로 구분된 숫자 목록을 반환해야 합니다.

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

`number_pattern`는 정수가 아닌 값을 전달받으면 `Argument must be an integer value.`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern`는 0 이하의 정수를 전달받으면 `Argument must be an integer greater than 0.`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
