---
id: 67d83df6f82eda3868dd2a84
title: Construir um Personagem de RPG
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

Neste laboratório, você praticará o básico de Python construindo um pequeno aplicativo que cria um personagem para uma aventura de RPG.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função chamada `create_character`.
1. A função deve aceitar, na ordem, um nome de personagem seguido por três atributos: **força**, **inteligência** e **carisma**.
1. O nome do personagem deve ser validado:
   - Se o nome do personagem não for uma string, a função deve retornar `The character name should be a string`.
   - Se o nome do personagem tiver mais de 10 caracteres, a função deve retornar `The character name is too long`.
   - Se o nome do personagem contiver espaços, a função deve retornar `The character name should not contain spaces`.
1. As estatísticas também devem ser validadas:
   - Se uma ou mais estatísticas não forem inteiros, a função deve retornar `All stats should be integers`.
   - Se uma ou mais estatísticas forem menores que 1, a função deve retornar `All stats should be no less than 1`.
   - Se uma ou mais estatísticas forem maiores que 4, a função deve retornar `All stats should be no more than 4`.
   - Se a soma de todas as estatísticas for diferente de 7, a função deve retornar `The character should start with 7 points`.
1. Se todos os valores passarem na verificação, a função deve retornar uma string com quatro linhas:
   - a primeira linha deve conter o nome do personagem
   - As linhas 2-4 devem começar com a abreviação da característica, `STR`, `INT` ou `CHA` (nessa ordem), seguida de um espaço e depois um número de pontos cheios (`●`) igual ao valor da característica, e um número de pontos vazios (`○`) para completar 10. Exemplo: se o valor de strength for 3, deve haver 3 pontos cheios seguidos de 7 pontos vazios. Os pontos são fornecidos no editor.

Aqui está a string `create_character("ren", 4, 2, 1)` deve retornar:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

NOTA: enquanto `str` e `int` são abreviações comuns para as estatísticas, lembre-se de que esses são palavras-chave reservadas em Python e não devem ser usados como nomes de variáveis.

# --hints--

Você deve ter uma função chamada `create_character`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

Quando `create_character` é chamado com um primeiro argumento que não é uma string, ele deve retornar `O nome do personagem deve ser uma string`.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {"a": 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'

# even if wrong numbers, wrong type has precedence

assert create_character(1, 1, 1, 1) == 'The character name should be a string'
`
    )
})
```

Quando `create_character` é chamado com um primeiro argumento que possui mais de 10 caracteres, ele deve retornar `O nome do personagem é muito longo`.

```js
({
    test: () => runPython(
        `
assert create_character("aaabbbcccdd", 4, 2, 1) == 'The character name is too long'
`
    )
})
```

Quando `create_character` é chamado com um primeiro argumento que contém um espaço, ele deve retornar `O nome do personagem não deve conter espaços`.

```js
({
    test: () => runPython(
        `
assert create_character("cha cha", 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

Quando `create_character` é chamado com um segundo, terceiro ou quarto argumento que não seja um inteiro, ele deve retornar `All stats should be integers`.

```js
({
    test: () => runPython(
        `
wrong_types = ["3", "5", {"a": 3}, 3.1, [1, 2], (1, 2), {1, 2}, "str", "friend"]
for wrong_type in wrong_types:
    assert create_character("friend", wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character("friend", 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character("friend", 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

Quando `create_character` é chamado com um segundo, terceiro ou quarto argumento menor que `1`, ele deve retornar `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
expected = "All stats should be no less than 1"

assert create_character("ren", 0, 4, 3) == expected
assert create_character("ren", 4, 0, 3) == expected
assert create_character("ren", 4, 3, 0) == expected
assert create_character("ren", -1, 4, 4) == expected
assert create_character("ren", 4, -1, 4) == expected
assert create_character("ren", 4, 4, -1) == expected
`
    )
})
```

Quando `create_character` é chamado com um segundo, terceiro ou quarto argumento que seja maior que `4`, ele deve retornar `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character("ren", 1, 5, 1) == 'All stats should be no more than 4'
assert create_character("ren", 5, 1, 1) == 'All stats should be no more than 4'
assert create_character("ren", 1, 1, 5) == 'All stats should be no more than 4'
assert create_character("ren", 1, 1, 7) == 'All stats should be no more than 4'
assert create_character("ren", 1, 7, 1) == 'All stats should be no more than 4'
assert create_character("ren", 7, 1, 1) == 'All stats should be no more than 4'
`
    )
})
```

Quando `create_character` é chamado com um segundo, terceiro ou quarto argumento que não somam `7`, ele deve retornar `O personagem deve começar com 7 pontos`.

```js
({
    test: () => runPython(
        `
assert create_character("ren", 4, 4, 4) == 'The character should start with 7 points'
assert create_character("ren", 1, 1, 1) == 'The character should start with 7 points'
assert create_character("ren", 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

`create_character("ren", 4, 2, 1)` deve retornar `ren\nFOR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCAR ●○○○○○○○○○`.

```js
({
    test: () => runPython(
        `
assert create_character("ren", 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

Quando `create_character` é chamado com valores válidos, ele deve exibir as estatísticas do personagem conforme necessário.

```js
({
    test: () => runPython(
        `
assert create_character("Bob", 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character("aaabbbcccd", 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return "The character name should be a string"
    if len(character_name) > 10:
        return "The character name is too long"
    if ' ' in character_name:
        return "The character name should not contain spaces"

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return "All stats should be integers"
    if strength < 1 or intelligence < 1 or charisma < 1:
        return "All stats should be no less than 1"
    if strength > 4 or intelligence > 4 or charisma > 4:
        return "All stats should be no more than 4"
    if strength + intelligence + charisma != 7:
        return "The character should start with 7 points"

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
