---
id: 67d83df6f82eda3868dd2a84
title: Construir um personagem de RPG
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

Neste laboratório, você praticará o básico de Python construindo um pequeno app que cria um personagem para uma aventura de RPG.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função chamada `create_character`.
1. A função deve aceitar, na ordem, um nome de personagem seguido por três atributos: strength, intelligence e charisma.
1. O nome do personagem deve ser validado:
   - Se o nome do personagem não for uma string, a função deve retornar `The character name should be a string`.
   - Se o nome do personagem for uma string vazia, a função deve retornar `The character should have a name`.
   - Se o nome do personagem tiver mais de 10 caracteres, a função deve retornar `The character name is too long`.
   - Se o nome do personagem contiver espaços, a função deve retornar `The character name should not contain spaces`.
1. Os atributos também devem ser validados:
   - Se um ou mais atributos não forem inteiros, a função deve retornar `All stats should be integers`.
   - Se um ou mais atributos forem menores que 1, a função deve retornar `All stats should be no less than 1`.
   - Se um ou mais atributos forem maiores que 4, a função deve retornar `All stats should be no more than 4`.
   - Se a soma de todos os atributos não for igual a 7, a função deve retornar `The character should start with 7 points`.
1. Se todos os valores passarem na verificação, a função deve retornar uma string com quatro linhas:
   - A primeira linha deve conter o nome do personagem.
   - As linhas 2 a 4 devem começar com a abreviação do atributo, `STR`, `INT` ou `CHA` (nessa ordem), depois um espaço, e então um número de pontos cheios (`●`) igual ao valor do atributo, e um número de pontos vazios (`○`) para completar 10. Exemplo: se o valor de strength for 3, deve haver 3 pontos cheios seguidos de 7 pontos vazios. Os pontos já estão dados no editor.

Aqui está a string que deve ser retornada por `create_character('ren', 4, 2, 1)`:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

**Nota:** Embora `str` e `int` sejam abreviações comuns para os atributos, Python também usa esses nomes para seus tipos string e inteiro, então é melhor evitar usá-los como nomes de variáveis.

# --hints--

Você deve ter uma função chamada `create_character`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

Quando `create_character` for chamado com um primeiro argumento que **não** seja uma string, deve retornar `The character name should be a string`.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

Quando `create_character` for chamado com um primeiro argumento que seja uma string, **não deve** retornar `The character name should be a string`.

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

Quando `create_character` é chamado com um primeiro argumento que é uma string vazia, ele deve retornar `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

Quando `create_character` é chamado com um primeiro argumento que **não** é uma string vazia, ele **não deve** retornar `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

Quando `create_character` for chamado com um primeiro argumento que tenha mais de 10 caracteres, deve retornar `The character name is too long`.

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

A função `create_character` **não deve** dizer que o nome do personagem é muito longo quando ele não tiver mais de 10 caracteres.

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

Quando `create_character` for chamado com um primeiro argumento que contenha um espaço, deve retornar `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

Quando `create_character` for chamado com um primeiro argumento que *não* contenha espaço, *não deve* retornar `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

Quando `create_character` for chamado com um segundo, terceiro ou quarto argumento que **não** seja um inteiro, deve retornar `All stats should be integers`.

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

Quando `create_character` for chamado com um segundo, terceiro e quarto argumento que sejam todos inteiros, **não deve** retornar `All stats should be integers`.

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

Quando `create_character` for chamado com um segundo, terceiro ou quarto argumento que seja menor que `1`, deve retornar `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

Quando `create_character` for chamado com um segundo, terceiro e quarto argumento que sejam todos iguais ou maiores que `1`, **não deve** retornar `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

Quando `create_character` for chamado com um segundo, terceiro ou quarto argumento que seja maior que `4`, deve retornar `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

Quando `create_character` for chamado com um segundo, terceiro e quarto argumento que sejam todos iguais ou menores que `4`, **não deve** retornar `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

Quando o segundo, terceiro e quarto argumentos passados para `create_character` **não** somarem `7`, `create_character` deve retornar `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

Quando `create_character` for chamado com um segundo, terceiro e quarto argumento que somem `7`, **não deve** retornar `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` deve retornar `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

Quando `create_character` for chamado com valores válidos, deve exibir os atributos do personagem conforme o esperado.

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
        return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
