---
id: 6735a7370e0ae93a4577c689
title: 设计一个杂志版面
challengeType: 25
dashedName: design-a-magazine-layout
demoType: onClick
---

# --description--

在本实验室中，你将使用 CSS Grid 设计杂志布局，包括网格行和网格列等概念。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个类为 `magazine-cover` 的 `main`元素。

2. 你的 `.magazine-cover` 应包含一个 `header` 元素，其类为 `title`。

3. 你的页面中应该有四个 `section` 元素。 前三个部分用于杂志文章，第四个 `section` 元素用于封面图片。

4. 四个 `section` 元素的 `class` 属性分别为 `feature-article`、`small-article1`、`small-article2` 和 `cover-image`。

5. 你应该将 `.magazine-cover` 的 `display` 属性设置为 `grid`。

6. `.magazine-cover` 应包含一个 `grid-template-areas` 属性，用于定义网格布局：

   - 标题横跨顶行。
   - 在第二行，专题文章横跨两栏，然后是封面图片。
   - 第三行是两篇小文章和封面图片。

7. `.magazine-cover` 应具有一个 `grid-template-columns` 属性，可将空间分成三个相等的部分。

8. `.magazine-cover` 应具有一个 `grid-template-rows` 属性。 对于第一行，应将其设置为 `auto`，并平分其余两行所占的空间。

9. 你应在网格项之间添加间隙，并将 `.magazine-cover` 的宽度设置为占据大部分视口宽度。

**注意：** 确保在 HTML 中链接样式表以应用 CSS。

# --hints--

你应该有一个 `main` 元素。

```js
assert.exists(document.querySelector('main'));
```

你的 `main` 元素应该有一个 `class`，值为 `magazine-cover`。

```js
assert.exists(document.querySelector('main.magazine-cover'));
```

你的 `main` 元素中应该有一个 `header` 元素。

```js
assert.exists(document.querySelector("main > header"));
```

你的 `header` 应该有一个 `class`，值为 `title`，并包含一些文本。

```js
const header = document.querySelector("header.title");
assert.exists(header);
assert.isAbove(header.innerText.length, 0);
```

你应该为每篇文章和封面图片设置一个 `section` 元素, 类分别为 `feature-article`、`small-article1`、`small-article2` 和 `cover-image`。

```js
const featureArticle = document.querySelector('section.feature-article');
const smallArticle1 = document.querySelector('section.small-article1');
const smallArticle2 = document.querySelector('section.small-article2');
const coverImage = document.querySelector('section.cover-image');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(coverImage);
```

`.magazine-cover` 的 `display` 属性应该设置为 `grid`。

```js
assert.equal(getComputedStyle(document.querySelector('.magazine-cover')).display, 'grid');
```

你的 `.magazine-cover` 应该有一个 `grid-template-areas` 属性， 设置 `title` 占据第一行的全部三列。

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /^("|')title title title\1/);
```

网格模板的第二行应包含一个跨两列的 `feature article` 和 一个 `cover image`。

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /.("|')feature-article feature-article cover-image\1./);
```

网格模板的第三行应包含 `small-article1`、`small-article2` 和 `cover-image`。

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /("|')small-article1 small-article2 cover-image\1$/);
```

你的 `.magazine-cover` 选择器应使用 `grid-template-columns` 属性将空间划分为三个相等的部分。 请记住，你可以使用 `repeat` 函数或手动定义每个分数单位。

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

你的 `.magazine-cover` 类应将 `grid-template-rows` 属性设置为 `auto 1fr 1fr` 。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gridTemplateRows, 'auto 1fr 1fr');
```

你应该在网格项之间添加 `10px` 的间隙。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gap, '10px');
```

`.title` 类应有一个同名的网格区域。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` 类应有一个同名的网格区域。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` 类应有一个同名的网格区域。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Magazine Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Adventure Magazine</title>
  <link rel="stylesheet" href="styles.css">
</head>
<body>
  <main class="magazine-cover">
      <header class="title">
          <h1>Adventure Awaits</h1>
      </header>
      <section class="feature-article">
          <h2>Top 10 Exotic Destinations for 2024</h2>
          <p>Discover the most breathtaking places to visit this year, from hidden beaches to mountain retreats. Our guide takes you through the best spots for your next adventure.</p>
      </section>
      <section class="small-article1">
          <h3>Gear Up: Must-Have Gadgets</h3>
          <p>Check out the latest tech and gear to make your travels more exciting and comfortable.</p>
      </section>
      <section class="small-article2">
          <h3>Meet the Explorers</h3>
          <p>Get to know the modern adventurers who are pushing the boundaries of exploration.</p>
      </section>
      <section class="cover-image">
          <img src="https://cdn.freecodecamp.org/curriculum/labs/magazine-cover.png" alt="Adventure Cover Image">
      </section>
  </main>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 0;
  padding: 0;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.magazine-cover {
  display: grid;
  grid-template-areas: 
      "title title title"
      "feature-article feature-article cover-image"
      "small-article1 small-article2 cover-image";
  grid-template-columns: 1fr 1fr 1fr;
  grid-template-rows: auto 1fr 1fr;
  gap: 10px;
  width: 80%;
  background-color: #fff;
}

.title {
  grid-area: title;
  background-color: #1B1B32;
  color: white;
  padding: 10px;
  text-align: center;
}

.feature-article {
  grid-area: feature-article;
  background-color: #f9f9f9;
  padding: 10px;
}

.small-article1, .small-article2 {
  background-color: #e9e9e9;
  padding: 10px;
}

.cover-image {
  grid-area: cover-image;
  display: flex;
  justify-content: center;
  align-items: center;
  background-color: #ddd;
  padding: 10px;
}

.cover-image img {
  max-width: 100%;
  height: auto;
} 
```
