---
id: 5895f70cf9fc0f352b528e67
title: Passport 사용자 직렬화 구현하기
challengeType: 2
forumTopicId: 301556
dashedName: implement-the-serialization-of-a-passport-user
---

# --description--

데이터베이스가 설정되어 있지 않아서 실제 사용자 객체를 불러오지 않습니다. 서버를 시작할 때 한 번 데이터베이스에 연결하고, 앱의 전체 수명 주기 동안 지속적인 연결을 유지하세요. 이를 위해 데이터베이스 연결 문자열(예: `mongodb+srv://<username>:<password>@cluster0-jvwxi.mongodb.net/?retryWrites=true&w=majority`)을 환경 변수 `MONGO_URI`에 추가하세요. 이 값은 `connection.js` 파일에서 사용됩니다.

*MongoDB Atlas에서 무료 데이터베이스 설정에 문제가 있다면, 이 <a href="https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/" target="_blank" rel="noopener noreferrer nofollow">튜토리얼</a>을 참고하세요.*

이제 데이터베이스에 연결한 후 요청을 수신 대기하려고 합니다. 이렇게 하는 목적은 데이터베이스가 연결되기 전이나 데이터베이스 오류가 있을 때 요청을 허용하지 않기 위함입니다. 이를 위해 직렬화와 앱 라우트를 다음 코드로 감싸세요:

```javascript
myDB(async client => {
  const myDataBase = await client.db('database').collection('users');

  // Be sure to change the title
  app.route('/').get((req, res) => {
    // Change the response to render the Pug template
    res.render('index', {
      title: 'Connected to Database',
      message: 'Please login'
    });
  });

  // Serialization and deserialization here...

  // Be sure to add this...
}).catch(e => {
  app.route('/').get((req, res) => {
    res.render('index', { title: e, message: 'Unable to connect to database' });
  });
});
// app.listen out here...
```

`myDataBase`에서 `deserializeUser` 코드를 주석 해제하고, `done(null, null)`에 `doc`를 포함하도록 수정하세요.

정확히 구현했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implement-the-serialization-of-a-passport-user-5" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트</a>를 확인할 수 있습니다.

# --hints--

데이터베이스 연결이 있어야 합니다.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Connected to Database/gi,
    'You successfully connected to the database!'
  );
```

역직렬화는 이제 데이터베이스를 올바르게 사용해야 하며, `done(null, null)`가 `doc`와 함께 호출되어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /null,\s*doc/gi,
    'The callback in deserializeUser of (null, null) should be altered to (null, doc)'
  );
```
