---
id: 5895f70cf9fc0f352b528e65
title: Passport 설정하기
challengeType: 2
forumTopicId: 301565
dashedName: set-up-passport
---

# --description--

이제 *Passport*를 설정하여 사용자가 계정을 등록하거나 로그인할 수 있게 하세요. Passport 외에도 세션을 처리하기 위해 Express-session을 사용할 것입니다. Express-session에는 다양한 고급 기능이 있지만, 지금은 기본 기능만 사용할 것입니다. 이 미들웨어는 세션 ID를 클라이언트의 쿠키로 저장하고, 서버에서 해당 ID를 사용해 세션 데이터를 접근할 수 있게 합니다. 이렇게 하면 개인 계정 정보를 클라이언트가 사용하는 쿠키에 저장하지 않고, 클라이언트가 인증되었음을 서버에 알리고 서버에 저장된 데이터에 접근할 수 있는 *키*를 유지할 수 있습니다.

`passport@~0.4.1`와 `express-session@~1.17.1`는 이미 설치되어 있으며, 둘 다 `package.json` 파일의 의존성 목록에 포함되어 있습니다.

세션 설정을 구성하고 Passport를 초기화해야 합니다. 먼저, `session`와 `passport` 변수를 각각 `express-session`와 `passport`를 require 하도록 만드세요.

그다음, 다음 옵션을 정의하여 Express 앱이 세션을 사용하도록 설정하세요:

```javascript
app.use(session({
  secret: process.env.SESSION_SECRET,
  resave: true,
  saveUninitialized: true,
  cookie: { secure: false }
}));
```

`SESSION_SECRET`를 `.env` 파일에 추가하고 임의의 값을 지정하는 것을 잊지 마세요. 이 값은 쿠키 암호화에 사용되는 해시를 계산하는 데 쓰입니다!

모든 설정을 마친 후, Express 앱에 `passport.initialize()`와 `passport.session()`를 **use** 하도록 하세요.

설정이 완료되었다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-passport-3" target="_blank" rel="noopener noreferrer nofollow">지금까지 완성된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

Passport와 Express-session이 의존성에 포함되어 있어야 합니다.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport',
    'Your project should list "passport" as a dependency'
  );
  assert.property(
    packJson.dependencies,
    'express-session',
    'Your project should list "express-session" as a dependency'
  );
```

의존성을 올바르게 require 해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport("|')/,
    'You should have required passport'
  );
  assert.match(
    data,
    /require.*("|')express-session("|')/,
    'You should have required express-session'
  );
```

Express 앱이 새 의존성을 사용해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /passport\.initialize/, 'Your express app should use "passport.initialize()"');
  assert.match(data, /passport\.session/, 'Your express app should use "passport.session()"');
```

세션과 세션 비밀이 올바르게 설정되어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /secret *:\s*process\.env(\.SESSION_SECRET|\[(?<q>"|')SESSION_SECRET\k<q>\])/,
    'Your express app should have express-session set up with your secret as process.env.SESSION_SECRET'
  );
```
