---
id: 5895f70bf9fc0f352b528e64
title: 템플릿 엔진의 기능 활용하기
challengeType: 2
forumTopicId: 301567
dashedName: use-a-template-engines-powers
---

# --description--

템플릿 엔진을 사용할 때 가장 큰 기능 중 하나는 서버에서 템플릿 파일로 변수를 전달하여 HTML로 렌더링하기 전 사용할 수 있다는 점입니다.

Pug 파일에서는 `#{variable_name}`와 같이 변수 이름을 참조하여 요소 내 다른 텍스트와 인라인으로 변수를 사용할 수 있고, `p=variable_name`처럼 요소에 공백 없이 등호를 사용해 변수 값을 p 요소의 텍스트로 할당할 수도 있습니다.

Pug는 들여쓰기와 탭을 사용해 중첩된 요소를 표현하고, 아름다운 사이트를 만들기 위해 필요한 코드 양을 줄이는 데 중점을 둡니다.

다음 Pug 코드를 예로 들어 보세요: 
 
```pug
head
  script(type='text/javascript').
    if (foo) bar(1 + 5);
body
  if youAreUsingPug
      p You are amazing
    else
      p Get on it!
```

위 코드는 다음과 같은 HTML을 생성합니다:

```html
<head>
  <script type="text/javascript">
    if (foo) bar(1 + 5);
  </script>
</head>
<body>
  <p>You are amazing</p>
</body>
```

프로젝트에 포함된 `index.pug` 파일은 `title`와 `message` 변수를 사용합니다.

서버에서 이 변수들과 그 값을 객체로 묶어 `res.render` 호출의 두 번째 인수로 전달하세요. `title`에는 `Hello` 값을, `message`에는 `Please log in` 값을 할당하세요.

다음과 같이 보일 것입니다:

```javascript
res.render('index', { title: 'Hello', message: 'Please log in' });
```

이제 페이지를 새로 고침하면 `index.pug` 파일에 배치된 대로 해당 값들이 뷰에 올바른 위치에 렌더링된 것을 볼 수 있습니다!

정확히 했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#use-a-template-engines-power-2" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트</a>를 참고할 수 있습니다.

# --hints--

Pug가 변수를 올바르게 렌더링하는지 확인하세요.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /pug-variable("|')>Please log in/gi,
    'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
  );
```
