---
id: 5895f70cf9fc0f352b528e66
title: 사용자 객체 직렬화
challengeType: 2
forumTopicId: 301563
dashedName: serialization-of-a-user-object
---

# --description--

직렬화와 역직렬화는 인증과 관련된 중요한 개념입니다. 객체를 직렬화한다는 것은 그 내용을 작은 *키*로 변환하는 것을 의미하며, 이 키는 다시 원래 객체로 역직렬화할 수 있습니다. 이렇게 하면 매번 새로운 페이지 요청 시 사용자 이름과 비밀번호 같은 인증 데이터를 보내지 않고도 누가 서버와 통신했는지 알 수 있습니다.

이를 제대로 설정하려면 직렬화 함수와 역직렬화 함수가 필요합니다. Passport에서는 다음과 같이 만들 수 있습니다:

```javascript
passport.serializeUser(cb);
passport.deserializeUser(cb);
```

`serializeUser`에 전달된 콜백 함수는 두 개의 인수와 함께 호출됩니다: 전체 사용자 객체와 passport에서 사용하는 콜백입니다.

콜백은 두 개의 인수를 기대합니다: 오류가 있으면 그 오류, 그리고 콜백에서 반환해야 하는 사용자를 식별하는 고유 키입니다. 객체 내에서 사용자의 `_id`을 사용할 것입니다. 이 값은 MongoDB에서 생성되므로 고유함이 보장됩니다.

마찬가지로 `deserializeUser`는 두 개의 인수와 함께 호출됩니다: 고유 키와 콜백 함수입니다.

이 콜백은 두 개의 인수를 기대합니다: 오류가 있으면 그 오류, 그리고 전체 사용자 객체입니다. 전체 사용자 객체를 얻으려면 아래와 같이 Mongo `_id`를 쿼리 검색하세요:

```javascript
passport.serializeUser((user, done) => {
  done(null, user._id);
});

passport.deserializeUser((id, done) => {
  myDataBase.findOne({ _id: new ObjectID(id) }, (err, doc) => {
    done(null, null);
  });
});
```

위 두 함수를 서버에 추가하세요. `ObjectID` 클래스는 `mongodb` 패키지에서 가져옵니다. `mongodb@~3.6.0`는 이미 의존성으로 추가되어 있습니다. 이 클래스를 다음과 같이 선언하세요:

```javascript
const { ObjectID } = require('mongodb');
```

`deserializeUser`는 데이터베이스 연결을 설정할 때까지 오류를 발생시킬 것입니다. 그러니 지금은 `myDatabase.findOne` 호출을 주석 처리하고 `done(null, null)` 콜백 함수 내에서 `deserializeUser`만 호출하세요.

정확히 했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#serialization-of-a-user-object-4" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

사용자 객체를 올바르게 직렬화해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.serializeUser/gi,
    'You should have created your passport.serializeUser function'
  );
  assert.match(
    data,
    /null,\s*user._id/gi,
    'There should be a callback in your serializeUser with (null, user._id)'
  );
```

사용자 객체를 올바르게 역직렬화해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.deserializeUser/gi,
    'You should have created your passport.deserializeUser function'
  );
  assert.match(
    data,
    /null,\s*null/gi,
    'There should be a callback in your deserializeUser with (null, null) for now'
  );
```

MongoDB가 의존성으로 포함되어야 합니다.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'mongodb',
    'Your project should list "mongodb" as a dependency'
  );
```

ObjectId를 포함하여 Mongodb를 제대로 가져와야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')mongodb\1/gi,
    'You should have required mongodb'
  );
  assert.match(
    data,
    /new ObjectID.*id/gi,
    'Even though the block is commented out, you should use new ObjectID(id) for when we add the database'
  );
```
