---
id: 5895f70df9fc0f352b528e69
title: Passport 전략 사용 방법
challengeType: 2
forumTopicId: 301555
dashedName: how-to-use-passport-strategies
---

# --description--

제공된 `index.pug` 파일에는 로그인 폼이 있습니다. 인라인 JavaScript `if showLogin` 때문에 숨겨져 있으며, 그 뒤에 폼이 들여쓰기 되어 있습니다.

해당 페이지의 `res.render`에 객체 `showLogin: true`에 새 변수를 추가하세요. 페이지를 새로 고침하면 폼이 보일 것입니다! 이 폼은 `/login`에 **POST** 하도록 설정되어 있습니다. 따라서 POST 요청을 받아 인증을 처리할 위치가 바로 여기입니다.

이번 도전 과제에서는 POST 요청을 받기 위해 `/login` 라우트를 추가해야 합니다. 이 라우트에서 인증하려면 응답을 보내기 전에 미들웨어를 추가해야 합니다. 이는 응답 전에 미들웨어를 인수로 전달하는 방식으로 구현합니다. 사용할 미들웨어는 `passport.authenticate('local')`입니다.

`passport.authenticate`은 `{ failureRedirect: '/' }` 같은 인수를 옵션으로 받을 수 있는데, 매우 유용하니 꼭 추가하세요. 미들웨어를 사용한 후(인증 미들웨어가 통과한 경우에만 호출됨) 사용자에게 `/profile`로 리다이렉트하는 응답을 추가하세요. 해당 라우트도 추가하고 `profile.pug` 뷰를 렌더링하도록 만드세요.

인증이 성공하면 사용자 객체가 `req.user`에 저장됩니다.

이 시점에서 폼에 사용자 이름과 비밀번호를 입력하면 홈 페이지 `/`로 리다이렉트되고, 현재 등록되지 않은 사용자는 로그인할 수 없으므로 서버 콘솔에 `'User {USERNAME} attempted to log in.'`가 표시될 것입니다.

정확히 구현했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-use-passport-strategies-7" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트</a>를 확인할 수 있습니다.

# --hints--

모든 단계는 `server.js`에 올바르게 구현되어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /showLogin:( |)true/,
    'You should be passing the variable "showLogin" as true to your render function for the homepage'
  );
  assert.match(
    data,
    /failureRedirect:( |)('|")\/('|")/,
    'Your code should include a failureRedirect to the "/" route'
  );
  assert.match(
    data,
    /login[^]*post[^]*local/,
    'You should have a route for login which accepts a POST and passport.authenticates local'
  );
```

`/login`에 대한 POST 요청은 `/`로 올바르게 리다이렉트해야 합니다.

```js
  const url = new URL("/login", code);
  const res = await fetch(url, { method: 'POST' });
  const data = await res.text();
  assert.match(
    data,
    /Looks like this page is being rendered from Pug into HTML!/,
    'A login attempt at this point should redirect to the homepage since we do not have any registered users'
  );
```
