---
id: 5895f70df9fc0f352b528e68
title: 인증 전략
challengeType: 2
forumTopicId: 301547
dashedName: authentication-strategies
---

# --description--

전략은 사용자를 인증하는 방법입니다. 사용자가 먼저 등록한 경우 로컬에 저장된 정보를 기반으로 인증하거나 Google, GitHub 같은 다양한 제공자를 통해 인증할 수 있습니다. 이 프로젝트에서는 Passport 미들웨어를 사용할 것입니다. Passport는 사용자 이름과 비밀번호, GitHub, Google 등을 사용한 인증을 지원하는 포괄적인 전략 세트를 제공합니다.

`passport-local@~1.0.0`가 이미 의존성으로 추가되어 있습니다. 다음과 같이 서버에 추가하세요:

```javascript
const LocalStrategy = require('passport-local');
```

몇 가지 설정이 정의된 인스턴스화된 `LocalStrategy` 객체를 사용하도록 passport에 지시하세요. 이 부분(그리고 이후 모든 부분)은 데이터베이스 연결 안에 포함되어야 합니다. 왜냐하면 이 설정이 데이터베이스에 의존하기 때문입니다!:

```javascript
passport.use(new LocalStrategy((username, password, done) => {
  myDataBase.findOne({ username: username }, (err, user) => {
    console.log(`User ${username} attempted to log in.`);
    if (err) return done(err);
    if (!user) return done(null, false);
    if (password !== user.password) return done(null, false);
    return done(null, user);
  });
}));
```

이것은 로컬에서 사용자를 인증하려 할 때 사용할 과정을 정의합니다. 먼저 입력한 사용자 이름으로 데이터베이스에서 사용자를 찾으려 시도합니다. 그다음 비밀번호가 일치하는지 확인합니다. 마지막으로, (예: 비밀번호가 틀린 경우와 같은) 확인한 오류가 없으면 `user` 객체가 반환되어 인증됩니다.

많은 전략이 서로 다른 설정으로 구성됩니다. 일반적으로 해당 전략 저장소의 README를 참고하면 쉽게 설정할 수 있습니다. 좋은 예가 GitHub 전략인데, 사용자 이름이나 비밀번호를 걱정할 필요가 없습니다. 사용자가 GitHub 인증 페이지로 보내져 인증을 진행하기 때문입니다. 사용자가 로그인되어 동의하면 GitHub가 프로필을 반환해 사용 가능합니다.

다음 단계에서는 폼 데이터를 기반으로 사용자를 검증하기 위해 실제로 인증 전략을 호출하는 방법을 설정할 것입니다.

정확하다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#authentication-strategies-6" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

Passport-local이 의존성으로 있어야 합니다.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport-local',
    'Your project should list "passport-local " as a dependency'
  );
```

Passport-local이 올바르게 require되고 설정되어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport-local("|')/,
    'You should have required passport-local'
  );
  assert.match(
    data,
    /new LocalStrategy/,
    'You should have told passport to use a new strategy'
  );
  assert.match(
    data,
    /findOne/,
    'Your new local strategy should use the findOne query to find a username based on the inputs'
  );
```
