---
id: 5895f70ef9fc0f352b528e6b
title: 프로필을 만드는 방법
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

이제 `/profile`에 접근하는 사용자가 인증된 것을 확인할 수 있으니, 페이지에 있는 `req.user`에 포함된 정보를 사용할 수 있습니다.

프로필 뷰의 `username` 메서드에 두 번째 인수로 `req.user.username` 속성과 `render` 값을 포함하는 객체를 전달하세요.

그다음 `profile.pug` 뷰로 가서, 기존 `h1` 요소 아래, 같은 들여쓰기 수준에 다음 줄을 추가하세요:

```pug
h2.center#welcome Welcome, #{username}!
```

이 코드는 `h2` 요소를 만들며, 클래스는 `center`, 아이디는 `welcome`이고, 텍스트 `Welcome, `과 사용자 이름이 뒤따릅니다.

또한 `profile.pug`에, 사용자를 인증 해제하는 로직을 호스트할 `/logout` 라우트를 가리키는 링크를 추가하세요:

```pug
a(href='/logout') Logout
```

정확히 작성했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">지금까지 완성된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

`/profile`에 Pug 렌더 변수로 올바르게 추가해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );
```
