---
id: 5895f700f9fc0f352b528e63
title: 템플릿 엔진 설정하기
challengeType: 2
forumTopicId: 301564
dashedName: set-up-a-template-engine
---

# --description--

이 도전 과제들은 다음 방법 중 하나를 사용해 코드를 작성하는 것을 포함합니다:

- <a href="https://github.com/freeCodeCamp/boilerplate-advancednode/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 도전 과제를 완료하세요.
- 원하는 사이트 빌더를 사용해 프로젝트를 완료하세요. 반드시 GitHub 저장소의 모든 파일을 포함해야 합니다.

템플릿 엔진은 앱에서 정적 템플릿 파일(예: *Pug*로 작성된 파일)을 사용할 수 있게 해줍니다. 실행 시 템플릿 엔진은 템플릿 파일 내 변수들을 서버가 제공하는 실제 값으로 대체합니다. 그런 다음 템플릿을 정적 HTML 파일로 변환해 클라이언트에 전송합니다. 이 방법은 HTML 페이지 디자인을 쉽게 하고, 클라이언트에서 API 호출 없이도 페이지에 변수를 표시할 수 있게 합니다.

`pug@~3.0.0`는 이미 설치되어 있으며 `package.json` 파일의 의존성으로 등록되어 있습니다.

Express는 어떤 템플릿 엔진을 사용하는지 알아야 합니다. `set` 메서드를 사용해 `pug`를 `view engine` 속성 값으로 할당하세요:

```javascript
app.set('view engine', 'pug');
```

그다음, 또 다른 `set` 메서드를 추가해 `views`의 `app` 속성을 `./views/pug` 디렉터리로 지정하세요. 이렇게 하면 Express가 모든 뷰를 해당 디렉터리를 기준으로 렌더링합니다.

마지막으로, 홈 페이지 라우트에서 `res.render()`를 사용해 첫 번째 인수로 `index`를 전달하세요. 그러면 `pug` 템플릿이 렌더링됩니다.

모든 것이 계획대로 되었다면, 앱 홈 페이지가 더 이상 빈 화면이 아니고 Pug 템플릿이 성공적으로 렌더링되었다는 메시지를 표시할 것입니다!

정상적으로 작동한다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-a-template-engine-1" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

Pug가 의존성으로 포함되어 있어야 합니다.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'pug',
    'Your project should list "pug" as a dependency'
  );
```

뷰 엔진은 Pug여야 합니다.

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.['view engine'], "pug");
```

애플리케이션의 `views` 속성을 `./views/pug`로 설정해야 합니다.

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.views, "./views/pug");
```

응답에서 인덱스 페이지를 렌더링하려면 올바른 ExpressJS 메서드를 사용하세요.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /FCC Advanced Node and Express/gi,
        'You successfully rendered the Pug template!'
      );
```

Pug가 정상 작동해야 합니다.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /pug-success-message/gi,
        'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
      );
```
