---
id: 5895f70df9fc0f352b528e6a
title: 새 미들웨어 만들기
challengeType: 2
forumTopicId: 301551
dashedName: create-new-middleware
---

# --description--

현재 상태로는 사용자가 인증 여부와 상관없이 URL을 입력해 `/profile`에 바로 접근할 수 있습니다. 프로필 페이지를 렌더링하기 전에 사용자가 인증되었는지 확인해서 이를 방지하고 싶습니다. 이런 경우에 미들웨어를 만드는 것이 가장 적합합니다.

이번 도전 과제는 `ensureAuthenticated(req, res, next)`라는 미들웨어 함수를 만드는 것입니다. 이 함수는 Passport의 `isAuthenticated` 메서드를 `request`에서 호출해 `req.user`가 정의되어 있는지 확인합니다. 만약 정의되어 있다면 `next()`를 호출해야 합니다. 그렇지 않으면 로그인할 수 있도록 홈페이지로 리다이렉트하는 응답을 요청에 보내면 됩니다.

이 미들웨어 구현 예시는 다음과 같습니다:

```javascript
function ensureAuthenticated(req, res, next) {
  if (req.isAuthenticated()) {
    return next();
  }
  res.redirect('/');
};
```

위 미들웨어 함수를 만든 후, GET 요청의 인수 앞에 `ensureAuthenticated`를 프로필 페이지 요청에 대한 미들웨어로 전달하세요:

```javascript
app
 .route('/profile')
 .get(ensureAuthenticated, (req,res) => {
    res.render('profile');
 });
```

정확히 구현했다고 생각되면 페이지를 제출하세요. 오류가 발생한다면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#create-new-middleware-8" target="_blank" rel="noopener noreferrer nofollow">지금까지 완성된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

미들웨어 `ensureAuthenticated`를 구현하고 `/profile` 라우트에 연결하세요.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /ensureAuthenticated[^]*req.isAuthenticated/,
    'Your ensureAuthenticated middleware should be defined and utilize the req.isAuthenticated function'
  );
  assert.match(
    data,
    /profile[^]*get[^]*ensureAuthenticated/,
    'Your ensureAuthenticated middleware should be attached to the /profile route'
  );
```

인증되지 않은 GET 요청이 `/profile`에 들어오면 `/`으로 올바르게 리다이렉트되어야 합니다.

```js
  const url = new URL("/profile", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/,
    'An attempt to go to the profile at this point should redirect to the homepage since we are not logged in'
  );
```
