---
id: 58966a17f9fc0f352b528e6d
title: 새 사용자 등록
challengeType: 2
forumTopicId: 301561
dashedName: registration-of-new-users
---

# --description--

이제 사이트에서 새 사용자가 계정을 등록할 수 있도록 하세요. 홈 페이지의 `res.render`에 전달되는 객체에 새 변수를 추가하세요 - `showRegistration: true`. 페이지를 새로 고침하면 `index.pug` 파일에 이미 만들어진 등록 양식이 보여야 합니다. 이 양식은 `/register`에 **POST**하도록 설정되어 있으니, 해당 라우트를 만들고 아래 논리에 따라 사용자 객체를 데이터베이스에 추가하세요.

등록 라우트의 논리는 다음과 같아야 합니다:

1. 새 사용자 등록
2. 새 사용자 인증
3. `/profile`로 리다이렉트

1단계의 논리는 다음과 같아야 합니다:

1. `findOne`로 데이터베이스 쿼리
2. 오류가 있으면 오류와 함께 `next` 호출
3. 사용자가 반환되면 홈으로 리다이렉트
4. 사용자가 없고 오류가 없으면 사용자 이름과 비밀번호로 데이터베이스에 `insertOne`. 오류가 없으면 `next`를 호출해 2단계인 새 사용자 인증으로 이동하세요. 이 인증 논리는 이미 `POST /login` 라우트에 작성되어 있습니다.

```js
app.route('/register')
  .post((req, res, next) => {
    myDataBase.findOne({ username: req.body.username }, (err, user) => {
      if (err) {
        next(err);
      } else if (user) {
        res.redirect('/');
      } else {
        myDataBase.insertOne({
          username: req.body.username,
          password: req.body.password
        },
          (err, doc) => {
            if (err) {
              res.redirect('/');
            } else {
              // The inserted document is held within
              // the ops property of the doc
              next(null, doc.ops[0]);
            }
          }
        )
      }
    })
  },
    passport.authenticate('local', { failureRedirect: '/' }),
    (req, res, next) => {
      res.redirect('/profile');
    }
  );
```

정확하다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#registration-of-new-users-11" target="_blank" rel="noopener noreferrer nofollow">지금까지 완성된 프로젝트를 확인할 수 있습니다</a>.

**참고:** 이 시점부터 *picture-in-picture* 브라우저 사용과 관련된 문제가 발생할 수 있습니다. 에디터 내에서 앱 미리보기를 제공하는 온라인 IDE를 사용 중이라면, 이 미리보기를 새 탭에서 여는 것을 권장합니다.

# --hints--

홈 페이지에 `/register` 라우트가 있어야 하며 등록 양식을 표시해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /showRegistration:( |)true/gi,
    'You should be passing the variable showRegistration as true to your render function for the homepage'
  );
  assert.match(
    data,
    /register[^]*post[^]*findOne[^]*username:( |)req.body.username/gi,
    'You should have a route that accepts a POST request on /register that queries the db with findOne and the query being username: req.body.username'
  );
```

등록이 작동해야 합니다.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttp = new XMLHttpRequest();
  xhttp.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      test(this);
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttp.open('POST', url + '/register', true);
  xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
  xhttp.send(`username=${user}&password=${user}`);
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Register should work, and redirect successfully to the profile.'
    );
  }
```

로그인이 작동해야 합니다.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttpReg = new XMLHttpRequest();
  xhttpReg.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      login();
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttpReg.open('POST', url + '/register', true);
  xhttpReg.setRequestHeader(
    'Content-type',
    'application/x-www-form-urlencoded'
  );
  xhttpReg.send(`username=${user}&password=${user}`);
  function login() {
    const xhttp = new XMLHttpRequest();
    xhttp.onreadystatechange = function () {
      if (this.readyState == 4 && this.status == 200) {
        test(this);
      } else {
        throw new Error(`${this.status} ${this.statusText}`);
      }
    };
    xhttp.open('POST', url + '/login', true);
    xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
    xhttp.send(`username=${user}&password=${user}`);
  }
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Login should work if previous test was done successfully and redirect successfully to the profile.'
    );
    assert.match(
      data,
      new RegExp(user, 'g'),
      'The profile should properly display the welcome to the user logged in'
    );
  }
```

로그아웃이 작동해야 합니다.

```js
  $.ajax({
    url: code + '/logout',
    type: 'GET',
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(data, /Home/gi, 'Logout should redirect to home');
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

로그아웃 후에는 프로필이 더 이상 작동하지 않아야 합니다.

```js
  $.ajax({
    url: code + '/profile',
    type: 'GET',
    crossDomain: true,
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(
        data,
        /Home/gi,
        'Profile should redirect to home when we are logged out now again'
      );
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```
