---
id: 589690e6f9fc0f352b528e6e
title: 모듈로 프로젝트 정리하기
challengeType: 2
forumTopicId: 301549
dashedName: clean-up-your-project-with-modules
---

# --description--

지금은 모든 것이 `server.js` 파일에 들어 있습니다. 이렇게 하면 관리하기 어렵고 확장성이 떨어지는 코드가 될 수 있습니다. `routes.js`과 `auth.js` 두 개의 새 파일을 만드세요.

두 파일 모두 다음 코드를 포함해서 시작하세요:

```js
module.exports = function (app, myDataBase) {

}
```

이제 서버 파일 상단에서 이 파일들을 다음과 같이 require 하세요: `const routes = require('./routes.js');` 데이터베이스와 성공적으로 연결한 직후에 각각을 다음과 같이 인스턴스화하세요: `routes(app, myDataBase)`

마지막으로, 서버에 있는 모든 라우트를 새 파일로 옮기고 서버 파일에서는 삭제하세요. 라우팅을 위해 특별히 만든 `ensureAuthenticated` 함수도 옮기세요. 그리고 `const passport = require('passport');` 파일의 export 라인 위, 가장 상단에 `routes.js` 같은 사용되는 의존성을 올바르게 추가해야 합니다.

오류가 더 이상 없고 서버 파일에 라우팅이 남아 있지 않을 때까지 계속 추가하세요 (**catch 블록의 라우트는 제외**).

`auth.js` 파일에서도 인증과 관련된 모든 것들, 예를 들어 직렬화와 로컬 전략 설정 등을 서버 파일에서 지우고 옮기세요. 의존성도 추가하고 서버에서 같은 위치에 `auth(app, myDataBase)` 을 호출하는 것을 잊지 마세요.

모두 제대로 했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#clean-up-your-project-with-modules-2" target="_blank" rel="noopener noreferrer nofollow">완성된 프로젝트 예시</a>를 확인할 수 있습니다.

# --hints--

모듈이 있어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\s*\(('|")\.\/routes(\.js)?\1\)/gi,
    'You should have required your new files'
  );
  assert.match(
    data,
    /client\s*\.db[^]*routes/gi,
    'Your new modules should be called after your connection to the database'
  );
```
