---
id: 58965611f9fc0f352b528e6c
title: 사용자 로그아웃 처리하기
challengeType: 2
forumTopicId: 301560
dashedName: logging-a-user-out
---

# --description--

로그아웃 로직을 만드는 것은 쉽습니다. 라우트는 단순히 사용자의 인증을 해제하고, 어떤 뷰도 렌더링하지 않고 홈 페이지로 리다이렉트하면 됩니다.

passport에서는 사용자의 인증을 해제하는 것이 `req.logout()`를 호출한 후 리다이렉트하는 것만큼 간단합니다. 이를 위해 `/logout` 라우트를 추가하세요:

```js
app.route('/logout')
  .get((req, res) => {
    req.logout();
    res.redirect('/');
});
```

404 페이지(없는 페이지)를 처리하지 않는다는 점을 눈치챘을 겁니다. Node에서 이를 처리하는 일반적인 방법은 다음 미들웨어를 사용하는 것입니다. 다른 모든 라우트 뒤에 이 코드를 추가하세요:

```js
app.use((req, res, next) => {
  res.status(404)
    .type('text')
    .send('Not Found');
});
```

코드가 제대로 됐다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#logging-a-user-out-10" target="_blank" rel="noopener noreferrer nofollow">지금까지 완성된 프로젝트</a>를 확인할 수 있습니다.

# --hints--

`req.logout()`는 `/logout` 라우트에서 호출되어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /req.logout/gi,
    'You should be calling req.logout() in your /logout route'
  );
```

`/logout`는 홈 페이지로 리다이렉트해야 합니다.

```js
  const url = new URL("/logout", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/gi,
    'When a user logs out they should be redirected to the homepage'
  );
```
