---
id: 589a8eb3f9fc0f352b528e72
title: 소셜 인증 구현 III
challengeType: 2
forumTopicId: 301558
dashedName: implementation-of-social-authentication-iii
---

# --description--

전략의 마지막 부분은 GitHub에서 반환된 프로필을 처리하는 것입니다. 사용자의 데이터베이스 객체가 있으면 불러오고, 없으면 새로 만들어 프로필에서 필드를 채운 다음 사용자의 객체를 반환해야 합니다. GitHub는 각 프로필 내에서 고유한 *id*를 제공하며, 이를 사용해 직렬화할 사용자를 검색할 수 있습니다(이미 구현됨). 아래는 프로젝트에서 사용할 수 있는 예시 구현으로, 새 전략의 두 번째 인수인 함수 내에서 `console.log(profile);` 바로 아래에 넣으세요:

```js
myDataBase.findOneAndUpdate(
  { id: profile.id },
  {
    $setOnInsert: {
      id: profile.id,
      username: profile.username,
      name: profile.displayName || 'John Doe',
      photo: profile.photos[0].value || '',
      email: Array.isArray(profile.emails)
        ? profile.emails[0].value
        : 'No public email',
      created_on: new Date(),
      provider: profile.provider || ''
    },
    $set: {
      last_login: new Date()
    },
    $inc: {
      login_count: 1
    }
  },
  { upsert: true, new: true },
  (err, doc) => {
    return cb(null, doc.value);
  }
);
```

`findOneAndUpdate`은 객체를 검색하고 업데이트할 수 있게 해줍니다. 객체가 없으면 삽입되어 콜백 함수에서 사용할 수 있게 됩니다. 이 예시에서는 항상 `last_login`를 설정하고, `login_count`을 `1`만큼 증가시키며, 새 객체(새 사용자)가 삽입될 때만 대부분의 필드를 채웁니다. 기본값 사용에 주목하세요. 반환된 프로필에 모든 정보가 채워져 있지 않거나 사용자가 비공개로 설정할 수도 있기 때문입니다. 이 경우 오류를 방지하도록 처리합니다.

이제 앱에 로그인할 수 있을 것입니다. 시도해 보세요!

정확히 구현했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-iii-5" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

GitHub 전략 설정이 완료되어야 합니다.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /GitHubStrategy[^]*myDataBase/gi,
    'Strategy should use now use the database to search for the user'
  );
  assert.match(
    data,
    /GitHubStrategy[^]*cb/gi,
    'Strategy should return the callback function "cb"'
  );
```
