---
id: 589a69f5f9fc0f352b528e70
title: 소셜 인증 구현
challengeType: 2
forumTopicId: 301559
dashedName: implementation-of-social-authentication
---

# --description--

이런 종류의 인증은 다음과 같은 흐름으로 진행됩니다:

1. 사용자가 특정 전략(e.g. GitHub)을 사용해 인증하기 위해 라우트로 보내는 버튼이나 링크를 클릭합니다.
2. 라우트가 `passport.authenticate('github')`를 호출하여 사용자를 GitHub로 리다이렉트합니다.
3. 사용자가 도착하는 GitHub 페이지에서 로그인하지 않은 경우 로그인할 수 있습니다. 그런 다음 애플리케이션이 프로필에 접근하는 것을 승인할지 묻습니다.
4. 승인이 되면 사용자는 특정 콜백 URL로 프로필과 함께 애플리케이션으로 돌아옵니다.
5. 이제 인증된 상태이며, 애플리케이션은 이전에 방문한 프로필인지 확인하거나 데이터베이스에 저장해야 합니다.

OAuth 전략은 최소한 *클라이언트 ID*와 *클라이언트 시크릿*이 필요합니다. 이는 서비스가 인증 요청이 어디서 왔는지, 그리고 유효한지 확인하는 방법입니다. 이 값들은 GitHub 같은 인증을 구현하려는 사이트에서 얻으며, 애플리케이션마다 고유합니다—**절대 공유해서는 안 되며** 공개 저장소에 올리거나 코드에 직접 작성해서도 안 됩니다. 일반적으로 `.env` 파일에 넣고 다음과 같이 참조합니다: `process.env.GITHUB_CLIENT_ID`. 이번 도전 과제에서는 GitHub 전략을 사용할 것입니다.

GitHub에서 *클라이언트 ID와 시크릿*을 얻으려면 <a href="https://www.freecodecamp.org/news/how-to-set-up-a-github-oauth-application/" target="_blank" rel="noopener noreferrer nofollow">이 지침</a>을 따르세요. 홈페이지 URL은 프로젝트 코드 URL이 아닌 자신의 홈페이지 URL로 설정하고, 콜백 URL은 홈페이지 URL 끝에 `/auth/github/callback`를 붙여서 설정하세요. 클라이언트 ID와 시크릿은 프로젝트의 `.env` 파일에 `GITHUB_CLIENT_ID`와 `GITHUB_CLIENT_SECRET`로 저장하세요.

`routes.js` 파일에서 `showSocialAuth: true` 뒤에 홈페이지 라우트에 `showRegistration: true`를 추가하세요. 이제 GET 요청을 받는 두 개의 라우트 `/auth/github`과 `/auth/github/callback`을 만드세요. 첫 번째는 passport를 호출해 `'github'`를 인증만 해야 합니다. 두 번째는 passport를 호출해 `'github'`를 인증하며 실패 시 `/`로 리다이렉트하고, 성공하면 `/profile`로 리다이렉트해야 합니다(이전 프로젝트와 유사합니다).

`/auth/github/callback`가 어떻게 보여야 하는지 예시는 일반 로그인 처리 방식과 비슷합니다:

```js
app.route('/login')
  .post(passport.authenticate('local', { failureRedirect: '/' }), (req,res) => {
    res.redirect('/profile');
  });
```

정확하다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-3" target="_blank" rel="noopener noreferrer nofollow">지금까지의 프로젝트를 확인할 수 있습니다</a>.

# --hints--

라우트 `/auth/github`가 올바른지 확인하세요.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
          data.replace(/\s/g, ''),
          /passport.authenticate.*?github/g,
          'Route auth/github should only call passport.authenticate with github'
        );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```

라우트 `/auth/github/callback`가 올바른지 확인하세요.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
        data.replace(/\s/g, ''),
        /failureRedirect:("|')\/\1/g,
        'Route auth/github/callback should accept a get request and call passport.authenticate for github with a failure redirect to home'
      );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github/callback');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github/callback"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```
