---
id: 589a69f5f9fc0f352b528e71
title: 소셜 인증 구현 II
challengeType: 2
forumTopicId: 301557
dashedName: implementation-of-social-authentication-ii
---

# --description--

GitHub 인증 설정의 마지막 단계는 전략 자체를 만드는 것입니다. `passport-github@~1.1.0`는 이미 의존성으로 추가되어 있으니, `auth.js` 파일에서 `GitHubStrategy`로 이렇게 require 하세요: `const GitHubStrategy = require('passport-github').Strategy;`. 환경 변수를 사용하도록 `dotenv`도 require 하고 설정하는 것을 잊지 마세요.

GitHub 전략을 설정하려면, Passport에게 인스턴스화된 `GitHubStrategy`를 사용하라고 알려야 합니다. 이 `GitHubStrategy`는 두 개의 인수를 받는데, 하나는 `clientID`, `clientSecret`, `callbackURL`를 포함하는 객체이고, 다른 하나는 사용자가 인증에 성공했을 때 호출되는 함수입니다. 이 함수는 사용자가 새로운지 여부와 사용자 데이터베이스 객체에 처음 저장할 필드를 결정합니다. 이는 많은 전략에서 공통적이지만, 특정 전략의 GitHub README에 명시된 대로 더 많은 정보가 필요할 수도 있습니다. 예를 들어, Google은 요청하는 정보 종류를 결정하는 *scope*도 필요하며, 사용자가 해당 접근을 승인하도록 요청합니다.

현재 구현 중인 전략은 GitHub 계정과 OAuth 2.0 토큰을 사용해 사용자를 인증합니다. 애플리케이션 생성 시 얻은 클라이언트 ID와 비밀키는 전략 생성 시 옵션으로 전달됩니다. 이 전략은 또한 `verify` 콜백이 필요한데, 이 콜백은 액세스 토큰과 선택적 새로 고침 토큰, 그리고 인증된 사용자의 GitHub 프로필이 담긴 `profile`을 받습니다. `verify` 콜백은 인증을 완료하기 위해 사용자 정보를 제공하며 `cb`를 호출해야 합니다.

현재 시점에서 새 전략은 다음과 같이 보여야 합니다:

```js
passport.use(new GitHubStrategy({
  clientID: process.env.GITHUB_CLIENT_ID,
  clientSecret: process.env.GITHUB_CLIENT_SECRET,
  callbackURL: /*INSERT CALLBACK URL ENTERED INTO GITHUB HERE*/
},
  function(accessToken, refreshToken, profile, cb) {
    console.log(profile);
    //Database logic here with callback containing your user object
  }
));
```

아직 인증이 성공하지는 않을 것이고, 데이터베이스 로직과 콜백이 없으면 실제로 오류가 발생할 것입니다. 하지만 시도하면 GitHub 프로필이 콘솔에 기록되는 것을 볼 수 있을 것입니다!

정확히 구현했다고 생각되면 페이지를 제출하세요. 오류가 발생한다면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-ii-4" target="_blank" rel="noopener noreferrer nofollow">지금까지 완성된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

`passport-github` 의존성이 추가되어야 합니다.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport-github',
    'Your project should list "passport-github" as a dependency'
  );
```

`passport-github`를 require 해야 합니다.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport-github("|')/gi,
    'You should have required passport-github'
  );
```

GitHub 전략이 지금까지 올바르게 설정되어야 합니다.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport\.use.*new GitHubStrategy/gis,
    'Passport should use a new GitHubStrategy'
  );
  assert.match(
    data,
    /callbackURL:\s*("|').*("|')/gi,
    'You should have a callbackURL'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_SECRET|\[(?<q>"|')GITHUB_CLIENT_SECRET\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_SECRET'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_ID|\[(?<q>"|')GITHUB_CLIENT_ID\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_ID'
  );
```
