---
id: 589fc831f9fc0f352b528e77
title: Socket.IO로 인증하기
challengeType: 2
forumTopicId: 301548
dashedName: authentication-with-socket-io
---

# --description--

현재 웹 소켓에 누가 연결되어 있는지 알 수 없습니다. `req.user`에는 사용자 객체가 포함되어 있지만, 이는 사용자가 웹 서버와 상호작용할 때만 해당하며, 웹 소켓에서는 `req` (요청)이 없기 때문에 사용자 데이터를 얻을 수 없습니다. 웹 소켓에 누가 연결되어 있는지 알 수 있는 문제를 해결하는 한 가지 방법은 passport 세션이 포함된 쿠키를 파싱하고 디코딩한 다음 역직렬화하여 사용자 객체를 얻는 것입니다. 다행히도, 이 복잡한 작업을 간단하게 만들어 주는 NPM 패키지가 있습니다!

`passport.socketio@~3.7.0`, `connect-mongo@~3.2.0`, 그리고 `cookie-parser@~1.4.5`가 이미 의존성으로 추가되어 있습니다. 각각 `passportSocketIo`, `MongoStore`, 그리고 `cookieParser`로 require 하세요. 또한, 이전에 require 했던 `express-session`에서 새로운 메모리 저장소를 초기화해야 합니다. 다음과 같이 보일 것입니다:

```js
const MongoStore = require('connect-mongo')(session);
const URI = process.env.MONGO_URI;
const store = new MongoStore({ url: URI });
```

이제 Socket.IO가 이를 사용하도록 설정하고 옵션을 지정하기만 하면 됩니다. 기존 소켓 코드 전에 추가하고 기존 연결 리스너 안에는 넣지 마세요. 서버에서는 다음과 같이 보일 것입니다:

```js
io.use(
  passportSocketIo.authorize({
    cookieParser: cookieParser,
    key: 'express.sid',
    secret: process.env.SESSION_SECRET,
    store: store,
    success: onAuthorizeSuccess,
    fail: onAuthorizeFail
  })
);
```

Socket.IO용 Passport 인증을 구성하는 것은 API용 `session` 미들웨어를 구성한 방식과 매우 비슷하다는 점에 주목하세요. 이는 동일한 인증 방식을 사용하기 때문입니다 — 쿠키에서 세션 ID를 가져와 검증하는 방식입니다.

이전에 `session` 미들웨어를 구성할 때 세션용 쿠키 이름(`key`)을 명시적으로 설정하지 않았습니다. 이는 `session` 패키지가 기본값을 사용했기 때문입니다. 이제 쿠키에서 동일한 값을 접근해야 하는 다른 패키지를 추가했으므로, 두 구성 객체 모두에 `key` 값을 명시적으로 설정해야 합니다.

Socket.IO 키와 일치하는 `key` 쿠키 이름을 `session` 미들웨어에 반드시 추가하세요. 또한, `store` 참조를 옵션에 `saveUninitialized: true`을 설정한 근처에 추가하세요. 이는 Socket.IO가 어떤 세션과 연관 지어야 하는지 알려주기 위해 필요합니다.

<hr>

이제 `success`와 `fail` 콜백 함수를 정의하세요:

```js
function onAuthorizeSuccess(data, accept) {
  console.log('successful connection to socket.io');

  accept(null, true);
}

function onAuthorizeFail(data, message, error, accept) {
  if (error) throw new Error(message);
  console.log('failed connection to socket.io:', message);
  accept(null, false);
}
```

사용자 객체는 이제 `socket.request.user`로 소켓 객체에서 접근할 수 있습니다. 예를 들어, 다음을 추가할 수 있습니다:

```js
console.log('user ' + socket.request.user.username + ' connected');
```

누가 연결되었는지 서버 콘솔에 기록될 것입니다!

정확히 했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#authentication-with-socketio-9" target="_blank" rel="noopener noreferrer nofollow">지금까지의 프로젝트를 확인할 수 있습니다</a>.

# --hints--

`passport.socketio`는 의존성이어야 합니다.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport.socketio',
    'Your project should list "passport.socketio" as a dependency'
  );
```

`cookie-parser`는 의존성이어야 합니다.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'cookie-parser',
    'Your project should list "cookie-parser" as a dependency'
  );
```

passportSocketIo가 제대로 require 되어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\((['"])passport\.socketio\1\)/gi,
    'You should correctly require and instantiate "passport.socketio"'
  );
```

passportSocketIo가 제대로 설정되어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io\.use\(\s*\w+\.authorize\(/,
    'You should register "passport.socketio" as socket.io middleware and provide it correct options'
  );
```
