---
id: 589fc831f9fc0f352b528e76
title: 연결 끊김 처리
challengeType: 2
forumTopicId: 301552
dashedName: handle-a-disconnect
---

# --description--

지금까지 사용자 수를 늘리는 것만 해왔다는 것을 알 수 있을 겁니다. 사용자가 연결을 끊는 것을 처리하는 것은 초기 연결을 처리하는 것만큼 쉽지만, 전체 서버가 아니라 각 소켓에서 이를 감지해야 합니다.

이를 위해 기존 `'connect'` 리스너 안에 또 다른 리스너를 추가해 소켓에서 `'disconnect'` 이벤트를 데이터 없이 감지하세요. 이 기능은 사용자가 연결을 끊었다는 메시지를 콘솔에 기록하는 것으로 테스트할 수 있습니다.

```js
socket.on('disconnect', () => {
  /*anything you want to do on disconnect*/
});
```

클라이언트가 현재 사용자 수를 계속 최신 상태로 유지하도록 하려면, 연결이 끊길 때 `currentUsers` 값을 1만큼 줄이고 업데이트된 수를 `'user count'` 이벤트로 내보내야 합니다.

**참고:** `'disconnect'`와 마찬가지로, 소켓이 서버로 내보낼 수 있는 다른 모든 이벤트는 'socket'이 정의된 연결 리스너 안에서 처리해야 합니다.

정확히 했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#handle-a-disconnect-8" target="_blank" rel="noopener noreferrer nofollow">지금까지 완성된 프로젝트</a>를 확인할 수 있습니다.

# --hints--

서버는 소켓에서 발생하는 disconnect 이벤트를 처리해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /socket.on.*('|")disconnect('|")/s, '');
```

클라이언트는 `'user count'` 이벤트를 수신 대기해야 합니다.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
