---
id: 589fc831f9fc0f352b528e75
title: 방출로 소통하기
challengeType: 2
forumTopicId: 301550
dashedName: communicate-by-emitting
---

# --description--

<dfn>Emit</dfn>은 가장 흔히 사용하는 소통 방법입니다. 서버에서 'io'로 무언가를 방출할 때, 연결된 모든 소켓에 이벤트 이름과 데이터를 보냅니다. 이 개념의 좋은 예는 새 사용자가 연결될 때마다 현재 연결된 사용자 수를 방출하는 것입니다!

현재 연결을 듣기 바로 전에 사용자를 추적할 변수를 추가하는 것부터 시작하세요.

```js
let currentUsers = 0;
```

이제 누군가 연결되면, 방출하기 전에 카운트를 증가시켜야 합니다. 따라서 연결 리스너 안에 증가시키는 코드를 추가하세요.

```js
++currentUsers;
```

마지막으로, 카운트를 증가시킨 후에는 이벤트를 방출해야 합니다(여전히 연결 리스너 안에 있어야 합니다). 이벤트 이름은 'user count'로 하고, 데이터는 `currentUsers`이어야 합니다.

```js
io.emit('user count', currentUsers);
```

이제 클라이언트가 이 이벤트를 듣도록 구현할 수 있습니다! 서버에서 연결을 듣는 것과 비슷하게, `on` 키워드를 사용하세요.

```js
socket.on('user count', function(data) {
  console.log(data);
});
```

이제 앱을 실행하고 인증하면, 클라이언트 콘솔에 현재 사용자 수를 나타내는 '1'이 표시될 것입니다! 클라이언트를 더 열고 인증하여 숫자가 올라가는 것을 확인하세요.

정확하다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#communicate-by-emitting-7" target="_blank" rel="noopener noreferrer nofollow">지금까지 완성된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

`currentUsers`가 정의되어 있어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /currentUsers/s,
    'You should have variable currentUsers defined'
  );
```

서버는 새 연결이 있을 때마다 현재 사용자 수를 방출해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.emit.*('|")user count('|").*currentUsers/s,
    'You should emit "user count" with data currentUsers'
  );
```

클라이언트는 `'user count'` 이벤트를 듣고 있어야 합니다.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
