---
id: 589fc832f9fc0f352b528e79
title: 채팅 메시지 보내고 표시하기
challengeType: 2
forumTopicId: 301562
dashedName: send-and-display-chat-messages
---

# --description--

클라이언트가 서버에 채팅 메시지를 보내고 모든 클라이언트에게 전송하도록 할 시간입니다! `client.js` 파일에서 메시지 폼이 제출될 때 처리하는 코드 블록이 이미 있는 것을 확인할 수 있습니다.

```js
$('form').submit(function() {
  /*logic*/
});
```

폼 제출 코드 내에서 `messageToSend`을 정의한 후, 텍스트 상자를 `#m` 하기 전에 이벤트를 발생시켜야 합니다. 이벤트 이름은 `'chat message'`이고, 데이터는 `messageToSend`이어야 합니다.

```js
socket.emit('chat message', messageToSend);
```

이제 서버에서는 `'chat message'` 이벤트를 `message`라는 데이터 이름으로 소켓에서 수신 대기해야 합니다. 이벤트를 받으면 `'chat message'`을 사용해 모든 소켓에 `io.emit` 이벤트를 발생시키고, `username`와 `message`을 포함하는 데이터 객체를 전송해야 합니다.

`client.js`에서는 이제 `'chat message'` 이벤트를 수신 대기하고, 이벤트가 수신되면 사용자 이름, 콜론, 메시지를 포함하는 리스트 항목을 `#messages`에 추가해야 합니다!

이 시점에서 채팅은 완전히 작동하며 모든 클라이언트 간에 메시지를 전송할 수 있어야 합니다!

정확히 구현했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#send-and-display-chat-messages-11" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트</a>를 확인할 수 있습니다.

# --hints--

서버는 `'chat message'` 이벤트를 수신 대기하고 적절히 발생시켜야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*io.emit.*('|")chat message('|").*username.*message/s,
    'Your server should listen to the socket for "chat message" then emit to all users "chat message" with name and message in the data object'
  );
```

클라이언트는 `'chat message'` 이벤트에서 새 데이터를 올바르게 처리하고 표시해야 합니다.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*messages.*li/s,
    'You should append a list item to #messages on your client within the "chat message" event listener to display the new message'
  );
```
