---
id: 589fc830f9fc0f352b528e74
title: 환경 설정하기
challengeType: 2
forumTopicId: 301566
dashedName: set-up-the-environment
---

# --description--

다음 도전 과제들은 `chat.pug` 파일을 사용합니다. 그러니 `routes.js` 파일에 `/chat`를 사용하고 `ensureAuthenticated`를 응답 인수로 전달하여 `chat.pug`를 렌더링하는 `{ user: req.user }`를 가리키는 GET 라우트를 추가하세요. 이제 기존 `/auth/github/callback` 라우트를 수정하여 `req.session.user_id = req.user.id`를 설정하고 `/chat`로 리다이렉트하세요.

`socket.io@~2.3.0`는 이미 의존성으로 추가되어 있으니, 다음과 같이 서버에서 `http` (Node.js에 내장되어 있음)과 함께 require/인스턴스화하세요:

```javascript
const http = require('http').createServer(app);
const io = require('socket.io')(http);
```

*http* 서버가 *express 앱*에 마운트되었으니, *http* 서버에서 listen해야 합니다. `app.listen`가 있는 줄을 `http.listen`로 변경하세요.

처음 처리해야 할 것은 클라이언트의 새 연결을 듣는 것입니다. <dfn>on</dfn> 키워드는 특정 이벤트를 듣는 역할을 합니다. 두 개의 인수를 필요로 하는데, 첫 번째는 발생한 이벤트 제목을 담은 문자열이고, 두 번째는 데이터를 전달받는 함수입니다. 연결 리스너의 경우 두 번째 인수에 데이터를 정의하기 위해 `socket`를 사용하세요. 소켓은 연결된 개별 클라이언트를 의미합니다.

서버에 대한 연결을 듣기 위해 데이터베이스 연결 내에 다음을 추가하세요:

```javascript
io.on('connection', socket => {
  console.log('A user has connected');
});
```

이제 클라이언트가 연결하려면 인증 후 페이지에서 로드되는 `client.js`에 다음을 추가하면 됩니다:

```js
/*global io*/
let socket = io();
```

주석은 'io'가 파일 내에 정의되지 않아 보통 발생하는 오류를 억제합니다. `chat.pug`에 페이지에서 Socket.IO 라이브러리의 신뢰할 수 있는 CDN을 이미 추가했습니다.

이제 앱을 로드하고 인증하면 서버 콘솔에 `A user has connected`가 표시될 것입니다.

**참고:** `io()`는 같은 url/서버에서 호스트되는 소켓에 연결할 때만 작동합니다. 다른 곳에 호스트된 외부 소켓에 연결하려면 `io.connect('URL');`를 사용하세요.

정확히 했다고 생각하면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-the-environment-6" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

`socket.io`는 의존성이어야 합니다.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'socket.io',
    'Your project should list "socket.io" as a dependency'
  );
```

`http`를 `http`로 올바르게 require하고 인스턴스화해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /http.*=.*require.*('|")http\1/s,
    'Your project should list "http" as a dependency'
  );
```

`socket.io`를 `io`로 올바르게 require하고 인스턴스화해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.*=.*require.*('|")socket.io\1.*http/s,
    'You should correctly require and instantiate socket.io as io.'
  );
```

Socket.IO가 연결을 듣고 있어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.on.*('|")connection\1.*socket/s,
    'io should listen for "connection" and socket should be the 2nd arguments variable'
  );
```

클라이언트가 서버에 연결해야 합니다.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.*=.*io/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
