---
id: 589fc832f9fc0f352b528e78
title: 새 사용자 알림
challengeType: 2
forumTopicId: 301546
dashedName: announce-new-users
---

# --description--

많은 채팅방에서는 사용자가 연결되거나 연결이 끊어질 때 이를 모든 사용자에게 알릴 수 있습니다. 이미 연결 및 연결 해제 시 이벤트를 발생시키고 있으므로, 이 이벤트를 수정하여 그런 기능을 지원하면 됩니다. 가장 논리적인 방법은 이벤트와 함께 세 가지 데이터를 보내는 것입니다: 연결되거나 연결이 끊어진 사용자의 사용자 이름, 현재 사용자 수, 그리고 그 사용자가 연결했는지 연결이 끊어졌는지 여부입니다.

이벤트 이름을 `'user'`로 변경하고, `username`, `currentUsers`, `connected` 필드를 포함하는 객체를 전달하세요 (연결 시에는 `true`, 연결 해제 시에는 `false`가 됩니다). `'user count'` 이벤트 두 곳 모두 변경하고, 연결 해제 이벤트에서는 `false` 대신 `connected`을 `true` 필드에 보내도록 설정하세요.

```js
io.emit('user', {
  username: socket.request.user.username,
  currentUsers,
  connected: true
});
```

이제 클라이언트는 현재 사용자 수를 올바르게 표시하고 사용자가 연결되거나 연결이 끊어질 때 알릴 수 있는 모든 필요한 정보를 갖게 됩니다! 클라이언트 측에서 이 이벤트를 처리하려면 `'user'` 이벤트를 수신 대기하고, jQuery를 사용해 `#num-users`의 텍스트를 `'{NUMBER} users online'`로 변경하여 현재 사용자 수를 업데이트하며, `<li>` 아이디를 가진 순서 없는 목록에 `messages`를 `'{NAME} has {joined/left} the chat.'`와 함께 추가하세요.

이 구현 예시는 다음과 같을 수 있습니다:

```js
socket.on('user', data => {
  $('#num-users').text(data.currentUsers + ' users online');
  let message =
    data.username +
    (data.connected ? ' has joined the chat.' : ' has left the chat.');
  $('#messages').append($('<li>').html('<b>' + message + '</b>'));
});
```

정확히 구현했다고 생각되면 페이지를 제출하세요. 오류가 발생하면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135/3#announce-new-users-10" target="_blank" rel="noopener noreferrer nofollow">지금까지 완료된 프로젝트</a>를 확인할 수 있습니다.

# --hints--

이벤트 `'user'`는 `username`, `currentUsers`, `connected`와 함께 발생해야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  // Regex is lenient to match both `username` and `name` as the key on purpose.
  assert.match(
    data,
    /io.emit.*('|")user\1.*name.*currentUsers.*connected/s,
    'You should have an event emitted named user sending name, currentUsers, and connected'
  );
```

클라이언트는 `'user'` 이벤트에서 새 데이터를 올바르게 처리하고 표시해야 합니다.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*num-users/s,
    'You should change the text of "#num-users" within on your client within the "user" event listener to show the current users connected'
  );
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*messages.*li/s,
    'You should append a list item to "#messages" on your client within the "user" event listener to announce a user came or went'
  );
```
