---
id: 58a25c98f9fc0f352b528e7f
title: 비밀번호 해싱하기
challengeType: 2
forumTopicId: 301553
dashedName: hashing-your-passwords
---

# --description--

정보 보안 섹션으로 돌아가 보면, 평문 비밀번호를 저장하는 것은 *절대* 허용되지 않는다는 점을 기억해주세요. 이제 이 문제를 해결하기 위해 BCrypt를 구현할 차례입니다.

`bcrypt@~5.0.0`가 이미 의존성으로 추가되어 있으니, 서버에서 require 하세요. 해싱을 처리해야 하는 주요 영역은 두 군데입니다: 새 계정을 등록/저장할 때와 로그인 시 비밀번호가 올바른지 확인할 때입니다.

현재 등록 라우트에서는 사용자의 평문 비밀번호를 다음과 같이 데이터베이스에 삽입합니다: `password: req.body.password`. 대신 비밀번호를 해싱하려면 데이터베이스 로직 전에 다음을 추가하세요: `const hash = bcrypt.hashSync(req.body.password, 12);`, 그리고 데이터베이스 저장 부분의 `req.body.password`를 `password: hash`로 바꾸세요.

인증 전략에서는 프로세스를 완료하기 전에 코드에서 다음을 확인합니다: `if (password !== user.password) return done(null, false);`. 이전 변경을 적용한 후, 이제 `user.password`는 해시입니다. 기존 코드를 변경하기 전에, 해당 문장이 비밀번호가 **같지 않으면** 인증 실패를 반환하는지 확인하세요. 이를 염두에 두고, 비밀번호를 해시와 올바르게 비교하도록 코드를 다음과 같이 바꾸세요:

```js
if (!bcrypt.compareSync(password, user.password)) { 
  return done(null, false);
}
```

비밀번호를 저장해야 할 때 가장 중요한 보안 기능 중 하나를 구현하는 데 필요한 전부입니다.

정확히 구현했다고 생각되면 페이지를 제출하세요. 오류가 발생한다면 <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#hashing-your-passwords-1" target="_blank" rel="noopener noreferrer nofollow">지금까지 완성된 프로젝트를 확인할 수 있습니다</a>.

# --hints--

BCrypt는 의존성으로 있어야 합니다.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json()
  assert.property(
    packJson.dependencies,
    'bcrypt',
    'Your project should list "bcrypt" as a dependency'
  );
```

BCrypt가 올바르게 require 되고 구현되어야 합니다.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')bcrypt\1/gi,
    'You should have required bcrypt'
  );
  assert.match(
    data,
    /bcrypt.hashSync/gi,
    'You should use hash the password in the registration'
  );
  assert.match(
    data,
    /bcrypt.compareSync/gi,
    'You should compare the password to the hash in your strategy'
  );
```
