---
id: 5f3313e74582ad9d063e3a38
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Add a `head` element within the `html` element, so you can add a `title` element. The `title` element's text should be `Cafe Menu`.

# --hints--

Dovresti avere un tag di apertura `<head>`.

```js
assert.match(code, /<head>/i);
```

Dovresti avere un tag di chiusura `</head>`.

```js
assert.match(code, /<head>/i);
```

Dovresti avere un tag di apertura `<title>`.

```js
assert.match(code, /<title>/i);
```

Dovresti avere un tag di chiusura `</title>`.

```js
assert.match(code, /<\/title>/i);
```

L'elemento `title` dovrebbe essere annidato nell'elemento `head`.

```js
assert.match(code, /<head>\s*<title>.*<\/title>\s*<\/head>/si);
```

L'elemento `title` dovrebbe avere il testo `Cafe Menu`. Controlla la tua ortografia.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
