---
id: 5f332a88dc25a0fd25c7687a
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

The name of the cafe is `CAMPER CAFE`. So, add an `h1` element within your `main` element. Give it the name of the cafe in capitalized letters to make it stand out.

# --hints--

Dovresti avere un tag di apertura `<h1>`.

```js
assert.match(code, /<h1>/i);
```

Dovresti avere un tag di chiusura `</h1>`.

```js
assert.match(code, /<\/h1>/i);
```

Non dovresti cambiare l'elemento `main` esistente.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

L'elemento `h1` dovrebbe essere annidato nell'elemento `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

L'elemento `h1` dovrebbe avere il testo `CAMPER CAFE`.

```js
assert.match(code, /<h1>CAMPER CAFE<\/h1>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

