---
id: 5f33294a6af5e9188dbdb8f3
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

It's time to add some menu content. Add a `main` element within the existing `body` element. It will eventually contain pricing information about coffee and desserts offered by the cafe.

# --hints--

Il codice dovrebbe avere un tag di apertura `<main>`.

```js
assert(code.match(/<main>/i));
```

Il codice dovrebbe avere un tag di chiusura `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Non dovresti cambiare l'elemento `body`. Assicurati di non aver eliminato accidentalmente il tag di chiusura.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Il tag `main` dovrebbe essere all'interno del tag `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
