---
id: 5f3477ae9675db8bb7655b30
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

In the previous step, you used a <dfn>type selector</dfn> to style the `h1` element. Center the content of the `h2` and the `p` elements by adding a new type selector for each one to the existing `style` element.

# --hints--

Non dovresti cambiare il selettore `h1` esistente.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasH1);
```

Non dovresti aggiungere un nuovo tag `style`. Aggiungi le nuove regole CSS al tag `style` esistente.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

Dovresti aggiungere un nuovo selettore `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert.exists(hasH2);
```

Dovresti aggiungere un nuovo selettore `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert.exists(hasP);
```

L'elemento `h1` dovrebbe avere una proprietà `text-align` con valore `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

L'elemento `h2` dovrebbe avere una proprietà `text-align` con valore `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert.equal(h2TextAlign, 'center');
```

L'elemento `p` dovrebbe avere una proprietà `text-align` con valore `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert.equal(pTextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

