---
id: a6b0bb188d873cb2c8729495
title: Converter entidades HTML
challengeType: 1
forumTopicId: 16007
dashedName: convert-html-entities
---

# --description--

Converta os caracteres `&`, `<`, `>`, `"` (aspas duplas) e `'` (apóstrofo), em uma string para suas respectivas entidades HTML.

# --hints--

`convertHTML("Dolce & Gabbana")` deve retornar a string `Dolce & Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` deve retornar a string `Hamburgers < Pizza < Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sessenta > doze")` deve retornar a string `Sessenta > doze`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` deve retornar a string `Stuff in "quotation marks"`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("A Lista de Schindler")` deve retornar a string `Schindler's List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` deve retornar a string `<>`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` deve retornar a string `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js
function convertHTML(str) {
  return str;
}

convertHTML("Dolce & Gabbana");
```

# --solutions--

```js
var MAP = { '&': '&amp;',
            '<': '&lt;',
            '>': '&gt;',
            '"': '&quot;',
            "'": '&apos;'};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(c) {
    return MAP[c];
  });
}
```
