---
id: 66ec4c8e9878d8441956516f
title: Створіть каталог книг
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Історія користувача:**

1. Створіть елемент `table`, який містить інформацію про книги.
1. Таблиця повинна мати елемент `thead`, який містить один рядок.
1. Рядок в елементі `thead` має містити чотири заголовки таблиці з текстом `Title`, `Author`, `Genre` та `Publication Year` (в такому ж порядку).
1. Таблиця повинна мати елемент `tbody`, який містить принаймні п’ять рядків.
1. Рядки в елементі `tbody` мають містити елементи з табличними даними, в яких вказано назву, автора, жанр та рік публікації книги.
1. Таблиця повинна мати елемент `tfoot`, який містить один рядок.
1. Рядок в елементі `tfoot` має містити елемент з табличними даними, який охоплює чотири стовпці і містить текст `Total Books: N`, де `N` позначає кількість книг в таблиці.

# --hints--

Ви повинні мати елемент `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

В межах елемента `table` має бути лише один елемент `thead`.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

В межах елемента `thead` має бути лише один елемент `tr`.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

В рядку елемента `thead` має бути чотири елементи `th`.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

Елементи `th` повинні мати текст `Title`, `Author`, `Genre` та `Publication Year` (в такому ж порядку).

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent, 'Title');
assert.equal(ths[1]?.textContent, 'Author');
assert.equal(ths[2]?.textContent, 'Genre');
assert.equal(ths[3]?.textContent, 'Publication Year');
```

В межах елемента `table` має бути лише один елемент `tbody`.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

Елемент `tbody` повинен мати принаймні п’ять рядків.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

Кожен рядок в елементі `tbody` повинен мати чотири дочірні елементи `td`.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

Кожен елемент `td` в `tbody` повинен містити текст з інформацією про книгу.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.length, 1);
});
```

В межах елемента `table` має бути лише один елемент `tfoot`.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

В межах елемента `tfoot` має бути лише один елемент `tr`.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

Елемент `td` в рядку елемента `tfoot` повинен мати атрибут `colspan` зі значенням `4`.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

Елемент `td` в рядку елемента `tfoot` повинен мати текст `Total Books: [number of books in your table]`.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent, `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
