---
id: 66c06d618d075c7f7f1b890a
title: Baue einen Wahrsager
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

In diesem Labor definieren Sie fünf Wahrsagungen und wählen zufällig eine davon aus, um sie dem Benutzer anzuzeigen.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten die fünf Variablen `fortune1`, `fortune2`, `fortune3`, `fortune4` und `fortune5` mit einem String-Wert Ihrer Wahl initialisieren. Sie können die untenstehenden Optionen verwenden, wenn Sie möchten:

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. Sie sollten eine Zufallszahl zwischen 1 und 5, einschließlich, auswählen und sie der Variablen `randomNumber` zuweisen.

3. Sie sollten eine `selectedFortune`-Variable erstellen und die passende Wahrsagung basierend auf diesen Regeln zuweisen:

   - Wenn `randomNumber` 1 ist, weisen Sie `fortune1` den Wert von `selectedFortune` zu.
   - Wenn `randomNumber` 2 ist, weisen Sie `fortune2` den Wert von `selectedFortune` zu.
   - Wenn `randomNumber` 3 ist, weisen Sie `fortune3` den Wert von `selectedFortune` zu.
   - Wenn `randomNumber` 4 ist, weisen Sie `fortune4` den Wert von `selectedFortune` zu.
   - Wenn `randomNumber` 5 ist, weisen Sie `fortune5` den Wert von `selectedFortune` zu.

4. Sie sollten `selectedFortune` in der Konsole ausgeben.

# --hints--

Sie sollten `fortune1` mit einem String-Wert initialisieren.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

Sie sollten `fortune2` mit einem String-Wert initialisieren.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

Sie sollten `fortune3` mit einem String-Wert initialisieren.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

Sie sollten `fortune4` mit einem String-Wert initialisieren.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

Sie sollten `fortune5` mit einem String-Wert initialisieren.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

Sie sollten die `Math.random()`-Methode verwenden, um eine Zufallszahl zu erzeugen.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

Sie sollten eine Zufallszahl zwischen 1 und 5, einschließlich, erzeugen und sie der Variablen `randomNumber` zuweisen.

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Sie sollten eine `selectedFortune`-Variable haben, der basierend auf dem Wert von `randomNumber` ein Wert zugewiesen wird.

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` sollte seiner Wahrsagung entsprechen. Zum Beispiel, wenn `randomNumber` 1 ist, sollte `selectedFortune` gleich `fortune1` sein und so weiter.

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Sie sollten `selectedFortune` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
