---
id: 68c16924498728d259e4781c
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Ahora crea una instancia de la clase `Employee` pasando las cadenas `Charlie Brown` y `trainee`. Asigna la instancia a una variable llamada `charlie_brown`.

# --hints--

Debes tener una variable llamada `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

Debes asignar `Employee('Charlie Brown', 'trainee')` a `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level


--fcc-editable-region--
```
