---
id: 68c15a2a1f0e70ca7154aba5
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, vas a construir un sistema de seguimiento de salarios para empleados.

Comienza creando una clase llamada `Employee`. Dentro de ella crea un método `__init__` con los parámetros `self`, `name` y `level`. Dentro del método `__init__`, asigna `name` y `level` a los atributos de la instancia con el mismo nombre.

# --hints--

Debes tener una clase llamada `Employee`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

Tu clase `Employee` debe tener un método `__init__` con tres parámetros `self`, `name` y `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

Tu método `__init__` debe establecer `self.name` en `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Tu método `__init__` debe establecer `self.level` en `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
