---
id: 68c18b1d0fce081f85312d97
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

En una lección anterior, aprendiste que un atributo prefijado con un solo guion bajo está destinado para uso interno por convención.

Modifica ambos atributos `name` y `level` a `_name` y `_level`, ya que no se supone que estos deban ser modificados desde fuera de su clase.

Ten en cuenta que esto no impide que el atributo sea accedido o modificado fuera de la clase. Además, en Python siempre hay una manera de acceder a los atributos privados (prefijados con un doble guion bajo).

# --hints--

Tu método `__init__` debe establecer `self._name` en `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

Tu método `__init__` debería establecer `self._level` en `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
