---
id: 68c28d659d345270d4a61084
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

El método `__repr__` es un método especial que se supone debe devolver una representación en cadena del objeto que pueda usarse para instanciarlo.

Por ejemplo, el método `__repr__` de `Employee('Charlie', 'developer')` debería devolver la cadena `Employee('Charlie', 'developer')`, que es la misma cadena usada para crear el objeto.

Dale a tu clase `Employee` un método `__repr__` con un parámetro `self`, y haz que devuelva una cadena que pueda usarse para instanciar el objeto.

# --hints--

Tu clase `Employee` debe tener un método `__repr__` con un parámetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

Tu método `__repr__` debe devolver la cadena usada para instanciar el objeto.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
