---
id: 68c2d31a86d5672ca49d4521
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

Ahora es momento de agregar algo de validación al método `__init__`. Al inicio del método, crea una sentencia `if` que verifique si `name` o `level` no son instancias de `str`.

Dentro de la sentencia `if`, lanza un `TypeError` con el mensaje `'name' and 'level' attribute must be of type 'str'.`.

# --hints--

Deberías tener una estructura `if` en tu método `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0]`)) })
```

Tu sentencia `if` debe verificar si `name` o `level` no son instancias de `str`.

```js
({ test: () => runPython(`
  cond = _Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_conditions()[0]
  conditions = [
    "not (isinstance(name, str) and isinstance(level, str))",
    "not isinstance(name, str) or not isinstance(level, str)",
    "not (isinstance(level, str) and isinstance(name, str))",
    "not isinstance(level, str) or not isinstance(name, str)",
  ]
  assert any(cond.is_equivalent(c) for c in conditions)
`) })
```

Tu sentencia `if` debe lanzar un `TypeError` con el mensaje `'name' and 'level' attribute must be of type 'str'.`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError("\\'name\\' and \\'level\\' attribute must be of type \\'str\\'.")')`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
--fcc-editable-region--
    def __init__(self, name, level):
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
