---
id: 68c28b7cc3ab8f4daf144b5e
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Siguiendo lo que hiciste en los pasos anteriores, crea un getter para el atributo `_level`.

# --hints--

Tu clase `Employee` debe tener un método `level` con un parámetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

Tu método `level` debe estar decorado con `@property`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

Tu método `level` debería devolver `self._level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
