---
id: 68c27fa1697ceb7c578b2c9f
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

El decorador `@property` se usa en Python para convertir un método en una propiedad. Normalmente se usa para definir métodos getter, que son métodos usados para obtener el valor de un atributo:

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

Crea un método llamado `name` con un parámetro `self` y decóralo con `@property`. Dentro del método, devuelve `self._name`.
    
# --hints--

Tu clase `Employee` debe tener un método `name` con un parámetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

Tu método `name` debe estar decorado con `@property`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

Tu método `name` debería devolver `self._name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
