---
id: 68c281cecda94ba2c6838e17
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Ahora que definiste un getter para `name`, puedes acceder al atributo `_name` a través de la propiedad `name`. Así que imprime `charlie_brown.name` en la consola. 

# --hints--

Debes imprimir `charlie_brown.name` en la consola.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
