---
id: 68c2d2414f5e111c02f4db5e
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

El método `__repr__` se llama internamente cuando llamas a la función `repr`. Para verlo en acción, imprime el resultado de llamar a `repr(charlie_brown)` al final de tu código.

# --hints--

Debes imprimir `repr(charlie_brown)`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(repr(charlie_brown))")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

--fcc-editable-region--
```
