---
id: 68c27eb292bfec6b9402d58f
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Agrega un método `__str__` a la clase `Employee`. Haz que retorne un f-string con el formato `name: level`, reemplazando `name` y `level` con los atributos correspondientes.

Después de eso, imprime `charlie_brown` en la consola.

# --hints--

Tu clase `Employee` debe tener un método `__str__` con un parámetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

Tu método `__str__` debe devolver un f-string con el formato `name: level`, reemplazando `name` y `level` con los valores de `self._name` y `self._level`, respectivamente.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

Debes imprimir `charlie_brown` en la consola.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--
```
