---
id: 68c28c6517e7de5e85f8ebe5
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Ahora que tienes getters para los atributos `_name` y `_level`, actualiza la cadena que devuelve `__str__` para usar `self.name` y `self.level`. Esto llamará a los getters en lugar de acceder directamente a los atributos.

# --hints--

Tu método `__str__` debería devolver `f'{self.name}: {self.level}'`

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_return("f'{self.name}: {self.level}'")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)
print(charlie_brown.level)
```
