---
id: 68c7f7012a700243eff1cbc0
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

El nivel usado para instanciar a un empleado debe elegirse entre niveles específicos. Usarás `_base_salaries` para validar el nivel y establecer el salario correcto para el empleado.

Después de tu `if` existente, crea otro `if` que verifique si `level` no está en `Employee._base_salaries`.

Dentro de la nueva estructura `if`, lanza un `ValueError` con el mensaje `Invalid value '{level}' for 'level' attribute.`, donde `{level}` debe ser reemplazado por el valor del argumento `level`.

# --hints--

Cuando `level` no está en `Employee._base_salaries`, debes lanzar un `ValueError` con el mensaje `Invalid value '{level}' for 'level' attribute.`, donde `{level}` debe ser reemplazado por el valor del argumento `level`.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

No debes generar ninguna excepción cuando `level` esté en `Employee._base_salaries`.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
