---
id: 68c8002beb157fe47d2e7699
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Un setter es un método usado para asignar el valor de un atributo, permitiendo realizar verificaciones de validación y restricciones. Puedes crear un setter usando el decorador `@propertyName.setter`, donde `propertyName` debe coincidir con el nombre de la propiedad que quieres asignar:

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, value):
        self._name = value

p = Person('Alice')
p.name = 'Abigail' # Calls the setter
print(p.name) # Abigail
```

Después de tu método getter, crea un método `name` con los parámetros `self` y `new_name`. Decora el método con `@name.setter`. Dentro del método, asigna `self._name` a `new_name`.

# --hints--

Tu clase `Employee` debe tener un método `name` con los parámetros `self` y `new_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_args("self, new_name")`)) })
```

Tu método `name` debe estar decorado con `@name.setter`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_decorators("name.setter")`)) })
```

Tu método `name` debe establecer `self._name` a `new_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_stmt("self._name = new_name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary


charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
