---
id: 68c9cab4b1118da59eecfc56
title: Paso 31
challengeType: 20
dashedName: step-31
---

# --description--

Después de las declaraciones `if` existentes, crea otra para lanzar un `ValueError` cuando `new_level` ya sea el nivel seleccionado.

Ten en cuenta que `_level` aún no existe durante la inicialización, así que usa `hasattr(self, '_level')` para verificar si existe antes de comparar. Esto evita un `AttributeError` cuando el objeto se crea por primera vez.

Para el mensaje, usa `'{level}' is already the selected level.`, donde `{level}` debe ser reemplazado por el nivel actual.

# --hints--

Deberías tener una tercera declaración `if` dentro de tu setter `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

Cuando `new_level` sea igual a `self.level`, deberías lanzar un `ValueError` con el mensaje `'{level}' is already the selected level.`, donde `{level}` debe ser reemplazado por el nivel actual.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

Tu declaración `if` debería usar `hasattr(self, '_level')` para verificar si `_level` existe antes de comparar.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
