---
id: 669fcb5106eeb13ab7afcd09
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafundishwa kufanya kazi na semantic HTML kwa kujenga ukurasa wa blogu uliotengwa kwa Bwana Whiskers paka. 

Sehemu ya kwanza utakayojenga ni <dfn>kichwa</dfn> cha ukurasa. 

Kielemeni cha `header` kinatumika kuwakilisha maudhui ya utangulizi kama vile urambazaji wa ukurasa na taarifa nyingine za utangulizi. 

Hapa kuna mfano unaotumia kielemeni cha `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Ndani ya kielemeni cha `body`, ongeza kielemeni cha `header`.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `header`.

```js
assert.match(code, /<header>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `header`.

```js
assert.match(code, /<\/header>/i);
```

Kielemeni chako cha `header` kinapaswa kuwa ndani ya kielemeni cha `body`.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
