---
id: 669fdc11c9b0263fe0814a7a
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Ndani ya kila kipengele cha `li`, utahitaji kuwa na kipengele cha nanga.

Kwa kipengele cha kwanza cha nanga, maandishi yanapaswa kuwa `About` na thamani ya sifa ya `href` inapaswa kuwa `"#about"`. Alama ya hash (#) mbele ya `about` inawakilisha jina la `id`, ambalo litaongezwa baadaye katika mradi.

Kwa kipengele cha pili cha nanga, maandishi yanapaswa kuwa `Posts` na thamani ya sifa ya `href` inapaswa kuwa `"#posts"`.

Kwa kipengele cha tatu cha nanga, maandishi yanapaswa kuwa `Contact` na thamani ya sifa ya `href` inapaswa kuwa `"#contact"`.

# --hints--

Kipengele chako cha kwanza cha nanga kinapaswa kuwa na maandishi ya `About` ndani ya kipengele chako cha kwanza cha `li`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

Kipengele chako cha kwanza cha nanga kinapaswa kuwa na sifa ya `href` iliyowekwa kuwa `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Kipengele chako cha pili cha nanga kinapaswa kuwa na maandishi ya `Posts` ndani ya kipengele chako cha pili cha `li`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

Kipengele chako cha pili cha nanga kinapaswa kuwa na sifa ya `href` iliyowekwa kuwa `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Kipengele chako cha tatu cha nanga kinapaswa kuwa na maandishi ya `Contact` ndani ya kipengele chako cha tatu cha `li`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

Kipengele chako cha tatu cha nanga kinapaswa kuwa na sifa ya `href` iliyowekwa kuwa `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
