---
id: 669fd003cb89ee3c2402e041
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Katika maudhui haya ya utangulizi, utataka kuonyesha picha ya Bwana Whiskers pamoja na maelezo ya picha.

Chini ya kipengele cha `h1`, anza kwa kuongeza kipengele cha `figure`.

Ndani ya kipengele cha `figure`, ongeza kipengele cha `img`.

Sifa ya kipengele cha `src` cha `img` inapaswa kuwa na thamani ya `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` na maandishi ya `alt` yanapaswa kuwa na thamani ya `"a cat in the garden"`.

Chini ya kipengele chako cha `img`, ongeza `figcaption` yenye maandishi ya `Mr. Whiskers in the Garden`.

# --hints--

Kipengele chako cha `figure` kinapaswa kuwa na lebo ya kufungua.

```js
assert.match(code, /<figure\>/);
```

Kipengele chako cha `figure` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/figure\>/);
```

Kipengele chako cha `figure` kinapaswa kuwa ndani ya kipengele chako cha `header`.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

Unapaswa kuwa na kipengele cha `img`.

```js
assert.exists(document.querySelector("img"));
```

Kipengele chako cha `img` kinapaswa kuwa ndani ya kipengele chako cha `figure`.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

Kipengele chako cha `img` kinapaswa kuwa na sifa ya kipengele cha `src` iliyowekwa kuwa `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Kipengele chako cha `img` kinapaswa kuwa na sifa ya kipengele cha `alt` iliyowekwa kuwa `"a cat in the garden"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

Kipengele chako cha `figcaption` kinapaswa kuwa na lebo ya kufungua.

```js
assert.match(code, /<figcaption\>/);
```

Kipengele chako cha `figcaption` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/figcaption\>/);
```

Kipengele chako cha `figcaption` kinapaswa kuwa ndani ya kipengele chako cha `figure`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Kipengele chako cha `figcaption` kinapaswa kuwa na maandishi ya `Mr. Whiskers in the Garden`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

`figcaption` yako inapaswa kuwa chini ya kipengele cha `img`. Umeviweka kwa mpangilio usio sahihi.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
