---
id: 669fdfc9a5e3da42d2376609
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Sasa umeongeza sehemu ya kuhusu, jaribu kubofya kiungo cha `About` kuona ukurasa unaruka hadi sehemu hiyo.

Sehemu inayofuata kwenye ukurasa wa blogu itakuwa orodha ya machapisho yanayozungumzia Bwana Whiskers.

Ongeza kipengele kingine cha `section` chenye `id` iliyowekwa kuwa `"posts"`.

Ndani ya kipengele cha `section`, ongeza kipengele cha `h2` chenye maandishi ya `Posts`.

# --hints--

Msimbo wako unapaswa kuwa na lebo ya ufunguzi ya `section`.

```js
assert.match(code, /<section.*>/i);
```

Msimbo wako unapaswa kuwa na lebo ya kufunga ya `section`.

```js
assert.match(code, /<\/section>/i);
```

Kipengele chako cha `section` kinapaswa kuwa na `id` iliyowekwa kuwa `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

Kipengele chako cha `section` kinapaswa kuwa ndani ya kipengele chako cha `main`.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Kipengele chako cha `h2` kinapaswa kuwa ndani ya kipengele cha `section`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

Kipengele chako cha `h2` kinapaswa kuwa na maandishi ya `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
