---
id: 669fd6fd12918e3de87854d4
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Kwa blogu yako, inapaswa kuwepo njia kwa watumizi kuvinjari sehemu tofauti kwenye ukurasa.

Kipengele cha `nav` kinatumika kutoa viungo vya urambazaji kwenda sehemu nyingine katika hati au sehemu nyingine kwenye tovuti. Mara nyingi utaona kipengele cha `nav` kinatumika kwa menyu au jedwali la maudhui.

Hapa kuna mfano wa kutumia kipengele cha `nav`:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

Chini ya kipengele chako cha `figure`, ongeza kipengele cha `nav` chenye kipengele cha `ul` kilichopangwa ndani.

Ndani ya kipengele cha `ul`, ongeza vipengele vitatu vya `li`.

# --hints--

Kipengele chako cha `nav` kinapaswa kuwa na lebo ya kufungua.

```js
assert.match(code, /<nav\>/);
```

Kipengele chako cha `nav` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/nav\>/);
```

Kipengele chako cha `nav` kinapaswa kuwa ndani ya kipengele chako cha `header`.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

Kipengele chako cha `ul` kinapaswa kuwa na lebo ya kufungua.

```js
assert.match(code, /<ul\>/);
```

Kipengele chako cha `ul` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/ul\>/);
```

Kipengele chako cha `ul` kinapaswa kuwa ndani ya kipengele chako cha `nav`.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

Kipengele chako cha `ul` kinapaswa kuwa na vipengele vitatu vya `li`.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
