---
id: 669fde2081f65141ad703fe4
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Sehemu ya kwanza kwenye ukurasa itakuwa sehemu ya kuhusu. Sehemu hiyo itamtambulisha Bwana Whiskers na kuwapa watumizi wazo la kile blogu hii inahusu.

Ndani ya kipengele chako cha `main`, ongeza kipengele cha `section` chenye sifa ya kipengele ya `id` iliyowekwa kuwa `"about"`.

Ndani ya kipengele cha `section`, ongeza `h2` yenye maandishi ya `About`.

# --hints--

Kipengele chako cha `section` kinapaswa kuwa na lebo ya kufungua.

```js
assert.match(code, /<section.*>/);
```

Kipengele chako cha `section` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/section\>/);
```

Kipengele chako cha `section` kinapaswa kuwa na sifa ya kipengele ya `id` iliyowekwa kuwa `"about"`.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

Kipengele chako cha `section` kinapaswa kuwa ndani ya kipengele chako cha `main`.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Kipengele chako cha `h2` kinapaswa kuwa na lebo ya kufungua.

```js
assert.match(code, /<h2\>/);
```

Kipengele chako cha `h2` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/h2\>/);
```

Kipengele chako cha `h2` kinapaswa kuwa na maandishi ya `About`. Hakikisha unakagua tahajia yako.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

Kipengele chako cha `h2` kinapaswa kuwa ndani ya kipengele chako cha `section`.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
