---
id: 66a335e5888ffb367633200a
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Chapisho hili la blogu litakuwa na aya chache zenye maandishi ya <dfn>lorem ipsum</dfn>.

Lorem ipsum hutumika sana katika ukuzaji wa mtandao kama maandishi ya kiwakilishi. Ni muhimu unapojikita katika kujenga muundo wa msingi wa kurasa zako za mtandao na usihangaike na maudhui halisi kwa sasa.

Hapa kuna mfano wa kutumia lorem ipsum:

```html
<p>
  Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
  quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
  quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
  quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
</p>
```

Chini ya kipengele chako cha `h3`, ongeza aya mbili za maandishi ya lorem ipsum.

# --hints--

Unapaswa kuwa na vipengele viwili vya aya ndani ya kipengele cha `article`.

```js
const paragraphElements = document.querySelectorAll('#posts article p');
assert.strictEqual(paragraphElements.length, 2);
```

Kipengele chako cha kwanza cha aya kinapaswa kuwa na maandishi ya lorem ipsum. Hakikisha hakikisha makosa ya tahajia na alama za uandishi.

```js
const paragraphElement = document.querySelector('article p:first-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim().replaceAll(/\s+/g, ' '), loremIpsum);
```

Kipengele chako cha pili cha aya kinapaswa kuwa na maandishi ya lorem ipsum. Hakikisha hakikisha makosa ya tahajia na alama za uandishi.

```js
const paragraphElement = document.querySelector('article p:last-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim().replaceAll(/\s+/g, ' '), loremIpsum);
```

Aya zako zinapaswa kuja baada ya kipengele cha `h3`.

```js
let h3 = document.querySelectorAll('article > p ~ h3');
assert.isEmpty(h3);
const paragraphs = document.querySelectorAll('article > h3 ~ p');
assert.lengthOf(paragraphs, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
        <article>
          <h3>Mr. Whiskers' First Day Home</h3>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </article>
      </section>
    </main>
  </body>
</html>
```
