---
id: 66a334a1a7cca6354999f9bf
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Kwa chapisho la kwanza la blogu, utatumia kipengele cha <dfn>article</dfn>.

Kipengele cha `article` kinawakilisha maudhui yaliyojitegemea kwenye ukurasa wa mtandao.

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

Chini ya kipengele cha `h2`, ongeza kipengele cha `article`.

Ndani ya kipengele cha `article`, ongeza kipengele cha `h3` chenye maandishi ya `Mr. Whiskers' First Day Home`.

Sababu ya kutumia `h3` hapa ni kwamba kudumisha mpangilio sahihi wa muundo kwa ajili ya vichwa vya sehemu ni muhimu. Kwa kuwa kichwa kidogo cha machapisho ni kipengele cha `h2`, kiwango kinachofuata chini katika muundo kitakuwa kipengele cha `h3`.

# --hints--

Msimbo wako unapaswa kuwa na lebo ya ufunguzi ya `article`.

```js
assert.match(code, /<article>/i);
```

Msimbo wako unapaswa kuwa na lebo ya kufunga ya `article`.

```js
assert.match(code, /<\/article>/i);
```

Kipengele chako cha `article` kinapaswa kuwa ndani ya kipengele chako cha `section` na `id` ikiwa imewekwa kuwa `posts`.

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

Msimbo wako unapaswa kuwa na lebo ya ufunguzi ya `h3`.

```js
assert.match(code, /<h3>/i);
```

Msimbo wako unapaswa kuwa na lebo ya kufunga ya `h3`.

```js
assert.match(code, /<\/h3>/i);
```

Kipengele chako cha `h3` kinapaswa kuwa na maandishi ya `Mr. Whiskers' First Day Home`.

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

Kipengele chako cha `h3` kinapaswa kuwa ndani ya kipengele chako cha `article`.

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
