---
id: 66a49f685961e997e337cab1
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

`header` itakuwa na jukumu la kuonyesha kichwa kikuu, picha na urambazaji wa ukurasa kwa blogu.

Ndani ya kipengele cha `header`, ongeza `h1` yenye maandishi ya `Welcome to Mr. Whiskers' Blog Page!`.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `h1`.

```js
assert.match(code, /<h1>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi ya `Welcome to Mr. Whiskers' Blog Page!`. Hakikisha upitie makosa yoyote ya tahajia.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

Kipengele chako cha `h1` kinapaswa kupangwa ndani ya kipengele chako cha `header`.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
