---
id: 66a7e72adf226c02626715a3
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Chini ya kipengele chako cha `h2`, ongeza kipengele cha aya chenye maandishi ya `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.`

Chini ya kipengele chako cha aya, ongeza kipengele kingine cha aya chenye maandishi ya `His playful nature and boundless energy keeps me on my toes. I love him so much.`

# --hints--

Aya yako ya kwanza inapaswa kuwa na maandishi ya `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.` Hakikisha unakagua tahajia yako kwa makini.

```js
const pElement = document.querySelector('body p:first-of-type');
const text = "Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers."
assert.strictEqual(pElement?.innerText.trim(), text);
```

Aya yako ya kwanza inapaswa kupangwa ndani ya kipengele cha `section`.

```js
const pElement = document.querySelector('#about p:first-of-type');
assert.strictEqual(pElement?.parentElement.tagName, 'SECTION');
```

Unapaswa kuwa na aya ya pili ndani ya kipengele chako cha `section`.

```js
const pElement = document.querySelectorAll('#about p');
assert.strictEqual(pElement?.length, 2);
```

Aya yako ya pili inapaswa kuwa na maandishi ya `His playful nature and boundless energy keeps me on my toes. I love him so much.` Hakikisha unakagua tahajia yako kwa makini.

```js
const pElement = document.querySelector('#about p:last-of-type');
const text = "His playful nature and boundless energy keeps me on my toes. I love him so much."
assert.strictEqual(pElement?.innerText.trim(), text);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
