---
id: 66ec4c8e9878d8441956516f
title: crea una tabella catalogo libri
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare un elemento `table` che elenchi le informazioni sui libri.
1. La tua tabella dovrebbe avere un elemento intestazione tabella con una riga al suo interno.
1. La riga nell'elemento intestazione tabella dovrebbe avere quattro elementi intestazione di tabella, con il testo `Title`, `Author`, `Genre` e `Publication Year`, in quest'ordine.
1. La tua tabella dovrebbe avere un elemento corpo tabella con almeno cinque righe al suo interno.
1. Ogni riga nel corpo tabella dovrebbe avere quattro elementi dati tabella che mostrano il titolo, l'autore, il genere e l'anno di pubblicazione del libro.
1. La tua tabella dovrebbe avere un elemento piè di pagina tabella con una riga al suo interno.
1. La riga nell'elemento piè di pagina tabella dovrebbe avere un elemento dati tabella che si estende su quattro colonne e ha il testo `Total Books: N` dove `N` dovrebbe essere sostituito dal numero di libri nella tua tabella.

# --hints--

Dovresti avere un elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Dovresti avere un elemento `thead` all'interno del tuo elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

Dovresti avere un elemento `tr` all'interno del tuo elemento `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

Dovresti avere quattro elementi `th` all'interno della riga del tuo elemento `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

I tuoi quattro elementi `th` dovrebbero avere il testo `Title`, `Author`, `Genre` e `Publication Year`, in quest'ordine.

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent.trim(), 'Title');
assert.equal(ths[1]?.textContent.trim(), 'Author');
assert.equal(ths[2]?.textContent.trim(), 'Genre');
assert.equal(ths[3]?.textContent.trim(), 'Publication Year');
```

Dovresti avere un elemento `tbody` all'interno del tuo elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

Il tuo elemento `tbody` dovrebbe avere almeno cinque righe.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

Ogni riga nel tuo elemento `tbody` dovrebbe avere esattamente quattro elementi `td` come figli.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

Ogni elemento `td` nel corpo tabella dovrebbe avere testo con informazioni sul libro.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.trim().length, 1);
});
```

Dovresti avere un elemento `tfoot` all'interno del tuo elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

Dovresti avere esattamente un elemento `tr` nel tuo elemento `tfoot`.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

L'elemento `td` nella riga del tuo elemento `tfoot` dovrebbe avere il suo attributo `colspan` impostato su `4`.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

L'elemento `td` nella riga del tuo elemento `tfoot` dovrebbe avere il testo `Total Books: N` dove `N` è il numero di libri nella tua tabella.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
