---
id: 66d6a7a3e1aa411e94bf2346
title: Ein Haus malen
challengeType: 25
dashedName: build-a-house-painting
demoType: onClick
---

# --description--

In diesem Labor verwenden Sie HTML, um die Struktur eines Hauses zu erstellen. Anschließend nutzen Sie CSS-Positionierung, um die Elemente Ihres Hauses wie Fenster und Türen anzuordnen.

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `div` mit einem `id` von `house` haben.
1. Ihr `#house` sollte eine `position` mit dem Wert `relative` haben, damit seine Kinder relativ zu ihm positioniert werden können.
1. Ihr `#house` sollte eine Breite von `500px` und eine Höhe von `400px` haben. Setzen Sie keine Minimum- oder Maximumwerte.
1. Ihr `#house` sollte eine Hintergrundfarbe und einen Rahmen haben.
1. Sie sollten fünf `div`-Elemente innerhalb von `#house` mit den folgenden `id`-Werten haben: `chimney`, `roof`, `window-1`, `window-2` und `door`.
1. Alle direkten Kinder von `#house` sollten eine `position` von `absolute` haben.
1. `#roof`, `#chimney`, `#window-1`, `#window-2` und `#door` sollten eine Breite, Höhe, einen Rahmen und eine Hintergrundfarbe haben.
1. Ihr `#roof` sollte einen oberen Wert von `0` haben.
1. Ihr `#door` sollte am unteren Rand Ihres Hauses platziert sein.
1. Ihre Fenster sollten unterhalb von `#roof` und mindestens höher als ein Drittel der Höhe von `#door` platziert sein.
1. Sowohl Ihre Fenster als auch Ihre Tür sollten entweder `left` oder `right` mit einem Wert haben, der sie innerhalb der Hausrahmen positioniert.
1. Ihr `#chimney` sollte einen oberen Wert haben, der es an den oberen Rand von `#house` setzt.
1. Ihr `#chimney` sollte eine `z-index` haben, die es hinter das Haus setzt.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verlinken, um Ihr CSS anzuwenden.

# --hints--

Sie sollten ein `div` mit einer `id` von `house` haben.

```js
assert.exists(document.querySelector("div#house"));
```

Sie sollten `#house` ansprechen und seine `position` auf `relative` setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("position"), "relative");
```

Sie sollten `#house` ansprechen und seine `border`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("border"));
```

Sie sollten `#house` ansprechen und seine `background-color`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("background-color"));
```

`#house` sollte keine `min-height`-Eigenschaft enthalten.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-height"));
```

`#house` sollte keine `min-width`-Eigenschaft enthalten.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-width"));
```

`#house` sollte keine `max-height`-Eigenschaft enthalten.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-height"));
```

`#house` sollte keine `max-width`-Eigenschaft enthalten.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-width"));
```

Sie sollten `#house` ansprechen und seine Breite auf `500px` setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("width"), "500px")
```

Sie sollten `#house` ansprechen und seine Höhe auf `400px` setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("height"), "400px")
```

Sie sollten ein `div` mit einer `id` von `chimney` innerhalb von `#house` haben.

```js
assert.exists(document.querySelector("div#house div#chimney"));
```

`#chimney` sollte die `position`-Eigenschaft auf `absolute` gesetzt haben.

```js
assert.equal(getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("position"), "absolute");
```

Sie sollten `#chimney` ansprechen und seine `width`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("width"));
```

Sie sollten `#chimney` ansprechen und seine `height`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("height"));
```

Sie sollten `#chimney` ansprechen und seine `border`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("border"));
```

Sie sollten `#chimney` ansprechen und seine `background-color`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("background-color"));
```

Sie sollten ein `div` mit einer `id` von `roof` innerhalb von `#house` haben.

```js
assert.exists(document.querySelector("div#house  div#roof"));
```

`#roof` sollte die `position`-Eigenschaft auf `absolute` gesetzt haben.

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("position"), "absolute");
```

Sie sollten `#roof` ansprechen und seine `width`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("width"));
```

Sie sollten `#roof` ansprechen und seine `height`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("height"));
```

Sie sollten `#roof` ansprechen und seine `border`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("border"));
```

Sie sollten `#roof` ansprechen und seine `background-color`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("background-color"));
```

Sie sollten ein `div` mit einer `id` von `window-1` innerhalb von `#house` haben.

```js
assert.exists(document.querySelector("div#house  div#window-1"));
```

`#window-1` sollte die `position`-Eigenschaft auf `absolute` gesetzt haben.

```js
assert.equal(getComputedStyle(document.querySelector("#window-1"))?.getPropertyValue("position"), "absolute");
```

Sie sollten `#window-1` ansprechen und seine `width`-Eigenschaft setzen.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window1 || windows);
```

Sie sollten `#window-1` ansprechen und seine `height`-Eigenschaft setzen.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window1 || windows);
```

Sie sollten `#window-1` ansprechen und seine `border`-Eigenschaft setzen.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window1 || windows);
```

Sie sollten `#window-1` ansprechen und seine `background-color`-Eigenschaft setzen.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window1 || windows);
```

Sie sollten ein `div` mit einer `id` von `window-2` innerhalb von `#house` haben.

```js
assert.exists(document.querySelector("div#house  div#window-2"));
```

`#window-2` sollte die `position`-Eigenschaft auf `absolute` gesetzt haben.

```js
assert.equal(getComputedStyle(document.querySelector("#window-2"))?.getPropertyValue("position"), "absolute");
```

Sie sollten `#window-2` ansprechen und seine `width`-Eigenschaft setzen.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window2 || windows);
```

Sie sollten `#window-2` ansprechen und seine `height`-Eigenschaft setzen.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window2 || windows);
```

Sie sollten `#window-2` ansprechen und seine `border`-Eigenschaft setzen.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window2 || windows);
```

Sie sollten `#window-2` ansprechen und seine `background-color`-Eigenschaft setzen.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window2 || windows);
```

Sie sollten ein `div` mit einer `id` von `door` innerhalb von `#house` haben.

```js
assert.exists(document.querySelector("div#house  div#door"));
```

`#door` sollte die `position`-Eigenschaft auf `absolute` gesetzt haben.

```js
assert.equal(getComputedStyle(document.querySelector("#door"))?.getPropertyValue("position"), "absolute");
```

Sie sollten `#door` ansprechen und seine `width`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("width"));
```

Sie sollten `#door` ansprechen und seine `height`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("height"));
```

Sie sollten `#door` ansprechen und seine `border`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("border"));
```

Sie sollten `#door` ansprechen und seine `background-color`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("background-color"));
```

Ihr `#roof` sollte die `top`-Eigenschaft auf `0` gesetzt haben.

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("top"), "0px");
```

Ihr `#door` sollte am unteren Rand des Hauses platziert sein.

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorBottom = Number(door.getPropertyValue("bottom").replace("px", ""));
assert.strictEqual(doorBottom, 0);
```

Sie sollten entweder `left` oder `right` auf Ihrem `#door` setzen und es innerhalb der Hausrahmen anordnen.

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorLeft = Number(door.getPropertyValue("left").replace("px", ""));
const doorRight = Number(door.getPropertyValue("right").replace("px", ""));
assert.isTrue(doorLeft > 0 || doorRight > 0);
```

Ihr `#chimney` sollte einen `top`-Wert haben, der es an den oberen Rand von `#house` setzt.

```js
const chimney = getComputedStyle(document.querySelector("#chimney"));
assert.approximately(parseFloat(chimney?.getPropertyValue("top")), - parseFloat(chimney?.getPropertyValue("height")), 2);
```

Ihr `#chimney` sollte eine `z-index` haben, die es hinter das Haus setzt.

```js
const houseZIndex = getComputedStyle(document.querySelector("#house"))?.getPropertyValue("z-index");
const chimneyZIndex = getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("z-index");
if (houseZIndex === "auto") {
    assert.isBelow(Number(chimneyZIndex), 0)
} else {
    assert.isBelow(Number(chimneyZIndex), Number(houseZIndex));
}
```

Ihre Fenster sollten unterhalb von `#roof` und mindestens höher als ein Drittel der Höhe von `#door` platziert sein.

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const door = getComputedStyle(document.querySelector("#door"))
const roof = getComputedStyle(document.querySelector("#roof"))
const w1Top = Number(window1.getPropertyValue("top").replace("px", ""));
const w1Bottom = Number(window1.getPropertyValue("bottom").replace("px", ""));
const w2Top = Number(window2.getPropertyValue("top").replace("px", ""));
const w2Bottom = Number(window2.getPropertyValue("bottom").replace("px", ""));
const doorHeight = Number(door.getPropertyValue("height").replace("px", ""));
const roofHeight = Number(roof.getPropertyValue("height").replace("px", ""));
assert.isAbove(w1Top, roofHeight);
assert.isAbove(w1Bottom, doorHeight/3);
assert.isAbove(w2Top, roofHeight);
assert.isAbove(w2Bottom, doorHeight/3);
```

Sie sollten entweder `left` oder `right` auf Ihren Fenstern setzen und sie innerhalb der Hausrahmen anordnen.

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const w1Left = Number(window1.getPropertyValue("left").replace("px", ""));
const w1Right = Number(window1.getPropertyValue("right").replace("px", ""));
const w2Left = Number(window2.getPropertyValue("left").replace("px", ""));
const w2Right = Number(window2.getPropertyValue("right").replace("px", ""));
assert.isTrue(w1Left > 0 || w1Right > 0);
assert.isTrue(w2Left > 0 || w2Right > 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div id="house">
        <div id="chimney">
          <div id="smoke"></div>
        </div>
        <div id="roof"></div>
        <div id="window-1"></div>
        <div id="window-2"></div>
        <div id="door">
          <div id="door-knob"></div>
        </div>
        <div id="welcome">WELCOME</div>
    </div>
</body>
</html>
```

```css
* {
    box-sizing: border-box;
  }

  body {
    margin: 0;
    padding: 0;
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    background: linear-gradient(
      #b3e6ff,
      #b3e6ff 60%,
      green 60%
    );
  }

  #house {
    flex: 0 0 auto;
    width: 500px;
    height: 400px;
    position: relative;
    background-color: #ff9980;
    border: 6px solid #b35900;
  }

  #chimney {
    position: absolute;
    width: 90px;
    height: 100px;
    top: -100px;
    left: 65%;
    background: repeating-linear-gradient(
      #e6e6e6,
      #e6e6e6 10%,
      black 10%,
      black 12%
    );
    border: 2px solid black;
    z-index: -1;
  }

  #smoke {
    width: 40px;
    height: 40px;
    background-color: rgba(0,0,0,0.3);
    position: absolute;
    top: -20px;
    left: 20px;
    border-radius: 50%;
    box-shadow: 0 -10px 10px 10px rgba(0,0,0,0.3);
    border: none;
  }

  #roof {
    border: 2px;
    width: 100%;
    height: 110px;
    background-color: green;
    position: absolute;
    top: 0;
    background: repeating-linear-gradient(
      45deg,
      #b35900,
      #b35900 2%,
      transparent 2%,
      transparent 5%
    ), repeating-linear-gradient(
      -45deg,
      #b35900,
      #b35900 2%,
      #ff9980 2%,
      #ff9980 5%
    )
  }

  #window-1, #window-2 {
    width: 100px;
    height: 100px;
    background: linear-gradient(
      to right,
      #ffffb3,
      #ffffb3 48%,
      #b35900 48%,
      #b35900 52%,
      #ffffb3 52%,
      #ffffb3 100%
    );
    position: absolute;
    top: 40%;
    border: 6px solid #b35900;
  }

  #window-1 {
    left: 7.5%;
  }

  #window-2 {
    right: 7.5%;
  }

  #door {
    width: 130px;
    height: 180px;
    position: absolute;
    bottom: 0%;
    left: 185px;
    background-color: #e6e6e6;
    border: 6px solid #b35900;
    transform: translateY(6px);
  }

  #door-knob {
    width: 20px;
    height: 20px;
    background-color: #b35900;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    right: 5%;
  }

  #welcome {
    font-family: sans-serif;
    font-weight: bold;
    width: 130px;
    height: 40px;
    position: absolute;
    bottom: -47px;
    left: 178px;
    background-color: #85e085;
    border: 2px solid black;
    transform: skewX(-20deg);
    display: flex;
    align-items: center;
    justify-content: center;
  }
```
