---
id: bad87fee1348bd9acdd08826
title: Impara come funzionano i tag script e document ready
challengeType: 6
forumTopicId: 18224
dashedName: learn-how-script-tags-and-document-ready-work
---

# --description--

Ora siamo pronti per imparare jQuery, lo strumento JavaScript più popolare di tutti i tempi.

Prima di poter iniziare a usare jQuery, dobbiamo aggiungere alcune cose al nostro HTML.

Per prima cosa, aggiungi un elemento `script` all'inizio della tua pagina. Assicurati di chiuderlo nella riga successiva.

Il tuo browser eseguirà qualsiasi JavaScript all'interno di un elemento `script`, incluso jQuery.

All'interno del tuo elemento `script`, aggiungi questo codice: `$(document).ready(function() {` al tuo `script`. Poi chiudilo nella riga successiva (ancora all'interno del tuo elemento `script`) con: `});`

Impareremo di più su `functions` più avanti. La cosa importante da sapere è che il codice che metti dentro questo `function` verrà eseguito non appena il browser avrà caricato la tua pagina.

Questo è importante perché senza il tuo `document ready function`, il codice potrebbe essere eseguito prima che il tuo HTML venga renderizzato, causando bug.

# --hints--

Dovresti creare un elemento `script` assicurandoti che sia valido e abbia un tag di chiusura.

```js
assert(
  code.match(/<\/script\s*>/g) &&
    code.match(
      /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
    ) &&
    code.match(/<\/script\s*>/g).length ===
      code.match(
        /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
      ).length
);
```

Dovresti aggiungere `$(document).ready(function() {` all'inizio del tuo elemento `script`.

```js
assert(
  code.match(
    /\$\s*?\(\s*?document\s*?\)\.ready\s*?\(\s*?function\s*?\(\s*?\)\s*?\{/g
  )
);
```

Dovresti chiudere la tua funzione `$(document).ready(function() {` con `});`

```js
assert(code.match(/\n*?\s*?\}\s*?\);/g));
```

# --seed--

## --seed-contents--

```html
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
  });
</script>
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
