---
id: bad87fee1348bd9bedc08826
title: Selezionare elementi HTML con selettori usando jQuery
challengeType: 6
forumTopicId: 18319
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-html-elements-with-selectors-using-jquery
---

# --description--

Ora abbiamo una funzione `document ready`.

Scriviamo la nostra prima istruzione jQuery. Tutte le funzioni jQuery iniziano con un `$`, solitamente chiamato operatore del dollaro, o bling.

jQuery spesso seleziona un elemento HTML con un <dfn>selettore</dfn>, poi fa qualcosa a quell'elemento.

Per esempio, facciamo rimbalzare tutti i tuoi elementi `button`. Basta aggiungere questo codice dentro la funzione document ready:

```js
$("button").addClass("animated bounce");
```

Nota che abbiamo già incluso sia la libreria jQuery sia la libreria Animate.css in background così puoi usarle nell'editor. Quindi stai usando jQuery per applicare la classe `bounce` di Animate.css ai tuoi elementi `button`.

# --hints--

Dovresti usare la funzione `addClass()` di jQuery per dare le classi `animated` e `bounce` ai tuoi elementi `button`.

```js
assert($('button').hasClass('animated') && $('button').hasClass('bounce'));
```

Dovresti usare jQuery solo per aggiungere queste classi all'elemento.

```js
assert(!code.match(/class.*animated/g));
```

Il tuo codice jQuery dovrebbe essere all'interno della funzione `$(document).ready();`.

```js
assert(
  code.replace(/\s/g, '').match(/\$\(document\)\.ready\(function\(\)\{\$/g)
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
