---
id: 69a012c6224fac85832247dd
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Haben Sie die Ergebnisse `true`, `{ head: null, length: 0 }` und `true` in der Konsole bemerkt? Das passiert, weil Sie der Liste `myList` noch nichts zugewiesen haben und die Eigenschaft `length` nicht inkrementiert wurde.

Weisen Sie `node` der `list.head` zu, und erhöhen Sie darunter die Eigenschaft `length` von `list`.

# --hints--

Sie sollten `node` `list.head` zuweisen.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.head\s*=\s*node;?/);
```

Sie sollten die Eigenschaft `length` von `list` inkrementieren.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.length\+\+;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
