---
id: 69a04712866a37ea8f38b6bd
title: Schritt 21
challengeType: 1
dashedName: step-21
---

# --description--

Jetzt sollten Sie die `remove`-Funktion testen.

Entfernen Sie eine der Zahlen aus der Liste, indem Sie die `remove`-Funktion mit `myList` und der zu entfernenden Zahl aufrufen, und geben Sie dann `myList` mit `console.log(JSON.stringify(myList, null, 2))` in der Konsole aus, um die vollständige aktualisierte Liste zu sehen.

Damit ist Ihr Workshop zur verketteten Liste abgeschlossen!

# --hints--

Sie sollten eine der vorhandenen Zahlen mit der `remove`-Funktion entfernen.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
const removalTargets = ['42', '43', '44'];
const removedOne = removalTargets.some(n => strippedCode.includes(`remove(myList,${n})`));
assert.isTrue(removedOne);
```

Sie sollten `myList` erneut in der Konsole ausgeben, diesmal mit `JSON.stringify(myList, null, 2)`.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(JSON.stringify(myList,null,2))');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    } else {
        list.head = current.next;
    }

    list.length--;
    
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    } else {
        list.head = current.next;
    }

    list.length--;
    
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
remove(myList, 43);
console.log(JSON.stringify(myList, null, 2));
```
