---
id: 69a014bd9b1730e94136279f
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Wenn die Liste nicht leer ist, möchten Sie ihr Ende finden, das neue Element dort hinzufügen und es auf `null` zeigen lassen. Das bedeutet, dass Sie die Liste ab `head` durchlaufen müssen.

Um damit zu beginnen, öffnen Sie eine `else`-Klausel. Verwenden Sie darin `let`, um eine `current`-Variable zu erstellen und setzen Sie sie auf `list.head`.

# --hints--

Sie sollten eine `else`-Klausel haben.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'else{');
```

Sie sollten eine `current`-Variable erstellen, die auf `list.head` gesetzt wird, innerhalb der `else`-Klausel.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.match(strippedAdd, /else\{(?:let|var)current=list\.head/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
