---
id: 69a00f673f171e29159e3ad8
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Dieser Knoten benötigt zwei Eigenschaften: das hinzuzufügende Element und einen Zeiger auf den nächsten Knoten.

Erstellen Sie innerhalb der Funktion add ein Objekt namens `node`. Es sollte eine Eigenschaft namens `element` besitzen und eine `next`-Eigenschaft, die auf `null` gesetzt ist, den Zeiger auf das nächste Element der Liste.

# --hints--

Sie sollten innerhalb Ihrer `node`-Funktion ein Objekt namens `add` erstellen.

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

Ihr `node`-Objekt sollte eine Eigenschaft namens `element` besitzen, die auf `element` gesetzt ist.

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

Ihr `node`-Objekt sollte eine Eigenschaft namens `next` besitzen, die auf `null` gesetzt ist.

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
