---
id: 69a00cd4ec513345d40e95ec
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Als Nächstes arbeiten Sie an einer Funktion, mit der Sie überprüfen können, ob die Liste leer ist.

Erstellen Sie dazu eine Funktion namens `isEmpty`, die einen `list`-Parameter entgegennimmt. Geben Sie in der Funktion eine Überprüfung zurück, ob die `length`-Eigenschaft von `list` strikt gleich `0` ist.

# --hints--

Sie sollten eine `isEmpty`-Funktion erstellen.

```js
assert.isFunction(isEmpty)
```

Ihre `isEmpty`-Funktion sollte einen `list`-Parameter haben.

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

Ihre `isEmpty`-Funktion sollte `list.length === 0` zurückgeben.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
