---
id: 69a012c6224fac85832247dd
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Haben Sie die Ergebnisse `true`, `{ head: null, length: 0 }` und `true` in der Konsole bemerkt? Das passiert, weil Sie der Liste `myList` noch nichts zugewiesen haben und die Eigenschaft `length` nicht inkrementiert wurde.

Weisen Sie `node` der `list.head` zu, und erhöhen Sie darunter die Eigenschaft `length` von `list`.

# --hints--

Sie sollten `node` `list.head` zuweisen.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.head=node');
```

Sie sollten die Eigenschaft `length` von `list` inkrementieren.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.length++');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
