---
id: 69a03aaa1e6fe8152788ce6a
title: Schritt 14
challengeType: 1
dashedName: step-14
---

# --description--

Innerhalb der `remove`-Funktion verwenden Sie `let`, um eine `previous`-Variable mit dem Wert `null` zu erstellen, und erneut `let`, um eine `current`-Variable mit dem Wert `list.head` zu erstellen. Diese Variablen verwenden Sie in der Schleife, die die Liste durchläuft.

`current` verfolgt den Knoten, der gerade in der Liste überprüft wird, beginnend am Kopf und jeweils einen Knoten weiter. `previous` hingegen verfolgt den Knoten direkt vor dem aktuellen.

Das ist wichtig beim Entfernen eines Knotens, weil Sie den next-Wert des vorherigen Knotens aktualisieren müssen, sodass er auf den Knoten nach dem zu entfernenden zeigt, um die Liste verbunden zu halten.

# --hints--

Sie sollten `let` verwenden, um innerhalb Ihrer `previous`-Funktion eine `null`-Variable mit dem Wert `remove` zu erstellen.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

Sie sollten `let` verwenden, um innerhalb Ihrer `current`-Funktion eine `list.head`-Variable mit dem Wert `remove` zu erstellen.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
```
