---
id: 69a04466917c66b132323a1b
title: Schritt 18
challengeType: 1
dashedName: step-18
---

# --description--

Wenn `previous` nicht `null` ist, bedeutet das, dass das zu entfernende Element nicht am Anfang, sondern irgendwo innerhalb der Liste steht. Daher sollte der aktuelle Knoten übersprungen werden, indem der vorherige Knoten direkt mit dem nächsten Knoten verbunden wird.

Dazu erstellen Sie eine weitere `if`-Anweisung, die überprüft, ob previous strikt nicht `null` ist. Falls das zutrifft, setzen Sie `previous.next` auf `current.next`.

# --hints--

Sie sollten eine weitere `if`-Anweisung mit der Bedingung `previous !== null` erstellen.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

Sie sollten `previous.next` auf `current.next` setzen.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
