---
id: 69a030fbe487f185f55b8627
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Jetzt arbeiten Sie an einer Funktion, um einen Knoten aus der Liste zu entfernen.

Erstellen Sie eine `remove`-Funktion mit den Parametern `list` und `element`.

# --hints--

Sie sollten eine `remove`-Funktion erstellen.

```js
assert.isFunction(remove)
```

Ihre `remove`-Funktion sollte die Parameter `list` und `element` haben.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.remove.parameters;
assert.equal(parameters[0].toString(), "list")
assert.equal(parameters[1].toString(), "element")
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

--fcc-editable-region--

--fcc-editable-region--

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
```
