---
id: 69a00eaa43458e618331873c
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Jetzt arbeiten Sie an der Funktion, die einen Knoten zur Liste hinzufügt.

Erstellen Sie eine `add`-Funktion mit den Parametern `list` und `element`.

# --hints--

Sie sollten eine `add`-Funktion erstellen.

```js
assert.isFunction(add)
```

Ihre `add`-Funktion sollte die Parameter `list` und `element` haben.

```js
const addParams = __helpers.getFunctionParams(add.toString());
const addParamsNames = addParams.map(param => param.name)
assert.include(addParamsNames, "list");
assert.include(addParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
