---
id: 69a045fd236d0d973a59f103
title: Schritt 19
challengeType: 1
dashedName: step-19
---

# --description--

Wenn `previous` `null` ist, bedeutet das, dass das zu entfernende Element am Anfang der Liste steht. Daher sollten Sie den `head` aktualisieren, damit er auf den nächsten Knoten zeigt, um das erste Element zu entfernen.

Fügen Sie dazu eine `else`-Klausel hinzu. Setzen Sie darin `list.head` auf `current.next`.

# --hints--

Sie sollten eine `else`-Klausel zu Ihrer `if`-Anweisung hinzufügen.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

Sie sollten `list.head` innerhalb der `current.next`-Klausel auf `else` setzen.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
