---
id: 69a03ddce82b0626f329b001
title: Schritt 15
challengeType: 1
dashedName: step-15
---

# --description--

Erstellen Sie nun eine `while`-Schleife, die überprüft, ob `current` nicht strikt `null` ist und `current.element` nicht strikt gleich dem Element ist.

Das bedeutet, die Schleife soll die Liste weiter durchlaufen, bis das Ende der Liste erreicht ist oder das zu entfernende Element gefunden wurde.

# --hints--

Sie sollten eine `while`-Schleife mit der Bedingung `current !== null && current.element !== element` erstellen.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /while\s+\(current\s*!==\s*null\s*&&\s*current\.element\s*!==\s*element\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    --fcc-editable-region--
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
