---
id: 69a0179a85b00b6127e0af45
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Innerhalb der `while`-Schleife aktualisieren Sie `current` zu `current.next`, sodass die Schleife die Liste jeweils um einen Knoten vorwärts durchläuft und stoppt, wenn sie den Knoten mit einer `next`-Eigenschaft von `null` erreicht, also das Ende der Liste.

# --hints--

Sie sollten den Wert Ihrer `current`-Variablen innerhalb der `current.next`-Schleife auf `while` aktualisieren.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
