---
id: 69a04088e600bd8074a9852f
title: Schritt 17
challengeType: 1
dashedName: step-17
---

# --description--

Wenn `current` `null` ist, bedeutet das, dass das Ende der Liste erreicht ist, und die Funktion sollte aufhören zu laufen.

Erstellen Sie eine `if`-Anweisung, die überprüft, ob `current` strikt gleich `null` ist. Verwenden Sie innerhalb der `if`-Anweisung `return`, um die Funktion zu beenden.

# --hints--

Sie sollten eine `if`-Anweisung mit der Bedingung `current === null` erstellen.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*/);
```

Sie sollten das Schlüsselwort `return` innerhalb Ihrer `if`-Anweisung verwenden.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*return;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
