---
id: 69a00a051b9a18694d00c16b
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop implementieren Sie eine verkettete Liste.

Erstellen Sie zunächst eine Funktion namens `initList`. Geben Sie darin ein Objekt zurück, das eine `head`-Eigenschaft mit dem Wert `null` und eine `length`-Eigenschaft mit dem Wert `0` enthält.

Dies ist die Funktion, die Sie aufrufen, um eine verkettete Liste zu initialisieren. Die `head` stellt das erste Element in der Liste dar, anfänglich ist sie `null`, da sich kein Element in der Liste befindet. Die `length` verfolgt die Anzahl der Knoten (Elemente) in der Liste, die Sie beim Hinzufügen oder Entfernen eines Knotens erhöhen oder verringern.

# --hints--

Sie sollten eine `initList`-Funktion erstellen.

```js
assert.isFunction(initList)
```

Ihre `initList`-Funktion sollte ein Objekt zurückgeben.

```js
assert.isObject(initList());
```

Ihre `initList`-Funktion sollte ein Objekt mit einer `head` zurückgeben, die auf `null` gesetzt ist.

```js
assert.propertyVal(initList(), "head", null);
```

Ihre `initList`-Funktion sollte ein Objekt mit einer `length` zurückgeben, die auf `0` gesetzt ist.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
