---
id: 69a0102e761a245bb1cbd8e7
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

An diesem Punkt sollten Sie die verkettete Liste testen, um sie besser zu verstehen.

Erstellen Sie eine `myList`-Variable und setzen Sie sie auf `initList()`. Darunter protokollieren Sie einen Aufruf von `isEmpty()` mit `myList` als Parameter in der Konsole. Rufen Sie danach `add` mit `myList` und einer Zahl Ihrer Wahl auf, und protokollieren Sie dann `myList` in der Konsole. Darunter protokollieren Sie erneut einen Aufruf von `isEmpty()` in der Konsole, wiederum mit `myList` als Parameter.

# --hints--

Sie sollten eine `myList`-Variable erstellen, die auf `initList()` gesetzt ist.

```js
assert.deepEqual(myList, initList());
```

Sie sollten einen Aufruf von `isEmpty(myList)` in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

Sie sollten `add` mit `myList` und einer Zahl Ihrer Wahl als Parameter aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

Sie sollten Ihre `myList`-Variable in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

Sie sollten `isEmpty(myList)` erneut in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
