---
id: 69a013a86c97dac3784b2636
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Jetzt sollte die Ausgabe in der Konsole etwas wie `{ head: { element: 42, next: null }, length: 1 }` anzeigen.

Dies ist jedoch nicht die richtige Vorgehensweise, wenn der hinzugefügte Knoten das erste Element in der Liste ist, also der Kopf.

Um das korrekt zu machen, verwenden Sie eine if-Anweisung, um zu prüfen, ob die Liste leer ist, indem Sie die `isEmpty`-Funktion mit dem `list` aufrufen. Wenn das zutrifft, verschieben Sie Ihre `list.head = node`-Zeile in den Körper der `if`-Anweisung.

# --hints--

Sie sollten eine `if`-Anweisung mit der Bedingung `isEmpty(list)` haben.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

Sie sollten Ihre `list.head = node`-Zeile in die `if`-Anweisung verschieben.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
