---
id: 69a02ee82bbeaee00e44c67d
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Das ist alles, was Sie brauchen, um ein Element zur Liste hinzuzufügen. Um zu sehen, wie das aussieht, fügen Sie der `myList`-Liste zwei weitere Zahlen Ihrer Wahl hinzu.

# --hints--

Sie sollten der `myList`-Liste zwei weitere Zahlen Ihrer Wahl hinzufügen.

```js
const matches = __helpers.removeJSComments(code).match(/add\s*\(\s*myList\s*,\s*\d+\s*\);?/g) || [];

assert.isAtLeast(matches.length, 3);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
--fcc-editable-region--

--fcc-editable-region--
console.log(myList);
console.log(isEmpty(myList));
```
