---
id: 66ad9c633fa26d3c1475eae3
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

حان الوقت الآن لإضافة رسالة أخرى من الروبوت.

أضف بيان `console` آخر إلى الكود يسجل الرسالة `"Allow me to introduce myself."`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب أن يحتوي كودك على بيان `console.log()` ثالث.

```js
assert.lengthOf(logSpy.calls, 3);
```

يجب أن يحتوي بيان `console` على الرسالة `"Allow me to introduce myself."`. لا تنسَ علامات الاقتباس.

```js
assert.isTrue(
  logSpy.calls.some(
    (call) => /Allow\s+me\s+to\s+introduce\s+myself\s*\./.test(call[0])
  )
);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
