---
id: 66ad8ab945d266383f318cbf
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

الآن يجب أن ترى الرسالة الأولى من الروبوت في وحدة التحكم.

حان الوقت لإضافة رسالة ثانية من الروبوت.

أضف بيان `console.log` آخر لإخراج الرسالة `"I am excited to talk to you."` في وحدة التحكم.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب أن يحتوي كودك على جملة `console.log()` ثانية.

```js
assert.lengthOf(logSpy.calls, 2);
```

يجب أن يحتوي بيان `console` على الرسالة `"I am excited to talk to you."`. لا تنسَ علامات الاقتباس.

```js
assert.isTrue(
  logSpy.calls.some(
    (call) => /I\s+am\s+excited\s+to\s+talk\s+to\s+you\s*\./.test(call[0])
  )
);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
