---
id: 66ad9c633fa26d3c1475eae3
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

حان الوقت الآن لإضافة رسالة أخرى من الروبوت.

أضف بيان `console` آخر إلى الشفرة يسجل الرسالة `"Allow me to introduce myself."`.

# --hints--

يجب أن يكون لديك بيان `console.log()` ثالث في الشفرة.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

يجب أن يحتوي بيان `console` على الرسالة `"Allow me to introduce myself."`. لا تنسَ علامات الاقتباس.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
