---
id: 66ad998285df023b23dacdd3
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

حان الوقت الآن لتهيئة المتغير `botLocation`.

عندما تحتاج إلى إعلان متغيرات تتكون من عدة كلمات، يمكنك استخدام اصطلاح التسمية <dfn>تنسيق سنام الجمل</dfn>.

عند استخدام `camelCase`، تكون الكلمة الأولى كلها بأحرف صغيرة وتُكتب الحرف الأول من كل كلمة لاحقة بحرف كبير.

إليك مثالاً:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

أعلن عن المتغير `botLocation` وعيّن له السلسلة النصية `"the universe"` في نفس السطر باستخدام الكلمة المفتاحية `let`.

# --hints--

استخدم الكلمة المفتاحية `let` لإعلان متغير.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

يجب أن يكون اسم متغيرك `botLocation`.

```js
assert.isNotNull(botLocation);
```

يجب أن تعيّن سلسلة نصية إلى متغير `botLocation` الخاص بك.

```js
assert.isString(botLocation);
```

يجب أن يُعيَّن لمتغير `botLocation` السلسلة النصية `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
