---
id: 66adb844118ba74107ce771f
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

الرسالة التالية من الروبوت ستتعلق بموقع الروبوت.

أنشئ متغيرًا اسمه `botLocationSentence`.

ثم استخدم ربط السلاسل النصية مع عامل `+` لضم السلسلة النصية `"I live in "` مع المتغير `botLocation` متبوعًا بنقطة (`.`).

عيّن هذه القيمة إلى المتغير `botLocationSentence`.

ثم سجّل قيمة `botLocationSentence` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير `botLocationSentence`.

```js
assert.isNotNull(botLocationSentence);
```

يجب أن يحمل متغير `botLocationSentence` قيمة سلسلة نصية.

```js
assert.isString(botLocationSentence);
```

يجب أن تستخدم ربط السلاسل النصية مع عامل `+` لضم السلسلة النصية `"I live in "` مع المتغير `botLocation` متبوعًا بنقطة (`.`). انتبه للمسافات.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

يجب أن يكون لديك عبارة `console` في شفرتك.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

يجب أن تُخرج عبارة `console` المتغير `botLocationSentence`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
