---
id: 66ad984e92ce213a66f7e51c
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

حان الوقت الآن لتعيين قيمة للمتغير `bot`.

في الدروس السابقة، تعلمت كيف تعيّن القيم للمتغيرات بهذه الطريقة:

```js
variableName = "Here is the value";
```

تذكر أن ما على الجانب الأيمن من معامل التعيين `=` هو القيمة التي تعيّنها للمتغير على الجانب الأيسر.

عيّن السلسلة النصية `"teacherBot"` للمتغير `bot`.

# --hints--

يجب أن تعيّن سلسلة نصية للمتغير `bot`.

```js
assert.isString(bot);
```

يجب أن يكون المتغير `bot` مُعيَّنًا بالسلسلة النصية `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
