---
id: 66adc6046acc18453decf577
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

لرؤية اللقب الجديد للروبوت، يجب تسجيل رسالة جديدة في وحدة التحكم.

أنشئ متغيرًا جديدًا باسم `newNicknameGreeting`.

ثم استخدم ربط السلاسل النصية لضم السلسلة النصية `"I love my nickname but I wish people would call me "` مع المتغير `bot` متبوعًا بنقطة.

عيّن النتيجة إلى المتغير `newNicknameGreeting`.

ثم سجّل قيمة `newNicknameGreeting` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير باسم `newNicknameGreeting`.

```js
assert.isNotNull(newNicknameGreeting);
```

يجب أن يكون المتغير `newNicknameGreeting` سلسلة نصية.

```js
assert.isString(newNicknameGreeting);
```

يجب أن تستخدم ربط السلاسل النصية لضم السلسلة النصية `"I love my nickname but I wish people would call me "` مع المتغير `bot` متبوعًا بنقطة.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

يجب أن يحتوي شفرتك على عبارة `console`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

يجب أن تسجل قيمة `newNicknameGreeting` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
