---
id: 66adc82170acb1464bd348eb
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

الرسائل القليلة الأخيرة من الروبوت ستركز على الموضوع المفضل للروبوت.

ابدأ بإنشاء متغير جديد اسمه `favoriteSubject` وقم بتعيين السلسلة النصية `"Computer Science"` له.

# --hints--

يجب أن يكون لديك متغير اسمه `favoriteSubject`.

```js
assert.isNotNull(favoriteSubject);
```

يجب أن يكون المتغير `favoriteSubject` سلسلة نصية.

```js
assert.isString(favoriteSubject);
```

يجب تعيين السلسلة النصية `"Computer Science"` إلى المتغير `favoriteSubject`.

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
