---
id: 66adcf383276814776aba3ca
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

بعد ذلك، أنشئ متغيرًا يسمى `favoriteSubjectSentence`.

استخدم دمج السلاسل النصية لربط السلسلة النصية `"My favorite subject is "` مع المتغير `favoriteSubject` متبوعًا بنقطة.

عيّن النتيجة إلى المتغير `favoriteSubjectSentence`.

ثم، سجّل قيمة `favoriteSubjectSentence` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير يسمى `favoriteSubjectSentence`.

```js
assert.isNotNull(favoriteSubjectSentence);
```

يجب أن يكون المتغير `favoriteSubjectSentence` سلسلة نصية.

```js
assert.isString(favoriteSubjectSentence);
```

يجب أن تستخدم دمج السلاسل النصية لربط السلسلة النصية `"My favorite subject is "` مع المتغير `favoriteSubject` متبوعًا بنقطة. انتبه إلى الفراغات.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

يجب أن يكون لديك عبارة `console` في شفرتك.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

يجب أن تسجل قيمة `favoriteSubjectSentence` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
