---
id: 66adc42868cab843ccee87d9
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

حان الوقت الآن لرؤية القيمة الجديدة لـ `bot`.

ابدأ بإنشاء متغير جديد اسمه `nicknameIntroduction`.

استخدم ربط السلاسل النصية لضم السلسلة النصية `"My nickname is "` مع المتغير `bot` متبوعًا بنقطة (`.`).

عيّن السلسلة النصية الناتجة إلى المتغير `nicknameIntroduction`.

ثم، سجّل قيمة `nicknameIntroduction` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير اسمه `nicknameIntroduction`.

```js
assert.isNotNull(nicknameIntroduction);
```

يجب أن يكون المتغير `nicknameIntroduction` سلسلة نصية.

```js
assert.isString(nicknameIntroduction);
```

يجب أن تستخدم ربط السلاسل النصية لضم السلسلة النصية `"My nickname is "` مع المتغير `bot` متبوعًا بنقطة (`.`). انتبه إلى الفراغات.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

يجب أن يكون لديك عبارة `console` في شفرتك.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

يجب أن تسجل قيمة `nicknameIntroduction` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
