---
id: 66ada3f46945763dd97f43f8
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

في السابق، أنشأت المتغيرين `bot` و `botLocation`. الآن، ستستخدمهما لإخراج رسائل جديدة إلى وحدة التحكم.

في الدروس السابقة، تعلمت كيفية العمل مع ربط السلاسل النصية باستخدام العامل `+` لربط السلاسل معًا مثل هذا:

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

تذكر أنه يجب الانتباه إلى الفراغات عند ربط السلاسل النصية مع المتغيرات.

أنشئ متغيرًا يسمى `botIntroduction`.

ثم استخدم ربط السلاسل النصية مع العامل `+` لضم السلسلة النصية `"My name is "` متبوعة بالمتغير `bot` ثم نقطة (`.`).

عيّن هذه القيمة إلى المتغير `botIntroduction`.

بعد ذلك، اطبع المتغير `botIntroduction` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير يسمى `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

يجب أن يحمل المتغير `botIntroduction` قيمة من نوع سلسلة نصية.

```js
assert.isString(botIntroduction);
```

يجب أن تستخدم ربط السلاسل النصية مع العامل `+` لضم السلسلة النصية `"My name is "` مع المتغير `bot` متبوعًا بنقطة (`.`). انتبه للفراغات. عيّن هذه السلسلة المربوطة إلى المتغير `botIntroduction`.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

يجب أن يكون لديك بيان رابع `console.log()` في شفرتك.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

يجب أن يطبع بيان وحدة التحكم `console` المتغير `botIntroduction`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
