---
id: 66ad8294a0ad902f1b31b612
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة ستتعلم كيفية العمل مع أساسيات JavaScript من خلال بناء بوت ترحيب.

في هذه الخطوة الأولى، ستحتاج إلى إخراج رسالة إلى وحدة التحكم من بوت الترحيب.

تذكر أنك تعلمت عن `console.log()` والسلاسل النصية في الدروس السابقة.

إليك تذكير بكيفية استخدام `console.log()` مع السلاسل النصية:

```js
console.log("Hello, World!");
```

أضف عبارة `console.log()` تُخرج السلسلة النصية `"Hi there!"` إلى وحدة التحكم. لا تنسَ علامات الاقتباس حول الرسالة!


# --hints--

يجب أن تحتوي شفرتك على عبارة `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

يجب أن تحتوي عبارة `console` على الرسالة `"Hi there!"`. لا تنسَ علامات الاقتباس.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
