---
id: 66ad8d150264db3926eccfeb
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

في الدروس السابقة، تعلمت عن الكلمة المفتاحية `let` وكيفية إعلان المتغيرات في JavaScript.

إليك تذكير بكيفية إعلان متغير باستخدام الكلمة المفتاحية `let`:

```js
let greeting;
```

استخدم الكلمة المفتاحية `let` لإعلان متغير اسمه `bot`.

**ملاحظة**: تستخدم `let` هنا لأنك في ورشة العمل لاحقًا ستغير قيمة المتغير `bot`.

# --hints--

يجب أن تحتوي شفرتك على الكلمة المفتاحية `let`.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

يجب إعلان متغير اسمه `bot`.

```js
assert.isNotNull(bot);
```

يجب استخدام الكلمة المفتاحية `let` لإعلان المتغير.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
