---
id: 695774002591bbc5f8cf3e53
title: Criar um Calculador de Descontos
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

Neste laboratório você escreverá uma função que calcula o preço final de um item após aplicar um desconto percentual.

Por exemplo, se o preço de um item é `50` e um desconto de `20` é aplicado, o valor do desconto é `10` e o preço final é `40`.

**Objetivo:** Cumprir as user stories abaixo para completar o laboratório.

**Histórias de Usuário:**

1. Você deve definir uma função chamada `apply_discount`.

2. A função `apply_discount` deve receber exatamente dois parâmetros: `price` e `discount`.

3. Se `price` não for um número (`int` ou `float`), a função deve retornar a string `The price should be a number`.

4. Se `discount` não for um número (`int` ou `float`), a função deve retornar a string `The discount should be a number`.

5. Se `price` for menor ou igual a `0`, a função deve retornar a string `The price should be greater than 0`.

6. Se `discount` for menor que `0` ou maior que `100`, a função deve retornar a string `The discount should be between 0 and 100`.

7. Se ambas as entradas forem válidas, a função deve calcular o desconto como uma porcentagem do preço.

8. A função deve retornar o preço final após aplicar o desconto.

# --hints--

Você deve ter uma função chamada `apply_discount`.

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

Sua função `apply_discount` deve receber dois parâmetros: `price` e `discount`.

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

Quando `apply_discount` é chamado com um `price` (primeiro argumento) que não é um número (`int` ou `float`) ele deve retornar `The price should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

Quando `apply_discount` é chamado com um `discount` (segundo argumento) que não é um número (`int` ou `float`) ele deve retornar `The discount should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

Quando `apply_discount` é chamado com um `price` menor ou igual a `0`, ele deve retornar `The price should be greater than 0`.

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

Quando `apply_discount` é chamado com um `discount` menor que `0` ou maior que `100`, ele deve retornar `The discount should be between 0 and 100`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` deve retornar `80`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` deve retornar `100`.

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` deve retornar `50`.

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

Quando `apply_discount` é chamado com um desconto de `100`, ele deve retornar `0`.

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` deve retornar `59.6`.

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
