---
id: 68e6a99a0fd7386fb66e156f
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você praticará a estilização de checkboxes construindo um painel de configurações do jogo.

Todo o boilerplate HTML (`DOCTYPE`, `html`, `head` e `body`) foi fornecido para você.

Dentro do `body`, crie um elemento `div` com um elemento `h1` aninhado dentro. Dê ao elemento `h1` o texto `Game Settings`

# --hints--

Você deve ter um elemento `div`.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

Você deve ter um elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Seu elemento `h1` deve estar dentro do seu elemento `div`.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

Seu elemento `h1` deve ter o texto `Game Settings`.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
