---
id: 68e9a3194160f8f5620ddf21
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Imediatamente após o seu elemento `h1`, crie quatro elementos `label`.

Dentro de cada um dos elementos `label`, crie um elemento `input` com o atributo `type` definido como `"checkbox"`.

Após cada elemento `input`, dê aos seus elementos `label` os seguintes textos, na ordem: `Sound Effects`, `Background Music`, `Hard Mode` e `Haptic Feedback`.

# --hints--

Você deve ter quatro elementos `label`.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Seu primeiro elemento `label` deve ter um elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

Seu primeiro elemento `input` deve ter o atributo `type` definido como `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

Seu primeiro elemento `label` deve ter o texto `Sound Effects`.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

O seu texto `Sound Effects` deve vir após o elemento `input`.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

Seu segundo elemento `label` deve ter um elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

Seu segundo elemento `input` deve ter o atributo `type` definido como `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

Seu segundo elemento `label` deve ter o texto `Background Music`.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

O seu texto `Background Music` deve vir depois do elemento `input`.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

Seu terceiro elemento `label` deve ter um elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

Seu terceiro elemento `input` deve ter o atributo `type` definido como `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

Seu terceiro elemento `label` deve ter o texto `Hard Mode`.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

Seu texto `Hard Mode` deve vir depois do elemento `input`.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

Seu quarto elemento `label` deve ter um elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

Seu quarto elemento `input` deve ter o atributo `type` definido como `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

Seu quarto elemento `label` deve ter o texto `Haptic Feedback`.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

O seu texto de `Haptic Feedback` deve vir depois do elemento `input`.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
