---
id: 68e9a834dfef752dc35f6f90
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

Agora que suas caixas de seleção mudam de cor quando o usuário clica nelas, você vai aplicar um pouco de estilo para tornar um pouco mais óbvio que a caixa de seleção foi marcada.

Pseudo-classes CSS podem ser encadeadas. Você pode adicionar uma pseudo-classe ou pseudo-elemento após outro para direcionar elementos que atendem a múltiplas condições.

Crie um novo seletor para `input[type="checkbox"]`. Adicione a pseudo-classe `:checked` seguida do pseudo-elemento `::after`. Isso permite que você adicione um indicador visual quando a checkbox estiver marcada.

Em seguida, adicione `content` com o valor `"✓"`.

# --hints--

Você deve ter um seletor de tipo combinado com um seletor de pseudo-classe e um seletor de pseudo-elemento de `input[type="checkbox"]:checked::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

Você deve adicionar uma propriedade `content` ao seu seletor de tipo combinado com um seletor de pseudo-classe e seletor de pseudo-elemento de `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

Você deve ter uma propriedade `content` com o valor `"✓"`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```

