---
id: 68e9a73075dc26207320811a
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Agora que você pode ver sua checkbox novamente, você vai finalizar algumas opções de estilo para ela. Dê bordas arredondadas adicionando um `border-radius` de `4px` ao seu seletor `input[type="checkbox"]`.

Então dê a ele um `background-color` de `white` para fazer o centro da sua checkbox se destacar do fundo do container.

Como você vai configurar uma transição personalizada para quando um usuário clicar nas caixas de seleção, defina um `transition` com o valor `all` e `0.3s` para que a transição aconteça suavemente ao longo de 0,3 segundos em vez de instantaneamente.

# --hints--

Você deve adicionar a propriedade `border-radius` ao seu seletor `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius);
```

Você deve ter uma propriedade `border-radius` com um valor de `4px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius, "4px");
```

Você deve adicionar uma propriedade `background-color` ao seu seletor `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor);
```

Você deve ter uma propriedade `background-color` com um valor de `white`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor, ["white", "rgb(255,255,255)", "#ffffff"]);
```

Você deve adicionar a propriedade `transition` ao seu seletor `input[type="checkbox]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.transition);
```

Você deve ter `transition` com um valor de `all 0.3s`.

```js
/**
 * Because CSS is simply the most lovely thing to ever exist, the behaviour of the `all` shorthand does wonky things and will not show up
 * in the computed stylesheet. Instead, we need to query the actual style element and extract the selector, checking the properties as written
 * by the camper.
 * I really hate CSS.
 * <3 Naomi
 */
const { props } = document.querySelector('.fcc-injected-styles')?.innerText?.match(/input\[type=('|")checkbox\1\]\s*\{(?<props>.*)\}/s)?.groups;
assert.match(props, /transition\s*\:\s*all\s+0.3\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
