---
id: 68e9a6201e2efa142d62bfe1
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

A seguir, você vai aumentar o tamanho das caixas de seleção para melhor visibilidade.

Comece configurando um seletor para `input`, mas especificamente direcionando seu `[type="checkbox"]`.

Dentro do seu seletor, defina o `width` e o `height` para `20px`. Isso o torna maior do que era antes.

E para conformar com sua configuração de `cursor` que foi definida nos rótulos, adicione `cursor` e atribua o valor `pointer` a ele. Depois disso, quando você passar o mouse sobre as caixas de seleção, será exibido um ponteiro.

# --hints--

Você deve ter um seletor no seu CSS de `input[type="checkbox]`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

Você deve adicionar uma propriedade `width` ao seu seletor `input[type="checkbox]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

Você deve ter uma propriedade `width` com um valor de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

Você deve adicionar uma propriedade `height` ao seu seletor `input[type="checkbox]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

Você deve ter uma propriedade `height` com um valor de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

Você deve adicionar uma propriedade `cursor` ao seu seletor `input[type="checkbox]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

Você deve ter uma propriedade `cursor` com um valor de `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
