---
id: 68e9a6bf1a74721a4e43e217
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Agora você vai remover a checkbox padrão que é aplicada pelos navegadores. Dentro do seu seletor `input[type='checkbox']`, adicione `appearance` com o valor `none`.

Definir a propriedade `appearance` como `none` irá remover a aparência que o navegador aplica às caixas de seleção, permitindo que você exiba o estilo desejado.

Depois de fazer isso, como a caixa de seleção não estará mais visível, defina um `border` com `2px` `solid` e um código hexadecimal `#f1be32`.

# --hints--

Você deve adicionar uma propriedade `appearance` ao seu seletor `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

Você deve ter uma propriedade `appearance` com o valor `none`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

Você deve adicionar a propriedade `border` ao seu seletor `input[type="checkbox]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

Você deve ter uma propriedade `border` com um valor de `2px solid #f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
