---
id: 68ee86e3ee07b01626a392cc
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Dentro do seu seletor `.settings-card` defina a propriedade `margin` como `auto`.

Definir a propriedade `margin` como `auto` ajusta automaticamente as margens de um elemento para distribuir uniformemente o espaço restante em seu contêiner, comumente usado para centralizar elementos de nível bloco horizontalmente.

E por último, defina uma propriedade `text-align` com o valor `left`. Isso alinhará o conteúdo inline, como texto, ao lado esquerdo do seu elemento contêiner.

# --hints--

Você deve adicionar a propriedade `margin` ao seletor de classe `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

Você deve ter uma propriedade `margin` com um valor de `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

Você deve adicionar a propriedade `text-align` ao seletor de classe `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

Você deve ter uma propriedade `text-align` com o valor `left`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
