---
id: 68f15cd397503429f0345262
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

Na declaração para o seletor `input[type="checkbox"]:checked::after`, defina o `display` como `block`. Definir sua propriedade `display` como `block` faz com que o elemento seja um elemento de nível bloco, o que significa que ele ocupa toda a largura do seu contêiner e começa em uma nova linha, permitindo que você controle sua largura, altura e espaçamento com mais facilidade.

Em seguida, defina `text-align` como `center`. Isso irá centralizar o conteúdo inline (como texto e elementos inline) horizontalmente dentro do bloco. O bloco ainda ocupará toda a largura do seu contêiner mas tudo dentro dele estará alinhado ao centro.

# --hints--

Você deve adicionar uma propriedade `display` ao seu seletor de tipo combinado com um seletor de pseudo-classe e seletor de pseudo-elemento de `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

Você deve ter uma propriedade `display` com o valor `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

Você deve adicionar uma propriedade `text-align` ao seu seletor de tipo combinado com um seletor de pseudo-classe e seletor de pseudo-elemento `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

Você deve ter uma propriedade `text-align` com um valor de `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
