---
id: 68f15ef0446464325ce1aa45
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Você precisa de algum espaçamento entre as caixas de seleção e os rótulos. Comece usando o seletor para `label` no seu CSS e defina uma propriedade `display` como `block`.

Em seguida, defina uma propriedade `margin` de `8px auto` para adicionar espaçamento vertical entre os elementos. Isso criará um espaçamento consistente acima e abaixo de cada elemento, ajudando a separar o conteúdo e melhorar a legibilidade.

Finalmente, defina a propriedade `cursor` como `pointer` nos elementos `label`. Isso mudará o cursor para um ícone de mão quando o usuário passar o mouse sobre um `label`, sinalizando que o elemento é clicável e melhorando a experiência geral do usuário.

# --hints--

Você deve ter um seletor de tipo no seu CSS para `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

Você deve adicionar uma propriedade `display` ao seu seletor `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

Você deve ter uma propriedade `display` com um valor de `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

Você deve adicionar a propriedade `margin` ao seu seletor `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

Você deve ter uma propriedade `margin` com um valor de `8px auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

Você deve adicionar uma propriedade `cursor` ao seu seletor `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

Você deve ter uma propriedade `cursor` com um valor de `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
