---
id: 68f29753fd5524530c6e65c8
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Em seguida, defina uma propriedade `vertical-align` com o valor `middle`.

A propriedade `vertical-align` controla como os elementos `inline` ou `inline-block` se alinham verticalmente com o texto ao redor ou outros elementos inline. Ela é frequentemente usada para ajustar a posição vertical de elementos como imagens, ícones ou texto dentro de uma linha.

Então finalize sua checkbox com uma propriedade `margin` de `15px`.

# --hints--

Você deve adicionar a propriedade `vertical-align` ao seu seletor `input[type="checkbox]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign);
```

Você deve ter uma propriedade `vertical-align` com um valor de `middle`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign, "middle");
```

Você deve adicionar a propriedade `margin-right` ao seu seletor `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight);
```

Você deve ter uma propriedade `margin-right` com um valor de `15px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight, "15px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
--fcc-editable-region--

--fcc-editable-region--
}
```
