---
id: 68f29dbf9125476b4808d6fb
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Agora é hora de estilizar o container `settings-card`. Crie um seletor de classe para `settings-card`.

Aqui será onde toda a sua formatação para o container será colocada. Defina o `max-width` para `250px` para definir o tamanho geral do seu container.

Em seguida, defina seu `padding` para `20px` para que seu conteúdo tenha espaço entre ele e a borda do contêiner.

Depois disso, crie uma borda arredondada definindo seu `border-radius` para `10px`.

Então defina um `box-shadow` com os valores `0 2px 6px rgba(0,0,0,0.2)`. Isso criará um efeito sutil de "elevação" que dará profundidade ao container.

# --hints--

Você deve ter um seletor de classe no seu css para `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

Você deve adicionar a propriedade `max-width` ao seletor de classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

Você deve ter uma propriedade `max-width` com um valor de `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

Você deve adicionar a propriedade `padding` ao seletor de classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

Você deve ter uma propriedade `padding` com um valor de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

Você deve adicionar a propriedade `border-radius` ao seletor de classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

Você deve ter uma propriedade `border-radius` com um valor de `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

Você deve adicionar a propriedade `box-shadow` ao seletor de classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

Você deve ter um `box-shadow` com o valor `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
