---
id: 67c562286b29447da020d407
title: Tengeneza kizalishaji nambari ya siri ya mara moja
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

Katika maabara hii, utatengeneza OTP (Nenosiri la Mara Moja) lenye tarakimu 6 na kuonyesha kwa mtumizi. OTP itatoweka baada ya sekunde 5, na OTP mpya itatengenezwa wakati mtumizi atabofya kitufe cha `"Generate OTP"`.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kutumia hook ya `useEffect` kusimamia kielekezi cha kuhesabu chini.
2. Sehemu yako ya `OTPGenerator` inapaswa kurudisha kipengele cha `div` chenye jina la darasa la vitu `container`.
3. `div` yako yenye darasa la vitu `container` inapaswa kujumuisha vipengele vifuatavyo:

    - Kipengele cha `h1` chenye kitambulisho `otp-title` na maandishi `"OTP Generator"`.
    - Kipengele cha `h2` chenye kitambulisho `otp-display` ambacho kinaonyesha ujumbe `"Click 'Generate OTP' to get a code"` au kinaonyesha OTP iliyotengenezwa ikiwa ipo.
    - Kipengele cha `p` chenye kitambulisho `otp-timer` na sifa ya kipengele `aria-live` iliyowekwa kwa thamani halali ambayo:
        - Huanza kuwa tupu.
        - Kinaonyesha `"Expires in: X seconds"` baada ya kitufe kubofyanwa, ambapo `X` ni muda uliobaki kabla ya OTP kutoweka.
        - Kinaonyesha ujumbe `"OTP expired. Click the button to generate a new OTP."` mara kielekezi cha kuhesabu chini kinapofikia `0`.
    - Kipengele cha `button` chenye kitambulisho `generate-otp-button` chenye lebo `"Generate OTP"`. Kitakapobofyanwa, kinapaswa kutengeneza OTP mpya na kuanzisha kielekezi cha kuhesabu chini cha sekunde 5.
    - Kitufe cha `"Generate OTP"` kinapaswa kuzimwa wakati kielekezi cha kuhesabu chini kinaendelea.

4. Unapaswa kuhakikisha kielekezi cha kuhesabu chini kinasimama kiotomatiki mara kinapofikia sekunde `0` ili kuzuia masasisho yasiyo ya lazima.
5. OTP yako iliyotengenezwa inapaswa kuwa na tarakimu 6.

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --after-all--

```js
clock.uninstall();
```

# --hints--

Haupaswi kuondoa mgawo wa thamani wa `const { useState, useEffect, useRef } = React;` uliopo kwenye msimbo.

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

Unapaswa onyesha kipengele cha `div` chenye darasa la `.container`.

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

`.container` yako inapaswa kuwa na kipengele cha `h1` chenye kitambulisho `otp-title` na maandishi `"OTP Generator"`.

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

`.container` yako inapaswa kuwa na kipengele cha `h2` chenye kitambulisho `otp-display` kinachoonyesha OTP wakati imetengenezwa.

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

`h2` yako ndani ya `.container` inapaswa kuonyesha ujumbe `"Click 'Generate OTP' to get a code"` mwanzoni.

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

Kipengele chako cha `.container` kinapaswa kuwa na kipengele cha `p` chenye kitambulisho `otp-timer`.

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = container.querySelectorAll('p');
assert.strictEqual(pElements.length, 1);
assert.strictEqual(pElements[0].id, "otp-timer", "The p element should have the id otp-timer");
```

Kipengele chako cha `p` kinapaswa kuwa na sifa ya kipengele `aria-live` iliyowekwa kuwa `assertive` au `polite`.

```js
const pElement = document.querySelector('.container p');
assert.exists(pElement?.ariaLive);
assert.oneOf(pElement?.ariaLive, ["assertive", "polite"]);
```

Kipengele chako cha `p` kinapaswa kuwa tupu mwanzoni.

```js
const pElement = document.querySelector('.container p');
assert.strictEqual(pElement?.innerText, "", "The `p` element should be empty");
```

`.container` yako inapaswa kuwa na kipengele cha `button` chenye kitambulisho `generate-otp-button` na maandishi `"Generate OTP"`.

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

Kipengele chako cha `h2` chenye kitambulisho `otp-display` kinapaswa kuonyesha OTP mpya wakati kitufe kinapobofyanwa.

```js
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

OTP yako iliyotengenezwa inapaswa kuwa na tarakimu 6.

```js
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```


Kipengele chako cha `p` chenye kitambulisho `otp-timer` kinapaswa kuonyesha kielekezi cha kuhesabu chini cha sekunde 5 wakati kitufe kinapobofyanwa.

```js
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

Kipengele chako cha `p` chenye kitambulisho `otp-timer` kinapaswa kusasisha ujumbe wake kila sekunde kuonyesha muda uliobaki.

```js
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
```

Kitufe chako cha `"Generate OTP"` kinapaswa kuzimwa wakati kielekezi cha kuhesabu chini kinaendelea.

```js
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);
  
  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
```

Kitufe chako cha `"Generate OTP"` kinapaswa kuwashwa mara kielekezi cha kuhesabu chini kinapofikia 0.

```js

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);
  
assert.isFalse(button.disabled);
```

Unapaswa kuonyesha ujumbe `OTP expired. Click the button to generate a new OTP.` wakati kielekezi cha kuhesabu chini kinapofikia `0`.

```js
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
```

Unapaswa kusafirisha sehemu ya `OTPGenerator`.

```js
assert.isFunction(window.index.OTPGenerator);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      <p id='otp-timer' aria-live='assertive'>
        {isCounting
          ? `Expires in: ${secondsLeft} seconds`
          : hasGeneratedOtp && 
            "OTP expired. Click the button to generate a new OTP."}
      </p>
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
