---
id: 68e2c945cc1d8e778152be31
title: "Costruisci un'interfaccia per il player"
challengeType: 27
dashedName: build-a-player-interface
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi definire una classe astratta chiamata `Player` che eredita dalla classe `abc.ABC`.
1. La classe `Player` deve avere un metodo `__init__` che imposta:

   - L'attributo `moves` a una lista vuota.
   - L'attributo `position` a `(0, 0)`.
   - L'attributo `path` a una lista contenente la posizione iniziale.

1. La classe `Player` deve avere un metodo chiamato `make_move` che:

   - Usa `random.choice` per ottenere una mossa casuale dall'attributo `moves` (definito nella classe concreta).
   - Aggiunge i valori della mossa selezionata alla posizione corrente e aggiorna l'attributo `position`.
   - Aggiunge la nuova tupla `position` all'attributo `path`.
   - Restituisce la nuova `position`.

1. La classe `Player` deve avere un metodo astratto chiamato `level_up` da implementare nelle classi concrete.
1. Devi definire una classe `Pawn` che eredita dalla classe `Player`.
1. La classe `Pawn` deve usare `super()` per chiamare il metodo `__init__` del genitore e poi impostare l'attributo `moves` a una lista di tuple che rappresentano coordinate `x, y`.
1. Ogni tupla di coordinate deve rappresentare un movimento di `1` unitĂ  nelle seguenti direzioni: su, giĂą, sinistra, destra.
1. La classe `Pawn` deve implementare un metodo concreto `level_up` aggiungendo altre mosse all'attributo `moves`. Le mosse aggiunte devono rappresentare i quattro movimenti diagonali (per esempio, `1` unitĂ  giĂą piĂą `1` unitĂ  a sinistra).

**Nota:** I moduli della libreria standard devono essere importati senza usare alias. I test relativi alla classe `Player` falliranno finchĂ© la classe `Pawn` non sarĂ  istanziabile.

# --hints--

Devi avere una classe chiamata `Player`.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Player")`) })
```

La classe `Player` deve ereditare dalla classe `ABC` del modulo `abc`.

```js
({ test: () => runPython(`
from abc import ABC
import ast
cls = _Node(_code).find_class("Player")
assert cls.inherits_from("ABC") or cls.inherits_from("abc.ABC")
`) })
```

La classe `Player` deve avere un metodo `__init__`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("__init__")`) })
```

Il metodo `Player` di `__init__` deve avere un singolo parametro `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("__init__").has_args("self")`) })
```

Il metodo `Player` di `__init__` deve impostare `self.moves` a una lista vuota.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.moves = []") or init.has_stmt("self.moves = list()")
`) })
```

Il metodo `Player` di `__init__` deve impostare `self.position` a `(0, 0)`.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.position = (0, 0)")
`) })
```

Il metodo `Player` di `__init__` deve impostare `self.path` a una lista contenente la posizione iniziale.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.path = [self.position]")
`) })
```

La classe `Player` deve avere un metodo `make_move`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("make_move")`) })
```

Il metodo `Player` di `make_move` deve avere un singolo parametro `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("make_move").has_args("self")`) })
```

Il metodo `Player` di `make_move` deve usare `random.choice` per ottenere una mossa casuale dall'attributo `moves`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").block_has_call("choice", "make_move")`) })
```

Il metodo `Player` di `make_move` deve aggiornare l'attributo `position` aggiungendo le coordinate della mossa selezionata casualmente.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.position == (1, 2)
  p.make_move()
  assert p.position == (2, 4)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Il metodo `Player` di `make_move` deve aggiungere la nuova tupla `position` all'attributo `path`.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.path == [(0, 0), (1, 2)]
  p.make_move()
  assert p.path == [(0, 0), (1, 2), (2, 4)]

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Il metodo `Player` di `make_move` deve restituire l'attributo `position` aggiornato.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  assert p.make_move() == (1, 2)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

La classe `Player` deve avere un metodo `level_up`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("level_up")`) })
```

Il metodo `Player` di `level_up` deve avere un singolo parametro `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("level_up").has_args("self")`) })
```

Il metodo `Player` di `level_up` deve essere un metodo astratto.

```js
({ test: () => runPython(`
target = _Node(_code).find_class("Player").find_function("level_up")
assert target.has_decorators("abstractmethod") or target.has_decorators("abc.abstractmethod")
`) })
```

La classe `Player` deve essere una classe astratta.

```js
({ test: () => runPython(`
try:
  Player()
except TypeError as e:
  assert str(e) == "Can't instantiate abstract class Player with abstract method level_up"
else:
  assert False, "Player class should not be instantiable"
`) })
```

Devi avere una classe chiamata `Pawn`.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Pawn")`) })
```

La classe `Pawn` deve ereditare dalla classe `Player`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").inherits_from("Player")`) })
```

La classe `Pawn` deve avere un metodo `__init__`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("__init__")`) })
```

Il metodo `Pawn` di `__init__` deve avere un singolo parametro `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_args("self")`) })
```

Il metodo `Pawn` di `__init__` deve chiamare il metodo `__init__` del genitore usando la funzione `super`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_stmt("super().__init__()")`) })
```

Il metodo `Pawn` di `__init__` deve impostare l'attributo `moves` a una lista di tuple che rappresentano coordinate `x, y`, dove ogni tupla di coordinate rappresenta un movimento di `1` unitĂ  nelle seguenti direzioni: su, giĂą, sinistra, destra.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]
assert isinstance(p.moves, list)
assert len(p.moves) == 4
assert all(move in p.moves for move in moves)
`) })
```

La classe `Pawn` deve avere un metodo `level_up`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("level_up")`) })
```

Il metodo `Pawn` di `level_up` deve aggiungere i quattro movimenti diagonali di `1` unitĂ  all'attributo `moves`.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1), (1, 1), (1, -1), (-1, 1), (-1, -1)]
p.level_up()
assert isinstance(p.moves, list)
assert len(p.moves) == 8
assert all(move in p.moves for move in moves)
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
from abc import ABC, abstractmethod
from random import choice


class Player(ABC):
    def __init__(self):
        self.moves = []
        self.position = (0, 0)
        self.path = [self.position]

    def make_move(self):
        x, y = self.position
        move = choice(self.moves)
        new_x = move[0] + x
        new_y = move[1] + y
        self.position = (new_x, new_y)
        self.path.append(self.position)
        return self.position

    @abstractmethod
    def level_up(self):
        pass


class Pawn(Player):
    def __init__(self):
        super().__init__()
        self.moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]

    def level_up(self):
        self.moves.extend([(1, 1), (1, -1), (-1, 1), (-1, -1)])
```
