---
id: 66e2edb23013ca596f8d9568
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Ahora es el momento de probar su función `getPunctuationCount`.

Cree una variable `punctuationCount` y asígnela al resultado de llamar a la función `getPunctuationCount` con el argumento de `"WHAT?!?!?!?!?"`.

Después de eso, registre lo siguiente en la consola: `"Punctuation Count: [Punctuation count goes here]"`. Reemplace `[Punctuation count goes here]` con el nombre real de la variable. Puede elegir usar cadenas de plantillas o concatenación de cadenas con el operador `+` aquí.

# --hints--

Debe crear una variable `punctuationCount`.

```js
assert.isNotNull(punctuationCount)
```

Su variable `punctuationCount` debe establecerse en el resultado de `getPunctuationCount("WHAT?!?!?!?!?")`.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

Debe registrar lo siguiente en la consola: `"Punctuation Count: [Punctuation count goes here]"`. Reemplace `[Punctuation count goes here]` con el nombre real de la variable. Asegúrese de usar la sintaxis adecuada para la concatenación de cadenas aquí.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

