---
id: 66e2d680e129e1423116a541
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, construirás un analizador de oraciones que tomará una oración y obtendrá el conteo del número de palabras, vocales, consonantes y signos de puntuación.

Para comenzar, crea una función `getVowelCount` con un parámetro llamado `sentence`. Tu función debe devolver el número total de vocales en una oración.

# --hints--

Deberías crear una función `getVowelCount`.

```js
assert.isFunction(getVowelCount);
```

Tu función `getVowelCount` debe tener un parámetro `sentence`.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

Tu función `getVowelCount` debe devolver un número.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

Cuando la oración es `"Apples are tasty fruits"`, la función `getVowelCount` debería devolver `7`.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

Cuando la oración es `"Hello, World!"`, la función debería devolver `3`.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

Tu conteo de vocales debe ser insensible a mayúsculas y minúsculas.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

Tu función `getVowelCount` debe devolver el conteo correcto de vocales para cualquier oración.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
