---
id: 68e51d8f163d3f8b70f9c4e5
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將練習使用 `includes()` 和 `slice()` 方法。

首先，創建一個名為 `fccSentence` 的變數，並將字串 `"freeCodeCamp is a great place to learn web development."` 指定給它。

# --hints--

你應該宣告一個名為 `fccSentence` 的變數。

```js
assert.exists(fccSentence);
```

你的 `fccSentence` 變數應該是一個字串。

```js
assert.isString(fccSentence);
```

你的 `fccSentence` 變數應該被指定為字串 `"freeCodeCamp is a great place to learn web development."`。

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
