---
id: 68e6172975008fea742a63ca
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

現在使用模板字面值或字串串接，將訊息 `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` 記錄到主控台。將 `<hasFreeCodeCamp>` 取代為變數的實際值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你應該記錄訊息 `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`，其中 `<hasFreeCodeCamp>` 應該被取代為變數的實際值。

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
