---
id: 68e6172975008fea742a63cb
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

現在是時候看看當子字串未被找到時會發生什麼。

創建一個名為 `hasJavaScript` 的變數。然後將使用 `includes()` 方法檢查 `fccSentence` 是否包含 `"JavaScript"` 的結果指定給它。

然後使用模板字面值或字串串接將訊息 `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` 記錄到主控台。將 `<hasJavaScript>` 取代為變數的實際值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你應該宣告一個名為 `hasJavaScript` 的變數。

```js
assert.exists(hasJavaScript);
```

你應該使用 `fccSentence.includes("JavaScript")`。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

你應該將 `fccSentence.includes("JavaScript")` 的結果指定到你的 `hasJavaScript` 變數。

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

`fccSentence.includes("JavaScript")` 的結果應該是 `false`。 

```js
assert.isFalse(hasJavaScript);
```

你應該記錄訊息 `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`，其中 `<hasJavaScript>` 應該被取代為變數的實際值。

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
