---
id: 68e6172975008fea742a63d0
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

你可以使用 `slice()` 從字串的開頭擷取。

創建一個名為 `greetingWord` 的變數。然後將使用 `message.slice(0, 7)` 擷取第一個字 `"Welcome"` 的結果指定給它。

然後使用模板字面值或字串串接將訊息 `The first word is "<greetingWord>".` 記錄到主控台。將 `<greetingWord>` 取代為變數的實際值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你應該宣告一個名為 `greetingWord` 的變數。

```js
assert.exists(greetingWord);
```

你的 `greetingWord` 變數應該是一個字串。

```js
assert.isString(greetingWord);
```

你應該使用 `message.slice(0, 7)`。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

你應該將 `message.slice(0, 7)` 的結果指定到你的 `greetingWord` 變數。

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

你的 `greetingWord` 變數應該相等於 `"Welcome"`。

```js
assert.strictEqual(greetingWord, "Welcome");
```

你應該記錄訊息 `The first word is "<greetingWord>".`，其中 `<greetingWord>` 應該被實際的變數值取代。

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
