---
id: 68e6172975008fea742a63ce
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

`slice()` 方法會擷取字串的一部分。它接受兩個引數：起始索引（包含）和結束索引（不包含）。

這裡有一個範例：

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

字串索引從 0 開始。在 `message` 字串中，索引 11 的字元是 `f`（"freeCodeCamp" 的開始）。

創建一個名為 `platform` 的變數。然後將使用 `message.slice(11, 23)` 擷取 `"freeCodeCamp"` 的結果指定給它。

# --hints--

你應該宣告一個名為 `platform` 的變數。

```js
assert.exists(platform);
```

你的 `platform` 變數應該是一個字串。

```js
assert.isString(platform);
```

你應該使用 `message.slice(11, 23)`。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

你應該將 `message.slice(11, 23)` 的結果指定到你的 `platform` 變數。

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

你的 `platform` 變數應該相等於 `"freeCodeCamp"`。

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
