---
id: 68e6172975008fea742a63cc
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

`includes()` 方法區分大小寫，表示它會區分大寫和小寫字母。

創建一個名為 `hasLowercaseFCC` 的變數。然後將使用 `includes()` 方法檢查 `fccSentence` 是否包含 `"freecodecamp"`（全部小寫）的結果指定給它。

然後使用模板字面值或字串串接將訊息 `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` 記錄到主控台。將 `<hasLowercaseFCC>` 取代為變數的實際值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你應該宣告一個名為 `hasLowercaseFCC` 的變數。

```js
assert.exists(hasLowercaseFCC);
```

你應該使用 `fccSentence.includes("freecodecamp")`。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

你應該將 `fccSentence.includes("freecodecamp")` 的結果指定到你的 `hasLowercaseFCC` 變數。

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

`fccSentence.includes("freecodecamp")` 的結果應該是 `false`。 

```js
assert.isFalse(hasLowercaseFCC);
```

你應該記錄包含 `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` 的訊息，其中 `<hasLowercaseFCC>` 應該被取代為該變數的實際值。

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
