---
id: 68e6172875008fea742a63c9
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

如你從先前的課程中回想，`includes()` 方法會檢查字串是否包含特定子字串，並傳回 `true` 或 `false`。

以下是使用 `includes()` 方法的範例：

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

創建一個名為 `hasFreeCodeCamp` 的變數。然後將使用 `includes()` 方法在 `fccSentence` 上檢查是否包含 `"freeCodeCamp"` 的結果指定給它。

# --hints--

你應該宣告一個名為 `hasFreeCodeCamp` 的變數。

```js
assert.exists(hasFreeCodeCamp);
```

你應該使用 `fccSentence.includes("freeCodeCamp")`。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

你應該將 `fccSentence.includes("freeCodeCamp")` 的結果指定給你的 `hasFreeCodeCamp` 變數。

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

`fccSentence.includes("freeCodeCamp")` 的結果應該是 `true`。 

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
