---
id: 68e6172975008fea742a63cf
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

現在是時候記錄使用 `slice()` 方法的結果。

使用模板字面值或字串串接將訊息 `The word "<platform>" was sliced from the message.` 記錄到主控台。將 `<platform>` 取代為變數的實際值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你應該記錄訊息 `The word "<platform>" was sliced from the message.`，其中 `<platform>` 應該被取代為變數的實際值。

```js
assert.deepInclude(spy.calls, ['The word "freeCodeCamp" was sliced from the message.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
--fcc-editable-region--

--fcc-editable-region--
```
