---
id: 68e6172975008fea742a63d1
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

`slice()` 方法可以使用負索引從字串的末尾開始計數。`-1` 取用最後一個字元。

這裡有一個範例：

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

創建一個名為 `endPunctuation` 的變數。然後將使用 `message.slice(-1)` 擷取最後一個字元的結果指定給它。

然後使用模板字面值或字串串接，將訊息 `The ending punctuation mark is a "<endPunctuation>"` 記錄到主控台。將 `<endPunctuation>` 取代為變數的實際值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你應該宣告一個名為 `endPunctuation` 的變數。

```js
assert.exists(endPunctuation);
```

你的 `endPunctuation` 變數應該是一個字串。

```js
assert.isString(endPunctuation);
```

你應該使用 `message.slice(-1)`。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

你應該將 `message.slice(-1)` 的結果指定到你的 `endPunctuation` 變數。

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

你的 `endPunctuation` 變數應該相等於 `"!"`。

```js
assert.strictEqual(endPunctuation, "!");
```

你應該記錄訊息 `The ending punctuation mark is a "<endPunctuation>"`，其中 `<endPunctuation>` 應該被取代為該變數的實際值。

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
