---
id: 68e6172975008fea742a63cd
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

現在你將探索 `slice()` 方法，它會擷取字串的一部分。

創建一個名為 `message` 的變數，並將字串 `"Welcome to freeCodeCamp!"` 指定給它。

然後使用 `console.log()` 將 `"Here are some examples of the slice() method:"` 記錄到主控台。

# --hints--

你應該宣告一個名為 `message` 的變數。

```js
assert.exists(message);
```

你的 `message` 變數應該是一個字串。

```js
assert.isString(message);
```

你的 `message` 變數應該是 `"Welcome to freeCodeCamp!"`。

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

你應該將 `"Here are some examples of the slice() method:"` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
