---
id: 66b102ee0bdbad34a9f85ba0
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستواصل تعلم السلاسل النصية من خلال بناء روبوت دردشة للمعلم.

لبدء العمل، أضف عبارة `console` مع رسالة `"Hi there!"`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب أن تحتوي على عبارة `console`.

```js
assert.isAbove(logSpy.calls.length, 0);
```

يجب أن تُخرج عبارة وحدة التحكم الرسالة `"Hi there!"`.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
