---
id: 66b5928b86e907fdfbf34e56
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

حان الوقت الآن لتعيين اسم الروبوت.

أنشئ متغيرًا اسمه `botName` وقم بتعيين القيمة النصية `"teacherBot"` له.

# --hints--

يجب أن يكون لديك متغير اسمه `botName`.

```js
assert.isNotNull(botName);
```

يجب أن يكون المتغير `botName` سلسلة نصية.

```js
assert.isString(botName);
```

يجب أن يحتوي المتغير `botName` على القيمة `"teacherBot"`.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
