---
id: 66b59b12a745e10011158f7b
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

الخطوة التالية هي إنشاء بعض المتغيرات الإضافية التي ستُستخدم في رسائل الروبوت المستقبلية.

أنشئ متغيرًا اسمه `subject` وقم بتعيين القيمة النصية `"JavaScript"` له.

ثم أنشئ متغيرًا اسمه `topic` وقم بتعيين القيمة النصية `"strings"` له.

# --hints--

يجب أن يكون لديك متغير اسمه `subject`.

```js
assert.isNotNull(subject);
```

يجب تعيين سلسلة نصية إلى متغير `subject` الخاص بك.

```js
assert.isString(subject);
```

يجب تعيين السلسلة النصية `"JavaScript"` إلى متغير `subject` الخاص بك.

```js
assert.equal(subject, "JavaScript");
```

يجب أن يكون لديك متغير اسمه `topic`.

```js
assert.isNotNull(topic);
```

يجب تعيين سلسلة نصية إلى متغير `topic` الخاص بك.

```js
assert.isString(topic);
```

يجب تعيين السلسلة النصية `"strings"` إلى متغير `topic` الخاص بك.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
