---
id: 66b59829dba144ff1351220f
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

حان الوقت الآن لإنشاء تحية باستخدام المتغير `botName`.

في الدروس السابقة، تعلمت كيفية ربط السلاسل النصية باستخدام قوالب حرفية مثل هذا:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

ابدأ بإنشاء متغير يسمى `greeting`.

بعد ذلك، باستخدام صيغة القوالب الحرفية، عيّن سلسلة نصية تقول `My name is`، تليها المتغير `botName`، وتنتهي بنقطة (`.`).

أخيرًا، سجّل المتغير `greeting` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير يسمى `greeting`.

```js
assert.isNotNull(greeting);
```

يجب أن تعيّن سلسلة نصية إلى متغير `greeting` الخاص بك.

```js
assert.isString(greeting);
```

يجب أن تستخدم القوالب الحرفية لربط السلسلة النصية `My name is` مع المتغير `botName` متبوعًا بنقطة (`.`).

```js
assert.equal(greeting, "My name is teacherBot.");
```

يجب أن يكون لديك جملة `console` تسجل المتغير `greeting` في وحدة التحكم.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
