---
id: 66b59be6ab830800c4df9146
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

حان الوقت الآن لاستخدام المتغيرات التي أنشأتها في الخطوة السابقة.

في الدروس السابقة، تعلمت كيفية العمل مع قوالب حرفية مثل هذه:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

ابدأ بإنشاء متغير `sentence`.

باستخدام صيغة القوالب الحرفية، قم بتعيين السلسلة النصية `Today, you will learn about [topic variable goes here] in [subject variable goes here].` إلى المتغير `sentence`.

ستستبدل عناصر النائب `[topic variable goes here]` و`[subject variable goes here]` بالمتغيرين `topic` و`subject` وبصيغة `${}`.

أخيرًا، سجّل المتغير `sentence` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير `sentence`.

```js
assert.isNotNull(sentence);
```

يجب أن يكون المتغير `sentence` سلسلة نصية.

```js
assert.isString(sentence);
```

يجب أن تستخدم صيغة القوالب الحرفية لتعيين السلسلة النصية `Today, you will learn about [topic variable goes here] in [subject variable goes here].` إلى المتغير `sentence`. ارجع إلى المثال إذا احتجت مساعدة.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

يجب أن تسجل المتغير `sentence` في وحدة التحكم.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
