---
id: 66b6ffb42b88e33943788abf
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

في الجزء الأخير من ورشة العمل، ستراجع كيفية إيجاد موقع الفهرس للجزء الفرعي في سلسلة نصية. تذكر أن الجزء الفرعي هو جزء من سلسلة نصية.

ابدأ بإنشاء متغير يسمى `learningIsFunSentence` وقم بتعيين القيمة النصية `"Learning is fun."` له.

# --hints--

يجب أن يكون لديك متغير يسمى `learningIsFunSentence`.

```js
assert.isNotNull(learningIsFunSentence);
```

يجب أن يحمل المتغير `learningIsFunSentence` قيمة من نوع سلسلة نصية.

```js
assert.isString(learningIsFunSentence);
```

يجب تعيين القيمة `"Learning is fun."` للمتغير `learningIsFunSentence`.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
