---
id: 66b6e62423e8a031d6c1f03d
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

حان الوقت الآن للحصول على طول سلسلة النص `topic`.

يمكنك استخدام قوالب حرفية داخل عبارات `console` بهذا الشكل:

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

ابدأ بإخراج الرسالة `Here is an example of using the length property on the word [topic].` إلى وحدة التحكم.

تذكر استبدال `[topic]` بالمتغير `topic`، واستخدم صيغة القوالب الحرفية الصحيحة كما فعلت في الخطوات السابقة.

ثم أضف عبارة `console.log` ثانية تُخرج طول سلسلة النص `topic` إلى وحدة التحكم.

# --hints--

يجب استخدام صيغة القوالب النصية لإخراج الرسالة `Here is an example of using the length property on the word [topic].` إلى `console`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

يجب أن يكون لديك بيان `console` ثاني لإخراج طول المتغير `topic` إلى الـ `console`.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
