---
id: 66b6efddeca35833cd6f0b03
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

في الدروس السابقة، تعلمت كيف تصل إلى الأحرف في سلسلة نصية بهذه الطريقة:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

تذكر أن أرقام الفهارس تبدأ من `0`، لذلك سيكون الحرف الأول في السلسلة النصية دائمًا عند الفهرس `0`.

ابدأ بإضافة بيان `console` آخر.

داخل بيان `console`، أخرج الحرف الأول من المتغير `subject` باستخدام ترميز الأقواس والرقم الصحيح للفهرس.

# --hints--

يجب أن يخرج بيان `console` الحرف الأول من المتغير `subject` باستخدام ترميز الأقواس والفهرس `0`.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
