---
id: 66b6d482bbb9e12f2e5ee1ae
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

في هذا الجزء التالي من المشروع، سيعلّم الروبوت كيفية العمل مع خاصية `length` للسلسلة النصية.

ابدأ بإنشاء متغير جديد اسمه `strLengthIntro`.

ثم باستخدام صيغة القالب النصي، عيّن السلسلة النصية `Here is an example of using the length property on the word [subject].` إلى المتغير `strLengthIntro`.

استبدل `[subject]` بالمتغير `subject` كما فعلت سابقًا.

أخيرًا، سجّل المتغير `strLengthIntro` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير اسمه `strLengthIntro`.

```js
assert.isNotNull(strLengthIntro);
```

يجب أن يكون المتغير `strLengthIntro` سلسلة نصية.

```js
assert.isString(strLengthIntro);
```

يجب أن تستخدم صيغة القالب النصي لتعيين السلسلة النصية `Here is an example of using the length property on the word [subject goes here].` إلى المتغير `strLengthIntro`.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

يجب أن تسجل المتغير `strLengthIntro` في وحدة التحكم.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
