---
id: 66b6fdb76441c738719039fa
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

في الدروس السابقة، تعلمت كيف تصل إلى آخر حرف في سلسلة نصية بهذه الطريقة:

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` سيعطيك دائمًا رقم الفهرس الأخير لسلسلة نصية.

أنشئ متغيرًا جديدًا اسمه `lastCharacter` وعيّن له قيمة آخر حرف في المتغير `subject`.

ثم، سجّل قيمة المتغير `lastCharacter` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير اسمه `lastCharacter`.

```js
assert.isNotNull(lastCharacter);
```

يجب أن تعيّن قيمة آخر حرف في المتغير `subject` إلى المتغير `lastCharacter`. ارجع إلى المثال إذا احتجت مساعدة.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

يجب أن تسجل قيمة المتغير `lastCharacter` في وحدة التحكم `console`.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
