---
id: 66b6e39031393e30f2c48d0f
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

للحصول على طول سلسلة نصية، يمكنك استخدام الخاصية <dfn>length</dfn>. تعيد هذه الخاصية عدد الأحرف في السلسلة النصية.

إليك مثالاً:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

ابدأ باستخدام الخاصية `length` للحصول على طول السلسلة النصية `subject` وسجل هذه القيمة في وحدة التحكم.

# --hints--


يجب أن تستخدم الخاصية `length` للحصول على طول السلسلة النصية `subject`.

```js
assert.match(code, /subject\.length/);
```

يجب أن تسجل طول السلسلة النصية `subject` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
