---
id: 66b6f586767a1534f3097353
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

حان الوقت الآن للوصول إلى الحرف الثاني من المتغير `subject`.

ابدأ بإضافة عبارة `console` تُخرج الرسالة `Here is an example of accessing the second letter in the word [subject].` تذكر استبدال `[subject]` بالقيمة الفعلية للمتغير `subject` واستخدم صيغة القالب النصي الصحيحة.

ثم أضف عبارة `console` أخرى تُخرج الحرف الثاني من المتغير `subject` باستخدام ترميز الأقواس والرقم الصحيح للفهرس.

# --hints--

يجب أن تُخرج عبارة `console` الخاصة بك الرسالة `Here is an example of accessing the second letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

يجب أن يكون لديك عبارة `console` أخرى تُخرج الحرف الثاني من المتغير `subject` باستخدام ترميز الأقواس مع الفهرس `1`.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
