---
id: 68e448ec4ac4205968c0a805
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Sería bueno si la palabra `"love"` se repitiera tres veces en la oración `"I love learning!"`. 

Como recuerdas de una lección anterior, puedes repetir una cadena un número específico de veces usando el método `repeat()`. Aquí tienes un ejemplo:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

Dado que las cadenas son inmutables, este método no modificará la cadena original. Devolverá una nueva cadena con el contenido repetido. 

Crea una variable llamada `repeatedLove` y asígnale `"love ".repeat(3)`. Luego muestra en la consola la variable `repeatedLove` para que puedas ver el resultado. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes tener una variable llamada `repeatedLove`.

```js
assert.exists(repeatedLove);
```

Tu variable `repeatedLove` debe ser una cadena.

```js
assert.isString(repeatedLove);
```

Debes asignar `"love ".repeat(3)` a la variable `repeatedLove`.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

Debes registrar la variable `repeatedLove` en la consola.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
