---
id: 68e44abea678a65b877864dc
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Sería bueno ver esa palabra repetida en una oración. Crea una nueva variable llamada `newSentence` y asígnale el resultado de concatenar la palabra `"I "` con la variable `repeatedLove` seguida de concatenar la cadena `" learning."`. Luego, muestra la variable `newSentence` en la consola. Puedes elegir usar el operador `+` o literales de plantilla para la concatenación de cadenas. 

Ahora deberías ver el texto `"I love love love  learning."`. 

*NOTA*: Hay un espacio extra después del último `"love"` que es correcto. En el último paso, eliminarás ese espacio extra. 

# --hints--

Deberías tener una variable llamada `newSentence`.

```js
assert.exists(newSentence);
```

Tu variable `newSentence` debe ser una cadena.

```js
assert.isString(newSentence);
```

Debes concatenar la palabra `"I "` con la variable `repeatedLove` seguida de concatenar la cadena `" learning."`. Este resultado debe asignarse a la variable `newSentence`. Ten cuidado con los espacios. 

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

NO debes asignar la cadena literal `"I love love love  learning."` a la variable `newSentence`. Por favor, vuelve a leer las instrucciones para el enfoque correcto.

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
