---
id: 68e42b62d7093345c6975f22
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Debajo de tu variable `originalString`, agrega un `console.log()` con la cadena `"Original string:"`. 

Luego, debajo de ese `console.log()`, agrega otro `console.log()` con la variable `originalString`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes registrar la cadena `"Original string:"` en la consola.

```js
assert.equal(getLogs()[0], "Original string:");
```

Debes registrar la variable `originalString` en la consola. 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
