---
id: 68e43e884ef09d4d2ab05631
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora es momento de ver el resultado de usar el método `replace()`.

Comienza agregando un `console.log()` con el valor `"After using the replace() method:"`.

Luego, debajo de ese `console.log()`, agrega otro `console.log()` con la variable `replacedString`.

Mira la consola y deberías ver que la nueva oración dice `"I love dogs."` en lugar de la original `"I love cats."`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes registrar `"After using the replace() method:"` en la consola.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

Debes registrar la variable `replacedString` en la consola. 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--
const replacedString = originalString.replace("cats", "dogs");

--fcc-editable-region--
```
