---
id: 68e43f94365a964f26c88f78
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora es momento de trabajar con un nuevo método de cadenas. 

Comienza creando una nueva variable llamada `exampleSentence` y asígnale la cadena `"I love cats and cats are so much fun!"`.

Luego, debajo de esa variable, registra en la consola `"Original sentence:"`. Luego, debajo de ese `console.log()`, agrega otro `console.log()` con la variable `exampleSentence`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Deberías tener una variable `exampleSentence`.

```js
assert.exists(exampleSentence);
```

Tu variable `exampleSentence` debería ser una cadena.

```js
assert.isString(exampleSentence);
```

Debes asignar la cadena `"I love cats and cats are so much fun!"` a la variable `exampleSentence`.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

Debes registrar `"Original sentence:"` en la consola.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

Debes registrar la variable `exampleSentence` en la consola.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
