---
id: 68e4417b7317835208145323
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

A continuación, crea una nueva variable llamada `dogsOnlySentence` y asígnale `exampleSentence.replace("cats", "dogs")`.

Luego, debajo de esa variable, registra en la consola `"Replacing all occurrences of cats with dogs:"`. Debajo de ese `console.log()`, agrega otro `console.log()` para la variable `dogsOnlySentence`. 

Ahora deberías ver el texto `"I love dogs and cats are so much fun!"` registrado en la consola. Pero ese no es el resultado deseado. En el siguiente paso, aprenderás cómo corregirlo para que el texto diga `"I love dogs and dogs are so much fun!"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes tener una variable `dogsOnlySentence`.

```js
assert.exists(dogsOnlySentence);
```

Tu variable `dogsOnlySentence` debe ser una cadena.

```js
assert.isString(dogsOnlySentence);
```

Debes asignar `exampleSentence.replace("cats", "dogs")` a tu variable `dogsOnlySentence`.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

Debes registrar `"Replacing all occurrences of cats with dogs:"` en la consola.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

Debes registrar la variable `dogsOnlySentence` en la consola.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
