---
id: 68e42de6396f5c49d90d1951
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

En una lección anterior, aprendiste sobre el método `replace()` que se usa para localizar una subcadena en una cadena y reemplazarla con otro valor.

Aquí hay un ejemplo:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

Recuerda que las cadenas son inmutables, lo que significa que el `text` original no se modifica en ese ejemplo. En su lugar, se crea una nueva cadena. 

Crea una nueva variable llamada `replacedString` y asígnale `originalString.replace("cats", "dogs")`. 

# --hints--

Debes tener una variable llamada `replacedString`.

```js
assert.exists(replacedString);
```

Tu variable `replacedString` debe ser una cadena.

```js
assert.isString(replacedString);
```

Debes asignar `originalString.replace("cats", "dogs")` a la variable `replacedString`.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
