---
id: 68e4465fba0bd256868b7deb
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Ahora es momento de trabajar con un nuevo método de cadenas.

Comienza creando una variable llamada `learningSentence` y asigna la cadena `"I love learning!"`. 

Luego, debajo de esa variable, registra en la consola `"Original learning sentence:"`. Debajo de ese `console.log()`, agrega otro `console.log()` para la variable `learningSentence`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes tener una variable `learningSentence`.

```js
assert.exists(learningSentence);
```

Tu variable `learningSentence` debe ser una cadena.

```js
assert.isString(learningSentence);
```

Debes asignar la cadena `"I love learning!"` a la variable `learningSentence`.

```js
assert.equal(learningSentence, "I love learning!");
```

Debes registrar `"Original learning sentence:"` en la consola.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

Debes registrar la variable `learningSentence` en la consola.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
