---
id: 68e44519c769bf545bb3f22a
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Para obtener el resultado deseado de `"I love dogs and dogs are so much fun!"`, necesitarás usar el método `replaceAll()` en lugar del método `replace()`. Este método devuelve una nueva cadena para todas las coincidencias con la subcadena. 

Actualiza el `exampleSentence.replace("cats", "dogs")` para usar el método `replaceAll()` en su lugar. Ahora deberías ver el texto correcto en la consola.

# --hints--

Debes actualizar la línea de código `exampleSentence.replace("cats", "dogs")` para usar el método `replaceAll()`.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
