---
id: 68e41084bb275539f72de9fe
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, practicarás trabajando con los métodos `replace()`, `replaceAll()` y `repeat()` construyendo una aplicación transformadora de cadenas. 

Para empezar, crea una variable llamada `originalString` y asígnale el valor de `"I love cats."`.

# --hints--

Debes tener una variable llamada `originalString`.

```js
assert.exists(originalString);
```

Tu variable `originalString` debe ser una cadena.

```js
assert.isString(originalString);
```

Debes asignar la cadena `"I love cats."` a la variable `originalString`.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
