---
id: 68a47dcdbce12193caf77012
title: تنفيذ محدد فهرس مرتب
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

في هذا المختبر ستُنشئ دالة تُرجع أدنى فهرس يجب إدخال قيمة فيه داخل مصفوفة بعد فرزها بترتيب تصاعدي.

**الهدف**: حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تمتلك دالة `getIndexToIns` تأخذ معلمتين: مصفوفة ورقم.  
2. يجب أن تستخدم طريقة `sort` لفرز المصفوفة بترتيب تصاعدي.
3. يجب أن تُرجع دالة `getIndexToIns` أدنى فهرس يجب إدخال الرقم فيه باستخدام طريقة `findIndex`.
4. يجب أن تُرجع دالة `getIndexToIns` دائمًا رقمًا.

**تلميح:**  
طريقة `findIndex` هي طريقة مصفوفة مدمجة في JavaScript. تأخذ دالة استدعاء راجع وتُرجع فهرس أول عنصر يحقق الشرط. كل من `findIndex` و `sort` هما دالتان رتبة عليا.

**أمثلة:**

- يجب أن تُرجع `getIndexToIns([1, 2, 3, 4], 1.5)` القيمة `1` لأن `1.5` أكبر من `1` (الفهرس `0`) وأصغر من `2` (الفهرس `1`).
- يجب أن تُرجع `getIndexToIns([20, 3, 5], 19)` القيمة `2` لأنه بعد الفرز إلى `[3, 5, 20]`، `19` أصغر من `20` (الفهرس `2`) وأكبر من `5` (الفهرس `1`).


# --hints--

يجب أن تمتلك دالة `getIndexToIns`.

```js
assert.isFunction(getIndexToIns);
```

يجب أن تُرجع `getIndexToIns` دائمًا رقمًا.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

يجب أن تستخدم دالتك طريقة `sort`.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


يجب أن تستخدم دالتك طريقة `findIndex`.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

يجب أن تُرجع `getIndexToIns([10, 20, 30, 40, 50], 35)` القيمة `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

يجب أن تُرجع `getIndexToIns([10, 20, 30, 40, 50], 30)` القيمة `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

يجب أن تُرجع `getIndexToIns([40, 60], 50)` القيمة `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

يجب أن تُرجع `getIndexToIns([3, 10, 5], 3)` القيمة `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

يجب أن تُرجع `getIndexToIns([5, 3, 20, 3], 5)` القيمة `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

يجب أن تُرجع `getIndexToIns([2, 20, 10], 19)` القيمة `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

يجب أن تُرجع `getIndexToIns([3, 10, 5], 11)` القيمة `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

يجب أن تُرجع `getIndexToIns([], 5)` القيمة `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
``` 

