---
id: 56104e9e514f539506016a5c
title: Rudia nambari zisizo za kawaida kwa kutumia mzunguko wa for
challengeType: 1
forumTopicId: 18212
dashedName: iterate-odd-numbers-with-a-for-loop
---

# --description--

Mizunguko ya for haipaswi kurudia hatua moja kwa moja. Kwa kubadilisha `final-expression`, tunaweza kuhesabu kwa nambari za hata.

Tutaanza kwa `i = 0` na kurudia mzunguko wakati `i < 10`. Tutazidisha `i` kwa 2 kila mzunguko kwa kutumia `i += 2`.

```js
const ourArray = [];

for (let i = 0; i < 10; i += 2) {
  ourArray.push(i);
}
```

`ourArray` sasa itakuwa na `[0, 2, 4, 6, 8]`. Tubadilishe `initialization` ili tuweze kuhesabu kwa nambari zisizo za kawaida.

# --instructions--

Tuma nambari zisizo za kawaida kutoka 1 hadi 9 kwenye `myArray` kwa kutumia mzunguko wa `for`.

# --hints--

Unapaswa kutumia mzunguko wa `for` kwa hili.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

`myArray` inapaswa kuwa sawa na `[1, 3, 5, 7, 9]`.

```js
assert.deepEqual(myArray, [1, 3, 5, 7, 9]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 10; i += 2) {
  myArray.push(i);
}
```
