---
id: 56105e7b514f539506016a5e
title: Hesabu kwa nyuma kwa kutumia mzunguko wa for
challengeType: 1
forumTopicId: 16808
dashedName: count-backwards-with-a-for-loop
---

# --description--

Mzunguko wa for pia unaweza kuhesabu kwa nyuma, mradi tu tuweze kufafanua masharti sahihi.

Ili kupunguza kwa mbili kila mzunguko, tutahitaji kubadilisha uanzishaji wetu, sharti, na usemi wa mwisho.

Tutaanza na `i = 10` na kuendelea mzunguko wakati `i > 0`. Tutapunguza `i` kwa 2 kila mzunguko kwa kutumia `i -= 2`.

```js
const ourArray = [];

for (let i = 10; i > 0; i -= 2) {
  ourArray.push(i);
}
```

`ourArray` sasa itakuwa na `[10, 8, 6, 4, 2]`. Tubadilishe uanzishaji wetu na usemi wa mwisho ili tuweze kuhesabu kwa nyuma kwa viwili kuunda safu ya data ya nambari zisizo sawa zinazopungua.

# --instructions--

Sukuma nambari zisizo sawa kutoka 9 hadi 1 kwenye `myArray` kwa kutumia mzunguko wa `for`.

# --hints--

Unapaswa kutumia mzunguko wa `for` kwa hili.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

Unapaswa kutumia njia ya safu ya data `push`.

```js
assert(__helpers.removeJSComments(code).match(/myArray.push/));
```

`myArray` inapaswa kuwa sawa na `[9, 7, 5, 3, 1]`.

```js
assert.deepEqual(myArray, [9, 7, 5, 3, 1]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 9; i > 0; i -= 2) {
  myArray.push(i);
}
```
