---
id: 56533eb9ac21ba0edf2244de
title: Kuongeza chaguo la msingi katika taarifa za kubadilisha
challengeType: 1
forumTopicId: 16653
dashedName: adding-a-default-option-in-switch-statements
---

# --description--

Katika taarifa ya `switch` huenda huwezi kubainisha thamani zote zinazowezekana kama taarifa za `case`. Badala yake, unaweza kuongeza taarifa ya `default` ambayo itatekelezwa ikiwa hakuna taarifa za `case` zinazolingana zitakazopatikana. Fikiria kama taarifa ya mwisho ya `else` katika mnyororo wa `if/else`.

Taarifa ya `default` inapaswa kuwa kesi ya mwisho.

```js
switch (num) {
  case value1:
    statement1;
    break;
  case value2:
    statement2;
    break;
...
  default:
    defaultStatement;
    break;
}
```

# --instructions--

Andika taarifa ya kubadilisha ili kuweka `answer` kwa masharti yafuatayo:  
`a` - `apple`  
`b` - `bird`  
`c` - `cat`  
`default` - `stuff`

# --hints--

`switchOfStuff("a")` inapaswa kurudisha mfuatano wa herufi `apple`

```js
assert(switchOfStuff('a') === 'apple');
```

`switchOfStuff("b")` inapaswa kurudisha mfuatano wa herufi `bird`

```js
assert(switchOfStuff('b') === 'bird');
```

`switchOfStuff("c")` inapaswa kurudisha mfuatano wa herufi `cat`

```js
assert(switchOfStuff('c') === 'cat');
```

`switchOfStuff("d")` inapaswa kurudisha mfuatano wa herufi `stuff`

```js
assert(switchOfStuff('d') === 'stuff');
```

`switchOfStuff(4)` inapaswa kurudisha mfuatano wa herufi `stuff`

```js
assert(switchOfStuff(4) === 'stuff');
```

Haupaswi kutumia taarifa za `if` au `else`

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Unapaswa kutumia taarifa ya `default`

```js
assert(switchOfStuff('string-to-trigger-default-case') === 'stuff');
```

Unapaswa kuwa na angalau taarifa 3 za `break`

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function switchOfStuff(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

switchOfStuff(1);
```

# --solutions--

```js
function switchOfStuff(val) {
  let answer = "";

  switch(val) {
    case "a":
      answer = "apple";
      break;
    case "b":
      answer = "bird";
      break;
    case "c":
      answer = "cat";
      break;
    default:
      answer = "stuff";
  }
  return answer;
}
```
