---
id: 56533eb9ac21ba0edf2244d8
title: Mlinganisho na Kiwango cha Mantiki cha Na
challengeType: 1
forumTopicId: 16799
dashedName: comparisons-with-the-logical-and-operator
---

# --description--

Wakati mwingine utahitaji kujaribu zaidi ya kitu kimoja kwa wakati mmoja. Kiwango cha <dfn>mantiki cha na</dfn> (`&&`) hurejesha `true` ikiwa na tu ikiwa <dfn>vipengele</dfn> vilivyo kushoto na kulia yake ni kweli.

Athari ile ile inaweza kupatikana kwa kupanga ndani taarifa ya `if` ndani ya nyingine ya `if`.

```js
if (num > 5) {
  if (num < 10) {
    return "Yes";
  }
}
return "No";
```

Msimbo huu utarejesha `Yes` ikiwa `num` ni kubwa kuliko `5` na ndogo kuliko `10`. Mantiki ile ile inaweza kuandikwa kwa kutumia kiwango cha <dfn>mantiki cha na</dfn>.

```js
if (num > 5 && num < 10) {
  return "Yes";
}
return "No";
```

# --instructions--

Badilisha taarifa mbili za if na taarifa moja, ukitumia kiwango cha `&&` ambacho kitarejesha mfuatano wa herufi `Yes` ikiwa `val` ni ndogo au sawa na `50` na kubwa au sawa na `25`. Vinginevyo, kitarejesha mfuatano wa herufi `No`.

# --hints--

Unapaswa kutumia kiwango cha `&&` mara moja

```js
assert(__helpers.removeJSComments(code).match(/&&/g).length === 1);
```

Unapaswa kuwa na taarifa moja tu ya `if`

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalAnd(0)` inapaswa kurejesha mfuatano wa herufi `No`

```js
assert(testLogicalAnd(0) === 'No');
```

`testLogicalAnd(24)` inapaswa kurejesha mfuatano wa herufi `No`

```js
assert(testLogicalAnd(24) === 'No');
```

`testLogicalAnd(25)` inapaswa kurejesha mfuatano wa herufi `Yes`

```js
assert(testLogicalAnd(25) === 'Yes');
```

`testLogicalAnd(30)` inapaswa kurejesha mfuatano wa herufi `Yes`

```js
assert(testLogicalAnd(30) === 'Yes');
```

`testLogicalAnd(50)` inapaswa kurejesha mfuatano wa herufi `Yes`

```js
assert(testLogicalAnd(50) === 'Yes');
```

`testLogicalAnd(51)` inapaswa kurejesha mfuatano wa herufi `No`

```js
assert(testLogicalAnd(51) === 'No');
```

`testLogicalAnd(75)` inapaswa kurejesha mfuatano wa herufi `No`

```js
assert(testLogicalAnd(75) === 'No');
```

`testLogicalAnd(80)` inapaswa kurejesha mfuatano wa herufi `No`

```js
assert(testLogicalAnd(80) === 'No');
```

# --seed--

## --seed-contents--

```js
function testLogicalAnd(val) {
  // Only change code below this line

  if (val) {
    if (val) {
      return "Yes";
    }
  }

  // Only change code above this line
  return "No";
}

testLogicalAnd(10);
```

# --solutions--

```js
function testLogicalAnd(val) {
  if (val >= 25 && val <= 50) {
    return "Yes";
  }
  return "No";
}
```
