---
id: 56533eb9ac21ba0edf2244ba
title: Elewa kutokubadilika kwa mfuatano wa herufi
challengeType: 1
forumTopicId: 18331
dashedName: understand-string-immutability
---

# --description--

Katika JavaScript, thamani za `String` ni <dfn>zisizobadilika</dfn>, ambayo inamaanisha kuwa haziwezi kubadilishwa mara tu zinapotengenezwa.

Kwa mfano, msimbo ufuatao utatoa kosa kwa sababu herufi `B` katika mfuatano wa herufi `Bob` haiwezi kubadilishwa kuwa herufi `J`:

```js
let myStr = "Bob";
myStr[0] = "J";
```

Kumbuka hii haimaanishi kuwa `myStr` haiwezi kupewa thamani mpya. Njia pekee ya kubadilisha `myStr` ni kuipatia thamani mpya, kama ifuatavyo:

```js
let myStr = "Bob";
myStr = "Job";
```

# --instructions--

Sahihisha mgawo wa thamani kwa `myStr` ili iwe na thamani ya mfuatano wa herufi `Hello World` kwa kutumia njia iliyoonyeshwa katika mfano hapo juu.

# --hints--

`myStr` inapaswa kuwa na thamani ya mfuatano wa herufi `Hello World`.

```js
assert(myStr === 'Hello World');
```

Haupaswi kubadilisha msimbo uliopo juu ya maoni yaliyotajwa.

```js
assert(/myStr = "Jello World"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let myStr = "Jello World";

// Only change code below this line
myStr[0] = "H"; // Change this line
// Only change code above this line
```

# --solutions--

```js
let myStr = "Jello World";
myStr = "Hello World";
```
