---
id: 56533eb9ac21ba0edf2244da
title: Utangulizi wa tamko la else
challengeType: 1
forumTopicId: 18207
dashedName: introducing-else-statements
---

# --description--

Wakati hali ya tamko la `if` ni kweli, kipande cha msimbo kinachofuata hufanyika. Je, inatokea nini wakati hali hiyo ni potofu? Kawaida hakuna kinachotokea. Kwa tamko la `else`, kipande mbadala cha msimbo kinaweza kufanyika.

```js
if (num > 10) {
  return "Bigger than 10";
} else {
  return "10 or Less";
}
```

# --instructions--

Unganisha tamko za `if` kuwa tamko moja la `if/else`.

# --hints--

Unapaswa kuwa na tamko moja tu la `if` katika mhariri

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

Unapaswa kutumia tamko la `else`

```js
assert(/else/g.test(__helpers.removeJSComments(code)));
```

`testElse(4)` inapaswa kurudisha mfuatano wa herufi `5 or Smaller`

```js
assert(testElse(4) === '5 or Smaller');
```

`testElse(5)` inapaswa kurudisha mfuatano wa herufi `5 or Smaller`

```js
assert(testElse(5) === '5 or Smaller');
```

`testElse(6)` inapaswa kurudisha mfuatano wa herufi `Bigger than 5`

```js
assert(testElse(6) === 'Bigger than 5');
```

`testElse(10)` inapaswa kurudisha mfuatano wa herufi `Bigger than 5`

```js
assert(testElse(10) === 'Bigger than 5');
```

Haupaswi kubadilisha msimbo uliopo juu au chini ya maoni yaliyotajwa.

```js
assert(/let result = "";/.test(__helpers.removeJSComments(code)) && /return result;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function testElse(val) {
  let result = "";
  // Only change code below this line

  if (val > 5) {
    result = "Bigger than 5";
  }

  if (val <= 5) {
    result = "5 or Smaller";
  }

  // Only change code above this line
  return result;
}

testElse(4);
```

# --solutions--

```js
function testElse(val) {
  let result = "";
  if(val > 5) {
    result = "Bigger than 5";
  } else {
    result = "5 or Smaller";
  }
  return result;
}
```
