---
id: 56533eb9ac21ba0edf2244be
title: Wigo wa kimataifa na vitendakazi
challengeType: 1
forumTopicId: 18193
dashedName: global-scope-and-functions
---

# --description--

Katika JavaScript, <dfn>wigo</dfn> unahusu uonekano wa vigezo. Vigezo vinavyotangazwa nje ya sehemu ya kitendakazi vina wigo wa <dfn>kimataifa</dfn>. Hii ina maana, vinaweza kuonekana kila mahali katika msimbo wako wa JavaScript.

Vigezo vinavyotangazwa bila maneno muhimu ya `let` au `const` huundwa moja kwa moja katika wigo wa `global`. Hii inaweza kusababisha matokeo yasiyotakiwa sehemu nyingine katika msimbo wako au wakati wa kuendesha kitendakazi tena. Unapaswa kila mara kutangaza vigezo vyako kwa kutumia `let` au `const`.

# --instructions--

Kwa kutumia `let` au `const`, tangaza kigezo cha kimataifa kinachoitwa `myGlobal` nje ya kitendakazi chochote. Kianzishe kwa thamani ya `10`.

Ndani ya kitendakazi `fun1`, weka thamani ya `5` kwa `oopsGlobal` ***bila*** kutumia maneno muhimu ya `var`, `let` au `const`.

# --before-each--

```js
var oopsGlobal;
```

# --hints--

`myGlobal` inapaswa kutangazwa

```js
assert(typeof myGlobal != 'undefined');
```

`myGlobal` inapaswa kuwa na thamani ya `10`

```js
assert(myGlobal === 10);
```

`myGlobal` inapaswa kutangazwa kwa kutumia maneno muhimu ya `let` au `const`

```js
assert(/(let|const)\s+myGlobal/.test(__helpers.removeJSComments(code)));
```

`oopsGlobal` inapaswa kuwa kigezo cha kimataifa na kuwa na thamani ya `5`

```js
fun1();
assert(typeof oopsGlobal != 'undefined');
```

# --seed--

## --seed-contents--

```js
// Declare the myGlobal variable below this line


function fun1() {
  // Assign 5 to oopsGlobal here

}

// Only change code above this line

function fun2() {
  let output = "";
  if (typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if (typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```

# --solutions--

```js
const myGlobal = 10;

function fun1() {
  oopsGlobal = 5;
}

function fun2() {
  let output = "";
  if(typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if(typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```
