---
id: 56533eb9ac21ba0edf2244ac
title: Ongeza nambari kwa JavaScript
challengeType: 1
forumTopicId: 18201
dashedName: increment-a-number-with-javascript
---

# --description--

Unaweza kwa urahisi <dfn>ongeza</dfn> au kuongeza moja kwenye kigezo kwa opereta ya `++`.

```js
i++;
```

ni sawa na

```js
i = i + 1;
```

**Kumbuka:** Mstari mzima unakuwa `i++;`, hivyo hauna haja ya ishara ya usawa.

# --instructions--

Badilisha msimbo kutumia opereta ya `++` kwenye `myVar`.

# --hints--

`myVar` inapaswa kuwa sawa na `88`.

```js
assert(myVar === 88);
```

Haupaswi kutumia opereta ya mgawo wa thamani.

```js
assert(
  /let\s+myVar\s*=\s*87;\s*\/*.*\s*([+]{2}\s*myVar|myVar\s*[+]{2})/.test(__helpers.removeJSComments(code))
);
```

Unapaswa kutumia opereta ya `++`.

```js
assert(/[+]{2}\s*myVar|myVar\s*[+]{2}/.test(__helpers.removeJSComments(code)));
```

Haupaswi kubadilisha msimbo juu ya maoni yaliyotajwa.

```js
assert(/let myVar = 87;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 87;

// Only change code below this line
myVar = myVar + 1;
```

# --solutions--

```js
let myVar = 87;
myVar++;
```
