---
id: 56533eb9ac21ba0edf2244d9
title: Mlinganisho na Operator ya Logical Or
challengeType: 1
forumTopicId: 16800
dashedName: comparisons-with-the-logical-or-operator
---

# --description--

Operator ya <dfn>logical or</dfn> (`||`) hurejesha `true` ikiwa mojawapo ya <dfn>operands</dfn> ni `true`. Vinginevyo, hurejesha `false`.

Operator ya <dfn>logical or</dfn> ina alama mbili za bomba: (`||`). Hii kawaida hupatikana kati ya funguo zako za Backspace na Enter.

Mfano uliopo hapa chini unapaswa kuonekana wa kawaida kutoka sehemu za awali.

```js
if (num > 10) {
  return "No";
}
if (num < 5) {
  return "No";
}
return "Yes";
```

Msimbo huu utarudisha `Yes` ikiwa `num` iko kati ya `5` na `10` (`5` na `10` zikiwemo). Mantiki ile ile inaweza kuandikwa kwa kutumia operator ya <dfn>logical or</dfn>.

```js
if (num > 10 || num < 5) {
  return "No";
}
return "Yes";
```

# --instructions--

Unganisha taarifa mbili za `if` kuwa taarifa moja ambayo itarudisha mfuatano wa herufi `Outside` ikiwa `val` haiko kati ya `10` na `20`, zikiwemo. Vinginevyo, rudisha mfuatano wa herufi `Inside`.

# --hints--

Unapaswa kutumia operator ya `||` mara moja

```js
assert(__helpers.removeJSComments(code).match(/\|\|/g).length === 1);
```

Unapaswa kuwa na taarifa moja tu ya `if`

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalOr(0)` inapaswa kurudisha mfuatano wa herufi `Outside`

```js
assert(testLogicalOr(0) === 'Outside');
```

`testLogicalOr(9)` inapaswa kurudisha mfuatano wa herufi `Outside`

```js
assert(testLogicalOr(9) === 'Outside');
```

`testLogicalOr(10)` inapaswa kurudisha mfuatano wa herufi `Inside`

```js
assert(testLogicalOr(10) === 'Inside');
```

`testLogicalOr(15)` inapaswa kurudisha mfuatano wa herufi `Inside`

```js
assert(testLogicalOr(15) === 'Inside');
```

`testLogicalOr(19)` inapaswa kurudisha mfuatano wa herufi `Inside`

```js
assert(testLogicalOr(19) === 'Inside');
```

`testLogicalOr(20)` inapaswa kurudisha mfuatano wa herufi `Inside`

```js
assert(testLogicalOr(20) === 'Inside');
```

`testLogicalOr(21)` inapaswa kurudisha mfuatano wa herufi `Outside`

```js
assert(testLogicalOr(21) === 'Outside');
```

`testLogicalOr(25)` inapaswa kurudisha mfuatano wa herufi `Outside`

```js
assert(testLogicalOr(25) === 'Outside');
```

# --seed--

## --seed-contents--

```js
function testLogicalOr(val) {
  // Only change code below this line

  if (val) {
    return "Outside";
  }

  if (val) {
    return "Outside";
  }

  // Only change code above this line
  return "Inside";
}

testLogicalOr(15);
```

# --solutions--

```js
function testLogicalOr(val) {
  if (val < 10 || val > 20) {
    return "Outside";
  }
  return "Inside";
}
```
