---
id: 56533eb9ac21ba0edf2244c3
title: Mgawo wa thamani yenye thamani ya return
challengeType: 1
forumTopicId: 16658
dashedName: assignment-with-a-returned-value
---

# --description--

Kama utakumbuka kutoka kwenye mazungumzo yetu kuhusu <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/storing-values-with-the-assignment-operator" target="_blank" rel="noopener noreferrer nofollow">Kuhifadhi Thamani kwa Opereta ya Mgawo</a>, kila kitu kilicho upande wa kulia wa ishara ya usawa kinatatuliwa kabla ya thamani kuwekwa. Hii inamaanisha tunaweza kuchukua thamani ya return ya kitendakazi na kuiweka kwenye kigezo.

Tuchukulie tumetangaza kitendakazi `sum` ambacho hujumlisha nambari mbili pamoja.

```js
ourSum = sum(5, 12);
```

Kupiga wito kitendakazi `sum` kwa hoja za `5` na `12` huleta thamani ya return ya `17`. Thamani hii ya return imepewa thamani kigezo cha `ourSum`.

# --instructions--

Piga wito kitendakazi `processArg` kwa hoja ya `7` na weka thamani yake ya return kwenye kigezo cha `processed`.

# --hints--

`processed` inapaswa kuwa na thamani ya `2`

```js
assert(processed === 2);
```

Unapaswa kuweka `processArg` kwa `processed`

```js
assert(/processed\s*=\s*processArg\(\s*7\s*\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

// Only change code below this line

```

# --solutions--

```js
var processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

processed = processArg(7);
```
