---
id: 56533eb9ac21ba0edf2244e0
title: Kubadilisha minyororo ya If Else kuwa Switch
challengeType: 1
forumTopicId: 18266
dashedName: replacing-if-else-chains-with-switch
---

# --description--

Ikiwa una chaguzi nyingi za kuchagua, usemi wa `switch` unaweza kuwa rahisi kuandika kuliko minyororo mingi ya `if`/`else if`. Ifuatayo:

```js
if (val === 1) {
  answer = "a";
} else if (val === 2) {
  answer = "b";
} else {
  answer = "c";
}
```

inaweza kubadilishwa na:

```js
switch (val) {
  case 1:
    answer = "a";
    break;
  case 2:
    answer = "b";
    break;
  default:
    answer = "c";
}
```

# --instructions--

Badilisha minyororo ya `if`/`else if` kuwa usemi wa `switch`.

# --hints--

Haupaswi kutumia usemi wowote wa `else` mahali popote katika mhariri

```js
assert(!/else/g.test(__helpers.removeJSComments(code)));
```

Haupaswi kutumia usemi wowote wa `if` mahali popote katika mhariri

```js
assert(!/if/g.test(__helpers.removeJSComments(code)));
```

Unapaswa kuwa na angalau usemi nne wa `break`

```js
assert(__helpers.removeJSComments(code).match(/break/g).length >= 4);
```

`chainToSwitch("bob")` inapaswa kurudisha mfuatano wa herufi `Marley`

```js
assert(chainToSwitch('bob') === 'Marley');
```

`chainToSwitch(42)` inapaswa kurudisha mfuatano wa herufi `The Answer`

```js
assert(chainToSwitch(42) === 'The Answer');
```

`chainToSwitch(1)` inapaswa kurudisha mfuatano wa herufi `There is no #1`

```js
assert(chainToSwitch(1) === 'There is no #1');
```

`chainToSwitch(99)` inapaswa kurudisha mfuatano wa herufi `Missed me by this much!`

```js
assert(chainToSwitch(99) === 'Missed me by this much!');
```

`chainToSwitch(7)` inapaswa kurudisha mfuatano wa herufi `Ate Nine`

```js
assert(chainToSwitch(7) === 'Ate Nine');
```

`chainToSwitch("John")` inapaswa kurudisha `""` (mfuatano wa herufi tupu)

```js
assert(chainToSwitch('John') === '');
```

`chainToSwitch(156)` inapaswa kurudisha `""` (mfuatano wa herufi tupu)

```js
assert(chainToSwitch(156) === '');
```

# --seed--

## --seed-contents--

```js
function chainToSwitch(val) {
  let answer = "";
  // Only change code below this line

  if (val === "bob") {
    answer = "Marley";
  } else if (val === 42) {
    answer = "The Answer";
  } else if (val === 1) {
    answer = "There is no #1";
  } else if (val === 99) {
    answer = "Missed me by this much!";
  } else if (val === 7) {
    answer = "Ate Nine";
  }

  // Only change code above this line
  return answer;
}

chainToSwitch(7);
```

# --solutions--

```js
function chainToSwitch(val) {
  let answer = "";

  switch (val) {
    case "bob":
      answer = "Marley";
      break;
    case 42:
      answer = "The Answer";
      break;
    case 1:
      answer = "There is no #1";
      break;
    case 99:
      answer = "Missed me by this much!";
      break;
    case 7:
      answer = "Ate Nine";
  }
  return answer;
}
```
