---
id: 56533eb9ac21ba0edf2244aa
title: Kuelewa Vigezo Visivyotangazwa Awali
challengeType: 1
forumTopicId: 18335
dashedName: understanding-uninitialized-variables
---

# --description--

Wakati vigezo vya JavaScript vinapotangazwa, vina thamani ya awali ya `undefined`. Ikiwa utafanya operesheni ya kihesabu kwa kigezo cha `undefined`, matokeo yako yatakuwa `NaN` ambayo inamaanisha <dfn>"Sio Nambari"</dfn>. Ikiwa utaunganisha mfuatano wa herufi na kigezo cha `undefined`, utapata <dfn>mfuatano wa herufi</dfn> wa `undefined`.

# --instructions--

Anzisha vigezo vitatu `a`, `b`, na `c` kwa `5`, `10`, na `"I am a"` mtawalia ili visiwe `undefined`.

# --hints--

`a` inapaswa kutangazwa na kuwa na thamani ya mwisho ya `6`.

```js
assert(typeof a === 'number' && a === 6);
```

`b` inapaswa kutangazwa na kuwa na thamani ya mwisho ya `15`.

```js
assert(typeof b === 'number' && b === 15);
```

`c` haipaswi kuwa na `undefined` na inapaswa kuwa na thamani ya mwisho ya mfuatano wa herufi `I am a String!`

```js
assert(!/undefined/.test(c) && c === 'I am a String!');
```

Haupaswi kubadilisha msimbo chini ya maoni yaliyotangazwa.

```js
assert(
  /a = a \+ 1;/.test(__helpers.removeJSComments(code)) &&
    /b = b \+ 5;/.test(__helpers.removeJSComments(code)) &&
    /c = c \+ " String!";/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
var a;
var b;
var c;
// Only change code above this line

a = a + 1;
b = b + 5;
c = c + " String!";
```

# --solutions--

```js
var a = 5;
var b = 10;
var c = "I am a";
a = a + 1;
b = b + 5;
c = c + " String!";
```
