---
id: 56533eb9ac21ba0edf2244dc
title: Uunganishaji wa Kauli za If Else
challengeType: 1
forumTopicId: 16772
dashedName: chaining-if-else-statements
---

# --description--

Kauli za `if/else` zinaweza kuunganishwa pamoja kwa mantiki tata. Hapa kuna <dfn>msimbo wa bandia</dfn> wa kauli nyingi za `if` / `else if` zilizounganishwa:

```js
if (condition1) {
  statement1
} else if (condition2) {
  statement2
} else if (condition3) {
  statement3
. . .
} else {
  statementN
}
```

# --instructions--

Andika kauli za `if`/`else if` zilizounganishwa ili kutimiza masharti yafuatayo:

`num < 5` - rudisha `Tiny`  
`num < 10` - rudisha `Small`  
`num < 15` - rudisha `Medium`  
`num < 20` - rudisha `Large`  
`num >= 20` - rudisha `Huge`

# --hints--

Unapaswa kuwa na angalau kauli nne za `else`

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 3);
```

Unapaswa kuwa na angalau kauli nne za `if`

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 3);
```

Unapaswa kuwa na angalau kauli moja ya `return`

```js
assert(__helpers.removeJSComments(code).match(/return/g).length >= 1);
```

`testSize(0)` inapaswa kurudisha mfuatano wa herufi `Tiny`

```js
assert(testSize(0) === 'Tiny');
```

`testSize(4)` inapaswa kurudisha mfuatano wa herufi `Tiny`

```js
assert(testSize(4) === 'Tiny');
```

`testSize(5)` inapaswa kurudisha mfuatano wa herufi `Small`

```js
assert(testSize(5) === 'Small');
```

`testSize(8)` inapaswa kurudisha mfuatano wa herufi `Small`

```js
assert(testSize(8) === 'Small');
```

`testSize(10)` inapaswa kurudisha mfuatano wa herufi `Medium`

```js
assert(testSize(10) === 'Medium');
```

`testSize(14)` inapaswa kurudisha mfuatano wa herufi `Medium`

```js
assert(testSize(14) === 'Medium');
```

`testSize(15)` inapaswa kurudisha mfuatano wa herufi `Large`

```js
assert(testSize(15) === 'Large');
```

`testSize(17)` inapaswa kurudisha mfuatano wa herufi `Large`

```js
assert(testSize(17) === 'Large');
```

`testSize(20)` inapaswa kurudisha mfuatano wa herufi `Huge`

```js
assert(testSize(20) === 'Huge');
```

`testSize(25)` inapaswa kurudisha mfuatano wa herufi `Huge`

```js
assert(testSize(25) === 'Huge');
```

# --seed--

## --seed-contents--

```js
function testSize(num) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

testSize(7);
```

# --solutions--

```js
function testSize(num) {
  if (num < 5) {
    return "Tiny";
  } else if (num < 10) {
    return "Small";
  } else if (num < 15) {
    return "Medium";
  } else if (num < 20) {
    return "Large";
  } else {
    return "Huge";
  }
}
```
