---
id: 56533eb9ac21ba0edf2244c8
title: Kupata sifa za kitu kwa kutumia uandishi wa mabano
challengeType: 1
forumTopicId: 16163
dashedName: accessing-object-properties-with-bracket-notation
---

# --description--

Njia ya pili ya kupata sifa za kitu ni kwa kutumia uandishi wa mabano (`[]`). Ikiwa sifa ya kitu unayojaribu kupata ina nafasi katika jina lake, utahitaji kutumia uandishi wa mabano.

Hata hivyo, bado unaweza kutumia uandishi wa mabano kwa sifa za kitu zisizo na nafasi.

Hapa kuna mfano wa kutumia uandishi wa mabano kusoma sifa ya kitu:

```js
const myObj = {
  "Space Name": "Kirk",
  "More Space": "Spock",
  "NoSpace": "USS Enterprise"
};

myObj["Space Name"];
myObj['More Space'];
myObj["NoSpace"];
```

`myObj["Space Name"]` itakuwa mfuatano wa herufi `Kirk`, `myObj['More Space']` itakuwa mfuatano wa herufi `Spock`, na `myObj["NoSpace"]` itakuwa mfuatano wa herufi `USS Enterprise`.

Kumbuka kwamba majina ya sifa yenye nafasi ndani yao lazima yawe katika nukuu (nukuu moja au mbili).

# --instructions--

Soma thamani za sifa `an entree` na `the drink` za `testObj` kwa kutumia uandishi wa mabano na ziweke katika `entreeValue` na `drinkValue` mtawalia.

# --hints--

`entreeValue` inapaswa kuwa mfuatano wa herufi

```js
assert(typeof entreeValue === 'string');
```

Thamani ya `entreeValue` inapaswa kuwa mfuatano wa herufi `hamburger`

```js
assert(entreeValue === 'hamburger');
```

`drinkValue` inapaswa kuwa mfuatano wa herufi

```js
assert(typeof drinkValue === 'string');
```

Thamani ya `drinkValue` inapaswa kuwa mfuatano wa herufi `water`

```js
assert(drinkValue === 'water');
```

Unapaswa kutumia uandishi wa mabano mara mbili

```js
assert(__helpers.removeJSComments(code).match(/testObj\s*?\[('|")[^'"]+\1\]/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};

// Only change code below this line
const entreeValue = testObj;   // Change this line
const drinkValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};
const entreeValue = testObj["an entree"];
const drinkValue = testObj['the drink'];
```
