---
id: 56533eb9ac21ba0edf2244b5
title: Kutoroka alama za nukuu halisi katika mfuatano wa herufi
challengeType: 1
forumTopicId: 17568
dashedName: escaping-literal-quotes-in-strings
---

# --description--

Unapokuwa unaeleza mfuatano wa herufi lazima uanze na kumaliza kwa nukuu moja au alama za nukuu maradufu. Nini kinatokea unapotaka nukuu halisi: `"` au `'` ndani ya mfuatano wako wa herufi?

Katika JavaScript, unaweza <dfn>kutoroka</dfn> nukuu ili isizingatiwe kama nukuu ya kumalizia mfuatano kwa kuweka <dfn>mkwaju wa nyuma</dfn> (`\`) mbele ya nukuu hiyo.

```js
const sampleStr = "Alan said, \"Peter is learning JavaScript\".";
```

Hii inaashiria kwa JavaScript kwamba nukuu inayofuata si mwisho wa mfuatano wa herufi, bali inapaswa kuonekana ndani ya mfuatano. Hivyo kama ungetaka kuichapisha kwenye konsoli, ungepata:

```js
Alan said, "Peter is learning JavaScript".
```

# --instructions--

Tumia <dfn>mikwaju ya nyuma</dfn> kuweka thamani ya mfuatano wa herufi kwenye variable `myStr` ili kama ungetaka kuichapisha kwenye konsoli, ungaona:

```js
I am a "double quoted" string inside "double quotes".
```

# --hints--

Unapaswa kutumia alama za nukuu maradufu mbili (`"`) na alama za nukuu maradufu nne zilizokotoroka (`\"`).

```js
assert(__helpers.removeJSComments(code).match(/\\"/g).length === 4 && __helpers.removeJSComments(code).match(/[^\\]"/g).length === 2);
```

Variable `myStr` inapaswa kuwa na mfuatano wa herufi: `I am a "double quoted" string inside "double quotes".`

```js
assert(/I am a "double quoted" string inside "double quotes(\."|"\.)$/.test(myStr));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "I am a \"double quoted\" string inside \"double quotes\".";
```
