---
id: 56533eb9ac21ba0edf2244b1
title: Mgawo mchanganyiko na kuzidisha kwa kuongezea
challengeType: 1
forumTopicId: 16662
dashedName: compound-assignment-with-augmented-multiplication
---

# --description--

Kipengele cha `*=` huzidisha kigezo kwa nambari.

```js
myVar = myVar * 5;
```

kitazidisha `myVar` kwa `5`. Hii inaweza kuandikwa upya kama:

```js
myVar *= 5;
```

# --instructions--

Badilisha migawo wa thamani ya `a`, `b`, na `c` ili kutumia kipengele cha `*=`.

# --hints--

`a` inapaswa kuwa sawa na `25`.

```js
assert(a === 25);
```

`b` inapaswa kuwa sawa na `36`.

```js
assert(b === 36);
```

`c` inapaswa kuwa sawa na `46`.

```js
assert(c === 46);
```

Unapaswa kutumia kipengele cha `*=` kwa kila kigezo.

```js
assert(__helpers.removeJSComments(code).match(/\*=/g).length === 3);
```

Hupaswi kubadilisha msimbo ulio juu ya maoni yaliyotajwa.

```js
assert(
  /let a = 5;/.test(__helpers.removeJSComments(code)) &&
    /let b = 12;/.test(__helpers.removeJSComments(code)) &&
    /let c = 4\.6;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 12;
let c = 4.6;

// Only change code below this line
a = a * 5;
b = 3 * b;
c = c * 10;
```

# --solutions--

```js
let a = 5;
let b = 12;
let c = 4.6;

a *= 5;
b *= 3;
c *= 10;
```
