---
id: 56533eb9ac21ba0edf2244c2
title: Rudisha Thamani Kutoka kwa Kitendakazi kwa Kutumia Return
challengeType: 1
forumTopicId: 18271
dashedName: return-a-value-from-a-function-with-return
---

# --description--

Tunaweza kupitisha thamani ndani ya kitendakazi kwa kutumia <dfn>hoja</dfn>. Unaweza kutumia tamko la `return` kutuma thamani kurudi kutoka kwa kitendakazi.

**Mfano**

```js
function plusThree(num) {
  return num + 3;
}

const answer = plusThree(5);
```

`answer` ina thamani ya `8`.

`plusThree` hupokea <dfn>hoja</dfn> ya `num` na hurudisha thamani sawa na `num + 3`.

# --instructions--

Tengeneza kitendakazi `timesFive` kinachokubali hoja moja, ikazidisha kwa `5`, na kurudisha thamani mpya.

# --hints--

`timesFive` inapaswa kuwa kitendakazi

```js
assert(typeof timesFive === 'function');
```

`timesFive(5)` inapaswa kurudisha `25`

```js
assert(timesFive(5) === 25);
```

`timesFive(2)` inapaswa kurudisha `10`

```js
assert(timesFive(2) === 10);
```

`timesFive(0)` inapaswa kurudisha `0`

```js
assert(timesFive(0) === 0);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function timesFive(num) {
  return num * 5;
}
timesFive(10);
```
