---
id: 56533eb9ac21ba0edf2244a9
title: Kuanza Variables kwa opereta ya mgawo
challengeType: 1
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

Ni kawaida kuanzisha (initialize) variable kwa thamani ya awali katika mstari mmoja na inapotangazwa.

```js
var myVar = 0;
```

Hutengeneza variable mpya inayoitwa `myVar` na kuwekea thamani ya awali ya `0`.

# --instructions--

Tangaza variable `a` na `var` na uianzishe kwa thamani ya `9`.

# --hints--

Unapaswa kuanzisha `a` kwa thamani ya `9`.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
