---
id: 56533eb9ac21ba0edf2244b7
title: Kuunganisha mfuatano wa herufi kwa kutumia operator ya plus
challengeType: 1
forumTopicId: 16802
dashedName: concatenating-strings-with-plus-operator
---

# --description--

Katika JavaScript, wakati operator ya `+` inapotumika na thamani ya `String`, huitwa operator ya <dfn>kuunganisha</dfn>. Unaweza kutengeneza mfuatano mpya wa herufi kutoka kwa mfuatano mingine kwa <dfn>kuwaunganisha</dfn> pamoja.

**Mfano**

```js
'My name is Alan,' + ' I concatenate.'
```

**Kumbuka:** Angalia nafasi. Kuunganisha hakutoi nafasi kati ya mfuatano unaounganishwa, kwa hivyo utahitaji kuziweka nafasi mwenyewe.

Mfano:

```js
const ourStr = "I come first. " + "I come second.";
```

Mfuatano wa herufi `I come first. I come second.` utaonyeshwa kwenye konsoli.
# --instructions--

Tengeneza `myStr` kutoka kwa mfuatano wa herufi `This is the start.` na `This is the end.` kwa kutumia operator ya `+`. Hakikisha unaweka nafasi kati ya mfuatano huo miwili.

# --hints--

`myStr` inapaswa kuwa na herufi moja ya nafasi kati ya mfuatano wa herufi miwili.

```js
assert(/start\. This/.test(myStr));
```

`myStr` inapaswa kuwa na thamani ya mfuatano wa herufi `This is the start. This is the end.`

```js
assert(myStr === 'This is the start. This is the end.');
```

Unapaswa kutumia operator ya `+` kutengeneza `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/(["']).*\1\s*\+\s*(["']).*\2/g));
```

`myStr` inapaswa kuundwa kwa kutumia neno kuu la `const`.

```js
assert(/const\s+myStr/.test(__helpers.removeJSComments(code)));
```

Unapaswa kuweka thamani hiyo kwenye mabadiliko ya `myStr`.

```js
assert(/myStr\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "This is the start. " + "This is the end.";
```
