---
id: 56533eb9ac21ba0edf2244db
title: Utangulizi wa taarifa za else if
challengeType: 1
forumTopicId: 18206
dashedName: introducing-else-if-statements
---

# --description--

Ikiwa una masharti mengi yanayohitaji kushughulikiwa, unaweza kuunganisha taarifa za `if` pamoja na taarifa za `else if`.

```js
if (num > 15) {
  return "Bigger than 15";
} else if (num < 5) {
  return "Smaller than 5";
} else {
  return "Between 5 and 15";
}
```

# --instructions--

Badilisha mantiki ili itumie taarifa za `else if`.

# --hints--

Unapaswa kuwa na angalau taarifa mbili za `else`

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 1);
```

Unapaswa kuwa na angalau taarifa mbili za `if`

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 1);
```

Unapaswa kuwa na mabano ya mviringo ya kufunga na kufungua kwa kila kipande cha msimbo cha `if else`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s+if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s*\{[\s\S]+\s*\}/
  )
);
```

`testElseIf(0)` inapaswa kurudisha mfuatano wa herufi `Smaller than 5`

```js
assert(testElseIf(0) === 'Smaller than 5');
```

`testElseIf(5)` inapaswa kurudisha mfuatano wa herufi `Between 5 and 10`

```js
assert(testElseIf(5) === 'Between 5 and 10');
```

`testElseIf(7)` inapaswa kurudisha mfuatano wa herufi `Between 5 and 10`

```js
assert(testElseIf(7) === 'Between 5 and 10');
```

`testElseIf(10)` inapaswa kurudisha mfuatano wa herufi `Between 5 and 10`

```js
assert(testElseIf(10) === 'Between 5 and 10');
```

`testElseIf(12)` inapaswa kurudisha mfuatano wa herufi `Greater than 10`

```js
assert(testElseIf(12) === 'Greater than 10');
```

# --seed--

## --seed-contents--

```js
function testElseIf(val) {
  if (val > 10) {
    return "Greater than 10";
  }

  if (val < 5) {
    return "Smaller than 5";
  }

  return "Between 5 and 10";
}

testElseIf(7);
```

# --solutions--

```js
function testElseIf(val) {
  if(val > 10) {
    return "Greater than 10";
  } else if(val < 5) {
    return "Smaller than 5";
  } else {
    return "Between 5 and 10";
  }
}
```
