---
id: 56533eb9ac21ba0edf2244b9
title: Kujenga mfuatano wa herufi kwa kutumia vigezo
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

Wakati mwingine utahitaji kujenga mfuatano wa herufi. Kwa kutumia opereta ya kuunganisha (`+`), unaweza kuweka kigezo kimoja au zaidi ndani ya mfuatano wa herufi unaojenga.

Mfano:

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

`ourStr` ingekuwa na thamani ya mfuatano wa herufi `Hello, our name is freeCodeCamp, how are you?`.

# --instructions--

Weka `myName` kuwa mfuatano wa herufi unaolingana na jina lako na jenga `myStr` ukiwa na `myName` kati ya mfuatano wa herufi `My name is ` na ` and I am well!`

# --hints--

`myName` inapaswa kuwekwa kuwa mfuatano wa herufi wenye angalau herufi 3.

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

Unapaswa kutumia opereta mbili za `+` kujenga `myStr` ukiwa na `myName` ndani yake.

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
