---
id: 56533eb9ac21ba0edf2244bd
title: Kupitisha Thamani kwa Vitendakazi kwa Kutumia Hoja
challengeType: 1
forumTopicId: 18254
dashedName: passing-values-to-functions-with-arguments
---

# --description--

<dfn>Vigezo</dfn> ni vigezo vinavyotumika kama viwakilishi kwa thamani zinazotakiwa kuingizwa kwenye kitendakazi wakati kinapoitwa. Kitendakazi kinapofafanuliwa, kawaida hufafanuliwa pamoja na kigezo kimoja au zaidi. Thamani halisi zinazotolewa (au <dfn>"kupitishwa"</dfn>) kwenye kitendakazi wakati kinapoitwa zinajulikana kama <dfn>hoja</dfn>.

Hapa kuna kitendakazi chenye vigezo viwili, `param1` na `param2`:

```js
function testFun(param1, param2) {
  console.log(param1, param2);
}
```

Kisha tunaweza kuita `testFun` hivi: `testFun("Hello", "World");`. Tumepitisha hoja mbili za mfuatano wa herufi, `Hello` na `World`. Ndani ya kitendakazi, `param1` itakuwa sawa na mfuatano wa herufi `Hello` na `param2` itakuwa sawa na mfuatano wa herufi `World`. Kumbuka unaweza kuita `testFun` tena kwa hoja tofauti na vigezo vitachukua thamani ya hoja mpya.

# --instructions--

<ol><li>Tengeneza kitendakazi kinachoitwa <code>functionWithArgs</code> kinachokubali hoja mbili na kutoa jumla yao kwenye konsoli ya msanidi programu.</li><li>Ita kitendakazi hicho kwa kutumia nambari mbili kama hoja.</li></ol>

# --before-each--

```js
var logOutput;
```

# --hints--

`functionWithArgs` inapaswa kuwa kitendakazi.

```js
assert(typeof functionWithArgs === 'function');
```

`functionWithArgs(1,2)` inapaswa kutoa `3`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(1, 2);
assert(logOutput == 3);
```

`functionWithArgs(7,9)` inapaswa kutoa `16`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(7, 9);
assert(logOutput == 16);
```

Unapaswa kuita `functionWithArgs` kwa nambari mbili baada ya kuifafanua.

```js
assert(
  /functionWithArgs\([-+]?\d*\.?\d*,[-+]?\d*\.?\d*\)/.test(
    __helpers.removeJSComments(code).replace(/\s/g, '')
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function functionWithArgs(a, b) {
  console.log(a + b);
}
functionWithArgs(10, 5);
```
