---
id: 56533eb9ac21ba0edf2244d2
title: Ulinganisho na opereta ya kutofautiana
challengeType: 1
forumTopicId: 16787
dashedName: comparison-with-the-inequality-operator
---

# --description--

Opereta ya kutofautiana (`!=`) ni kinyume cha opereta ya usawa. Kutofautiana maana yake si sawa. Opereta ya kutofautiana hurejesha `false` wakati opereta ya usawa itakaporejesha `true` na *kivilevile*. Kama opereta ya usawa, opereta ya kutofautiana itabadilisha aina za data za thamani wakati wa kulinganisha.

**Mifano**

```js
1 !=  2    // true
1 != "1"   // false
1 != '1'   // false
1 != true  // false
0 != false // false
```

# --instructions--

Ongeza opereta ya kutofautiana `!=` katika usemi wa `if` ili kitendakazi kirejeshe mfuatano wa herufi `Not Equal` wakati `val` si sawa na `99`.

# --hints--

`testNotEqual(99)` inapaswa kurejesha mfuatano wa herufi `Equal`

```js
assert(testNotEqual(99) === 'Equal');
```

`testNotEqual("99")` inapaswa kurejesha mfuatano wa herufi `Equal`

```js
assert(testNotEqual('99') === 'Equal');
```

`testNotEqual(12)` inapaswa kurejesha mfuatano wa herufi `Not Equal`

```js
assert(testNotEqual(12) === 'Not Equal');
```

`testNotEqual("12")` inapaswa kurejesha mfuatano wa herufi `Not Equal`

```js
assert(testNotEqual('12') === 'Not Equal');
```

`testNotEqual("bob")` inapaswa kurejesha mfuatano wa herufi `Not Equal`

```js
assert(testNotEqual('bob') === 'Not Equal');
```

Unapaswa kutumia opereta ya `!=`

```js
assert(__helpers.removeJSComments(code).match(/(?!!==)!=/));
```

# --seed--

## --seed-contents--

```js
// Setup
function testNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testNotEqual(10);
```

# --solutions--

```js
function testNotEqual(val) {
  if (val != 99) {
    return "Not Equal";
  }
  return "Equal";
}
```
