---
id: 56533eb9ac21ba0edf2244b8
title: Kuunganisha mfuatano wa herufi kwa kutumia operator ya plus equals
challengeType: 1
forumTopicId: 16803
dashedName: concatenating-strings-with-the-plus-equals-operator
---

# --description--

Tunaweza pia kutumia operator ya `+=` ku<dfn>unganisha</dfn> mfuatano wa herufi mwishoni mwa mfuatano wa herufi uliopo tayari. Hii inaweza kusaidia sana kugawanya mfuatano mrefu wa herufi katika mistari kadhaa.

**Kumbuka:** Angalia nafasi. Kuunganisha hakutoi nafasi kati ya mfuatano wa herufi unaounganishwa, kwa hivyo utahitaji kuziweka nafasi hizo mwenyewe.

Mfano:

```js
let ourStr = "I come first. ";
ourStr += "I come second.";
```

`ourStr` sasa ina thamani ya mfuatano wa herufi `I come first. I come second.`.

# --instructions--

Jenga `myStr` katika mistari kadhaa kwa kuunganisha mfuatano huu wa herufi mbili: `This is the first sentence.` na `This is the second sentence.` kwa kutumia operator ya `+=`. Tumia operator ya `+=` kama inavyoonyeshwa katika mfano na hakikisha unaweka nafasi kati ya mfuatano wa herufi mbili. Anza kwa kuweka thamani ya mfuatano wa herufi wa kwanza kwa `myStr`, kisha ongeza mfuatano wa herufi wa pili.

# --hints--

`myStr` inapaswa kuwa na tabia ya nafasi moja kati ya mfuatano wa herufi mbili.

```js
assert(/sentence\. This/.test(myStr));
```

`myStr` inapaswa kuwa na thamani ya mfuatano wa herufi `This is the first sentence. This is the second sentence.`

```js
assert(myStr === 'This is the first sentence. This is the second sentence.');
```

Unapaswa kutumia operator ya `+=` kujenga `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*(["']).*\1/g));
```

# --seed--

## --seed-contents--

```js
let myStr;
```

# --solutions--

```js
let myStr = "This is the first sentence. ";
myStr += "This is the second sentence.";
```
