---
id: 56533eb9ac21ba0edf2244df
title: Chaguzi nyingi zinazofanana katika taarifa za kubadilisha
challengeType: 1
forumTopicId: 18242
dashedName: multiple-identical-options-in-switch-statements
---

# --description--

Kama taarifa ya `break` itakosekana katika `switch` ya taarifa ya `case`, taarifa zifuatazo za `case` zatekelezwa hadi `break` itakapokutana. Ikiwa unaingizo nyingi zenye matokeo sawa, unaweza kuzawakilisha katika taarifa ya `switch` kama hii:

```js
let result = "";
switch (val) {
  case 1:
  case 2:
  case 3:
    result = "1, 2, or 3";
    break;
  case 4:
    result = "4 alone";
}
```

Matukio ya 1, 2, na 3 yote yatatoa matokeo sawa.

# --instructions--

Andika taarifa ya kubadilisha kuweka `answer` kwa anuwai zifuatazo:  
`1-3` - `Low`  
`4-6` - `Mid`  
`7-9` - `High`

**Kumbuka:** Utahitaji kuwa na taarifa ya `case` kwa kila nambari katika anuwai hiyo.

# --hints--

`sequentialSizes(1)` inapaswa kurudisha mfuatano wa herufi `Low`

```js
assert(sequentialSizes(1) === 'Low');
```

`sequentialSizes(2)` inapaswa kurudisha mfuatano wa herufi `Low`

```js
assert(sequentialSizes(2) === 'Low');
```

`sequentialSizes(3)` inapaswa kurudisha mfuatano wa herufi `Low`

```js
assert(sequentialSizes(3) === 'Low');
```

`sequentialSizes(4)` inapaswa kurudisha mfuatano wa herufi `Mid`

```js
assert(sequentialSizes(4) === 'Mid');
```

`sequentialSizes(5)` inapaswa kurudisha mfuatano wa herufi `Mid`

```js
assert(sequentialSizes(5) === 'Mid');
```

`sequentialSizes(6)` inapaswa kurudisha mfuatano wa herufi `Mid`

```js
assert(sequentialSizes(6) === 'Mid');
```

`sequentialSizes(7)` inapaswa kurudisha mfuatano wa herufi `High`

```js
assert(sequentialSizes(7) === 'High');
```

`sequentialSizes(8)` inapaswa kurudisha mfuatano wa herufi `High`

```js
assert(sequentialSizes(8) === 'High');
```

`sequentialSizes(9)` inapaswa kurudisha mfuatano wa herufi `High`

```js
assert(sequentialSizes(9) === 'High');
```

Haupaswi kutumia taarifa za `if` au `else`

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Unapaswa kuwa na taarifa tisa za `case`

```js
assert(__helpers.removeJSComments(code).match(/case/g).length === 9);
```

# --seed--

## --seed-contents--

```js
function sequentialSizes(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

sequentialSizes(1);
```

# --solutions--

```js
function sequentialSizes(val) {
  let answer = "";

  switch (val) {
    case 1:
    case 2:
    case 3:
      answer = "Low";
      break;
    case 4:
    case 5:
    case 6:
      answer = "Mid";
      break;
    case 7:
    case 8:
    case 9:
      answer = "High";
  }

  return answer;
}
```
