---
id: 56533eb9ac21ba0edf2244bb
title: Maneno ya kujaza
challengeType: 1
forumTopicId: 18377
dashedName: word-blanks
---

# --description--

Unapewa sentensi zenye maneno kadhaa yaliyokosekana, kama vile majina, vitenzi, vivumishi na vitenzi kivumishi. Kisha unajaza sehemu zilizokosekana kwa maneno unayochagua kwa njia ambayo sentensi iliyokamilishwa ina maana.

Fikiria sentensi hii:

```md
It was really ____, and we ____ ourselves ____.
```

Sentensi hii ina sehemu tatu zilizokosekana - kivumishi, kitenzi na kitenzi kivumishi, na tunaweza kuongeza maneno tunayochagua ili kuimaliza. Kisha tunaweza kuweka sentensi iliyokamilishwa kwenye kigezo kama ifuatavyo:

```js
const sentence = "It was really " + "hot" + ", and we " + "laughed" + " ourselves " + "silly" + ".";
```

# --instructions--

Katika changamoto hii, tunakupa jina, kitenzi, kivumishi na kitenzi kivumishi. Unahitaji kuunda sentensi kamili ukitumia maneno unayochagua, pamoja na maneno tunayokupa.

Utahitaji kutumia opereta ya kuunganisha mfuatano wa herufi `+` kujenga mfuatano mpya wa herufi, ukitumia vigezo vilivyotolewa: `myNoun`, `myAdjective`, `myVerb`, na `myAdverb`. Kisha utaweka mfuatano ulioundwa kwenye kigezo `wordBlanks`. Haufai kubadilisha maneno yaliyotolewa kwa vigezo.

Pia utahitaji kuzingatia nafasi katika mfuatano wako wa herufi, ili sentensi ya mwisho iwe na nafasi kati ya maneno yote. Matokeo yanapaswa kuwa sentensi kamili.

# --before-each--

```js
const removeAssignments = str => str
  .replace(/myNoun\s*=\s*["']dog["']/g, '')
  .replace(/myAdjective\s*=\s*["']big["']/g, '')
  .replace(/myVerb\s*=\s*["']ran["']/g, '')
  .replace(/myAdverb\s*=\s*["']quickly["']/g, '');
```

# --hints--

`wordBlanks` inapaswa kuwa mfuatano wa herufi.

```js
assert(typeof wordBlanks === 'string');
```

Haufai kubadilisha thamani zilizowekwa kwa `myNoun`, `myVerb`, `myAdjective` au `myAdverb`.

```js
assert(
  myNoun === 'dog' &&
    myVerb === 'ran' &&
    myAdjective === 'big' &&
    myAdverb === 'quickly'
);
```

Haufai kutumia moja kwa moja thamani za `dog`, `ran`, `big`, au `quickly` kuunda `wordBlanks`.

```js
const newCode = removeAssignments(__helpers.removeJSComments(code));
assert(
  !/dog/.test(newCode) &&
    !/ran/.test(newCode) &&
    !/big/.test(newCode) &&
    !/quickly/.test(newCode)
);
```

`wordBlanks` inapaswa kuwa na maneno yote yaliyowekwa kwa vigezo `myNoun`, `myVerb`, `myAdjective` na `myAdverb` yakiwa yamegawanyika kwa herufi zisizo za maneno (na maneno mengine yoyote unayochagua).

```js
assert(
  /\bdog\b/.test(wordBlanks) &&
    /\bbig\b/.test(wordBlanks) &&
    /\bran\b/.test(wordBlanks) &&
    /\bquickly\b/.test(wordBlanks)
);
```

# --seed--

## --seed-contents--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

// Only change code below this line
const wordBlanks = ""; // Change this line
// Only change code above this line
```

# --solutions--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

let wordBlanks = "Once there was a " + myNoun + " which was very " + myAdjective + ". ";
wordBlanks += "It " + myVerb + " " + myAdverb + " around the yard.";
```
