---
id: 56533eb9ac21ba0edf2244d0
title: Ulinganishi na opereta ya usawa
challengeType: 1
forumTopicId: 16784
dashedName: comparison-with-the-equality-operator
---

# --description--

Kuna opereta nyingi za <dfn>ulinganishi</dfn> katika JavaScript. Opereta zote hizi hurejesha thamani ya kweli au siyo kweli `true` au `false`.

Opereta ya msingi kabisa ni opereta ya usawa `==`. Opereta ya usawa hulinganisha thamani mbili na kurejesha `true` ikiwa ni sawa au `false` ikiwa si sawa. Kumbuka kuwa usawa ni tofauti na mgawo wa thamani (`=`), ambao huweka thamani iliyoko upande wa kulia wa opereta kwenye kigezo upande wa kushoto.

```js
function equalityTest(myVal) {
  if (myVal == 10) {
    return "Equal";
  }
  return "Not Equal";
}
```

Ikiwa `myVal` ni sawa na `10`, opereta ya usawa hurejesha `true`, hivyo msimbo ulioko ndani ya mabano ya mviringo utaendeshwa, na kitendakazi kitarudisha `Equal`. Vinginevyo, kitendakazi kitarudisha `Not Equal`. Ili JavaScript ilinganishe aina mbili tofauti za <dfn>data</dfn> (kwa mfano, `numbers` na `strings`), lazima ibadilishe aina moja kuwa nyingine. Hii inajulikana kama kulazimisha aina. Hata hivyo, mara inapoifanya, inaweza kulinganisha maneno kama ifuatavyo:

```js
1   ==  1  // true
1   ==  2  // false
1   == '1' // true
"3" ==  3  // true
```

# --instructions--

Ongeza opereta ya usawa kwenye mstari ulioonyeshwa ili kitendakazi kirejeshe mfuatano wa herufi `Equal` wakati `val` ni sawa na `12`.

# --hints--

`testEqual(10)` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(testEqual(10) === 'Not Equal');
```

`testEqual(12)` inapaswa kurudisha mfuatano wa herufi `Equal`

```js
assert(testEqual(12) === 'Equal');
```

`testEqual("12")` inapaswa kurudisha mfuatano wa herufi `Equal`

```js
assert(testEqual('12') === 'Equal');
```

Unapaswa kutumia opereta `==`

```js
assert(__helpers.removeJSComments(code).match(/==/g) && !__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function testEqual(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testEqual(10);
```

# --solutions--

```js
function testEqual(val) {
  if (val == 12) {
    return "Equal";
  }
  return "Not Equal";
}
```
