---
id: 56533eb9ac21ba0edf2244d7
title: Ulinganisho na opereta ya chini ya au sawa na
challengeType: 1
forumTopicId: 16788
dashedName: comparison-with-the-less-than-or-equal-to-operator
---

# --description--

Opereta ya chini ya au sawa na (`<=`) inalinganisha thamani za nambari mbili. Ikiwa nambari upande wa kushoto ni ndogo au sawa na nambari upande wa kulia, inarudisha `true`. Ikiwa nambari upande wa kushoto ni kubwa kuliko nambari upande wa kulia, inarudisha `false`. Kama opereta ya usawa, opereta ya chini ya au sawa na hubadilisha aina za data.

**Mifano**

```js
4   <= 5 // true
'7' <= 7 // true
5   <= 5 // true
3   <= 2 // false
'8' <= 4 // false
```

# --instructions--

Ongeza opereta ya chini ya au sawa na kwenye mistari iliyotajwa ili taarifa za return ziwe na maana.

# --hints--

`testLessOrEqual(0)` inapaswa kurudisha mfuatano wa herufi `Smaller Than or Equal to 12`

```js
assert(testLessOrEqual(0) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(11)` inapaswa kurudisha mfuatano wa herufi `Smaller Than or Equal to 12`

```js
assert(testLessOrEqual(11) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(12)` inapaswa kurudisha mfuatano wa herufi `Smaller Than or Equal to 12`

```js
assert(testLessOrEqual(12) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(23)` inapaswa kurudisha mfuatano wa herufi `Smaller Than or Equal to 24`

```js
assert(testLessOrEqual(23) === 'Smaller Than or Equal to 24');
```

`testLessOrEqual(24)` inapaswa kurudisha mfuatano wa herufi `Smaller Than or Equal to 24`

```js
assert(testLessOrEqual(24) === 'Smaller Than or Equal to 24');
```

`testLessOrEqual(25)` inapaswa kurudisha mfuatano wa herufi `More Than 24`

```js
assert(testLessOrEqual(25) === 'More Than 24');
```

`testLessOrEqual(55)` inapaswa kurudisha mfuatano wa herufi `More Than 24`

```js
assert(testLessOrEqual(55) === 'More Than 24');
```

Unapaswa kutumia opereta `<=` angalau mara mbili

```js
assert(__helpers.removeJSComments(code).match(/val\s*<=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessOrEqual(val) {
  if (val) {  // Change this line
    return "Smaller Than or Equal to 12";
  }

  if (val) {  // Change this line
    return "Smaller Than or Equal to 24";
  }

  return "More Than 24";
}

testLessOrEqual(10);
```

# --solutions--

```js
function testLessOrEqual(val) {
  if (val <= 12) {  // Change this line
    return "Smaller Than or Equal to 12";
  }

  if (val <= 24) {  // Change this line
    return "Smaller Than or Equal to 24";
  }

  return "More Than 24";
}
```
