---
id: 56533eb9ac21ba0edf2244d1
title: Ulinganishaji na opereta ya usawa mkali
challengeType: 1
forumTopicId: 16790
dashedName: comparison-with-the-strict-equality-operator
---

# --description--

Usawa mkali (`===`) ni sawa na opereta ya usawa (`==`). Hata hivyo, tofauti na opereta ya usawa, ambayo hujaribu kubadilisha thamani zote mbili zinazolinganishwa kuwa aina moja, opereta ya usawa mkali haifanyi ubadilishaji wa aina.

Ikiwa thamani zinazolinganishwa zina aina tofauti, zinachukuliwa kuwa si sawa, na opereta ya usawa mkali itarudisha false.

**Mifano**

```js
3 ===  3  // true
3 === '3' // false
```

Katika mfano wa pili, `3` ni aina ya `Number` na `'3'` ni aina ya `String`.

# --instructions--

Tumia opereta ya usawa mkali katika tamko la `if` ili kitendakazi kirejee mfuatano wa herufi `Equal` wakati `val` ni sawa kabisa na `7`.

# --hints--

`testStrict(10)` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(testStrict(10) === 'Not Equal');
```

`testStrict(7)` inapaswa kurudisha mfuatano wa herufi `Equal`

```js
assert(testStrict(7) === 'Equal');
```

`testStrict("7")` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(testStrict('7') === 'Not Equal');
```

Unapaswa kutumia opereta ya `===`

```js
assert(__helpers.removeJSComments(code).match(/(val\s*===\s*\d+)|(\d+\s*===\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrict(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testStrict(10);
```

# --solutions--

```js
function testStrict(val) {
  if (val === 7) {
    return "Equal";
  }
  return "Not Equal";
}
```
