---
id: 56533eb9ac21ba0edf2244a8
title: Kuhifadhi Thamani kwa Opereta ya Mgawo
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

Katika JavaScript, unaweza kuhifadhi thamani katika kigezo kwa kutumia opereta ya <dfn>mgawo wa thamani</dfn> (`=`).

```js
myVariable = 5;
```

Hii inaweka thamani ya `Number` kwa `5` katika `myVariable`.

Kama kuna mahesabu yoyote upande wa kulia wa opereta ya `=`, hayo hufanywa kabla thamani kuwekwa katika kigezo upande wa kushoto wa opereta.

```js
var myVar;
myVar = 5;
```

Kwanza, msimbo huu unaunda kigezo kinachoitwa `myVar`. Kisha, msimbo unaweka `5` kwa `myVar`. Sasa, kama `myVar` itaonekana tena katika msimbo, programu itaitambua kana kwamba ni `5`.

# --instructions--

Weka thamani `7` kwa kigezo `a`.

# --hints--

Haupaswi kubadilisha msimbo ulio juu ya maoni yaliyotajwa.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a` inapaswa kuwa na thamani ya 7.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
