---
id: 56533eb9ac21ba0edf2244d5
title: Ulinganishi na opereta ya kubwa kuliko au sawa na
challengeType: 1
forumTopicId: 16785
dashedName: comparison-with-the-greater-than-or-equal-to-operator
---

# --description--

Opereta ya kubwa kuliko au sawa na (`>=`) inalinganisha thamani za nambari mbili. Ikiwa nambari upande wa kushoto ni kubwa kuliko au sawa na nambari upande wa kulia, inarudisha `true`. Vinginevyo, inarudisha `false`.

Kama opereta ya usawa, opereta ya kubwa kuliko au sawa na itabadilisha aina za data wakati wa kulinganisha.

**Mifano**

```js
6   >=  6  // true
7   >= '3' // true
2   >=  3  // false
'7' >=  9  // false
```

# --instructions--

Ongeza opereta ya kubwa kuliko au sawa na kwenye mistari iliyotajwa ili taarifa za return ziwe na maana.

# --hints--

`testGreaterOrEqual(0)` inapaswa kurudisha mfuatano wa herufi `Less than 10`

```js
assert(testGreaterOrEqual(0) === 'Less than 10');
```

`testGreaterOrEqual(9)` inapaswa kurudisha mfuatano wa herufi `Less than 10`

```js
assert(testGreaterOrEqual(9) === 'Less than 10');
```

`testGreaterOrEqual(10)` inapaswa kurudisha mfuatano wa herufi `10 or Over`

```js
assert(testGreaterOrEqual(10) === '10 or Over');
```

`testGreaterOrEqual(11)` inapaswa kurudisha mfuatano wa herufi `10 or Over`

```js
assert(testGreaterOrEqual(11) === '10 or Over');
```

`testGreaterOrEqual(19)` inapaswa kurudisha mfuatano wa herufi `10 or Over`

```js
assert(testGreaterOrEqual(19) === '10 or Over');
```

`testGreaterOrEqual(100)` inapaswa kurudisha mfuatano wa herufi `20 or Over`

```js
assert(testGreaterOrEqual(100) === '20 or Over');
```

`testGreaterOrEqual(21)` inapaswa kurudisha mfuatano wa herufi `20 or Over`

```js
assert(testGreaterOrEqual(21) === '20 or Over');
```

Unapaswa kutumia opereta `>=` angalau mara mbili

```js
assert(__helpers.removeJSComments(code).match(/val\s*>=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterOrEqual(val) {
  if (val) {  // Change this line
    return "20 or Over";
  }

  if (val) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}

testGreaterOrEqual(10);
```

# --solutions--

```js
function testGreaterOrEqual(val) {
  if (val >= 20) {  // Change this line
    return "20 or Over";
  }

  if (val >= 10) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}
```
