---
id: 56533eb9ac21ba0edf2244c4
title: Mfano wa kurudi mapema kwa vitendakazi
challengeType: 1
forumTopicId: 18272
dashedName: return-early-pattern-for-functions
---

# --description--

Wakati tamko la `return` linapofikiwa, utekelezaji wa kitendakazi kinachoendelea unasimama na udhibiti unarudi mahali panapoitwa.

**Mfano**

```js
function myFun() {
  console.log("Hello");
  return "World";
  console.log("byebye")
}
myFun();
```

Hili litaonyesha mfuatano wa herufi `Hello` kwenye konsoli, na kurudisha mfuatano wa herufi `World`. Mfuatano wa herufi `byebye` hautaonyeshwa kamwe kwenye konsoli, kwa sababu kitendakazi kinatoka kwenye tamko la `return`.

# --instructions--

Badilisha kitendakazi `abTest` ili ikiwa `a` au `b` ni ndogo kuliko `0` kitendakazi kitoke mara moja na thamani ya `undefined`.

**Kibukisho**  
Kumbuka kwamba <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-javascript/understanding-uninitialized-variables" target="_blank" rel="noopener noreferrer nofollow"><code>undefined</code> ni neno kuu</a>, si mfuatano wa herufi.

# --hints--

`abTest(2, 2)` inapaswa kurudisha nambari

```js
assert(typeof abTest(2, 2) === 'number');
```

`abTest(2, 2)` inapaswa kurudisha `8`

```js
assert(abTest(2, 2) === 8);
```

`abTest(-2, 2)` inapaswa kurudisha `undefined`

```js
assert(abTest(-2, 2) === undefined);
```

`abTest(2, -2)` inapaswa kurudisha `undefined`

```js
assert(abTest(2, -2) === undefined);
```

`abTest(2, 8)` inapaswa kurudisha `18`

```js
assert(abTest(2, 8) === 18);
```

`abTest(3, 3)` inapaswa kurudisha `12`

```js
assert(abTest(3, 3) === 12);
```

`abTest(0, 0)` inapaswa kurudisha `0`

```js
assert(abTest(0, 0) === 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function abTest(a, b) {
  // Only change code below this line



  // Only change code above this line

  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}

abTest(2,2);
```

# --solutions--

```js
function abTest(a, b) {
  if(a < 0 || b < 0) {
    return undefined;
  }
  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}
```
