---
id: 56533eb9ac21ba0edf2244c6
title: Simama katika mstari
challengeType: 1
forumTopicId: 18307
dashedName: stand-in-line
---

# --description--

Katika Sayansi ya Kompyuta <dfn>queue</dfn> ni <dfn>muundo wa data</dfn> wa dhahania ambapo vitu huhifadhiwa kwa mpangilio. Vitu vipya vinaweza kuongezwa mwishoni mwa queue na vitu vya zamani huondolewa kutoka mwanzo wa queue.

# --instructions--

Andika kitendakazi `nextInLine` kinachopokea safu ya data (`arr`) na nambari (`item`) kama hoja.

Ongeza nambari mwishoni mwa safu ya data, kisha ondoa kipengele cha kwanza cha safu hiyo.

Kitendakazi `nextInLine` kinapaswa kurudisha kipengele kilichotolewa.

# --hints--

`nextInLine([], 5)` inapaswa kurudisha nambari.

```js
assert.isNumber(nextInLine([], 5));
```

`nextInLine([], 1)` inapaswa kurudisha `1`

```js
assert(nextInLine([], 1) === 1);
```

`nextInLine([2], 1)` inapaswa kurudisha `2`

```js
assert(nextInLine([2], 1) === 2);
```

`nextInLine([5,6,7,8,9], 1)` inapaswa kurudisha `5`

```js
assert(nextInLine([5, 6, 7, 8, 9], 1) === 5);
```

# --seed--

## --seed-contents--

```js
function nextInLine(arr, item) {
  // Only change code below this line
  
  return item;
  // Only change code above this line
}

// Setup
let testArr = [1, 2, 3, 4, 5];

// Display code
console.log("Before: " + JSON.stringify(testArr));
console.log(nextInLine(testArr, 6));
console.log("After: " + JSON.stringify(testArr));
```

# --solutions--

```js
let testArr = [1, 2, 3, 4, 5];

function nextInLine(arr, item) {
    arr.push(item);
    return arr.shift();
}
```
