---
id: 56533eb9ac21ba0edf2244bc
title: Orodha ya manunuzi
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

Tengeneza orodha ya manunuzi katika variable `myList`. Orodha hiyo inapaswa kuwa safu ya data yenye vipengele vingi (multi-dimensional) vinavyobeba safu ndogo kadhaa.

Kipengele cha kwanza katika kila safu ndogo kinapaswa kuwa mfuatano wa herufi unaowakilisha jina la bidhaa. Kipengele cha pili kinapaswa kuwa nambari inayowakilisha kiasi, kwa mfano

```js
["Chocolate Bar", 15]
```

Inapaswa kuwepo angalau safu ndogo 5 katika orodha hiyo.

# --hints--

`myList` inapaswa kuwa safu ya data.

```js
assert(Array.isArray(myList));
```

Vipengele vya kwanza katika kila safu ndogo vinapaswa kuwa mfuatano wa herufi.

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

Vipengele vya pili katika kila safu ndogo vinapaswa kuwa nambari.

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

Unapaswa kuwa na angalau bidhaa 5 katika orodha yako.

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
