---
id: 56533eb9ac21ba0edf2244ae
title: Kupata Mabaki katika JavaScript
challengeType: 1
forumTopicId: 18184
dashedName: finding-a-remainder-in-javascript
---

# --description--

Kipengele cha <dfn>mabaki</dfn> `%` hutoa mabaki ya mgawanyo wa nambari mbili.

**Mfano**

<pre>
5 % 2 = 1
5 / 2 = 2 mabaki 1
2 * 2 = 4
5 - 4 = 1
</pre>

**Matumizi**  
Katika hesabu, nambari inaweza kukaguliwa kama ni nambari jamaa au nambari tofauti kwa kukagua mabaki ya mgawanyo wa nambari hiyo kwa `2`. Nambari jamaa zina mabaki ya `0`, wakati nambari tofauti zina mabaki ya `1`.

<pre>
17 % 2 = 1
48 % 2 = 0
</pre>

**Kumbuka:** Kipengele cha <dfn>mabaki</dfn> mara nyingine huitwa vibaya kama kipengele cha moduli. Ni sawa sana na moduli, lakini hakifanyi kazi vizuri na nambari hasi.

# --instructions--

Weka `remainder` sawa na mabaki ya `11` kugawanywa na `3` kwa kutumia kipengele cha <dfn>mabaki</dfn> (`%`).

# --hints--

Kigezo `remainder` kinapaswa kuanzishwa

```js
assert(/(const|let|var)\s+?remainder/.test(__helpers.removeJSComments(code)));
```

Thamani ya `remainder` inapaswa kuwa `2`

```js
assert(remainder === 2);
```

Unapaswa kutumia kipengele cha `%`

```js
assert(/\s+?remainder\s*?=\s*?.*%.*;?/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const remainder = 0;
```

# --solutions--

```js
const remainder = 11 % 3;
```
