---
id: 56533eb9ac21ba0edf2244cd
title: Kupata safu za data zilizopangwa ndani
challengeType: 1
forumTopicId: 16160
dashedName: accessing-nested-arrays
---

# --description--

Kama tulivyoona katika mifano ya awali, vitu vinaweza kuwa na vitu vilivyopangwa ndani na safu za data zilizopangwa ndani. Kama vile unavyopata vitu vilivyopangwa ndani, uandishi wa mabano wa safu ya data unaweza kuunganishwa kupata safu za data zilizopangwa ndani.

Hapa kuna mfano wa jinsi ya kupata safu ya data iliyopangwa ndani:

```js
const ourPets = [
  {
    animalType: "cat",
    names: [
      "Meowzer",
      "Fluffy",
      "Kit-Cat"
    ]
  },
  {
    animalType: "dog",
    names: [
      "Spot",
      "Bowser",
      "Frankie"
    ]
  }
];

ourPets[0].names[1];
ourPets[1].names[0];
```

`ourPets[0].names[1]` itakuwa mfuatano wa herufi `Fluffy`, na `ourPets[1].names[0]` itakuwa mfuatano wa herufi `Spot`.

# --instructions--

Ukitegemea uandishi wa nukta na uandishi wa mabano, weka thamani ya `secondTree` kuwa kipengele cha pili katika safu ya data ya `list` kutoka kwa kitu cha pili katika safu ya data ya `myPlants`.

# --hints--

`secondTree` inapaswa kuwa sawa na mfuatano wa herufi `pine`.

```js
assert(secondTree === 'pine');
```

Msimbo wako unapaswa kutumia uandishi wa nukta na uandishi wa mabano kupata `myPlants`.

```js
assert(/=\s*myPlants\[1\].list\[1\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = "";
```

# --solutions--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = myPlants[1].list[1];
```
