---
id: 56533eb9ac21ba0edf2244c9
title: Kupata sifa za kitu kwa kutumia vigezo
challengeType: 1
forumTopicId: 16165
dashedName: accessing-object-properties-with-variables
---

# --description--

Matumizi mengine ya uandishi wa mabano kwenye vitu ni kupata sifa ambayo imehifadhiwa kama thamani ya kigezo. Hii inaweza kuwa muhimu sana kwa kurudia kupitia sifa za kitu au wakati wa kupata thamani kutoka kwa jedwali la kutafuta.

Hapa kuna mfano wa kutumia kigezo kupata sifa:

```js
const dogs = {
  Fido: "Mutt",
  Hunter: "Doberman",
  Snoopie: "Beagle"
};

const myDog = "Hunter";
const myBreed = dogs[myDog];
console.log(myBreed);
```

Mfuatano wa herufi `Doberman` utaonyeshwa kwenye konsoli.

Kumbuka kwamba hatutumi *nakala* kuzunguka jina la kigezo tunapotumia kupata sifa kwa sababu tunatumia *thamani* ya kigezo, si *jina*.

# --instructions--

Weka thamani ya `playerNumber` kwenye kigezo `16`. Kisha, tumia kigezo hicho kupata jina la mchezaji na uweke kwenye `player`.

# --hints--

`playerNumber` inapaswa kuwa nambari

```js
assert(typeof playerNumber === 'number');
```

Kigezo `player` kinapaswa kuwa mfuatano wa herufi

```js
assert(typeof player === 'string');
```

Thamani ya `player` inapaswa kuwa mfuatano wa herufi `Montana`

```js
assert(player === 'Montana');
```

Unapaswa kutumia uandishi wa mabano kupata `testObj`

```js
assert(/testObj\s*?\[.*?\]/.test(__helpers.removeJSComments(code)));
```

Hupaswi kuweka thamani ya `Montana` moja kwa moja kwenye kigezo `player`

```js
assert(!__helpers.removeJSComments(code).match(/player\s*=\s*"|\'\s*Montana\s*"|\'\s*;/gi));
```

Unapaswa kutumia kigezo `playerNumber` katika uandishi wako wa mabano

```js
assert(/testObj\s*?\[\s*playerNumber\s*\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};

// Only change code below this line
const playerNumber = 42;  // Change this line
const player = testObj;   // Change this line
```

# --solutions--

```js
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};
const playerNumber = 16;
const player = testObj[playerNumber];
```
