---
id: 56533eb9ac21ba0edf2244ab
title: Kuelewa Uwezo wa Tofauti ya Herufi Ndogo na Kubwa katika Vigezo
challengeType: 1
forumTopicId: 18334
dashedName: understanding-case-sensitivity-in-variables
---

# --description--

Katika JavaScript vigezo vyote na majina ya vitendakazi vinaweza kutofautiana kwa herufi ndogo na kubwa. Hii ina maana kwamba matumizi ya herufi kubwa na ndogo ni muhimu.

`MYVAR` si sawa na `MyVar` wala `myvar`. Inawezekana kuwa na vigezo vingi tofauti vyenye jina moja lakini tofauti kwa herufi. Inapendekezwa sana kwa ajili ya uwazi, usitumie kipengele hiki cha lugha.

**Mbinu bora**

Andika majina ya vigezo katika JavaScript kwa kutumia <dfn>camelCase</dfn>. Katika <dfn>camelCase</dfn>, majina ya vigezo yenye maneno mengi huanza na neno la kwanza kwa herufi ndogo na herufi ya kwanza ya kila neno linalofuata huwa kubwa.

**Mifano:**

```js
var someVariable;
var anotherVariableName;
var thisVariableNameIsSoLong;
```

# --instructions--

Badilisha tamko na mgawo wa thamani uliopo ili majina yao yatumie <dfn>camelCase</dfn>.

Usitengeneze vigezo vipya vyovyote.

# --hints--

`studlyCapVar` inapaswa kutamkwa na kuwa na thamani ya `10`.

```js
assert(typeof studlyCapVar !== 'undefined' && studlyCapVar === 10);
```

`properCamelCase` inapaswa kutamkwa na kuwa na thamani ya mfuatano wa herufi `A String`.

```js
assert(
  typeof properCamelCase !== 'undefined' && properCamelCase === 'A String'
);
```

`titleCaseOver` inapaswa kutamkwa na kuwa na thamani ya `9000`.

```js
assert(typeof titleCaseOver !== 'undefined' && titleCaseOver === 9000);
```

`studlyCapVar` inapaswa kutumia camelCase katika tamko na mgawo wa thamani.

```js
assert(__helpers.removeJSComments(code).match(/studlyCapVar/g).length === 2);
```

`properCamelCase` inapaswa kutumia camelCase katika tamko na mgawo wa thamani.

```js
assert(__helpers.removeJSComments(code).match(/properCamelCase/g).length === 2);
```

`titleCaseOver` inapaswa kutumia camelCase katika tamko na mgawo wa thamani.

```js
assert(__helpers.removeJSComments(code).match(/titleCaseOver/g).length === 2);
```

# --seed--

## --seed-contents--

```js
// Variable declarations
var StUdLyCapVaR;
var properCamelCase;
var TitleCaseOver;

// Variable assignments
STUDLYCAPVAR = 10;
PRoperCAmelCAse = "A String";
tITLEcASEoVER = 9000;
```

# --solutions--

```js
var studlyCapVar;
var properCamelCase;
var titleCaseOver;

studlyCapVar = 10;
properCamelCase = "A String";
titleCaseOver = 9000;
```
