---
id: 56533eb9ac21ba0edf2244ca
title: Kutumia vitu kwa ajili ya kutafuta thamani
challengeType: 1
forumTopicId: 18373
dashedName: using-objects-for-lookups
---

# --description--

Vitu vinaweza kufikiriwa kama hifadhi ya funguo/thamani, kama kamusi. Ikiwa una data ya jedwali, unaweza kutumia kitu kutafuta thamani badala ya taarifa ya `switch` au mnyororo wa `if/else`. Hii ni muhimu zaidi wakati unajua kuwa data yako ya ingizo ni mdogo kwa anuwai fulani.

Hapa kuna mfano wa kitu cha makala:

```js
const article = {
  "title": "How to create objects in JavaScript",
  "link": "https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/",
  "author": "Kaashan Hussain",
  "language": "JavaScript",
  "tags": "TECHNOLOGY",
  "createdAt": "NOVEMBER 28, 2018"
};

const articleAuthor = article["author"];
const articleLink = article["link"];

const value = "title";
const valueLookup = article[value];
```

`articleAuthor` ni mfuatano wa herufi `Kaashan Hussain`, `articleLink` ni mfuatano wa herufi `https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/`, na `valueLookup` ni mfuatano wa herufi `How to create objects in JavaScript`.

# --instructions--

Badilisha taarifa ya switch kuwa kitu kinachoitwa `lookup`. Kitumie kutafuta `val` na weka mfuatano wa herufi unaohusiana kwenye kigezo cha `result`.

# --hints--

`phoneticLookup("alpha")` inapaswa kuwa sawa na mfuatano wa herufi `Adams`

```js
assert(phoneticLookup('alpha') === 'Adams');
```

`phoneticLookup("bravo")` inapaswa kuwa sawa na mfuatano wa herufi `Boston`

```js
assert(phoneticLookup('bravo') === 'Boston');
```

`phoneticLookup("charlie")` inapaswa kuwa sawa na mfuatano wa herufi `Chicago`

```js
assert(phoneticLookup('charlie') === 'Chicago');
```

`phoneticLookup("delta")` inapaswa kuwa sawa na mfuatano wa herufi `Denver`

```js
assert(phoneticLookup('delta') === 'Denver');
```

`phoneticLookup("echo")` inapaswa kuwa sawa na mfuatano wa herufi `Easy`

```js
assert(phoneticLookup('echo') === 'Easy');
```

`phoneticLookup("foxtrot")` inapaswa kuwa sawa na mfuatano wa herufi `Frank`

```js
assert(phoneticLookup('foxtrot') === 'Frank');
```

`phoneticLookup("")` inapaswa kuwa sawa na `undefined`

```js
assert(typeof phoneticLookup('') === 'undefined');
```

Haupaswi kubadilisha taarifa ya `return`

```js
assert(__helpers.removeJSComments(code).match(/return\sresult;/));
```

Haupaswi kutumia taarifa za `case`, `switch`, au `if`

```js
assert(
  !/case|switch|if/g.test(__helpers.removeJSComments(code).replace(/([/]{2}.*)|([/][*][^/*]*[*][/])/g, ''))
);
```

# --seed--

## --seed-contents--

```js
// Setup
function phoneticLookup(val) {
  let result = "";

  // Only change code below this line
  switch(val) {
    case "alpha":
      result = "Adams";
      break;
    case "bravo":
      result = "Boston";
      break;
    case "charlie":
      result = "Chicago";
      break;
    case "delta":
      result = "Denver";
      break;
    case "echo":
      result = "Easy";
      break;
    case "foxtrot":
      result = "Frank";
  }

  // Only change code above this line
  return result;
}

phoneticLookup("charlie");
```

# --solutions--

```js
function phoneticLookup(val) {
  let result = "";

  const lookup = {
    alpha: "Adams",
    bravo: "Boston",
    charlie: "Chicago",
    delta: "Denver",
    echo: "Easy",
    foxtrot: "Frank"
  };

  result = lookup[val];

  return result;
}
```
