---
id: 56533eb9ac21ba0edf2244b0
title: Mgawo mchanganyiko kwa kutumia kutoa thamani iliyoongezwa
challengeType: 1
forumTopicId: 16660
dashedName: compound-assignment-with-augmented-subtraction
---

# --description--

Kama operator `+=`, `-=` hutenganisha nambari kutoka kwa kigezo.

```js
myVar = myVar - 5;
```

itatenganisha `5` kutoka `myVar`. Hii inaweza kuandikwa upya kama:

```js
myVar -= 5;
```

# --instructions--

Badilisha migawo wa thamani ya `a`, `b`, na `c` kutumia operator `-=`.

# --hints--

`a` inapaswa kuwa sawa na `5`.

```js
assert(a === 5);
```

`b` inapaswa kuwa sawa na `-6`.

```js
assert(b === -6);
```

`c` inapaswa kuwa sawa na `2`.

```js
assert(c === 2);
```

Unapaswa kutumia operator `-=` kwa kila kigezo.

```js
assert(__helpers.removeJSComments(code).match(/-=/g).length === 3);
```

Hupaswi kubadilisha msimbo ulio juu ya maoni yaliyotajwa.

```js
assert(
  /let a = 11;/.test(__helpers.removeJSComments(code)) && /let b = 9;/.test(__helpers.removeJSComments(code)) && /let c = 3;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 11;
let b = 9;
let c = 3;

// Only change code below this line
a = a - 6;
b = b - 15;
c = c - 1;
```

# --solutions--

```js
let a = 11;
let b = 9;
let c = 3;

a -= 6;
b -= 15;
c -= 1;
```
