---
id: 56533eb9ac21ba0edf2244af
title: Mgawo mchanganyiko na kuongeza kwa njia ya kuongezwa
challengeType: 1
forumTopicId: 16661
dashedName: compound-assignment-with-augmented-addition
---

# --description--

Katika programu, ni kawaida kutumia migawo ya thamani kubadilisha yaliyomo katika kigezo. Kumbuka kuwa kila kitu kilicho upande wa kulia wa alama ya usawa kinahesabiwa kwanza, kwa hivyo tunaweza kusema:

```js
myVar = myVar + 5;
```

kuongeza `5` kwa `myVar`. Kwa kuwa hii ni mtindo wa kawaida sana, kuna vionyeshi vinavyofanya operesheni ya kihesabu na mgawo wa thamani kwa hatua moja.

Moja ya vionyeshi hivyo ni kionyeshi cha `+=`.

```js
let myVar = 1;
myVar += 5;
console.log(myVar);
```

`6` itaonyeshwa kwenye konsoli.

# --instructions--

Badilisha migawo ya thamani ya `a`, `b`, na `c` ili itumie kionyeshi cha `+=`.

# --hints--

`a` inapaswa kuwa sawa na `15`.

```js
assert(a === 15);
```

`b` inapaswa kuwa sawa na `26`.

```js
assert(b === 26);
```

`c` inapaswa kuwa sawa na `19`.

```js
assert(c === 19);
```

Unapaswa kutumia kionyeshi cha `+=` kwa kila kigezo.

```js
assert(__helpers.removeJSComments(code).match(/\+=/g).length === 3);
```

Hupaswi kubadilisha msimbo uliopo juu ya maoni yaliyotajwa.

```js
assert(
  /let a = 3;/.test(__helpers.removeJSComments(code)) &&
    /let b = 17;/.test(__helpers.removeJSComments(code)) &&
    /let c = 12;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 3;
let b = 17;
let c = 12;

// Only change code below this line
a = a + 12;
b = 9 + b;
c = c + 7;
```

# --solutions--

```js
let a = 3;
let b = 17;
let c = 12;

a += 12;
b += 9;
c += 7;
```
