---
id: 56533eb9ac21ba0edf2244d4
title: Ulinganisho na opereta ya kubwa kuliko
challengeType: 1
forumTopicId: 16786
dashedName: comparison-with-the-greater-than-operator
---

# --description--

Opereta ya kubwa kuliko (`>`) inalinganisha thamani za nambari mbili. Ikiwa nambari upande wa kushoto ni kubwa kuliko nambari upande wa kulia, inarudisha `true`. Vinginevyo, inarudisha `false`.

Kama opereta ya usawa, opereta ya kubwa kuliko itabadilisha aina za data za thamani wakati wa kulinganisha.

**Mifano**

```js
5   >  3  // true
7   > '3' // true
2   >  3  // false
'1' >  9  // false
```

# --instructions--

Ongeza opereta ya kubwa kuliko kwenye mistari iliyoonyeshwa ili taarifa za return ziwe na maana.

# --hints--

`testGreaterThan(0)` inapaswa kurudisha mfuatano wa herufi `10 or Under`

```js
assert(testGreaterThan(0) === '10 or Under');
```

`testGreaterThan(10)` inapaswa kurudisha mfuatano wa herufi `10 or Under`

```js
assert(testGreaterThan(10) === '10 or Under');
```

`testGreaterThan(11)` inapaswa kurudisha mfuatano wa herufi `Over 10`

```js
assert(testGreaterThan(11) === 'Over 10');
```

`testGreaterThan(99)` inapaswa kurudisha mfuatano wa herufi `Over 10`

```js
assert(testGreaterThan(99) === 'Over 10');
```

`testGreaterThan(100)` inapaswa kurudisha mfuatano wa herufi `Over 10`

```js
assert(testGreaterThan(100) === 'Over 10');
```

`testGreaterThan(101)` inapaswa kurudisha mfuatano wa herufi `Over 100`

```js
assert(testGreaterThan(101) === 'Over 100');
```

`testGreaterThan(150)` inapaswa kurudisha mfuatano wa herufi `Over 100`

```js
assert(testGreaterThan(150) === 'Over 100');
```

Unapaswa kutumia opereta `>` angalau mara mbili

```js
assert(__helpers.removeJSComments(code).match(/val\s*>\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterThan(val) {
  if (val) {  // Change this line
    return "Over 100";
  }

  if (val) {  // Change this line
    return "Over 10";
  }

  return "10 or Under";
}

testGreaterThan(10);
```

# --solutions--

```js
function testGreaterThan(val) {
  if (val > 100) {  // Change this line
    return "Over 100";
  }
  if (val > 10) {  // Change this line
    return "Over 10";
  }
  return "10 or Under";
}
```
