---
id: 56533eb9ac21ba0edf2244b2
title: Mgawo mchanganyiko na mgawo ulioboreshwa
challengeType: 1
forumTopicId: 16659
dashedName: compound-assignment-with-augmented-division
---

# --description--

Operator `/=` hugawanya kigezo kwa nambari nyingine.

```js
myVar = myVar / 5;
```

Itagawanya `myVar` kwa `5`. Hii inaweza kuandikwa upya kama:

```js
myVar /= 5;
```

# --instructions--

Badilisha migawo wa thamani ya `a`, `b`, na `c` ili kutumia operator `/=`.

# --hints--

`a` inapaswa kuwa sawa na `4`.

```js
assert(a === 4);
```

`b` inapaswa kuwa sawa na `27`.

```js
assert(b === 27);
```

`c` inapaswa kuwa sawa na `3`.

```js
assert(c === 3);
```

Unapaswa kutumia operator `/=` kwa kila kigezo.

```js
assert(__helpers.removeJSComments(code).match(/\/=/g).length === 3);
```

Hupaswi kubadilisha msimbo uliopo juu ya maoni yaliyotajwa.

```js
assert(
  /let a = 48;/.test(__helpers.removeJSComments(code)) &&
    /let b = 108;/.test(__helpers.removeJSComments(code)) &&
    /let c = 33;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 48;
let b = 108;
let c = 33;

// Only change code below this line
a = a / 12;
b = b / 4;
c = c / 11;
```

# --solutions--

```js
let a = 48;
let b = 108;
let c = 33;

a /= 12;
b /= 4;
c /= 11;
```
