---
id: 56533eb9ac21ba0edf2244c7
title: Kupata sifa za kitu kwa kutumia uandishi wa nukta
challengeType: 1
forumTopicId: 16164
dashedName: accessing-object-properties-with-dot-notation
---

# --description--

Kuna njia mbili za kupata sifa za kitu: uandishi wa nukta (`.`) na uandishi wa mabano (`[]`), kama safu ya data.

Uandishi wa nukta ndilo unalotumia unapojua jina la sifa unayotaka kupata kabla.

Hapa kuna mfano wa kutumia uandishi wa nukta (`.`) kusoma sifa ya kitu:

```js
const myObj = {
  prop1: "val1",
  prop2: "val2"
};

const prop1val = myObj.prop1;
const prop2val = myObj.prop2;
```

`prop1val` itakuwa na thamani ya mfuatano wa herufi `val1`, na `prop2val` itakuwa na thamani ya mfuatano wa herufi `val2`.

# --instructions--

Soma thamani za sifa za `testObj` kwa kutumia uandishi wa nukta. Weka thamani ya vigezo `hatValue` sawa na sifa ya kitu `hat` na weka thamani ya vigezo `shirtValue` sawa na sifa ya kitu `shirt`.

# --hints--

`hatValue` inapaswa kuwa mfuatano wa herufi

```js
assert(typeof hatValue === 'string');
```

Thamani ya `hatValue` inapaswa kuwa mfuatano wa herufi `ballcap`

```js
assert(hatValue === 'ballcap');
```

`shirtValue` inapaswa kuwa mfuatano wa herufi

```js
assert(typeof shirtValue === 'string');
```

Thamani ya `shirtValue` inapaswa kuwa mfuatano wa herufi `jersey`

```js
assert(shirtValue === 'jersey');
```

Unapaswa kutumia uandishi wa nukta mara mbili

```js
assert(__helpers.removeJSComments(code).match(/testObj\.\w+/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

// Only change code below this line
const hatValue = testObj;      // Change this line
const shirtValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

const hatValue = testObj.hat;
const shirtValue = testObj.shirt;
```
