---
id: 56533eb9ac21ba0edf2244d3
title: Ulinganishi na opereta ya utofauti mkali
challengeType: 1
forumTopicId: 16791
dashedName: comparison-with-the-strict-inequality-operator
---

# --description--

Opereta ya utofauti mkali (`!==`) ni kinyume cha mantiki cha opereta ya usawa mkali. Inamaanisha "Sio sawa kabisa" na hurudisha `false` pale opereta ya usawa mkali itakaporudisha `true` na *kivyo hivyo*. Opereta ya utofauti mkali haitabadilisha aina za data.

**Mifano**

```js
3 !==  3  // false
3 !== '3' // true
4 !==  3  // true
```

# --instructions--

Ongeza opereta ya utofauti mkali kwenye usemi wa `if` ili kitendakazi kirudishe mfuatano wa herufi `Not Equal` wakati `val` si sawa kabisa na `17`

# --hints--

`testStrictNotEqual(17)` inapaswa kurudisha mfuatano wa herufi `Equal`

```js
assert(testStrictNotEqual(17) === 'Equal');
```

`testStrictNotEqual("17")` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(testStrictNotEqual('17') === 'Not Equal');
```

`testStrictNotEqual(12)` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(testStrictNotEqual(12) === 'Not Equal');
```

`testStrictNotEqual("bob")` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(testStrictNotEqual('bob') === 'Not Equal');
```

Unapaswa kutumia opereta ya `!==`

```js
assert(__helpers.removeJSComments(code).match(/(val\s*!==\s*\d+)|(\d+\s*!==\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrictNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testStrictNotEqual(10);
```

# --solutions--

```js
function testStrictNotEqual(val) {
  if (val !== 17) {
    return "Not Equal";
  }
  return "Equal";
}
```
