---
id: 56533eb9ac21ba0edf2244e1
title: Kupanga ndani kwa kutumia mizunguko
challengeType: 1
forumTopicId: 18248
dashedName: nesting-for-loops
---

# --description--

Ikiwa una safu ya data yenye vipimo vingi, unaweza kutumia mantiki ile ile kama sehemu ya awali kupita kupitia safu ya data na pia safu ndogo yoyote. Hapa kuna mfano:

```js
const arr = [
  [1, 2], [3, 4], [5, 6]
];

for (let i = 0; i < arr.length; i++) {
  for (let j = 0; j < arr[i].length; j++) {
    console.log(arr[i][j]);
  }
}
```

Hii inatoa kila kipengele kidogo katika `arr` moja kwa moja kwa wakati. Kumbuka kuwa kwa mzunguko wa ndani, tunakagua `.length` ya `arr[i]`, kwa sababu `arr[i]` ni safu ya data yenyewe.

# --instructions--

Badilisha kitendakazi `multiplyAll` ili kirudishe bidhaa ya nambari zote katika safu ndogo za `arr`.

# --hints--

`multiplyAll([[1], [2], [3]])` inapaswa kurudisha `6`

```js
assert(multiplyAll([[1], [2], [3]]) === 6);
```

`multiplyAll([[1, 2], [3, 4], [5, 6, 7]])` inapaswa kurudisha `5040`

```js
assert(
  multiplyAll([
    [1, 2],
    [3, 4],
    [5, 6, 7]
  ]) === 5040
);
```

`multiplyAll([[5, 1], [0.2, 4, 0.5], [3, 9]])` inapaswa kurudisha `54`

```js
assert(
  multiplyAll([
    [5, 1],
    [0.2, 4, 0.5],
    [3, 9]
  ]) === 54
);
```

# --seed--

## --seed-contents--

```js
function multiplyAll(arr) {
  let product = 1;
  // Only change code below this line

  // Only change code above this line
  return product;
}

multiplyAll([[1, 2], [3, 4], [5, 6, 7]]);
```

# --solutions--

```js
function multiplyAll(arr) {
  let product = 1;
  for (let i = 0; i < arr.length; i++) {
    for (let j = 0; j < arr[i].length; j++) {
      product *= arr[i][j];
    }
  }
  return product;
}
```
