---
id: 56533eb9ac21ba0edf2244dd
title: Kuchagua kutoka kwa chaguzi nyingi kwa kutumia taarifa za kubadilisha
challengeType: 1
forumTopicId: 18277
dashedName: selecting-from-many-options-with-switch-statements
---

# --description--

Ikiwa unahitaji kulinganisha thamani moja dhidi ya chaguzi nyingi, unaweza kutumia taarifa ya <dfn>kubadilisha</dfn>. Taarifa ya `switch` inalinganisha thamani na taarifa za <dfn>kesi</dfn> ambazo zinaeleza thamani mbalimbali zinazowezekana. Taarifa zozote halali za JavaScript zinaweza kutekelezwa ndani ya sehemu ya <dfn>kesi</dfn> na zitafanya kazi kuanzia thamani ya kwanza iliyolingana ya `case` hadi `break` itakapokutana.

Hapa kuna mfano wa taarifa ya `switch`:

```js
switch (fruit) {
  case "apple":
    console.log("The fruit is an apple");
    break;
  case "orange":
    console.log("The fruit is an orange");
    break;
}
```

Thamani za `case` zinajaribiwa kwa usawa mkali (`===`). `break` inamwambia JavaScript kuacha kutekeleza taarifa. Ikiwa `break` itakosekana, taarifa inayofuata itatekelezwa.

# --instructions--

Andika taarifa ya kubadilisha ambayo inajaribu `val` na kuweka `answer` kwa masharti yafuatayo:  
`1` - `alpha`  
`2` - `beta`  
`3` - `gamma`  
`4` - `delta`

# --hints--

`caseInSwitch(1)` inapaswa kuwa na thamani ya mfuatano wa herufi `alpha`

```js
assert(caseInSwitch(1) === 'alpha');
```

`caseInSwitch(2)` inapaswa kuwa na thamani ya mfuatano wa herufi `beta`

```js
assert(caseInSwitch(2) === 'beta');
```

`caseInSwitch(3)` inapaswa kuwa na thamani ya mfuatano wa herufi `gamma`

```js
assert(caseInSwitch(3) === 'gamma');
```

`caseInSwitch(4)` inapaswa kuwa na thamani ya mfuatano wa herufi `delta`

```js
assert(caseInSwitch(4) === 'delta');
```

Haupaswi kutumia taarifa za `if` au `else`

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Unapaswa kuwa na angalau taarifa 3 za `break`

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function caseInSwitch(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

caseInSwitch(1);
```

# --solutions--

```js
function caseInSwitch(val) {
  let answer = "";

  switch (val) {
    case 1:
      answer = "alpha";
      break;
    case 2:
      answer = "beta";
      break;
    case 3:
      answer = "gamma";
      break;
    case 4:
      answer = "delta";
  }
  return answer;
}
```
