---
id: 56533eb9ac21ba0edf2244c0
title: Wigo wa kimataifa dhidi ya wigo wa ndani katika vitendakazi
challengeType: 1
forumTopicId: 18194
dashedName: global-vs--local-scope-in-functions
---

# --description--

Inawezekana kuwa na vigezo vya <dfn>ndani</dfn> na <dfn>kimataifa</dfn> vyenye jina moja. Unapofanya hivi, kigezo cha ndani kinapewa kipaumbele zaidi kuliko kigezo cha kimataifa.

Katika mfano huu:

```js
const someVar = "Hat";

function myFun() {
  const someVar = "Head";
  return someVar;
}
```

Kitendakazi `myFun` kitarudisha mfuatano wa herufi `Head` kwa sababu toleo la ndani la kigezo lipo.

# --instructions--

Ongeza kigezo cha ndani katika kitendakazi `myOutfit` ili ghairi thamani ya `outerWear` kwa mfuatano wa herufi `sweater`.

# --hints--

Haupasi kubadilisha thamani ya `outerWear` ya kimataifa.

```js
assert(outerWear === 'T-Shirt');
```

`myOutfit` inapaswa kurudisha mfuatano wa herufi `sweater`.

```js
assert(myOutfit() === 'sweater');
```

Haupasi kubadilisha taarifa ya return.

```js
assert(/return outerWear/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const outerWear = "T-Shirt";

function myOutfit() {
  // Only change code below this line

  // Only change code above this line
  return outerWear;
}

myOutfit();
```

# --solutions--

```js
const outerWear = "T-Shirt";
function myOutfit() {
  const outerWear = "sweater";
  return outerWear;
}
```
