---
id: 56533eb9ac21ba0edf2244ed
title: Kuongeza Vigezo kwa Mfuatano wa herufi
challengeType: 1
forumTopicId: 16656
dashedName: appending-variables-to-strings
---

# --description--

Kama tunavyoweza kujenga mfuatano wa herufi katika mistari mingi kwa kutumia <dfn>maneno halisi ya mfuatano wa herufi</dfn>, tunaweza pia kuongeza vigezo kwenye mfuatano wa herufi kwa kutumia kiendeshaji cha plus equals (`+=`).

Mfano:

```js
const anAdjective = "awesome!";
let ourStr = "freeCodeCamp is ";
ourStr += anAdjective;
```

`ourStr` ingekuwa na thamani ya `freeCodeCamp is awesome!`.

# --instructions--

Weka `someAdjective` kuwa mfuatano wa herufi wenye angalau herufi 3 na uongeze kwenye `myStr` kwa kutumia kiendeshaji cha `+=`.

# --hints--

`someAdjective` inapaswa kuwekwa kuwa mfuatano wa herufi wenye angalau herufi 3.

```js
assert(typeof someAdjective !== 'undefined' && someAdjective.length > 2);
```

Unapaswa kuongeza `someAdjective` kwenye `myStr` kwa kutumia kiendeshaji cha `+=`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*someAdjective\s*/).length > 0);
```

# --seed--

## --seed-contents--

```js
// Change code below this line
const someAdjective = "";
let myStr = "Learning to code is ";
```

# --solutions--

```js
const someAdjective = "neat";
let myStr = "Learning to code is ";
myStr += someAdjective;
```
