---
id: 56533eb9ac21ba0edf2244d6
title: Ulinganisho na opereta ya chini ya
challengeType: 1
forumTopicId: 16789
dashedName: comparison-with-the-less-than-operator
---

# --description--

Opereta ya chini ya (`<`) inalinganisha thamani za nambari mbili. Ikiwa nambari upande wa kushoto ni ndogo kuliko nambari upande wa kulia, inarudisha `true`. Vinginevyo, inarudisha `false`. Kama opereta ya usawa, opereta ya chini ya hubadilisha aina za data wakati wa kulinganisha.

**Mifano**

```js
2   < 5 // true
'3' < 7 // true
5   < 5 // false
3   < 2 // false
'8' < 4 // false
```

# --instructions--

Ongeza opereta ya chini ya kwenye mistari iliyotajwa ili taarifa za return ziwe na maana.

# --hints--

`testLessThan(0)` inapaswa kurudisha mfuatano wa herufi `Under 25`

```js
assert(testLessThan(0) === 'Under 25');
```

`testLessThan(24)` inapaswa kurudisha mfuatano wa herufi `Under 25`

```js
assert(testLessThan(24) === 'Under 25');
```

`testLessThan(25)` inapaswa kurudisha mfuatano wa herufi `Under 55`

```js
assert(testLessThan(25) === 'Under 55');
```

`testLessThan(54)` inapaswa kurudisha mfuatano wa herufi `Under 55`

```js
assert(testLessThan(54) === 'Under 55');
```

`testLessThan(55)` inapaswa kurudisha mfuatano wa herufi `55 or Over`

```js
assert(testLessThan(55) === '55 or Over');
```

`testLessThan(99)` inapaswa kurudisha mfuatano wa herufi `55 or Over`

```js
assert(testLessThan(99) === '55 or Over');
```

Unapaswa kutumia opereta ya `<` angalau mara mbili

```js
assert(__helpers.removeJSComments(code).match(/val\s*<\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessThan(val) {
  if (val) {  // Change this line
    return "Under 25";
  }

  if (val) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}

testLessThan(10);
```

# --solutions--

```js
function testLessThan(val) {
  if (val < 25) {  // Change this line
    return "Under 25";
  }

  if (val < 55) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}
```
