---
id: 56533eb9ac21ba0edf2244ad
title: Punguza nambari kwa kutumia JavaScript
challengeType: 1
forumTopicId: 17558
dashedName: decrement-a-number-with-javascript
---

# --description--

Unaweza kwa urahisi <dfn>kupunguza</dfn> au kupunguza thamani ya kigezo kwa moja kwa kutumia opereta `--`.

```js
i--;
```

ni sawa na

```js
i = i - 1;
```

**Kumbuka:** Mstari mzima unakuwa `i--;`, hivyo hauna haja ya ishara ya usawa.

# --instructions--

Badilisha msimbo kutumia opereta `--` kwenye `myVar`.

# --hints--

`myVar` inapaswa kuwa sawa na `10`.

```js
assert(myVar === 10);
```

`myVar = myVar - 1;` inapaswa kubadilishwa.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*myVar\s*[-]\s*1.*?;?/));
```

Haupaswi kuweka thamani ya `myVar` kwa `10`.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*10.*?;?/));
```

Unapaswa kutumia opereta `--` kwenye `myVar`.

```js
assert(/[-]{2}\s*myVar|myVar\s*[-]{2}/.test(__helpers.removeJSComments(code)));
```

Haupaswi kubadilisha msimbo juu ya maoni yaliyotajwa.

```js
assert(/let myVar = 11;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 11;

// Only change code below this line
myVar = myVar - 1;
```

# --solutions--

```js
let myVar = 11;
myVar--;
```
