---
id: 56592a60ddddeae28f7aa8e1
title: Upatikanaji wa safu za data za vipimo vingi kwa kutumia vielezo
challengeType: 1
forumTopicId: 16159
dashedName: access-multi-dimensional-arrays-with-indexes
---

# --description--

Njia moja ya kufikiria safu ya data ya <dfn>vipimo vingi</dfn>, ni kama *safu ya data ya safu za data*. Unapotumia mabano kupata upatikanaji wa safu yako ya data, seti ya kwanza ya mabano inahusu maingizo katika safu ya data ya ngazi ya kwanza (ya nje kabisa), na kila jozi ya ziada ya mabano inahusu ngazi inayofuata ya maingizo ndani yake.

**Mfano**

```js
const arr = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14]
];

const subarray = arr[3];
const nestedSubarray = arr[3][0];
const element = arr[3][0][1];
```

Katika mfano huu, `subarray` ina thamani ya `[[10, 11, 12], 13, 14]`,
`nestedSubarray` ina thamani ya `[10, 11, 12]`, na `element` ina thamani ya `11` .

**Kumbuka:** Hakuna nafasi zinapaswa kuwepo kati ya jina la safu ya data na mabano ya mraba, kama `array [0][0]` na hata hii `array [0] [0]` haikubaliki. Ingawa JavaScript inaweza kusindika hili kwa usahihi, hii inaweza kuwachanganya programu wengine wanaosoma msimbo wako.

# --instructions--

Kutumia uandishi wa mabano chagua kipengee kutoka `myArray` ili `myData` iwe sawa na `8`.

# --hints--

`myData` inapaswa kuwa sawa na `8`.

```js
assert(myData === 8);
```

Unapaswa kutumia uandishi wa mabano kusoma thamani sahihi kutoka `myArray`.

```js
assert(/myData=myArray\[2\]\[1\]/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --seed-contents--

```js
const myArray = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14],
];

const myData = myArray[0][0];
```

# --solutions--

```js
const myArray = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [[10, 11, 12], 13, 14]];
const myData = myArray[2][1];
```
