---
id: 5664820f61c48e80c9fa476c
title: Msimbo wa Golf
challengeType: 1
forumTopicId: 18195
dashedName: golf-code
---

# --description--

Katika mchezo wa Golf, kila shimo lina `par`, ikimaanisha, wastani wa nambari ya `strokes` ambayo mchezaji wa golf anatarajiwa kufanya ili kuingiza mpira kwenye shimo na kumaliza mchezo. Kulingana na umbali wa juu au chini ya `par` ambapo `strokes` zako ziko, kuna jina tofauti la utani.

Kitendakazi chako kitapokea hoja za `par` na `strokes`. Rudisha mfuatano wa herufi sahihi kulingana na jedwali hili linaloorodhesha michoro kwa kipaumbele; juu (kipaumbele cha juu) hadi chini (kipaumbele cha chini):

<table><thead><tr><th>Michoro</th><th>Rudisha</th></tr></thead><tbody><tr><td>1</td><td>"Hole-in-one!"</td></tr><tr><td>&#x3C;= par - 2</td><td>"Eagle"</td></tr><tr><td>par - 1</td><td>"Birdie"</td></tr><tr><td>par</td><td>"Par"</td></tr><tr><td>par + 1</td><td>"Bogey"</td></tr><tr><td>par + 2</td><td>"Double Bogey"</td></tr><tr><td>>= par + 3</td><td>"Go Home!"</td></tr></tbody></table>

`par` na `strokes` zitakuwa nambari na chanya kila wakati. Tumeongeza safu ya data ya majina yote kwa urahisi wako.

# --hints--

`golfScore(4, 1)` inapaswa kurudisha mfuatano wa herufi `Hole-in-one!`

```js
assert(golfScore(4, 1) === 'Hole-in-one!');
```

`golfScore(4, 2)` inapaswa kurudisha mfuatano wa herufi `Eagle`

```js
assert(golfScore(4, 2) === 'Eagle');
```

`golfScore(5, 2)` inapaswa kurudisha mfuatano wa herufi `Eagle`

```js
assert(golfScore(5, 2) === 'Eagle');
```

`golfScore(4, 3)` inapaswa kurudisha mfuatano wa herufi `Birdie`

```js
assert(golfScore(4, 3) === 'Birdie');
```

`golfScore(4, 4)` inapaswa kurudisha mfuatano wa herufi `Par`

```js
assert(golfScore(4, 4) === 'Par');
```

`golfScore(1, 1)` inapaswa kurudisha mfuatano wa herufi `Hole-in-one!`

```js
assert(golfScore(1, 1) === 'Hole-in-one!');
```

`golfScore(5, 5)` inapaswa kurudisha mfuatano wa herufi `Par`

```js
assert(golfScore(5, 5) === 'Par');
```

`golfScore(4, 5)` inapaswa kurudisha mfuatano wa herufi `Bogey`

```js
assert(golfScore(4, 5) === 'Bogey');
```

`golfScore(4, 6)` inapaswa kurudisha mfuatano wa herufi `Double Bogey`

```js
assert(golfScore(4, 6) === 'Double Bogey');
```

`golfScore(4, 7)` inapaswa kurudisha mfuatano wa herufi `Go Home!`

```js
assert(golfScore(4, 7) === 'Go Home!');
```

`golfScore(5, 9)` inapaswa kurudisha mfuatano wa herufi `Go Home!`

```js
assert(golfScore(5, 9) === 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

function golfScore(par, strokes) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

golfScore(5, 4);
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
