---
id: 5675e877dbd60be8ad28edc6
title: Rudia kupitia safu ya data kwa kutumia mzunguko wa for
challengeType: 1
forumTopicId: 18216
dashedName: iterate-through-an-array-with-a-for-loop
---

# --description--

Zoezi la kawaida katika JavaScript ni rudia kupitia yaliyomo katika safu ya data. Njia moja ya kufanya hivyo ni kwa kutumia mzunguko wa `for`. Msimbo huu utaonyesha kila kipengele cha safu ya data `arr` kwenye konsoli:

```js
const arr = [10, 9, 8, 7, 6];

for (let i = 0; i < arr.length; i++) {
   console.log(arr[i]);
}
```

Kumbuka kuwa safu za data zina kielezo kuanzia sufuri, ambayo ina maana kielezo cha mwisho cha safu ni `length - 1`. Masharti yetu kwa mzunguko huu ni `i < arr.length`, ambayo huzuia mzunguko wakati `i` inakuwa sawa na `length`. Katika kesi hii mzunguko wa mwisho ni `i === 4` yaani wakati `i` inakuwa sawa na `arr.length - 1` na kuonyesha `6` kwenye konsoli. Kisha `i` inaongezeka hadi `5`, na mzunguko unamalizika kwa sababu `i < arr.length` ni `false`.

# --instructions--

Tangaza na anzisha kigezo `total` kwa `0`. Tumia mzunguko wa `for` kuongeza thamani ya kila kipengele cha safu ya data `myArr` kwa `total`.

# --hints--

`total` inapaswa kutangazwa na kuanzishwa kwa 0.

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s*?total\s*=\s*0.*?;?/));
```

`total` inapaswa kuwa sawa na 20.

```js
assert(total === 20);
```

Unapaswa kutumia mzunguko wa `for` rudia kupitia `myArr`.

```js
assert(/for\s*\(/g.test(__helpers.removeJSComments(code)) && /myArr\s*\[/g.test(__helpers.removeJSComments(code)));
```

Hupaswi kujaribu kuweka thamani 20 moja kwa moja kwa `total`.

```js
assert(!__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/total[=+-]0*[1-9]+/gm));
```

# --seed--

## --seed-contents--

```js
// Setup
const myArr = [2, 3, 4, 5, 6];

// Only change code below this line

```

# --solutions--

```js
const myArr = [2, 3, 4, 5, 6];
let total = 0;

for (let i = 0; i < myArr.length; i++) {
  total += myArr[i];
}
```
