---
id: 5679ceb97cbaa8c51670a16b
title: Kurudisha Thamani za Kweli au Siyo Kweli Kutoka kwa Vitendakazi
challengeType: 1
forumTopicId: 18273
dashedName: returning-boolean-values-from-functions
---

# --description--

Unaweza kukumbuka kutoka <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/comparison-with-the-equality-operator" target="_blank" rel="noopener noreferrer nofollow">Kulinganisha kwa Opereta ya Usawa</a> kwamba opereta zote za kulinganisha hurudisha thamani ya kweli au siyo kweli `true` au `false`.

Wakati mwingine watu hutumia tamko la `if/else` kufanya kulinganisha, kama hii:

```js
function isEqual(a, b) {
  if (a === b) {
    return true;
  } else {
    return false;
  }
}
```

Lakini kuna njia bora ya kufanya hivi. Kwa kuwa `===` hurudisha `true` au `false`, tunaweza kurudisha matokeo ya kulinganisha:

```js
function isEqual(a, b) {
  return a === b;
}
```

# --instructions--

Rekebisha kitendakazi `isLess` kuondoa tamko la `if/else`.

# --hints--

`isLess(10, 15)` inapaswa kurudisha `true`

```js
assert(isLess(10, 15) === true);
```

`isLess(15, 10)` inapaswa kurudisha `false`

```js
assert(isLess(15, 10) === false);
```

Haupaswi kutumia tamko lolote la `if` au `else`

```js
assert(!/if|else/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function isLess(a, b) {
  // Only change code below this line
  if (a < b) {
    return true;
  } else {
    return false;
  }
  // Only change code above this line
}

isLess(10, 15);
```

# --solutions--

```js
function isLess(a, b) {
  return a < b;
}
```
