---
id: 567af2437cbaa8c51670a16c
title: Kupima vitu kwa sifa zao
challengeType: 1
forumTopicId: 18324
dashedName: testing-objects-for-properties
---

# --description--

Ili kuangalia kama sifa fulani ipo kwenye kitu fulani au la, unaweza kutumia njia ya `.hasOwnProperty()`. `someObject.hasOwnProperty(someProperty)` hurejesha `true` au `false` kulingana na kama sifa hiyo imetambuliwa kwenye kitu hicho au la.

**Mfano**

```js
function checkForProperty(object, property) {
  return object.hasOwnProperty(property);
}

checkForProperty({ top: 'hat', bottom: 'pants' }, 'top'); // true
checkForProperty({ top: 'hat', bottom: 'pants' }, 'middle'); // false
```

Wito wa kwanza wa kitendakazi `checkForProperty` hurejesha `true`, wakati wito wa pili hurejesha `false`.

# --instructions--

Badilisha kitendakazi `checkObj` kupima kama kitu kilichotumwa kwenye kigezo cha kitendakazi `obj` kina sifa maalum iliyotumwa kwenye kigezo cha kitendakazi `checkProp`. Ikiwa sifa iliyotumwa kwa `checkProp` inapatikana kwenye `obj`, rejesha thamani ya sifa hiyo. Ikiwa haipo, rejesha `Not Found`.

# --hints--

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "gift")` inapaswa kurejesha mfuatano wa herufi `pony`.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'gift') === 'pony'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "pet")` inapaswa kurejesha mfuatano wa herufi `kitten`.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'pet') === 'kitten'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "house")` inapaswa kurejesha mfuatano wa herufi `Not Found`.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'house') ===
    'Not Found'
);
```

`checkObj({city: "Seattle"}, "city")` inapaswa kurejesha mfuatano wa herufi `Seattle`.

```js
assert(checkObj({ city: 'Seattle' }, 'city') === 'Seattle');
```

`checkObj({city: "Seattle"}, "district")` inapaswa kurejesha mfuatano wa herufi `Not Found`.

```js
assert(checkObj({ city: 'Seattle' }, 'district') === 'Not Found');
```

`checkObj({pet: "kitten", bed: "sleigh"}, "gift")` inapaswa kurejesha mfuatano wa herufi `Not Found`.

```js
assert(checkObj({ pet: 'kitten', bed: 'sleigh' }, 'gift') === 'Not Found');
```

# --seed--

## --seed-contents--

```js
function checkObj(obj, checkProp) {
  // Only change code below this line
  return "Change Me!";
  // Only change code above this line
}
```

# --solutions--

```js
function checkObj(obj, checkProp) {
  if(obj.hasOwnProperty(checkProp)) {
    return obj[checkProp];
  } else {
    return "Not Found";
  }
}
```
