---
id: 5688e62ea601b2482ff8422b
title: Utafutaji wa wasifu
challengeType: 1
forumTopicId: 18259
dashedName: profile-lookup
---

# --description--

Tuna safu ya vitu vinavyowakilisha watu tofauti katika orodha zetu za mawasiliano.

Kitendakazi `lookUpProfile` ambacho kinapokea `name` na sifa (`prop`) kama hoja kimeandikwa tayari kwa ajili yako.

Kitendakazi kinapaswa kukagua kama `name` ni `firstName` halisi wa mtu katika mawasiliano na sifa iliyotolewa (`prop`) ni sifa ya mtu huyo.

Kama zote mbili ni kweli, basi rudisha "thamani" ya sifa hiyo.

Kama `name` hailingani na mawasiliano yoyote, basi rudisha mfuatano wa herufi `No such contact`.

Kama `prop` hailingani na sifa halali yoyote ya mtu aliyepatikana kuendana na `name`, basi rudisha mfuatano wa herufi `No such property`.

# --hints--

`lookUpProfile("Kristian", "lastName")` inapaswa kurudisha mfuatano wa herufi `Vos`

```js
assert(lookUpProfile('Kristian', 'lastName') === 'Vos');
```

`lookUpProfile("Sherlock", "likes")` inapaswa kurudisha `["Intriguing Cases", "Violin"]`

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` inapaswa kurudisha safu ya data

```js
assert(typeof lookUpProfile('Harry', 'likes') === 'object');
```

`lookUpProfile("Bob", "number")` inapaswa kurudisha mfuatano wa herufi `No such contact`

```js
assert(lookUpProfile('Bob', 'number') === 'No such contact');
```

`lookUpProfile("Bob", "potato")` inapaswa kurudisha mfuatano wa herufi `No such contact`

```js
assert(lookUpProfile('Bob', 'potato') === 'No such contact');
```

`lookUpProfile("Akira", "address")` inapaswa kurudisha mfuatano wa herufi `No such property`

```js
assert(lookUpProfile('Akira', 'address') === 'No such property');
```

# --seed--

## --seed-contents--

```js
// Setup
const contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

function lookUpProfile(name, prop) {
  // Only change code below this line

  // Only change code above this line
}

lookUpProfile("Akira", "likes");
```

# --solutions--

```js
const contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];
function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
