---
id: 5690307fddb111c6084545d7
title: Mpangilio wa mantiki katika kauli za if else
challengeType: 1
forumTopicId: 18228
dashedName: logical-order-in-if-else-statements
---

# --description--

Mpangilio ni muhimu katika kauli za `if`, `else if`.

Kitendakazi hufanyika kutoka juu hadi chini kwa hivyo utahitaji kuwa makini na kauli gani inakuja kwanza.

Chukua mifano hii miwili ya vitendakazi.

Huu ni wa kwanza:

```js
function foo(x) {
  if (x < 1) {
    return "Less than one";
  } else if (x < 2) {
    return "Less than two";
  } else {
    return "Greater than or equal to two";
  }
}
```

Na wa pili hubadilisha tu mpangilio wa kauli:

```js
function bar(x) {
  if (x < 2) {
    return "Less than two";
  } else if (x < 1) {
    return "Less than one";
  } else {
    return "Greater than or equal to two";
  }
}
```

Ingawa vitendakazi hivi viwili vinaonekana karibu sawa, tukipitisha nambari kwa vyote tunapata matokeo tofauti.

```js
foo(0)
bar(0)
```

`foo(0)` itarudisha mfuatano wa herufi `Less than one`, na `bar(0)` itarudisha mfuatano wa herufi `Less than two`.

# --instructions--

Badilisha mpangilio wa mantiki katika kitendakazi ili kitarudisha kauli sahihi katika kila hali.

# --hints--

`orderMyLogic(4)` inapaswa kurudisha mfuatano wa herufi `Less than 5`

```js
assert(orderMyLogic(4) === 'Less than 5');
```

`orderMyLogic(6)` inapaswa kurudisha mfuatano wa herufi `Less than 10`

```js
assert(orderMyLogic(6) === 'Less than 10');
```

`orderMyLogic(11)` inapaswa kurudisha mfuatano wa herufi `Greater than or equal to 10`

```js
assert(orderMyLogic(11) === 'Greater than or equal to 10');
```

# --seed--

## --seed-contents--

```js
function orderMyLogic(val) {
  if (val < 10) {
    return "Less than 10";
  } else if (val < 5) {
    return "Less than 5";
  } else {
    return "Greater than or equal to 10";
  }
}

orderMyLogic(7);
```

# --solutions--

```js
function orderMyLogic(val) {
  if(val < 5) {
    return "Less than 5";
  } else if (val < 10) {
    return "Less than 10";
  } else {
    return "Greater than or equal to 10";
  }
}
```
