---
id: 56bbb991ad1ed5201cd392d3
title: Futa sifa kutoka kwa kitu cha JavaScript
challengeType: 1
forumTopicId: 17560
dashedName: delete-properties-from-a-javascript-object
---

# --description--

Tunaweza pia kufuta sifa kutoka kwa vitu kama hivi:

```js
delete ourDog.bark;
```

Mfano:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"],
  "bark": "bow-wow"
};

delete ourDog.bark;
```

Baada ya mstari wa mwisho ulioonyeshwa hapo juu, `ourDog` inaonekana kama ifuatavyo:

```js
{
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
}
```

# --instructions--

Futa sifa ya `tails` kutoka `myDog`. Unaweza kutumia uandishi wa nukta au uandishi wa mabano.

# --hints--

Unapaswa kufuta sifa ya `tails` kutoka `myDog`.

```js
assert(typeof myDog === 'object' && myDog.tails === undefined);
```

Haupaswi kubadilisha usanidi wa `myDog`.

```js
assert(__helpers.removeJSComments(code).match(/"tails": 1/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};
delete myDog.tails;
```
