---
id: 56bbb991ad1ed5201cd392cd
title: Badilisha safu za data kwa njia ya shift
challengeType: 1
forumTopicId: 18238
dashedName: manipulate-arrays-with-shift
---

# --description--

`pop()` huondoa kipengele cha mwisho cha safu ya data kila wakati. Je, unataka kuondoa cha kwanza?

Hapo ndipo `.shift()` inapoingia. Inafanya kazi kama `.pop()`, isipokuwa huondoa kipengele cha kwanza badala ya cha mwisho.

Mfano:

```js
const ourArray = ["Stimpson", "J", ["cat"]];
const removedFromOurArray = ourArray.shift();
```

`removedFromOurArray` ingekuwa na thamani ya mfuatano wa herufi `Stimpson`, na `ourArray` ingekuwa na `["J", ["cat"]]`.

# --instructions--

Tumia kitendakazi cha `.shift()` kuondoa kipengele cha kwanza kutoka `myArray` na weka thamani iliyotolewa "shifted off" katika kigezo kipya, `removedFromMyArray`.

# --hints--

`myArray` sasa inapaswa kuwa sawa na `[["dog", 3]]`.

```js
assert(
  (function (d) {
    if (d[0][0] == 'dog' && d[0][1] === 3 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`removedFromMyArray` inapaswa kuwa na `["John", 23]`.

```js
assert(
  (function (d) {
    if (
      d[0] == 'John' &&
      d[1] === 23 &&
      typeof removedFromMyArray === 'object'
    ) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line
const removedFromMyArray = myArray.shift();
```
