---
id: 56bbb991ad1ed5201cd392ce
title: Badilisha safu za data kwa njia ya unshift
challengeType: 1
forumTopicId: 18239
dashedName: manipulate-arrays-with-unshift
---

# --description--

Si tu unaweza `shift` vipengele mwanzoni mwa safu ya data, pia unaweza `unshift` vipengele mwanzoni mwa safu ya data, yaani ongeza vipengele mbele ya safu ya data.

`.unshift()` hufanya kazi sawa kabisa na `.push()`, lakini badala ya kuongeza kipengele mwishoni mwa safu ya data, `unshift()` huongeza kipengele mwanzoni mwa safu ya data.

Mfano:

```js
const ourArray = ["Stimpson", "J", "cat"];
ourArray.shift();
ourArray.unshift("Happy");
```

Baada ya `shift`, `ourArray` itakuwa na thamani ya `["J", "cat"]`. Baada ya `unshift`, `ourArray` itakuwa na thamani ya `["Happy", "J", "cat"]`.

# --instructions--

Ongeza `["Paul", 35]` mwanzoni mwa variable ya `myArray` kwa kutumia `unshift()`.

# --hints--

`myArray` inapaswa sasa kuwa na `[["Paul", 35], ["dog", 3]]`.

```js
assert(
  (function (d) {
    if (
      typeof d[0] === 'object' &&
      d[0][0] == 'Paul' &&
      d[0][1] === 35 &&
      d[1][0] != undefined &&
      d[1][0] == 'dog' &&
      d[1][1] != undefined &&
      d[1][1] == 3
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();
myArray.unshift(["Paul", 35]);
```
