---
id: 56bbb991ad1ed5201cd392cb
title: Badilisha safu za data kwa njia ya push
challengeType: 1
forumTopicId: 18237
dashedName: manipulate-arrays-with-push
---

# --description--

Njia rahisi ya kuongeza data mwishoni mwa safu ya data ni kupitia njia ya `push()`.

Njia ya `push()` inachukua hoja moja au zaidi na kuziunganisha mwishoni mwa safu ya data, kwa mpangilio zinavyoonekana. Inarudisha urefu mpya wa safu ya data.

Mifano:

```js
const arr1 = [1, 2, 3];
arr1.push(4, 5);

const arr2 = ["Stimpson", "J", "cat"];
arr2.push(["happy", "joy"]);
```

`arr1` sasa ina thamani ya `[1, 2, 3, 4, 5]` na `arr2` ina thamani ya `["Stimpson", "J", "cat", ["happy", "joy"]]`.

# --instructions--

Sukuma `["dog", 3]` mwishoni mwa kigezo cha `myArray`.

# --hints--

`myArray` inapaswa sasa kuwa sawa na `[["John", 23], ["cat", 2], ["dog", 3]]`.

```js
assert(
  (function (d) {
    if (
      d[2] != undefined &&
      d[0][0] == 'John' &&
      d[0][1] === 23 &&
      d[2][0] == 'dog' &&
      d[2][1] === 3 &&
      d[2].length == 2
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
myArray.push(["dog",3]);
```
