---
id: 56bbb991ad1ed5201cd392d0
title: Jenga vitu vya JavaScript
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

Huenda umewahi kusikia neno `object` hapo awali.

Vitu ni sawa na `arrays`, isipokuwa badala ya kutumia vielezo kufikia na kubadilisha data zao, unapata data katika vitu kupitia kile kinachoitwa `properties`.

Vitu ni muhimu kuhifadhi data kwa njia iliyopangwa, na vinaweza wakilisha vitu halisi vya dunia, kama paka.

Hapa kuna mfano wa kitu cha paka:

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

Katika mfano huu, sifa zote zimehifadhiwa kama mfuatano wa herufi, kama `name`, `legs`, na `tails`. Hata hivyo, unaweza pia kutumia nambari kama sifa. Unaweza hata kuacha nukuu kwa sifa za mfuatano wa herufi zenye neno moja, kama ifuatavyo:

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

Hata hivyo, ikiwa kitu chako kina sifa zozote zisizo za mfuatano wa herufi, JavaScript itabadilisha aina zao moja kwa moja kuwa mfuatano wa herufi.

# --instructions--

Tengeneza kitu kinachowakilisha mbwa anayeitwa `myDog` ambacho kina sifa `name` (mfuatano wa herufi), `legs`, `tails` na `friends`.

Unaweza kuweka sifa hizi za kitu kwa thamani yoyote unayotaka, mradi tu `name` ni mfuatano wa herufi, `legs` na `tails` ni nambari, na `friends` ni safu ya data.

# --hints--

`myDog` inapaswa kuwa na sifa `name` na inapaswa kuwa `string`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` inapaswa kuwa na sifa `legs` na inapaswa kuwa `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` inapaswa kuwa na sifa `tails` na inapaswa kuwa `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` inapaswa kuwa na sifa `friends` na inapaswa kuwa safu ya data `array`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` inapaswa kuwa na sifa zote zilizotolewa tu.

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
