---
id: 56bbb991ad1ed5201cd392ca
title: Upatikanaji wa data ya safu ya data kwa kutumia vielezo
challengeType: 1
forumTopicId: 16158
dashedName: access-array-data-with-indexes
---

# --description--

Tunaweza kupata data ndani ya safu za data kwa kutumia <dfn>vielezo</dfn>.

Vielezo vya safu ya data vinaandikwa kwa uandishi wa mabano kama vile mfuatano wa herufi unavyotumia, isipokuwa badala ya kubainisha herufi, vinaelezea kipengee katika safu ya data. Kama mfuatano wa herufi, safu za data hutumia uandishi wa vielezo <dfn>kuanzia sufuri</dfn>, kwa hivyo kipengee cha kwanza katika safu ya data kina kielezo cha `0`.

<br>

**Mfano**

```js
const array = [50, 60, 70];
console.log(array[0]);
const data = array[1];
```

`console.log(array[0])` inachapisha `50`, na `data` ina thamani ya `60`.

# --instructions--

Tengeneza kigezo kinachoitwa `myData` na uweke sawa na thamani ya kwanza ya `myArray` kwa kutumia uandishi wa mabano.

# --hints--

Kigezo `myData` kinapaswa kuwa sawa na thamani ya kwanza ya `myArray`.

```js
assert(
  (function () {
    if (
      typeof myArray !== 'undefined' &&
      typeof myData !== 'undefined' &&
      myArray[0] === myData
    ) {
      return true;
    } else {
      return false;
    }
  })()
);
```

Data katika kigezo `myArray` inapaswa kufikiliwa kwa kutumia uandishi wa mabano.

```js
assert(
  (function () {
    if (__helpers.removeJSComments(code).match(/\s*=\s*myArray\[0\]/g)) {
      return true;
    } else {
      return false;
    }
  })()
);
```

# --seed--

## --seed-contents--

```js
const myArray = [50, 60, 70];


```

# --solutions--

```js
const myArray = [50, 60, 70];
const myData = myArray[0];
```
