---
id: 56bbb991ad1ed5201cd392d2
title: Ongeza sifa mpya kwa kitu cha JavaScript
challengeType: 1
forumTopicId: 301169
dashedName: add-new-properties-to-a-javascript-object
---

# --description--

Unaweza kuongeza sifa mpya kwa vitu vya JavaScript vilivyopo kwa njia ile ile unavyobadilisha.

Hapa ni jinsi tunavyoongeza sifa ya `bark` kwa `ourDog`:

```js
ourDog.bark = "bow-wow";
```

au

```js
ourDog["bark"] = "bow-wow";
```

Sasa tunapopima `ourDog.bark`, tutapata sauti yake ya kubweka, `bow-wow`.

Mfano:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};

ourDog.bark = "bow-wow";
```

# --instructions--

Ongeza sifa ya `bark` kwa `myDog` na uiweke kuwa sauti ya mbwa, kama "woof". Unaweza kutumia uandishi wa nukta au uandishi wa mabano.

# --hints--

Unapaswa kuongeza sifa `bark` kwa `myDog`.

```js
assert(myDog.bark !== undefined);
```

Haupaswi kuongeza `bark` katika uanzishaji wa `myDog`.

```js
assert(!/bark[^\n]:/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};


```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.bark = "Woof Woof";
```
