---
id: 56bbb991ad1ed5201cd392d1
title: Kusahihisha sifa za kitu
challengeType: 1
forumTopicId: 18336
dashedName: updating-object-properties
---

# --description--

Baada ya kuunda kitu cha JavaScript, unaweza kusahihisha sifa zake wakati wowote kama unavyofanya kwa aina nyingine yoyote ya mabadiliko. Unaweza kutumia uandishi wa nukta au uandishi wa mabano kusahihisha.

Kwa mfano, tuchukulie `ourDog`:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```

Kwa kuwa ni mbwa mwenye furaha sana, tubadilishe jina lake kuwa mfuatano wa herufi `Happy Camper`. Hapa ni jinsi tunavyosahihisha sifa ya jina la kitu chake: `ourDog.name = "Happy Camper";` au `ourDog["name"] = "Happy Camper";` Sasa tunapopima `ourDog.name`, badala ya kupata `Camper`, tutapata jina lake jipya, `Happy Camper`.

# --instructions--

Sahihisha sifa ya jina la kitu cha `myDog`. Tubadilishe jina lake kutoka `Coder` kuwa `Happy Coder`. Unaweza kutumia uandishi wa nukta au uandishi wa mabano.

# --hints--

Unapaswa kusahihisha sifa ya `myDog` ya `name` ili iwe sawa na mfuatano wa herufi `Happy Coder`.

```js
assert(/happy coder/gi.test(myDog.name));
```

Hauruhusiwi kuhariri ufafanuzi wa `myDog`.

```js
assert(/"name": "Coder"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.name = "Happy Coder";
```
