---
id: 56bbb991ad1ed5201cd392cf
title: Andika JavaScript inayoweza kutumika tena kwa kutumia kitendakazi
challengeType: 1
forumTopicId: 18378
dashedName: write-reusable-javascript-with-functions
---

# --description--

Katika JavaScript, tunaweza kugawanya msimbo wetu katika sehemu zinazoweza kutumika tena zinazojulikana kama <dfn>kitendakazi</dfn>.

Hapa kuna mfano wa kitendakazi:

```js
function functionName() {
  console.log("Hello World");
}
```

Unaweza kuita au <dfn>ita</dfn> kitendakazi hiki kwa kutumia jina lake ikifuatiwa na mabano ya kawaida, kama hivi: `functionName();` Kila wakati kitendakazi kinapoitwa kitachapisha ujumbe `Hello World` kwenye konsoli ya maendeleo. Msimbo wote ulioko kati ya mabano ya mviringo utaendeshwa kila wakati kitendakazi kinapoitwa.

# --instructions--

<ol>
  <li>
    Unda kitendakazi kinachoitwa <code>reusableFunction</code> ambacho kinachapisha mfuatano wa herufi <code>Hi World</code> kwenye konsoli ya maendeleo.
  </li>
  <li>
    Ita kitendakazi hicho.
  </li>
</ol>

# --before-each--

```js
function testConsole() {
  var originalConsole = console;
  var nativeLog = console.log;
  var hiWorldWasLogged = false;
  console.log = function (message) {
    if (message === 'Hi World') {
      console.warn(message);
      hiWorldWasLogged = true;
    }
    if (nativeLog.apply) {
      nativeLog.apply(originalConsole, arguments);
    } else {
      var nativeMsg = Array.prototype.slice.apply(arguments).join(' ');
      nativeLog(nativeMsg);
    }
  };
  reusableFunction();
  console.log = nativeLog;
  return hiWorldWasLogged;
}
```

# --hints--

`reusableFunction` inapaswa kuwa kitendakazi.

```js
assert(typeof reusableFunction === 'function');
```

Ikiwa `reusableFunction` itaitwa, inapaswa kutoa mfuatano wa herufi `Hi World` kwenye konsoli.

```js
assert(testConsole());
```

Unapaswa kuita `reusableFunction` mara moja itakapobainishwa.

```js
const functionStr = reusableFunction && __helpers.removeWhiteSpace(reusableFunction.toString());
const codeWithoutFunction = __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).replace(/reusableFunction\(\)\{/g, '');
assert(/reusableFunction\(\)/.test(codeWithoutFunction));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reusableFunction() {
  console.log("Hi World");
}
reusableFunction();
```
