---
id: 56bbb991ad1ed5201cd392cc
title: Badilisha safu za data kwa njia ya pop
challengeType: 1
forumTopicId: 18236
dashedName: manipulate-arrays-with-pop
---

# --description--

Njia nyingine ya kubadilisha data katika safu ya data ni kwa kutumia kitendakazi cha `.pop()`.

`.pop()` hutumika kuondoa thamani kutoka mwisho wa safu ya data. Tunaweza kuhifadhi thamani hii iliyotolewa kwa kuweka thamani yake katika kigezo kipya. Kwa maneno mengine, `.pop()` huondoa kipengele cha mwisho kutoka katika safu ya data na hurudisha kipengele hicho.

Aina yoyote ya kipengele inaweza kuondolewa kutoka katika safu ya data - nambari, mfuatano wa herufi, hata safu za data zilizopangwa ndani.

```js
const threeArr = [1, 4, 6];
const oneDown = threeArr.pop();
console.log(oneDown);
console.log(threeArr);
```

`console.log` ya kwanza itaonyesha thamani ya `6`, na ya pili itaonyesha thamani ya `[1, 4]`.

# --instructions--

Tumia kitendakazi cha `.pop()` kuondoa kipengele cha mwisho kutoka `myArray` na weka thamani iliyotolewa katika kigezo kipya, `removedFromMyArray`.

# --hints--

`myArray` inapaswa kuwa na `[["John", 23]]` tu.

```js
assert(
  (function (d) {
    if (d[0][0] == 'John' && d[0][1] === 23 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

Unapaswa kutumia `pop()` kwenye `myArray`.

```js
assert(/removedFromMyArray\s*=\s*myArray\s*.\s*pop\s*(\s*)/.test(__helpers.removeJSComments(code)));
```

`removedFromMyArray` inapaswa kuwa na `["cat", 2]` tu.

```js
assert(
  (function (d) {
    if (d[0] == 'cat' && d[1] === 2 && d[2] == undefined) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
const removedFromMyArray = myArray.pop();
```
