---
id: 587d7b7e367417b2b2512b24
title: Tumia opereta ya kikokotoo cha masharti cha pande tatu
challengeType: 1
forumTopicId: 301181
dashedName: use-the-conditional-ternary-operator
---

# --description--

<dfn>Opereta ya masharti</dfn>, inayojulikana pia kama <dfn>opereta ya ternary</dfn>, inaweza kutumika kama usemi wa if-else wa mstari mmoja.

Sintaksia ni `a ? b : c`, ambapo `a` ni sharti, `b` ni msimbo unaotekelezwa wakati sharti linarudisha `true`, na `c` ni msimbo unaotekelezwa wakati sharti linarudisha `false`.

Kitendakazi kilicho hapa chini kinatumia tamko la `if/else` kukagua sharti:

```js
function findGreater(a, b) {
  if(a > b) {
    return "a is greater";
  }
  else {
    return "b is greater or equal";
  }
}
```

Hii inaweza kuandikwa upya kwa kutumia opereta ya masharti:

```js
function findGreater(a, b) {
  return a > b ? "a is greater" : "b is greater or equal";
}
```

# --instructions--

Tumia opereta ya masharti katika kitendakazi cha `checkEqual` kukagua kama nambari mbili ni sawa au la. Kitendakazi kinapaswa kurudisha mfuatano wa herufi `Equal` au mfuatano wa herufi `Not Equal`.

# --hints--

`checkEqual` inapaswa kutumia opereta ya masharti

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?/.test(__helpers.removeJSComments(code)));
```

`checkEqual(1, 2)` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(checkEqual(1, 2) === 'Not Equal');
```

`checkEqual(1, 1)` inapaswa kurudisha mfuatano wa herufi `Equal`

```js
assert(checkEqual(1, 1) === 'Equal');
```

`checkEqual(1, -1)` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(checkEqual(1, -1) === 'Not Equal');
```

# --seed--

## --seed-contents--

```js
function checkEqual(a, b) {

}

checkEqual(1, 2);
```

# --solutions--

```js
function checkEqual(a, b) {
  return a === b ? "Equal" : "Not Equal";
}
```
