---
id: 587d7b7e367417b2b2512b22
title: Tumia kitendakazi cha parseInt na msingi wa nambari
challengeType: 1
forumTopicId: 301182
dashedName: use-the-parseint-function-with-a-radix
---

# --description--

Kitendakazi cha `parseInt()` huchanganuzi mfuatano wa herufi na kurudisha nambari kamili. Huchukua hoja ya pili kwa ajili ya msingi wa nambari, ambayo inaelezea msingi wa nambari katika mfuatano wa herufi. Msingi wa nambari unaweza kuwa nambari kamili kati ya 2 na 36.

Wito wa kitendakazi unaonekana kama ifuatavyo:

```js
parseInt(string, radix);
```

Na hapa kuna mfano:

```js
const a = parseInt("11", 2);
```

Kigezo cha msingi wa nambari kinaonyesha kuwa `11` iko katika mfumo wa binary, au msingi wa nambari 2. Mfano huu hubadilisha mfuatano wa herufi `11` kuwa nambari kamili `3`.

# --instructions--

Tumia `parseInt()` katika kitendakazi cha `convertToInteger` ili kibadilishe nambari ya binary kuwa nambari kamili na kikirudishe.

# --hints--

`convertToInteger` inapaswa kutumia kitendakazi cha `parseInt()`

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("10011")` inapaswa kurudisha nambari

```js
assert(typeof convertToInteger('10011') === 'number');
```

`convertToInteger("10011")` inapaswa kurudisha 19

```js
assert(convertToInteger('10011') === 19);
```

`convertToInteger("111001")` inapaswa kurudisha 57

```js
assert(convertToInteger('111001') === 57);
```

`convertToInteger("JamesBond")` inapaswa kurudisha `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("10011");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str, 2);
}
```
