---
id: 587d7b7e367417b2b2512b21
title: Tumia waendeshaji wa masharti (ternary) wengi
challengeType: 1
forumTopicId: 301179
dashedName: use-multiple-conditional-ternary-operators
---

# --description--

Katika changamoto iliyopita, ulitumia waendeshaji mmoja wa masharti. Pia unaweza kuunganisha waendeshaji hawa pamoja ili kukagua masharti mengi.

Kitendakazi kifuatacho kinatumia `if`, `else if`, na `else` ili kukagua masharti mengi:

```js
function findGreaterOrEqual(a, b) {
  if (a === b) {
    return "a and b are equal";
  }
  else if (a > b) {
    return "a is greater";
  }
  else {
    return "b is greater";
  }
}
```

Kitendakazi hicho hapo juu kinaweza kuandikwa upya kwa kutumia waendeshaji wa masharti wengi:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" 
    : (a > b) ? "a is greater" 
    : "b is greater";
}
```

Inachukuliwa kuwa mbinu bora kupanga waendeshaji wa masharti wengi ili kila sharti liwe kwenye mstari wake, kama ilivyoonyeshwa hapo juu. Kutumia waendeshaji wa masharti wengi bila kupanga vizuri inaweza kufanya msimbo wako usomeke kwa urahisi. Kwa mfano:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" : (a > b) ? "a is greater" : "b is greater";
}
```

# --instructions--

Katika kitendakazi cha `checkSign`, tumia waendeshaji wa masharti wengi - ukifuata mpangilio uliopendekezwa unaotumika kwenye `findGreaterOrEqual` - ili kukagua kama nambari ni chanya, hasi au sifuri. Kitendakazi kinapaswa kurejesha `positive`, `negative` au `zero`.

# --hints--

`checkSign` inapaswa kutumia waendeshaji wa masharti wengi

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?\s*?\?\s*?.+?\s*?:\s*?.+?/gi.test(__helpers.removeJSComments(code)));
```

`checkSign(10)` inapaswa kurejesha mfuatano wa herufi `positive`. Kumbuka kwamba herufi kubwa na ndogo ni muhimu

```js
assert(checkSign(10) === 'positive');
```

`checkSign(-12)` inapaswa kurejesha mfuatano wa herufi `negative`. Kumbuka kwamba herufi kubwa na ndogo ni muhimu

```js
assert(checkSign(-12) === 'negative');
```

`checkSign(0)` inapaswa kurejesha mfuatano wa herufi `zero`. Kumbuka kwamba herufi kubwa na ndogo ni muhimu

```js
assert(checkSign(0) === 'zero');
```

# --seed--

## --seed-contents--

```js
function checkSign(num) {

}

checkSign(10);
```

# --solutions--

```js
function checkSign(num) {
  return (num > 0) ? 'positive' : (num < 0) ? 'negative' : 'zero';
}
```
