---
id: 587d7b7e367417b2b2512b23
title: Tumia kitendakazi cha parseInt
challengeType: 1
forumTopicId: 301183
dashedName: use-the-parseint-function
---

# --description--

Kitendakazi cha `parseInt()` huchanganuzi mfuatano wa herufi na kurudisha nambari kamili. Hapa kuna mfano:

```js
const a = parseInt("007");
```

Kitendakazi kilicho hapo juu hubadilisha mfuatano wa herufi `007` kuwa nambari kamili `7`. Ikiwa herufi ya kwanza katika mfuatano hawezi kubadilishwa kuwa nambari, basi hurudisha `NaN`.

# --instructions--

Tumia `parseInt()` katika kitendakazi cha `convertToInteger` ili kibadilishe mfuatano wa herufi wa ingizo `str` kuwa nambari kamili, na kikirudishe.

# --hints--

`convertToInteger` inapaswa kutumia kitendakazi cha `parseInt()`

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("56")` inapaswa kurudisha nambari

```js
assert(typeof convertToInteger('56') === 'number');
```

`convertToInteger("56")` inapaswa kurudisha 56

```js
assert(convertToInteger('56') === 56);
```

`convertToInteger("77")` inapaswa kurudisha 77

```js
assert(convertToInteger('77') === 77);
```

`convertToInteger("JamesBond")` inapaswa kurudisha `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("56");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str);
}
```
