---
id: 598e8944f009e646fc236146
title: Kuelewa thamani isiyoainishwa inayorejeshwa kutoka kwa kitendakazi
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

Kitendakazi kinaweza kujumuisha tamko la `return` lakini si lazima kiwe nacho. Katika hali ambapo kitendakazi hakina tamko la `return`, unapoifanya wito, kitendakazi huchakata msimbo wa ndani lakini thamani inayorejeshwa ni `undefined`.

**Mfano**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum` ni kitendakazi kisicho na tamko la `return`. Kitendakazi kitabadilisha kigezo cha `sum` cha ulimwengu mzima lakini thamani inayorejeshwa ya kitendakazi ni `undefined`.

# --instructions--

Tengeneza kitendakazi cha `addFive` kisicho na hoja yoyote. Kitendakazi hiki kinaongeza 5 kwenye kigezo cha `sum`, lakini thamani inayorejeshwa ni `undefined`.

# --hints--

`addFive` inapaswa kuwa kitendakazi.

```js
assert(typeof addFive === 'function');
```

Mara zote kitendakazi zote zikiendeshwa, `sum` inapaswa kuwa sawa na `8`.

```js
assert(sum === 8);
```

Thamani inayorejeshwa kutoka `addFive` inapaswa kuwa `undefined`.

```js
assert(addFive() === undefined);
```

Ndani ya kitendakazi cha `addFive`, unapaswa kuongeza `5` kwenye kigezo cha `sum`.

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
