---
id: 599a789b454f2bbd91a3ff4d
title: Mazoezi ya kulinganisha maadili tofauti
challengeType: 1
forumTopicId: 301174
dashedName: practice-comparing-different-values
---

# --description--

Katika changamoto mbili za mwisho, tulijifunza kuhusu opereta ya usawa (`==`) na opereta ya usawa mkali (`===`). Hebu tufanye kagua haraka na kufanya mazoezi ya kutumia opereta hizi zaidi.

Ikiwa maadili yanayolinganishwa si ya aina moja, opereta ya usawa itafanya uongofu wa aina, kisha itahesabu maadili. Hata hivyo, opereta ya usawa mkali italinganisha aina ya data na thamani kama zilivyo, bila kubadilisha aina moja kuwa nyingine.

**Mifano**

`3 == '3'` inarudisha `true` kwa sababu JavaScript hufanya uongofu wa aina kutoka mfuatano wa herufi kwenda nambari. `3 === '3'` inarudisha `false` kwa sababu aina ni tofauti na uongofu wa aina haufanyiki.

**Kumbuka:** Katika JavaScript, unaweza kubaini aina ya kigezo au thamani kwa kutumia opereta `typeof`, kama ifuatavyo:

```js
typeof 3
typeof '3'
```

`typeof 3` inarudisha mfuatano wa herufi `number`, na `typeof '3'` inarudisha mfuatano wa herufi `string`.

# --instructions--

Kitendakazi `compareEquality` katika mhariri kinalinganisha maadili mawili kwa kutumia opereta ya usawa. Badilisha kitendakazi ili kirudishe mfuatano wa herufi `Equal` tu wakati maadili ni sawa kabisa kwa usawa mkali.

# --hints--

`compareEquality(10, "10")` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(compareEquality(10, '10') === 'Not Equal');
```

`compareEquality("20", 20)` inapaswa kurudisha mfuatano wa herufi `Not Equal`

```js
assert(compareEquality('20', 20) === 'Not Equal');
```

Unapaswa kutumia opereta `===`

```js
assert(__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function compareEquality(a, b) {
  if (a == b) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

compareEquality(10, "10");
```

# --solutions--

```js
function compareEquality(a,b) {
  if (a === b) {
    return "Equal";
  }
  return "Not Equal";
}
```
