---
id: 5ee127a03c3b35dd45426493
title: Kuweka thamani ya kigezo kimoja kwa kingine
challengeType: 1
forumTopicId: 418265
dashedName: assigning-the-value-of-one-variable-to-another
---

# --description--

Baada ya thamani kuwekwa kwa kigezo kwa kutumia kiendeshaji cha <dfn>mgawo wa thamani</dfn>, unaweza kuweka thamani ya kigezo hicho kwa kigezo kingine kwa kutumia kiendeshaji cha <dfn>mgawo wa thamani</dfn>.

```js
var myVar;
myVar = 5;
var myNum;
myNum = myVar;
```

Hapo juu inatangaza kigezo cha `myVar` bila thamani, kisha inakiweka thamani `5`. Ifuatayo, kigezo kinachoitwa `myNum` kinatangazwa bila thamani. Kisha, yaliyomo ya `myVar` (ambayo ni `5`) yanawekwa kwa kigezo `myNum`. Sasa, `myNum` pia ina thamani ya `5`.

# --instructions--

Weka yaliyomo ya `a` kwa kigezo `b`.

# --hints--

Haupasi kubadilisha msimbo juu ya maoni yaliyotajwa.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)) && /a = 7;/.test(__helpers.removeJSComments(code)) && /var b;/.test(__helpers.removeJSComments(code)));
```

`b` inapaswa kuwa na thamani ya `7`.

```js
assert(typeof b === 'number' && b === 7);
```

`a` inapaswa kuwekwa kwa `b` kwa kutumia `=`.

```js
assert(/b\s*=\s*a\s*/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
var a;
a = 7;
var b;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
var b;
b = a;
```
