---
id: bd7123c9c441eddfaeb5bdef
title: Kuelewa thamani za kweli au siyo kweli
challengeType: 1
forumTopicId: 301176
dashedName: understanding-boolean-values
---

# --description--

Aina nyingine ya data ni <dfn>thamani ya kweli au siyo kweli</dfn>. Thamani za kweli au siyo kweli zinaweza kuwa moja kati ya thamani mbili tu: `true` au `false`. Kwa msingi, ni vitufe vidogo vya kuwasha-kuzima, ambapo `true` ni kuwasha na `false` ni kuzima. Hali hizi mbili ni za kipekee na haziziruhusu kutokea pamoja.

**Kumbuka:** Thamani za kweli au siyo kweli hazinaandikwa kamwe kwa nukuu. Mfuatano wa herufi `"true"` na `"false"` si thamani za kweli au siyo kweli na hazina maana maalum katika JavaScript.

# --instructions--

Badilisha kitendakazi cha `welcomeToBooleans` ili kirudishe `true` badala ya `false`.

# --hints--

Kitendakazi cha `welcomeToBooleans()` kinapaswa kurudisha thamani ya kweli au siyo kweli (`true` au `false`).

```js
assert(typeof welcomeToBooleans() === 'boolean');
```

`welcomeToBooleans()` inapaswa kurudisha `true`.

```js
assert(welcomeToBooleans() === true);
```

# --seed--

## --seed-contents--

```js
function welcomeToBooleans() {
  // Only change code below this line

  return false; // Change this line

  // Only change code above this line
}
```

# --solutions--

```js
function welcomeToBooleans() {
  return true; // Change this line
}
```
