---
id: bd7993c9c69feddfaeb7bdef
title: Zidisha desimali mbili kwa JavaScript
challengeType: 1
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

Katika JavaScript, unaweza pia kufanya mahesabu na nambari za desimali, kama vile nambari kamili.

Tuzidishe desimali mbili pamoja kupata bidhaa yao.

# --instructions--

Badilisha `0.0` ili bidhaa iwe sawa na `5.0`.

# --hints--

Kigezo `product` kinapaswa kuwa sawa na `5.0`.

```js
assert(product === 5.0);
```

Unapaswa kutumia opereta wa `*`

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
