---
id: cf1111c1c11feddfaeb6bdef
title: Gawanya nambari moja kwa nyingine kwa kutumia JavaScript
challengeType: 1
forumTopicId: 17566
dashedName: divide-one-number-by-another-with-javascript
---

# --description--

Tunaweza pia kugawanya nambari moja kwa nyingine.

JavaScript hutumia alama ya `/` kwa ajili ya mgawanyo.

**Mfano**

```js
const myVar = 16 / 2;
```

`myVar` sasa ina thamani ya `8`.
# --instructions--

Badilisha `0` ili `quotient` iwe sawa na `2`.

# --hints--

Kigezo `quotient` kinapaswa kuwa sawa na 2.

```js
assert(quotient === 2);
```

Unapaswa kutumia opereta wa `/`.

```js
assert(/\d+\s*\/\s*\d+/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const quotient = 66 / 0;
```

# --solutions--

```js
const quotient = 66 / 33;
```
