---
id: cf1111c1c11feddfaeb3bdef
title: Ongeza nambari mbili kwa kutumia JavaScript
challengeType: 1
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` ni aina ya data katika JavaScript inayowakilisha data ya nambari.

Sasa jaribu kuongeza nambari mbili kwa kutumia JavaScript.

JavaScript hutumia alama ya `+` kama kiendeshaji cha kuongeza wakati imewekwa kati ya nambari mbili.

**Mfano:**

```js
const myVar = 5 + 10;
```

`myVar` sasa ina thamani ya `15`.

# --instructions--

Badilisha `0` ili jumla iwe sawa na `20`.

# --hints--

`sum` inapaswa kuwa sawa na `20`.

```js
assert(sum === 20);
```

Unapaswa kutumia kiendeshaji cha `+`.

```js
assert(/\+/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
