---
id: cf1111c1c11feddfaeb4bdef
title: Punguza nambari moja kutoka kwa nyingine kwa kutumia JavaScript
challengeType: 1
forumTopicId: 18314
dashedName: subtract-one-number-from-another-with-javascript
---

# --description--

Tunaweza pia kupunguza nambari moja kutoka kwa nyingine.

JavaScript hutumia alama ya `-` kwa ajili ya upunguzaji.

**Mfano**

```js
const myVar = 12 - 6;
```

`myVar` itakuwa na thamani ya `6`.
# --instructions--

Badilisha `0` ili tofauti iwe `12`.

# --hints--

Kigezo `difference` kinapaswa kuwa sawa na `12`.

```js
assert(difference === 12);
```

Unapaswa kupunguza nambari moja tu kutoka `45`.

```js
assert(/difference=45-33;?/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --seed-contents--

```js
const difference = 45 - 0;
```

# --solutions--

```js
const difference = 45 - 33;
```
