---
id: cf1231c1c11feddfaeb5bdef
title: Zidisha nambari mbili kwa JavaScript
challengeType: 1
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

Tunaweza pia kuzidisha nambari moja kwa nyingine.

JavaScript hutumia alama ya `*` kwa ajili ya kuzidisha nambari mbili.

**Mfano**

```js
const myVar = 13 * 13;
```

`myVar` itakuwa na thamani ya `169`.

# --instructions--

Badilisha `0` ili bidhaa iwe sawa na `80`.

# --hints--

Kigezo `product` kinapaswa kuwa sawa na 80.

```js
assert(product === 80);
```

Unapaswa kutumia opereta wa `*`.

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 8 * 0;
```

# --solutions--

```js
const product = 8 * 10;
```
