---
id: cf1391c1c11feddfaeb4bdef
title: Tengeneza nambari za sehemu za mfululizo na JavaScript
challengeType: 1
forumTopicId: 16826
dashedName: create-decimal-numbers-with-javascript
---

# --description--

Tunaweza pia kuhifadhi nambari za sehemu za mfululizo katika vigezo. Nambari za sehemu za mfululizo mara nyingine huitwa nambari za <dfn>floating point</dfn> au <dfn>floats</dfn>.

**Kumbuka:** unapoendesha hesabu za nambari, hufanywa kwa usahihi wa kikomo. Operesheni zinazotumia floating points zinaweza kusababisha matokeo tofauti na yaliyotarajiwa. Ikiwa unapata mojawapo ya matokeo haya, fungua mada kwenye <a href="https://forum.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">jukwaa la freeCodeCamp</a>.

# --instructions--

Tengeneza kigezo `myDecimal` na umpe thamani ya sehemu za mfululizo yenye sehemu ya sehemu (mfano `5.7`).

# --hints--

`myDecimal` inapaswa kuwa nambari.

```js
assert(typeof myDecimal === 'number');
```

`myDecimal` inapaswa kuwa na nukta ya sehemu

```js
assert(myDecimal % 1 != 0);
```

# --seed--

## --seed-contents--

```js
const ourDecimal = 5.7;

// Only change code below this line

```

# --solutions--

```js
const myDecimal = 9.9;
```
