---
id: 66a207974c806a19d6607073
title: 책 재고 관리 앱 만들기
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 텍스트가 `h1`인 `Book Inventory` 요소가 있어야 합니다.
1. `table`, `Title`, `Author`, `Category`, `Status`라는 이름의 열을 가진 `Rate` 요소가 있어야 합니다.
1. 표에는 최소 네 개의 행이 있어야 하며, 첫 번째는 열 제목이고 나머지는 책 정보로 채워져야 합니다.
1. 표 본문 안의 각 행은 `read`, `to-read`, 또는 `in-progress` 클래스 중 하나를 가져야 합니다.
1. `td` 열의 `Status` 요소는 `span` 요소를 포함해야 하며, 이 요소는 `class`의 `status`로 텍스트 `Read`, `To Read`, 또는 `In Progress`을 감싸야 하며, 이는 해당 행의 클래스에 따라 다릅니다.
1. `td` 열의 `Rate` 요소는 `span` 요소를 포함해야 하며, 이 요소는 세 개의 빈 `class` 요소를 감싸는 `rate`의 `span`를 가져야 합니다.
1. `.rate` 행 안에 배치된 `.read` 요소는 개인 평점에 따라 `one`, `two`, 또는 `three` 값 중 하나를 추가 클래스로 가져야 하며, 이 값은 `rate` 다음에 와야 합니다.
1. `read`, `to-read`, `in-progress` 클래스를 가진 요소를 대상으로 하는 세 개의 속성 선택자를 만들어야 하며, 이들의 `background-image` 속성을 원하는 `linear-gradient`로 설정해야 합니다.
1. 각 `display` 요소의 `span` 속성을 `inline-block`로 설정해야 합니다.
1. `span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `to-read` 요소를 대상으로 하는 속성 선택자를 사용하여 `border`와 `background-image` 속성을 설정해야 합니다.
1. `span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `read` 요소를 대상으로 하는 속성 선택자를 사용하여 `border`와 `background-image` 속성을 설정해야 합니다.
1. `span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `in-progress` 요소를 대상으로 하는 속성 선택자를 사용하여 `border`와 `background-image` 속성을 설정해야 합니다.
1. `span` 클래스를 가진 `status` 요소와 `span`으로 시작하는 클래스 값을 가진 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `height`, `width`, `padding` 속성을 설정해야 합니다.
1. `span` 요소의 `span` 값이 `class`으로 시작하는 직계 자식인 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `border`, `border-radius`, `margin`, `height`, `width`, `background-color` 속성을 설정해야 합니다.
1. `span` 요소의 `one` 값에 `class`라는 단어가 포함된 첫 번째 자손을 대상으로 하는 속성 선택자를 사용하여 `background-image` 속성을 `linear-gradient`로 설정해야 합니다.
1. `span` 요소의 `two` 값에 `class`라는 단어가 포함된 첫 번째와 두 번째 자손을 대상으로 하는 속성 선택자를 사용하여 `background-image` 속성을 `linear-gradient`로 설정해야 합니다.
1. `span` 요소의 `span` 값에 `three`라는 단어가 포함된 세 개의 `class` 요소를 대상으로 하는 속성 선택자를 사용하여 `background-image` 속성을 `linear-gradient`로 설정해야 합니다.

# --hints--

텍스트가 `h1`인 `Book Inventory` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

`h1` 요소는 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`table` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

`thead` 안에 `tbody` 요소 하나와 `table` 요소 하나가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

`thead` 안에는 5개의 `tr` 요소를 가진 `th`가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

첫 번째 열 제목은 `Title`이어야 합니다.

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

두 번째 열 제목은 `Author`이어야 합니다.

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

세 번째 열 제목은 `Category`이어야 합니다.

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

네 번째 열 제목은 `Status`이어야 합니다.

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

다섯 번째 열 제목은 `Rate`이어야 합니다.

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

표에는 최소 네 개의 행이 있어야 합니다.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

각 행은 항상 5개의 열을 가져야 합니다.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

제목 행을 제외한 각 표 행은 `read`, `to-read`, 또는 `in-progress` 클래스 중 하나를 가져야 합니다.

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

`td` 열의 `Status` 요소는 `span` 요소를 포함해야 합니다.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

`span` 열의 각 `Status` 요소는 `status` 클래스를 가져야 합니다.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

각 `.status` 요소는 행의 클래스에 따라 `Read`, `To Read`, 또는 `In Progress` 텍스트를 가져야 합니다.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

`td` 열의 `Rate` 요소는 `span` 요소를 포함해야 합니다.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

`span` 열의 `td` 요소의 직계 자식인 각 `Rate` 요소는 첫 번째 클래스로 `rate` 클래스를 가져야 합니다.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

각 `.rate` 요소는 세 개의 빈 `span` 요소를 포함해야 합니다.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

`.rate` 행 안에 배치된 `.read` 요소는 개인 평점에 따라 `rate`, `one`, 또는 `two` 값 중 하나를 `three` 클래스 다음에 추가 클래스로 가져야 합니다.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

`read` 클래스를 가진 행을 대상으로 하는 속성 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

`read` 클래스를 가진 행을 대상으로 하는 속성 선택자를 사용하여 원하는 선형 그라디언트로 `background-image` 속성을 설정해야 합니다.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

`to-read` 클래스를 가진 행을 대상으로 하는 속성 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

`to-read` 클래스를 가진 행을 대상으로 하는 속성 선택자를 사용하여 원하는 선형 그라디언트로 `background-image` 속성을 설정해야 합니다.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

`in-progress` 클래스를 가진 행을 대상으로 하는 속성 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

`in-progress` 클래스를 가진 행을 대상으로 하는 속성 선택자를 사용하여 원하는 선형 그라디언트로 `background-image` 속성을 설정해야 합니다.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

각 `display` 요소의 `span` 속성을 `inline-block`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

`span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `to-read` 요소를 대상으로 하는 속성 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

`span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `to-read` 요소를 대상으로 하는 속성 선택자를 사용하여 `border` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

`span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `to-read` 요소를 대상으로 하는 속성 선택자를 사용하여 `background-image` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

`span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `read` 요소를 대상으로 하는 속성 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

`span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `read` 요소를 대상으로 하는 속성 선택자를 사용하여 `border` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

`span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `read` 요소를 대상으로 하는 속성 선택자를 사용하여 `background-image` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

`span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `in-progress` 요소를 대상으로 하는 속성 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

`span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `in-progress` 요소를 대상으로 하는 속성 선택자를 사용하여 `border` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

`span` 요소의 `status` 클래스를 가진 자손인 `tr` 클래스의 `in-progress` 요소를 대상으로 하는 속성 선택자를 사용하여 `background-image` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

`span` 클래스를 가진 `status` 요소와 `span`으로 시작하는 클래스 값을 가진 `rate` 요소를 대상으로 하는 속성 선택자가 있어야 합니다.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

`span` 클래스를 가진 `status` 요소와 `span`으로 시작하는 클래스 값을 가진 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `height` 속성을 설정해야 합니다.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

`span` 클래스를 가진 `status` 요소와 `span`으로 시작하는 클래스 값을 가진 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `width` 속성을 설정해야 합니다.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

`span` 클래스를 가진 `status` 요소와 `span`으로 시작하는 클래스 값을 가진 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `padding` 속성을 설정해야 합니다.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

`span` 요소의 `span` 값이 `class`으로 시작하는 직계 자식인 `rate` 요소를 대상으로 하는 속성 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

`span` 요소의 `span` 값이 `class`으로 시작하는 직계 자식인 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `border` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

`span` 요소의 `span` 값이 `class`으로 시작하는 직계 자식인 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `border-radius` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

`span` 요소의 `span` 값이 `class`으로 시작하는 직계 자식인 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `margin` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

`span` 요소의 `span` 값이 `class`으로 시작하는 직계 자식인 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `height` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

`span` 요소의 `span` 값이 `class`으로 시작하는 직계 자식인 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `width` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

`span` 요소의 `span` 값이 `class`으로 시작하는 직계 자식인 `rate` 요소를 대상으로 하는 속성 선택자를 사용하여 `background-color` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

`span` 요소의 `one` 값에 `class`라는 단어가 포함된 첫 번째 자손을 대상으로 하는 속성 선택자가 있어야 합니다.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

`span` 요소의 `one` 값에 `class`라는 단어가 포함된 첫 번째 자손을 대상으로 하는 속성 선택자를 사용하여 `background-image` 속성을 `linear-gradient`로 설정해야 합니다.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

`span` 요소의 `two` 값에 `class`라는 단어가 포함된 첫 번째와 두 번째 자손을 대상으로 하는 속성 선택자가 있어야 합니다.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

`span` 요소의 `two` 값에 `class`라는 단어가 포함된 첫 번째와 두 번째 자손을 대상으로 하는 속성 선택자를 사용하여 `background-image` 속성을 `linear-gradient`로 설정해야 합니다.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

`span` 요소의 `span` 값에 `three`라는 단어가 포함된 자손인 `class` 요소를 대상으로 하는 속성 선택자가 있어야 합니다.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

`span` 요소의 `span` 값에 `three`라는 단어가 포함된 자손인 `class` 요소를 대상으로 하는 속성 선택자를 사용하여 `background-image` 속성을 `linear-gradient`로 설정해야 합니다.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
