---
id: 5a24c314108439a4d4036147
title: Collegare Redux a React
challengeType: 6
forumTopicId: 301426
dashedName: connect-redux-to-react
---

# --description--

Ora che hai scritto sia le funzioni `mapStateToProps()` che `mapDispatchToProps()`, puoi usarle per mappare `state` e `dispatch` al `props` di uno dei tuoi componenti React. Il metodo `connect` di React Redux può gestire questo compito. Questo metodo prende due argomenti opzionali, `mapStateToProps()` e `mapDispatchToProps()`. Sono opzionali perché potresti avere un componente che ha bisogno solo di consultare `state` ma non deve dispatchare azioni, o viceversa.

Per usare questo metodo, passa le funzioni come argomenti e chiama immediatamente il risultato con il tuo componente. Questa sintassi è un po' insolita e appare così:

```js
connect(mapStateToProps, mapDispatchToProps)(MyComponent)
```

**Nota:** Se vuoi omettere uno degli argomenti del metodo `connect`, passa `null` al suo posto.

# --instructions--

L'editor del codice ha le funzioni `mapStateToProps()` e `mapDispatchToProps()` e un nuovo componente React chiamato `Presentational`. Collega questo componente a Redux con il metodo `connect` dell'oggetto globale `ReactRedux`, e chiamalo immediatamente sul componente `Presentational`. Assegna il risultato a una nuova variabile `const` chiamata `ConnectedComponent` che rappresenta il componente collegato. È tutto, ora sei collegato a Redux! Prova a cambiare uno qualsiasi dei due argomenti di `connect` in `null` e osserva i risultati dei test.

# --hints--

Il componente `Presentational` dovrebbe renderizzare.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('Presentational').length === 1;
  })()
);
```

Il componente `Presentational` dovrebbe ricevere una prop `messages` tramite `connect`.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return props.messages === '__INITIAL__STATE__';
  })()
);
```

Il componente `Presentational` dovrebbe ricevere una prop `submitNewMessage` tramite `connect`.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return typeof props.submitNewMessage === 'function';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line

const ConnectedComponent = connect(mapStateToProps, mapDispatchToProps)(Presentational);
```
