---
id: 5a24c314108439a4d4036144
title: Usa Provider per collegare Redux a React
challengeType: 6
forumTopicId: 301435
dashedName: use-provider-to-connect-redux-to-react
---

# --description--

Nell'ultima sfida, hai creato un archivio Redux per gestire l'array dei messaggi e creato un'azione per aggiungere nuovi messaggi. Il passaggio successivo è fornire a React la possibilità di consultare l'archivio Redux e le azioni che deve inviare per aggiornare lo stato. React Redux offre il suo pacchetto `react-redux` per aiutarti a svolgere questi compiti.

React Redux fornisce una piccola API con due funzionalità chiave: `Provider` e `connect`. Un'altra sfida tratta `connect`. Il `Provider` è un componente wrapper di React Redux che avvolge la tua app React. Questo wrapper ti permette di consultare l'archivio Redux e le funzioni `store` e `dispatch` in tutto l'albero dei componenti. `Provider` prende due props, l'archivio Redux e i componenti figli della tua app. Definire il `Provider` per un componente App potrebbe essere così:

```jsx
<Provider store={store}>
  <App/>
</Provider>
```

# --instructions--

L'editor del codice mostra ora tutto il tuo codice Redux e React delle ultime sfide. Include l'archivio Redux, le azioni e il componente `DisplayMessages`. L'unica novità è il componente `AppWrapper` in fondo. Usa questo componente di livello superiore per renderizzare il `Provider` da `ReactRedux`, e passa l'archivio Redux come prop. Poi renderizza il componente `DisplayMessages` come figlio. Quando hai finito, dovresti vedere il tuo componente React renderizzato nella pagina.

**Nota:** React Redux è disponibile qui come variabile globale, quindi puoi consultare Provider con la notazione a punto. Il codice nell'editor sfrutta questo e lo assegna a una costante `Provider` che puoi usare nel metodo render di `AppWrapper`.

# --hints--

Il `AppWrapper` dovrebbe essere renderizzato.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('AppWrapper').length === 1;
  })()
);
```

Il componente wrapper `Provider` dovrebbe avere una prop `store` passata, uguale all'archivio Redux.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
      return __helpers
        .removeWhiteSpace(code)
        .includes('<Providerstore={store}>');
    })()
  );
```

`DisplayMessages` dovrebbe essere renderizzato come figlio di `AppWrapper`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('AppWrapper').find('DisplayMessages').length === 1
    );
  })()
);
```

Il componente `DisplayMessages` dovrebbe renderizzare un elemento `h2`, `input`, `button` e `ul`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h2').length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('ul').length === 1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};



const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
    this.handleChange = this.handleChange.bind(this);
    this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {  
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Render the Provider below this line

  // Change code above this line
};
```

# --solutions--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
 this.handleChange = this.handleChange.bind(this);
 this.submitMessage = this.submitMessage.bind(this);
 }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };  
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Change code below this line
  render() {
    return (
      <Provider store = {store}>
        <DisplayMessages/>
      </Provider>
    );
  }
  // Change code above this line
};
```
