---
id: 5a24c314108439a4d403614a
title: Andare avanti da qui
challengeType: 6
forumTopicId: 301434
dashedName: moving-forward-from-here
---

# --description--

Congratulazioni! Hai completato le lezioni su React e Redux. C'è un ultimo aspetto importante da sottolineare prima di andare avanti. Di solito, non scriverai app React in un editor di codice come questo. Questa sfida ti offre un'anteprima di come appare la sintassi se lavori con un file system sulla tua macchina. Il codice dovrebbe essere simile, a eccezione dell'uso delle istruzioni `import` (che importano tutte le dipendenze fornite nelle sfide).

Infine, scrivere codice React e Redux richiede generalmente un po' di configurazione. Questo può diventare rapidamente complicato. Se vuoi sperimentare sulla tua macchina, <a href="https://www.freecodecamp.org/news/install-react-with-create-react-app/" target="_blank" rel="noopener noreferrer nofollow">Create React App</a> è già configurato e pronto all'uso.

In alternativa, puoi abilitare Babel come preprocessore JavaScript in CodePen, aggiungere React e ReactDOM come risorse JavaScript esterne e lavorare anche lì.

# --instructions--

Registra il messaggio `'Now I know React and Redux!'` nella console.

# --hints--

Il messaggio `Now I know React and Redux!` dovrebbe essere registrato nella console.

```js
assert(
    /console\s*\.\s*log\s*\(\s*('|"|`)Now I know React and Redux!\1\s*\)/.test(
      code)
  );
```

# --seed--

## --seed-contents--

```jsx
/*
import React from 'react'
import ReactDOM from 'react-dom'
import { Provider, connect } from 'react-redux'
import { createStore, combineReducers, applyMiddleware } from 'redux'
import thunk from 'redux-thunk'

import rootReducer from './redux/reducers'
import App from './components/App'

const store = createStore(
  rootReducer,
  applyMiddleware(thunk)
);

ReactDOM.render(
  <Provider store={store}>
    <App/>
  </Provider>,
  document.getElementById('root')
);
*/

// Only change code below this line
```

# --solutions--

```jsx
console.log('Now I know React and Redux!');
```
