---
id: 5a24c314108439a4d4036143
title: estrarre la logica dello stato in Redux
challengeType: 6
forumTopicId: 301429
dashedName: extract-state-logic-to-redux
---

# --description--

Ora che hai finito il componente React, devi spostare la logica che esegue localmente nel suo `state` in Redux. Questo è il primo passaggio per collegare la semplice app React a Redux. L'unica funzionalità della tua app è aggiungere nuovi messaggi dall'utente a una lista non ordinata. L'esempio è semplice per mostrare come React e Redux lavorano insieme.

# --instructions--

Per prima cosa, definisci un tipo di azione `ADD` e impostalo su una costante `ADD`. Poi, definisci un creatore di azioni `addMessage()` che crea l'azione per aggiungere un messaggio. Dovrai passare un `message` a questo creatore di azioni e includere il messaggio nell'oggetto `action` restituito.

Quindi crea un reducer chiamato `messageReducer()` che gestisce lo stato dei messaggi. Lo stato iniziale dovrebbe essere un array vuoto. Questo reducer dovrebbe aggiungere un messaggio all'array di messaggi contenuti nello stato, oppure restituire lo stato corrente. Infine, crea il tuo archivio Redux e passagli il reducer.

# --hints--

La costante `ADD` dovrebbe esistere e contenere un valore uguale alla stringa `ADD`

```js
assert(ADD === 'ADD');
```

Il creatore di azioni `addMessage` dovrebbe restituire un oggetto con `type` uguale a `ADD` e `message` uguale al messaggio passato.

```js
assert(
  (function () {
    const addAction = addMessage('__TEST__MESSAGE__');
    return addAction.type === ADD && addAction.message === '__TEST__MESSAGE__';
  })()
);
```

`messageReducer` dovrebbe essere una funzione.

```js
assert(typeof messageReducer === 'function');
```

L'archivio dovrebbe esistere e avere uno stato iniziale impostato a un array vuoto.

```js
assert(
  (function () {
    const initialState = store.getState();
    return typeof store === 'object' && initialState.length === 0;
  })()
);
```

Eseguire `addMessage` sull'archivio dovrebbe aggiungere in modo immutabile un nuovo messaggio all'array di messaggi contenuti nello stato.

```js
assert(
  (function () {
    const initialState = store.getState();
    const isFrozen = DeepFreeze(initialState);
    store.dispatch(addMessage('__A__TEST__MESSAGE'));
    const addState = store.getState();
    return isFrozen && addState[0] === '__A__TEST__MESSAGE';
  })()
);
```

`messageReducer` dovrebbe restituire lo stato corrente se chiamato con altre azioni.

```js
assert(
  (function () {
    const addState = store.getState();
    store.dispatch({ type: 'FAKE_ACTION' });
    const testState = store.getState();
    return addState === testState;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Define ADD, addMessage(), messageReducer(), and store here:
```

# --solutions--

```jsx
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);
```
