---
id: 587d8251367417b2b2512c62
title: 連結リストクラスを作成する
challengeType: 1
forumTopicId: 301628
dashedName: create-a-linked-list-class
---

# --description--

`linked list` クラスを作成しましょう。すべての連結リストは基本的なプロパティから始まります：`head`（リストの最初の項目）と `length`（リスト内の項目数）です。連結リストの実装でリストの最後の要素のための `tail` を含むものもありますが、今回はこの2つだけにします。連結リストに要素を追加するたびに、`length` プロパティは1ずつ増やす必要があります。

連結リストに項目を追加する方法が必要なので、最初に作成したいメソッドは `add` メソッドです。

リストが空の場合、連結リストに要素を追加するのは簡単です：その要素を `Node` クラスでラップし、そのノードを連結リストの `head` に代入します。

しかし、リストにすでに1つ以上のメンバーがいる場合はどうでしょう？要素をリストに追加するにはどうすればよいでしょう？連結リストの各ノードには `next` プロパティがあることを思い出してください。ノードをリストに追加するには、リストの最後のノードを見つけ、その最後のノードの `next` プロパティを新しいノードに向けます。（ヒント：ノードの `next` プロパティが `null` のとき、連結リストの終わりに到達したことがわかります。）

# --instructions--

add メソッドを書いて、連結リストに最初に追加したノードを `head` に代入してください。その後は、ノードを追加するたびに、すべてのノードは前のノードの `next` プロパティで参照されるようにしてください。

注意

連結リストに要素を追加するたびに、リストの `length` は1ずつ増やしてください。

# --hints--

`LinkedList` クラスには `add` メソッドが必要です。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.add === 'function';
  })()
);
```

`LinkedList` クラスは最初に追加されたノードを `head` に代入する必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    return test.head().element === 'cat';
  })()
);
```

`node` クラスの前の `LinkedList` は、作成された最新のノードを参照している必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('fish');
    return test.head().next.element === 'dog' && test.head().next.next.element === 'fish';
  })()
);
```

`size` クラスの `LinkedList` は、連結リスト内のノード数と等しくなります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.size() === 2;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() { 
  var length = 0; 
  var head = null; 

  var Node = function(element){
    this.element = element; 
    this.next = null; 
  }; 

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line
    if (head == null) {
      head = new Node(element);
    } 
    else {
      let currentNode = head;
      while (currentNode.next != null) {
        // currentNode.next will be last node of linked list after loop
        currentNode = currentNode.next;
      }
      currentNode.next = new Node(element);
    }
    length++;
    // Only change code above this line
  };
}
```
