---
id: 587d8255367417b2b2512c72
title: ES6 Set の .has と .size を使う
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

ES6 Set オブジェクトで使える .has と .size メソッドを見てみましょう。

まず、ES6 Set を作成します。

```js
var set = new Set([1,2,3]);
```

.has メソッドは、その値がセットに含まれているかどうかをチェックします。

```js
var hasTwo = set.has(2);
```

.size メソッドは、Set のサイズを表す整数を返します。

```js
var howBig = set.size;
```

# --instructions--

この演習では、配列と値を checkSet() 関数に渡します。関数は配列の引数から ES6 Set を作成してください。セットに値の引数が含まれているかどうかを調べてください。セットのサイズも調べてください。そして、その2つの値を配列で返してください。

# --hints--

`checkSet([4, 5, 6], 3)` は [ false, 3 ] を返すべきです。

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
