---
id: 587d825c367417b2b2512c90
title: 幅優先探索
challengeType: 1
forumTopicId: 301622
dashedName: breadth-first-search
---

# --description--

これまでに、グラフの表現方法をいくつか学びました。では次に何をするでしょうか？自然な疑問は、グラフ内の任意の2つのノード間の距離はどれくらいかということです。ここで登場するのが<dfn>グラフトラバーサルアルゴリズム</dfn>です。

<dfn>トラバーサルアルゴリズム</dfn>とは、グラフ内のノードを巡回または訪問するためのアルゴリズムです。その一種が幅優先探索アルゴリズムです。

このアルゴリズムは1つのノードから始め、そのノードから1エッジ離れたすべての隣接ノードを訪問します。次にそれらの隣接ノードの隣接ノードを訪問し、すべてのノードに到達するまで続けます。

幅優先探索アルゴリズムの実装に役立つ重要なデータ構造はキューです。これは一端に要素を追加し、反対の端から要素を取り出せる配列です。これは<dfn>FIFO</dfn>または<dfn>先入れ先出し</dfn>のデータ構造としても知られています。

視覚的には、アルゴリズムは次のように動作します：
<img alt="幅優先探索アルゴリズムを示すアニメーション" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/breadth-first-search.gif'>

灰色のシェーディングはノードがキューに追加されることを表し、黒色のシェーディングはノードがキューから取り出されることを表します。ノードがキューから取り出されるたびに（ノードが黒くなる）、その隣接ノードがすべてキューに追加される（ノードが灰色になる）様子がわかります。

このアルゴリズムを実装するには、グラフ構造と開始したいノードを入力として受け取る必要があります。

まず、開始ノードからの距離、つまりエッジの数を把握しておく必要があります。すべての距離は`Infinity`のような大きな数で初期化してください。これは開始ノードから到達できないノードがある場合のカウントの問題を防ぎます。次に、開始ノードから隣接ノードへ進みます。これらの隣接ノードは1エッジ離れているため、追跡している距離に1単位を加えます。

# --instructions--

隣接行列グラフ（二次元配列）とノードのラベルrootをパラメータに取る関数`bfs()`を書いてください。ノードのラベルは`0`と`n - 1`の間の整数値で、`n`はグラフ内のノードの総数です。

関数は、ノードとrootからの距離をキーと値のペアで持つJavaScriptオブジェクトを出力します。ノードに到達できなかった場合は距離を`Infinity`にしてください。

# --hints--

入力グラフ`[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`で開始ノードが`1`の場合、`{0: 1, 1: 0, 2: 1, 3: 2}`を返すはずです

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 1);
assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: 2 });
```

入力グラフ`[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]`で開始ノードが`1`の場合、`{0: 1, 1: 0, 2: 1, 3: Infinity}`を返すはずです

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 0],
  [0, 0, 0, 0]
];
var results = bfs(graph, 1);
 assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: Infinity });
```

入力グラフ`[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`で開始ノードが`0`の場合、`{0: 0, 1: 1, 2: 2, 3: 3}`を返すはずです

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 0);
 assert.deepEqual(results, { 0: 0, 1: 1, 2: 2, 3: 3 });
```

入力グラフ`[[0, 1], [1, 0]]`で開始ノードが`0`の場合、`{0: 0, 1: 1}`を返すはずです

```js
var graph = [
  [0, 1],
  [1, 0]
];
var results = bfs(graph, 0);
assert.deepEqual(results, { 0: 0, 1: 1 });
```

# --seed--

## --seed-contents--

```js
function bfs(graph, root) {
  var nodesLen = {};

  return nodesLen;
};

var exBFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(bfs(exBFSGraph, 3));
```

# --solutions--

```js
function bfs(graph, root) {
  var nodesLen = {};
  // Set all distances to infinity
  for (var i = 0; i < graph.length; i++) {
    nodesLen[i] = Infinity;
  }
  nodesLen[root] = 0; // ...except root node
  var queue = [root]; // Keep track of nodes to visit
  var current; // Current node traversing
  // Keep on going until no more nodes to traverse
  while (queue.length !== 0) {
    current = queue.shift();
    // Get adjacent nodes from current node
    var curConnected = graph[current]; // Get layer of edges from current
    var neighborIdx = []; // List of nodes with edges
    var idx = curConnected.indexOf(1); // Get first edge connection
    while (idx !== -1) {
      neighborIdx.push(idx); // Add to list of neighbors
      idx = curConnected.indexOf(1, idx + 1); // Keep on searching
    }
    // Loop through neighbors and get lengths
    for (var j = 0; j < neighborIdx.length; j++) {
      // Increment distance for nodes traversed
      if (nodesLen[neighborIdx[j]] === Infinity) {
        nodesLen[neighborIdx[j]] = nodesLen[current] + 1;
        queue.push(neighborIdx[j]); // Add new neighbors to queue
      }
    }
  }
  return nodesLen;
}
```
