---
id: 587d8254367417b2b2512c70
title: ES6でセットを作成して追加する
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--

ES5を学習したので、今度はES6で同様のことを行います。こちらの方がかなり簡単です。ES6には組み込みのデータ構造`Set`があり、これまで手作業で書いていた多くの操作がすでに用意されています。見てみましょう。

新しい空のセットを作成するには：

```js
var set = new Set();
```

値を持つセットを作成することもできます：

```js
var set = new Set(1);
```

配列を使ってセットを作成することもできます：

```js
var set = new Set([1, 2, 3]);
```

セットを作成したら、`add`メソッドを使って追加したい値を加えられます：

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

セットは重複する値を含められないデータ構造であることを思い出してください：

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set contains [1, 2, 3] only
```

# --instructions--

この演習では、次の値を持つセットを返してください：`1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--

`Set`には`1, 2, 3, Taco, Cat, Awesome`の値だけを含めるようにしてください。

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```
