---
id: 587d825d367417b2b2512c96
title: 深さ優先探索
challengeType: 1
forumTopicId: 301640
dashedName: depth-first-search
---

# --description--

<dfn>幅優先探索</dfn>と似ていますが、ここでは<dfn>深さ優先探索</dfn>と呼ばれる別のグラフ探索アルゴリズムについて学びます。

幅優先探索がソースノードからの辺の長さを段階的に探索するのに対し、<dfn>深さ優先探索</dfn>はまずできるだけ深く辺の経路をたどります。

経路の端に到達すると、未訪問の辺のある最後のノードまで戻り、探索を続けます。

下のアニメーションはアルゴリズムの動作を示しています。アルゴリズムは上のノードから始まり、番号順にノードを訪問します。

<img alt="animation showing the depth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/depth-first-search.gif'>

幅優先探索とは異なり、ノードを訪問するたびにすべての隣接ノードを訪問するのではなく、まず1つの隣接ノードを訪問し、その経路をたどって訪問可能なノードがなくなるまで進みます。

このアルゴリズムを実装するには、スタックを使うのがお勧めです。スタックは最後に追加された要素が最初に取り出される配列で、<dfn>後入れ先出し</dfn>のデータ構造として知られています。スタックは深さ優先探索で役立ちます。なぜなら、隣接ノードをスタックに追加するとき、最新に追加された隣接ノードから先に訪問し、スタックから取り出したいからです。

このアルゴリズムの単純な出力は、あるノードから到達可能なノードのリストです。したがって、訪問したノードを記録しておく必要があります。

# --instructions--

無向の隣接行列 `dfs()` とノードのラベル `graph` をパラメータに取る関数 `root` を書いてください。ノードのラベルは `0` と `n - 1` の間の数値で、`n` はグラフ内のノードの総数です。

関数は `root` から到達可能なすべてのノードを配列で返してください。

# --hints--

入力グラフ `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` で開始ノードが `1` の場合、配列に `0`、`1`、`2`、`3` が含まれるはずです。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })(),
  [0, 1, 2, 3]
);
```

入力グラフ `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` で開始ノードが `3` の場合、配列に `3`、`2`、`1`、`0` が含まれるはずです。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3, 2, 1, 0]
);
```

入力グラフ `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` で開始ノードが `1` の場合、4つの要素を含む配列が返されるはずです。

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })().length === 4
);
```

入力グラフ `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` で開始ノードが `3` の場合、配列に `3` が含まれるはずです。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3]
);
```

入力グラフ `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` で開始ノードが `3` の場合、1つの要素を含む配列が返されるはずです。

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })().length === 1
);
```

入力グラフ `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` で開始ノードが `3` の場合、配列に `2` と `3` が含まれるはずです。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [2, 3]
);
```

入力グラフ `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` で開始ノードが `3` の場合、2つの要素を含む配列が返されるはずです。

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })().length === 2
);
```

入力グラフ `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` で開始ノードが `0` の場合、配列に `0` と `1` が含まれるはずです。

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })(),
  [0, 1]
);
```

入力グラフ `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` で開始ノードが `0` の場合、2つの要素を含む配列が返されるはずです。

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })().length === 2
);
```

# --seed--

## --seed-contents--

```js
function dfs(graph, root) {

}

var exDFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(dfs(exDFSGraph, 3));
```

# --solutions--

```js
function dfs(graph, root) {
    var stack = [];
    var tempV;
    var visited = [];
    var tempVNeighbors = [];
    stack.push(root);
    while (stack.length > 0) {
        tempV = stack.pop();
        if (visited.indexOf(tempV) == -1) {
            visited.push(tempV);
            tempVNeighbors = graph[tempV];
            for (var i = 0; i < tempVNeighbors.length; i++) {
                if (tempVNeighbors[i] == 1) {
                    stack.push(i);
                }
            }
        }
    }
    return visited;
}
```
