---
id: 587d8252367417b2b2512c67
title: 連結リストの特定のインデックスに要素を追加する
challengeType: 1
forumTopicId: 301619
dashedName: add-elements-at-a-specific-index-in-a-linked-list
---

# --description--

addAt(index,element) メソッドを作成して、指定されたインデックスに要素を追加しましょう。指定されたインデックスの要素を削除する方法と同様に、連結リストをたどる際に currentIndex を追跡する必要があります。currentIndex が指定されたインデックスと一致したとき、前のノードの next プロパティを新しく追加されたノードを参照するように再代入する必要があります。そして新しいノードは currentIndex の次のノードを参照するべきです。コンガラインの例に戻ると、新しい人が列に加わりたいのですが、真ん中に入りたいのです。あなたは列の真ん中にいるので、前の人の手を離します。新しい人がやってきて、あなたがかつて手をつないでいた人に手を置き、あなたは新しい人に手を置きます。

# --instructions--

指定されたインデックスに要素を追加する `addAt(index,element)` メソッドを作成してください。要素を追加できなかった場合は false を返してください。**注意：** 指定されたインデックスが負の値か連結リストの長さより大きいかどうかを必ず確認してください。

# --hints--

`addAt` メソッドは、指定されたインデックスが 0 の場合、新しいノードに `head` を再代入する必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'fish');
    return test.head().element === 'fish' && test.head().next.element === 'cat';
  })()
);
```

`addAt` メソッドは、連結リストに新しいノードが追加されるたびに長さを 1 増やす必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'cat');
    return test.size() === 3;
  })()
);
```

`addAt` メソッドは、ノードを追加できなかった場合に `false` を返す必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.addAt(4, 'cat') === false;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element) {
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function() {
    return head;
  };

  this.add = function(element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if (head === null){
        head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };
  this.addAt = function (index, element) {
    if (index > length || index < 0) {
      return false;
    }
    var newNode = new Node(element);
    var currentNode = head;
    if (index === 0) {
      head = newNode;
    } else {
      var previousNode = null;
      var i = 0;
      while (currentNode && i < index) {
        previousNode = currentNode;
        currentNode = currentNode.next;
        i++;
      }
      previousNode.next = newNode;
    }
    newNode.next = currentNode;
    length++;
  }
}
```
