---
id: 587d825b367417b2b2512c8c
title: 最小ヒープでヒープソートを実装する
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--

要素の追加と削除ができるようになったので、ヒープが使われるいくつかの応用例を見てみましょう。ヒープは常に最大または最小の値を持つ要素を先頭にストアするため、優先キューの実装によく使われます。さらに、ヒープソートというソートアルゴリズムの実装にも使われます。ここではその方法を見ていきます。ヒープソートは最大ヒープの逆である最小ヒープを使います。最小ヒープは常に最小の値を持つ要素をルート位置にストアします。

ヒープソートは、未ソートの配列を取り、その配列の各要素を最小ヒープに追加し、次に最小ヒープからすべての要素を新しい配列に取り出すことで動作します。最小ヒープの構造により、新しい配列には元の要素が小さい順に並びます。これは平均および最悪の場合の性能がO(nlog(n))である、最も効率的なソートアルゴリズムの一つです。

# --instructions--

最小ヒープでヒープソートを実装しましょう。ここでは最大ヒープのコードを適宜流用してかまいません。`MinHeap`オブジェクトを作成し、`insert`、`remove`、`sort`のメソッドを持たせてください。`sort`メソッドは、最小ヒープ内のすべての要素を小さい順に並べた配列を返すようにしてください。

# --hints--

`MinHeap`データ構造が存在している必要があります。

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MinHeap`は`insert`というメソッドを持つべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MinHeap`は`remove`というメソッドを持つべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`MinHeap`は`sort`というメソッドを持つべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

`sort`メソッドは、最小ヒープに追加されたすべての要素をソート済みの配列として返す必要があります。

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
