---
id: 587d825b367417b2b2512c8d
title: ES6 JavaScript Map を作成する
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

JavaScript の新しいバージョンでは、組み込みの Map オブジェクトが提供されており、前のチャレンジで手作業で書いた多くの機能を備えています。この Map オブジェクトは通常の JavaScript オブジェクトに似ていますが、通常のオブジェクトにはない便利な機能を提供します。たとえば、ES6 の Map は追加されたアイテムの挿入順序を追跡します。以下はそのメソッドのより完全な概要です：

- `.has(key)` はキーの存在に基づいて true または false を返します
- `.get(key)` はキーに関連付けられた値を返します
- `.set(key, value)` は新しいキーと値のペアを設定します
- `.delete(key)` はキーと値のペアを削除します
- `.clear()` はすべてのキーと値のペアを削除します
- `.keys()` は挿入順にすべてのキーを含む新しいマップイテレーターオブジェクトを返します
- `.values()` は挿入順にすべての値を含む新しいマップイテレーターオブジェクトを返します
- `.entries()` は挿入順にすべてのキーと値のペアを含む新しいマップイテレーターオブジェクトを返します

# --instructions--

JavaScript の Map オブジェクトを定義し、それを myMap という変数に代入してください。キーと値のペア `freeCodeCamp`、`Awesome!` を追加してください。

# --hints--

`myMap` オブジェクトが存在しているはずです。

```js
assert(typeof myMap === 'object');
```

`myMap` はキーと値のペア `freeCodeCamp`、`Awesome!` を含んでいるはずです。

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
