---
id: 587d8250367417b2b2512c5e
title: スタックの仕組みを学ぼう
challengeType: 1
forumTopicId: 301705
dashedName: learn-how-a-stack-works
---

# --description--

テーブルの上に積まれた本の山には馴染みがあると思います。テキストエディターの元に戻す機能を使ったこともあるでしょう。スマホのアプリで戻るボタンを押して前の画面に戻ることにも慣れているはずです。

これらに共通していることは何でしょうか？どれもデータを後ろ向きにたどれるように保存しています。

スタックの一番上にある本は最後に置かれたものです。その本を取り除くと、その前に置かれた本が現れます。これが繰り返されます。

よく考えると、これらの例はすべて「後入れ先出し（Last-In-First-Out）」の仕組みを使っています。これをコードで再現してみましょう。

このデータ保存の仕組みは「スタック」と呼ばれます。特に、スタックの一番上にJavaScriptオブジェクトを追加する`push()`メソッドと、現在スタックの一番上にあるJavaScriptオブジェクトを取り除く`pop()`メソッドを実装する必要があります。

# --instructions--

ここに配列で表された宿題のスタックがあります。`"BIO12"`が底にあり、`"PSY44"`がスタックの一番上にあります。

与えられた配列を修正し、上記のJavaScriptメソッドを使って`stack`のように扱ってください。スタックの一番上の要素`"PSY44"`を取り除きます。その後、`"CS50"`を新しい一番上の要素として追加してください。

# --hints--

`homeworkStack`には4つの要素だけが含まれている必要があります。

```js
assert(homeworkStack.length === 4);
```

`homeworkStack`の最後の要素は`"CS50"`であるべきです。

```js
assert(homeworkStack[3] === 'CS50');
```

`homeworkStack`には`"PSY44"`が含まれていてはいけません。

```js
assert(homeworkStack.indexOf('PSY44') === -1);
```

`homeworkStack`の最初の宣言は変更しないでください。

```js
assert(
  __helpers.removeJSComments(code).match(/=/g).length === 1 &&
    /homeworkStack\s*=\s*\["BIO12"\s*,\s*"HIS80"\s*,\s*"MAT122"\s*,\s*"PSY44"\]/.test(
      __helpers.removeJSComments(code)
    )
);
```

# --seed--

## --seed-contents--

```js
var homeworkStack = ["BIO12","HIS80","MAT122","PSY44"];
// Only change code below this line
```

# --solutions--

```js
// solution required
```
