---
id: 587d8255367417b2b2512c73
title: ES5のSet()統合にスプレッド演算子とメモを使う
challengeType: 1
forumTopicId: 301720
dashedName: use-spread-and-notes-for-es5-set-integration
---

# --description--

ES6のスプレッド演算子`...`を覚えていますか？

`...`はES6でイテラブルなオブジェクトを配列に変換できます。

Setを作成して、スプレッド関数を試してみましょう。

```js
var set = new Set([1,2,3]);
var setToArr = [...set]
console.log(setToArr) // returns [ 1, 2, 3 ]
```

# --instructions--

このチャレンジでは、Setオブジェクトを`checkSet`関数に渡します。Setの値を含む配列を返す必要があります。

これでES6の`Set()`オブジェクトの使い方をマスターできました。よくできました！

# --hints--

`checkSet(new Set([1,2,3,4,5,6,7])`は`[1, 2, 3, 4, 5, 6, 7]`を返すべきです。

```js
var test = checkSet(new Set([1, 2, 3, 4, 5, 6, 7]));
assert.deepEqual(test, [1, 2, 3, 4, 5, 6, 7]);
```

# --seed--

## --seed-contents--

```js
function checkSet(set){
   // Only change code below this line

   // Only change code above this line
}
```

# --solutions--

```js
function checkSet(set){
return [...set];}
```
