---
id: 587d8255367417b2b2512c74
title: 優先キューのクラスを作成する
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--

このチャレンジでは優先キューを作成します。優先キューは特別なタイプのキューで、アイテムに優先度を指定する追加情報が付く場合があります。これは単純に整数で表すことができます。アイテムの優先度は、アイテムが取り出される順序を決める際に配置順序よりも優先されます。もし優先度の高いアイテムが優先度の低いアイテムの後にキューに追加された場合、その優先度の高いアイテムが他のすべてのアイテムより先に取り出されます。

例えば、3つのアイテムがある優先キューを考えてみましょう：

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

ここで2番目の値（整数）がアイテムの優先度を表しています。もし`['human', 1]`を優先度`1`でキューに追加した場合（優先度が低いほど優先されると仮定すると）、それは最初に取り出されるアイテムになります。コレクションは次のようになります：

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

コードエディターには`PriorityQueue`の作成を開始しています。優先度付きでアイテムを追加する`enqueue`メソッド、アイテムを取り出して返す`dequeue`メソッド、キュー内のアイテム数を返す`size`メソッド、キューの先頭の要素を返す`front`メソッド、そして最後にキューが空の場合は`isEmpty`を、そうでなければ`true`を返す`false`メソッドを追加する必要があります。

`enqueue`は上記の形式（`['human', 1]`）のアイテムを受け入れる必要があり、`1`が優先度を表します。`dequeue`と`front`はアイテムの名前のみを返し、優先度は返しません。

# --hints--

あなたの`PriorityQueue`クラスには`enqueue`メソッドが必要です。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

あなたの`PriorityQueue`クラスには`dequeue`メソッドが必要です。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

あなたの`PriorityQueue`クラスには`size`メソッドが必要です。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

あなたの`PriorityQueue`クラスには`front`メソッドが必要です。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

あなたの`PriorityQueue`クラスには`isEmpty`メソッドが必要です。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

あなたの`PriorityQueue`クラスは、アイテムがキューに追加および取り出される際に、`size`メソッドを使って現在のアイテム数を正しく管理する必要があります。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

`front`メソッドは、アイテムがキューに追加および取り出される際に、キューの先頭の正しいアイテムを返す必要があります。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

`isEmpty`メソッドは、キューが空の場合に`true`を返す必要があります。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

優先キューは、優先度の高いアイテムを優先度の低いアイテムより先に返し、それ以外の場合は先入れ先出しの順序でアイテムを返す必要があります。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```
