---
id: 587d8251367417b2b2512c64
title: 連結リスト内の検索
challengeType: 1
forumTopicId: 301715
dashedName: search-within-a-linked-list
---

# --description--

連結リストのクラスに、さらに便利なメソッドをいくつか追加しましょう。`Stack`や`Queue`のクラスのように、リストが空かどうかを判別できたら便利ですよね？

また、連結リスト内の特定の要素を見つけられるようにする必要があります。データ構造をたどる操作は、たくさん練習して慣れておきたいものです！`indexOf`を引数に取り、その要素の`element`を連結リスト内で返す`index`メソッドを作りましょう。もし連結リスト内にその要素が見つからなければ、`-1`を返してください。

逆の操作をするメソッドも実装しましょう。`elementAt`を引数に取り、指定された`index`の位置にある`element`を返す`index`メソッドです。もし`element`が見つからなければ、`undefined`を返してください。

# --instructions--

連結リストが空かどうかをチェックする`isEmpty`メソッド、指定した要素の`indexOf`を返す`index`メソッド、指定された`elementAt`の位置にある`element`を返す`index`メソッドを書いてください。

# --hints--

あなたの`LinkedList`クラスには`isEmpty`メソッドが必要です。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.isEmpty === 'function';
  })()
);
```

`isEmpty`メソッドは、連結リストに少なくとも1つ要素がある場合に`false`を返すべきです。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.isEmpty() === false;
  })()
);
```

`isEmpty`メソッドは、連結リストに要素がない場合に`true`を返すべきです。

```js
assert(
  (function () {
    var test = new LinkedList();
    return test.isEmpty() === true;
  })()
);
```

あなたの`LinkedList`クラスには`indexOf`メソッドが必要です。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.indexOf === 'function';
  })()
);
```

`indexOf`メソッドは、連結リスト内で見つかった指定した要素のインデックスを返すべきです。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('cat') === 0;
  })()
);
```

`indexOf`メソッドは、指定した要素が連結リスト内に見つからなければ`-1`を返すべきです。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('pony') === -1;
  })()
);
```

あなたの`LinkedList`クラスには`elementAt`メソッドが必要です。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.elementAt === 'function';
  })()
);
```

`elementAt`メソッドは、連結リスト内の指定されたインデックスにある要素を返すべきです。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(1) === 'dog';
  })()
);
```

`elementAt`メソッドは、指定されたインデックスに要素が見つからなければ`undefined`を返すべきです。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(5) === undefined;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
      head = currentNode.next;
    } else {
      while(currentNode.element !== element) {
        previousNode = currentNode;
        currentNode = currentNode.next;
      }

      previousNode.next = currentNode.next;
    }

    length --;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
        head = currentNode.next;
    } else {
        while(currentNode.element !== element) {
            previousNode = currentNode;
            currentNode = currentNode.next;
        }

        previousNode.next = currentNode.next;
    }

    length --;
  };

  this.indexOf = function(element) {
    if (head === null) return -1

    let current = head;
    let index = 0;

    while (current.element !== element && current.next !== null) {
      current = current.next;
      index++
    }

    if (current.element !== element && current.next === null) {
      return -1
    }

    return index;
  }

  this.elementAt = function(index) {
    if (head === null) return undefined;

    let current = head;
    let currentIndex = 0;

    while (currentIndex !== index && current.next !== null) {
      current = current.next;
      currentIndex++
    }

    if (currentIndex !== index && current.next === null) {
      return undefined;
    }

    return current.element;
  }

  this.isEmpty = function() {
    return length === 0;
  }
}
```
