---
id: 587d8250367417b2b2512c5f
title: スタッククラスを作成する
challengeType: 1
forumTopicId: 301633
dashedName: create-a-stack-class
---

# --description--

前のセクションでは、スタックとは何か、そして配列を使ってスタックを表す方法について話しました。このセクションでは、自分自身のスタッククラスを作成します。配列を使ってスタックを作ることもできますが、スタックの操作を制限したほうが良い場合もあります。`push` と `pop` メソッド以外にも、スタックには便利なメソッドがあります。スタッククラスに `peek`、`isEmpty`、`clear` メソッドを追加しましょう。

# --instructions--

スタックの一番上に要素を追加する `push` メソッド、スタックの一番上の要素を取り除いて返す `pop` メソッド、スタックの一番上の要素を確認する `peek` メソッド、スタックが空かどうかをチェックする `isEmpty` メソッド、そしてスタックのすべての要素を削除する `clear` メソッドを書いてください。通常スタックにはありませんが、コレクションをコンソールに表示する `print` ヘルパーメソッドも追加しています。

# --hints--

あなたの `Stack` クラスには `push` メソッドが必要です。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.push === 'function';
  })()
);
```

あなたの `Stack` クラスには `pop` メソッドが必要です。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.pop === 'function';
  })()
);
```

あなたの `Stack` クラスには `peek` メソッドが必要です。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.peek === 'function';
  })()
);
```

あなたの `Stack` クラスには `isEmpty` メソッドが必要です。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.isEmpty === 'function';
  })()
);
```

あなたの `Stack` クラスには `clear` メソッドが必要です。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.clear === 'function';
  })()
);
```

`peek` メソッドはスタックの一番上の要素を返す必要があります。

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.peek() === 'CS50' && test.peek() === 'CS50';
  })()
);
```

`pop` メソッドはスタックの一番上の要素を取り除いて返す必要があります。

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.pop() === 'CS50' && test.pop() === 'CS61';
  })()
);
```

`isEmpty` メソッドはスタックに要素が含まれていなければ true を返す必要があります。

```js
assert(
  (function () {
    var test = new Stack();
    return test.isEmpty();
  })()
);
```

`clear` メソッドはスタックのすべての要素を削除する必要があります。

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    test.clear();
    return test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Stack() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
class Stack {
  constructor() {
    this.collection = [];
  }
  print() {
    console.log(this.collection);
  }
  push(val) {
    this.collection.push(val);
  }
  pop() {
    return this.collection.pop();
  }
  peek() {
    return this.collection[this.collection.length - 1];
  }
  isEmpty() {
    return this.collection.length === 0;
  }
  clear() {
    return (this.collection.length = 0);
  }
}
```
