---
id: 587d8257367417b2b2512c7c
title: 二分探索木に要素が存在するか確認する
challengeType: 1
forumTopicId: 301623
dashedName: check-if-an-element-is-present-in-a-binary-search-tree
---

# --description--

二分探索木がどのようなものか大まかに理解できたところで、もう少し詳しく説明しましょう。二分探索木は、平均的なケースでの検索、挿入、削除といった基本操作に対して対数時間を提供し、最悪の場合は線形時間になります。なぜでしょうか？これらの基本操作はすべて、木の中でアイテムを見つける（挿入の場合はどこに入れるべきかを見つける）必要があり、親ノードごとに左か右に分岐して残りの木の半分を効果的に除外していくためです。これにより、探索は木のノード数の対数に比例し、平均的なケースでこれらの操作は対数時間になります。では、最悪の場合はどうでしょう？次の値を左から右へ順に追加して木を構築するとします：`10`、`12`、`17`、`25`。二分探索木のルールに従い、`12`は`10`の右に、`17`はその右に、`25`はさらにその右に追加されます。こうしてできた木は連結リストのようになり、`25`を探すにはすべてのアイテムを線形にたどる必要があります。したがって、最悪の場合は線形時間になります。この問題は木がアンバランスであることに起因します。この意味については、次のチャレンジでさらに詳しく見ていきます。

# --instructions--

このチャレンジでは、木のためのユーティリティを作成します。整数値を入力として受け取り、その値が二分探索木に存在するかどうかを示すブール値を返すメソッド`isPresent`を書いてください。

# --hints--

`BinarySearchTree`というデータ構造が存在しているはずです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

二分探索木には`isPresent`というメソッドがあるべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isPresent == 'function';
  })()
);
```

`isPresent`メソッドは、木に追加された要素の存在・不在を正しくチェックする必要があります。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    test.add(4);
    test.add(7);
    test.add(411);
    test.add(452);
    return (
      test.isPresent(452) &&
      test.isPresent(411) &&
      test.isPresent(7) &&
      !test.isPresent(100)
    );
  })()
);
```

`isPresent`は木が空の場合も処理できるようにしてください。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    return test.isPresent(5) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.isPresent = function (value) {
    var current = this.root
    while (current) {
      if (value === current.value) {
        return true;
      }
      current = value < current.value ? current.left : current.right;
    }
    return false;
  }
}
```
