---
id: 587d8257367417b2b2512c7b
title: 二分探索木に新しい要素を追加する
challengeType: 1
forumTopicId: 301618
dashedName: add-a-new-element-to-a-binary-search-tree
---

# --description--

この一連のチャレンジでは、ツリーというデータ構造を紹介します。ツリーはコンピューターサイエンスで重要かつ多用途なデータ構造です。もちろん、その名前は、視覚化すると自然界の木に非常によく似ていることに由来します。ツリーのデータ構造は通常、ルートと呼ばれる1つのノードから始まり、そこから枝分かれして追加のノードが続き、それぞれがさらに子ノードを持つことがあり、という形で広がっていきます。このデータ構造は通常、ルートノードが上にある状態で視覚化されます。自然の木を逆さにしたように考えることができます。

まず、ツリーでよく使われる用語を説明しましょう。ルートノードはツリーの最上部にあります。ツリー内のデータポイントはノードと呼ばれます。枝が他のノードに繋がっているノードは、その枝が繋がるノードの親ノードと呼ばれ（子ノード）、他にも予想通りの複雑な親族関係の用語があります。サブツリーは特定のノードのすべての子孫要素を指し、枝はエッジとも呼ばれ、葉ノードは子ノードを持たないツリーの末端のノードです。最後に、ツリーは本質的に再帰的なデータ構造であることに注意してください。つまり、ノードの子はそれぞれ自分のサブツリーの親ノードであり、これが繰り返されます。ツリーの再帰的な性質は、一般的なツリー操作のアルゴリズムを設計する際に理解が重要です。

まず、特定の種類のツリーである二分木について説明します。実際には、特定の二分木である二分探索木について説明します。これが何を意味するか説明しましょう。ツリーのデータ構造は1つのノードに任意の数の枝を持つことができますが、二分木は各ノードに2つの枝しか持てません。さらに、二分探索木は子のサブツリーに関して順序付けられており、左のサブツリーの各ノードの値は親ノードの値以下であり、右のサブツリーの各ノードの値は親ノードの値以上です。この関係を視覚化すると理解しやすいです：

<div style='width: 100%; display: flex; justify-content: center; align-items: center;'><img alt="an example of a binary search tree" style='width: 100%; max-width: 350px; background-color: var(--gray-05);' src='https://user-images.githubusercontent.com/18563015/32136009-1e665d98-bbd6-11e7-9133-63184f9f9182.png'></div>

この順序関係は非常にわかりやすいです。ルートノードの8の左側のすべての値は8未満であり、右側のすべての値は8より大きいことに注意してください。また、この関係は各サブツリーにも適用されます。例えば、最初の左の子はサブツリーです。3が親ノードで、正確に2つの子ノードがあります。二分探索木のルールにより、このノードの左の子（およびその子孫）は3未満であり、右の子（およびその子孫）は3より大きく（ただし構造のルート値より小さい）ことがわかります。

二分探索木は非常に一般的で便利なデータ構造であり、探索、挿入、削除などの一般的な操作で平均的に対数時間を提供します。

# --instructions--

まずは簡単に始めましょう。ここでは二分探索木の構造の骨組みと、ツリーのノードを作成する関数を定義しています。各ノードは左と右の値を持つことができます。これらは子のサブツリーが存在すれば代入されます。二分探索木では、新しい値をツリーに追加するメソッドを作成してください。このメソッドは`add`と呼ばれ、ツリーに追加する整数値を受け取ります。二分探索木の不変条件を維持するように注意してください：左の子の値は親の値以下であり、右の子の値は親の値以上であるべきです。ここでは、ツリーに重複した値を持たせないようにしましょう。すでに存在する値を追加しようとした場合、メソッドは`null`を返すべきです。そうでなければ、追加が成功した場合は`undefined`を返してください。

**ヒント：** ツリーは本質的に再帰的なデータ構造です！

# --hints--

`BinarySearchTree`というデータ構造が存在するはずです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

二分探索木は`add`というメソッドを持つべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

追加メソッドは二分探索木のルールに従って要素を追加するべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    const expectedResult = [1, 4, 7, 8, 34, 45, 73, 87];
    const result = test.inOrder();
    return expectedResult.toString() === result.toString();
  })()
);
```

すでに存在する要素を追加しようとした場合は`null`を返すべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    return test.add(4) == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    inOrder() {
      if (!this.root) {
        return null;
      }
      var result = new Array();
      function traverseInOrder(node) {
        node.left && traverseInOrder(node.left);
        result.push(node.value);
        node.right && traverseInOrder(node.right);
      }
      traverseInOrder(this.root);
      return result;
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.add = function(element) {
    let current = this.root;
    if (!current) {
      this.root = new Node(element);
      return;
    } else {
      const searchTree = function(current) {
        if (current.value > element) {
          if (current.left) {
            return searchTree(current.left);
          } else {
            current.left = new Node(element);
            return;
          }
        } else if (current.value < element) {
          if (current.right) {
            return searchTree(current.right);
          } else {
            current.right = new Node(element);
            return;
          }
        } else {
          return null;
        }
      };
      return searchTree(current);
    }
  };
}
```
