---
id: 587d8251367417b2b2512c65
title: インデックスで連結リストから要素を削除する
challengeType: 1
forumTopicId: 301711
dashedName: remove-elements-from-a-linked-list-by-index
---

# --description--

別のデータ構造に進む前に、連結リストで最後の練習をしましょう。

指定された `removeAt` の `element` を削除する `index` メソッドを書きましょう。このメソッドは `removeAt(index)` と呼ばれます。特定の `element` にある `index` を削除するには、連結リストを進みながら各ノードのカウントを保持する必要があります。

連結リストの要素を繰り返し処理する一般的な手法として、<dfn>ランナー</dfn>またはセンチネルと呼ばれるものがあり、これはコードが比較しているノードを「指し示す」役割をします。今回の場合、リストの `head` から始めて、`currentIndex` から始まる `0` 変数を使います。`currentIndex` は通過したノードごとに1ずつ増やします。

以前のレッスンで扱った `remove(element)` メソッドと同様に、`removeAt(index)` メソッドでノードを削除するときにリストの残りの部分が孤立しないように注意が必要です。削除されたノードを参照していたノードが次のノードを参照するようにして、ノードを連続させます。

# --instructions--

指定された `removeAt(index)` のノードを削除して返す `index` メソッドを書いてください。指定された `null` が負の値か、連結リストの `index` 以上の場合は `length` を返す必要があります。

**注意:** `currentIndex` のカウントを忘れずに行ってください。

# --hints--

`LinkedList` クラスには `removeAt` メソッドが必要です。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.removeAt === 'function';
  })()
);
```

`removeAt` メソッドは連結リストの `length` を1減らす必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.removeAt(1);
    return test.size() === 2;
  })()
);
```

`removeAt` メソッドは指定されたインデックスの要素を連結リストから削除する必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.add('bird');
    test.removeAt(1);
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"kitten","next":{"element":"bird","next":null}}}'
    );
  })()
);
```

連結リストに要素が1つだけの場合、`removeAt` メソッドは指定されたインデックスの要素を削除して返し、連結リストの長さを減らす必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    var removedItem = test.removeAt(0);
    return test.head() === null && test.size() === 0 && removedItem === 'cat';
  })()
);
```

`removeAt` メソッドは削除されたノードの要素を返す必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.removeAt(1) === 'dog';
  })()
);
```

`removeAt` メソッドは指定されたインデックスが `null` 未満の場合、`0` を返し、連結リストは変更されない必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(-1);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

`removeAt` メソッドは指定されたインデックスがリストの `null` 以上の場合、`length` を返し、連結リストは変更されない必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(3);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
      head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function (element) {
    this.element = element;
    this.next = null;
  };

  this.size = function () {
    return length;
  };

  this.head = function () {
    return head;
  };

  this.add = function (element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.removeAt = function (index) {
    var currentNode = head;
    var previous = head;
    var count = 0;
    if (index >= length || index < 0) {
      return null;
    }
    if (index === 0) {
      var removed = head.element;
      head = currentNode.next;
    } else {
      while (count < index) {
        previous = currentNode;
        currentNode = currentNode.next;
        count++;
      }
      var removed = previous.next.element;
      previous.next = currentNode.next;
    }
    length--;
    return removed;
  };
}
```
