---
id: 587d8251367417b2b2512c61
title: 連結リストのノードを操作する
challengeType: 1
forumTopicId: 301721
dashedName: work-with-nodes-in-a-linked-list
---

# --description--

コンピューターサイエンスでよく出会うもう一つの一般的なデータ構造は<dfn>連結リスト</dfn>です。連結リストは「ノード」と呼ばれるデータ要素の線形コレクションで、それぞれが次のノードを指しています。連結リストの各<dfn>ノード</dfn>は2つの重要な情報を持っています：`element`自体と、次の`node`への参照です。

コンガラインにいると想像してください。あなたは前の人の手に触れていて、後ろの人はあなたの手に触れています。あなたは前の人を見ることができますが、その人が前の他の人の視界を遮っています。ノードはコンガラインの人と同じで、自分が誰かを知っていて、次の人だけを見ることができますが、前後の他の人は認識していません。

# --instructions--

コードエディターでは、すでに`Kitten`と`Puppy`の2つのノードを作成し、`Kitten`ノードを`Puppy`ノードに手動で接続しています。

`Cat`と`Dog`のノードを作成し、手動でラインに追加してください。

# --hints--

あなたの`Puppy`ノードは`Cat`ノードへの参照を持つべきです。

```js
assert(Puppy.next.element === 'Cat');
```

あなたの`Cat`ノードは`Dog`ノードへの参照を持つべきです。

```js
assert(Cat.next.element === 'Dog');
```

# --seed--

## --seed-contents--

```js
var Node = function(element) {
  this.element = element;
  this.next = null;
};
var Kitten = new Node('Kitten');
var Puppy = new Node('Puppy');

Kitten.next = Puppy;
// Only change code below this line
```

# --solutions--

```js
// solution required
```
