---
id: 587d8254367417b2b2512c71
title: ES6でセットからアイテムを削除する
challengeType: 1
forumTopicId: 301713
dashedName: remove-items-from-a-set-in-es6
---

# --description--

`delete` メソッドを使ってES6のSetからアイテムを削除する練習をしましょう。

まず、ES6のSetを作成します。

```js
var set = new Set([1,2,3]);
```

次に、`delete` メソッドでSetからアイテムを削除してください。

```js
set.delete(1);
console.log([...set]) // should return [ 2, 3 ]
```

# --instructions--

整数の1、2、3、4、5を含むセットを作成してください。

値の2と5を削除し、その後セットを返してください。

# --hints--

セットには値の1、3、4が含まれているはずです。

```js
assert(
  (function () {
    var test = checkSet();
    return test.has(1) && test.has(3) && test.has(4) && test.size === 3;
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet(){
  // Only change code below this line
  var set = null;

  // Only change code above this line
  return set;   
}
```

# --solutions--

```js
function checkSet(){
var set = new Set([1,2,3,4,5]);
set.delete(2);
set.delete(5);
return set;}
```
