---
id: 587d8255367417b2b2512c75
title: 循環キューを作成する
challengeType: 1
forumTopicId: 301625
dashedName: create-a-circular-queue
---

# --description--

このチャレンジでは循環キューを作成します。循環キューとは、コレクションの末尾に書き込みを行い、その後コレクションの先頭から上書きを始めるキューのことです。この種のデータ構造は特定の状況で役立ちます。例えば、循環キューはストリーミングメディアに使えます。キューがいっぱいになると、新しいメディアデータが古いデータを上書きします。

この概念を説明する良い方法は、長さ`5`の配列を使うことです。

```js
[null, null, null, null, null]
 ^Read @ 0
 ^Write @ 0
```

ここで読み取りと書き込みはどちらも位置`0`にあります。次にキューに3つの新しいレコード`a`、`b`、`c`が入ります。キューは次のようになります。

```js
[a, b, c, null, null]
 ^Read @ 0
          ^Write @ 3
```

読み取りヘッドが読み取ると、値を削除するか保持するか選べます。

```js
[null, null, null, null, null]
                   ^Read @ 3
                   ^Write @ 3
```

次に値`d`、`e`、`f`を書き込みます。書き込みが配列の末尾に達すると、先頭に戻ります。

```js
[f, null, null, d, e]
                ^Read @ 3
    ^Write @ 1
```

この方法は一定のメモリ量で済みますが、より大きなサイズのファイルを処理できます。

# --instructions--

このチャレンジでは循環キューを実装します。循環キューは`enqueue`と`dequeue`のメソッドを提供し、キューから読み書きできるようにしてください。クラス自体は整数の引数を受け取り、作成時にキューのサイズを指定できるようにしてください。コードエディターにはこのクラスの開始バージョンを用意しています。

キューにアイテムをエンキューするとき、書き込みポインターは前方に進み、キューの末尾に達したら先頭に戻るようにしてください。`enqueue`メソッドは成功した場合、エンキューしたアイテムを返し、そうでなければ`null`を返します。

同様に、読み取りポインターはデキューするときに前方に進みます。アイテムをデキューすると、そのアイテムを返してください。デキューできない場合は`null`を返してください。

書き込みポインターは読み取りポインターを越えて進めないようにしてください（まだ読み取っていないデータを上書きできないようにします）。読み取りポインターも書き込んだデータを越えて進めません。

# --hints--

`enqueue`メソッドは循環キューにアイテムを追加するためのものです。

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

読み取りポインターを越えてエンキューしてはいけません。

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    test.enqueue(13);
    test.enqueue(25);
    test.enqueue(59);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

`dequeue`メソッドはキューからアイテムをデキューするためのものです。

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 && test.dequeue() === 32 && test.dequeue() === 591
    );
  })()
);
```

アイテムをデキューした後、その位置は`null`にリセットしてください。

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(672);
    test.dequeue();
    test.dequeue();
    var print = test.print();
    return print[0] === null && print[1] === null && print[2] === 672;
  })()
);
```

書き込みポインターを越えてデキューしようとすると、`null`を返し、書き込みポインターは進みません。

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 &&
      test.dequeue() === 32 &&
      test.dequeue() === 591 &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.enqueue(100) === 100 &&
      test.dequeue() === 100
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
class CircularQueue {
  constructor(size) {

    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line

    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line

    // Only change code above this line
  }
}
```

# --solutions--

```js
class CircularQueue {
  constructor(size) {
    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line
    console.log(this.write, this.max);
    if (this.queue[this.write] === null) {
      this.queue[this.write++] = item;

      if (this.write > this.max) {
        this.write = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line
    if (this.queue[this.read] !== null) {
      let item = this.queue[this.read];
      this.queue[this.read++] = null;
      if (this.read > this.max) {
        this.read = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }
}
```
