---
id: 587d825a367417b2b2512c8a
title: 最大ヒープに要素を挿入する
challengeType: 1
forumTopicId: 301703
dashedName: insert-an-element-into-a-max-heap
---

# --description--

次に、別の木構造のデータ構造である二分ヒープに進みます。二分ヒープはヒープ特性を満たす部分的に順序付けられた二分木です。ヒープ特性は親ノードと子ノードの関係を指定します。最大ヒープではすべての親ノードが子ノード以上の値を持ち、最小ヒープではその逆になります。二分ヒープは完全二分木でもあります。これは、木のすべてのレベルが完全に埋まっており、最後のレベルが部分的に埋まっている場合は左から右に埋まっていることを意味します。

二分ヒープは左と右の参照を持つノードを使った木構造として実装される場合もありますが、ヒープ特性に基づく部分的な順序付けにより、配列でヒープを表すことができます。親子関係に注目し、簡単な算術計算で任意の親の子ノードや任意の子ノードの親を計算できます。

例えば、次の配列で表された二分最小ヒープを考えてみましょう。

```js
[ 6, 22, 30, 37, 63, 48, 42, 76 ]
```

ルートノードは最初の要素である `6` です。その子ノードは `22` と `30` です。これらの値の配列インデックスの関係を見ると、インデックス `i` の子ノードは `2 * i + 1` と `2 * i + 2` です。同様に、インデックス `0` の要素はインデックス `1` と `2` の子ノードの親です。一般的に、任意のインデックスのノードの親は次の式で求められます： `Math.floor((i - 1) / 2)` 。これらのパターンは二分木がどんな大きさに成長しても成り立ちます。最後に、この算術をさらに簡単にするために配列の最初の要素をスキップするという調整を加えることができます。これにより、任意のインデックス `i` の要素に対して次の関係が成り立ちます。

例の配列表現：

```js
[ null, 6, 22, 30, 37, 63, 48, 42, 76 ]
```

要素の左の子ノード： `i * 2`

要素の右の子ノード： `i * 2 + 1`

要素の親ノード： `Math.floor(i / 2)`

この算術を理解すれば、配列表現は非常に便利です。ノードの位置をすばやく計算でき、子ノードへの参照を保持する必要がないためメモリ使用量も減ります。

# --instructions--

指示：ここでは最大ヒープを作成します。まず、要素をヒープに追加する `insert` メソッドを作成してください。挿入時には常にヒープ特性を維持することが重要です。最大ヒープの場合、ルート要素は常に木の中で最大の値を持ち、すべての親ノードは子ノードより大きい必要があります。配列で実装されたヒープでは、通常次の3つのステップでこれを実現します：

<ol>
  <li>新しい要素を配列の末尾に追加します。</li>
  <li>要素が親より大きければ、両者を入れ替えます。</li>
  <li>新しい要素が親より小さくなるか、ルートに到達するまで入れ替えを続けます。</li>
</ol>

最後に、ヒープに追加されたすべてのアイテムを配列で返す `print` メソッドを追加してください。

# --hints--

`MaxHeap` データ構造が存在するはずです。

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` は `insert` というメソッドを持つべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` は `print` というメソッドを持つべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`insert` メソッドは最大ヒープ特性に従って要素を追加する必要があります。

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    test.insert(50);
    test.insert(100);
    test.insert(700);
    test.insert(32);
    test.insert(51);
    test.insert(800);
    const result = test.print();
    const solution = JSON.stringify([null,800,51,700,32,50,100]);
    const solutionWithoutNull = JSON.stringify([800,51,700,32,50,100]);

    return (result.length == 6) ? (JSON.stringify(result) == solutionWithoutNull) : (JSON.stringify(result) == solution);
  })()
);
```

# --seed--

## --seed-contents--

```js
var MaxHeap = function() {
  // Only change code below this line
    
  // Only change code above this line
};
```

# --solutions--

```js
var MaxHeap = function() {
    // Only change code below this line
    this.heap = [];
    this.parent = index => {
      return Math.floor((index - 1) / 2);
    }
    this.insert = element => {
      this.heap.push(element);
      this.heapifyUp(this.heap.length - 1);
    }
    this.heapifyUp = index => {
      let currentIndex = index,
      parentIndex = this.parent(currentIndex);
      while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
        this.swap(currentIndex, parentIndex);
        currentIndex = parentIndex;
        parentIndex = this.parent(parentIndex);
      }
    }
    this.swap = (index1, index2) => {
      [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
    }
    this.print = () => {
      return this.heap;
    }
    // Only change code above this line
};
```
