---
id: 587d825b367417b2b2512c8e
title: ハッシュテーブルを作成する
challengeType: 1
forumTopicId: 301627
dashedName: create-a-hash-table
---

# --description--

このチャレンジではハッシュテーブルについて学びます。ハッシュテーブルは連想配列やキーと値のペアのマッピングを実装するために使われます。これは、私たちが学んだオブジェクトやMapsのようなものです。例えば、JavaScriptのオブジェクトはハッシュテーブルとして実装されることがあります（実際の実装は実行環境によって異なります）。ハッシュテーブルの仕組みは、キーの入力を受け取り、そのキーを決定的な方法で数値にハッシュ化します。この数値が実際のキーとして使われ、対応する値が保存されます。そして、同じキーに再度アクセスしようとすると、ハッシュ関数がキーを処理し、同じ数値を返します。その数値を使って対応する値を検索します。これにより、平均して非常に効率的なO(1)の検索時間が実現されます。

ハッシュテーブルは、ハッシュ関数が指定された範囲内の配列のインデックスを生成する配列として実装できます。この方法では、配列のサイズの選択やハッシュ関数が重要になります。例えば、ハッシュ関数が異なる2つのキーに対して同じ値を生成したらどうなるでしょうか？これを衝突と呼びます。衝突を処理する方法の一つは、そのインデックスに両方のキーと値のペアを保存することです。そして、どちらかのキーで検索するときは、そのバケット内のアイテムを繰り返し処理して探す必要があります。良いハッシュ関数は衝突を最小限に抑え、効率的な検索時間を維持します。

ここではハッシュ化やハッシュテーブルの実装の詳細にはこだわらず、どのように動作するかの大まかな感覚を掴みましょう。

# --instructions--

基本的なハッシュテーブルの機能を作成しましょう。あなたが使える単純なハッシュ関数を用意しました。文字列の値を`hash`に渡すと、ストレージのキーとして使えるハッシュ値が返されます。このハッシュ値を使って`this.collection`オブジェクトにアイテムを保存してください。次の3つのメソッドを作成しましょう：`add`、`remove`、`lookup`。最初のメソッドはキーと値のペアをハッシュテーブルに追加します。2つ目はキーを渡すとキーと値のペアを削除します。3つ目はキーを受け取り、対応する値を返すか、キーが存在しない場合は`null`を返します。

衝突が起きる場合を考慮してコードを書いてください！

**注意:** `remove`メソッドのテストは、`add`と`lookup`メソッドが正しく実装されるまで通りません。

# --hints--

`hash`関数は有効である必要があります。

```js
let calls = 0;
const ourHash = string => {
  calls++;
  let hashCode = 0;
  for (let i = 0; i < string.length; i++) {
    hashCode += string.charCodeAt(i);
  }
  return hashCode;
};

assert.strictEqual(hash('yek'),ourHash('yek'));

assert.strictEqual(hash('key'),ourHash('key'));
assert.strictEqual(hash('key1'),ourHash('key1'));
assert.strictEqual(hash('key2'),ourHash('key2'));
assert.strictEqual(hash('key3'),ourHash('key3'));

assert.strictEqual(hash('1key'),ourHash('1key'));
assert.strictEqual(hash('ke1y'),ourHash('ke1y'));
assert.strictEqual(hash('altKey'),ourHash('altKey'));

assert.strictEqual(called,calls); 
```

`HashTable`データ構造が存在する必要があります。

```js
assert.isDefined(HashTable); 
let newHashTable = new HashTable(); 
assert.isObject(newHashTable); 
```

`HashTable`は`add`メソッドを持つべきです。

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.add); 
```

`HashTable`は`lookup`メソッドを持つべきです。

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.lookup); 
```

`HashTable`は`remove`メソッドを持つべきです。

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.remove); 
```

`add`メソッドはキーと値のペアを追加し、`lookup`メソッドは指定されたキーに対応する値を返すべきです。

```js
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
assert.strictEqual(newHashTable.lookup('key'),'value'); 
```

`remove`メソッドはキーを入力として受け取り、対応するキーと値のペアを削除するべきです。

```js
let hashValue = hash('key');
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
newHashTable.remove('key'); 
assert.notProperty(newHashTable.collection,hashValue); 
```

`remove`メソッドは正しいキーと値のペアのみを削除するべきです。

```js
let newHashTable = new HashTable(); 
let hashValue = hash('key');
newHashTable.add('key','value');
newHashTable.add('yek','value');
newHashTable.add('altKey','value');
newHashTable.remove('yek');

assert.notExists(newHashTable.lookup('yek'));
assert.exists(newHashTable.lookup('altKey'));
assert.exists(newHashTable.lookup('key'));

newHashTable.remove('key');

assert.notProperty(newHashTable.collection,hashValue); 
assert.exists(newHashTable.lookup('altKey'));

```

アイテムはハッシュ関数を使って追加されるべきです。

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1', 'value1')
newHashTable.add('key2', 'value2');
newHashTable.add('key3', 'value3');
assert.strictEqual(called,3); 
```

ハッシュテーブルは衝突を処理するべきです。

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1','value1');
newHashTable.add('1key', 'value2'); 
newHashTable.add('ke1y', 'value3'); 
assert.strictEqual(newHashTable.lookup('key1'),'value1');
assert.strictEqual(newHashTable.lookup('1key'),'value2');
assert.strictEqual(newHashTable.lookup('ke1y'),'value3');
```

# --seed--

## --seed-contents--

```js
var called = 0;
var hash = string => {
  called++;
  var hashed = 0;
  for (var i = 0; i < string.length; i++) {
    hashed += string.charCodeAt(i);
  }
  return hashed;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
var called = 0;
var hash = (string) => {
  called++;
  var hash = 0;
  for (var i = 0; i < string.length; i++) { hash += string.charCodeAt(i); }
  return hash;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line

  this.add = function(key, val) {
    var theHash = hash(key);
    if (!this.collection.hasOwnProperty(theHash)) {
      this.collection[theHash] = {};
    }
    this.collection[theHash][key] = val;
  }

  this.remove = function(key) {
    var theHash = hash(key);
    var hashedObj = this.collection[theHash];
    if (hashedObj.hasOwnProperty(key)) {
      delete hashedObj[key];
    }
    if (!Object.keys(hashedObj).length) {
      delete this.collection[theHash];
    }
  }

  this.lookup = function(key) {
    var theHash = hash(key);
    if (this.collection.hasOwnProperty(theHash)) {
      return this.collection[theHash][key];
    }
    return null
  }
  // Only change code above this line
};
```
