---
id: 587d8250367417b2b2512c60
title: キュークラスを作成する
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--

スタックと同様に、キューは要素の集合です。しかしスタックとは異なり、キューはFIFO（先入れ先出し）原則に従います。キューに追加された要素はキューの末尾にプッシュされ、キューの先頭にある要素だけが削除されます。

キューを表すのに配列を使うこともできますが、スタックと同様に、キューに対する操作の制御を制限したいです。

キュークラスの主なメソッドはenqueueとdequeueです。enqueueメソッドは要素をキューの末尾にプッシュし、dequeueメソッドはキューの先頭の要素を削除して戻り値として返します。他に便利なメソッドとしてfront、size、isEmptyがあります。

# --instructions--

要素をキューの末尾にプッシュする`enqueue`メソッド、先頭の要素を削除して戻り値として返す`dequeue`メソッド、先頭の要素を確認できる`front`メソッド、長さを表示する`size`メソッド、キューが空かどうかをチェックする`isEmpty`メソッドを書いてください。

# --hints--

`Queue`クラスには`enqueue`メソッドが必要です。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

`Queue`クラスには`dequeue`メソッドが必要です。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

`Queue`クラスには`front`メソッドが必要です。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

`Queue`クラスには`size`メソッドが必要です。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

`Queue`クラスには`isEmpty`メソッドが必要です。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

`dequeue`メソッドはキューの先頭の要素を削除して戻り値として返す必要があります。

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

`front`メソッドはキューの先頭の要素の値を返す必要があります。

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

`size`メソッドはキューの長さを返す必要があります。

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

`isEmpty`メソッドはキューに要素がある場合に`false`を返す必要があります。

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```
