---
id: 587d825a367417b2b2512c87
title: Create a Doubly Linked List
challengeType: 1
forumTopicId: 301626
dashedName: create-a-doubly-linked-list
---

# --description--

これまで作成してきた連結リストはすべて単方向連結リストです。ここでは、<dfn>双方向連結リスト</dfn>を作成します。名前の通り、双方向連結リストのノードはリスト内の次のノードと前のノードの両方を参照しています。

これによりリストを両方向にたどることができますが、各ノードが前のノードへの参照を追加で持つ必要があるため、より多くのメモリを使用します。

# --instructions--

`Node` オブジェクトを用意し、`DoublyLinkedList` を開始しました。双方向連結リストに `add` と `remove` という2つのメソッドを追加しましょう。`add` メソッドは与えられた要素をリストに追加し、`remove` メソッドはリスト内の指定された要素のすべての出現を削除する必要があります。

これらのメソッドを書く際には、最初や最後の要素の削除など、あらゆる可能なエッジケースに注意してください。また、空のリストからアイテムを削除しようとした場合は `null` を返すようにしてください。

# --hints--

`DoublyLinkedList` のデータ構造が存在している必要があります。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

`DoublyLinkedList` は `add` というメソッドを持っているはずです。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.add == undefined) {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

`DoublyLinkedList` は `remove` というメソッドを持っているはずです。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.remove == undefined) {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

空のリストからアイテムを削除すると `null` を返すべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.remove(100) == null;
  })()
);
```

`add` メソッドはリストにアイテムを追加するべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(5);
    test.add(6);
    test.add(723);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '56723';
  })()
);
```

各ノードは前のノードを追跡しているべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(50);
    test.add(68);
    test.add(73);
    var result = [];
    var curr = test.tail;
    while (curr) {
      result.push(curr.data);
      curr = curr.prev;
    }
    return result.join('') == '736850';
  })()
);
```

最初のアイテムはリストから削除可能であるべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(25);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '3560';
  })()
);
```

最後のアイテムはリストから削除可能であるべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(60);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '2535';
  })()
);
```


# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
