---
id: 587d8258367417b2b2512c80
title: 二分探索木で葉ノードを削除する
challengeType: 1
forumTopicId: 301637
dashedName: delete-a-leaf-node-in-a-binary-search-tree
---

# --description--

これは二分探索木でより難しい操作である削除を実装する3つのチャレンジの最初です。削除が難しいのは、ノードを削除すると木のリンクが切れてしまうためです。これらのリンクは二分木の構造を維持するために慎重に再構築する必要があります。削除によっては木の再配置が必要になる場合もあります。一般的に、ノードを削除しようとすると次の3つのケースのいずれかに遭遇します。葉ノード：削除対象のノードに子ノードが0個。子ノードが1つ：削除対象のノードに子ノードが1つだけ。子ノードが2つ：削除対象のノードに子ノードが2つある。葉ノードの削除は簡単で、単純に削除すればよいです。子ノードが1つのノードの削除も比較的簡単で、削除したノードの親をその子ノードに繋げばよいです。しかし子ノードが2つあるノードの削除はより難しく、親の木に再接続すべき子ノードが2つできてしまいます。このケースの対処は3つ目のチャレンジで見ていきます。さらに、削除を扱う際にはいくつかの例外的なケースにも注意が必要です。木が空の場合は？削除対象のノードがルートノードの場合は？木に要素が2つだけの場合は？今回はまず葉ノードを削除する最初のケースを扱いましょう。

# --instructions--

二分木に`remove`というメソッドを作成してください。ここに削除操作のロジックを構築します。まず、remove内に現在の木から削除対象のノードを探す関数を作成してください。ノードが木に存在しない場合、`remove`は`null`を返すべきです。次に、削除対象のノードが子ノードを持たない葉ノードであれば、その親の参照を`null`に設定してください。これでノードは木から削除されます。そのためには削除対象のノードの親も追跡する必要があります。また、削除対象のノードの子ノードの数を追跡する方法を作ると便利です。これによりどのケースに該当するか判断できます。2つ目と3つ目のケースは次のチャレンジで扱います。頑張ってください！

# --hints--

`BinarySearchTree`というデータ構造が存在するはずです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

二分探索木には`remove`というメソッドがあるべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

空の木から要素を削除しようとすると`null`を返すべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

存在しない要素を削除しようとすると`null`を返すべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(30);
    return test.remove(100) == null;
  })()
);
```

ルートノードに子ノードがなければ、削除するとルートは`null`に設定されるべきです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

`remove`メソッドは葉ノードを木から削除するはずです。

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
}
```

# --solutions--

```js
// solution required
```
