---
id: 587d8251367417b2b2512c63
title: 連結リストから要素を削除する
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--

連結リストの実装で必要となる次の重要なメソッドは`remove`メソッドです。このメソッドは削除したい要素を引数として受け取り、その要素を含むノードを探して削除します。

連結リストからノードを削除するときは、リストの残りの部分が孤立しないようにすることが重要です。すべてのノードの`next`プロパティはリスト内の次のノードを指していることを思い出してください。たとえば、真ん中の要素を削除する場合は、その要素の前のノードの`next`プロパティが、真ん中の要素の`next`プロパティ（つまりリストの次のノード）を指すように接続を確保する必要があります。

少しわかりにくいかもしれませんので、コンガラインの例に戻って概念モデルを確認しましょう。コンガラインに並んでいて、あなたのすぐ前にいる人がラインを離れたと想像してください。ラインを離れた人はもう誰の手も握っておらず、あなたも離れた人の手を握っていません。あなたは前に進み、次にいる人の手を握ります。

もし削除したい要素が`head`要素であれば、`head`を連結リストの2番目のノードに再代入します。

# --instructions--

要素を受け取り、それを連結リストから削除する`remove`メソッドを書いてください。

**注意:** 要素が連結リストから削除されるたびに、リストの`length`は1つ減少する必要があります。

# --hints--

あなたの`LinkedList`クラスには`remove`メソッドが必要です。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

あなたの`remove`メソッドは、最初のノードが削除されたときに`head`を2番目のノードに再代入する必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

あなたの`remove`メソッドは、削除されたノードごとに連結リストの`length`を1つ減らす必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

あなたの`remove`メソッドは、削除されたノードの前のノードの参照を削除されたノードの`next`参照に再代入する必要があります。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

あなたの`remove`メソッドは、要素が連結リストに存在しない場合はリストを変更しないでください。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```
