---
id: 691f7bbf76229bb97c60827d
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、`for...of` ループを使って、文字列の配列の中にある文字列が何回現れるかを数える関数を作る練習をします。

その関数を書く前に、文字列の各文字をコンソールに記録する、より簡単な関数を作成します。

まず、`printCharacters` というパラメータを持つ、`str` という空の関数を定義してください。

# --hints--

`printCharacters` 関数を作成してください。

```js
assert.isFunction(printCharacters);
```

`printCharacters` 関数は `str` という名前のパラメータを1つ持つ必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.printCharacters;
assert.lengthOf(parameters, 1);
assert.equal(parameters[0].toString(), 'str');
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
