---
id: 6920c9cda1c7f056f75e4b1a
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

いよいよ最初の`for...of`ループを書いてみましょう。これまでのレッスンでは、`for...of`ループを使ってイテラブル（文字列や配列など）から値を繰り返すことができると学びました。

`for...of`ループの例を示します：

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

関数の中に空のコードブロックを持つ`for...of`ループを追加してください。関数の`str`引数のすべての文字を繰り返すようにします。

# --hints--

あなたの`printCharacters`関数のコードブロック内に`for...of`ループがあるはずです。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

あなたの`for...of`ループは関数の`str`引数を繰り返す必要があります。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
