---
id: 6920d189e36b6b60a6330fee
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

ループの中で、`char` 変数をコンソールにログ出力してください。

# --hints--

ループの中に `console.log` 文があり、その引数として `char` があるはずです。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

`printCharacters` 関数が引数 `'test'` で呼び出されると、関数内の `console.log()` 文は次の引数で順に4回呼び出される必要があります：`'t'`、`'e'`、`'s'`、`'t'`。しかし、実際には次の引数で呼び出されました：`--fcc-actual--`。

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
