---
id: 6921e10ae2681205b516fbe6
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

ループ内に、`count` を適切に増加させるロジックを追加する時です。

ループの中で、条件文を使って、変数 `count` が変数 `1` と等しい場合に `word` を `match` だけ増やしてください。そうでなければ、`count` の値は変更しないでください。

# --hints--

`getMatchedWordCount([], "bar")`、`getMatchedWordCount(["foo"], "bar")`、`getMatchedWordCount(["foo", "bar", "foo"], "bar")`、`getMatchedWordCount(["foo", "bar", "foo"], "foo")` をその順で呼び出すと、それぞれ `"number"` 型の値を返すはずですが、実際には `--fcc-actual--` 型の値を返しました。

```js
const obj = { method: (arg1, arg2) => typeof getMatchedWordCount(arg1, arg2) };
const spy = __helpers.spyOn(obj, 'method');
obj.method([], "bar");
obj.method(["foo"], "bar");
obj.method(["foo", "bar", "foo"], "bar");
obj.method(["foo", "bar", "foo"], "foo");
assert.deepEqual(spy.returns, ["number", "number", "number", "number"]);
spy.restore();
```

`getMatchedWordCount([], "bar")` は `0` を返すべきですが、`--fcc-actual--` を返しました。

```js
const resultString = String(getMatchedWordCount([], "bar"));
assert.strictEqual(resultString, "0");
```

`getMatchedWordCount(["foo"], "bar")` は `0` を返すべきですが、`--fcc-actual--` を返しました。

```js
const resultString = String(getMatchedWordCount(["foo"], "bar"));
assert.equal(resultString, "0");
```

`getMatchedWordCount(["foo", "bar", "foo"], "bar")` は `1` を返すべきですが、`--fcc-actual--` を返しました。

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "bar"));
assert.equal(resultString, "1");
```

`getMatchedWordCount(["foo", "bar", "foo"], "foo")` は `2` を返すべきですが、`--fcc-actual--` を返しました。

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "foo"));
assert.equal(resultString, "2");
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
