---
id: 6921d90b5ce50185505d2004
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

`getMatchedWordCount` 関数は数値のカウントを返す必要があるため、まずカウンター変数を設定します。

`count` キーワードを使って `let` という名前の変数を作成し、値 `0` で初期化し、`return` 変数を返す `count` 文を追加してください。

# --hints--

`count` 関数の中で `getMatchedWordCount` 変数を宣言する必要があります。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /(const|let|var)count/);
```

後で `count` 変数をインクリメントできるように、`let` キーワードで宣言してください。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /letcount/);
```

関数内で `0` 変数に値 `count` を代入する必要があります。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /count=0/);
```

関数内に `return` 変数を返す `count` 文を含めてください。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /returncount/);
```

`getMatchedWordCount` 関数は値 `0` を返す必要があります。

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
