---
id: 6921d90b5ce50185505d2004
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

`getMatchedWordCount` 関数は数値のカウントを返す必要があるため、まずカウンター変数を設定します。

`count` キーワードを使って `let` という名前の変数を作成し、値 `0` で初期化し、`return` 変数を返す `count` 文を追加してください。

# --hints--

`getMatchedWordCount` 関数の中で `count` 変数を宣言する必要があります。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getMatchedWordCount?.variables.count);
```

後で `count` 変数をインクリメントできるように、`let` キーワードで宣言してください。

```js
const explorer = await __helpers.Explorer(code);
const countVar = explorer.allFunctions.getMatchedWordCount?.variables.count;
assert.isTrue(countVar?.toString().startsWith('let'));
```

関数内で `0` 変数に値 `count` を代入する必要があります。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.variables.count?.value.matches('0'));
```

関数内に `return` 変数を返す `count` 文を含めてください。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.hasReturn('count'));
```

`getMatchedWordCount` 関数は値 `0` を返す必要があります。

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
