---
id: 6921d688b8d58581d34459b9
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

次の数ステップで、文字列の配列の中に文字列が何回現れるかを数える関数を作成します。

まず、`getMatchedWordCount` という名前で、パラメータ `sentence` と `match` をその順に持つ空の関数を定義してください。

# --hints--

`getMatchedWordCount` 関数を作成してください。

```js
assert.isFunction(getMatchedWordCount);
```

`getMatchedWordCount` 関数は、パラメータ `sentence` と `match` をその順に持つ必要があります。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
