---
id: 6921dae10700e300f26994e3
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

現在の状態の`getMatchedWordCount`関数の出力を確認するには、その下に`console.log`文を追加してください。`console.log()`の中で、`["I", "really", "really", "really", "like", "to", "code"]`パラメータには`sentence`を、`"really"`パラメータには`match`を引数として関数を呼び出してください。

# --hints--

`getMatchedWordCount`関数は`console.log`文の中で呼び出す必要があります。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

`getMatchedWordCount`関数は`["I", "really", "really", "really", "like", "to", "code"]`パラメータに`sentence`、`"really"`パラメータに`match`を引数として呼び出してください。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
