---
id: 68e7e729f6d14f69a23786ef
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

接著，在你的 `main` 元素內，為你的 `h1` 元素添加文字 `Parent Teacher Conference Form` 以及類別 `title` 和 `center`。 

# --hints--

你應該有一個 `h1` 元素。

```js
assert.exists(document.querySelector("h1"));
```

你的 `h1` 元素應該有文字 `Parent Teacher Conference Form`。

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

你的 `h1` 元素應該有 `title` 和 `center` 這兩個類別。

```js
assert.exists(document.querySelector("h1.title.center"));
```

你的 `h1` 元素應該位於 `main` 元素內。

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>

</html>
```
