---
id: 68e707257ea91c2e12d489fb
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將練習如何透過建置家長會表單，為圓鈕添加訂製樣式。已為你提供 HTML 樣板。

首先為你的 `main` 元素添加一個名為 `container` 的類別。

# --hints--

你應該有一個 `main` 元素。

```js
assert.exists(document.querySelector("main"));
```

你的 `main` 元素應該有一個名為 `container` 的類別。

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
