---
id: 68ea9f01c51b66b937793611
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

接著，為你的標籤添加另一個類別為 `"contact-method"` 且 `for` 屬性設為 `"phone"` 的標籤。標籤文字應為 `Phone: `。

在你的 `label` 元素下方，為你的圓鈕添加 `id` 和 `value` 屬性，並設為 `"phone"`。`name` 屬性應設為 `"contact-method"`，類別應設為 `"contact-method-radio-btn"`。

# --hints--

你應該在最後一個 `fieldset` 中有兩個 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type label"), 2);
```

`label` 元素應該有一個類別為 `"contact-method"`。

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

`label` 元素的 `for` 屬性應該設為 `"phone"`。

```js
assert.exists(document.querySelector("fieldset:last-of-type label[for='phone']"));
```

`label` 元素應該有文字 `Phone: `。

```js
const label = document.querySelector("fieldset:last-of-type label[for='phone']");
assert.match(label?.innerText, /Phone:\s?/);
```

你應該在最後一個 `fieldset` 裡有兩個 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type input"), 2);
```

你的 `input` 元素應該有一個 `type` 屬性設定為 `radio`。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

`input` 元素應該有一個 `id` 屬性設為 `"phone"`。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("id"), "phone");
```

`input` 元素應該有一個 `value` 屬性設定為 `"phone"`。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("value"), "phone");
```

`input` 元素應該有一個 `name` 屬性設定為 `"contact-method"`。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

`input` 元素應該有一個類別為 `"contact-method-radio-btn"`。

```js
assert.exists(document.querySelector("fieldset:last-of-type input[id='phone'].contact-method-radio-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
