---
id: 68ea8666b8b473455f5b33ef
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

表單的第一部分將專注於學生的資訊，例如他們的姓名和年級。

在 `form` 元素內，為你的 `form` 添加一個 `fieldset` 元素。在該 `fieldset` 元素內，為你的 `fieldset` 添加一個帶有文字 `Student Information` 的 `legend` 元素。 

在 `legend` 元素下方，添加一個帶有文字 `Full Name: ` 和 `for` 屬性值為 `student-name` 的 `label` 元素。


# --hints--

你應該有一個 `fieldset` 元素。

```js
assert.exists(document.querySelector("fieldset"));
```

你的 `fieldset` 元素應該在 `form` 元素內。

```js
assert.exists(document.querySelector("form fieldset"));
```

你應該有一個 `legend` 元素。

```js
assert.exists(document.querySelector("legend"));
```

你的 `legend` 元素應該在 `fieldset` 元素內。

```js
assert.exists(document.querySelector("fieldset legend"));
```

你的 `legend` 元素應該有文字 `Student Information`。

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

你應該有一個 `label` 元素。

```js
assert.exists(document.querySelector("label"));
```

你的 `label` 元素應該有文字 `Full Name: `。別忘了冒號後面的空格。

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

你的 `label` 元素應該有一個 `for` 屬性。

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

你的 `label` 元素的 `for` 屬性應該有值 `student-name`。

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

你的 `label` 元素應該位於 `legend` 元素下方。

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
