---
id: 68ea8c330ed07b60879f3d38
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

現在是時候為你的表單添加元素以收集學生的成績資訊。

首先為你的表單添加另一個 `label` 元素，其 `文字` 為 `Student Grade: `，並將 `for` 屬性設為 `"grade"`。

然後，在你的 `label` 元素下方，添加一個 `input` 元素，並將 `type` 屬性設為 `"number"`。`name` 和 `id` 屬性應設為 `"grade"`。`placeholder` 屬性應設為 `"E.g., 4"`。最後，你的數字 `input` 應該是必填的。 

# --hints--

你應該在頁面上有總共兩個 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

你的第二個 `label` 元素應該有文字 `Student Grade: `。別忘了冒號後面的空格。

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

你的第二個 `label` 元素應該有一個 `for` 屬性。

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

你的第二個 `label` 元素的 `for` 屬性應該有值 `"grade"`。

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

你應該在頁面上有兩個 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

你的第二個 `input` 元素應該位於你的第二個 `label` 元素下方。

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

你的第二個 `input` 元素應該有一個 `type` 屬性設定為 `"number"`。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

你的第二個 `input` 元素應該有一個 `name` 屬性設為 `"grade"`。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

你的第二個 `input` 元素應該有一個 `id` 屬性設定為 `"grade"`。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

你的第二個 `input` 元素應該有一個 `placeholder` 屬性，設定為 `"E.g., 4"`。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

你的第二個 `input` 元素應該有一個 `required` 屬性。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
