---
id: 68ea8fdf2c9b167307e7f671
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

現在是時候為收集家長的資訊添加表單元素。

首先為你的文件添加一個 `label` 元素，其 `文字` 為 `Parent/Guardian Name: `，並將 `for` 屬性設為 `"parent-name"`。

然後，在你的 `label` 元素下方，添加一個 `input` 元素，並將 `type` 屬性設為 `"text"`。`name` 和 `id` 屬性應設為 `"parent-name"`。`placeholder` 屬性應設為 `"E.g., Nancy Doe"`。最後，你的 `input` 應該是必填的。 

# --hints--

你應該在頁面上有總共三個 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

你的第三個 `label` 元素應該有文字 `Parent/Guardian Name: `。別忘了冒號後面的空格。

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

你的第三個 `label` 元素應該有一個 `for` 屬性。

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

你的第三個 `label` 元素的 `for` 屬性應該有值 `"parent-name"`。

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

你應該在頁面上有三個 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

你的第三個 `input` 元素應該位於你的第三個 `label` 元素下方。

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

你的第三個 `input` 元素應該有一個 `type` 屬性設定為 `"text"`。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

你的第三個 `input` 元素應該有一個 `name` 屬性設定為 `"parent-name"`。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

你的第三個 `input` 元素應該有一個 `id` 屬性設為 `"parent-name"`。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

你的第三個 `input` 元素應該有一個 `placeholder` 屬性，設定為 `"E.g., Nancy Doe"`。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

你的第三個 `input` 元素應該有一個 `required` 屬性。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
