---
id: 68eaa668d4d9fbe529a0e16b
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

現在，為你的 `label` 元素添加文字 `Any specific concerns or topics you'd like to discuss?`，並將 `for` 屬性設為 `notes`。

在此 `label` 下方，添加一個 `textarea` 元素。將其 `id` 設為 `notes`、`name` 設為 `notes`、`rows` 設為 4，並將 `cols` 設為 50。

# --hints--

你的最後一個 `fieldset` 應該包含一個 `label` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

`label` 元素應該有文字 `Any specific concerns or topics you'd like to discuss?`。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

`label` 元素的 `for` 屬性應該設為 `"notes"`。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

你的最後一個 `fieldset` 應該包含一個 `textarea` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

你的 `textarea` 元素應該位於 `label` 元素下方。

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

`textarea` 元素應該設定 `id` 屬性為 `"notes"`。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

`textarea` 元素應該有一個 `name` 屬性設定為 `"notes"`。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

`textarea` 元素應該有一個 `rows` 屬性設為 `4`。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

`textarea` 元素應該有一個 `cols` 屬性設定為 `50`。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
