---
id: 68ea81da9f687732368a5322
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

現在是時候添加 `form` 和 `input` 元素，這些元素將表述家長和學生的資訊。

首先在 `p` 元素下方添加一個 `form` 元素。

# --hints--

你應該有一個 `form` 元素。

```js
assert.exists(document.querySelector(".container form"));
```

表單元素應該位於 `p` 元素下方。

```js
assert.exists(document.querySelector(".container > p+form"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
