---
id: 68ea975c7cbbb894a386e402
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

下一步是為電子郵件聯絡方法添加 `label` 和 `input` 元素。

首先為你的標籤添加一個類別為 `"contact-method"` 且 `for` 屬性設為 `"email"`。標籤文字應為 `Email: `。

在你的 `label` 元素下方，添加一個 `id` 和 `value` 屬性設為 `"email"` 的圓鈕。`name` 屬性應設為 `"contact-method"`，類別應設為 `"contact-method-radio-btn"`。 

最後，請確保此單選按鈕預設為已勾選。

# --hints--


你的最後一個欄位集應該包含一個 `label` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

這個新的 `label` 元素（用於電子郵件）應該有一個 `"contact-method"` 的類別。

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

新的 `label` 元素的 `for` 屬性應該設為 `"email"`。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

`label` 元素應該有文字 `Email: `。請確保冒號後有空格。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

你應該在最後一個 `fieldset` 裡面有一個 `input` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

你的 `input` 元素應該有一個 `type` 屬性設定為 `radio`。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

`input` 元素應該有一個 `id` 屬性設為 `"email"`。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

`input` 元素應該有一個 `value` 屬性設定為 `"email"`。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

`input` 元素應該有一個 `name` 屬性設定為 `"contact-method"`。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

`input` 元素應該有一個類別為 `"contact-method-radio-btn"`。

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

你的 `radio` 按鈕預設應該是 `checked`。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
