---
id: 68ea7c41f2555e16174e7cfa
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

接著，為你的段落元素添加文字 `Please fill out the form below to help schedule your parent-teacher conference.`。你的段落元素也應該有類別 `description` 和 `center`。

# --hints--

你應該有一個 `p` 元素。

```js
assert.exists(document.querySelector("p"));
```

你的 `p` 元素應該有文字 `Please fill out the form below to help schedule your parent-teacher conference.`

```js
assert.strictEqual(document.querySelector("p")?.innerText, "Please fill out the form below to help schedule your parent-teacher conference.");
```

你的 `p` 元素應該有 `description` 和 `center` 類別。

```js
assert.exists(document.querySelector("p.description.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
