---
id: 68ea8a5db73b155708992687
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

下一步是為學生的資訊添加相關的 `input` 元素。

首先為你的 `input` 元素添加一個 `type` 屬性，設定為 `"text"`。然後添加一個 `name` 和 `id` 屬性，兩者皆設定為 `student-name`。 

接著，為你的元素添加一個 `placeholder` 屬性，並設為 `E.g., Jane Doe`。最後，添加一個 `required` 屬性。 

# --hints--

你應該有一個 `input` 元素。

```js
assert.exists(document.querySelector("fieldset input"));
```

你的 `input` 元素應該有一個 `type` 屬性設定為 `"text"`。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

你的 `input` 元素應該有一個 `name` 屬性設為 `"student-name"`。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

你的 `input` 元素應該有一個 `id` 屬性設定為 `"student-name"`。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

你的 `input` 元素應該有一個 `placeholder` 屬性，設定為 `E.g., Jane Doe`。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

你的 `input` 元素應該有一個 `required` 屬性。

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
