---
id: 68eb058081504963447bfa60
title: 第 17 步
challengeType: 0
dashedName: step-17
---

# --description--

接著，你應該為主要容器設計樣式。選取具有類別 `container` 的元素，並將其 `background-color` 設為十六進位碼 `#ffffff1a`。

十六進位碼是由六位數字和字母組成的組合，用於 HTML 和 CSS 中表述顏色。它以 `#` 開頭，後面接著三個對組：

- 前兩位數字表述紅色
- 接下來兩個表述綠色
- 最後兩個表現藍色

例如，`#ffffff` 是純白色，因為它的紅色、綠色和藍色值都達到最大值。

你也可以在結尾添加兩個額外的數字來控制不透明度（稱為 alpha 通道）。在 `#ffffff1a` 中，`1a` 使白色變成半透通的。alpha 值越低，顏色看起來越透明。你稍後會學到更多關於十六進位碼的知識！

# --hints--

你應該將標的設定為具有 `container` 類別的元素。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

將 `background-color` 屬性設為 `#ffffff1a`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.container')
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```


```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

--fcc-editable-region--

--fcc-editable-region--
```
