---
id: 68eb33a0e109d026cf545b43
title: 步驟 33
challengeType: 0
dashedName: step-33
---

# --description--

你需要定位內圈並使其初始時不可見。

仍然針對 `.contact-method-radio-btn::before`，將 `transform` 設為 `translate(3px, 3px) scale(0)`。`translate` 函式（程式）會稍微移動圓形，而 `scale(0)` 則將其縮小到零大小，因此預設會隱藏。 

此外，為你的 `transition` 添加 `all 0.3s ease-in`，以在出現時創建平滑的動畫。這表示任何樣式變化都會在 0.3 秒內動畫呈現，起初緩慢然後加速。

# --hints--

你應該繼續將標的設定為 `.contact-method-radio-btn::before` 偽元素。

```js
assert(new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before'));
```

將 `transform` 屬性設為 `translate(3px, 3px) scale(0)`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.transform, 'translate(3px, 3px) scale(0)');
```

將 `transition` 屬性設為 `all 0.3s ease-in`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.oneOf(style?.transition, ['0.3s ease-in', 'all 0.3s ease-in 0s'])
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;  
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

fieldset {
  border: 1px solid gray;
  border-radius: 5px;
  margin: 20px 0;
  padding: 20px;
}

fieldset legend {
  font-size: 1.3rem;
  font-weight: 600;
}

label {
  font-size: 1.2rem;
}

label:not(.contact-method) {
  display: block;
  margin: 10px 0;
}

input:not(.contact-method-radio-btn),
textarea {
  background-color: #ffffff1a;
  width: 95%;
  border: 1px solid gray;
  border-radius: 5px;
  padding: 10px;
}

input,
input::placeholder,
textarea {
  color: whitesmoke;
}

.contact-method-radio-btn {
  appearance: none;
  width: 20px;
  height: 20px;
  border-radius: 50%;
  border: 2px solid gray;
  vertical-align: bottom;
}


.contact-method-radio-btn::before {
  display: block;
  content: " ";
  width: 10px;
  height: 10px;
  border-radius: 50%;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
