---
id: 68eb288f63d818e6e2ff246b
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

現在，你應該為 `input` 和 `textarea` 元素設計樣式。通常先為 `input` 和 `textarea` 設定一般樣式，然後再針對特定型別進行調整，是個不錯的做法。

選取 `input:not(.contact-method-radio-btn), textarea`。將 `background-color` 設為 `#ffffff1a`，`width` 設為 `95%`，`border` 設為 `1px solid gray`，`border-radius` 設為 `5px`，以及 `padding` 設為 `10px`。

`:not(.contact-method-radio-btn)` 這部分排除具有該類別的圓鈕不受此規則的樣式影響。這樣，我們可以將一般樣式套用到大多數輸入元件，同時保持圓鈕的樣式獨立。

這些樣式會讓你的表單欄位更寬且更易讀。試著預覽表單以查看版面配置如何變化。

# --hints--

你應該標的 `input:not(.contact-method-radio-btn), textarea`。

```js
assert(new __helpers.CSSHelp(document).getStyle('input:not(.contact-method-radio-btn), textarea'));
```

將 `background-color` 屬性設為 `#ffffff1a`。

```js
const style = new __helpers.CSSHelp(document).getStyle('input:not(.contact-method-radio-btn), textarea');
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)');
```

將 `width` 屬性設為 `95%`。

```js
const style = new __helpers.CSSHelp(document).getStyle('input:not(.contact-method-radio-btn), textarea');
assert.strictEqual(style?.width, '95%');
```

將 `border` 屬性設為 `1px solid gray`。

```js
const style = new __helpers.CSSHelp(document).getStyle('input:not(.contact-method-radio-btn), textarea');
assert.strictEqual(style?.border, '1px solid gray');
```

將 `border-radius` 屬性設為 `5px`。

```js
const style = new __helpers.CSSHelp(document).getStyle('input:not(.contact-method-radio-btn), textarea');
assert.strictEqual(style?.borderRadius, '5px');
```

將 `padding` 屬性設為 `10px`。

```js
const style = new __helpers.CSSHelp(document).getStyle('input:not(.contact-method-radio-btn), textarea');
assert.strictEqual(style?.padding, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;  
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

fieldset {
  border: 1px solid gray;
  border-radius: 5px;
  margin: 20px 0;
  padding: 20px;
}

fieldset legend {
  font-size: 1.3rem;
  font-weight: 600;
}

label {
  font-size: 1.2rem;
}

label:not(.contact-method) {
  display: block;
  margin: 10px 0;
}

--fcc-editable-region--

--fcc-editable-region--
```

