---
id: a8e512fbe388ac2f9198f0fa
title: implementare un filtro per oggetti corrispondenti
challengeType: 26
dashedName: implement-a-matching-object-filter
---

# --description--

In questo laboratorio, creerai una funzione che filtra un array di oggetti e restituisce solo quegli oggetti che corrispondono a tutte le coppie chiave-valore di un oggetto sorgente dato.

**Obiettivo**: soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione `whatIsInAName` che accetta due argomenti, un array di oggetti e un oggetto sorgente.
1. La funzione `whatIsInAName` dovrebbe restituire un nuovo array contenente solo gli oggetti della collezione che hanno tutte le coppie chiave-valore presenti nell'oggetto sorgente.
1. Se nessun oggetto corrisponde a tutte le coppie chiave-valore dell'oggetto sorgente, la funzione dovrebbe restituire un array vuoto. Per esempio:

```js
whatIsInAName(
  [
    { first: "Romeo", last: "Montague" },
    { first: "Mercutio", last: null },
    { first: "Tybalt", last: "Capulet" }
  ],
  { last: "Capulet" }
);
// should return [{ first: "Tybalt", last: "Capulet" }]
```

# --hints--

Dovresti avere una funzione `whatIsInAName`.

```js
assert.isFunction(whatIsInAName);
```

`whatIsInAName([{ first: "Romeo", last: "Montague" }, { first: "Mercutio", last: null }, { first: "Tybalt", last: "Capulet" }], { last: "Capulet" })` dovrebbe restituire `[{ first: "Tybalt", last: "Capulet" }]`.

```js
assert.deepEqual(
  whatIsInAName(
    [{ first: "Romeo", last: "Montague" }, { first: "Mercutio", last: null }, { first: "Tybalt", last: "Capulet" }],
    { last: "Capulet" }
  ),
  [{ first: "Tybalt", last: "Capulet" }]
);
```

`whatIsInAName([{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }], { "apple": 1 })` dovrebbe restituire `[{"apple": 1}, {"apple": 1}, {"apple": 1, "bat": 2}]`.

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }],
    { "apple": 1 }
  ),
  [{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }], { "apple": 1, "bat": 2 })` dovrebbe restituire `[{"apple": 1, "bat": 2}, {"apple": 1, "bat": 2, "cookie": 2}]`.

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }],
    { "apple": 1, "bat": 2 }
  ),
  [{ "apple": 1, "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }], { "apple": 1, "cookie": 2 })` dovrebbe restituire `[{"apple": 1, "bat": 2, "cookie": 2}]`.

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }],
    { "apple": 1, "cookie": 2 }
  ),
  [{ "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }, { "bat": 2 }], { "apple": 1, "bat": 2 })` dovrebbe restituire `[{"apple": 1, "bat": 2}, {"apple": 1, "bat": 2, "cookie": 2}]`

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }, { "bat": 2 }],
    { "apple": 1, "bat": 2 }
  ),
  [{ "apple": 1, "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{"a": 1, "b": 2, "c": 3}], {"a": 1, "b": 9999, "c": 3})` dovrebbe restituire `[]`.

```js
assert.deepEqual(
  whatIsInAName([{ "a": 1, "b": 2, "c": 3 }], { "a": 1, "b": 9999, "c": 3 }),
  []
);
```

`whatIsInAName([{"a": 1, "b": 2, "c": 3, "d": 9999}], {"a": 1, "b": 9999, "c": 3})` dovrebbe restituire `[]`.

```js
assert.deepEqual(
  whatIsInAName([{ "a": 1, "b": 2, "c": 3, "d": 9999 }], { "a": 1, "b": 9999, "c": 3 }),
  []
);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function whatIsInAName(collection, source) {
  const sourceKeys = Object.keys(source);

  return collection.filter(obj =>
    sourceKeys.every(key => obj.hasOwnProperty(key) && obj[key] === source[key])
  );
}
```
