---
id: 698c4f4e965bd853ad5084e1
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Chini ya `h1`, ongeza kipengele cha `h2` chenye maandishi `Can't Stay Down`, huu ni kichwa cha wimbo wa kwanza.

Chini ya kipengele cha `h2` ongeza kipengele cha `p` chenye maandishi `Artist: Quincy Larson`.

# --hints--

Unapaswa kuwa na kipengele cha `h2`.

```js
assert.exists(document.querySelector('h2'));
```

Kipengele chako cha `h2` kinapaswa kuwa chini ya kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1 + h2'));
```

Kipengele chako cha `h2` kinapaswa kuwa na maandishi `Can't Stay Down`.

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

Unapaswa kuwa na kipengele cha `p`.

```js
assert.exists(document.querySelector('p'));
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi `Artist: Quincy Larson`.

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
