---
id: 698c5229e47d57c6dfaa91bd
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Sifa ya kipengele ya `controls` hutoa vidhibiti vya kucheza ikijumuisha kusitisha/kurejesha kucheza, kutafuta, na udhibiti wa sauti kwa kipengele cha `audio`. 

Sifa ya kipengele ya `controls` ni sifa ya kipengele ya thamani ya kweli au siyo kweli na haihitaji thamani.

Ongeza sifa ya kipengele ya `controls` kwa kipengele cha `audio`.

Sasa unapaswa kuona kipengele cha `audio` kikiwa kinaonyeshwa kwenye ukurasa.

# --hints--

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele ya `controls`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
