---
id: 698c5710136539797f296098
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Chini ya kipengele cha `p`, ongeza kipengele cha `audio`, kipa sifa ya kipengele cha `src` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`, pamoja na sifa za kipengele za `loop` na `controls`.

# --hints--

Unapaswa kuwa na kipengele cha `audio` cha pili.

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

Kipengele kipya cha `audio` kinapaswa kuwa chini ya kipengele cha `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele cha `src`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele cha `src` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele cha `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele cha `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
