---
id: 698c4e69abe14a9746c72312
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utaunda mchezaji wa muziki wa HTML. Muundo wa msingi wa HTML umeandaliwa kwa ajili yako.

Tengeneza kipengele cha `h1` na umpe maandishi ya `freeCodeCamp Tunes`.

# --hints--

Unapaswa kuwa na kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi ya `freeCodeCamp Tunes`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'freeCodeCamp Tunes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
