---
id: 698c50535b54b2d75fa985fb
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Sifa ya kipengele `loop` itaanzisha tena sauti mara tu uchezaji utakapokamilika. Fikiria meme ya mtandao inayorudia uchezaji. Kukosa sifa ya kipengele `loop` kutafanya sauti icheze mara moja tu.

Sifa ya kipengele `loop` ni sifa ya kipengele ya thamani ya kweli au siyo kweli na haihitaji thamani.

Ongeza sifa ya kipengele `loop` kwenye kipengele cha `audio`.

# --hints--

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele `loop`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
