---
id: 698c55e7aaacfae48a227459
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Ongeza kipengele kipya cha `h2` chini ya kipengele cha `audio` chenye maandishi ya `Cruising for a Musing`.

Chini ya `h2`, ongeza kipengele cha `p` chenye maandishi ya `Artist: Quincy Larson`.

# --hints--

Unapaswa kuwa na kipengele cha pili cha `h2`.

```js
assert.exists(document.querySelector('h2:nth-of-type(2)'));
```

Kipengele chako cha pili cha `h2` kinapaswa kuwa na maandishi ya `Cruising for a Musing`.

```js
const h2 = document.querySelector('h2:nth-of-type(2)');
assert.strictEqual(h2?.textContent.trim(), 'Cruising for a Musing');
```

Unapaswa kuwa na kipengele cha pili cha `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(2)'));
```

Kipengele chako cha pili cha `p` kinapaswa kuwa na maandishi ya `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(2)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
