---
id: 698c4fe24f1f1a17be74a23e
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Ifuatayo, tengeneza kipengele cha `audio` chini ya kipengele cha `p`. Katika hatua chache zijazo, utaongeza sifa za kipengele zinazohitajika ili kufanya kipengele cha `audio` kucheza muziki.

# --hints--

Unapaswa kuwa na kipengele cha `audio`.

```js
assert.exists(document.querySelector('audio'));
```

Kipengele kipya cha `audio` kinapaswa kuwa chini ya kipengele cha `p`.

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
