---
id: 698c58745b73289bb6c89ece
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Wimbo wa mwisho unaitwa `Scratching the Surface`, msanii ni Quincy Larson, na faili ya kutumia ni `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

Ongeza wimbo wa mwisho ili kukamilisha mchezaji wa muziki ukifuata muundo uleule uliotumia kwa nyimbo mbili za kwanza.

# --hints--

Unapaswa kuwa na kipengele cha tatu cha `h2`.

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

Kipengele chako cha tatu cha `h2` kinapaswa kuwa na maandishi ya `Scratching the Surface`.

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

Unapaswa kuwa na kipengele cha tatu cha `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

Kipengele chako cha tatu cha `p` kinapaswa kuwa na maandishi `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

Unapaswa kuwa na kipengele cha tatu cha `audio`.

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

Kipengele kipya cha `audio` kinapaswa kuwa chini ya kipengele cha `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele ya `src`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele ya `src` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele ya `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Kipengele chako cha `audio` kinapaswa kuwa na sifa ya kipengele ya `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
