---
id: 669fc7e141e4703748c558bf
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

Inside the `html` element, add a `head` element.

# --hints--

你应该有一个 `head` 开始标签。

```js
assert.match(code, /<head>/i);
```

You should have a closing `head` tag.

```js
assert.match(code, /<\/head>/i);
```

你的 `head` 开始标签应在 `head` 结束标签之前。

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
