---
id: 587d7fb7367417b2b2512c0d
title: استخدام model.findById() للبحث في قاعدة بياناتك بواسطة _id
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

عند حفظ وثيقة، تضيف MongoDB تلقائيًا الحقل `_id` وتعيّنه إلى مفتاح أبجدي رقمي فريد. البحث بواسطة `_id` هو عملية متكررة جدًا، لذا يوفر Mongoose طريقة مخصصة لذلك.

# --instructions--

عدّل `findPersonById` لإيجاد الشخص الوحيد الذي يملك `_id` معين، باستخدام `Model.findById() -> Person`. استخدم المعلمة `personId` كمعيار للبحث.

# --hints--

يجب أن ينجح البحث عن عنصر بواسطة Id

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```

