---
id: 587d7fb9367417b2b2512c12
title: تسلسل مساعدي استعلام البحث لتضييق نتائج البحث
challengeType: 2
forumTopicId: 301533
dashedName: chain-search-query-helpers-to-narrow-search-results
---

# --description--

إذا لم تمرر استدعاء راجع كآخر معلمة إلى `Model.find()` (أو إلى طرق البحث الأخرى)، فلن يُنفذ الاستعلام. يمكنك تخزين الاستعلام في متغير لاستخدامه لاحقًا. هذا النوع من الكائنات يتيح لك بناء استعلام باستخدام صيغة نحوية للتسلسل. يُنفذ البحث الفعلي في قاعدة البيانات عندما تسلسل أخيرًا الطريقة `.exec()`. يجب دائمًا تمرير استدعاء الراجع إلى هذه الطريقة الأخيرة. هناك العديد من مساعدي الاستعلام، وسنستخدم هنا الأكثر شيوعًا.

# --instructions--

عدّل دالة `queryChain` للعثور على الأشخاص الذين يحبون الطعام المحدد في المتغير المسمى `foodToSearch`. رتبهم حسب `name`، حدّد النتائج إلى وثيقتين، وأخفِ عمرهم. سلّس `.find()`، `.sort()`, `.limit()`, `.select()`، ثم `.exec()`. مرر استدعاء الراجع `done(err, data)` إلى `exec()`.

# --hints--

يجب أن ينجح تسلسل مساعدي الاستعلام

```js
  const response = await fetch(code + '/_api/query-tools', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Pablo', age: 26, favoriteFoods: ['burrito', 'hot-dog'] },
      { name: 'Bob', age: 23, favoriteFoods: ['pizza', 'nachos'] },
      { name: 'Ashley', age: 32, favoriteFoods: ['steak', 'burrito'] },
      { name: 'Mario', age: 51, favoriteFoods: ['burrito', 'prosciutto'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data.length,
    2,
    'the data array length is not what expected'
  );
  assert.notProperty(
    data[0],
    'age',
    'The returned first item has too many properties'
  );
  assert.equal(
    data[0].name,
    'Ashley',
    'The returned first item name is not what expected'
  );
  assert.notProperty(
    data[1],
    'age',
    'The returned second item has too many properties'
  );
  assert.equal(
    data[1].name,
    'Mario',
    'The returned second item name is not what expected'
  );
```

