---
id: 587d7fb7367417b2b2512c0a
title: إنشاء سجلات متعددة باستخدام model.create()
challengeType: 2
forumTopicId: 301537
dashedName: create-many-records-with-model-create
---

# --description--

أحيانًا تحتاج إلى إنشاء عدة نسخ من النماذج الخاصة بك، مثلًا عند تعبئة قاعدة بيانات ببيانات أولية. تأخذ الدالة `Model.create()` مصفوفة من الكائنات مثل `[{name: 'John', ...}, {...}, ...]` كمعلمة أولى، وتحفظها جميعًا في قاعدة البيانات.

# --instructions--

عدّل الدالة `createManyPeople` لإنشاء عدة أشخاص باستخدام `Model.create()` مع المعلمة `arrayOfPeople`.

**ملاحظة:** يمكنك إعادة استخدام النموذج الذي أنشأته في التمرين السابق.

# --hints--

يجب أن ينجح إنشاء عدة عناصر في قاعدة البيانات دفعة واحدة

```js
  const response = await fetch(code + '/_api/create-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'John', age: 24, favoriteFoods: ['pizza', 'salad'] },
      { name: 'Mary', age: 21, favoriteFoods: ['onions', 'chicken'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an array');
  assert.equal(
    data.length,
    2,
    'the response does not contain the expected number of items'
  );
  assert.equal(data[0].name, 'John', 'The first item is not correct');
  assert.equal(
    data[0].__v,
    0,
    'The first item should be not previously edited'
  );
  assert.equal(data[1].name, 'Mary', 'The second item is not correct');
  assert.equal(
    data[1].__v,
    0,
    'The second item should be not previously edited'
  );
```

