---
id: 587d7fb7367417b2b2512c0c
title: استخدام model.findOne() لإرجاع مستند مطابق واحد من قاعدة بياناتك
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()` يتصرف مثل `Model.find()`، لكنه يُرجع مستندًا واحدًا فقط (ليس مصفوفة)، حتى لو وُجدت عناصر متعددة. هو مفيد بشكل خاص عند البحث باستخدام خصائص أعلنتها فريدة.

# --instructions--

عدّل دالة `findOneByFood` لتجد شخصًا واحدًا فقط لديه طعام معين في المفضلات، باستخدام `Model.findOne() -> Person`. استخدم المعلمة `food` في الدالة كمفتاح للبحث.

# --hints--

يجب أن ينجح العثور على عنصر واحد

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```

