---
id: 587d7fb6367417b2b2512c06
title: تثبيت وإعداد Mongoose
challengeType: 2
forumTopicId: 301540
dashedName: install-and-set-up-mongoose
---

# --description--

ستتطلب منك هذه التحديات كتابة شفرتك باستخدام إحدى الطرق التالية:

- استنسخ <a href="https://github.com/freeCodeCamp/boilerplate-mongomongoose/" target="_blank" rel="noopener noreferrer nofollow">مستودع GitHub هذا</a> وأكمل هذه التحديات محليًا.
- استخدم منشئ مواقع من اختيارك لإكمال المشروع. تأكد من دمج جميع الملفات من مستودع GitHub الخاص بنا.

في هذا التحدٍّ، ستقوم بإعداد قاعدة بيانات MongoDB Atlas واستيراد الحزم المطلوبة للاتصال بها.

اتبع <a href='https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/' target="_blank" rel="noopener noreferrer nofollow">هذا الدليل</a> لإعداد قاعدة بيانات مستضافة على MongoDB Atlas.

# --instructions--

تمت إضافة `mongoose@^5.11.15` إلى ملف `package.json` الخاص بمشروعك. أولًا، استورد mongoose كـ `mongoose` في `myApp.js`. بعد ذلك، أنشئ ملف `.env` وأضف إليه متغير `MONGO_URI`. يجب أن تكون قيمته هي URI لقاعدة بيانات MongoDB Atlas الخاصة بك. تأكد من إحاطة URI بعلامات اقتباس مفردة أو مزدوجة، وتذكر أنه لا يمكنك استخدام مسافات حول علامة `=` في متغيرات البيئة. على سبيل المثال، `MONGO_URI='VALUE'`.

عندما تنتهي، اتصل بقاعدة البيانات باستدعاء طريقة `connect` داخل ملف `myApp.js` باستخدام الصيغة التالية:

```js
mongoose.connect(<Your URI>, { useNewUrlParser: true, useUnifiedTopology: true });
```

# --hints--

يجب أن تكون تبعية "mongoose version ^5.11.15" موجودة في package.json

```js
  const response = await fetch(code + '/_api/file/package.json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  var packJson = JSON.parse(data);
  assert.property(packJson.dependencies, 'mongoose');
  assert.match(
    packJson.dependencies.mongoose,
    /^\^5\.11\.15/,
    'Wrong version of "mongoose". It should be ^5.11.15'
  );
```

يجب أن يكون "mongoose" متصلًا بقاعدة بيانات

```js
  const response = await fetch(code + '/_api/is-mongoose-ok');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.isMongooseOk, 'mongoose is not connected');
```

