---
id: 587d7fb7367417b2b2512c0b
title: استخدام model.find() للبحث في قاعدة بياناتك
challengeType: 2
forumTopicId: 301543
dashedName: use-model-find-to-search-your-database
---

# --description--

في أبسط استخدام له، تقبل `Model.find()` مستند استعلام (كائن JSON) كالمعلمة الأولى، ثم استدعاء راجع. تُرجع مصفوفة من النتائج المطابقة. تدعم مجموعة واسعة جدًا من خيارات البحث. اقرأ المزيد في الوثائق.

# --instructions--

عدّل دالة `findPeopleByName` للعثور على جميع الأشخاص الذين يحملون اسمًا معينًا، باستخدام <code>Model.find() -\> [Person]</code>

استخدم المعلمة `personName` للدالة كمفتاح للبحث.

# --hints--

يجب أن ينجح العثور على جميع العناصر التي تطابق معيارًا معينًا

```js
  const response = await fetch(code + '/_api/find-all-by-name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'r@nd0mN4m3',
      age: 24,
      favoriteFoods: ['pizza']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data[0].name,
    'r@nd0mN4m3',
    'item.name is not what expected'
  );
  assert.equal(data[0].__v, 0, 'The item should be not previously edited');
```

