---
id: 587d7fb8367417b2b2512c0e
title: "أداء التحديثات الكلاسيكية عبر تشغيل Find و Edit ثم Save"
challengeType: 2
forumTopicId: 301541
dashedName: perform-classic-updates-by-running-find-edit-then-save
---

# --description--

في الأيام الخوالي، كان هذا ما تحتاج إلى فعله إذا أردت تعديل وثيقة، وأن تتمكن من استخدامها بطريقة ما (مثل إرسالها مرة أخرى في استجابة خادم). لدى Mongoose طريقة مخصصة للتحديث: `Model.update()`. ترتبط مباشرة بمحرك mongo منخفض المستوى. يمكنها تعديل العديد من الوثائق دفعة واحدة التي تطابق معايير معينة، لكنها لا ترسل الوثيقة المحدثة، بل ترسل فقط رسالة 'حالة'. علاوة على ذلك، تجعل التحقق من صحة النموذج صعبًا، لأنها تستدعي محرك mongo مباشرة.

# --instructions--

عدّل الدالة `findEditThenSave` للعثور على شخص بواسطة `_id` (استخدم أي من الطرق المذكورة أعلاه) باستخدام المعلمة `personId` كمفتاح بحث. أضف `"hamburger"` إلى قائمة `favoriteFoods` الخاصة بالشخص (يمكنك استخدام `Array.push()`). ثم - داخل استدعاء البحث - احفظ `Person` المحدثة باستخدام `save()`.

**ملاحظة:** قد يكون هذا معقدًا إذا كنت قد أعلنت في مخططك أن `favoriteFoods` هي مصفوفة، دون تحديد النوع (أي `[String]`). في هذه الحالة، تصبح `favoriteFoods` من النوع Mixed، ويجب عليك تعليمها يدويًا كمعدلة باستخدام `document.markModified('edited-field')`. اطلع على <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">مقالة Mongoose</a>.

# --hints--

يجب أن ينجح العثور على عنصر وتعديله وتحديثه

```js
  const response = await fetch(code + '/_api/find-edit-save', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Poldo',
      age: 40,
      favoriteFoods: ['spaghetti']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Poldo', 'item.name is not what is expected');
  assert.equal(data.age, 40, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['spaghetti', 'hamburger'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 1, 'The item should be previously edited');
```

