---
id: 587d7fb8367417b2b2512c11
title: حذف العديد من الوثائق باستخدام model.remove()
challengeType: 2
forumTopicId: 301538
dashedName: delete-many-documents-with-model-remove
---

# --description--

`Model.remove()` مفيد لحذف جميع الوثائق التي تطابق المعايير المحددة.

# --instructions--

عدّل الدالة `removeManyPeople` لحذف جميع الأشخاص الذين أسماؤهم موجودة في المتغير `nameToRemove`، باستخدام `Model.remove()`. مرّر لها وثيقة استعلام مع تعيين حقل `name`، واستدعاء راجع.

**ملاحظة:** لا تُعيد `Model.remove()` الوثيقة المحذوفة، بل تُعيد كائن JSON يحتوي على نتيجة العملية وعدد العناصر المتأثرة. لا تنسَ تمريرها إلى استدعاء الراجع `done()`، لأننا نستخدمه في الاختبارات.

# --hints--

يجب أن ينجح حذف العديد من العناصر دفعة واحدة

```js
  const response = await fetch(code + '/_api/remove-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Mary', age: 16, favoriteFoods: ['lollipop'] },
      { name: 'Mary', age: 21, favoriteFoods: ['steak'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(!!data.ok, 'The mongo stats are not what expected');
  assert.equal(
    data.n,
    2,
    'The number of items affected is not what expected'
  );
  assert.equal(data.count, 0, 'the db items count is not what expected');
```

