---
id: 587d7fb6367417b2b2512c07
title: إنشاء نموذج
challengeType: 2
forumTopicId: 301535
dashedName: create-a-model
---

# --description--

**C**RUD الجزء الأول - الإنشاء

أولًا، نحتاج إلى مخطط. كل مخطط يرتبط بمجموعة في MongoDB. يحدد شكل الوثائق داخل تلك المجموعة. المخططات هي اللبنات الأساسية للنماذج. يمكن تسكينها لإنشاء نماذج معقدة، لكن في هذه الحالة سنبقي الأمور بسيطة. النموذج يسمح لك بإنشاء نسخ من كائناتك، تُسمى وثائق.

في الخوادم، تحدث التفاعلات مع قاعدة البيانات داخل دوال المعالِج. تُنفذ هذه الدوال عند حدوث حدث معين (مثل أن يطلب أحدهم نقطة نهاية في واجهة برمجة التطبيقات الخاصة بك). سنتبع نفس النهج في هذه التمارين. الدالة `done()` هي استدعاء راجع يخبرنا أنه يمكننا المتابعة بعد إكمال عملية غير متزامنة مثل الإدراج أو البحث أو التحديث أو الحذف. تتبع هذه الدالة اصطلاح Node، ويجب استدعاؤها كـ `done(null, data)` عند النجاح، أو `done(err)` عند الخطأ.

تحذير - عند التفاعل مع خدمات بعيدة، قد تحدث أخطاء!

```js
/* Example */

const someFunc = function(done) {
  //... do something (risky) ...
  if (error) return done(error);
  done(null, result);
};
```

# --instructions--

أنشئ مخطط شخص يسمى `personSchema` بالشكل التالي:

* حقل `name` مطلوب من نوع `String`
* حقل `age` من نوع `Number`
* حقل `favoriteFoods` من نوع `[String]`

استخدم أنواع المخطط الأساسية في Mongoose. إذا أردت، يمكنك إضافة حقول أخرى، واستخدام محققين بسيطين مثل required أو unique، وتعيين قيم افتراضية. اطلع على <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">مقالة Mongoose</a>.

الآن، أنشئ نموذجًا من `personSchema` وعيّنه إلى المتغير الموجود `Person`.

# --hints--

يجب أن ينجح إنشاء نسخة من مخطط mongoose

```js
  const response = await fetch(code + '/_api/mongoose-model', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Mike',
      age: 28,
      favoriteFoods: ['pizza', 'cheese']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Mike', '"model.name" is not what expected');
  assert.equal(data.age, '28', '"model.age" is not what expected');
  assert.isArray(
    data.favoriteFoods,
    '"model.favoriteFoods" is not an Array'
  );
  assert.include(
    data.favoriteFoods,
    'pizza',
    '"model.favoriteFoods" does not include the expected items'
  );
  assert.include(
    data.favoriteFoods,
    'cheese',
    '"model.favoriteFoods" does not include the expected items'
  );
```

