---
id: 587d7fb6367417b2b2512c09
title: إنشاء وحفظ سجل لنموذج
challengeType: 2
forumTopicId: 301536
dashedName: create-and-save-a-record-of-a-model
---

# --description--

في هذا التحدٍّ يجب عليك إنشاء سجل لنموذج وحفظه.

# --instructions--

داخل الدالة `createAndSavePerson`، أنشئ كائن المستند باستخدام دالة البناء `Person` التي أنشأتها سابقًا. مرر إلى دالة البناء كائنًا يحتوي على الحقول `name` و`age` و`favoriteFoods`. يجب أن تتطابق أنواعها مع تلك الموجودة في `personSchema`. بعد ذلك، استدعِ الطريقة `document.save()` على كائن المستند المُعاد. مرر إليها استدعاء راجع باستخدام اصطلاح Node. هذا نمط شائع؛ كل طرق CRUD التالية تأخذ دالة استدعاء راجع كهذا كالمعلمة الأخيرة.

```js
/* Example */

// ...
person.save(function(err, data) {
  //   ...do your stuff here...
});
```

# --hints--

يجب أن ينجح إنشاء وحفظ عنصر قاعدة البيانات

```js
  const response = await fetch(code + '/_api/create-and-save-person');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isString(data.name, '"item.name" should be a String');
  assert.isNumber(data.age, '28', '"item.age" should be a Number');
  assert.isArray(
    data.favoriteFoods,
    '"item.favoriteFoods" should be an Array'
  );
  assert.equal(data.__v, 0, 'The db item should be not previously edited');
```

