---
id: 587d7fb8367417b2b2512c0f
title: إجراء تحديثات جديدة على وثيقة باستخدام model.findOneAndUpdate()
challengeType: 2
forumTopicId: 301542
dashedName: perform-new-updates-on-a-document-using-model-findoneandupdate
---

# --description--

تحتوي الإصدارات الحديثة من Mongoose على طرق لتبسيط تحديث الوثائق. بعض الميزات المتقدمة مثل pre/post hooks و التحقق تتصرف بشكل مختلف مع هذا الأسلوب، لذلك لا تزال الطريقة الكلاسيكية مفيدة في العديد من الحالات. يمكن استخدام `findByIdAndUpdate()` عند البحث بواسطة المعرف.

# --instructions--

عدّل دالة `findAndUpdate` للعثور على شخص بواسطة `Name` واضبط عمر الشخص على `20`. استخدم معلمة الدالة `personName` كمفتاح للبحث.

**ملاحظة:** يجب أن تُرجع الوثيقة المحدثة. للقيام بذلك، تحتاج إلى تمرير وثيقة الخيارات `{ new: true }` كالمعلمة الثالثة إلى `findOneAndUpdate()`. افتراضيًا، تُرجع هذه الطرق الكائن غير المعدل.

# --hints--

يجب أن ينجح تحديث عنصر باستخدام findOneAndUpdate

```js
  const response = await fetch(code + '/_api/find-one-update', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Dorian Gray',
      age: 35,
      favoriteFoods: ['unknown']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Dorian Gray', 'item.name is not what expected');
  assert.equal(data.age, 20, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['unknown'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(
    data.__v,
    0,
    'findOneAndUpdate does not increment version by design!'
  );
```

