---
id: 587d7fb8367417b2b2512c10
title: حذف وثيقة واحدة باستخدام model.findByIdAndRemove
challengeType: 2
forumTopicId: 301539
dashedName: delete-one-document-using-model-findbyidandremove
---

# --description--

تعمل `findByIdAndRemove` و`findOneAndRemove` مثل طرق التحديث السابقة. تمرر الوثيقة المحذوفة إلى قاعدة البيانات. كما هو معتاد، استخدم المعلمة `personId` كمعيار للبحث.

# --instructions--

عدّل الدالة `removeById` لحذف شخص واحد باستخدام `_id` الخاص به. يجب أن تستخدم إحدى الطريقتين `findByIdAndRemove()` أو `findOneAndRemove()`.

# --hints--

يجب أن ينجح حذف العنصر

```js
  const response = await fetch(code + '/_api/remove-one-person', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Jason Bourne',
      age: 36,
      favoriteFoods: ['apples']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Jason Bourne', 'item.name is not what expected');
  assert.equal(data.age, 36, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['apples'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0);
  assert.equal(data.count, 0, 'the db items count is not what expected');
```

