---
id: 68552464ef378e1051a55c28
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

All'interno del secondo `div`, crea due elementi àncora. Il primo deve avere il testo `Learn more`, e il secondo deve avere il testo `Start listening`.

Per l'attributo `href` dei due elementi àncora, impostali entrambi su `#`.

# --hints--

Dovresti creare due elementi `a` all'interno del secondo elemento `div`

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

Il tuo primo elemento `a` dovrebbe avere il testo `Learn more` e un attributo `href` con valore `#` all'interno del secondo elemento `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

Il tuo secondo elemento `a` dovrebbe avere il testo `Start listening` e un attributo `href` con valore `#` all'interno del secondo elemento `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
