---
id: 685523c05fb41d0848c4b52f
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

All'interno del primo elemento `div`, crea un elemento `span` con il testo `Soundflow`, un elemento `h1` con il testo `Discover New Music` e un elemento `p` con il testo `Stream your favorite tracks and discover new artists.`

# --hints--

Dovresti creare un elemento `span` con il testo `Soundflow` all'interno del primo elemento `div`.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

Dovresti creare un elemento `h1` con il testo `Discover New Music` all'interno del primo elemento `div`.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

Dovresti creare un elemento `p` con il testo `Stream your favorite tracks and discover new artists.` all'interno del primo elemento `div`.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
