---
id: 68551b677a5f6898af364b85
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop, imparerai di più su Tailwind CSS usando le sue classi utility flexbox per costruire un componente call to action (CTA) responsivo.

Il codice boilerplate HTML di base è già fornito, quindi devi prima costruire gli elementi HTML.

Inizia creando un elemento `div`. All’interno di quel `div`, crea altri due elementi `div` separati.

# --hints--

Dovresti creare un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Dovresti avere altri due elementi `div` all’interno dell’elemento `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
