---
id: 6858e7e23bf9e475e14cb931
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Ora che la struttura HTML è stata creata, è il momento di passare allo stile corretto con Tailwind.

Inizia con l'elemento `div` che avvolge l'intero contenuto. Imposta il suo colore di sfondo assegnandogli un `class` di `bg-indigo-600`.

# --hints--

Il tuo `div` esterno dovrebbe avere un attributo `class` impostato su `bg-indigo-600`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("bg-indigo-600"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
   --fcc-editable-region--
  </body>
</html>
```
