---
id: 6858eadd3144c89d45e02bc7
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Se noti, l'elemento esterno `div` è troppo attaccato in alto e il contenuto al suo interno è troppo addossato a sinistra. Quindi, ora dovresti lavorare un po' sullo spazio.

Aggiungi la classe `mt-8` a `div` per spingerlo un po' verso il basso e `p-4` per allontanare il testo dal lato sinistro.

# --hints--

Il tuo elemento esterno `div` dovrebbe avere la classe di utilità `mt-8`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

Il tuo elemento esterno `div` dovrebbe avere la classe di utilità `p-4`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
