---
id: 6858ebf2552046ab3de9f17c
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

A questo punto, noterai che c’è troppo spazio a destra del contenuto nel `div`. Puoi ridurlo aggiungendo una larghezza.

Aggiungi la classe `md:w-1/2` al `div` per ridurre la larghezza al 50% su schermi medi e più grandi. Cioè, su schermi di 768 pixel o più.

Per vedere l’effetto, espandi la scheda di anteprima.

# --hints--

Il tuo elemento esterno `div` dovrebbe avere la classe di utilità `md:w-1/2`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
