---
id: 6858f90c54647a48ddb8822d
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

A questo punto, il contenuto, cioè i due elementi `div` all'interno dell'elemento esterno `div`, saranno impilati uno sopra l'altro indipendentemente dalla dimensione dello schermo.

Questo succede perché hai la classe `flex-col` nell'elemento contenitore `div`. Per risolvere, usa il prefisso `lg:` per aggiungere la classe `flex-row` al `div`.

Questo significa che sui dispositivi con schermi più grandi, la direzione del flex sarà `row`. Per vedere di nuovo questo effetto, regola la larghezza del browser a un valore più ampio.

# --hints--

Il tuo elemento esterno `div` dovrebbe avere la classe di utilità `lg:flex-row`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
