---
id: 6858ee73ed8106ccb214942a
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Ora che tutto è allineato a sinistra, devi centrare l'intero contenuto aggiungendo un margine automatico sull'asse x.

Aggiungi la classe Tailwind corrispondente a questo al `div` esterno. Per vedere l'effetto, regola la scheda di anteprima finché l'intero contenuto non è centrato.

# --hints--

Il tuo `div` esterno dovrebbe avere la classe di utilità `mx-auto`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```
