---
id: 6858ef974c8ad5dcd7abb9f1
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Il contenuto all'interno del `div` può essere disposto e spaziato meglio. Per il resto della spaziatura e della disposizione, devi usare il flexbox.

Aggiungi la classe `flex` al `div` in modo che i due elementi figli `div` siano affiancati.

Inoltre, aggiungi `justify-around` affinché lo spazio intorno e tra i contenuti sia distribuito in modo uniforme, e `items-center` affinché l'intero contenuto sia finalmente centrato nel `div`. 

# --hints--

Il tuo elemento esterno `div` dovrebbe avere la classe di utilità `flex`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

Il tuo elemento esterno `div` dovrebbe avere la classe di utilità `justify-around`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

Il tuo elemento esterno `div` dovrebbe avere la classe di utilità `items-center`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
