---
id: 6858e6408a0999603b2a4fee
title: Passaggio 14
challengeType: 0
dashedName: step-14
---

# --description--

Per il `h1`, aggiungi le classi `font-bold` e `text-4xl`. Questo renderà il testo più in grassetto, più grande e aumenterà anche l'interlinea.

Inoltre, aggiungi la classe `my-4` per aumentare il margine superiore e inferiore del `h1` di 16 pixel, ovvero `1rem`.

# --hints--

Il tuo elemento `h1` dovrebbe avere la classe di utilità `font-bold`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

Il tuo elemento `h1` dovrebbe avere la classe di utilità `text-4xl`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

Il tuo elemento `h1` dovrebbe avere la classe di utilità `my-4`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
         --fcc-editable-region--
        <h1>Discover New Music</h1>
         --fcc-editable-region--
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
