---
id: 6858f3a67aa8580ff37016c7
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Se ridimensioni di nuovo la larghezza della scheda di anteprima a una larghezza identica a quella di un dispositivo mobile, noterai che il contenuto è sempre affiancato. Questo non è quello che vuoi sui dispositivi mobili.

Per risolvere, aggiungi la classe `flex-col` a `div`. Questo cambierà la direzione del flex in colonna, così gli elementi si impileranno uno sopra l'altro.

Ecco come usare la classe di utilità `flex-col`:

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

Il tuo elemento `div` di apertura dovrebbe avere la classe di utilità `flex-col`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
