---
id: 68597114730d7c6b8838d204
title: Passaggio 15
challengeType: 0
dashedName: step-15
---

# --description--

Il paragrafo sotto `h1` sembra a posto, quindi non ha bisogno di ulteriori stili. Puoi quindi passare agli elementi àncora all'interno del secondo elemento `div`.

Dai al primo elemento àncora le classi `bg-white`, `hover:bg-gray-100` e `text-indigo-600`. Questo gli conferisce uno sfondo bianco, un colore del testo indaco e uno sfondo leggermente più chiaro quando ci passi sopra con il mouse.

Per il secondo elemento àncora, dagli le classi `bg-pink-500`, `hover:bg-pink-600` e `text-white`.

# --hints--

Il tuo primo elemento àncora dovrebbe avere la classe di utilità `bg-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

Il tuo primo elemento àncora dovrebbe avere la classe di utilità `hover:bg-gray-100`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

Il tuo primo elemento àncora dovrebbe avere la classe di utilità `text-indigo-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

Il tuo secondo elemento àncora dovrebbe avere la classe di utilità `bg-pink-500`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

Il tuo secondo elemento àncora dovrebbe avere la classe di utilità `hover:bg-pink-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

Il tuo secondo elemento àncora dovrebbe avere la classe di utilità `text-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
