---
id: 6718d2d59337c822ecb697ff
title: Створіть програму для керування банківським рахунком
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні визначити клас з назвою `BankAccount` з конструктором. Конструктор має встановити властивість `balance` за замовчуванням у `0` та ініціалізувати порожній масив з назвою `transactions` для збереження записів транзакцій у вигляді об’єктів.

2. Кожна транзакція, що зберігається в масиві `transactions`, має бути об’єктом із двома властивостями: `type` і `amount`. Властивість `type` має бути або `deposit`, або `withdraw`, а властивість `amount` — сумою внеску або зняття.

3. Ви повинні визначити метод з назвою `deposit`, який приймає суму внеску як параметр. Коли сума внеску більша за `0`, він має:

   - Додати новий об’єкт у масив транзакцій із типом `deposit` та сумою внеску.
   - Оновити баланс.
   - Повернути `"Successfully deposited $[amount]. New balance: $[balance]"`.

4. Якщо сума менша або дорівнює `0`, метод `deposit` має повернути `"Deposit amount must be greater than zero."`.

5. Ви повинні визначити метод з назвою `withdraw`, який приймає суму як параметр. Цей метод має оновити поточний баланс відповідно до зняття коштів. Коли сума для зняття більша за `0` і менша або дорівнює поточному балансу, він має:

   - Додати новий об’єкт у масив транзакцій із типом `withdraw` та сумою зняття.
   - Оновити баланс.
   - Повернути `"Successfully withdrew $[amount]. New balance: $[balance]"`.

6. Якщо сума для зняття менша або дорівнює `0` або більша за поточний баланс, метод `withdraw` має повернути `"Insufficient balance or invalid amount."`.

7. Ви повинні визначити метод з назвою `checkBalance`, який повертає поточний баланс у форматі `"Current balance: $[balance]"`.

8. Ви повинні визначити метод з назвою `listAllDeposits`, який ітерує масив транзакцій і повертає всі внески у форматі `"Deposits: amount,amount,..."`.

9. Ви повинні визначити метод з назвою `listAllWithdrawals`, який ітерує масив транзакцій і повертає всі зняття у форматі `"Withdrawals: amount,amount,..."`.

10. Ви повинні створити новий екземпляр `BankAccount` з назвою `myAccount`.

11. Ваш банківський рахунок `myAccount` має містити щонайменше п’ять транзакцій.

12. Ваш банківський рахунок `myAccount` має містити щонайменше два внески.

13. Ваш банківський рахунок `myAccount` має містити щонайменше два зняття.

14. Баланс вашого банківського рахунку `myAccount` має бути більшим за `$100`.

# --hints--

Ви повинні визначити клас з назвою `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

Об’єкт `BankAccount` спочатку має мати властивість `balance` зі значенням `0` та порожній масив `transactions` для збереження записів транзакцій.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

Ви повинні мати метод `deposit`, який приймає суму внеску як параметр.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

Ви повинні мати метод `withdraw`, який приймає суму зняття як параметр.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

Ви повинні мати метод `checkBalance`, який перевіряє поточний баланс.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

Ви повинні мати метод `listAllDeposits`, який перераховує всі внески.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

Ви повинні мати метод `listAllWithdrawals`, який перераховує всі зняття.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` має повертати `"Successfully deposited $100. New balance: $100"`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` має повертати `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` має повертати `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

Коли баланс рахунку `100`, `BankAccount.withdraw(150)` має повертати `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` має повертати `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` має повертати `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

Коли баланс рахунку `200`, `BankAccount.withdraw(150)` має повертати `"Successfully withdrew $150. New balance: $50"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

Коли баланс рахунку `200`, `BankAccount.checkBalance()` має повертати `"Current balance: $200"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

Коли ви вносите `10`, `35`, `90`, метод `listAllDeposits` має повертати `"Deposits: 10,35,90"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

Коли ви знімаєте `20`, `50`, `100`, метод `listAllWithdrawals` має повертати `"Withdrawals: 20,50,100"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

Ви повинні мати екземпляр `BankAccount` з назвою `myAccount`.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

Ваш банківський рахунок `myAccount` має містити щонайменше п’ять транзакцій.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

Ваш банківський рахунок `myAccount` має містити щонайменше два внески.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

Ваш банківський рахунок `myAccount` має містити щонайменше два зняття.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```

Ваш банківський рахунок `myAccount` має мати баланс більший за `$100`.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
