---
id: 587d7fb0367417b2b2512bee
title: 作業用の Express サーバーを起動する
challengeType: 2
forumTopicId: 301519
dashedName: start-a-working-express-server
---

# --description--

In the first two lines of the file `myApp.js`, you can see how easy it is to create an Express app object. This object has several methods, and you will learn many of them in these challenges. One fundamental method is `app.listen(port)`. It tells your server to listen on a given port, putting it in running state. For testing reasons, the app needs to be running in the background, so this method has been added in the `server.js` file for you.

サーバーから最初の文字列を提供してみましょう！ Express では、ルートは `app.METHOD(PATH, HANDLER)` という構造をとります。 METHOD is an HTTP method in lowercase. PATH は、サーバー上の相対パスです (文字列、または正規表現でもかまいません) 。 HANDLER は、ルートが一致したときに Express が呼び出す関数です。 ハンドラーは `function(req, res) {...}` という形式をとります。req はリクエストオブジェクトで、res はレスポンスオブジェクトです。 たとえば、下記のハンドラー

```js
function(req, res) {
  res.send('Response String');
}
```

は、文字列「Response String」を提供します。

# --instructions--

`app.get()` メソッドを使用して、`/` (root) パスに一致する GET リクエストに対して文字列「Hello Express」を提供してください。 Be sure that your code works by looking at the logs, then see the results in the preview if you are using Gitpod.

**Note:** All the code for these lessons should be added in between the few lines of code that have been written for you.

# --hints--

アプリは、「Hello Express」という文字列を提供する必要があります。

```js
(getUserInput) =>
  $.get(getUserInput('url')).then(
    (data) => {
      assert.equal(
        data,
        'Hello Express',
        'Your app does not serve the text "Hello Express"'
      );
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

