---
id: 5dc1798ff86c76b9248c6eb3
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`h1`부터 `h6`까지의 제목 요소는 그 아래 내용의 중요도를 나타내는 데 사용됩니다. 숫자가 낮을수록 중요도가 높으므로 `h2` 요소는 `h1` 요소보다 중요도가 낮습니다.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

한 페이지에 `h1` 요소는 하나만 사용하고, 중요도가 낮은 제목은 중요도가 높은 제목 아래에 배치하세요.

`h1` 요소 아래에 다음 텍스트가 포함된 `h2` 요소를 추가하세요:

`Cat Photos`

# --hints--

`h1` 요소는 여는 태그가 있어야 합니다. 여는 태그의 구문은 `<elementName>`입니다.

```js
assert(document.querySelector('h1'));
```

`h1` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `<` 문자 바로 뒤에 `/`가 있습니다.

```js
assert(code.match(/<\/h1\>/));
```

`h1` 요소는 하나만 있어야 합니다. 중복된 요소를 제거하세요.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

`h1` 요소의 텍스트는 `CatPhotoApp`여야 합니다. 텍스트를 누락했거나 오타가 있습니다.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

`h2` 태그는 소문자로 작성해야 합니다. 관례상 모든 HTML 태그는 소문자로 씁니다.

```js
assert.notMatch(code, /<\/?H2>/);
```

`h2` 요소는 여는 태그가 있어야 합니다. 여는 태그의 구문은 `<elementName>`입니다.

```js
assert(document.querySelector('h2'));
```

`h2` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `<` 문자 바로 뒤에 `/`가 있습니다.

```js
assert(code.match(/<\/h2\>/));
```

`h2` 요소의 텍스트는 `Cat Photos`여야 합니다. 여는 `h2` 태그와 닫는 `h2` 태그 사이에는 `Cat Photos` 텍스트만 넣으세요.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

`h2` 요소는 `h1` 요소 아래에 있어야 합니다. `h1` 요소가 더 중요하므로 `h2` 요소 위에 있어야 합니다.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```
