---
id: 5dc17d3bf86c76b9248c6eb4
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`p` 요소는 웹사이트에서 텍스트 단락을 만들 때 사용합니다. `p` 요소 아래에 `h2` 요소를 만들고 다음 텍스트를 넣으세요:

`Everyone loves cute cats online!`

# --hints--

`p` 요소는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelector('p'));
```

`p` 태그는 소문자로 작성해야 합니다. 관례상 모든 HTML 태그는 소문자로 씁니다.

```js
assert.notMatch(code, /<\/?P>/);
```

`p` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 옵니다.

```js
assert(code.match(/<\/p\>/));
```

`p` 요소의 텍스트는 `Everyone loves cute cats online!`이어야 합니다. 텍스트를 빼먹었거나 오타가 있습니다.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

`p` 요소는 `h2` 요소 아래에 있어야 합니다. 순서가 잘못되었습니다.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```
