---
id: 5dc174fcf86c76b9248c6eb2
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML 요소는 `<h1>` 같은 여는 태그와 `</h1>` 같은 닫는 태그를 가집니다. 요소가 표시할 텍스트는 여는 태그와 닫는 태그 사이에 들어갑니다.

아래 `h1` 요소의 텍스트를 `Hello World`에서 `CatPhotoApp`로 바꾸고 브라우저 미리보기에서 변화를 확인하세요.

완료하면 "코드 확인" 버튼을 눌러 올바른지 확인하세요.

# --hints--

코드에 `CatPhotoApp` 텍스트가 있어야 합니다. 철자를 확인해 보세요.

```js
assert(code.match(/catphotoapp/i));
```

`h1` 요소에는 여는 태그가 있어야 합니다. 여는 태그 구문은 `<elementName>`입니다.

```js
assert(document.querySelector('h1'));
```

`h1` 태그는 소문자여야 합니다. 관례상 모든 HTML 태그는 소문자로 작성합니다.

```js
assert.notMatch(code, /<\/?H1>/);
```

`h1` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 옵니다.

```js
assert(code.match(/<\/h1\>/));
```

`h1` 요소가 여러 개 있습니다. 불필요한 `h1` 요소를 제거하세요.

```js
assert(document.querySelectorAll('h1').length === 1);
```

`h1` 요소의 텍스트는 `CatPhotoApp`여야 합니다. 텍스트가 빠졌거나 오타가 있거나 `h1` 요소의 여는 태그와 닫는 태그 사이에 있지 않습니다.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```
