---
id: 5dc23f9bf86c76b9248c6eba
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

웹사이트에 이미지를 추가하려면 `img` 요소를 사용하세요. `img` 요소는 닫는 태그 없이 여는 태그만 있습니다. 닫는 태그가 없는 요소를 <dfn>void element</dfn>라고 합니다.

`img` 요소 아래에 `p` 요소를 추가하세요. 이 시점에서는 브라우저에 이미지가 표시되지 않습니다.

# --hints--

`img` 요소는 여는 태그가 있어야 합니다. 여는 태그 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelector('img'));
```

`img` 요소는 닫는 태그 `</img>`가 없어야 합니다.

```js
assert(!code.match(/<\/img\>/));
```

`img` 요소는 하나만 있어야 합니다. 추가된 요소가 있다면 제거하세요.

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 요소는 `p` 요소 아래에 있어야 합니다. 순서가 잘못되어 있습니다.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
