---
id: 5dc24165f86c76b9248c6ebc
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

모든 `img` 요소에는 `alt` 속성이 있어야 합니다. `alt` 속성의 텍스트는 화면 읽기 프로그램의 접근성을 향상시키기 위해 사용되며, 이미지가 로드되지 않을 경우 표시됩니다.

다음은 `img` 요소에 `alt` 속성이 있는 예시입니다:

```html
<img src="cat.jpg" alt="A cat">
```

`img` 요소 안에 다음 텍스트를 가진 `alt` 속성을 추가하세요:

`A cute orange cat lying on its back`

# --hints--

코드에 `img` 요소가 있어야 합니다. 이전 단계에서 `img` 요소를 제거했습니다.

```js
assert(document.querySelector('img'));
```

`img` 요소에 `alt` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있거나 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

`img` 요소의 `alt` 속성 값이 `A cute orange cat lying on its back`가 아닌 다른 값으로 설정되어 있습니다. `alt` 속성 값이 따옴표로 감싸져 있는지 확인하세요.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
