---
id: 5dc24614f86c76b9248c6ebd
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

앵커(`a`) 요소로 다른 페이지에 링크할 수 있습니다.

다음은 `https://www.freecodecamp.org`에 링크하는 예시입니다:

```html
<a href="https://www.freecodecamp.org"></a>
```

단락 뒤에 `https://freecatphotoapp.com`에 링크하는 앵커 요소를 추가하세요. 이 시점에서는 미리보기에서 링크가 보이지 않습니다.

# --hints--

앵커(`a`) 요소는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelector('a'));
```

앵커(`a`) 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
assert(code.match(/<\/a\>/));
```

앵커(`a`) 요소는 `p` 요소 아래에 있어야 합니다. 순서가 잘못되었습니다.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

앵커(`a`) 요소에 `href` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있거나 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

앵커(`a`) 요소는 `https://freecatphotoapp.com`에 링크해야 합니다. URL을 빼먹었거나 오타가 있습니다.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

앵커(`a`) 요소의 `href` 속성을 올바른 링크로 설정했지만, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
