---
id: 5dfa3589eacea3f48c6300ae
title: 20단계
challengeType: 0
dashedName: step-20
---

# --description--

두 번째 `section` 요소 안에 텍스트 `h2`를 가진 새로운 `Cat Lists` 요소를 추가하세요.

# --hints--

`section` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert(
  document.querySelectorAll('section').length === 2 &&
    code.match(/<\/section>/g).length === 2
);
```

`h2` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelectorAll('h2').length === 2);
```

`h2` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
assert(code.match(/<\/h2\>/g).length === 2);
```

두 번째 `h2` 요소는 두 번째 `section` 요소의 닫는 태그 바로 위에 있어야 합니다. 현재 위치가 올바르지 않습니다.

```js
const secondSection = document.querySelectorAll('section')[1];
assert(secondSection.lastElementChild.nodeName === 'H2');
```

두 번째 `h2` 요소에는 텍스트 `Cat Lists`가 있어야 합니다. 텍스트를 빼먹었거나 오타가 있습니다.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase() === 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
