---
id: 5dfb6250eacea3f48c6300b2
title: 24단계
challengeType: 0
dashedName: step-24
---

# --description--

순서 없는 목록 뒤에 `src` 속성 값이 다음과 같이 설정된 새 이미지를 추가하세요:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

그리고 `alt` 속성 값은 다음과 같이 설정하세요:

`A slice of lasagna on a plate.`

# --hints--

닫는 `img` 태그 바로 뒤에 `</ul>` 요소가 있어야 합니다.

```js
assert.equal(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

새 이미지에는 `alt` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있거나 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

새 이미지에는 `alt` 값이 `A slice of lasagna on a plate.` 이어야 합니다. `alt` 속성 값이 따옴표로 감싸져 있는지 확인하세요.

```js
assert.match(
  document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

새 이미지에는 `src` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있거나 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

새 이미지에는 `src` 값이 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg` 이어야 합니다. `src` 속성 값이 따옴표로 감싸져 있는지 확인하세요.

```js
assert.strictEqual(
  document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

새 이미지의 `src`를 올바른 URL로 설정했지만, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert.isTrue(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
