---
id: 5ef9b03c81a63668521804db
title: 43단계
challengeType: 0
dashedName: step-43
---

# --description--

필수 정보가 누락된 상태에서 사용자가 폼을 제출하지 못하게 하려면 `required` 요소에 `input` 속성을 추가해야 합니다.

다음은 `required` 속성이 있는 입력 필드 예시입니다:

```html
<input type="text" name="firstName" required>
```

`required` 속성에 값을 설정할 필요는 없습니다. 대신 `required` 요소에 `input`라는 단어만 추가하고, 다른 속성과 사이에 공백이 있도록 하세요.

# --hints--

`input` 요소를 삭제했거나 구문이 잘못되었습니다. 모든 속성 값은 따옴표로 감싸야 합니다.

```js
assert(document.querySelectorAll('input').length);
```

`form`에는 `input` 요소만 포함되어야 합니다. `form` 요소 내의 추가 HTML 요소나 텍스트를 삭제하세요.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

`input` 요소에는 `required` 속성이 있어야 합니다. `required` 요소 태그 안에 `input`라는 단어만 추가하면 됩니다.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('required'));
```

`required` 속성에는 값을 주면 안 됩니다.

```js
assert.strictEqual(document.querySelector('input')?.getAttribute('required'), '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
