---
id: 5ef9b03c81a63668521804e1
title: 52단계
challengeType: 0
dashedName: step-52
---

# --description--

`fieldset` 요소는 웹 폼에서 관련된 입력과 레이블을 함께 그룹화하는 데 사용됩니다. `fieldset` 요소는 <dfn>블록 레벨 요소</dfn>로, 새 줄에 나타납니다.

`Indoor`와 `Outdoor` 라디오 버튼을 `fieldset` 요소 안에 중첩하고, 라디오 버튼 들을 들여쓰기하는 것을 잊지 마세요.

# --hints--

두 라디오 버튼 모두 여는 `label` 요소 태그와 닫는 태그 사이에 있어야 합니다.

```js
const labelChildNodes = [...document.querySelectorAll('label')].map((node) => [...node.childNodes]);
assert(
  labelChildNodes.filter(
    childNodes =>
      childNodes.filter(node => node.nodeName === 'INPUT').length === 1
  ).length === 2
);
```

`fieldset` 요소는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelector('fieldset'));
```

`fieldset` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
assert(code.match(/<\/fieldset\>/));
```

라디오 버튼과 관련 레이블 모두 `fieldset` 요소의 여는 태그와 닫는 태그 사이에 있어야 합니다.

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(
  radioButtons.every((btn) => btn.parentNode.parentNode.nodeName === 'FIELDSET')
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
        
          <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
          <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
        
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
