---
id: 5ef9b03c81a63668521804dd
title: 47단계
challengeType: 0
dashedName: step-47
---

# --description--

`label` 요소는 `input` 요소의 텍스트를 `input` 요소 자체와 연결하는 데 사용됩니다(특히 화면 읽기 프로그램 같은 보조 기술에 유용합니다).

다음은 `label` 요소 안에 `radio` 버튼이 있는 예시입니다:

```html
<label><input type="radio"> cat</label>
```

예시에서 `"cat"` 단어를 클릭하면 `radio` 버튼도 선택됩니다.

`radio` 버튼을 `label` 요소 안에 중첩하세요.

# --hints--

라디오 버튼이 여전히 존재하는지 확인하세요.

```js
assert(document.querySelector('input[type="radio"]'));
```

텍스트 ` Indoor`는 `radio` 버튼 바로 오른쪽에 위치해야 합니다. 요소와 텍스트 사이에 공백이 있는지 확인하세요. 텍스트를 빼먹었거나 오타가 있을 수 있습니다.

```js
const radioInputElem = document.querySelector('input');
assert.match(
  radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Indoor/i
);
```

`label` 요소는 여는 태그가 있어야 합니다. 여는 태그 구문은 `<elementName>`입니다.

```js
assert(document.querySelector('label'));
```

`label` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 있습니다 `<`.

```js
assert(code.match(/<\/label\>/));
```

라디오 버튼과 그 텍스트는 모두 `label` 요소의 여는 태그와 닫는 태그 사이에 있어야 합니다.

```js
const labelChildNodes = [...document.querySelector('form > label').childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="radio"> Indoor
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
