---
id: 5ef9b03c81a63668521804e3
title: 60단계
challengeType: 0
dashedName: step-60
---

# --description--

방금 추가한 체크박스 다음에 또 다른 체크박스를 추가하세요. `id` 속성 값은 `lazy`이어야 하고 `name` 속성 값은 마지막 체크박스와 같아야 합니다.

또한 새 체크박스 오른쪽에 `label` 요소를 추가하고 텍스트는 `Lazy`로 하세요. `label` 속성을 사용해 `for` 요소를 새 체크박스와 연결하는 것을 잊지 마세요.

# --hints--

새로운 체크박스를 추가해야 합니다.

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'), 2);
```

새 체크박스에는 `id` 속성이 `lazy` 값으로, `name` 속성이 `personality` 값으로 있어야 합니다. 여는 태그 이름 뒤에 공백이 있거나 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'lazy' && checkbox.getAttribute('name') === 'personality'
  )
);
```

새 체크박스는 첫 번째 체크박스 뒤에 있어야 합니다. 순서가 잘못되었습니다.

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('loving') < checkboxes.indexOf('lazy'));
```

새 체크박스 오른쪽에는 텍스트가 `label`인 `Lazy` 요소가 있어야 합니다.

```js
const nextElementSibling = document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling;
assert.isOk(nextElementSibling?.nodeName === 'LABEL' && nextElementSibling?.innerText?.replace(/\s+/g, '')?.match(/^Lazy$/i)
);
```

새 `label`에는 새 체크박스의 `for` 속성과 같은 값의 `id` 속성이 있어야 합니다. 값이 빠졌거나 오타가 있을 수 있습니다.

```js
assert.strictEqual(
  document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling?.getAttribute('for'),
    'lazy'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
