---
id: 5ef9b03c81a63668521804e2
title: 56단계
challengeType: 0
dashedName: step-56
---

# --description--

폼에서는 여러 답변이 가능한 질문에 체크박스를 자주 사용합니다. `input` 요소에 `type` 속성을 `checkbox`로 설정하면 체크박스가 만들어집니다.

방금 추가한 `legend` 요소 아래에 `input`를 추가하고, `type` 속성을 `checkbox`로 설정한 뒤 텍스트를 다음과 같이 지정하세요:

` Loving`

# --hints--

체크박스용 `input` 요소는 여는 태그만 있어야 하며 닫는 태그는 없어야 합니다.

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

체크박스용 input 요소는 하나만 추가해야 합니다. 추가로 넣은 요소가 있다면 제거하세요.

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

새로운 `input` 요소는 `legend` 요소 아래에 있어야 하며, 텍스트는 `What's your cat's personality?`이어야 합니다. 순서가 잘못되었습니다.

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

새로운 `input` 요소에 `type` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

새로운 `input` 요소에는 `type` 속성이 `checkbox` 값으로 있어야 합니다. 값이 빠졌거나 오타가 있습니다. 속성 값은 반드시 따옴표로 감싸야 합니다.

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

새로운 `input` 요소의 `type` 속성을 `checkbox`로 설정했지만, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

텍스트 ` Loving`는 체크박스 바로 오른쪽에 있어야 합니다. 요소와 텍스트 사이에 공백이 있는지 확인하세요. 텍스트가 빠졌거나 오타가 있습니다.

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>
            
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
