---
id: 5ef9b03c81a63668521804d1
title: 28단계
challengeType: 0
dashedName: step-28
---

# --description--

`figure` 요소 다음에 텍스트가 `h3` 인 또 다른 `Top 3 things cats hate:` 요소를 추가하세요.

# --hints--

두 번째 `h3` 요소의 닫는 태그 바로 위에 `section` 요소가 있어야 합니다. 여는 태그와 닫는 태그가 모두 있는지 확인하세요.

```js
assert(
  document.querySelectorAll('main > section')[1].lastElementChild.nodeName ===
    'H3' && code.match(/<\/h3\>/g).length === 2
);
```

새로운 `h3` 요소에는 텍스트 `Top 3 things cats hate:` 가 있어야 합니다. 텍스트 끝에 콜론이 포함되어 있는지 확인하세요.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'top 3 things cats hate:'
);
```

새로운 `figure` 요소 위에 `h3` 이 있어야 합니다. `figure` 요소를 실수로 삭제했을 수 있습니다.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  .lastElementChild;
assert(
  secondSectionLastElemNode.nodeName === 'H3' &&
    secondSectionLastElemNode.previousElementSibling.nodeName === 'FIGURE'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
