---
id: 5ef9b03c81a63668521804dc
title: 46단계
challengeType: 0
dashedName: step-46
---

# --description--

여러 선택지 중에서 하나의 답변만 선택하도록 할 때 라디오 버튼을 사용할 수 있습니다.

다음은 텍스트가 `cat`로 설정된 라디오 버튼 예시입니다:

```html
<input type="radio"> cat
```

`input` 요소는 빈 요소임을 기억하세요.

텍스트 입력 앞에 텍스트가 다음과 같이 설정된 라디오 버튼을 추가하세요:

`Indoor`

# --hints--

라디오 버튼을 위해 `input` 요소를 생성하세요. 구문을 확인하세요.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

`input` 요소는 여는 태그만 있어야 하며 닫는 태그는 없어야 합니다.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

라디오 버튼을 위해 `input` 요소는 하나만 추가해야 합니다. 추가된 것이 있다면 제거하세요.

```js
assert(document.querySelectorAll('form > input'), 2);
```

새로운 `input` 요소는 기존 `input` 위에 있어야 하며 `type` 속성은 `text`로 설정되어야 합니다. 순서가 잘못되었습니다.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

새로운 `input` 요소에 `type` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

새로운 `input` 요소는 `type` 속성을 하나만 가져야 합니다. 추가된 속성이 있다면 제거하세요.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

새로운 `input` 요소는 `type` 속성을 값 `radio`로 가져야 합니다. 값이 빠졌거나 오타가 있습니다. 속성 값은 반드시 따옴표로 감싸야 합니다.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

새로운 `input` 요소의 `type` 속성을 `radio`로 설정했지만, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

`radio` 버튼의 ` Indoor` 텍스트는 버튼 뒤에 위치해야 하며 앞에 있으면 안 됩니다.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

텍스트 ` Indoor`는 `radio` 버튼 바로 오른쪽에 위치해야 합니다. 텍스트가 빠졌거나 오타가 있습니다.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
