---
id: 5ef9b03c81a63668521804d5
title: 36단계
challengeType: 0
dashedName: step-36
---

# --description--

세 번째 `section` 요소 안에 텍스트가 `h2`인 `Cat Form` 요소를 추가하세요.

# --hints--

세 번째 `section` 요소를 찾을 수 없습니다. 실수로 삭제했거나 여는 태그 또는 닫는 태그가 없을 수 있습니다.

```js
assert(
  document.querySelectorAll('section').length === 3 &&
    code.match(/<\/section>/g).length === 3
);
```

`h2` 요소에는 여는 태그와 닫는 태그가 있어야 합니다. 필요한 태그 중 하나 또는 둘 다 빠졌을 수 있습니다.

```js
assert(
  document.querySelectorAll('h2').length >= 3 &&
    code.match(/<\/h2>/g).length >= 3
);
```

하나의 `h2` 요소만 추가해야 합니다. 추가된 요소가 있다면 삭제하세요.

```js
assert(document.querySelectorAll('h2').length === 3);
```

새로운 `h2` 요소는 세 번째 `section` 요소의 닫는 태그 바로 위에 위치해야 합니다.

```js
const thirdSection = document.querySelectorAll('section')[2];
assert(thirdSection.lastElementChild.nodeName === 'H2');
```

`h2` 요소의 텍스트는 `Cat Form`이어야 합니다.

```js
const thirdSection = document.querySelectorAll('section')[2];
assert(
  thirdSection
    .querySelector('h2')
    .innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'cat form'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>
        </figure>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
