---
id: 5ef9b03c81a63668521804e5
title: 63단계
challengeType: 0
dashedName: step-63
---

# --description--

체크박스를 기본적으로 선택하거나 라디오 버튼을 기본 선택 상태로 만들려면 `checked` 속성을 추가해야 합니다.

다음은 `checked` 속성이 있는 라디오 버튼 예시입니다:

```html
<input checked type="radio" name="meal" value="breakfast"> Breakfast
```

`checked` 속성에 값을 설정할 필요는 없습니다. 대신 `checked` 요소에 `input`라는 단어를 추가하고, 다른 속성과 띄어쓰기가 되어 있는지 확인하세요.

첫 번째 라디오 버튼과 첫 번째 체크박스를 기본 선택 상태로 만드세요.

# --hints--

각각의 `fieldset` 요소 안에 여전히 두 개의 라디오 버튼과 세 개의 체크박스가 중첩되어 있는지 확인하세요.

```js
assert.isTrue(
  document.querySelectorAll('input[type="radio"]').length === 2 &&
     document.querySelectorAll('fieldset > input[type="checkbox"]').length === 3
);
```

`id`가 `indoor`로 설정된 첫 번째 라디오 버튼에는 `checked` 속성이 있어야 합니다.

```js
assert.isTrue(document.querySelector('input[type="radio"]')?.hasAttribute('checked'));
```

두 번째 라디오 버튼에는 `checked` 속성이 없어야 합니다.

```js
assert.isFalse(document.querySelectorAll('input[type="radio"]')?.[1]?.hasAttribute('checked'));
```

첫 번째 체크박스에 `checked` 속성이 빠져 있습니다.

```js
assert.isTrue(document.querySelector('input[type="checkbox"]')?.hasAttribute('checked'));
```

두 번째 체크박스에는 `checked` 속성이 없어야 합니다.

```js
assert.isFalse(document.querySelectorAll('input[type="checkbox"]')?.[1]?.hasAttribute('checked'));
```

세 번째 체크박스에는 `checked` 속성이 없어야 합니다.

```js
assert.isFalse(document.querySelectorAll('input[type="checkbox"]')?.[2]?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic"> Energetic</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
