---
id: 5ef9b03c81a63668521804d9
title: 42단계
challengeType: 0
dashedName: step-42
---

# --description--

자리 표시자 텍스트는 입력에 어떤 종류의 정보를 입력해야 하는지 힌트를 주기 위해 사용합니다.

다음은 자리 표시자가 `input`로 설정된 `e.g. Jane Doe` 요소의 예입니다:

```html
<input type="text" placeholder="e.g. Jane Doe">
```

`cat photo URL` 요소에 자리 표시자 텍스트 `input`를 추가하세요.

# --hints--

`input` 요소를 삭제했거나 구문이 올바르지 않습니다. 모든 속성 값은 따옴표로 감싸야 합니다.

```js
assert(document.querySelectorAll('input').length);
```

`form`에는 `input` 요소만 포함되어야 합니다. `form` 요소 내에 추가 HTML 요소나 텍스트가 있다면 삭제하세요.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

`input` 요소에 `placeholder` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있는지, 그리고 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('placeholder'));
```

`input` 요소에는 값이 `placeholder`인 `cat photo URL` 속성이 있어야 합니다. 값을 누락했거나 오타가 있을 수 있습니다. 속성 값은 반드시 따옴표로 감싸야 한다는 점을 기억하세요.

```js
assert(
  document.querySelector('input')
    ?.getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

`input` 요소의 `placeholder` 속성을 `cat photo URL`로 설정했지만, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
