---
id: 5ef9b03c81a63668521804d7
title: 38단계
challengeType: 0
dashedName: step-38
---

# --description--

`action` 속성은 폼 데이터가 전송될 위치를 나타냅니다.

다음은 `form` 속성이 있는 `action` 요소의 예입니다:

```html
<form action="/submit-url"></form>
```

예제에서 `action="/submit-url"`는 브라우저에 폼 데이터를 `/submit-url` 경로로 전송해야 한다고 알려줍니다.

`action` 요소에 값이 `https://freecatphotoapp.com/submit-cat-photo`인 `form` 속성을 추가하세요.

# --hints--

`form` 요소는 올바른 순서로 여는 태그와 닫는 태그를 가져야 합니다. 필요한 태그 하나 또는 둘이 없거나 순서가 잘못되었을 수 있습니다.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

마지막 `form` 요소 안에 중첩된 `section` 요소는 `h2` 요소 아래에 있어야 합니다. `h2` 요소와 `form` 요소의 순서가 잘못되었습니다.

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

`form` 요소는 내용이 없어야 합니다. `form` 요소의 태그 사이에 있는 모든 HTML 요소나 텍스트를 제거하세요.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

`form` 요소에 `action` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있는지, 그리고 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

`form` 요소는 값이 `action`인 `https://freecatphotoapp.com/submit-cat-photo` 속성을 가져야 합니다.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

`action` 속성을 올바른 URL로 설정했지만, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
