---
id: 5ef9b03c81a63668521804d8
title: 39단계
challengeType: 0
dashedName: step-39
---

# --description--

`input` 요소는 웹 폼에서 데이터를 수집하는 여러 방법을 제공합니다. `img` 요소처럼 `input` 요소는 <dfn>빈 요소</dfn>이며 닫는 태그가 필요 없습니다.

`input` 요소를 `form` 요소 안에 중첩하세요.

# --hints--

`form` 요소는 올바른 순서로 여는 태그와 닫는 태그가 있어야 합니다. 필요한 태그 하나 또는 둘이 없거나 순서가 잘못되었을 수 있습니다.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

`form` 요소의 여는 태그에는 `action` 속성만 있어야 합니다. 그 외에 입력한 내용은 삭제하세요.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

`input` 요소를 생성해야 합니다. 구문을 확인하세요.

```js
assert(document.querySelector('input'));
```

`input` 요소는 여는 태그만 있어야 하며 닫는 태그는 없어야 합니다.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

`input` 요소는 `form` 요소 안에 중첩되어야 합니다.

```js
assert(document.querySelector('form > input'));
```

`form` 요소는 `input` 요소만 포함해야 합니다. `form` 요소 태그 사이에 있는 다른 HTML 요소나 텍스트는 삭제하세요.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
