---
id: 5ef9b03c81a63668521804d0
title: 27단계
challengeType: 0
dashedName: step-27
---

# --description--

특정 단어나 구절에 강조를 주려면 `em` 요소를 사용할 수 있습니다.

`love` 요소 안에서 `figcaption` 단어를 강조 `em` 요소로 감싸세요.

# --hints--

강조 `em` 요소는 여는 태그가 있어야 합니다. 여는 태그 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelector('em'));
```

강조 `em` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
assert(code.match(/<\/em\>/));
```

`figcaption` 요소를 삭제했거나 여는 태그 또는 닫는 태그가 없습니다.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

강조 `em` 요소는 `love` 텍스트를 감싸야 합니다. 텍스트를 빼먹었거나 오타가 있습니다.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

`figcaption`의 텍스트는 `Cats love lasagna`이어야 합니다. 오타가 없는지, `em` 요소의 여는 태그와 닫는 태그 주변에 필요한 공백이 있는지 확인하세요.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
