---
id: 5efb23e70dc218d6c85f89b1
title: 40단계
challengeType: 0
dashedName: step-40
---

# --description--

`type` 속성을 사용해 만들 수 있는 입력 유형이 다양합니다. 비밀번호 필드, 리셋 버튼, 또는 사용자가 컴퓨터에서 파일을 선택할 수 있는 컨트롤을 쉽게 만들 수 있습니다.

`type` 요소에 `text` 속성을 값 `input`로 추가해 사용자로부터 텍스트 입력을 받는 텍스트 필드를 만드세요.

# --hints--

`input` 요소를 삭제했거나 구문이 올바르지 않습니다. 속성을 추가했다면 값이 따옴표로 감싸져 있는지 확인하세요.

```js
assert(document.querySelectorAll('input').length);
```

`form` 요소는 `input` 요소만 포함해야 합니다. `form` 요소 태그 사이에 불필요한 HTML 요소나 텍스트가 있다면 삭제하세요.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

`input` 요소에 `type` 속성이 값 `text`으로 설정되어 있지 않습니다. 여는 태그 이름 뒤에 공백이 있는지, 그리고 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

`input` 요소에 `type` 속성이 값 `text`으로 있어야 합니다. 값이 빠졌거나 오타가 있을 수 있습니다. 속성 값은 반드시 따옴표로 감싸야 한다는 점을 기억하세요.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type').replace(/\s+/g, ' '),
  /^text$/i
);
```

`input` 요소의 `type` 속성을 `text`으로 설정했지만, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
