---
id: 5efb2c990dc218d6c85f89b2
title: 45단계
challengeType: 0
dashedName: step-45
---

# --description--

텍스트 입력 아래에 버튼을 추가했지만, 페이지에서는 버튼과 입력이 나란히 나타납니다. 이는 `input`와 `button` 요소가 모두 <dfn>인라인 요소</dfn>이기 때문에 새 줄에 나타나지 않기 때문입니다.

추가한 버튼은 기본적으로 폼을 제출합니다. 하지만 기본 동작에만 의존하면 혼란이 생길 수 있습니다. `type`에 `submit` 속성을 값 `button`으로 추가해 제출 버튼임을 명확히 하세요.

# --hints--

`button` 요소는 여는 태그가 있어야 합니다. 여는 태그 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelector('button'));
```

`button` 요소는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 있습니다.

```js
assert(code.match(/<\/button\>/));
```

`button` 요소에는 `type` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('type'));
```

`button` 요소에는 값이 `type`인 `submit` 속성이 있어야 합니다. 값이 빠졌거나 오타가 있을 수 있습니다. 속성 값은 반드시 따옴표로 감싸야 합니다.

```js
assert.match(
  document.querySelector('button')?.getAttribute('type'),
  /^submit$/i
);
```

`button` 요소의 `type` 속성을 `submit`으로 설정했지만, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
