---
id: 5efc4f528d6a74d05e68af74
title: 58단계
challengeType: 0
dashedName: step-58
---

# --description--

`input` 요소의 텍스트를 요소 자체와 연결하는 또 다른 방법이 있습니다. 텍스트를 `label` 요소 안에 중첩하고, `for` 요소의 `input` 속성과 같은 값을 가진 `id` 속성을 추가할 수 있습니다.

아래와 같은 `input` 요소가 주어졌을 때:

```html
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

이 `label` 요소와 연결된 `input` 요소의 예는 다음과 같습니다:

```html
<label for="breakfast">Breakfast</label>
```

체크박스와 텍스트 `Loving`를 연결하려면 텍스트 `Loving`만 `label` 요소 안에 중첩하고 적절한 `for` 속성을 부여하세요.

# --hints--

체크박스 요소가 여전히 존재하는지, 그리고 변경하지 않았는지 확인하세요. 시작 코드와 동일하게 유지되어야 합니다.

```js
const checkboxElementRegex = /\s*<input\s+id\s*=\s*"loving"\s*type\s*=\s*"checkbox"\s*>/;
assert.match(code, checkboxElementRegex);
```

체크박스에는 여전히 값이 `id`인 `loving` 속성이 있어야 합니다. `--fcc-expected--`가 예상되었으나 `--fcc-actual--`가 발견되었습니다.

```js
const checkboxElementId = document.querySelector('input[type="checkbox"]')?.getAttribute('id');
assert.strictEqual(checkboxElementId, 'loving');
```

체크박스 요소 뒤에 새 `label` 요소를 추가해야 합니다. 여는 태그와 닫는 태그가 모두 있는지 확인하세요.

```js
assert.lengthOf(document.querySelectorAll('label'), 3);
assert.lengthOf(code.match(/<\/label\>/g), 3);
```

`label` 요소는 체크박스 뒤에 위치해야 합니다.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElement = checkboxElement?.nextElementSibling?.nodeName;
assert.strictEqual(checkboxSiblingElement, 'LABEL');
```

`label` 요소에는 텍스트 `Loving`가 있어야 합니다.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementText = checkboxElement?.nextElementSibling?.innerText;

assert.strictEqual(checkboxSiblingElementText, 'Loving');
```

레이블 텍스트 `Loving`는 새 `label` 요소 안에만 있어야 합니다. `label` 요소 밖에 텍스트가 중복되어 있지 않은지 확인하세요.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingNode = checkboxElement?.nextSibling?.nodeValue ?? '';
const labelElement = checkboxElement?.nextElementSibling;
const labelSiblingNode = labelElement?.nextSibling?.nodeValue ?? '';

assert.notInclude(checkboxSiblingNode, 'Loving');
assert.notInclude(labelSiblingNode, 'Loving');
```

`label` 요소에 `for` 속성이 없습니다. 여는 `label` 태그 이름과 `for` 속성 사이에 공백이 있는지 확인하세요.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.isNotNull(checkboxSiblingElementAttr);
```

`label` 요소에는 값이 `for`인 `loving` 속성이 있어야 합니다. `--fcc-expected--`가 예상되었으나 `--fcc-actual--`가 발견되었습니다.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.strictEqual(checkboxSiblingElementAttr, 'loving');
```

체크박스와 새 `label` 요소 사이에 공백이 있어야 합니다.

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
