---
id: 5efc518e8d6a74d05e68af75
title: 59단계
challengeType: 0
dashedName: step-59
---

# --description--

체크박스 `name` 요소에 값이 `personality`인 `input` 속성을 추가하세요.

브라우저에서는 눈에 띄지 않지만, 이렇게 하면 특히 여러 개의 체크박스가 있을 때 서버가 웹 폼을 처리하기가 더 쉬워집니다.

# --hints--

체크박스가 여전히 존재하는지 확인하세요.

```js
assert(document.querySelector('input[type="checkbox"]'));
```

체크박스 `input` 요소에는 `name` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelector('input[type="checkbox"]')?.hasAttribute('name'));
```

체크박스 `input` 요소에는 값이 `name`인 `personality` 속성이 있어야 합니다. 값을 빼먹었거나 오타가 있을 수 있습니다. 속성 값은 반드시 따옴표로 감싸야 한다는 점을 기억하세요.

```js
assert.match(
  document.querySelector('input[type="checkbox"]')?.getAttribute('name'),
  /^personality$/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
