---
id: 5f05a1d8e233dff4a68508d8
title: 49단계
challengeType: 0
dashedName: step-49
---

# --description--

첫 번째 라디오 버튼 아래에 또 다른 라디오 버튼을 만드세요. `label` 요소 안에 중첩하고 `Outdoor`을 `label` 텍스트로 사용하세요. 라디오 버튼에 `id` 속성을 주고 값으로 `outdoor`를 지정하세요.

# --hints--

원래 `Indoor` 라디오 버튼 `input`만 `id`가 `indoor`로 설정되어 있어야 합니다. 필요하면 단계를 다시 시작해 원래 HTML을 복원하세요.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

`Indoor` 라디오 버튼에는 어떤 변경도 하지 마세요. 필요하면 단계를 다시 시작해 원래 HTML을 복원하세요.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

새로운 `input` 요소를 하나만 만들고 새 `label` 요소 안에 중첩하세요. 새 `label`는 여는 태그와 닫는 태그가 모두 있어야 합니다.

```js
assert(document.querySelectorAll('label input').length === 2);
```

어느 `label` 여는 태그에도 속성을 추가하지 마세요.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

새 `input`에는 `id` 속성이 있고 값은 `outdoor`이어야 합니다.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

새 `input`에는 `type` 속성이 있고 값은 `radio`이어야 합니다.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

`input`가 `label` 안에 중첩된 것 외에는 새 요소를 추가하지 마세요.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

새 `input` 왼쪽에는 텍스트가 없어야 합니다.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

새 라디오 버튼의 레이블 텍스트는 대소문자까지 정확히 `Outdoor`이어야 합니다.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

새 라디오 버튼과 레이블은 `Indoor` 라디오 버튼과 레이블 바로 아래/뒤에 있어야 하며, 그 사이에 다른 태그가 없어야 합니다.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
