---
id: 5f07fb1579dc934717801375
title: 35단계
challengeType: 0
dashedName: step-35
---

# --description--

새로운 섹션을 추가할 시간입니다. 두 번째 `section` 요소 아래에 세 번째 `section` 요소를 추가하세요.

# --hints--

`section` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 3);
```

여는 `section` 태그는 하나만 추가하세요. 추가된 태그가 있다면 제거하세요.

```js
assert(document.querySelectorAll('section').length === 3);
```

`section` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 옵니다.

```js
assert(code.match(/<\/section>/g).length >= 3);
```

닫는 `section` 태그는 하나만 추가하세요. 추가된 태그가 있다면 제거하세요.

```js
assert(code.match(/<\/section>/g).length === 3);
```

모든 `section` 요소는 `main` 요소의 여는 태그와 닫는 태그 사이에 있어야 합니다.

```js
const childrenOfMain = [...document.querySelector('main').children];
const sectionElemsFound = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(sectionElemsFound.length === 3);
```

마지막 `section` 요소는 내용이 없어야 합니다. `section` 요소 안의 모든 HTML 요소나 텍스트를 제거하세요.

```js
assert.lengthOf(document.querySelectorAll('main > section')?.[2]?.children, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
