---
id: 5f07c98cdb9413cbd4b16750
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

새로운 섹션을 추가할 시간입니다. 기존 `section` 요소 아래에 두 번째 `section` 요소를 추가하세요.

# --hints--

`section` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 2);
```

여는 `section` 태그는 하나만 추가하세요. 추가된 태그가 있으면 제거하세요.

```js
assert(document.querySelectorAll('section').length === 2);
```

`section` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 옵니다.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

닫는 `section` 태그는 하나만 추가하세요. 추가된 태그가 있으면 제거하세요.

```js
assert(code.match(/<\/section>/g).length === 2);
```

두 번째 `section` 요소는 첫 번째 `section` 요소 안에 중첩되어 있으면 안 됩니다.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

두 `section` 요소 모두 `main` 요소의 여는 태그와 닫는 태그 사이에 있어야 합니다.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
