---
id: 5f0d4ab1b435f13ab6550052
title: 54단계
challengeType: 0
dashedName: step-54
---

# --description--

다음으로, 새로운 form `input` 요소를 추가할 것입니다. 현재 `fieldset` 요소 바로 아래에 또 다른 `fieldset` 요소를 추가하세요.

# --hints--

새로운 `fieldset` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert(document.querySelectorAll('fieldset').length >= 2);
```

여는 `fieldset` 태그는 하나만 추가하세요. 추가된 태그가 있으면 제거하세요.

```js
assert(document.querySelectorAll('fieldset').length === 2);
```

새로운 `fieldset` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 바로 뒤에 `<` 문자가 옵니다.

```js
assert(code.match(/<\/fieldset>/g).length >= 2);
```

닫는 `fieldset` 태그는 하나만 추가하세요. 추가된 태그가 있으면 제거하세요.

```js
assert(code.match(/<\/fieldset>/g).length === 2);
```

두 번째 `fieldset` 요소는 첫 번째 `fieldset` 요소 안에 중첩되어 있으면 안 됩니다.

```js
const childrenOf1stFieldset = [
  ...document.querySelector('form > fieldset').children
];
const foundElems = childrenOf1stFieldset.filter((child) => {
  return child.nodeName === 'FIELDSET';
});
assert(foundElems.length === 0);
```

두 `fieldset` 요소는 텍스트 필드 위에 있어야 합니다. 순서가 잘못되어 있습니다.

```js
const formChildren = document.querySelector('form')?.children;
assert.isTrue(
  formChildren?.[0]?.nodeName === 'FIELDSET' &&
  formChildren?.[1]?.nodeName === 'FIELDSET' &&
  formChildren?.[2]?.nodeName === 'INPUT' &&
  formChildren?.[2]?.getAttribute('type') === 'text'
);
```

새로운 `fieldset` 요소는 기존 `fieldset` 요소 아래에 있어야 합니다. 순서가 잘못되어 있습니다.

```js
const fieldsetChildren = [...document.querySelectorAll('fieldset')].map(
  (elem) => elem.children
);
assert(fieldsetChildren[0].length > fieldsetChildren[1].length);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
