---
id: 5f0d4d04b435f13ab6550053
title: 55단계
challengeType: 0
dashedName: step-55
---

# --description--

두 번째 `legend` 요소 안에 텍스트 `What's your cat's personality?`를 가진 `fieldset` 요소를 추가하세요.

# --hints--

두 번째 `fieldset` 요소를 삭제했거나 여는 태그 또는 닫는 태그가 없을 수 있습니다.

```js
assert(
  document.querySelectorAll('fieldset').length === 2 &&
    code.match(/<\/fieldset>/g).length === 2
);
```

`legend` 요소에는 여는 태그가 있어야 합니다. 여는 태그 구문은 다음과 같습니다: `<elementName>`.

```js
const secondFieldset = document.querySelectorAll('fieldset')?.[1];
assert(
  secondFieldset &&
    [...secondFieldset.children].filter((child) => child.nodeName === 'LEGEND')
      .length
);
```

`legend` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그는 `/` 문자 바로 뒤에 `<`가 옵니다.

```js
assert(code.match(/<\/legend\>/g).length === 2);
```

`legend` 요소에는 텍스트 `What's your cat's personality?`가 있어야 합니다. 텍스트를 생략했거나 오타가 있을 수 있습니다.

```js
const secondFieldset = document.querySelectorAll('fieldset')?.[1];
assert(
  secondFieldset &&
    [...secondFieldset.children].filter((child) => {
      const extraSpacesRemoved = child.innerText.replace(/\s+/g, ' ');
      return (
        child.nodeName === 'LEGEND' &&
        extraSpacesRemoved.match(/What's your cat's personality\??$/i)
      );
    }).length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
--fcc-editable-region--
          <fieldset>
            
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
