---
id: 5f1a80975fc4bcae0edb3497
title: 51단계
challengeType: 0
dashedName: step-51
---

# --description--

`Indoor` 라디오 버튼을 선택하고 폼을 제출하면, 버튼의 폼 데이터는 `name`과 `value` 속성을 기반으로 합니다. 라디오 버튼에 `value` 속성이 없으므로, 폼 데이터에는 여러 버튼이 있을 때 유용하지 않은 `indoor-outdoor=on`가 포함됩니다.

두 라디오 버튼 모두에 `value` 속성을 추가하세요. 편의성을 위해 버튼의 `value` 속성을 `id` 속성과 동일한 값으로 설정하세요.

# --hints--

두 라디오 버튼 모두 여는 태그와 닫는 `label` 요소 태그 사이에 있어야 합니다.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')].map(
  (node) => node.childNodes
);
assert(
  labelChildNodes.filter((childNode) => childNode[0].nodeName === 'INPUT')
    .length === 2
);
```

두 라디오 버튼 모두 `value` 속성을 가져야 합니다. 여는 태그 이름 뒤에 공백이 있거나 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(radioButtons.every((btn) => btn.hasAttribute('value')));
```

두 라디오 버튼 모두 `type` 속성을 가져야 합니다. 여는 태그 이름 뒤에 공백이 있거나 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
const radioButtons = [...document.querySelectorAll('input')];
assert(radioButtons.every((btn) => btn.hasAttribute('type')));
```

`Indoor` 라디오 버튼의 `value` 속성은 `indoor`으로 설정되어야 합니다. 값을 빼먹었거나 오타가 있을 수 있습니다. 속성 값은 반드시 따옴표로 감싸야 한다는 점을 기억하세요.

```js
const indoorRadioButton = document.querySelector('#indoor');
assert(indoorRadioButton.getAttribute('value').match(/^indoor$/));
```

`Outdoor` 라디오 버튼의 `value` 속성은 `outdoor`으로 설정되어야 합니다. 값을 빼먹었거나 오타가 있을 수 있습니다. 속성 값은 반드시 따옴표로 감싸야 한다는 점을 기억하세요.

```js
const outdoorRadioButton = document.querySelector('#outdoor');
assert(outdoorRadioButton.getAttribute('value').match(/^outdoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
          <label><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
