---
id: 671141d8e32fe934c26fa1be
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

앵커 요소 앞에 `See more `를, 앵커 요소 뒤에 ` in our gallery`를 추가하세요.

# --hints--

기존 앵커 요소를 사용해야 하며, 새로 만들지 마세요.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

앵커 요소 안의 텍스트는 변경하지 말고 `cat photos`여야 합니다.

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

코드에 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery`가 있어야 합니다.

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
