---
id: 7cf9b03d81a65668421804c3
title: 41단계
challengeType: 0
dashedName: step-41
---

# --description--

폼의 데이터가 `action` 속성에 지정된 위치에서 접근되려면, 텍스트 필드에 `name` 속성을 주고 제출되는 데이터를 나타내는 값을 할당해야 합니다.

다음은 `input` 속성이 있는 `name` 요소의 예입니다:

```html
<input type="text" name="name">
```

텍스트 필드에 값이 `name`인 `catphotourl` 속성을 추가하세요.

# --hints--

`input` 요소를 삭제했거나 구문이 올바르지 않습니다. 모든 속성 값은 따옴표로 감싸야 합니다.

```js
assert(document.querySelectorAll('input').length);
```

`form`에는 `input` 요소만 포함되어야 합니다. `form` 요소 내의 추가 HTML 요소나 텍스트를 제거하세요.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

`input` 요소에 `name` 속성이 없습니다. 여는 태그 이름 뒤에 공백이 있는지, 그리고 모든 속성 이름 앞에 공백이 있는지 확인하세요.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

`input` 요소에는 값이 `name`인 `catphotourl` 속성이 있어야 합니다. 값을 빼먹었거나 오타가 있을 수 있습니다. 속성 값은 반드시 따옴표로 감싸야 한다는 점을 기억하세요.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

`input` 요소의 `name` 속성을 `catphotourl`로 설정했지만, 속성 값은 항상 따옴표로 감싸는 것이 좋습니다.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
