---
id: 587d7b7a367417b2b2512b12
title: 使用 slice() 複製陣列項目
challengeType: 1
forumTopicId: 301158
dashedName: copy-array-items-using-slice
---

# --description--

接下來我們將涵蓋的行為是 `slice()`。`slice()` 不會修改陣列，而是將指定數量的元素複製或*擷取*到一個新的陣列，並且不會改變被呼叫的陣列。`slice()` 只接受 2 個參數——第一個是開始擷取的索引，第二個是停止擷取的索引（擷取會進行到該索引之前，但不包含該索引的元素）。請參考以下範例：

```js
let weatherConditions = ['rain', 'snow', 'sleet', 'hail', 'clear'];

let todaysWeather = weatherConditions.slice(1, 3);
```

`todaysWeather` 會有值 `['snow', 'sleet']`，而 `weatherConditions` 仍會有 `['rain', 'snow', 'sleet', 'hail', 'clear']`。

實際上，我們是從現有的陣列中提取元素，並創建了一個新的陣列。

# --instructions--

我們已定義一個函式（程式）`forecast`，該函式以一個陣列作為引數。使用 `slice()` 修改該函式，從引數陣列中擷取資訊並傳回一個包含字串元素 `warm` 和 `sunny` 的新陣列。

# --hints--

`forecast` 應該傳回 `["warm", "sunny"]`

```js
assert.deepEqual(
  forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']),
  ['warm', 'sunny']
);
```

`forecast` 函式應該使用 `slice()` 方法。

```js
assert(/\.slice\(/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function forecast(arr) {
  // Only change code below this line

  return arr;
}

// Only change code above this line
console.log(forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']));
```

# --solutions--

```js
function forecast(arr) {
  return arr.slice(2,4);
}
```
