---
id: 587d7b7c367417b2b2512b18
title: 為你的 JavaScript 物件添加鍵值配對
challengeType: 1
forumTopicId: 301153
dashedName: add-key-value-pairs-to-javascript-objects
---

# --description--

在最基本的層次上，物件只是 <dfn>鍵值</dfn> 對組的群集。換句話說，它們是映射到稱為 <dfn>屬性</dfn>（<dfn>鍵</dfn>）的唯一識別符的資料（<dfn>值</dfn>）片段。看看以下範例：

```js
const tekkenCharacter = {
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true
};
```

上述程式碼定義了一個名為 `tekkenCharacter` 的 Tekken 遊戲字元物件。它有三個屬性，每個屬性都對應到特定的值。如果你想為物件添加額外的屬性，例如「origin」，可以透過將 `origin` 指定給該物件來完成：

```js
tekkenCharacter.origin = 'South Korea';
```

這使用點記法。如果你觀察 `tekkenCharacter` 物件，它現在會包含 `origin` 屬性。Hwoarang 也有獨特的橘色頭髮。你可以透過括號記法將此屬性添加進去，方法如下：

```js
tekkenCharacter['hair color'] = 'dyed orange';
```

如果你的屬性中有空格，或你想使用變數來命名屬性，則必須使用括號表示法。在上述情況中，屬性以引號括起來以表示它是字串，並且會如實添加。若沒有引號，則會被計算為變數，屬性的名稱將是該變數的值。以下是一個使用變數的範例：

```js
const eyes = 'eye color';

tekkenCharacter[eyes] = 'brown';
```

在添加所有範例後，該物件將會是這樣：

```js
{
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true,
  origin: 'South Korea',
  'hair color': 'dyed orange',
  'eye color': 'brown'
};
```

# --instructions--

已經創建了一個包含三個條目的 `foods` 物件。使用你選擇的語法，為它添加三個條目：`bananas`，值為 `13`，`grapes`，值為 `35`，以及 `strawberries`，值為 `27`。

# --hints--

`foods` 應該是一個物件。

```js
assert(typeof foods === 'object');
```

`foods` 物件應該有一個鍵為 `bananas`，其值為 `13`。

```js
assert(foods.bananas === 13);
```

`foods` 物件應該有一個鍵為 `grapes`，其值為 `35`。

```js
assert(foods.grapes === 35);
```

`foods` 物件應該有一個鍵為 `strawberries`，其值為 `27`。

```js
assert(foods.strawberries === 27);
```

`foods` 物件的定義不應該被更改。

```js
assert(
  __helpers.removeJSComments(code).search(/let foods/) === -1 &&
  __helpers.removeJSComments(code).search(/const\s+foods\s*=\s*{\s*apples:\s*25,\s*oranges:\s*32,\s*plums:\s*28\s*};/
) !== -1
);
```

# --seed--

## --seed-contents--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

foods['bananas'] = 13;
foods['grapes']  = 35;
foods['strawberries'] = 27;
```
