---
id: 587d7b7c367417b2b2512b1a
title: 使用中括號表示法存取屬性名稱
challengeType: 1
forumTopicId: 301150
dashedName: access-property-names-with-bracket-notation
---

# --description--

在第一個物件挑戰中，我們提到使用括號表示法作為一種透過變數評估來存取屬性值的方法。例如，想像我們的 `foods` 物件正在一個超市收銀機的程式中使用。我們有某個函式（程式）設定了 `selectedFood`，並且我們想檢查 `foods` 物件中是否存在該食物。這可能看起來像：

```js
let selectedFood = getCurrentFood(scannedItem);
let inventory = foods[selectedFood];
```

這段程式碼會計算儲存在 `selectedFood` 變數中的值，並傳回 `foods` 物件中該鍵的值，如果不存在則傳回 `undefined`。括號表示法非常有用，因為有時物件屬性在執行期之前並不清楚，或者我們需要以更動態的方式存取它們。

# --instructions--

我們已定義一個函式（程式）`checkInventory`，它接收一個掃描的條款作為引數。傳回 `foods` 物件中 `scannedItem` 鍵的當前值。你可以假設只會提供有效的鍵作為 `checkInventory` 的引數。

# --hints--

`checkInventory` 應該是一個函式（程式）。

```js
assert.strictEqual(typeof checkInventory, 'function');
```

`foods` 物件應該只包含以下鍵值配對：`apples: 25`、`oranges: 32`、`plums: 28`、`bananas: 13`、`grapes: 35`、`strawberries: 27`。

```js
assert.deepEqual(foods, {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
});
```

`checkInventory("apples")` 應該傳回 `25`。

```js
assert.strictEqual(checkInventory('apples'), 25);
```

`checkInventory("bananas")` 應該傳回 `13`。

```js
assert.strictEqual(checkInventory('bananas'), 13);
```

`checkInventory("strawberries")` 應該傳回 `27`。

```js
assert.strictEqual(checkInventory('strawberries'), 27);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(checkInventory("apples"));
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  return foods[scannedItem];
}
```
