---
id: 587d7b7b367417b2b2512b14
title: 使用 indexOf() 檢查元素是否存在
challengeType: 1
forumTopicId: 301154
dashedName: check-for-the-presence-of-an-element-with-indexof
---

# --description--

由於陣列可以隨時被更改或*突變*，因此無法保證特定的資料會出現在給定陣列的哪個位置，或該元素是否仍然存在。幸運的是，JavaScript 提供了另一個內建方法 `indexOf()`，讓我們能夠快速且輕鬆地檢查陣列中是否存在某個元素。`indexOf()` 會將一個元素作為參數，當被呼叫時，它會傳回該元素的位置或索引，如果該元素不存在於陣列中，則傳回 `-1`。

例如：

```js
let fruits = ['apples', 'pears', 'oranges', 'peaches', 'pears'];

fruits.indexOf('dates');
fruits.indexOf('oranges');
fruits.indexOf('pears');
```

`indexOf('dates')` 傳回 `-1`，`indexOf('oranges')` 傳回 `2`，而 `indexOf('pears')` 傳回 `1`（每個元素存在的第一個索引）。

# --instructions--

`indexOf()` 對於快速檢查陣列中是否存在某個元素非常有用。我們已經定義了一個函式（程式）`quickCheck`，它將一個陣列和一個元素作為引數。使用 `indexOf()` 修改該函式，使其在傳入的元素存在於陣列中時傳回 `true`，否則傳回 `false`。

# --hints--

`quickCheck` 函式（程式）應該傳回布林值（`true` 或 `false`），而不是字串（`"true"` 或 `"false"`）。

```js
assert.isBoolean(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

`quickCheck(["squash", "onions", "shallots"], "mushrooms")` 應該傳回 `false`

```js
assert.strictEqual(
  quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'),
  false
);
```

`quickCheck(["onions", "squash", "shallots"], "onions")` 應該傳回 `true`

```js
assert.strictEqual(
  quickCheck(['onions', 'squash', 'shallots'], 'onions'),
  true
);
```

`quickCheck([3, 5, 9, 125, 45, 2], 125)` 應該傳回 `true`

```js
assert.strictEqual(quickCheck([3, 5, 9, 125, 45, 2], 125), true);
```

`quickCheck([true, false, false], undefined)` 應該傳回 `false`

```js
assert.strictEqual(quickCheck([true, false, false], undefined), false);
```

`quickCheck` 函式（程式）應該使用 `indexOf()` 方法。

```js
assert.notStrictEqual(quickCheck.toString().search(/\.indexOf\(/), -1);
```

# --seed--

## --seed-contents--

```js
function quickCheck(arr, elem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

# --solutions--

```js
function quickCheck(arr, elem) {
  return arr.indexOf(elem) >= 0; 
}
```
