---
id: 587d7b7c367417b2b2512b19
title: 修改巢狀於物件中的物件
challengeType: 1
forumTopicId: 301164
dashedName: modify-an-object-nested-within-an-object
---

# --description--

現在讓我們來看看一個稍微複雜一點的物件。物件屬性可以巢狀到任意深度，而它們的值可以是 JavaScript 支援的任何型別的資料，包括陣列甚至其他物件。請參考以下內容：

```js
let nestedObject = {
  id: 28802695164,
  date: 'December 31, 2016',
  data: {
    totalUsers: 99,
    online: 80,
    onlineStatus: {
      active: 67,
      away: 13,
      busy: 8
    }
  }
};
```

`nestedObject` 有三個屬性：`id`（值為數字）、`date`（值為字串）和 `data`（值為具有其巢狀結構的物件）。雖然結構可能很快變得複雜，但我們仍然可以使用相同的符號來存取我們需要的資訊。若要將值 `10` 指定給巢狀的 `onlineStatus` 物件的 `busy` 屬性，我們使用點記法來取用該屬性：

```js
nestedObject.data.onlineStatus.busy = 10;
```

# --instructions--

這裡我們定義了一個物件 `userActivity`，其中包含另一個巢狀的物件。將 `online` 鍵的值設定為 `45`。

# --hints--

`userActivity` 應該具有 `id`、`date` 和 `data` 屬性。

```js
assert(
  'id' in userActivity && 'date' in userActivity && 'data' in userActivity
);
```

`userActivity` 應該有一個 `data` 鍵，該鍵設定為一個包含 `totalUsers` 和 `online` 鍵的物件。

```js
assert('totalUsers' in userActivity.data && 'online' in userActivity.data);
```

應將 `userActivity` 的 `data` 鍵中巢狀的 `online` 屬性設定為 `45`

```js
assert(userActivity.data.online === 45);
```

`online` 屬性應該使用點記法或括號記法設定。

```js
assert.strictEqual(__helpers.removeJSComments(code).search(/online: 45/), -1);
```

# --seed--

## --seed-contents--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

// Only change code below this line

// Only change code above this line

console.log(userActivity);
```

# --solutions--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

userActivity.data.online = 45;
```
