---
id: 587d7b7b367417b2b2512b17
title: 使用展開運算子合併陣列
challengeType: 1
forumTopicId: 301156
dashedName: combine-arrays-with-the-spread-operator
---

# --description--

<dfn>spread</dfn> 運算子另一個巨大的優點是能夠合併陣列，或將一個陣列的所有元素插入到另一個陣列的任意索引位置。使用較傳統的語法，我們可以串接陣列，但這僅允許我們在一個陣列的結尾和另一個陣列的開頭合併陣列。Spread 語法使以下操作變得非常簡單：

```js
let thisArray = ['sage', 'rosemary', 'parsley', 'thyme'];

let thatArray = ['basil', 'cilantro', ...thisArray, 'coriander'];
```

`thatArray` 的值會是 `['basil', 'cilantro', 'sage', 'rosemary', 'parsley', 'thyme', 'coriander']`。

使用展開語法，我們剛完成了一個如果使用傳統方法會更複雜且冗長的操作。

# --instructions--

我們已定義一個函式（程式）`spreadOut`，該函式會傳回變數 `sentence`。請使用 <dfn>spread</dfn> 運算子修改該函式，使其傳回陣列 `['learning', 'to', 'code', 'is', 'fun']`。

# --hints--

`spreadOut` 應該傳回 `["learning", "to", "code", "is", "fun"]`

```js
assert.deepEqual(spreadOut(), ['learning', 'to', 'code', 'is', 'fun']);
```

`spreadOut` 函式（程式）應該使用展開語法

```js
assert.notStrictEqual(spreadOut.toString().search(/[...]/), -1);
```

# --seed--

## --seed-contents--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence; // Change this line
  return sentence;
}

console.log(spreadOut());
```

# --solutions--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence = ['learning', ...fragment, 'is', 'fun'];
  return sentence;
}
```
