---
id: 587d7b7b367417b2b2512b13
title: 使用展開運算子複製陣列
challengeType: 1
forumTopicId: 301157
dashedName: copy-an-array-with-the-spread-operator
---

# --description--

雖然 `slice()` 讓我們可以選擇性地複製陣列中的元素，以及執行其他幾項有用的任務，ES6 新的 <dfn>展開運算子</dfn> 讓我們能夠輕鬆地以簡單且高度可讀的語法，依序複製陣列的*所有*元素。展開語法看起來就是這樣：`...`

在實務中，我們可以使用展開運算子來複製陣列，如下所示：

```js
let thisArray = [true, true, undefined, false, null];
let thatArray = [...thisArray];
```

`thatArray` 等於 `[true, true, undefined, false, null]`。`thisArray` 保持不變，且 `thatArray` 包含與 `thisArray` 相同的元素。

# --instructions--

我們已經定義了一個函式（程式）`copyMachine`，它接受 `arr`（一個陣列）和 `num`（一個數字）作為引數。該函式應該傳回一個由 `num` 個 `arr` 複製組成的新陣列。我們已經為你完成大部分工作，但它還沒完全正常運作。請使用展開語法修改該函式，使其能正確運作（提示：我們之前涵蓋過的另一個方法可能會派上用場！）。

# --hints--

`copyMachine([true, false, true], 2)` 應該傳回 `[[true, false, true], [true, false, true]]`

```js
assert.deepEqual(copyMachine([true, false, true], 2), [
  [true, false, true],
  [true, false, true]
]);
```

`copyMachine([1, 2, 3], 5)` 應該傳回 `[[1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3]]`

```js
assert.deepEqual(copyMachine([1, 2, 3], 5), [
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3]
]);
```

`copyMachine([true, true, null], 1)` 應該傳回 `[[true, true, null]]`

```js
assert.deepEqual(copyMachine([true, true, null], 1), [[true, true, null]]);
```

`copyMachine(["it works"], 3)` 應該傳回 `[["it works"], ["it works"], ["it works"]]`

```js
assert.deepEqual(copyMachine(['it works'], 3), [
  ['it works'],
  ['it works'],
  ['it works']
]);
```

`copyMachine` 函式（程式）應該使用陣列 `arr` 的 `spread operator`。

```js
assert(__helpers.removeJSComments(__helpers.removeJSComments(code)).match(/\.\.\.\s*arr/));
```

# --seed--

## --seed-contents--

```js
function copyMachine(arr, num) {
  let newArr = [];
  while (num >= 1) {
    // Only change code below this line

    // Only change code above this line
    num--;
  }
  return newArr;
}

console.log(copyMachine([true, false, true], 2));
```

# --solutions--

```js
function copyMachine(arr,num){
    let newArr=[];
    while(num >=1){
    newArr.push([...arr]);
    num--;
    }
    return newArr;
}
console.log(copyMachine([true, false, true], 2));
```
