---
id: 587d7b7d367417b2b2512b1d
title: 使用 for...in 述語迭代物件的鍵
challengeType: 1
forumTopicId: 301162
dashedName: iterate-through-the-keys-of-an-object-with-a-for
---

# --description--

有時你需要迭代物件中的所有鍵。你可以使用 <dfn>for...in</dfn> 迴圈來做到這一點。for...in 迴圈的語法如下：

```javascript
const refrigerator = {
  'milk': 1,
  'eggs': 12,
};

for (const food in refrigerator) {
  console.log(food, refrigerator[food]);
}
```

此程式碼記錄了 `milk 1` 和 `eggs 12`，每個鍵值對各自佔一行。

我們在迴圈開頭定義了變數 `food`，並且此變數在每次迭代時被設定為該物件的每個鍵，導致每個食物的名稱被列印到主控台。

**注意：** 物件不像陣列那樣維護已儲存鍵的排序；因此，在取用或存取該鍵時，鍵在物件中的位置或其出現的相對順序並不重要。

# --instructions--

我們定義了一個函式（程式）`countOnline`，它接受一個引數 `allUsers`。在這個函式（程式）內使用 <dfn>for...in</dfn> 述語來迴圈遍歷 `allUsers` 物件，並傳回 `online` 屬性設定為 `true` 的使用者數量。下面顯示了一個可以傳遞給 `countOnline` 的物件範例。每個使用者都會有一個 `online` 屬性，設定為 `true` 或 `false`。

```js
{
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
```

# --hints--

函式（程式）`countOnline` 應該使用 `for in` 述語來迭代傳入的物件的物件鍵。

```js
assert.match(
  __helpers.removeJSComments(code),
    /for\s*\(\s*(var|let|const)\s+[a-zA-Z_$]\w*\s+in\s+[a-zA-Z_$]\w*\s*\)/
);
```

當物件 `{ Alan: { online: false }, Jeff: { online: true }, Sarah: { online: false } }` 傳入時，函式 `countOnline` 應該傳回 `1`。

```js
const usersObj1 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj1), 1);
```

當物件 `{ Alan: { online: true }, Jeff: { online: false }, Sarah: { online: true } }` 傳入時，函式 `countOnline` 應該傳回 `2`。

```js
const usersObj2 = {
  Alan: {
    online: true
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: true
  }
}
assert.equal(countOnline(usersObj2), 2);
```

當物件 `{ Alan: { online: false }, Jeff: { online: false }, Sarah: { online: false } }` 傳入時，函式 `countOnline` 應該傳回 `0`。

```js
const usersObj3 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj3), 0);
```

# --seed--

## --seed-contents--

```js
const users = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}

function countOnline(allUsers) {
  // Only change code below this line

  // Only change code above this line
}

console.log(countOnline(users));
```

# --solutions--

```js
function countOnline(allUsers) {
  let numOnline = 0;
  for(const user in allUsers){
    if(allUsers[user].online) {
      numOnline++;
    }
  }
  return numOnline;
}
```
