---
id: 587d78b2367417b2b2512b0f
title: 使用 pop() 和 shift() 從陣列中移除條款
challengeType: 1
forumTopicId: 301165
dashedName: remove-items-from-an-array-with-pop-and-shift
---

# --description--

`push()` 和 `unshift()` 都有對應的幾乎功能相反的方法：`pop()` 和 `shift()`。如你現在可能已經猜到的，`pop()` 並非添加，而是從陣列的末端 *移除* 一個元素，而 `shift()` 則是從開頭移除一個元素。`pop()` 和 `shift()` 與它們的對應方法 `push()` 和 `unshift()` 之間的主要差異在於，這兩個方法都不接受參數，且每次只允許陣列被修改一個元素。

讓我們來看看：

```js
let greetings = ['whats up?', 'hello', 'see ya!'];

greetings.pop();
```

`greetings` 的值會是 `['whats up?', 'hello']`。

```js
greetings.shift();
```

`greetings` 的值會是 `['hello']`。

我們也可以用任一方法像這樣傳回被移除元素的值：

```js
let popped = greetings.pop();
```

`greetings` 的值會是 `[]`，而 `popped` 的值會是 `hello`。

# --instructions--

我們已定義一個函式（程式）`popShift`，它以一個陣列作為引數並傳回一個新的陣列。修改該函式，使用 `pop()` 和 `shift()`，以移除引數陣列的第一個和最後一個元素，並將移除的元素指定給對應的變數，使傳回的陣列包含它們的值。

# --hints--

`popShift(["challenge", "is", "not", "complete"])` 應該傳回 `["challenge", "complete"]`

```js
assert.deepEqual(popShift(['challenge', 'is', 'not', 'complete']), [
  'challenge',
  'complete'
]);
```

`popShift` 函式（程式）應該使用 `pop()` 方法。

```js
assert.notStrictEqual(popShift.toString().search(/\.pop\(/), -1);
```

`popShift` 函式應該使用 `shift()` 方法。

```js
assert.notStrictEqual(popShift.toString().search(/\.shift\(/), -1);
```

# --seed--

## --seed-contents--

```js
function popShift(arr) {
  let popped; // Change this line
  let shifted; // Change this line
  return [shifted, popped];
}

console.log(popShift(['challenge', 'is', 'not', 'complete']));
```

# --solutions--

```js
function popShift(arr) {
  let popped = arr.pop(); // Change this line
  let shifted = arr.shift(); // Change this line
  return [shifted, popped];
}
```
