---
id: 587d7b7b367417b2b2512b15
title: 使用 For 迴圈迭代整個陣列的所有項目
challengeType: 1
forumTopicId: 301161
dashedName: iterate-through-all-an-arrays-items-using-for-loops
---

# --description--

有時在處理陣列時，能夠迭代每個條款以尋找一個或多個我們可能需要的元素，或根據符合特定條件的資料條款來操作陣列，是非常方便的。JavaScript 提供了幾種內建的行為，每種行為都以稍微不同的方式迭代陣列以達成不同的結果（例如 `every()`、`forEach()`、`map()` 等），然而最靈活且提供我們最大控制權的技術是一個簡單的 `for` 迴圈。

請考慮以下事項：

```js
function greaterThanTen(arr) {
  let newArr = [];
  for (let i = 0; i < arr.length; i++) {
    if (arr[i] > 10) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}

greaterThanTen([2, 12, 8, 14, 80, 0, 1]);
```

使用 `for` 迴圈，這個函式（程式）迭代並存取陣列中的每個元素，並對其進行我們創建的簡單測試。透過這種方式，我們輕鬆且以程式化方式判斷哪些資料項目大於 `10`，並傳回包含這些項目的新陣列 `[12, 14, 80]`。

# --instructions--

我們已定義一個函式（程式）`filteredArray`，它接受 `arr`（一個巢狀陣列）和 `elem` 作為引數，並傳回一個新的陣列。`elem` 表述一個可能存在或不存在於 `arr` 中一個或多個巢狀陣列的元素。請修改該函式，使用 `for` 迴圈，傳回傳入陣列的篩選版本，使得任何包含 `elem` 的巢狀陣列都被移除。

# --hints--

`filteredArray([[10, 8, 3], [14, 6, 23], [3, 18, 6]], 18)` 應該傳回 `[[10, 8, 3], [14, 6, 23]]`

```js
assert.deepEqual(
  filteredArray(
    [
      [10, 8, 3],
      [14, 6, 23],
      [3, 18, 6]
    ],
    18
  ),
  [
    [10, 8, 3],
    [14, 6, 23]
  ]
);
```

`filteredArray([["trumpets", 2], ["flutes", 4], ["saxophones", 2]], 2)` 應該傳回 `[["flutes", 4]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['trumpets', 2],
      ['flutes', 4],
      ['saxophones', 2]
    ],
    2
  ),
  [['flutes', 4]]
);
```

`filteredArray([["amy", "beth", "sam"], ["dave", "sean", "peter"]], "peter")` 應該傳回 `[["amy", "beth", "sam"]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['amy', 'beth', 'sam'],
      ['dave', 'sean', 'peter']
    ],
    'peter'
  ),
  [['amy', 'beth', 'sam']]
);
```

`filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3)` 應該傳回 `[]`

```js
assert.deepEqual(
  filteredArray(
    [
      [3, 2, 3],
      [1, 6, 3],
      [3, 13, 26],
      [19, 3, 9]
    ],
    3
  ),
  []
);
```

`filteredArray` 函式（程式）應該使用 `for` 迴圈

```js
assert.notStrictEqual(filteredArray.toString().search(/for/), -1);
```

# --seed--

## --seed-contents--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  // Only change code below this line

  // Only change code above this line
  return newArr;
}

console.log(filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3));
```

# --solutions--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  for (let i = 0; i<arr.length; i++) {
    if (arr[i].indexOf(elem) < 0) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}
```
