---
id: 587d7b7c367417b2b2512b1b
title: 使用 delete 關鍵字來移除物件屬性
challengeType: 1
forumTopicId: 301168
dashedName: use-the-delete-keyword-to-remove-object-properties
---

# --description--

現在你知道什麼是物件以及它們的基本特性和優點。簡而言之，它們是鍵值儲存，提供一種靈活且直覺的方式來結構資料，***並且***，它們提供非常快速的查找時間。在接下來的挑戰中，我們將描述你可以對物件執行的幾個常見操作，讓你能夠熟悉在你的程式中應用這些有用的資料結構。

在先前的挑戰中，我們既為物件添加過，也修改過物件的鍵值對組。在這裡，我們將看到如何從物件中*移除*一個鍵值對組。

讓我們最後再回顧一次 `foods` 物件範例。如果我們想要移除 `apples` 鍵，可以使用 `delete` 關鍵字這樣移除：

```js
delete foods.apples;
```

# --instructions--

使用 delete 關鍵字從 `foods` 物件中刪除 `oranges`、`plums` 和 `strawberries` 鍵。

# --hints--

`foods` 物件應該只包含三個鍵：`apples`、`grapes` 和 `bananas`。

```js
assert(
  !foods.hasOwnProperty('oranges') &&
    !foods.hasOwnProperty('plums') &&
    !foods.hasOwnProperty('strawberries') &&
    Object.keys(foods).length === 3
);
```

應使用 `delete` 移除 `oranges`、`plums` 和 `strawberries` 鍵。

```js
assert(
  __helpers.removeJSComments(code).search(/oranges:/) !== -1 &&
    __helpers.removeJSComments(code).search(/plums:/) !== -1 &&
    __helpers.removeJSComments(code).search(/strawberries:/) !== -1
);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

delete foods.oranges;
delete foods.plums;
delete foods.strawberries;

console.log(foods);
```
