---
id: 587d78b3367417b2b2512b11
title: 使用 splice() 添加項目
challengeType: 1
forumTopicId: 301152
dashedName: add-items-using-splice
---

# --description--

還記得在上一個挑戰中我們提到過 `splice()` 最多可以接受三個參數嗎？你可以使用第三個參數，該參數由一個或多個元素組成，將它們添加到陣列中。這對於快速替換一個元素或一組元素非常有用。

```js
const numbers = [10, 11, 12, 12, 15];
const startIndex = 3;
const amountToDelete = 1;

numbers.splice(startIndex, amountToDelete, 13, 14);
console.log(numbers);
```

第二個出現的 `12` 被移除，並且我們在相同的索引位置添加 `13` 和 `14`。`numbers` 陣列現在會是 `[ 10, 11, 12, 13, 14, 15 ]`。

這裡，我們從一個數字陣列開始。然後，我們傳遞以下內容給 `splice()`：開始刪除元素的索引（3）、要刪除的元素數量（1），以及剩餘的引數（13、14）將從相同的索引開始插入。請注意，在 `amountToDelete` 之後可以有任意數量的元素（以逗號分隔），每個元素都會被插入。

# --instructions--

我們已定義一個函式（程式）`htmlColorNames`，它以一個 HTML 顏色的陣列作為引數。修改該函式，使用 `splice()` 移除陣列的前兩個元素，並在相應位置添加 `'DarkSalmon'` 和 `'BlanchedAlmond'`。

# --hints--

`htmlColorNames` 應該傳回 `["DarkSalmon", "BlanchedAlmond", "LavenderBlush", "PaleTurquoise", "FireBrick"]`

```js
assert.deepEqual(
  htmlColorNames([
    'DarkGoldenRod',
    'WhiteSmoke',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]),
  [
    'DarkSalmon',
    'BlanchedAlmond',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]
);
```

`htmlColorNames` 函式應該使用 `splice()` 方法。

```js
assert(/.splice/.test(__helpers.removeJSComments(code)));
```

你不應該使用 `shift()` 或 `unshift()`。

```js
assert(!/shift|unshift/.test(__helpers.removeJSComments(code)));
```

你不應該使用陣列括號表示法。

```js
assert(!/\[\d\]\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function htmlColorNames(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(htmlColorNames(['DarkGoldenRod', 'WhiteSmoke', 'LavenderBlush', 'PaleTurquoise', 'FireBrick']));
```

# --solutions--

```js
function htmlColorNames(arr) {
  arr.splice(0,2,'DarkSalmon', 'BlanchedAlmond');
  return arr;
}
```
