---
id: 587d7b7e367417b2b2512b20
title: 使用陣列來儲存資料群集
challengeType: 1
forumTopicId: 301167
dashedName: use-an-array-to-store-a-collection-of-data
---

# --description--

以下是陣列資料結構最簡單實作的範例。這稱為 <dfn>一維陣列</dfn>，表示它只有一個層級，或是沒有其他陣列巢狀其中。請注意它包含 <dfn>布林值</dfn>、<dfn>字串</dfn> 和 <dfn>數字</dfn>，以及其他有效的 JavaScript 資料型別：

```js
let simpleArray = ['one', 2, 'three', true, false, undefined, null];
console.log(simpleArray.length);
```

`console.log` 呼叫顯示 `7`。

所有陣列都有一個 length 屬性，如上所示，可以非常輕鬆地使用語法 `Array.length` 來存取。下面可以看到一個較為複雜的陣列實作。這被稱為 <dfn>多維陣列</dfn>，或是包含其他陣列的陣列。請注意，這個陣列也包含 JavaScript <dfn>物件</dfn>，我們將在下一節中仔細檢視，但目前你只需要知道陣列也能儲存複雜的物件。

```js
let complexArray = [
  [
    {
      one: 1,
      two: 2
    },
    {
      three: 3,
      four: 4
    }
  ],
  [
    {
      a: "a",
      b: "b"
    },
    {
      c: "c",
      d: "d"
    }
  ]
];
```

# --instructions--

我們已經定義了一個名為 `yourArray` 的變數。請完成述語，將長度至少為 5 個元素的陣列指定給 `yourArray` 變數。你的陣列應該至少包含一個 <dfn>string</dfn>、一個 <dfn>number</dfn> 和一個 <dfn>boolean</dfn>。

# --hints--

`yourArray` 應該是一個陣列。

```js
assert.strictEqual(Array.isArray(yourArray), true);
```

`yourArray` 應該至少有 5 個元素長。

```js
assert.isAtLeast(yourArray.length, 5);
```

`yourArray` 應該至少包含一個 `boolean`。

```js
assert(yourArray.filter((el) => typeof el === 'boolean').length >= 1);
```

`yourArray` 應該至少包含一個 `number`。

```js
assert(yourArray.filter((el) => typeof el === 'number').length >= 1);
```

`yourArray` 應該至少包含一個 `string`。

```js
assert(yourArray.filter((el) => typeof el === 'string').length >= 1);
```

# --seed--

## --seed-contents--

```js
let yourArray; // Change this line
```

# --solutions--

```js
let yourArray = ['a string', 100, true, ['one', 2], 'another string'];
```
