---
id: 587d7b7d367417b2b2512b1e
title: 使用 Object.keys() 產生所有物件鍵的陣列
challengeType: 1
forumTopicId: 301160
dashedName: generate-an-array-of-all-object-keys-with-object-keys
---

# --description--

我們也可以使用 `Object.keys()` 方法產生一個包含物件中所有鍵的陣列。這個方法以物件作為引數，並傳回一個字串陣列，表述物件中的每個屬性。陣列中的項目同樣不會有特定的順序。

# --instructions--

完成撰寫 `getArrayOfUsers` 函式（程式），使其傳回一個包含它作為引數接收的物件中所有屬性的陣列。

# --hints--

`users` 物件應該只包含鍵 `Alan`、`Jeff`、`Sarah` 和 `Ryan`。

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

`getArrayOfUsers` 函式（程式）應該傳回一個包含 `users` 物件中所有鍵的陣列。

```js
assert(
  (function () {
    users.Sam = {};
    users.Lewis = {};
    let R = getArrayOfUsers(users);
    return (
      R.indexOf('Alan') !== -1 &&
      R.indexOf('Jeff') !== -1 &&
      R.indexOf('Sarah') !== -1 &&
      R.indexOf('Ryan') !== -1 &&
      R.indexOf('Sam') !== -1 &&
      R.indexOf('Lewis') !== -1
    );
  })() === true
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  // Only change code below this line

  // Only change code above this line
}

console.log(getArrayOfUsers(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  return Object.keys(obj);
}

console.log(getArrayOfUsers(users));
```
