---
id: 587d7b7d367417b2b2512b1f
title: 修改儲存在物件中的陣列
challengeType: 1
forumTopicId: 301163
dashedName: modify-an-array-stored-in-an-object
---

# --description--

現在你已經看過所有 JavaScript 物件的基本操作。你可以為你的物件添加、修改和移除鍵值對組，檢查鍵是否存在，並對物件中的所有鍵進行迭代。隨著你繼續學習 JavaScript，你將會看到物件更靈活多變的應用。此外，課程中位於程式設計面試準備部分的資料結構課程也涵蓋了 ES6 的 <dfn>Map</dfn> 和 <dfn>Set</dfn> 物件，這兩者類似於常規的物件，但提供了一些額外的特性。現在你已經學會了陣列和物件的基礎，已經完全準備好開始使用 JavaScript 解決更複雜的問題！

# --instructions--

看看我們在程式碼編輯器中提供的物件。`user` 物件包含三個鍵。`data` 鍵包含五個鍵，其中一個包含 `friends` 陣列。由此，你可以看到物件作為資料結構的彈性。我們已經開始撰寫函式 `addFriend`。完成撰寫它，使其接受一個 `user` 物件，並將 `friend` 引數的名稱添加到儲存在 `user.data.friends` 的陣列中，然後傳回該陣列。

# --hints--

`user` 物件應該有 `name`、`age` 和 `data` 鍵。

```js
assert('name' in user && 'age' in user && 'data' in user);
```

`addFriend` 函式（程式）應該接受一個 `user` 物件和一個 `friend` 字串作為引數，並將該朋友添加到 `user` 物件中 `friends` 的陣列。

```js
assert(
  (function () {
    let L1 = user.data.friends.length;
    addFriend(user, 'Sean');
    let L2 = user.data.friends.length;
    return L2 === L1 + 1;
  })()
);
```

`addFriend(user, "Pete")` 應該傳回 `["Sam", "Kira", "Tomo", "Pete"]`。

```js
assert.deepEqual(
  (function () {
    delete user.data.friends;
    user.data.friends = ['Sam', 'Kira', 'Tomo'];
    return addFriend(user, 'Pete');
  })(),
  ['Sam', 'Kira', 'Tomo', 'Pete']
);
```

# --seed--

## --seed-contents--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  // Only change code below this line

  // Only change code above this line
}

console.log(addFriend(user, 'Pete'));
```

# --solutions--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  userObj.data.friends.push(friend);
  return userObj.data.friends;
}
```
