---
id: 587d7b7d367417b2b2512b1c
title: 檢查物件是否具有屬性
challengeType: 1
forumTopicId: 301155
dashedName: check-if-an-object-has-a-property
---

# --description--

現在我們可以為物件添加、修改和移除鍵。但如果我們只是想知道物件是否有特定的屬性呢？JavaScript 提供了兩種不同的方法來做到這一點。一種使用 `hasOwnProperty()` 方法，另一種使用 `in` 關鍵字。如果我們有一個名為 `users` 的物件，且它有一個 `Alan` 屬性，我們可以用以下任一種方式來檢查它是否存在：

```js
users.hasOwnProperty('Alan');
'Alan' in users;
```

這兩者都會傳回 `true`。

# --instructions--

完成撰寫函式，使其在傳入的物件包含所有四個名稱 `Alan`、`Jeff`、`Sarah` 和 `Ryan` 時傳回 `true`，否則傳回 `false`。

# --hints--

不應直接存取 `users` 物件

```js 

assert(__helpers.removeJSComments(code).match(/users/gm).length <= 2)

```

`users` 物件應該只包含鍵 `Alan`、`Jeff`、`Sarah` 和 `Ryan`。

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

如果傳入的物件具有 `Alan`、`Jeff`、`Sarah` 和 `Ryan` 這些屬性，函式 `isEveryoneHere` 應該傳回 `true`。 

```js
assert(isEveryoneHere(users) === true);
```

如果傳入的物件中沒有 `Alan` 這個屬性，函式 `isEveryoneHere` 應該傳回 `false`。

```js
assert(
  (function () {
    delete users.Alan;
    return isEveryoneHere(users);
  })() === false
);
```

如果傳入的物件中沒有 `Jeff` 這個屬性，函式（程式）`isEveryoneHere` 應該傳回 `false`。

```js
assert(
  (function () {
    delete users.Jeff;
    return isEveryoneHere(users);
  })() === false
);
```

如果傳入的物件沒有 `Sarah` 這個屬性，函式（程式）`isEveryoneHere` 應該傳回 `false`。

```js
assert(
  (function () {
    delete users.Sarah;
    return isEveryoneHere(users);
  })() === false
);
```

如果傳入的物件中沒有 `Ryan` 這個屬性，函式（程式）`isEveryoneHere` 應該傳回 `false`。

```js
assert(
  (function () {
    delete users.Ryan;
    return isEveryoneHere(users);
  })() === false
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  // Only change code below this line
  
  // Only change code above this line
}

console.log(isEveryoneHere(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  return [
    'Alan',
    'Jeff',
    'Sarah',
    'Ryan'
  ].every(user => userObj.hasOwnProperty(user));
}

console.log(isEveryoneHere(users));
```
