---
id: 587d78af367417b2b2512b03
title: Crea un modulo di sondaggio
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un titolo della pagina in un elemento `h1` con un `id` di `title`.
1. Dovresti avere una breve spiegazione in un elemento `p` con un `id` di `description`.
1. Dovresti avere un elemento `form` con un `id` di `survey-form`.
1. All'interno dell'elemento modulo dovresti avere un campo `input` obbligatorio per inserire il tuo nome che ha un `id` di `name` e un `type` di `text`.
1. All'interno dell'elemento modulo dovresti avere un campo `input` obbligatorio per inserire la tua email che ha un `id` di `email`.
1. Se inserisci un'email non formattata correttamente, dovresti vedere un errore di validazione HTML5.
1. All'interno dell'elemento modulo dovresti avere un campo `input` per inserire un numero che ha un `id` di `number`.
1. L'input numerico non dovrebbe accettare valori non numerici, sia impedendoti di digitarli sia mostrando un errore di validazione HTML5 (a seconda del browser).
1. Se inserisci numeri fuori dall'intervallo definito dagli attributi `min` e `max` del campo numerico, dovresti vedere un errore di validazione HTML5.
1. Per i campi di input nome, email e numero, dovresti avere elementi `label` corrispondenti nel modulo, che descrivono lo scopo di ogni campo con gli id: `name-label`, `email-label` e `number-label`.
1. Per i campi di input nome, email e numero, dovresti avere un testo segnaposto che fornisce una descrizione o istruzioni per ogni campo.
1. All'interno dell'elemento modulo, dovresti avere un menu a discesa `select` con un `id` di `dropdown` e almeno due opzioni tra cui scegliere.
1. All'interno dell'elemento modulo, puoi selezionare un'opzione da un gruppo di almeno due pulsanti di opzione raggruppati usando l'attributo `name`.
1. All'interno dell'elemento modulo, puoi selezionare più campi da una serie di caselle di spunta, ognuna delle quali deve avere un attributo `value`.
1. All'interno dell'elemento modulo, dovresti avere un `textarea` per commenti aggiuntivi.
1. All'interno dell'elemento modulo, dovresti avere un pulsante con `id` di `submit` per inviare tutti gli input.

# --hints--

Dovresti avere un elemento `h1` con un `id` di `title`.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

Il tuo `#title` non dovrebbe essere vuoto.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Dovresti avere un elemento `p` con un `id` di `description`.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

Il tuo `#description` non dovrebbe essere vuoto.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Dovresti avere un elemento `form` con un `id` di `survey-form`.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

Dovresti avere un elemento `input` con un `id` di `name`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Il tuo `#name` dovrebbe avere un `type` di `text`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

Il tuo `#name` dovrebbe richiedere un input.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Il tuo `#name` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

Dovresti avere un elemento `input` con un `id` di `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Il tuo `#email` dovrebbe avere un `type` di `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

Il tuo `#email` dovrebbe richiedere un input.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Il tuo `#email` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

Dovresti avere un elemento `input` con un `id` di `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Il tuo `#number` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

Il tuo `#number` dovrebbe avere un `type` di `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

Il tuo `#number` dovrebbe avere un attributo `min` con un valore numerico.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

Il tuo `#number` dovrebbe avere un attributo `max` con un valore numerico.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

Dovresti avere un elemento `label` con un `id` di `name-label`.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Dovresti avere un elemento `label` con un `id` di `email-label`.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Dovresti avere un elemento `label` con un `id` di `number-label`.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Il tuo `#name-label` dovrebbe contenere un testo che descrive l'input.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Il tuo `#email-label` dovrebbe contenere un testo che descrive l'input.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Il tuo `#number-label` dovrebbe contenere un testo che descrive l'input.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Il tuo `#name-label` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

Il tuo `#email-label` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

Il tuo `#number-label` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

Il tuo `#name` dovrebbe avere un attributo `placeholder` e un valore.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Il tuo `#email` dovrebbe avere un attributo `placeholder` e un valore.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Il tuo `#number` dovrebbe avere un attributo `placeholder` e un valore.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Dovresti avere un campo `select` con un `id` di `dropdown`.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

Il tuo `#dropdown` dovrebbe avere almeno due elementi `option` selezionabili (non disabilitati).

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

Il tuo `#dropdown` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

Dovresti avere almeno due elementi `input` con un `type` di `radio` (pulsanti di opzione).

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Dovresti avere almeno due pulsanti di opzione che sono discendenti di `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Tutti i tuoi pulsanti di opzione dovrebbero avere un attributo `value` e un valore.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Tutti i tuoi pulsanti di opzione dovrebbero avere un attributo `name` e un valore.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Ogni gruppo di pulsanti di opzione dovrebbe avere almeno 2 pulsanti.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

Dovresti avere almeno due elementi `input` con un `type` di `checkbox` (caselle di spunta) che sono discendenti di `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

Tutte le tue caselle di spunta dentro `#survey-form` dovrebbero avere un attributo `value` e un valore.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Dovresti avere almeno un elemento `textarea` che è un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

Dovresti avere un elemento `input` o `button` con un `id` di `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

Il tuo `#submit` dovrebbe avere un `type` di `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

Il tuo `#submit` dovrebbe essere un discendente di `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
