---
id: 6626a060c4006f793e10cb33
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Agora o professor precisa da sua ajuda para converter a pontuação do aluno em uma nota letra.

Complete a função `getGrade` que recebe um número `score` como parâmetro.Sua função deve retornar uma string que representa uma nota em letra com base na pontuação.

Aqui estão as pontuações e suas notas correspondentes:

| Score Range | Grade   |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**Dicas**

- Lembre-se de que você aprendeu sobre declarações condicionais (`if`, `else if` e `else`).
- Lembre-se de que você aprendeu sobre operadores de comparação (`>`, `<`, `>=`, `<=`, `===`).
- Lembre-se de que você aprendeu sobre o operador lógico AND (`&&`).

# --hints--

Sua função `getGrade` deve retornar `"A++"` se a pontuação for `100`.

```js
assert.strictEqual(getGrade(100), "A++");
```

Sua função `getGrade` deve retornar `"A"` se a pontuação for `94`.

```js
assert.strictEqual(getGrade(94), "A");
```

A função `getGrade` deve retornar `"B"` se a pontuação estiver entre `80` e `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Sua função `getGrade` deve retornar `"C"` se a pontuação for `78`.

```js
assert.strictEqual(getGrade(75), "C");
```

A função `getGrade` deve retornar `"D"` se a pontuação estiver entre `60` e `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Sua função `getGrade` deve retornar `"F"` se a pontuação for `57`.

```js
assert.strictEqual(getGrade(57), "F");
```

A função `getGrade` deve retornar `"F"` se a pontuação estiver entre `0` e `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
