---
id: 6626b8dcf5057f896f948440
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora que o professor tem todas as informações de que precisa, ele quer poder enviar uma mensagem para o aluno com os resultados. 

Complete a função `studentMsg` com `totalScores` e `studentScore` como parâmetros. A função deve retornar uma string que representa uma mensagem para o estudante.

Se o aluno passou no curso, a string deve seguir este formato:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

Se o aluno reprovou no curso, a string deve seguir este formato:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

Substitua `average-goes-here` pela média das pontuações totais.  
Substitua `grade-goes-here` pela nota do aluno.

**Dicas**

- Use a função `getAverage` para obter a média da turma.
- Use a função `getGrade` para obter a nota do aluno.
- Use concatenação de strings (`+`) para construir a mensagem.
- Tenha cuidado com a pontuação e os espaços na mensagem.

# --hints--

"`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` deve retornar a seguinte mensagem: `"Média da turma: 71,7."` Sua nota: F. Você reprovou no curso.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

"`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` deve retornar a seguinte mensagem: `"Média da turma: 50,8."` Your grade: A++. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

A função `studentMsg` deve retornar a mensagem correta com base na pontuação do estudante e na média da turma.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
