---
id: 662693f82c91a66be46c881b
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Um professor terminou de corrigir as provas dos alunos e precisa da sua ajuda para calcular a média da turma.

Complete a função `getAverage` que recebe um array de notas de teste e retorna a nota média.

A média é calculada somando todas as pontuações e dividindo pelo número total de pontuações.

```md
average = sum of all scores / total number of scores
```

Algumas chamadas de função foram fornecidas para que você possa testar seu código.

**Dicas**

- Você pode usar um loop para iterar sobre o array `scores` e somar todas as pontuações.
- Você pode usar a propriedade `length` para obter o número total de pontuações.

# --hints--

A função `getAverage` deve retornar um número.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` deve retornar `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` deve retornar `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` deve retornar `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

A função `getAverage` deve retornar a pontuação média das pontuações de teste.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
