---
id: 66c8ba975ee7230e29f6c4af
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Se il reddito annuo del richiedente è maggiore o uguale a `minIncomeForCondo` E se il suo punteggio di credito è maggiore o uguale a `minCreditScoreForCondo`, allora ha diritto a un prestito per un condominio e un'auto.

Verifica se questo è vero nella funzione `getLoanMessage`. Se lo è, restituisci la stringa `"You qualify for a condo and car loan."`

# --hints--

La tua funzione `getLoanMessage` dovrebbe restituire una stringa se il richiedente ha diritto a un prestito per un condominio e un'auto.

```js
assert.isString(getLoanMessage(45000, 680));
```

La tua funzione `getLoanMessage` dovrebbe restituire la stringa `"You qualify for a condo and car loan."`.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

La tua funzione `getLoanMessage` dovrebbe restituire `undefined` se il reddito annuo e il punteggio di credito del richiedente non soddisfano i requisiti per un prestito per un condominio.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
