---
id: 66c8ba975ee7230e29f6c4ad
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Per verificare per quale prestito un utente è qualificato in base a `annualIncome` e `creditScore`, devi usare un'istruzione `if/else if` o un ternario direttamente nella funzione `getLoanMessage`. Poi restituirai il messaggio appropriato nel blocco di ogni condizione.

Partendo dal prestito duplex, controlla se `annualIncome` è maggiore o uguale a `minIncomeForDuplex` E se `creditScore` è maggiore o uguale a `minCreditScoreForDuplex`.

Se questa condizione è vera, allora il richiedente è idoneo per un prestito duplex e per gli altri prestiti. Quindi, all'interno del controllo, restituisci la stringa `"You qualify for a duplex, condo, and car loan."`

# --hints--

La tua funzione `getLoanMessage` dovrebbe restituire una stringa.

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

La tua funzione `getLoanMessage` dovrebbe restituire la stringa `"You qualify for a duplex, condo, and car loan."`.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

La tua funzione `getLoanMessage` dovrebbe restituire `undefined` se il reddito annuo e il punteggio di credito del richiedente non soddisfano i requisiti per un prestito duplex.

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
