---
id: 66c8ba975ee7230e29f6c4b1
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Se il reddito annuo e il punteggio di credito del richiedente sono inferiori a `minIncomeForCar` e `minCreditScoreForCar`, allora non ha diritto a nessun prestito. Quindi, restituisci la stringa `"You don't qualify for any loans."`

# --hints--

La tua funzione `getLoanMessage` dovrebbe restituire una stringa se il richiedente non ha diritto a nessun prestito.

```js
assert.isString(getLoanMessage(25000, 550));
```

La tua funzione `getLoanMessage` dovrebbe restituire la stringa `"You don't qualify for any loans."` se il richiedente non soddisfa i requisiti per nessun prestito.

```js
assert.strictEqual(getLoanMessage(15000, 600), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(25000, 550), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(20000, 500), "You don't qualify for any loans.");
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
