---
id: 66c8ba975ee7230e29f6c4b0
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora dovresti verificare se il richiedente è qualificato solo per un prestito auto. Se lo è, restituisci la stringa `"You qualify for a car loan."`.

# --hints--

La tua funzione `getLoanMessage` dovrebbe restituire una stringa se il richiedente si qualifica per un prestito auto.

```js
assert.isString(getLoanMessage(30000, 650));
```

La tua funzione `getLoanMessage` dovrebbe restituire la stringa `"You qualify for a car loan."`.

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

La tua funzione `getLoanMessage` dovrebbe restituire `undefined` se il reddito annuo e il punteggio di credito del richiedente non soddisfano i requisiti per un prestito auto.

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 
  
--fcc-editable-region--
}
```
