---
id: 66c8ba975ee7230e29f6c4b2
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora è il momento di testare la tua funzione `getLoanMessage`.

Usa la tabella qui sotto per creare le variabili `4` e i loro valori:

| Nome variabile | Valore   |
| ----------- | ------- |
| `duplexLoanMsg`       | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)`   |
| `carLoanMsg`   | `getLoanMessage(45000, 660)`   |
| `noLoanMsg`   | `getLoanMessage(25000, 550)`   |

Dopodiché, stampa ogni variabile nella console per vedere i messaggi.

Con questo, il tuo progetto per il controllo della qualificazione del prestito è completo!

# --hints--

Dovresti avere una variabile chiamata `duplexLoanMsg`.

```js
assert.isNotNull(duplexLoanMsg);
```

La tua variabile `duplexLoanMsg` dovrebbe essere impostata al risultato di `getLoanMessage(85000, 850)`.

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

Dovresti stampare la variabile `duplexLoanMsg` nella console.

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

Dovresti avere una variabile chiamata `condoLoanMsg`.

```js
assert.isNotNull(condoLoanMsg);
```

La tua variabile `condoLoanMsg` dovrebbe essere impostata al risultato di `getLoanMessage(65000, 690)`.

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

Dovresti stampare la variabile `condoLoanMsg` nella console.

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

Dovresti avere una variabile chiamata `carLoanMsg`.

```js
assert.isNotNull(carLoanMsg);
```

La tua variabile `carLoanMsg` dovrebbe essere impostata al risultato di `getLoanMessage(45000, 660)`.

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

Dovresti stampare la variabile `carLoanMsg` nella console.

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

Dovresti avere una variabile chiamata `noLoanMsg`.

```js
assert.isNotNull(noLoanMsg);
```

La tua variabile `noLoanMsg` dovrebbe essere impostata al risultato di `getLoanMessage(25000, 550)`.

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

Dovresti stampare la variabile `noLoanMsg` nella console.

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
