---
id: 66c8ba41a77db20a93f9d7a1
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, rivedrai le condizioni in JavaScript costruendo un'app per verificare l'idoneità a un prestito.

L'app controllerà se l'utente è idoneo a un prestito per duplex, auto e condominio in base al suo reddito annuo e al punteggio di credito. Un punteggio di credito è un numero che rappresenta quanto sei bravo a gestire il denaro preso in prestito.

Per iniziare, crea le seguenti variabili e valori.

| Nome variabile | Valore   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

Dovresti creare una variabile `minIncomeForDuplex` impostata su `60000`.

```js
assert.equal(minIncomeForDuplex, 60000);
```

Dovresti creare un `minCreditScoreForDuplex` impostato su `700`.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

Dovresti creare una variabile `minIncomeForCar` impostata su `30000`.

```js
assert.equal(minIncomeForCar, 30000);
```

Dovresti creare una variabile `minCreditScoreForCar` impostata su `650`.

```js
assert.equal(minCreditScoreForCar, 650);
```

Dovresti creare una variabile `minIncomeForCondo` impostata su `45000`.

```js
assert.equal(minIncomeForCondo, 45000);
```

Dovresti creare una variabile `minCreditScoreForCondo` impostata su `680`.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
