---
id: 66c8ba975ee7230e29f6c4ac
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Quando l'utente è idoneo per un prestito, vorrai mostrargli un messaggio nella console.

Per farlo, creerai una funzione all'interno della quale farai dei controlli che restituiranno a quale prestito il richiedente ha diritto.

Crea una funzione `getLoanMessage` vuota con due parametri `annualIncome` e `creditScore`.

# --hints--

Dovresti creare una funzione `getLoanMessage`

```js
assert.isFunction(getLoanMessage)
```

La tua funzione `getLoanMessage` dovrebbe avere due parametri `annualIncome` e `creditScore`.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` dovrebbe essere una funzione vuota.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
