---
id: 69757cc0faae0152c1418aad
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Nesta oficina, você praticará trabalhar com números e operações matemáticas para construir um divisor de conta. Esta ferramenta calculará quanto cada pessoa deve após adicionar os custos da refeição e a gorjeta.

Para começar, você precisa de uma forma de acompanhar o valor total à medida que os custos são adicionados. Em Python, você pode usar uma variável para armazenar um inteiro (um número inteiro) que muda ao longo do tempo.

Por exemplo, você pode escrever:

```py
my_number = 2
```

Crie uma variável chamada `running_total` e atribua a ela o valor `0`.

# --hints--

Você deve ter uma variável chamada `running_total`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

Você deve atribuir o inteiro `0` à sua variável `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
