---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Cada curso tem um custo. Você precisa armazenar esses valores em variáveis para usá-los depois. Como esses valores incluem centavos, você usará o tipo float, que é usado para representar números decimais. Aqui está um exemplo de uma variável com um valor float:

```py
change = 2.35
```

Crie quatro variáveis: `appetizers` com valor `37.89`, `main_courses` com valor `57.34`, `desserts` com valor `39.39` e `drinks` com valor `64.21`. 

# --hints--

Você deve definir uma variável chamada `appetizers` e atribuir a ela o valor `37.89`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

Você deve definir uma variável chamada `main_courses` e atribuir a ela o valor `57.34`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

Você deve definir uma variável chamada `desserts` e atribuir a ela o valor `39.39`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

Você deve definir uma variável chamada `drinks` e atribuir a ela o valor `64.21`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
