---
id: 6976cc0b0c135686a4620b94
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Em seguida, você precisa considerar o número de pessoas dividindo a conta. Armazene esse valor em uma variável, como fez no passo anterior.

Crie uma variável chamada `num_of_friends` e atribua a ela o valor de `4`. Isso será usado mais tarde na oficina para calcular a divisão final.

# --hints--

Você deve ter uma variável chamada `num_of_friends`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

Você deve atribuir o inteiro `4` à sua variável `num_of_friends`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
