---
id: 69779dcb44016eccc05c15a4
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

O serviço foi excelente, então o grupo decide deixar uma gorjeta de 25%. Para calcular uma porcentagem em Python, você pode multiplicar o total pelo equivalente decimal da porcentagem.

Por exemplo, para buscar 10% de um valor, você multiplicaria ele por `0.10` usando o operador `*`:

```py
tax = total * 0.10
```

Crie uma variável chamada `tip` e atribua a ela o resultado de multiplicar `running_total` por `0.25`.

Finalmente, use `print()` para exibir a string `Tip amount:` seguida de um espaço e do valor da sua variável `tip`.

# --hints--

Você deve ter uma variável chamada `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

Sua variável `tip` deve ser o resultado de `running_total * 0.25`.

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

Você deve imprimir a string `Tip amount:` seguida de um espaço e da variável `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total = appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
