---
id: 6977a8bbcbdd557886bd0718
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Com a gorjeta agora incluída, você tem o valor final para todo o grupo. Você precisa determinar quanto cada pessoa deve dividindo a conta total pelo número de amigos. 

Em Python, você usa a barra `/` para realizar divisão. Por exemplo:

```py
half = 10 / 2
```

Crie uma variável chamada `final_bill` e atribua a ela o resultado da divisão de `running_total` por `num_of_friends`. 

Finalmente, use a `print()` função para exibir a string `Bill per person:` seguida de um espaço e do valor de `final_bill`.

# --hints--

Você deve definir uma variável chamada `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

Você deve usar o operador `/` para dividir `running_total` por `num_of_friends` e atribuir o resultado à sua variável `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

Você deve imprimir a string `Bill per person:` seguida por um espaço e a variável `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
