---
id: 6977a4306bb3e856690a4890
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Agora que você calculou a gorjeta, você precisa adicionar ela ao seu `running_total` para buscar o valor final da conta.

Em Python, você pode usar o operador de atribuição aumentada `+=` para adicionar um valor a uma variável e atualizar essa variável ao mesmo tempo. Por exemplo, `total += 5` é uma forma abreviada de escrever `total = total + 5`.

Use o operador `+=` para adicionar o valor de `tip` ao seu `running_total`. Finalmente, use `print()` para imprimir a string `Total with tip:` seguida de um espaço e do valor de `running_total`.

# --hints--

Você deve usar o operador `+=` para adicionar o valor de `tip` ao seu `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

Você deve imprimir a string `Total with tip:` seguida por um espaço e a variável `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total with tip:', running_total)") or _Node(_code).has_call("print(f'Total with tip: {running_total}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
