---
id: 6982684f3a25f379e195a5fc
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora que você armazenou os custos individuais, pode calcular o total. Em Python, você usa o operador de adição `+` para somar valores. 

Atualize a variável `running_total` adicionando `appetizers`, `main_courses`, `desserts` e `drinks` juntos. 

Finalmente, use `print()` para exibir a string `Total bill so far:` seguida de um espaço e do valor de `running_total`.

**Nota:** Você pode notar que a saída tem mais dígitos decimais do que o esperado. Como você aprendeu em uma lição anterior, isso acontece porque os números são armazenados em binário e muitos valores decimais não podem ser representados exatamente nesse formato, o que leva a erros de arredondamento.

# --hints--

Você deve atualizar `running_total` usando o operador `+` para somar o custo dos quatro cursos. 

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total = {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

Você deve imprimir a string `Total bill so far:` seguida por um espaço e o valor de `running_total`.

```js
({
    test: () => runPython(`
    sol1 = "print('Total bill so far:', running_total)"
    sol2 = "print(f'Total bill so far: {running_total}')"
    assert _Node(_code).has_call(sol1) or _Node(_code).has_call(sol2)`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
