---
id: 6982684f3a25f379e195a5fc
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora que você armazenou os custos individuais, pode calcular o total.

Lembre-se de que o operador `+=` adiciona um valor a uma variável existente e a atualiza ao mesmo tempo. Por exemplo:

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

Use o operador `+=` uma vez para adicionar `appetizers`, `main_courses`, `desserts` e `drinks` a `running_total`.

Finalmente, use `print()` para exibir a string `Total bill so far:` seguida de um espaço e do valor de `running_total`.

**Nota:** Você pode notar que a saída tem mais dígitos decimais do que o esperado. Como você aprendeu em uma lição anterior, isso acontece porque os números são armazenados em binário e muitos valores decimais não podem ser representados exatamente nesse formato, o que leva a erros de arredondamento.

# --hints--

Você deve atualizar `running_total` usando o operador `+=` uma vez para adicionar os custos dos quatro cursos.

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total += {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

Você deve imprimir a string `Total bill so far:` seguida por um espaço e o valor de `running_total`.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Total bill so far: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - (37.89 + 57.34 + 39.39 + 64.21)) < 1e-6
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
