---
id: 587d7fb5367417b2b2512c04
title: 從你的相依性中移除軟體包
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

你現在已經測試了幾種使用 `package.json` 的 dependencies 區段來管理你的專案依賴的方法。你也已經透過將外部軟體包添加到檔案中來包含它們，甚至使用像是波浪號或脫字符這類特殊字元告訴 `npm` 你想要的版本型別。

但是如果你想移除一個你不再需要的外部軟體包怎麼辦？你可能已經猜到了，只要從你的相依性中移除該軟體包對應的鍵值對組即可。

這同樣的方法也適用於移除你 `package.json` 中的其他欄。

# --instructions--

從你的相依性中移除 `@freecodecamp/example` 軟體包。

**注意：** 移除它之後，請確保你有正確數量的逗號。

# --hints--

`"dependencies"` 不應包含 `"@freecodecamp/example"`。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
