---
id: 587d7fb3367417b2b2512bfb
title: '如何使用 package.json，任何 Node.js 專案或 npm 軟體包的核心'
challengeType: 2
forumTopicId: 301528
dashedName: how-to-use-package-json-the-core-of-any-node-js-project-or-npm-package
---

# --description--

處理這些挑戰將涉及你使用以下其中一種方法來撰寫你的程式碼：

- 複製<a href="https://github.com/freeCodeCamp/boilerplate-npm/" target="_blank" rel="noopener noreferrer nofollow">這個 GitHub 倉庫</a>，並在區域端完成這些挑戰。
- 使用你選擇的網站建置工具來完成專案。務必將我們 GitHub 倉庫中的所有檔案納入其中。

`package.json` 檔案是任何 Node.js 專案或 npm 軟體包的核心。它儲存有關你的專案的資訊。它由一個單一的 JSON 物件組成，資訊以鍵值對組的形式儲存。只有兩個必填欄位；`name` 和 `version`，但提供額外的資訊是良好的做法。

你可以從終端機使用 `npm init` 指令創建 `package.json` 檔案。這將執行一個引導式設定。使用帶有 `-y` 旗標的 `npm init` 將在不詢問任何問題的情況下產生該檔案，`npm init -y`。

如果你查看你的專案的檔案樹，你會發現在樹的最上層有 `package.json` 檔案。這是你在接下來幾個挑戰中將要改進的檔案。

這個檔案中最常見的資訊之一是 `author` 欄位。它指定誰創建了專案，並且可以由字串或包含聯絡方式或其他細節的物件組成。對於較大的專案，建議使用物件，但像以下範例這樣的簡單字串對於此專案來說也足夠。

```json
"author": "Jane Doe",
```

# --instructions--

在 `package.json` 檔案中將你的名稱添加為專案的 `author`。

**注意：** 請記得你正在撰寫 JSON，因此所有欄位名稱必須使用雙引號（"）並以逗號（,）分隔。

# --hints--

`package.json` 應該有一個有效的 "author" 鍵。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.author, '"author" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
