---
id: 587d7fb4367417b2b2512c00
title: 使用來自 npm 的外部軟體包擴充你的專案
challengeType: 2
forumTopicId: 301527
dashedName: expand-your-project-with-external-packages-from-npm
---

# --description--

使用套件管理工具的最大原因之一，是它們強大的相依性管理。當你在新的電腦上設定專案時，不必手動確保取得所有相依項，`npm` 會自動為你安裝所有東西。但 `npm` 怎麼能準確知道你的專案需要什麼？請看你的 package.json 檔案中的 `dependencies` 區段。

在本節中，你的專案所需的軟體包將使用以下格式儲存：

```json
"dependencies": {
  "package-name": "version",
  "express": "4.14.0"
}

```

# --instructions--

將版本 `1.1.0` 的 `@freecodecamp/example` 軟體包添加到你的 `package.json` 檔案的 `dependencies` 欄。

**注意：** `@freecodecamp/example` 是一個用作學習工具的假軟體包。

# --hints--

`"dependencies"` 應該包含 `"@freecodecamp/example"`。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`"@freecodecamp/example"` 版本應該是 `"1.1.0"`。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^[\^\~]?1\.1\.0/,
        'Wrong version of "@freecodecamp/example" installed. It should be 1.1.0'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
