---
id: 587d7fb4367417b2b2512bfe
title: 為你的 `package.json` 添加授權許可
challengeType: 2
forumTopicId: 301523
dashedName: add-a-license-to-your-package-json
---

# --description--

`license` 欄是你告知使用者他們可以對你的專案執行哪些操作的地方。

一些常見的開源專案授權包括 MIT 和 BSD。授權資訊並非必要，大多數國家的著作權法會預設賦予你對你所創建內容的所有權。然而，明確說明使用者可以和不可以做什麼總是一個良好的做法。以下是授權欄位的範例：

```json
"license": "MIT",
```

# --instructions--

在你的專案的 package.json 檔案中，依照你的需求填寫 `license` 欄。

# --hints--

`package.json` 應該有一個有效的 "license" 鍵。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.license, '"license" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
