---
id: 587d7fb4367417b2b2512bff
title: 為你的 package.json 添加版本號
challengeType: 2
forumTopicId: 301525
dashedName: add-a-version-to-your-package-json
---

# --description--

`version` 是你的 `package.json` 檔案中必填的欄位之一。此欄位描述你的專案目前的版本。以下是一個範例：

```json
"version": "1.2.0",
```

# --instructions--

為你的專案的 package.json 檔案添加一個 `version`。

# --hints--

`package.json` 應該有一個有效的 "version" 鍵。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.version, '"version" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
