---
id: 587d7fb3367417b2b2512bfc
title: 為你的 `package.json` 添加描述
challengeType: 2
forumTopicId: 301522
dashedName: add-a-description-to-your-package-json
---

# --description--

良好 `package.json` 檔案的下一個部分是 `description` 欄，這裡放置關於你的專案的簡短但具資訊性的描述。

如果有一天你計劃將軟體包發佈到 npm，這就是當使用者決定是否安裝你的軟體包時，應該用來推銷你想法的字串。然而，描述的用途不僅限於此，它也是總結專案功能的絕佳方式。在任何 Node.js 專案中，它同樣重要，有助於其他開發者、未來的維護者甚至你未來的自己快速理解專案。

無論你為你的專案規劃什麼，建議一定要有描述。以下是一個範例：

```json
"description": "A project that does something awesome",
```

# --instructions--

為你的專案的 `package.json` 檔案添加 `description`。

**注意：** 請記得對欄位名稱使用雙引號（"）並用逗號（,）分隔欄位。

# --hints--

`package.json` 應該有一個有效的 "description" 鍵。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.description, '"description" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
