---
id: 587d7fb4367417b2b2512bfd
title: 為你的 `package.json` 添加 Keywords
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

`keywords` 欄是你可以使用相關關鍵字描述你的專案的地方。以下是一個範例：

```json
"keywords": [ "descriptive", "related", "words" ],
```

如你所見，這個欄位被結構為一個雙引號字串的陣列。

# --instructions--

為你的專案的 `package.json` 檔案中的 `keywords` 欄添加一個適當字串的陣列。

其中一個關鍵字應該是 "freecodecamp"。

# --hints--

`package.json` 應該有一個有效的 "keywords" 鍵。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`keywords` 欄位應該是一個陣列

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

「keywords」應該包含「freecodecamp」

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
