---
id: 587d7fb5367417b2b2512c02
title: 使用波浪號字元來始終使用相依性的最新修補版本
challengeType: 2
forumTopicId: 301532
dashedName: use-the-tilde-character-to-always-use-the-latest-patch-version-of-a-dependency
---

# --description--

在上一個挑戰中，你告訴 npm 只包含軟體包的特定版本。如果你需要確保專案中不同零件彼此相容，這是凍結相依性的有用方法。但在大多數使用情況下，你不想錯過錯誤修正，因為它們通常包含重要的安全修補程式，且（希望）不會因此破壞功能。

為了允許 npm 相依性更新到最新的 PATCH 版本，你可以在相依性的版本前置式加上波浪號（`~`）字元。以下是一個允許更新到任何 `1.3.x` 版本的範例。

```json
"package": "~1.3.8"
```

# --instructions--

在 `package.json` 檔案中，你目前對於 npm 如何升級 `@freecodecamp/example` 的規則是使用特定版本（`1.2.13`）。但現在，你想允許最新的 `1.2.x` 版本。

使用波浪號（`~`）字元作為你相依套件中 `@freecodecamp/example` 版本的前置式，並允許 npm 將其更新到任何新的 *patch* 版本。

**注意：** 版本號本身不應該被更改。

# --hints--

`"dependencies"` 應該包含 `"@freecodecamp/example"`。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`"@freecodecamp/example"` 版本應該符合 `"~1.2.13"`。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\~1\.2\.13/,
        'Wrong version of "@freecodecamp/example". It should be ~1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
