---
id: 587d7fb5367417b2b2512c03
title: 使用插入符號字元來使用相依性的最新次要版本
challengeType: 2
forumTopicId: 301531
dashedName: use-the-caret-character-to-use-the-latest-minor-version-of-a-dependency
---

# --description--

類似於我們在上一個挑戰中學到的波浪號允許 npm 安裝相依性的最新 PATCH，插入符號（`^`）也允許 npm 安裝未來的更新。不同之處在於插入符號會允許 MINOR 更新和 PATCH。

你目前的 `@freecodecamp/example` 版本應該是 `~1.2.13`，這允許 npm 安裝最新的 `1.2.x` 版本。如果你使用插入符號 (^) 作為版本前置式，npm 將允許更新到任何 `1.x.x` 版本。

```json
"package": "^1.3.8"
```

這將允許更新到該軟體包的任何 `1.x.x` 版本。

# --instructions--

使用脫字符號（`^`）為你相依的 `@freecodecamp/example` 版本加上前置式，並允許 npm 將其更新到任何新的次要版本。

**注意：** 版本號本身不應該被更改。

# --hints--

`"dependencies"` 應該包含 `"@freecodecamp/example"`。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`"@freecodecamp/example"` 版本應該符合 `"^1.x.x"`。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\^1\./,
        'Wrong version of "@freecodecamp/example". It should be ^1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
