---
id: 65c4ef8463f869d0eaf87c67
title: 步驟2
challengeType: 20
dashedName: step-2
---

# --description--

在 `TreeNode` 類中，將 `pass` 替換爲 `__init__` 方法，以便可以初始化對象的屬性。 暫時不要添加任何參數。

# --hints--

你應該從 `TreeNode` 類中刪除 `pass` 關鍵字，並將其移至 `__init__` 方法內。

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

你應該在 `TreeNode` 類中定義一個 `__init__` 方法。 請記住在其內部使用 `pass` 關鍵字。

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
