---
id: 5f33071498eb2472b87ddee4
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому проєкті ви вивчите основи CSS (каскадних таблиць стилів), створюючи меню кафе. CSS — це мова, яка використовується для стилізації документа HTML. Вона описує, як елементи HTML мають відтворюватись на екрані.

Як ви дізналися в останніх кількох кроках застосунку «Cat Photo», існує базова структура, яка необхідна для початку створення вебсторінки. Кожен HTML-документ повинен мати оголошення `DOCTYPE` та елемент `html`. `DOCTYPE` повідомляє браузеру, яку версію HTML використовує документ. А елемент `html` представляє кореневий елемент, який містить усі інші елементи.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Додайте теґ `<!DOCTYPE html>` та елемент `html` з атрибутом `lang` зі значенням `en`.

# --hints--

Ви повинні мати оголошення `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

У вас має бути початковий `<html lang="en">` з англійською мовою.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ви повинні мати кінцевий тег `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
