---
id: 5f331e55dfab7a896e53c3a1
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

The `title` is one of several elements that provide extra information not visible on the web page, but it is useful for search engines or how the page gets displayed.

Всередині елемента `head` вкладіть елемент `meta` з атрибутом `charset` зі значенням `utf-8`, щоб повідомити браузеру як кодувати символи сторінки.

# --hints--

Ви повинні мати теґ `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

Елемент `meta` є порожнім елементом, він не повинен мати кінцевий тег `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Тег `meta` повинен мати атрибут `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Your `charset` attribute should have a value of `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
