---
id: 5f332a88dc25a0fd25c7687a
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

The name of the cafe is `CAMPER CAFE`. Додайте елемент `h1` в межах елемента `main`. Напишіть назву кафе великими літерами, щоб вона виділялась.

# --hints--

Ви повинні мати початковий тег `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Ви повинні мати кінцевий тег `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Не змінюйте наявний елемент `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Елемент `h1` повинен бути вкладеним в елементі `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Елемент `h1` повинен містити текст `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

