---
id: 5f3326b143638ee1a09ff1e3
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Щоб підготуватись до створення вмісту, додайте елемент `body` під елементом `head`.

# --hints--

Ви повинні мати початковий тег `<body>`.

```js
assert(code.match(/<body>/i));
```

Ви повинні мати кінцевий тег `</body>`.

```js
assert(code.match(/<\/body>/i));
```

Не змінюйте елемент `head`. Переконайтеся, що не видалили кінцевий тег.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

Елемент `body` повинен йти після елемента `head`.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

