---
id: 5f33071498eb2472b87ddee4
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS is the language used to style an HTML document. It describes how HTML elements should be displayed on the screen.

As you learned in the last few steps of the Cat Photo App, there is a basic structure needed to start building your web page. Кожен HTML-документ повинен мати оголошення `DOCTYPE` та елемент `html`. `DOCTYPE` повідомляє браузеру, яку версію HTML використовує документ. А елемент `html` представляє кореневий елемент, який містить усі інші елементи.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Add the `<!DOCTYPE html>` tag, and an `html` element with a `lang` attribute of `en`.

# --hints--

Ви повинні мати оголошення `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

У вас має бути початковий `<html lang="en">` з англійською мовою.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ви повинні мати кінцевий тег `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
