---
id: 5f33294a6af5e9188dbdb8f3
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Час додати деякий вміст меню. Додайте елемент `main` в межі наявного елемента `body`. В кінцевому підсумку там буде ціна напоїв та десертів, які пропонує кафе.

# --hints--

Код повинен містити початковий тег `<main>`.

```js
assert(code.match(/<main>/i));
```

Код повинен містити кінцевий тег `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Не змінюйте наявний елемент `body`. Переконайтеся, що випадково не видалили кінцевий тег.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Тег `main` повинен бути в межах тегу `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
