---
id: 5f344fbc22624a2976425065
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Create an `h2` element in the `section` element and give it the text `Coffee`.

# --hints--

Ви повинні мати початковий тег `<h2>`.

```js
assert(code.match(/<h2\s*>/i));
```

Ви повинні мати кінцевий тег `</h2>`.

```js
assert(code.match(/<\/h2\s*>/i));
```

Не змінюйте наявний елемент `section`. Переконайтеся, що не видалили кінцевий тег.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

Елемент `h2` має бути в межах елемента `section`.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

Елемент `h2` повинен містити текст `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
