---
id: 5f3c866daec9a49519871816
title: Крок 30
challengeType: 0
dashedName: step-30
---

# --description--

Елементи `article` зазвичай містять декілька елементів з пов'язаною інформацією. В цьому разі він міститиме смак кави та його ціну. Вкладіть два елементи `p` всередині елемента `article`. Текст першого повинен бути `French Vanilla`, текст другого – `3.00`.

# --hints--

Не змінюйте наявний елемент `article`.

```js
assert.lengthOf(document.querySelectorAll('article'),1);
```

Ваш елемент `article` повинен мати два елементи `p`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphChildren,2);
```

Перший елемент `p` повинен містити текст `French Vanilla`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const firstP = paragraphChildren?.[0];
assert.match(firstP?.innerText, /French Vanilla/i);
```

Другий елемент `p` повинен містити текст `3.00`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const secondP = paragraphChildren?.[1];
assert.match(secondP?.innerText, /3\.00/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
