---
id: 6a89c69cbe1bb69845992d8b
title: بناء مولد نطاق أعداد
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

في هذا المختبر، ستبني دالة `range_of_numbers` تستخدم الاستدعاء الذاتي لتوليد قائمة أعداد ضمن نطاق محدد.

على سبيل المثال:

- يجب أن تُرجع الدالة `range_of_numbers(3, 9)` القائمة `[3, 4, 5, 6, 7, 8, 9]`.
- يجب أن تُرجع الدالة `range_of_numbers(5, 5)` القائمة `[5]`.

نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `range_of_numbers` بمعلمتين اسميهما `start_num` و `end_num`.
1. يجب أن تُرجع الدالة قائمة من أعداد صحيحة متتالية تبدأ بـ `start_num` وتنتهي بـ `end_num`، مع تضمين القيمتين.
1. ستكون قيمة `start_num` دائمًا أقل من أو تساوي `end_num`.
1. يجب أن تستخدم دالتك الاستدعاء الذاتي عبر استدعاء نفسها.
1. يجب ألا تستخدم دالتك حلقات `for` أو `while`، أو التراكيب التكرارية، أو دالة `range()`.
1. عندما تكون قيمة `start_num` مساوية لـ `end_num`، يجب أن تُرجع الدالة قائمة تحتوي على `start_num`. هذه هي الحالة الأساسية.
1. في الحالة الاستدعائية، يجب أن تستدعي الدالة نفسها بمعلمة تتحرك نحو الحالة الأساسية، ثم تضيف الرقم الحالي إلى القائمة المرجعة.
1. يجب ألا تستخدم الدالة متغيرات عامة لتخزين القائمة المرجعة.

# --hints--

يجب أن يكون لديك دالة باسم `range_of_numbers` بمعلمتين اسميهما `start_num` و `end_num`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

يجب أن تُرجع الدالة `range_of_numbers(1, 5)` القائمة `[1, 2, 3, 4, 5]`.

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

يجب أن تُرجع الدالة `range_of_numbers(6, 9)` القائمة `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

يجب أن تُرجع الدالة `range_of_numbers(4, 4)` القائمة `[4]`.

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

يجب أن تُرجع الدالة `range_of_numbers(10, 15)` القائمة `[10, 11, 12, 13, 14, 15]`.

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

يجب أن تستدعي دالتك `range_of_numbers` نفسها.

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

يجب ألا تستخدم دالتك `range_of_numbers` حلقات أو تراكيب تكرارية أو دالة `range()`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

استدعاء `range_of_numbers(1, 3)` يجب ألا يؤثر على استدعاء لاحق. بعد هذا الاستدعاء، يجب أن تُرجع `range_of_numbers(6, 9)` القائمة `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
