---
id: a202eed8fc186c8434cb6d61
title: 建置一個字串反轉器
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

在這個實驗中，你將建置一個簡單的字串反轉器，用來反轉給定字串的字元。

例如，`"hello"` 應該變成 `"olleh"`。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `reverseString` 的函式（程式），該函式（程式）接受一個字串作為引數。
2. 該函式（程式）應該傳回反轉的字串。

# --hints--

你應該有一個名為 `reverseString` 的函式（程式）。

```js
assert.isFunction(reverseString);
```

`reverseString` 應該接受一個字串作為引數。

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` 應該傳回一個字串。

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` 應該傳回字串 `olleh`。

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` 應該傳回字串 `ydwoH`。

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` 應該傳回字串 `htraE morf sgniteerG`。

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
