---
id: a202eed8fc186c8434cb6d61
title: 創建一個字符串反轉器
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

在本實驗室中，你將創建一個簡單的字符串反轉器，它可以反轉給定字符串的字符。

例如， `"hello"` 應該變成 `"olleh"`。

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶故事：**

1. 你應創建一個名爲 `reverseString` 的函數，該函數以字符串爲參數。
2. 函數應返回反轉的字符串。

# --hints--

你應該有一個名爲 `reverseString` 的函數。

```js
assert.isFunction(reverseString);
```

`reverseString` 應使用字符串作爲參數。

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` 應返回一個字符串。

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` 應返回 `olleh`。

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` 應返回 `ydwoH`。

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` 應返回 `htraE morf sgniteerG`。

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
