---
id: 62a115879a6d51422652cbfc
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

在你的角色扮演游戏中，用户将能够跟踪他们的统计数据，购买武器并与怪物战斗。 在继续游戏的交互式 JavaScript 部分之前，你需要先创建将显示游戏信息的 HTML 元素。

在 `#game` 元素中创建四个 `div` 元素。 给它们分配各自的 `id` 值，依次为：`stats`、`controls`、`monsterStats`和`text`。

# --hints--

应该创建四个新的 `div` 元素。

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

你应该给新的 `div` 元素之一一个 `id` 为 `stats`。

```js
assert.exists(document.querySelector('div#stats'));
```

你应该给新的 `div` 元素之一一个 `id` 为 `controls`。

```js
assert.exists(document.querySelector('div#controls'));
```

你应该给新的 `div` 元素之一一个 `id` 为 `monsterStats`。

```js
assert.exists(document.querySelector('div#monsterStats'));
```

你应该给新的 `div` 元素之一一个 `id` 为 `text`。

```js
assert.exists(document.querySelector('div#text'));
```

应该以正确的顺序放置新的 `div` 元素。

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

应该将新的 `div` 元素放在 `#game` 元素中。

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
