---
id: 62a23c1d505bfa13747c8a9b
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

将数字 `0`、`100` 和 `50` 包含在 `span` 元素中，并将这些新的 `span` 元素包含在 `strong` 元素中。 然后分别为新的 `span` 元素设置 `id` 值为 `xpText`、`healthText` 和 `goldText` 。

你的答案应该遵循这个基本结构：

```html
<span class="stat">TEXT <strong><span id="VALUE">TEXT</span></strong></span>
```

# --hints--

你应该在第一个 `.stat` 元素中添加一个 `strong` 元素。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

你的第一个新的 `strong` 元素应该有一个 `span` 元素。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

你的第一个嵌套 `span` 元素应该具有 `xpText` 的 `id`。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'xpText');
```

你的第一个 `span` 元素应该包裹在文本 `0` 周围。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '0');
```

你的第一个 `.stat` 元素的文本仍应为 `XP: 0`。

```js
const stat = document.querySelectorAll('.stat')[0];
assert.equal(stat.innerText.trimEnd(), 'XP: 0');
```

你应该在第二个 `.stat` 元素中添加一个 `strong` 元素。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

你的第二个新 `strong` 元素应该有一个 `span` 元素。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

你的第二个嵌套 `span` 元素应该具有 `healthText` 的 `id`。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'healthText');
```

你的第二个 `span` 元素应包裹文本 `100`。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '100');
```

你的第二个 `.stat` 元素的文本仍应为 `Health: 100`。

```js
const stat = document.querySelectorAll('.stat')[1];
assert.equal(stat.innerText.trimEnd(), 'Health: 100');
```

你应该在第三个 `.stat` 元素中添加一个 `strong` 元素。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

你的第三个新 `strong` 元素应该有一个 `span` 元素。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

你的第三个嵌套 `span` 元素应该具有 `goldText` 的 `id`。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'goldText');
```

你的第三个 `span` 元素应包裹文本 `50`。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '50');
```

你的第三个 `.stat` 元素的文本仍应为 `Gold: 50`。

```js
const stat = document.querySelectorAll('.stat')[2];
assert.equal(stat.innerText.trimEnd(), 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
--fcc-editable-region--
        <span class="stat">XP: 0</span>
        <span class="stat">Health: 100</span>
        <span class="stat">Gold: 50</span>
--fcc-editable-region--
      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
