---
id: 62a3bec30ea7f941412512dc
title: 步骤 29
challengeType: 0
dashedName: step-29
---

# --description--

就像你对按钮所做的那样，为以下 `id` 创建变量并使用 `querySelector()` 将元素作为值赋予它们：

`text`、`xpText`、`healthText`、`goldText`、`monsterStats` 和 `monsterName`。

请记住使用 `const` 关键字声明这些，并将变量命名为与 `id` 相匹配。

# --hints--

你应该使用关键字 `const` 声明一个名叫 `text` 变量。

```js
assert.match(code, /const\s+text/);
```

你的 `text` 变量应该具有 `#text` 元素的值。

```js
assert.match(code, /const\s+text\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#text\1\s*\)\s*/);
```

你应该使用关键字 `const` 声明一个名为 `xpText` 的变量。

```js
assert.match(code, /const\s+xpText/);
```

你的 `xpText` 变量应该具有 `#xpText` 元素的值。

```js
assert.match(code, /const\s+xpText\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#xpText\1\s*\)\s*/);
```

你应该使用关键字 `const` 声明一个名为 `healthText` 变量。

```js
assert.match(code, /const\s+healthText/);
```

你的 `healthText` 变量应该具有 `#healthText` 元素的值。

```js
assert.match(code, /const\s+healthText\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#healthText\1\s*\)\s*/);
```

你应该使用关键字 `const` 声明一个 `goldText` 变量。

```js
assert.match(code, /const\s+goldText/);
```

你的 `goldText` 变量应该具有 `#goldText` 元素的值。

```js
assert.match(code, /const\s+goldText\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#goldText\1\s*\)\s*/);
```

你应该使用关键字 `const` 声明一个名为 `monsterStats` 变量。

```js
assert.match(code, /const\s+monsterStats/);
```

你的 `monsterStats` 变量应该具有 `#monsterStats` 元素的值。

```js
assert.match(code, /const\s+monsterStats\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#monsterStats\1\s*\)\s*/);
```

你应该使用关键字 `const` 声明一个变量名为 `monsterName` 的变量。

```js
assert.match(code, /const\s+monsterName/);
```

你的 `monsterName` 变量应该具有 `#monsterName` 元素的值。

```js
assert.match(code, /const\s+monsterName\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#monsterName\1\s*\)\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
