---
id: 62a3b3eab50e193608c19fc6
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

声明一个名为 `fighting` 的变量，但不要用值初始化它。 

# --hints--

你应该使用 `let` 来声明变量 `fighting` 。

```js
assert.match(code, /let\s+fighting/);
```

你的 `fighting` 变量不应该有值。

```js
assert.isUndefined(fighting);
```

你不应该为你的 `fighting` 变量赋值。 

```js
assert.match(code, /let\s+fighting\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
--fcc-editable-region--
```
