---
id: 62a3b41c9494f937560640ab
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

再声明两个名为 `monsterHealth` 和 `inventory` 的变量。

对于你的 `inventory` 变量，为其分配包含字符串 `"stick"` 的数组的值。

想想你在之前的项目中学会了像这样数组方法：

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

你应该使用 `let` 关键词声明你的 `monsterHealth` 变量。

```js
assert.match(code, /let\s+monsterHealth/i);
```

你应该使用驼峰命名法（camelCase）来命名 `monsterHealth` 变量。

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` 不应该有值。

```js
assert.isUndefined(monsterHealth);
```

你不应该为你的 `monsterHealth` 变量赋值。

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

你应该使用 `let` 关键词声明你的 `inventory` 变量。

```js
assert.match(code, /let\s+inventory/i);
```

你的 `inventory` 变量仍然应该是一个数组。

```js
assert.isArray(inventory);
```

你的 `inventory` 变量只有1个值。

```js
assert.lengthOf(inventory, 1);
```

你的 `inventory` 变量应存放着 `"stick"` 字符串。

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
