---
id: 62a8d81f539f004776dd9b1e
title: 步骤 110
challengeType: 0
dashedName: step-110
---

# --description--

现在你可以启动代码来对抗怪物了。 为了使你的代码保持井然有序，你的 `fightDragon` 函数已被移动，以便靠近其他 `fight` 函数。

在 `weapons` 数组下方，定义一个 `monsters` 变量并为其分配一个数组。 将该数组设置为具有三个对象，每个对象均具有 `name`、`level` 和 `health` 属性。 第一个对象的值应该按顺序为 `"slime"`、`2` 和 `15`。 第二个应该是 `"fanged beast"`、`8` 和 `60`。 第三个应该是 `"dragon"`、`20` 和 `300`。

# --hints--

你应该使用 `const` 来声明 `monsters` 变量。

```js
assert.match(code, /const\s+monsters/);
```

你的 `monsters` 变量应该是一个数组。

```js
assert.isArray(monsters);
```

你的 `monsters` 变量应该有 3 个值。

```js
assert.lengthOf(monsters, 3);
```

你的 `monsters` 数组应该有 3 个对象。

```js
assert(monsters.every(val => typeof val === "object"));
```

`monsters` 数组中的第一个值应该是一个对象，其 `name` 属性设置为 `"slime"`。

```js
assert.equal(monsters[0].name, "slime");
```

`monsters` 数组中的第一个值应该是一个对象，其 `level` 属性设置为 `2`。

```js
assert.equal(monsters[0].level, 2);
```

`monsters` 数组中的第一个值应该是一个对象，其 `health` 属性设置为 `15`。

```js
assert.equal(monsters[0].health, 15);
```

`monsters` 数组中的第二个值应该是一个对象，其 `name` 属性设置为 `"fanged beast"`。

```js
assert.equal(monsters[1].name, "fanged beast");
```

`monsters` 数组中的第二个值应该是一个对象，其 `level` 属性设置为 `8`。

```js
assert.equal(monsters[1].level, 8);
```

`monsters` 数组中的第二个值应该是一个对象，其 `health` 属性设置为 `60`。

```js
assert.equal(monsters[1].health, 60);
```

`monsters` 数组中的第三个值应该是一个对象，其 `name` 属性设置为 `"dragon"`。

```js
assert.equal(monsters[2].name, "dragon");
```

`monsters` 数组中的第三个值应该是一个对象，其 `level` 属性设置为 `20`。

```js
assert.equal(monsters[2].level, 20);
```

`monsters` 数组中的第三个值应该是一个对象，其 `health` 属性设置为 `300`。

```js
assert.equal(monsters[2].health, 300);
```

你不应该向你的对象添加任何额外的属性。

```js
assert.deepEqual(monsters, [
  {name: "slime", level: 2, health: 15},
  {name: "fanged beast", level: 8, health: 60},
  {name: "dragon", level: 20, health: 300}
])
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const weapons = [
  { name: 'stick', power: 5 },
  { name: 'dagger', power: 30 },
  { name: 'claw hammer', power: 50 },
  { name: 'sword', power: 100 }
];
--fcc-editable-region--

--fcc-editable-region--
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
  {
    name: "cave",
    "button text": ["Fight slime", "Fight fanged beast", "Go to town square"],
    "button functions": [fightSlime, fightBeast, goTown],
    text: "You enter the cave. You see some monsters."
  }
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  update(locations[2]);
}

function buyHealth() {
  if (gold >= 10) {
    gold -= 10;
    health += 10;
    goldText.innerText = gold;
    healthText.innerText = health;
  } else {
    text.innerText = "You do not have enough gold to buy health.";
  }
}

function buyWeapon() {
  if (currentWeaponIndex < weapons.length - 1) {
    if (gold >= 30) {
      gold -= 30;
      currentWeaponIndex++;
      goldText.innerText = gold;
      let newWeapon = weapons[currentWeaponIndex].name;
      text.innerText = "You now have a " + newWeapon + ".";
      inventory.push(newWeapon);
      text.innerText += " In your inventory you have: " + inventory;
    } else {
      text.innerText = "You do not have enough gold to buy a weapon.";
    }
  } else {
    text.innerText = "You already have the most powerful weapon!";
    button2.innerText = "Sell weapon for 15 gold";
    button2.onclick = sellWeapon;
  }
}

function sellWeapon() {
  if (inventory.length > 1) {
    gold += 15;
    goldText.innerText = gold;
    let currentWeapon = inventory.shift();
    text.innerText = "You sold a " + currentWeapon + ".";
    text.innerText += " In your inventory you have: " + inventory;
  } else {
    text.innerText = "Don't sell your only weapon!";
  }
}

function fightSlime() {

}

function fightBeast() {

}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
