---
id: 657efce98e958b75df86b305
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Innerhalb des `if`-Anweisungskörpers fügen Sie das konvertierte Zeichen zu der zuvor erstellten Liste hinzu.

Dafür wird die `.append()`-Methode verwendet. Diese Methode fügt ein gegebenes Objekt ans Ende der Liste hinzu, auf der sie aufgerufen wird.

Verwenden Sie die `.append()`-Methode auf dem `snake_cased_char_list`, um das `converted_character` zur Liste hinzuzufügen.

# --hints--

Sie sollten die `.append()`-Methode verwenden, um das konvertierte Zeichen zum `snake_cased_char_list` hinzuzufügen.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
