---
id: 657effaa2a5e0277d71f9cbe
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Bis zu diesem Punkt enthält die Variable `snake_cased_char_list` die Liste der umgewandelten Zeichen. Um diese Zeichen zu einem einzigen String zu verbinden, können Sie die `.join()`-Methode verwenden.

Die `join`-Methode funktioniert, indem sie jedes Element einer Liste zu einem String zusammenfügt, getrennt durch einen festgelegten String, der als Trennzeichen bezeichnet wird.

```py
result_string = ''.join(characters)
```

Das obige Beispiel verbindet die Elemente der `characters`-Liste zu einem einzigen String, wobei jedes Element mit einem leeren String als Trennzeichen aneinandergereiht wird.

Verwenden Sie nun direkt nach der `for`-Schleife die `.join()`-Methode, um die Elemente in `snake_cased_char_list` mit einem leeren String als Trennzeichen zu verbinden. Weisen Sie das Ergebnis einer neuen Variablen namens `snake_cased_string` zu.


# --hints--

Sie sollten die Zeichen in der `snake_cased_char_list`-Liste mit der `.join()`-Methode zu einem String verbinden und den String nach Ihrer `snake_cased_string`-Schleife einer Variablen namens `for` zuweisen.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
