---
id: 657ed53c19461d4b95c4757a
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Sie müssen eine leere Liste hinzufügen, die die Zeichen des Strings aufnehmen wird, nachdem Sie sie in snake case umgewandelt haben.

Ersetzen Sie innerhalb der Funktion die `pass`-Anweisung, indem Sie eine leere Liste mit dem Namen `snake_cased_char_list` erstellen.

# --hints--

Sie sollten `pass` durch eine leere Liste mit dem Namen `snake_cased_char_list` innerhalb von `convert_to_snake_case` ersetzen.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const commentless_code = __helpers.python.removeComments(transformedCode)
        const convert_to_snake_case = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;
        assert.notMatch(function_body, /\bpass\b/);
        assert.match(function_body, / +snake_cased_char_list\s*=\s*\[\s*\]\s*/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    pass    
--fcc-editable-region--
```
