---
id: 657efa642593c5746acc5c81
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Innerhalb des `if`-Anweisungsblocks müssen Sie jeden Großbuchstaben in Kleinbuchstaben umwandeln und diesem Kleinbuchstaben einen Unterstrich voranstellen.

Verwenden Sie die `.lower()`-Stringmethode, um Großbuchstaben in Kleinbuchstaben umzuwandeln. Fügen Sie dann dem Zeichen einen Unterstrich voran. Weisen Sie die Ergebnisse einer Variablen mit dem Namen `converted_character` zu.

# --hints--

Sie sollten das veränderte Zeichen einer Variablen mit dem Namen `converted_character` zuweisen.

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

Sie sollten `pass` nicht in Ihrer `if`-Anweisung haben.

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
