---
id: 657e928716b77b2277980276
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Projekt erstellen Sie ein Programm, das einen `camelCase` oder `PascalCase` formatierten String als Eingabe erhält und diesen mit zwei Methoden in einen `snake_case` formatierten String umwandelt. Zuerst verwenden Sie eine for-Schleife, anschließend List Comprehension, um dasselbe Ergebnis zu erzielen. Sie werden sehen, wie List Comprehension Ihren Code prägnanter machen kann.

Beginnen Sie damit, eine neue Funktion namens `convert_to_snake_case()` zu definieren, die einen String namens `pascal_or_camel_cased_string` als Eingabe akzeptiert. Fügen Sie vorerst eine `pass`-Anweisung innerhalb der Funktion hinzu.

# --hints--

Sie sollten eine neue Funktion namens `convert_to_snake_case()` definieren. Vergessen Sie nicht den Doppelpunkt am Ende und verwenden Sie `pass`, um den Funktionskörper zu füllen.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Ihre Funktion sollte einen Parameter namens `pascal_or_camel_cased_string` annehmen.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
