---
id: 657ef2a86d4e545cec9a85fb
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Sowohl im camel case als auch im pascal case markieren Großbuchstaben den Beginn neuer Wörter. Um den Eingabe-String in snake case umzuwandeln, müssen Sie überprüfen, ob die Zeichen im Eingabe-String Großbuchstaben sind.

Sie können die `.isupper()` String-Methode verwenden, um zu prüfen, ob ein Zeichen ein Großbuchstabe ist. Diese Methode gibt `True` zurück, wenn das Zeichen ein Großbuchstabe ist, und `False`, wenn es keiner ist.

Fügen Sie innerhalb der `for`-Schleife eine `if`-Anweisung hinzu, um zu prüfen, ob das aktuelle Zeichen ein Großbuchstabe ist. Verschieben Sie die `pass`-Anweisung in die neue `if`-Anweisung.

# --hints--

Sie sollten eine neue if-Anweisung mit char.isupper() als Bedingung schreiben. Denken Sie daran, am Ende einen Doppelpunkt zu setzen und das Schlüsselwort pass zu verwenden, um den Rumpf der if-Anweisung zu füllen.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

Sie sollten die `pass`-Anweisung im Schleifenrumpf durch die `if`-Anweisung ersetzen.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
