---
id: 657ee28cefc4945568287673
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Mit der leeren Liste können Sie nun beginnen, den Eingabe-String durchzulaufen und ihn in snake case umzuwandeln.

Fügen Sie innerhalb der Funktion unterhalb der gerade erstellten Liste eine `for`-Schleife hinzu, um den `pascal_or_camel_cased_string` zu durchlaufen. Achten Sie darauf, die Zielvariable `char` zu nennen. Fügen Sie vorerst eine `pass`-Anweisung als Platzhalter in den Schleifenkörper ein.

# --hints--

Sie sollten eine neue `for`-Schleife schreiben, deren Zielvariable `char` heißt. Vergessen Sie nicht den Doppelpunkt am Ende und verwenden Sie `pass`, um den Schleifenkörper zu füllen.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
