---
id: 657f465f8e718b19c5105ae5
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

List Comprehensions akzeptieren bedingte Anweisungen, um den angegebenen Ausdruck nur dann auszuwerten, wenn bestimmte Bedingungen erfüllt sind:

```py
spam = [i * 2 for i in iterable if i > 0]
```

Wie Sie an der Ausgabe sehen können, wurde die Liste der Zeichen, die aus `pascal_or_camel_cased_string` generiert wurde, zusammengefügt. Da der Ausdruck innerhalb der List Comprehension für jedes Zeichen ausgewertet wird, ist das Ergebnis ein String in Kleinbuchstaben, bei dem alle Zeichen durch einen Unterstrich getrennt sind.

Folgen Sie dem obigen Beispiel und fügen Sie Ihrer List Comprehension eine `if`-Klausel hinzu, sodass der Ausdruck nur ausgeführt wird, wenn das Zeichen Großbuchstaben ist.

# --hints--

Sie sollten Ihrer List Comprehension eine `if`-Klausel mit der Bedingung `char.isupper()` hinzufügen.

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
