---
id: 657f4add33ea4b1f61ba3dc8
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Probieren Sie diese neue Implementierung schließlich aus, indem Sie das Programm ausführen. Ändern Sie den Eingabe-String zu `'IAmAPascalCasedString'` und prüfen Sie, ob er als `'i_am_a_pascal_cased_string'` ausgegeben wird, auch wenn das eine Lüge ist.

Wenn Sie alles richtig gemacht haben, sollten Sie den Eingabe-String wie zuvor in snake case umgewandelt sehen.

Herzlichen Glückwunsch! Jetzt ist Ihre `convert_to_snake_case()`-Funktion bereit.

# --hints--

Sie sollten den Eingabe-String von `'aLongAndComplexString'` zu `'IAmAPascalCasedString'` innerhalb der `main()`-Funktion ändern.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
