---
id: 657f456223b8c1187b461987
title: Schritt 19
challengeType: 20
dashedName: step-19
---

# --description--

In Python ist eine <dfn>list comprehension</dfn> eine Konstruktion, mit der Sie eine neue Liste erzeugen können, indem Sie einen Ausdruck auf jedes Element einer vorhandenen iterierbaren Struktur anwenden und optional Elemente mit einer Bedingung filtern. Abgesehen davon, dass sie kürzer ist, laufen list comprehensions oft schneller.

Eine einfache list comprehension besteht aus einem Ausdruck, gefolgt von einer `for`-Klausel:

```py
spam = [i * 2 for i in iterable]
```

Oben wird die Variable `i` verwendet, um `iterable` zu durchlaufen. Jedes Element der resultierenden Liste wird durch Auswerten des Ausdrucks `i * 2` bei der aktuellen Iteration erhalten.

In diesem Schritt müssen Sie die leere Liste `snake_cased_char_list` mit der Syntax der list comprehension füllen.

Verwandeln Sie Ihre leere Liste in eine list comprehension, die jedes Zeichen in `pascal_or_camel_cased_string` in Kleinbuchstaben umwandelt und ihm einen Unterstrich voranstellt (der Code, den Sie zuvor auskommentiert haben, kann Ihnen beim Schreiben des Ausdrucks helfen). Verwenden Sie `char`, um `pascal_or_camel_cased_string` zu durchlaufen.

# --hints--

Sie sollten `snake_cased_char_list` in eine list comprehension umwandeln, die `pascal_or_camel_cased_string` durchläuft.

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

Ihre list comprehension sollte `char` verwenden, um `pascal_or_camel_cased_string` zu durchlaufen.

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

Ihre list comprehension sollte `'_' + char.lower()` für jedes `char` in `pascal_or_camel_cased_string` auswerten.

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
