---
id: 657f0044be09db790b1eb1c5
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Im Pascal-Case beginnen Strings mit einem Großbuchstaben. Nachdem alle Zeichen in Kleinbuchstaben umgewandelt und Unterstriche hinzugefügt wurden, kann es passieren, dass am Anfang Ihres Strings ein zusätzlicher Unterstrich steht.

Der einfachste Weg, dies zu beheben, ist die Verwendung der `.strip()`-String-Methode, die aus einem String alle führenden oder abschließenden Zeichen entfernt, die zu einer Menge von Zeichen gehören, die als Argument übergeben werden. Zum Beispiel:

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

Die `strip()`-Methode wird auf `original_string` angewendet. Dadurch werden alle führenden und abschließenden Unterstriche entfernt. Das Ergebnis des obigen Beispiels wäre der String `'example_string'`.

Deklarieren Sie eine neue Variable namens `clean_snake_cased_string` und weisen Sie ihr das Ergebnis der `.strip()`-Methode zu, die auf `snake_cased_string` angewendet wird, wobei `'_'` als Argument an die Methode übergeben wird.

# --hints--

Sie sollten den resultierenden String von `snake_cased_string.strip('_')` einer Variablen namens `clean_snake_cased_string` zuweisen.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
