---
id: 657f04ed0035f47ed04d0f1f
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Um die Ausgabe der `convert_to_snake_case()`-Funktion anzuzeigen, müssen Sie die `main()`-Funktion aufrufen.

Fügen Sie auf derselben Ebene wie die beiden vorhandenen Funktionen einen Aufruf der `main()`-Funktion hinzu. Beim Ausführen sollte der gegebene camel- oder pascal-cased String in snake case umgewandelt werden.

# --hints--

Sie sollten einen Aufruf der `main()`-Funktion hinzufügen.

```js
({
  test: () => {
    const pythonCode = `_Node(_code.replace('\\\\r', '')).has_call("main()")`;
    const result = runPython(pythonCode);
    assert(result);
  }
});
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

--fcc-editable-region--
```
