---
id: 657f025ec86c3d7c4177b6be
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Mit der fertigen Funktion können Sie sie jetzt innerhalb einer anderen Funktion verwenden.

Erstellen Sie eine neue Funktion namens `main()` mit `pass` als Funktionskörper.

# --hints--

Sie sollten eine neue Funktion namens `main()` definieren. Vergessen Sie nicht den Doppelpunkt am Ende und verwenden Sie `pass`, um den Funktionskörper auszufüllen.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
