---
id: 657f425dbab54e11993c80f0
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Bisher haben Sie in diesem Projekt eine `for`-Schleife verwendet, um Ihren Eingabe-String zu durchlaufen und in die gewünschte Ausgabe umzuwandeln. Nun beginnen Sie den Übergang von einer `for`-Schleife zu einer List Comprehension.

Kommentieren Sie zunächst alle Codezeilen innerhalb der `convert_to_snake_case()`-Funktion aus. Löschen Sie diese nicht, da sie hilfreich sein werden, wenn Sie die Logik mit einer List Comprehension implementieren.

Denken Sie daran, das Schlüsselwort `pass` im Funktionskörper hinzuzufügen, um zu verhindern, dass der Code während der Tests fehlschlägt.

# --hints--

Sie sollten alle 10 Codezeilen innerhalb der `convert_to_snake_case()`-Funktion auskommentieren und sonst nichts. Fügen Sie `pass` hinzu, um den Funktionskörper vorübergehend zu füllen und einen Fehler zu vermeiden.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
