---
id: 663b16e62fee463b4caf46e9
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Nachdem Sie die Elemente der Liste `snake_cased_char_list` verbunden haben, müssen Sie alle führenden oder abschließenden Unterstriche aus dem resultierenden String entfernen. Verwenden Sie dafür die `strip`-Methode mit dem Unterstrichzeichen `_` als Argument.

Methodenaufrufe können verkettet werden, das bedeutet, dass das Ergebnis eines Methodenaufrufs als Objekt für einen weiteren Methodenaufruf verwendet werden kann.

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

Im obigen Beispiel ist die `.upper()`-Methode an `' '.join(words_list)` verkettet, daher wird `.upper()` auf das Ergebnis des `.join()`-Aufrufs angewendet.

Ändern Sie die `return`-Anweisung, indem Sie an `''.join(snake_cased_char_list)` einen Aufruf der `.strip()`-Methode anhängen, um alle führenden oder abschließenden Unterstriche zu entfernen.

# --hints--

Sie sollten die return-Anweisung so ändern, dass die `strip`-Methode mit dem Unterstrichzeichen `_` als Argument enthalten ist.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
