---
id: 587d7fa8367417b2b2512bca
title: 棒グラフの表示を変更する
challengeType: 6
forumTopicId: 301481
dashedName: change-the-presentation-of-a-bar-chart
---

# --description--

前回のチャレンジでは棒グラフを作成しましたが、いくつかのフォーマット変更でより良くできます：

1. 各棒の間にスペースを追加して視覚的に区切ること。これは`bar`クラスのCSSにマージンを追加することで行います。

2. 棒の高さを大きくして値の違いをよりわかりやすくすること。これは値に数値を乗算して高さをスケールすることで行います。

# --instructions--

まず、`margin`タグ内の`2px`クラスに`bar`の`style`を追加してください。次に、`style()`メソッドのコールバック関数を変更して、元のデータ値に`10`を乗算した値（さらに`px`を加えたもの）を返すようにしてください。

**注意：** 各データポイントを同じ定数で乗算するだけではスケールが変わるだけで、ズームインしたようなもので、元のデータの意味は変わりません。

# --hints--

最初の`div`は`height`ピクセルの`120`と`margin`ピクセルの`2`を持つべきです。

```js
const firstDiv = document.querySelector('div');
assert.exists(firstDiv); 
const height = window.getComputedStyle(firstDiv)['height'];
const margin = window.getComputedStyle(firstDiv)['margin-right'];
assert.strictEqual(height, '120px');
assert.strictEqual(margin, '2px');
```

2番目の`div`は`height`ピクセルの`310`と`margin`ピクセルの`2`を持つべきです。

```js
const secondDiv = document.querySelectorAll('div')[1];
assert.exists(secondDiv); 
const height = window.getComputedStyle(secondDiv)['height'];
const margin = window.getComputedStyle(secondDiv)['margin-right'];
assert.strictEqual(height, '310px');
assert.strictEqual(margin, '2px');
```

3番目の`div`は`height`ピクセルの`220`と`margin`ピクセルの`2`を持つべきです。

```js
const thirdDiv = document.querySelectorAll('div')[2];
assert.exists(thirdDiv); 
const height = window.getComputedStyle(thirdDiv)['height'];
const margin = window.getComputedStyle(thirdDiv)['margin-right'];
assert.strictEqual(height, '220px');
assert.strictEqual(margin, '2px');
```

4番目の`div`は`height`ピクセルの`170`と`margin`ピクセルの`2`を持つべきです。

```js
const fourthDiv = document.querySelectorAll('div')[3];
assert.exists(fourthDiv); 
const height = window.getComputedStyle(fourthDiv)['height'];
const margin = window.getComputedStyle(fourthDiv)['margin-right'];
assert.strictEqual(height, '170px');
assert.strictEqual(margin, '2px');
```

5番目の`div`は`height`ピクセルの`250`と`margin`ピクセルの`2`を持つべきです。

```js
const fifthDiv = document.querySelectorAll('div')[4];
assert.exists(fifthDiv); 
const height = window.getComputedStyle(fifthDiv)['height'];
const margin = window.getComputedStyle(fifthDiv)['margin-right'];
assert.strictEqual(height, '250px');
assert.strictEqual(margin, '2px');
```

6番目の`div`は`height`ピクセルの`180`と`margin`ピクセルの`2`を持つべきです。

```js
const sixthDiv = document.querySelectorAll('div')[5];
assert.exists(sixthDiv); 
const height = window.getComputedStyle(sixthDiv)['height'];
const margin = window.getComputedStyle(sixthDiv)['margin-right'];
assert.strictEqual(height, '180px');
assert.strictEqual(margin, '2px');
```

7番目の`div`は`height`ピクセルの`290`と`margin`ピクセルの`2`を持つべきです。

```js
const seventhDiv = document.querySelectorAll('div')[6];
assert.exists(seventhDiv); 
const height = window.getComputedStyle(seventhDiv)['height'];
const margin = window.getComputedStyle(seventhDiv)['margin-right'];
assert.strictEqual(height, '290px');
assert.strictEqual(margin, '2px');
```

8番目の`div`は`height`ピクセルの`140`と`margin`ピクセルの`2`を持つべきです。

```js
const eighthDiv = document.querySelectorAll('div')[7];
assert.exists(eighthDiv); 
const height = window.getComputedStyle(eighthDiv)['height'];
const margin = window.getComputedStyle(eighthDiv)['margin-right'];
assert.strictEqual(height, '140px');
assert.strictEqual(margin, '2px');
```

9番目の`div`は`height`ピクセルの`90`と`margin`ピクセルの`2`を持つべきです。

```js
const ninthDiv = document.querySelectorAll('div')[8];
assert.exists(ninthDiv); 
const height = window.getComputedStyle(ninthDiv)['height'];
const margin = window.getComputedStyle(ninthDiv)['margin-right'];
assert.strictEqual(height, '90px');
assert.strictEqual(margin, '2px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    /* Add your code below this line */

    
    /* Add your code above this line */
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d + 'px'); // Change this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    margin: 2px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d * 10 + 'px');
  </script>
</body>
```
