---
id: 587d7fac367417b2b2512bdd
title: 動的スケールの使用
challengeType: 6
forumTopicId: 301495
dashedName: use-dynamic-scales
---

# --description--

D3の`min()`と`max()`メソッドは、スケールを設定するのに役立ちます。

複雑なデータセットがある場合、優先すべきことの1つは、可視化がSVGコンテナの幅と高さに収まるようにスケールを設定することです。すべてのデータがSVG内にプロットされて、ウェブページ上で見えるようにしたいです。

以下の例は、散布図データのx軸スケールを設定しています。`domain()`メソッドは、プロットの生データ値に関する情報をスケールに渡します。`range()`メソッドは、可視化のためのウェブページ上の実際のスペースに関する情報を与えます。

例では、ドメインは0からセット内の最大値までです。配列のx値に基づくコールバック関数を使って`max()`メソッドを使用しています。レンジはSVGの幅（`w`）を使いますが、パディングも含まれています。これにより、散布図の点とSVGの端の間にスペースができます。

```js
const dataset = [
  [34, 78],
  [109, 280],
  [310, 120],
  [79, 411],
  [420, 220],
  [233, 145],
  [333, 96],
  [222, 333],
  [78, 320],
  [21, 123]
];
const w = 500;
const h = 500;

const padding = 30;
const xScale = d3
  .scaleLinear()
  .domain([0, d3.max(dataset, d => d[0])])
  .range([padding, w - padding]);
```

パディングは最初はわかりにくいかもしれません。x軸を0から500（SVGの幅の値）までの水平線と考えてください。`range()`メソッドにパディングを含めることで、プロットはその線上の0ではなく30から始まり、500ではなく470で終わるようになります。

# --instructions--

`yScale`変数を使って線形のy軸スケールを作成してください。ドメインは0からセット内の最大`y`値までにします。レンジはSVGの高さ（`h`）を使い、パディングも含めてください。

**注意:** プロットが上下逆さまにならないようにしてください。y座標のレンジを設定するときは、高い値（高さからパディングを引いた値）が最初の引数で、低い値が2番目の引数です。

# --hints--

`h2`内のテキストは`30`であるべきです。

```js
assert.strictEqual(output, 30);
assert.strictEqual(document.querySelector('h2')?.textContent, '30');
```

yScaleの`domain()`は`[0, 411]`と同じであるべきです。

```js
assert.deepEqual(yScale.domain(), [0, 411]);
```

yScaleの`range()`は`[470, 30]`と同じであるべきです。

```js
assert.deepEqual(yScale.range(), [470, 30]);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    // Padding between the SVG boundary and the plot
    const padding = 30;

    // Create an x and y scale

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    // Add your code below this line

    const yScale = undefined;


    // Add your code above this line

    const output = yScale(411); // Returns 30
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const padding = 30;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const output = yScale(411);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
