---
id: 587d7fa7367417b2b2512bc7
title: データに基づいてスタイルを変更する
challengeType: 6
forumTopicId: 301479
dashedName: change-styles-based-on-data
---

# --description--

D3はデータの可視化と表示に関するものです。データに基づいて要素のスタイルを変更したくなることがよくあります。
例えば、値が20未満のデータポイントを青色にし、それ以外は赤色にしたい場合があります。`style()`メソッドでコールバック関数を使い、その中に条件ロジックを含めることができます。コールバック関数は`d`パラメータを使ってデータポイントを表します。

```js
selection.style('color', d => {});
```

`style()`メソッドは`color`の設定に限定されず、他のCSSプロパティにも使えます。

# --instructions--

エディター内のコードに`style()`メソッドを追加して、`color`要素の`h2`を条件付きで設定してください。データの値が20未満なら赤を返し、それ以外は緑を返すようにコールバック関数を書いてください。

**注意:** if-elseロジックや三項演算子を使うことができます。

# --hints--

最初の`h2`は`color`が赤であるべきです。

```js
const headingTwoFirst = document.querySelector('h2');
assert.exists(headingTwoFirst); 
const color = window.getComputedStyle(headingTwoFirst)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

2番目の`h2`は`color`が緑であるべきです。

```js
const headingTwoSecond = document.querySelectorAll('h2')[1];
assert.exists(headingTwoSecond); 
const color = window.getComputedStyle(headingTwoSecond)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

3番目の`h2`は`color`が緑であるべきです。

```js
const headingTwoThird = document.querySelectorAll('h2')[2];
assert.exists(headingTwoThird); 
const color = window.getComputedStyle(headingTwoThird)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

4番目の`h2`は`color`が赤であるべきです。

```js
const headingTwoFourth = document.querySelectorAll('h2')[3];
assert.exists(headingTwoFourth); 
const color = window.getComputedStyle(headingTwoFourth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

5番目の`h2`は`color`が緑であるべきです。

```js
const headingTwoFifth = document.querySelectorAll('h2')[4];
assert.exists(headingTwoFifth); 
const color = window.getComputedStyle(headingTwoFifth)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

6番目の`h2`は`color`が赤であるべきです。

```js
const headingTwoSixth = document.querySelectorAll('h2')[5];
assert.exists(headingTwoSixth); 
const color = window.getComputedStyle(headingTwoSixth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

7番目の`h2`は`color`が緑であるべきです。

```js
const headingTwoSeventh = document.querySelectorAll('h2')[6];
assert.exists(headingTwoSeventh); 
const color = window.getComputedStyle(headingTwoSeventh)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

8番目の`h2`は`color`が赤であるべきです。

```js
const headingTwoEighth = document.querySelectorAll('h2')[7];
assert.exists(headingTwoEighth); 
const color = window.getComputedStyle(headingTwoEighth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

9番目の`h2`は`color`が赤であるべきです。

```js
const headingTwoNinth = document.querySelectorAll('h2')[8];
assert.exists(headingTwoNinth); 
const color = window.getComputedStyle(headingTwoNinth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('color', d => (d < 20 ? 'red' : 'green'));
  </script>
</body>
```
