---
id: 587d7fad367417b2b2512bdf
title: ビジュアライゼーションに軸を追加する
challengeType: 6
forumTopicId: 301472
dashedName: add-axes-to-a-visualization
---

# --description--

散布図を改善するもう一つの方法は、x軸とy軸を追加することです。

D3には、y軸とx軸をそれぞれレンダリングするための`axisLeft()`と`axisBottom()`という2つのメソッドがあります。前のチャレンジの`xScale`に基づいてx軸を作成する例を示します。

```js
const xAxis = d3.axisBottom(xScale);
```

次のステップは、SVG上に軸をレンダリングすることです。そのために、一般的なSVGコンポーネントである`g`要素を使うことができます。`g`はグループを表します。`rect`、`circle`、`text`とは異なり、軸はレンダリングされると単なる直線です。単純な形状なので、`g`を使うのが適しています。最後のステップは、軸をSVGの正しい位置に配置するために`transform`属性を適用することです。そうしないと、線はSVGのボーダーに沿ってレンダリングされて見えなくなります。SVGはさまざまな種類の`transforms`をサポートしていますが、軸の位置決めには`translate`が必要です。これが`g`要素に適用されると、指定された量だけグループ全体が右と下に移動します。例を示します。

```js
const xAxis = d3.axisBottom(xScale);

svg
  .append('g')
  .attr('transform', 'translate(0, ' + (h - padding) + ')')
  .call(xAxis);
```

上記のコードはx軸をSVGの下部に配置します。それから`call()`メソッドの引数として渡されます。y軸も同じ方法で動作しますが、`translate`引数は`(x, 0)`の形式です。上記の`translate`メソッドの`attr()`が文字列なので、引数に変数の値を含めるために連結を使うことができます。

# --instructions--

散布図にはx軸が追加されました。`yAxis`メソッドを使って、`axisLeft()`という変数にy軸を作成してください。次に、`g`要素を使って軸をレンダリングしてください。軸を右にパディング単位分、下に`transform`単位分移動させるために`0`属性を使うことを忘れないでください。軸を`call()`することも忘れずに。

# --hints--

コードは`axisLeft()`メソッドを使い、引数に`yScale`を渡す必要があります。

```js
assert.match(code, /\.axisLeft\(yScale\)/g);
```

y軸の`g`要素には、`transform`だけ軸を移動させるための`(60, 0)`属性が必要です。

```js
const yAxis = document.querySelectorAll('g')[10];
assert.match(yAxis?.getAttribute('transform'), /translate\(60\s*,\s*0\)/g);
```

コードは`yAxis`を呼び出す必要があります。

```js
assert.match(code, /\.call\(\s*yAxis\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);
    // Add your code below this line
    const yAxis = undefined;
    // Add your code above this line

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);

    const yAxis = d3.axisLeft(yScale);

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    svg
      .append('g')
      .attr('transform', 'translate(' + padding + ',0)')
      .call(yAxis);
  </script>
</body>
```
