---
id: 587d7fab367417b2b2512bd7
title: SVGの円で散布図を作成する
challengeType: 6
forumTopicId: 301484
dashedName: create-a-scatterplot-with-svg-circles
---

# --description--

散布図は別の種類のチャートです。通常、円を使ってデータポイントをマッピングし、それぞれに2つの値があります。これらの値は`x`軸と`y`軸に対応し、円の位置を決めるために使われます。

SVGには円の形を作る`circle`タグがあります。これは棒グラフで使った`rect`要素とよく似ています。

# --instructions--

`data()`、`enter()`、`append()`メソッドを使って、`dataset`を新しい`circle`要素にバインドし、それらをSVGに追加してください。

**注意:** 円はまだ属性を設定していないので見えません。属性の設定は次のチャレンジで行います。

# --hints--

コードには10個の`circle`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
  </script>
</body>
```
