---
id: 587d7fa7367417b2b2512bc8
title: D3でクラスを追加する
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--

HTML要素に多くのインラインスタイルを使うと、小規模なアプリでも管理が難しくなります。要素にクラスを追加して、そのクラスに対してCSSルールで一度だけスタイルを設定するほうが簡単です。D3には、クラス名を含む任意のHTML属性を要素に追加できる`attr()`メソッドがあります。

`attr()`メソッドは`style()`と同じように動作します。カンマ区切りの値を受け取り、コールバック関数を使うこともできます。以下は、選択した要素に`container`というクラスを追加する例です。

```js
selection.attr('class', 'container');
```

`class`パラメータはクラスを追加するたびに同じままで、`container`パラメータだけが変わることに注意してください。

# --instructions--

エディターのコードに`attr()`メソッドを追加し、`bar`要素に`div`というクラスを付けてください。

# --hints--

`div`要素には`bar`というクラスが付いているはずです。

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

コードは`attr()`メソッドを使う必要があります。

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body>
```
