---
id: 587d7fa7367417b2b2512bc6
title: 要素にインラインスタイルを追加する
challengeType: 6
forumTopicId: 301475
dashedName: add-inline-styling-to-elements
---

# --description--

D3では、`style()` メソッドを使って動的な要素にインラインCSSスタイルを追加できます。

`style()` メソッドは、カンマ区切りのキーと値のペアを引数として受け取ります。以下は選択した要素の文字色を青に設定する例です。

```js
selection.style('color', 'blue');
```

# --instructions--

エディター内のコードに`style()` メソッドを追加して、表示されているすべてのテキストの`font-family`を`verdana`にしてください。

# --hints--

`h2` 要素は`font-family` が`verdana`である必要があります。

```js
const headingTwo = document.querySelector('h2');
assert.exists(headingTwo); 
const fontFamily = window.getComputedStyle(headingTwo)['font-family'];
assert.strictEqual(fontFamily, 'verdana');
```

コードは`style()` メソッドを使うべきです。

```js
assert.match(code, /\.style/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('font-family', 'verdana');
  </script>
</body>
```
