---
id: 587d7fab367417b2b2512bd8
title: 円要素に属性を追加する
challengeType: 6
forumTopicId: 301471
dashedName: add-attributes-to-the-circle-elements
---

# --description--

前のチャレンジでは、`circle` 要素を `dataset` の各点に対して作成し、SVGに追加しました。しかし、D3が正しく表示するには、各 `circle` の位置とサイズに関する詳細な情報が必要です。

SVGの `circle` には主に3つの属性があります。`cx` と `cy` の属性は座標で、形の中心をSVG上のどこに配置するかをD3に伝えます。半径（`r` 属性）は `circle` の大きさを示します。

`rect` の `y` 座標と同様に、`cy` の `circle` 属性はSVGの上端からの距離で測られ、下端からではありません。

これら3つの属性はすべて、値を動的に設定するためにコールバック関数を使うことができます。`data(dataset)` の後にチェーンされたすべてのメソッドは `dataset` の各アイテムに対して1回ずつ実行されます。コールバック関数の `d` パラメータは `dataset` の現在のアイテムを指し、これは各点の配列です。ブラケット記法（`d[0]` のように）を使ってその配列の値にアクセスします。

# --instructions--

`cx` 要素に `cy`、`r`、`circle` の属性を追加してください。`cx` の値は `dataset` の各アイテムの配列の最初の数字にしてください。`cy` の値は配列の2番目の数字に基づきますが、チャートが上下逆にならないようにしてください。`r` の値はすべての円に対して `5` にしてください。

# --hints--

コードには10個の `circle` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

最初の `circle` 要素は `cx` が `34`、`cy` が `422`、`r` が `5` であるべきです。

```js
const circle1 = document.querySelector('circle');
assert.strictEqual(circle1?.getAttribute('cx'), '34');
assert.strictEqual(circle1?.getAttribute('cy'), '422');
assert.strictEqual(circle1?.getAttribute('r'), '5');
```

2番目の `circle` 要素は `cx` が `109`、`cy` が `220`、`r` が `5` であるべきです。

```js
const circle2 = document.querySelectorAll('circle')[1];
assert.strictEqual(circle2?.getAttribute('cx'), '109');
assert.strictEqual(circle2?.getAttribute('cy'), '220');
assert.strictEqual(circle2?.getAttribute('r'), '5');
```

3番目の `circle` 要素は `cx` が `310`、`cy` が `380`、`r` が `5` であるべきです。

```js
const circle3 = document.querySelectorAll('circle')[2];
assert.strictEqual(circle3?.getAttribute('cx'), '310');
assert.strictEqual(circle3?.getAttribute('cy'), '380');
assert.strictEqual(circle3?.getAttribute('r'), '5');
```

4番目の `circle` 要素は `cx` が `79`、`cy` が `89`、`r` が `5` であるべきです。

```js
const circle4 = document.querySelectorAll('circle')[3];
assert.strictEqual(circle4?.getAttribute('cx'), '79');
assert.strictEqual(circle4?.getAttribute('cy'), '89');
assert.strictEqual(circle4?.getAttribute('r'), '5');
```

5番目の `circle` 要素は `cx` が `420`、`cy` が `280`、`r` が `5` であるべきです。

```js
const circle5 = document.querySelectorAll('circle')[4];
assert.strictEqual(circle5?.getAttribute('cx'), '420');
assert.strictEqual(circle5?.getAttribute('cy'), '280');
assert.strictEqual(circle5?.getAttribute('r'), '5');
```

6番目の `circle` 要素は `cx` が `233`、`cy` が `355`、`r` が `5` であるべきです。

```js
const circle6 = document.querySelectorAll('circle')[5];
assert.strictEqual(circle6?.getAttribute('cx'), '233');
assert.strictEqual(circle6?.getAttribute('cy'), '355');
assert.strictEqual(circle6?.getAttribute('r'), '5');
```

7番目の `circle` 要素は `cx` が `333`、`cy` が `404`、`r` が `5` であるべきです。

```js
const circle7 = document.querySelectorAll('circle')[6];
assert.strictEqual(circle7?.getAttribute('cx'), '333');
assert.strictEqual(circle7?.getAttribute('cy'), '404');
assert.strictEqual(circle7?.getAttribute('r'), '5');
```

8番目の `circle` 要素は `cx` が `222`、`cy` が `167`、`r` が `5` であるべきです。

```js
const circle8 = document.querySelectorAll('circle')[7];
assert.strictEqual(circle8?.getAttribute('cx'), '222');
assert.strictEqual(circle8?.getAttribute('cy'), '167');
assert.strictEqual(circle8?.getAttribute('r'), '5');
```

9番目の `circle` 要素は `cx` が `78`、`cy` が `180`、`r` が `5` であるべきです。

```js
const circle9 = document.querySelectorAll('circle')[8];
assert.strictEqual(circle9?.getAttribute('cx'), '78');
assert.strictEqual(circle9?.getAttribute('cy'), '180');
assert.strictEqual(circle9?.getAttribute('r'), '5');
```

10番目の `circle` 要素は `cx` が `21`、`cy` が `377`、`r` が `5` であるべきです。

```js
const circle10 = document.querySelectorAll('circle')[9];
assert.strictEqual(circle10?.getAttribute('cx'), '21');
assert.strictEqual(circle10?.getAttribute('cy'), '377');
assert.strictEqual(circle10?.getAttribute('r'), '5');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => d[0])
      .attr('cy', d => h - d[1])
      .attr('r', 5);
  </script>
</body>
```
