---
id: 587d7fa9367417b2b2512bce
title: 各棒の座標を動的に設定する
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--

前回のチャレンジでは、`svg` の各ポイントに対して `dataset` 要素に長方形を作成して追加し、棒を表現しました。しかし、すべての棒が重なって積み重なってしまいました。

長方形の配置は `x` と `y` 属性で制御されます。これらは D3 に対して `svg` エリアのどこから図形を描き始めるかを指示します。前回のチャレンジでは両方とも 0 に設定されていたため、すべての棒が左上隅に配置されていました。

棒グラフでは、すべての棒が同じ垂直レベルに並ぶ必要があり、つまり `y` の値はすべての棒で同じ（0のまま）です。一方で `x` の値は新しい棒を追加するごとに変わる必要があります。`x` の値が大きいほど右方向に押し出されることを覚えておいてください。`dataset` の配列要素を順に処理する際、`x` の値は増加していくべきです。

D3 の `attr()` メソッドは、属性を動的に設定するためのコールバック関数を受け入れます。コールバック関数は2つの引数を取り、1つはデータポイント自体（通常は `d`）、もう1つは配列内のデータポイントのインデックスです。インデックスの引数は省略可能です。形式は以下の通りです。

```js
selection.attr('property', (d, i) => {});
```

重要なのは、データセットのアイテムを繰り返すために `for` ループを書いたり `forEach()` を使ったりする必要はないことです。`data()` メソッドがデータセットをパースし、`data()` の後にチェーンされたメソッドはデータセットの各アイテムに対して1回ずつ実行されます。

# --instructions--

`x` 属性のコールバック関数を変更して、インデックスに30を掛けた値を返すようにしてください。

**注意:** 各棒の幅は25なので、`x` の値を30ずつ増やすことで棒の間に少しスペースができます。この例では25より大きい値ならどれでも機能します。

# --hints--

最初の `rect` は `x` の値が `0` であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

2番目の `rect` は `x` の値が `30` であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

3番目の `rect` は `x` の値が `60` であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

4番目の `rect` は `x` の値が `90` であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

5番目の `rect` は `x` の値が `120` であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

6番目の `rect` は `x` の値が `150` であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

7番目の `rect` は `x` の値が `180` であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

8番目の `rect` は `x` の値が `210` であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

9番目の `rect` は `x` の値が `240` であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
