---
id: 587d7fab367417b2b2512bd9
title: 散布図の円にラベルを追加する
challengeType: 6
forumTopicId: 301477
dashedName: add-labels-to-scatter-plot-circles
---

# --description--

散布図の点にラベルを作成するためにテキストを追加できます。

目標は、`x`の各アイテムの最初のフィールド（`y`）と2番目のフィールド（`dataset`）のカンマ区切りの値を表示することです。

`text`ノードには、SVG上の位置を決めるために`x`と`y`属性が必要です。このチャレンジでは、高さを決める`y`の値は、`circle`が`cy`属性に使う値と同じ値を使えます。`x`の値は、`cx`の`circle`の値より少し大きくして、ラベルが見えるようにします。これにより、ラベルはプロットされた点の右側に配置されます。

# --instructions--

散布図の各点に`text`要素を使ってラベルを付けてください。ラベルのテキストは、カンマとスペースで区切られた2つの値にしてください。例えば、最初の点のラベルは`34, 78`です。`x`属性は、`5`の`cx`属性に使った値より`circle`単位だけ大きく設定してください。`y`属性は、`cy`の`circle`値と同じように設定してください。

# --hints--

コードには10個の`text`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

最初のラベルはテキストが`34, 78`、`x`の値が`39`、`y`の値が`422`です。

```js
const labelOne = document.querySelector('text');
assert.strictEqual(labelOne?.textContent, '34, 78');
assert.strictEqual(labelOne?.getAttribute('x'), '39');
assert.strictEqual(labelOne?.getAttribute('y'), '422');
```

2番目のラベルはテキストが`109, 280`、`x`の値が`114`、`y`の値が`220`です。

```js
const labelTwo = document.querySelectorAll('text')[1];
assert.strictEqual(labelTwo?.textContent, '109, 280');
assert.strictEqual(labelTwo?.getAttribute('x'), '114');
assert.strictEqual(labelTwo?.getAttribute('y'), '220');
```

3番目のラベルはテキストが`310, 120`、`x`の値が`315`、`y`の値が`380`です。

```js
const labelThree = document.querySelectorAll('text')[2];
assert.strictEqual(labelThree?.textContent, '310, 120');
assert.strictEqual(labelThree?.getAttribute('x'), '315');
assert.strictEqual(labelThree?.getAttribute('y'), '380');
```

4番目のラベルはテキストが`79, 411`、`x`の値が`84`、`y`の値が`89`です。

```js
const labelFour = document.querySelectorAll('text')[3];
assert.strictEqual(labelFour?.textContent, '79, 411');
assert.strictEqual(labelFour?.getAttribute('x'), '84');
assert.strictEqual(labelFour?.getAttribute('y'), '89');
```

5番目のラベルはテキストが`420, 220`、`x`の値が`425`、`y`の値が`280`です。

```js
const labelFive = document.querySelectorAll('text')[4];
assert.strictEqual(labelFive?.textContent, '420, 220');
assert.strictEqual(labelFive?.getAttribute('x'), '425');
assert.strictEqual(labelFive?.getAttribute('y'), '280');
```

6番目のラベルはテキストが`233, 145`、`x`の値が`238`、`y`の値が`355`です。

```js
const labelSix = document.querySelectorAll('text')[5];
assert.strictEqual(labelSix?.textContent, '233, 145');
assert.strictEqual(labelSix?.getAttribute('x'), '238');
assert.strictEqual(labelSix?.getAttribute('y'), '355');
```

7番目のラベルはテキストが`333, 96`、`x`の値が`338`、`y`の値が`404`です。

```js
const labelSeven = document.querySelectorAll('text')[6];
assert.strictEqual(labelSeven?.textContent, '333, 96');
assert.strictEqual(labelSeven?.getAttribute('x'), '338');
assert.strictEqual(labelSeven?.getAttribute('y'), '404');
```

8番目のラベルはテキストが`222, 333`、`x`の値が`227`、`y`の値が`167`です。

```js
const labelEight = document.querySelectorAll('text')[7];
assert.strictEqual(labelEight?.textContent, '222, 333');
assert.strictEqual(labelEight?.getAttribute('x'), '227');
assert.strictEqual(labelEight?.getAttribute('y'), '167');
```

9番目のラベルはテキストが`78, 320`、`x`の値が`83`、`y`の値が`180`です。

```js
const labelNine = document.querySelectorAll('text')[8];
assert.strictEqual(labelNine?.textContent, '78, 320');
assert.strictEqual(labelNine?.getAttribute('x'), '83');
assert.strictEqual(labelNine?.getAttribute('y'), '180');
```

10番目のラベルはテキストが`21, 123`、`x`の値が`26`、`y`の値が`377`です。

```js
const labelTen = document.querySelectorAll('text')[9];
assert.strictEqual(labelTen?.textContent, '21, 123');
assert.strictEqual(labelTen?.getAttribute('x'), '26');
assert.strictEqual(labelTen?.getAttribute('y'), '377');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg.selectAll('text').data(dataset).enter().append('text');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', d => d[0] + 5)
      .attr('y', d => h - d[1])
      .text(d => d[0] + ', ' + d[1]);
  </script>
</body>
```
