---
id: 587d7fa8367417b2b2512bc9
title: 要素の高さを動的に更新する
challengeType: 6
forumTopicId: 301493
dashedName: update-the-height-of-an-element-dynamically
---

# --description--

前のチャレンジでは、配列からデータを表示する方法とCSSクラスを追加する方法を学びました。これらのレッスンを組み合わせて、シンプルな棒グラフを作成できます。手順は2つあります：

1. 配列の各データポイントに対して`div`を作成する

2. `div`メソッドのコールバック関数を使って、各`style()`に動的な高さを設定し、高さをデータ値に等しくする

コールバック関数を使ってスタイルを設定する形式を思い出してください：

```js
selection.style('cssProperty', d => d);
```

# --instructions--

エディター内のコードに`style()`メソッドを追加して、各要素の`height`プロパティを設定してください。コールバック関数を使って、データポイントの値に文字列`px`を追加して返してください。

# --hints--

最初の`div`は`height`ピクセルの`12`を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('div')[0]?.style.height, '12px');
```

2番目の`div`は`height`ピクセルの`31`を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('div')[1]?.style.height, '31px');
```

3番目の`div`は`height`ピクセルの`22`を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('div')[2]?.style.height, '22px');
```

4番目の`div`は`height`ピクセルの`17`を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('div')[3]?.style.height, '17px');
```

5番目の`div`は`height`ピクセルの`25`を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('div')[4]?.style.height, '25px');
```

6番目の`div`は`height`ピクセルの`18`を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('div')[5]?.style.height, '18px');
```

7番目の`div`は`height`ピクセルの`29`を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('div')[6]?.style.height, '29px');
```

8番目の`div`は`height`ピクセルの`14`を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('div')[7]?.style.height, '14px');
```

9番目の`div`は`height`ピクセルの`9`を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('div')[8]?.style.height, '9px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => `${d}px`);
  </script>
</body>
```
