---
id: 587d7fa9367417b2b2512bcf
title: 各バーの高さを動的に変更する
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--

各バーの高さは、配列内のデータポイントの値に設定できます。これは、`x` の値を動的に設定した方法と似ています。

```js
selection.attr('property', (d, i) => {});
```

ここで `d` はデータポイントの値、`i` は配列内のデータポイントのインデックスになります。

# --instructions--

`height` 属性のコールバック関数を変更して、データの値に3を乗算した値を返すようにしてください。

**注意:** すべてのデータポイントを同じ定数で乗算するとデータがスケーリングされます（ズームインのような効果です）。この例ではバーの値の違いを見やすくするのに役立ちます。

# --hints--

最初の `rect` は `height` の `36` を持つべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

2番目の `rect` は `height` の `93` を持つべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

3番目の `rect` は `height` の `66` を持つべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

4番目の `rect` は `height` の `51` を持つべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

5番目の `rect` は `height` の `75` を持つべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

6番目の `rect` は `height` の `54` を持つべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

7番目の `rect` は `height` の `87` を持つべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

8番目の `rect` は `height` の `42` を持つべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

9番目の `rect` は `height` の `27` を持つべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body>
```
