---
id: 587d7faa367417b2b2512bd2
title: D3要素にラベルを追加する
challengeType: 6
forumTopicId: 301476
dashedName: add-labels-to-d3-elements
---

# --description--

D3では、棒グラフなどのグラフ要素にSVGの`text`要素を使ってラベルを付けることができます。

`rect`要素と同様に、`text`要素にはSVG上に配置するための`x`属性と`y`属性が必要です。また、表示する値のデータにアクセスする必要もあります。

D3は棒グラフにラベルを付ける方法を細かく制御できます。

# --instructions--

エディター内のコードはすでに各新しい`text`要素にデータをバインドしています。まず、`text`に`svg`ノードを追加してください。次に、`x`と`y`の座標属性を追加します。これらは`rect`の座標と同じ方法で計算しますが、`y`の`text`値はラベルが棒グラフより3単位高くなるように設定してください。最後に、D3の`text()`メソッドを使ってラベルをデータポイントの値に設定してください。

**注意:** ラベルを棒グラフより高く配置するには、`y`の`text`値が棒グラフの`y`値より3大きいか3小さいかを決めてください。

# --hints--

最初の`text`要素はラベルが`12`で、`y`の値が`61`であるべきです。

```js
assert.strictEqual(document.querySelectorAll('text')[0]?.textContent, '12');
assert.strictEqual(
  document.querySelectorAll('text')[0]?.getAttribute('y'),
  '61'
);
```

2番目の`text`要素はラベルが`31`で、`y`の値が`4`であるべきです。

```js
assert.strictEqual(document.querySelectorAll('text')[1]?.textContent, '31');
assert.strictEqual(
  document.querySelectorAll('text')[1]?.getAttribute('y'),
  '4'
);
```

3番目の`text`要素はラベルが`22`で、`y`の値が`31`であるべきです。

```js
assert.strictEqual(document.querySelectorAll('text')[2]?.textContent, '22');
assert.strictEqual(
  document.querySelectorAll('text')[2]?.getAttribute('y'),
  '31'
);
```

4番目の`text`要素はラベルが`17`で、`y`の値が`46`であるべきです。

```js
assert.strictEqual(document.querySelectorAll('text')[3]?.textContent, '17');
assert.strictEqual(
  document.querySelectorAll('text')[3]?.getAttribute('y'),
  '46'
);
```

5番目の`text`要素はラベルが`25`で、`y`の値が`22`であるべきです。

```js
assert.strictEqual(document.querySelectorAll('text')[4]?.textContent, '25');
assert.strictEqual(
  document.querySelectorAll('text')[4]?.getAttribute('y'),
  '22'
);
```

6番目の`text`要素はラベルが`18`で、`y`の値が`43`であるべきです。

```js
assert.strictEqual(document.querySelectorAll('text')[5]?.textContent, '18');
assert.strictEqual(
  document.querySelectorAll('text')[5]?.getAttribute('y'),
  '43'
);
```

7番目の`text`要素はラベルが`29`で、`y`の値が`10`であるべきです。

```js
assert.strictEqual(document.querySelectorAll('text')[6]?.textContent, '29');
assert.strictEqual(
  document.querySelectorAll('text')[6]?.getAttribute('y'),
  '10'
);
```

8番目の`text`要素はラベルが`14`で、`y`の値が`55`であるべきです。

```js
assert.strictEqual(document.querySelectorAll('text')[7]?.textContent, '14');
assert.strictEqual(
  document.querySelectorAll('text')[7]?.getAttribute('y'),
  '55'
);
```

9番目の`text`要素はラベルが`9`で、`y`の値が`70`であるべきです。

```js
assert.strictEqual(document.querySelectorAll('text')[8]?.textContent, '9');
assert.strictEqual(
  document.querySelectorAll('text')[8]?.getAttribute('y'),
  '70'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg.selectAll('text').data(dataset).enter();
    // Add your code below this line




    // Add your code above this line
  </script>
  <body></body>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3)
      .text(d => d);
  </script>
  <body></body>
</body>
```
