---
id: 587d7faa367417b2b2512bd6
title: D3要素にツールチップを追加する
challengeType: 6
forumTopicId: 301470
dashedName: add-a-tooltip-to-a-d3-element
---

# --description--

ツールチップは、ユーザーがページ上のアイテムにカーソルを合わせたときに、そのアイテムに関する詳細情報を表示します。ビジュアライゼーションにツールチップを追加する方法はいくつかあります。このチャレンジではSVGの`title`要素を使います。

`title`は`text()`メソッドと組み合わせて、バーに動的にデータを追加します。

# --instructions--

各`title`ノードの下に`rect`要素を追加してください。その後、コールバック関数を使って`text()`メソッドを呼び出し、テキストにデータの値が表示されるようにしてください。

# --hints--

コードには9つの`title`要素が必要です。

```js
assert.lengthOf(document.querySelectorAll('title'), 9);
```

最初の`title`要素には`12`のツールチップテキストが必要です。

```js
assert.strictEqual(document.querySelectorAll('title')[0]?.textContent, '12');
```

2番目の`title`要素には`31`のツールチップテキストが必要です。

```js
assert.strictEqual(document.querySelectorAll('title')[1]?.textContent, '31');
```

3番目の`title`要素には`22`のツールチップテキストが必要です。

```js
assert.strictEqual(document.querySelectorAll('title')[2]?.textContent, '22');
```

4番目の`title`要素には`17`のツールチップテキストが必要です。

```js
assert.strictEqual(document.querySelectorAll('title')[3]?.textContent, '17');
```

5番目の`title`要素には`25`のツールチップテキストが必要です。

```js
assert.strictEqual(document.querySelectorAll('title')[4]?.textContent, '25');
```

6番目の`title`要素には`18`のツールチップテキストが必要です。

```js
assert.strictEqual(document.querySelectorAll('title')[5]?.textContent, '18');
```

7番目の`title`要素には`29`のツールチップテキストが必要です。

```js
assert.strictEqual(document.querySelectorAll('title')[6]?.textContent, '29');
```

8番目の`title`要素には`14`のツールチップテキストが必要です。

```js
assert.strictEqual(document.querySelectorAll('title')[7]?.textContent, '14');
```

9番目の`title`要素には`9`のツールチップテキストが必要です。

```js
assert.strictEqual(document.querySelectorAll('title')[8]?.textContent, '9');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar')
      .append('title')
      .text(d => d);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```
