---
id: 587d7fa6367417b2b2512bc3
title: D3で要素のグループを選択する
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--

D3には要素のグループを選択するための`selectAll()`メソッドもあります。これは、入力文字列に一致するドキュメント内のすべての項目のHTMLノードの配列を返します。以下は、ドキュメント内のすべてのアンカータグを選択する例です。

```js
const anchors = d3.selectAll('a');
```

`select()`メソッドと同様に、`selectAll()`はメソッドチェーンをサポートしており、他のメソッドと組み合わせて使うことができます。

# --instructions--

ドキュメント内のすべての`li`タグを選択し、`list item`メソッドをチェーンしてテキストを文字列`.text()`に変更してください。

# --hints--

ページには3つの`li`要素があり、それぞれのテキストは正確に`list item`と表示されている必要があります。大文字・小文字やスペースも正確に一致させてください。

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

コードは`d3`オブジェクトにアクセスしている必要があります。

```js
assert.match(code, /d3/g);
```

コードは`selectAll`メソッドを使用している必要があります。

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```
