---
id: 587d7fa8367417b2b2512bcd
title: セット内の各データポイントにバーを作成する
challengeType: 6
forumTopicId: 301482
dashedName: create-a-bar-for-each-data-point-in-the-set
---

# --description--

前回のチャレンジでは、バーを表すために`svg`要素に1つの長方形だけを追加しました。ここでは、これまでに学んだ`data()`、`enter()`、およびSVGシェイプに関する知識を組み合わせて、`dataset`内の各データポイントに対して長方形を作成し、追加します。

以前のチャレンジでは、`div`内の各アイテムに対して`dataset`を作成し追加する方法のフォーマットを示しました：

```js
d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
```

`rect`要素を扱う場合は、`div`要素の場合といくつか違いがあります。`rect`要素は`svg`要素に追加する必要があり、直接`body`に追加してはいけません。また、D3に対して`rect`エリア内の各`svg`の配置場所を指定する必要があります。バーの配置については次のチャレンジで扱います。

# --instructions--

`data()`、`enter()`、および`append()`メソッドを使って、`rect`内の各アイテムに対して`dataset`を作成し追加してください。バーはすべて重なって表示されますが、これは次のチャレンジで修正します。

# --hints--

ドキュメントには9つの`rect`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('rect'), 9);
```

コードでは`data()`メソッドを使う必要があります。

```js
assert.match(code, /\.data/g);
```

コードでは`enter()`メソッドを使う必要があります。

```js
assert.match(code, /\.enter/g);
```

コードでは`append()`メソッドを使う必要があります。

```js
assert.match(code, /\.append/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      // Add your code below this line

      // Add your code above this line
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
