---
id: 587d7fa9367417b2b2512bd0
title: SVG要素の反転
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--

棒グラフが上下逆さま、つまり反転しているように見えたかもしれません。これはSVGが(x, y)座標を使う方法によるものです。

SVGでは、座標の原点は左上隅にあります。`x`の座標が0だと、図形はSVGエリアの左端に配置されます。`y`の座標が0だと、図形はSVGエリアの上端に配置されます。`x`の値が大きくなると長方形は右に移動します。`y`の値が大きくなると長方形は下に移動します。

棒グラフを正しい向きにするには、`y`の座標の計算方法を変える必要があります。棒の高さとSVGエリア全体の高さの両方を考慮しなければなりません。

SVGエリアの高さは100です。データセットに0の値がある場合、棒はSVGエリアの下端から始まる必要があります（上端ではありません）。これを実現するには、`y`の座標に100の値が必要です。もしデータポイントの値が1なら、棒を下端に設定するために`y`の座標を100にします。次に棒の高さ1を考慮するので、最終的な`y`の座標は99になります。

`y`の座標、つまり`y = heightOfSVG - heightOfBar`が棒を正しい向きに配置します。

# --instructions--

`y`属性のコールバック関数を変更して、棒を正しい向きにしてください。棒の`height`はデータ値`d`の3倍であることを覚えておいてください。

**注意:** 一般的に、関係は`y = h - m * d`で、`m`はデータポイントをスケールする定数です。

# --hints--

最初の`rect`は`y`の値が`64`であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

2番目の`rect`は`y`の値が`7`であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

3番目の`rect`は`y`の値が`34`であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

4番目の`rect`は`y`の値が`49`であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

5番目の`rect`は`y`の値が`25`であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

6番目の`rect`は`y`の値が`46`であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

7番目の`rect`は`y`の値が`13`であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

8番目の`rect`は`y`の値が`58`であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

9番目の`rect`は`y`の値が`73`であるべきです。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```
