---
id: 587d7fac367417b2b2512bdc
title: >-
  データセット内の最小値と最大値を見つけるために d3.max と d3.min 関数を使う
challengeType: 6
forumTopicId: 301496
dashedName: >-
  use-the-d3-max-and-d3-min-functions-to-find-minimum-and-maximum-values-in-a-dataset
---

# --description--

D3 のメソッド `domain()` と `range()` は、データに基づいてスケールの情報を設定します。その作業を簡単にするいくつかのメソッドがあります。

ドメインを設定するとき、データセット内の最小値と最大値を使いたいことがよくあります。特に大きなデータセットでこれらの値を手動で探そうとすると、エラーが起こる場合があります。

D3 にはこの情報を返す `min()` と `max()` という2つのメソッドがあります。例を示します。

```js
const exampleData = [34, 234, 73, 90, 6, 52];
d3.min(exampleData);
d3.max(exampleData);
```

データセットにネストされた配列がある場合もあります。たとえば散布図の例にあった `[x, y]` の座標ペアのようなものです。その場合、D3 に最大値と最小値の計算方法を教える必要があります。幸い、`min()` と `max()` の両方のメソッドはコールバック関数を受け取ります。この例では、コールバック関数の引数 `d` は現在の内側の配列です。コールバック関数は、最大値または最小値を計算したい内側の配列の要素（`x` または `y` の値）を返す必要があります。配列の配列で最小値と最大値を見つける例を示します。

```js
const locationData = [
  [1, 7],
  [6, 3],
  [8, 3]
];
const minX = d3.min(locationData, d => d[0]);
```

`minX` は `1` の値を持ちます。

# --instructions--

`positionData` 配列は x、y、z 座標のサブ配列を保持しています。D3 のメソッドを使って、配列の中の z 座標（3番目の値）の最大値を見つけ、`output` 変数に保存してください。

# --hints--

`h2` のテキストは `8` であるべきです。

```js
assert.strictEqual(output, 8);
assert.strictEqual(document.querySelector('h2')?.textContent, '8');
```

コードは `max()` メソッドを使うべきです。

```js
assert.match(
  code,
  /\.max/g,
  'Your code should use the <code>max()</code> method.'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];
    // Add your code below this line

    const output = undefined; // Change this line

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];

    const output = d3.max(positionData, d => d[2]);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
