---
id: 587d7fa7367417b2b2512bc5
title: D3で動的データを扱う
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--

最後の2つのチャレンジでは、`data()` と `enter()` メソッドを使ってD3で動的にデータを表示する基本を扱います。これらのメソッドはデータセットを受け取り、`append()` メソッドと組み合わせて、データセットの各エントリーに対して新しいDOM要素を作成します。

前のチャレンジでは、`h2` 配列の各アイテムに対して新しい `dataset` 要素を作成しましたが、すべて同じテキスト `New Title` が含まれていました。これは、各 `h2` 要素にバインドされたデータを利用していなかったためです。

D3の `text()` メソッドは、引数として文字列またはコールバック関数を取ることができます：

```js
selection.text(d => d);
```

上の例では、パラメータ `d` は、選択にバインドされたデータセットの単一のエントリーを指します。

現在の例を文脈として使うと、最初の `h2` 要素は12にバインドされ、2番目の `h2` 要素は31にバインドされ、3番目の `h2` 要素は22にバインドされる、という具合です。

# --instructions--

`text()` メソッドを変更して、各 `h2` 要素が `dataset` 配列の対応する値と、半角スペースと文字列 `USD` を表示するようにしてください。例えば、最初の見出しは `12 USD` になるはずです。

# --hints--

最初の `h2` はテキスト `12 USD` を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

2番目の `h2` はテキスト `31 USD` を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

3番目の `h2` はテキスト `22 USD` を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

4番目の `h2` はテキスト `17 USD` を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

5番目の `h2` はテキスト `25 USD` を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

6番目の `h2` はテキスト `18 USD` を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

7番目の `h2` はテキスト `29 USD` を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

8番目の `h2` はテキスト `14 USD` を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

9番目の `h2` はテキスト `9 USD` を持つべきです。

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```
