---
id: 587d7fac367417b2b2512bde
title: 事前定義されたスケールを使って要素を配置する
challengeType: 6
forumTopicId: 301494
dashedName: use-a-pre-defined-scale-to-place-elements
---

# --description--

スケールが設定できたので、散布図を再びマッピングしましょう。スケールは、`x` と `y` の生データを、SVG上で正しくフィットしてレンダリングできる値に変換する処理関数のようなものです。これにより、データが画面のプロットエリア内に収まります。

SVGの図形の座標属性値はスケーリング関数で設定します。これは、`x` または `y` 要素の `rect` と `text` 属性、あるいは `cx` の `cy` と `circles` 属性を含みます。例を示します。

```js
shape.attr('x', d => xScale(d[0]));
```

スケールは図形の座標属性を設定して、データポイントをSVG上に配置します。ツールチップやラベルの `text()` メソッドで実際のデータ値を表示する場合は、スケールを適用する必要はありません。

# --instructions--

`xScale` と `yScale` を使って、`circle` と `text` の両方の図形をSVG上に配置してください。`circles` にはスケールを適用して、`cx` と `cy` の属性を設定します。半径は `5` 単位にしてください。

`text` 要素にはスケールを適用して、`x` と `y` の属性を設定します。ラベルは点の右側にオフセットしてください。そのために、`10` のデータ値に `x` 単位を加えてから `xScale` に渡してください。

# --hints--

コードには `circle` 要素が10個あるはずです。

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

最初の `circle` 要素は、スケール適用後に `cx` の値がおよそ `91`、`cy` の値がおよそ `368` であるべきです。また、`r` の値は `5` である必要があります。

```js
const firstCircle = document.querySelectorAll('circle')[0];
assert.strictEqual(Math.round(firstCircle?.getAttribute('cx')), 91);
assert.strictEqual(Math.round(firstCircle?.getAttribute('cy')), 368);
assert.strictEqual(Math.round(firstCircle?.getAttribute('r')), 5);
```

2番目の `circle` 要素は、スケール適用後に `cx` の値がおよそ `159`、`cy` の値がおよそ `181` であるべきです。また、`r` の値は `5` である必要があります。

```js
const secondCircle = document.querySelectorAll('circle')[1];
assert.strictEqual(Math.round(secondCircle?.getAttribute('cx')), 159);
assert.strictEqual(Math.round(secondCircle?.getAttribute('cy')), 181);
assert.strictEqual(Math.round(secondCircle?.getAttribute('r')), 5);
```

3番目の `circle` 要素は、スケール適用後に `cx` の値がおよそ `340`、`cy` の値がおよそ `329` であるべきです。また、`r` の値は `5` である必要があります。

```js
const thirdCircle = document.querySelectorAll('circle')[2];
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cx')), 340);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cy')), 329);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('r')), 5);
```

4番目の `circle` 要素は、スケール適用後に `cx` の値がおよそ `131`、`cy` の値がおよそ `60` であるべきです。また、`r` の値は `5` である必要があります。

```js
const fourthCircle = document.querySelectorAll('circle')[3];
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cy')), 60);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('r')), 5);
```

5番目の `circle` 要素は、スケール適用後に `cx` の値がおよそ `440`、`cy` の値がおよそ `237` であるべきです。また、`r` の値は `5` である必要があります。

```js
const fifthCircle = document.querySelectorAll('circle')[4];
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cx')), 440);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cy')), 237);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('r')), 5);
```

6番目の `circle` 要素は、スケール適用後に `cx` の値がおよそ `271`、`cy` の値がおよそ `306` であるべきです。また、`r` の値は `5` である必要があります。

```js
const sixthCircle = document.querySelectorAll('circle')[5];
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cx')), 271);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cy')), 306);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('r')), 5);
```

7番目の `circle` 要素は、スケール適用後に `cx` の値がおよそ `361`、`cy` の値がおよそ `351` であるべきです。また、`r` の値は `5` である必要があります。

```js
const seventhCircle = document.querySelectorAll('circle')[6];
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cx')), 361);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cy')), 351);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('r')), 5);
```

8番目の `circle` 要素は、スケール適用後に `cx` の値がおよそ `261`、`cy` の値がおよそ `132` であるべきです。また、`r` の値は `5` である必要があります。

```js
const eighthCircle = document.querySelectorAll('circle')[7];
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cx')), 261);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cy')), 132);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('r')), 5);
```

9番目の `circle` 要素は、スケール適用後に `cx` の値がおよそ `131`、`cy` の値がおよそ `144` であるべきです。また、`r` の値は `5` である必要があります。

```js
const ninthCircle = document.querySelectorAll('circle')[8];
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cy')), 144);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('r')), 5);
```

10番目の `circle` 要素は、スケール適用後に `cx` の値がおよそ `79`、`cy` の値がおよそ `326` であるべきです。また、`r` の値は `5` である必要があります。

```js
const tenthCircle = document.querySelectorAll('circle')[9];
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cx')), 79);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cy')), 326);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('r')), 5);
```

コードには `text` 要素が10個あるはずです。

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

最初のラベルは、スケール適用後に `x` の値がおよそ `100`、`y` の値がおよそ `368` であるべきです。

```js
const firstLabel = document.querySelectorAll('text')[0];
assert.strictEqual(Math.round(firstLabel?.getAttribute('x')), 100);
assert.strictEqual(Math.round(firstLabel?.getAttribute('y')), 368);
```

2番目のラベルは、スケール適用後に `x` の値がおよそ `168`、`y` の値がおよそ `181` であるべきです。

```js
const secondLabel = document.querySelectorAll('text')[1];
assert.strictEqual(Math.round(secondLabel?.getAttribute('x')), 168);
assert.strictEqual(Math.round(secondLabel?.getAttribute('y')), 181);
```

3番目のラベルは、スケール適用後に `x` の値がおよそ `350`、`y` の値がおよそ `329` であるべきです。

```js
const thirdLabel = document.querySelectorAll('text')[2];
assert.strictEqual(Math.round(thirdLabel?.getAttribute('x')), 350);
assert.strictEqual(Math.round(thirdLabel?.getAttribute('y')), 329);
```

4番目のラベルは、スケール適用後に `x` の値がおよそ `141`、`y` の値がおよそ `60` であるべきです。

```js
const fourthLabel = document.querySelectorAll('text')[3];
assert.strictEqual(Math.round(fourthLabel?.getAttribute('x')), 141);
assert.strictEqual(Math.round(fourthLabel?.getAttribute('y')), 60);
```

5番目のラベルは、スケール適用後に `x` の値がおよそ `449`、`y` の値がおよそ `237` であるべきです。

```js
const fifthLabel = document.querySelectorAll('text')[4];
assert.strictEqual(Math.round(fifthLabel?.getAttribute('x')), 449);
assert.strictEqual(Math.round(fifthLabel?.getAttribute('y')), 237);
```

6番目のラベルは、スケール適用後に `x` の値がおよそ `280`、`y` の値がおよそ `306` であるべきです。

```js
const sixthLabel = document.querySelectorAll('text')[5];
assert.strictEqual(Math.round(sixthLabel?.getAttribute('x')), 280);
assert.strictEqual(Math.round(sixthLabel?.getAttribute('y')), 306);
```

7番目のラベルは、スケール適用後に `x` の値がおよそ `370`、`y` の値がおよそ `351` であるべきです。

```js
const seventhLabel = document.querySelectorAll('text')[6];
assert.strictEqual(Math.round(seventhLabel?.getAttribute('x')), 370);
assert.strictEqual(Math.round(seventhLabel?.getAttribute('y')), 351);
```

8番目のラベルは、スケール適用後に `x` の値がおよそ `270`、`y` の値がおよそ `132` であるべきです。

```js
const eighthLabel = document.querySelectorAll('text')[7];
assert.strictEqual(Math.round(eighthLabel?.getAttribute('x')), 270);
assert.strictEqual(Math.round(eighthLabel?.getAttribute('y')), 132);
```

9番目のラベルは、スケール適用後に `x` の値がおよそ `140`、`y` の値がおよそ `144` であるべきです。

```js
const ninthLabel = document.querySelectorAll('text')[8];
assert.strictEqual(Math.round(ninthLabel?.getAttribute('x')), 140);
assert.strictEqual(Math.round(ninthLabel?.getAttribute('y')), 144);
```

10番目のラベルは、スケール適用後に `x` の値がおよそ `88`、`y` の値がおよそ `326` であるべきです。

```js
const tenthLabel = document.querySelectorAll('text')[9];
assert.strictEqual(Math.round(tenthLabel?.getAttribute('x')), 88);
assert.strictEqual(Math.round(tenthLabel?.getAttribute('y')), 326);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1]);
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));
  </script>
</body>
```
