---
id: bad87fee1348bd9aed308826
title: Den übergeordneten Knoten eines Elements mit jQuery ansprechen
challengeType: 6
forumTopicId: 18321
dashedName: target-the-parent-of-an-element-using-jquery
---

# --description--

Jedes HTML-Element hat ein `parent`-Element, von dem es Eigenschaften `inherits`.

Zum Beispiel hat das `h3`-Element in Ihrem `jQuery Playground` das übergeordnete Element `<div class="container-fluid">`, das selbst das übergeordnete Element `body` hat.

jQuery hat eine Funktion namens `parent()`, mit der Sie auf das übergeordnete Element des Elements zugreifen können, das Sie ausgewählt haben.

Hier ist ein Beispiel, wie Sie die `parent()`-Funktion verwenden würden, wenn Sie dem übergeordneten Element des `left-well`-Elements eine Hintergrundfarbe Blau geben möchten:

```js
$("#left-well").parent().css("background-color", "blue")
```

Geben Sie dem übergeordneten Element des `#target1`-Elements eine Hintergrundfarbe Rot.

# --hints--

Ihr `left-well`-Element sollte einen roten Hintergrund haben.

```js
assert(
  $('#left-well').css('background-color') === 'red' ||
    $('#left-well').css('background-color') === 'rgb(255, 0, 0)' ||
    $('#left-well').css('background-color').toLowerCase() === '#ff0000' ||
    $('#left-well').css('background-color').toLowerCase() === '#f00'
);
```

Sie sollten die `.parent()`-Funktion verwenden, um dieses Element zu ändern.

```js
assert(code.match(/\.parent\s*\(\s*\)\s*\.css/g));
```

Die `.parent()`-Methode sollte auf dem `#target1`-Element aufgerufen werden.

```js
assert(
  code.match(/\$\s*?\(\s*?(?:'|")\s*?#target1\s*?(?:'|")\s*?\)\s*?\.parent/gi)
);
```

Sie sollten nur jQuery verwenden, um diese Klassen dem Element hinzuzufügen.

```js
assert(code.match(/<div class="well" id="left-well">/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");

  });
</script>

<!-- Only change code above this line -->

<body>
  <div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
    <div class="row">
      <div class="col-xs-6">
        <h4>#left-well</h4>
        <div class="well" id="left-well">
          <button class="btn btn-default target" id="target1">#target1</button>
          <button class="btn btn-default target" id="target2">#target2</button>
          <button class="btn btn-default target" id="target3">#target3</button>
        </div>
      </div>
      <div class="col-xs-6">
        <h4>#right-well</h4>
        <div class="well" id="right-well">
          <button class="btn btn-default target" id="target4">#target4</button>
          <button class="btn btn-default target" id="target5">#target5</button>
          <button class="btn btn-default target" id="target6">#target6</button>
        </div>
      </div>
    </div>
  </div>
</body>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
  });
</script>

<!-- Only change code above this line -->

<body>
  <div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
    <div class="row">
      <div class="col-xs-6">
        <h4>#left-well</h4>
        <div class="well" id="left-well">
          <button class="btn btn-default target" id="target1">#target1</button>
          <button class="btn btn-default target" id="target2">#target2</button>
          <button class="btn btn-default target" id="target3">#target3</button>
        </div>
      </div>
      <div class="col-xs-6">
        <h4>#right-well</h4>
        <div class="well" id="right-well">
          <button class="btn btn-default target" id="target4">#target4</button>
          <button class="btn btn-default target" id="target5">#target5</button>
          <button class="btn btn-default target" id="target6">#target6</button>
        </div>
      </div>
    </div>
  </div>
</body>
```
