---
id: bad87fee1348bd9bedc08826
title: HTML-Elemente mit Selektoren in jQuery ansprechen
challengeType: 6
forumTopicId: 18319
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-html-elements-with-selectors-using-jquery
---

# --description--

Jetzt haben wir eine `document ready`-Funktion.

Schreiben wir nun unsere erste jQuery-Anweisung. Alle jQuery-Funktionen beginnen mit einem `$`, der üblicherweise als Dollarzeichen-Operator oder als Bling bezeichnet wird.

jQuery wählt oft ein HTML-Element mit einem <dfn>Selektor</dfn> aus und führt dann etwas mit diesem Element aus.

Zum Beispiel lassen wir alle Ihre `button`-Elemente hüpfen. Fügen Sie einfach diesen Code in Ihre document ready-Funktion ein:

```js
$("button").addClass("animated bounce");
```

Beachten Sie, dass wir bereits sowohl die jQuery-Bibliothek als auch die Animate.css-Bibliothek im Hintergrund eingebunden haben, damit Sie sie im Editor verwenden können. Sie verwenden also jQuery, um die Animate.css-`bounce`-Klasse auf Ihre `button`-Elemente anzuwenden.

# --hints--

Sie sollten die jQuery-`addClass()`-Funktion verwenden, um die Klassen `animated` und `bounce` Ihren `button`-Elementen hinzuzufügen.

```js
assert($('button').hasClass('animated') && $('button').hasClass('bounce'));
```

Sie sollten jQuery nur verwenden, um diese Klassen dem Element hinzuzufügen.

```js
assert(!code.match(/class.*animated/g));
```

Ihr jQuery-Code sollte sich innerhalb der `$(document).ready();`-Funktion befinden.

```js
assert(
  code.replace(/\s/g, '').match(/\$\(document\)\.ready\(function\(\)\{\$/g)
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
