---
id: bad87fee1348bd9aeda08826
title: Elemente mit der id mit jQuery ansprechen
challengeType: 6
forumTopicId: 18317
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-id-using-jquery
---

# --description--

Sie können Elemente auch über ihre id-Attribute ansprechen.

Zuerst wählen Sie Ihr `button`-Element mit der id `target3` aus, indem Sie den `$("#target3")`-Selektor verwenden.

Beachten Sie, dass Sie, genau wie bei CSS-Deklarationen, ein `#` vor den Namen der id schreiben.

Dann verwenden Sie die jQuery-`.addClass()`-Funktion, um die Klassen `animated` und `fadeOut` hinzuzufügen.

So würden Sie das `button`-Element mit der id `target6` ausblenden lassen:

```js
$("#target6").addClass("animated fadeOut");
```

# --hints--

Sie sollten das `button`-Element mit dem `id` von `target3` auswählen und die jQuery-`addClass()`-Funktion verwenden, um ihm die Klasse `animated` zu geben.

```js
assert($('#target3').hasClass('animated'));
```

Sie sollten das Element mit der id `target3` ansprechen und die jQuery-`addClass()`-Funktion verwenden, um ihm die Klasse `fadeOut` zu geben.

```js
assert(
  ($('#target3').hasClass('fadeOut') || $('#target3').hasClass('fadeout')) &&
    code.match(/\$\(\s*.#target3.\s*\)/g)
);
```

Sie sollten nur jQuery verwenden, um diese Klassen dem Element hinzuzufügen.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
    $("#target3").addClass("animated fadeOut");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
