---
id: bad87fee1348bd9aedc08826
title: Elemente mit Klasse mit jQuery ansprechen
challengeType: 6
forumTopicId: 18316
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-class-using-jquery
---

# --description--

Sie sehen, wie wir alle Ihre `button` Elemente zum Hüpfen gebracht haben? Wir haben sie mit `$("button")` ausgewählt und ihnen dann mit `.addClass("animated bounce");` einige CSS-Klassen hinzugefügt.

Sie haben gerade die jQuery-`.addClass()`-Funktion verwendet, mit der Sie Klassen zu Elementen hinzufügen können.

Zuerst sprechen wir Ihre `div` Elemente mit der Klasse `well` an, indem wir den `$(".well")`-Selektor verwenden.

Beachten Sie, dass Sie, genau wie bei CSS-Deklarationen, ein `.` vor den Klassennamen schreiben.

Verwenden Sie dann die jQuery-`.addClass()`-Funktion, um die Klassen `animated` und `shake` hinzuzufügen.

Zum Beispiel könnten Sie alle Elemente mit der Klasse `text-primary` zum Wackeln bringen, indem Sie Folgendes zu Ihrem `document ready function` hinzufügen:

```js
$(".text-primary").addClass("animated shake");
```

# --hints--

Sie sollten die jQuery-`addClass()`-Funktion verwenden, um allen Ihren Elementen mit der Klasse `animated` die Klassen `shake` und `well` zu geben.

```js
assert($('.well').hasClass('animated') && $('.well').hasClass('shake'));
```

Sie sollten jQuery nur verwenden, um diese Klassen dem Element hinzuzufügen.

```js
assert(!code.match(/class\.\*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
