---
id: 65822773f4c60a479ec89f66
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

All'interno della funzione `print_expenses`, crea un loop `for` che itera su ogni elemento della lista `expenses`. Usa `expense` come variabile del loop e sposta `pass` all'interno del corpo del loop.

# --hints--

Dovresti creare un loop `for` che itera su `expenses` nella tua funzione `print_expenses`. Ricorda di usare `pass` all'interno del corpo del loop.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

Dovresti usare `expense` come variabile del loop.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

Dovresti avere `pass` solo all'interno del corpo del loop.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
