---
id: 6582267ef39f5047411c1743
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Successivamente, definisci una funzione chiamata `print_expenses` che prende un parametro `expenses`. Questa funzione sarà usata in seguito per mostrare ogni spesa nella tua lista.

Riempi il corpo della tua nuova funzione con un'istruzione `pass`.

# --hints--

Dovresti definire una nuova funzione chiamata `print_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

La tua funzione `print_expenses` dovrebbe prendere un singolo parametro `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
