---
id: 6582352cbdbcde4d0c4f7b0b
title: Passaggio 20
challengeType: 20
dashedName: step-20
---

# --description--

Per chiamare una funzione lambda puoi usare la sintassi usuale delle funzioni con una coppia di parentesi tonde dopo il nome della variabile.

Chiama la tua funzione lambda `test` e passa `3` come argomento. Poi, stampa il risultato.

# --hints--

Dovresti avere `print(test(3))` nel tuo codice.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
