---
id: 6582575b8089f85b8b92d7c8
title: Passaggio 42
challengeType: 20
dashedName: step-42
---

# --description--

Una volta ottenuti i dettagli della spesa, devi chiamare la funzione `add_expense` per aggiungere la nuova spesa alla lista `expenses`.

Dopo aver ottenuto `amount` e `category` usando `input()`, chiama la funzione `add_expense`, passando tre argomenti: `expenses`, `amount` e `category`.

- `expenses` è la lista vuota creata nella funzione principale all'inizio di questo progetto.
- `amount` è l'importo della spesa.
- `category` è la categoria della spesa.

# --hints--

Dovresti chiamare la funzione `add_expense()` all'interno della tua istruzione `if`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

Dovresti passare tre argomenti alla tua chiamata `add_expense()`: `expenses`, `amount` e `category`. L'ordine è importante.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

La tua chiamata `add_expense()` dovrebbe venire dopo l'assegnazione di `amount` e `category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
