---
id: 658238f7604f154ea9a23e1e
title: Passaggio 23
challengeType: 20
dashedName: step-23
---

# --description--

La funzione `sum()` restituisce la somma degli elementi nell'iterabile passato come argomento. Userai `sum()` insieme alle funzioni `map()` e `lambda` per ottenere l'importo totale delle spese.

Per ora, fai un piccolo test e modifica la tua chiamata attuale a `print()` sostituendo la chiamata a `list()` con una chiamata alla funzione `sum()` passando come argomento la chiamata attuale a `map()`.

# --hints--

Dovresti stampare `sum(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
