---
id: 658212ba98182f3e855e85f9
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

In questo progetto, costruirai un semplice ma funzionale tracker di spese in Python.

Inizia definendo una funzione chiamata `add_expense` che prende tre parametri: `expenses`, `amount` e `category`. Usa la parola chiave `pass` per compilare il corpo della funzione.

# --hints--

Dovresti usare la parola chiave `def` per dichiarare una funzione chiamata `add_expense`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

La tua funzione `add_expense` dovrebbe prendere tre parametri: `expenses`, `amount` e `category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
