---
id: 65822308ded1ba4632f66aa3
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Un dizionario è un altro tipo di dato integrato in Python. Un dizionario è una raccolta di dati sotto forma di coppie *chiave*-*valore*. I dizionari si definiscono con parentesi graffe (`{}`) e contengono coppie chiave-valore separate da virgole. Ogni chiave è seguita da due punti (`:`) e dal valore:

```py
{'amount': 50.0, 'category': 'Food'}
```

Nell'esempio sopra, `'amount'` e `'category'` sono chiavi, mentre `50.0` e `'Food'` sono i loro valori corrispondenti.

Crea un dizionario con una chiave `'amount'` e valore del parametro `amount` e passa il tuo nuovo dizionario alla chiamata `.append()`.

# --hints--

Dovresti passare un dizionario come argomento al metodo `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

Dovresti passare `{'amount': amount}` come argomento al metodo `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
