---
id: 65824872894f59546e3084e2
title: Passaggio 34
challengeType: 20
dashedName: step-34
---

# --description--

Il loop `while` che hai creato nel passaggio precedente è un loop infinito che permetterà al programma di presentare continuamente le opzioni del menu finché l’utente non deciderà di uscire.

Dopo la chiamata `print()`, aggiungi un’altra chiamata `print()` per stampare la stringa `'1. Add an expense'`.

# --hints--

Dovresti stampare `'1. Add an expense'` all’interno del tuo loop `while` dopo la chiamata `print()` esistente.

```js
({ test: () => assert(runPython(`
call_lst = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")
test_lst = ["print('\\\\nExpense Tracker')", "print('1. Add an expense')"]
all(call_lst[n].is_equivalent(i) for n, i in enumerate(test_lst))
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
--fcc-editable-region--
    while True:
        print('\nExpense Tracker')
        
--fcc-editable-region--
```
