---
id: 6582324b5d37fe4b622bc3c0
title: Passaggio 19
challengeType: 20
dashedName: step-19
---

# --description--

Le funzioni lambda sono funzioni brevi e anonime in Python, ideali per attività semplici e una tantum. Sono definite dalla parola chiave `lambda` e usano la seguente sintassi:

```py
lambda x: expr
```

Nell'esempio sopra, `x` rappresenta un parametro da usare nell'espressione `expr` e funziona come qualsiasi parametro in una funzione tradizionale. `expr` è l'espressione che viene valutata e restituita quando la funzione lambda viene chiamata.

Crea una variabile chiamata `test` e assegnale una funzione lambda che prende un parametro `x` e restituisce `x * 2`.

# --hints--

Dovresti dichiarare una variabile chiamata `test`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

Dovresti assegnare `lambda x: x * 2` alla tua variabile `test`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
