---
id: 65822934c8a0454803e3ef12
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Adesso, mostrerai i dettagli per ogni spesa.

All'interno del loop `for`, sostituisci `pass` con una chiamata `print()` e passale la seguente f-string: `f'Amount: {expense}, Category: {expense}'`.

# --hints--

Dovresti stampare `f'Amount: {expense}, Category: {expense}'` nel tuo loop `for`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

Non dovresti avere `pass` all'interno del corpo del tuo loop.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
