---
id: 65824111a09164518320088d
title: Passaggio 30
challengeType: 20
dashedName: step-30
---

# --description--

La funzione `filter()` ti permette di selezionare elementi da un iterabile, come una lista, in base all'output di una funzione:

```py
filter(my_function, my_list)
```

`filter()` prende una funzione come primo argomento e un iterabile come secondo argomento. Restituisce un iteratore, un oggetto speciale che ti consente di iterare sugli elementi di una collezione, come una lista.

Il risultato dell'esempio sopra è un iteratore contenente gli elementi di `my_list` per cui `my_function` restituisce `True`.

All'interno della funzione `filter_expenses_by_category`, chiama `filter()` passando la funzione `lambda` che hai scritto nel passaggio precedente come primo argomento e la lista `expenses` come secondo argomento.

# --hints--

Dovresti chiamare `filter()` all'interno della funzione `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

Dovresti passare `lambda expense: expense['category'] == category` come primo argomento alla chiamata `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

Dovresti passare `expenses` come secondo argomento alla chiamata `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
