---
id: 65823ff0d4b991510fade1a8
title: Passaggio 29
challengeType: 20
dashedName: step-29
---

# --description--

All'interno della funzione `filter_expenses_by_category`, sostituisci `pass` con una funzione `lambda`. Usa `expense` come parametro e valuta il confronto tra il valore della chiave `'category'` del dizionario `expense` e `category` usando l'operatore di uguaglianza.

# --hints--

Dovresti creare una funzione `lambda` che usa un parametro chiamato `expense` e valuta l'espressione `expense['category'] == category` all'interno della funzione `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

Non dovresti avere `pass` all'interno della funzione `filter_expenses_by_category`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
