---
id: 65824561f3478e5371a33ae5
title: Passaggio 32
challengeType: 20
dashedName: step-32
---

# --description--

Il passaggio successivo è definire la funzione principale, che sarà il punto di ingresso del programma interattivo per il monitoraggio delle spese.

Definisci una funzione chiamata `main` senza parametri. Riempi il corpo della funzione con la lista `expenses` che hai creato all'inizio di questo progetto. Userai questa lista per archiviare i record delle spese.

# --hints--

Dovresti definire una funzione chiamata `main()` senza parametri.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

Dovresti spostare la lista `expenses` all'interno del corpo della funzione `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
