---
id: 65823dde36392f505a39f7c7
title: Passaggio 27
challengeType: 20
dashedName: step-27
---

# --description--

Infine, passa la tua chiamata `map()` alla funzione `sum()` per ottenere le spese totali e restituisci il risultato.

# --hints--

Dovresti chiamare la funzione `sum()` passando come argomento la tua attuale chiamata `map()` e restituire il risultato dalla tua funzione `total_expenses`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
