---
id: 65821fcc010c3245718f2a06
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Il parametro `expenses` della tua funzione `add_expense` sarà una lista di spese. Vuoi poter aggiungere elementi alla fine della tua lista. Per questo userai il metodo `.append()` delle liste.

Aggiungi una chiamata al metodo `.append()` sulla lista `expenses`. Per ora non passare argomenti a `.append()`.

# --hints--

Dovresti avere `expenses.append()` nella tua funzione.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

Non dovresti avere `pass` nella tua funzione.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
