---
id: 65823b0167fd894f4a7ea60a
title: Passaggio 24
challengeType: 20
dashedName: step-24
---

# --description--

Adesso implementerai la stessa logica all'interno della funzione `total_expenses`.

Per ora, elimina sia la funzione `test` sia la chiamata `print()`.

# --hints--

Dovresti eliminare la variabile `test` e la chiamata `print()` correlata.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_variable("test")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(sum(map(test, [2, 3, 5, 8])))
--fcc-editable-region--

expenses = []
```
