---
id: 65822bd82d708c4895080c35
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

In Python, una cosa importante da sapere è che lo stesso tipo di citazione usato per definire una stringa non può essere usato al suo interno. Per esempio, la stringa `'I'm a string!'` non è valida. Per usare l'apostrofo all'interno di quella stringa dovresti:

- Fare escape dell'apostrofo anteponendo una barra rovesciata: `'I\'m a string!'`
- Oppure usare le virgolette doppie per definire la stringa: `"I'm a string!"` (preferito).

Puoi consultare i valori in un dizionario tramite le sue chiavi. Devi usare la notazione a parentesi e includere la chiave tra parentesi quadre:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

Attualmente stai interpolando il dizionario `expense` nella tua f-string. Modifica l'espressione della f-string per consultare il valore della chiave `'amount'` e della chiave `'category'` nel dizionario `expense`.

# --hints--

Dovresti passare `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` alla tua chiamata `print()`. Ricorda di usare le virgolette doppie all'interno della tua f-string con virgolette singole e viceversa.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
