---
id: 65823cfc74aa564ffc460489
title: Passaggio 26
challengeType: 20
dashedName: step-26
---

# --description--

Ora, chiama `map()` passando la tua funzione `lambda` come primo argomento e la lista `expenses` come secondo argomento.

# --hints--

Dovresti chiamare la funzione `map()` all'interno della funzione `total_expenses`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

Dovresti passare la tua funzione `lambda` come primo argomento alla chiamata `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

Dovresti passare la lista `expenses` come secondo argomento alla chiamata `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
