---
id: 65823e87c9741750a22085a7
title: Passaggio 28
challengeType: 20
dashedName: step-28
---

# --description--

Definisci una funzione chiamata `filter_expenses_by_category` che prende due parametri: `expenses` e `category`. Usa `pass` per completare il corpo della funzione.

# --hints--

Dovresti definire una funzione chiamata `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

La tua funzione `filter_expenses_by_category` dovrebbe prendere `expenses` e `category` come parametri.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
