---
id: 658246c28575d653d1f89b59
title: Passaggio 33
challengeType: 20
dashedName: step-33
---

# --description--

Un loop `while` è un altro tipo di loop che esegue una porzione di codice finché una condizione specificata è `True`. Il loop termina quando la condizione diventa `False`:

```py
while condition:
    <code>
```

Sotto la lista `expenses`, crea un loop `while`. Usa `True` come condizione e stampa la stringa `'\nExpense Tracker'` all’interno del corpo del loop per mostrare il titolo del programma.

# --hints--

Dovresti creare un loop `while` usando `True` come condizione all’interno della tua funzione `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Dovresti stampare `'\nExpense Tracker'` nel tuo loop `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
