---
id: 658225d560369446ece5518b
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Aggiungi un’altra coppia chiave-valore al dizionario che stai aggiungendo alla lista `expense`. Usa la stringa `'category'` come chiave e il parametro `category` come valore.

# --hints--

Dovresti aggiungere una seconda coppia chiave-valore al dizionario che stai passando come argomento al metodo `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

Dovresti passare `{'amount': amount, 'category': category}` come argomento al metodo `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
