---
id: 65823634a308cb4d6bcc5fa6
title: Passaggio 21
challengeType: 20
dashedName: step-21
---

# --description--

Le funzioni lambda possono essere utilmente combinate con la funzione `map()`, che esegue una funzione specificata per ogni elemento in una collezione di oggetti, come una lista:

```py
map(lambda x: x * 2, [1, 2, 3])
```

La funzione da eseguire viene passata come primo argomento, e l'iterabile viene passato come secondo argomento.

Il risultato dell'esempio sopra sarebbe `[2, 4, 6]`, dove ogni elemento nella lista passata a `map()` è stato raddoppiato dall'azione della funzione lambda.

Modifica la tua chiamata a `print()` per stampare il risultato della chiamata a `map()` con `test` come primo argomento e `[2, 3, 5, 8]` come secondo argomento. Non potrai ancora vedere un output leggibile. 

# --hints--

Dovresti avere `print(map(test, [2, 3, 5, 8]))` nel tuo codice.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
