---
id: 65821c1c40ccac44adf38e6b
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Crea una lista vuota chiamata `expenses`. La userai per archiviare ciascuna delle tue spese.

# --hints--

Dovresti dichiarare una nuova variabile chiamata `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

Dovresti assegnare una lista vuota alla variabile `expenses`.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
