---
id: 66314f2add0a8b57da5d6d91
title: Passaggio 22
challengeType: 20
dashedName: step-22
---

# --description--

Dovresti vedere qualcosa come `<map object at 0xd273a8>` stampato sulla console, che è la rappresentazione stringa dell’oggetto map restituito da `map()`.

Per ottenere un output leggibile devi trasformare l’oggetto map in una lista. Fallo passando la chiamata `map()` come argomento alla funzione `list()`.

# --hints--

Devi passare la chiamata `map()` come argomento alla funzione `list()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
