---
id: 66694eaaeb7d253049fa7883
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Per aggiungere elementi a una lista, puoi scriverli tra le parentesi quadre, separando i valori con le virgole.

Aggiungi i numeri `1` e `2` all'interno della lista.

# --hints--

Il primo elemento di `my_list` dovrebbe essere `1`.

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

Il secondo elemento di `my_list` dovrebbe essere `2`.

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` dovrebbe avere due elementi all'interno.

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

Dovresti assegnare `[1, 2]` a `my_list` direttamente.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
