---
id: 66694fc4bba24f33ca01fa5b
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Python ha diversi metodi per le liste. Come metodi per aggiungere o rimuovere elementi dalla lista.

Puoi aggiungere un elemento alla fine di una lista usando il metodo `append()`. Per esempio:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` ora è `[4, 5, 6, 7]`.

Prova a usare il metodo `append()` per aggiungere `3` a `my_list`. Poi stampa la lista.

# --hints--

L'ultimo elemento di `my_list` dovrebbe essere `3`.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` dovrebbe avere una lunghezza di 3.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

Non dovresti modificare la dichiarazione della lista.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Dovresti stampare la lista.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

Dovresti stampare la lista dopo aver usato `append()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
