---
id: 666951d2e4b27e3a3c72ca5f
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Puoi anche consultare un singolo elemento per ottenere il suo valore.

Le liste hanno indice zero come le stringhe. Questo significa che il primo elemento si trova all'indice `0`, il secondo elemento all'indice `1` e così via.

Per consultare un elemento usi la notazione a parentesi. Per esempio, `example_list[1]` consulta l'elemento all'indice `1`, il secondo elemento, di `example_list`.

Stampa il primo elemento di `my_list`.

# --hints--

Dovresti stampare il primo elemento della lista con `print(my_list[0])`.

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
