---
id: 6669539c1379793f9cb8917c
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Le liste in Python sono mutabili, il che significa che il valore degli elementi della lista può essere modificato. Puoi cambiare il valore di un elemento usando la notazione a parentesi.

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

Questo farà sì che `example_list` abbia il valore `[4, 'oh', 6, 7]`.

Cambia il primo elemento di `my_list` in `0`, poi stampa la lista per verificare il valore.

# --hints--

Non modificare la dichiarazione della lista.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Il primo elemento di `my_list` dovrebbe essere 0.

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

Dovresti cambiare il primo elemento usando `my_list[0] = 0`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

Dovresti stampare la lista.

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

Dovresti stampare la lista dopo aver cambiato il primo elemento.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
