---
id: 66694bfa435cef241f6f9a68
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Prima di iniziare con il progetto, imparerai cosa sono le liste.

A differenza di numeri e stringhe, una <dfn>lista</dfn> è un tipo di dato che funziona come un contenitore per altri valori:

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

La lista si caratterizza per le parentesi quadre che racchiudono tutti i valori, e per la virgola che separa i valori, come: `["A", "happy", "list"]`. Se la lista non contiene alcun valore, allora è una lista vuota: `[]`.

Una lista può contenere tipi di dati diversi: `[1, "Up", ["Down", "Twice"]]`. Questo include tutti i tipi di dati possibili. Può anche contenere un'altra lista!

Crea una variabile chiamata `my_list` e assegnale una lista vuota.

# --hints--

Dovresti avere una variabile chiamata `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

La variabile `my_list` dovrebbe avere il valore di una lista vuota.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
