---
id: 666ab82c874add3c95145213
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Il metodo `pop()` può essere usato per rimuovere un elemento da una lista. Per impostazione predefinita, rimuove l'ultimo elemento della lista. Puoi passare un indice come argomento al metodo, e rimuoverà l'elemento all'indice specificato.

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

In questo caso, `fruits_list.pop(2)` rimuove l'elemento all'indice `2` della lista.

Usa `pop()` per rimuovere l'ultimo elemento da `my_list`, poi stampa `my_list`.

# --hints--

Dovresti usare `pop()`.

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

`my_list` dovrebbe essere `[0, 1, 2]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

Dovresti stampare `my_list` dopo aver usato `pop()`.

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
