---
id: 666aae6a5d30a71f1fd7749f
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Il metodo `insert` può aggiungere un elemento in qualsiasi posizione di una lista. Il primo argomento è la posizione in cui l'elemento deve essere aggiunto, e il secondo argomento è l'elemento da aggiungere. Per esempio, ecco come aggiungere un nuovo elemento nella terza posizione di `example_list`:

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

Usando `insert()`, aggiungi `1` a `my_list` nella posizione corretta in modo che conti in ordine crescente, poi stampa la lista.

# --hints--

Dovresti aggiungere `1` a `my_list` all'indice `1` usando `insert()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list` dovrebbe avere 4 elementi.

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` dovrebbe essere uguale a `[0, 1, 2, 3]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

Dovresti stampare la lista dopo aver aggiunto il nuovo elemento.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
