---
id: 686e64be3cd9db25282a956d
title: Progetta una mostra fotografica
challengeType: 25
dashedName: lab-photography-exhibit
demoType: onClick
---

# --description--

In questo laboratorio, creerai il layout di una mostra fotografica usando le classi utility di Tailwind CSS.

Per le immagini, puoi usare i seguenti URL se vuoi:

- https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg
- https://cdn.freecodecamp.org/curriculum/labs/alps.jpg
- https://cdn.freecodecamp.org/curriculum/labs/sea.jpg

**Obiettivo:** Soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento con una classe chiamata `main-container`.
2. Il tuo elemento `.main-container` dovrebbe avere anche le seguenti classi utility di Tailwind CSS:
   - la classe utility corretta per creare un contenitore griglia CSS.
   - una classe utility per definire il numero di colonne all’interno della griglia. Dovresti usare la classe utility che utilizza un numero fisso di colonne. Es. `grid-cols-<number>`.
   - una classe utility per impostare lo spazio tra righe e colonne per gli elementi della griglia.
3. All’interno del tuo elemento `.main-container` dovrebbero esserci almeno tre elementi ciascuno con una classe chiamata `card`.
4. Ciascuno dei tuoi elementi `.card` dovrebbe avere anche le seguenti classi utility di Tailwind CSS:
   - una classe utility per impostare un bordo arrotondato a tua scelta.
   - una classe utility per impostare la spaziatura interna a tua scelta. Dovresti usare questo formato `p-<number>`.
5. All’interno di ciascuno dei tuoi elementi `.card` dovresti avere i seguenti elementi:
   - un elemento immagine con gli attributi `src` e `alt`.
   - un elemento con la classe chiamata `subheading`.
   - un elemento con la classe chiamata `description`.
6. Ciascuno dei tuoi elementi immagine dovrebbe avere una classe utility per impostare un bordo arrotondato a tua scelta.
7. Ciascuno dei tuoi elementi `.subheading` dovrebbe usare classi utility per impostare il peso e la dimensione del carattere tipografico a tua scelta.
8. Ciascuno dei tuoi elementi `.description` dovrebbe usare una classe utility per impostare la dimensione del carattere tipografico a tua scelta.

# --hints--

Dovresti avere un elemento con una classe chiamata `main-container`.

```js
assert.exists(document.querySelector(".main-container"));
```

Il tuo elemento `.main-container` dovrebbe usare la classe utility corretta di Tailwind CSS per creare un contenitore griglia CSS.

```js
const mainContainer = document.querySelector(".main-container");
assert.isTrue(mainContainer.classList.contains("grid"));
```

Il tuo elemento `.main-container` dovrebbe usare una classe utility per definire il numero di colonne all’interno della griglia.

```js
const mainContainer = document.querySelector(".main-container");
// Grid classes can vary, so this check is for the presence of any grid-cols class
const hasGridCols = [...mainContainer.classList].some(className => className.startsWith("grid-cols-"));
assert.isTrue(hasGridCols);
```

Il tuo elemento `.main-container` dovrebbe usare una classe utility per impostare lo spazio tra righe e colonne per gli elementi della griglia.

```js
const mainContainer = document.querySelector(".main-container");
// Check for any gap class including gap-x and gap-y

const hasGap = [...mainContainer.classList].some(c => c.startsWith("gap-")) ||
  (classList.some(c => c.startsWith("gap-x-")) &&
   classList.some(c => c.startsWith("gap-y-")));

assert.isTrue(hasGap);
```

Il tuo elemento `.main-container` dovrebbe contenere almeno tre elementi ciascuno con una classe chiamata `card`.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isAtLeast(cards.length, 3);
```

Ciascuno dei tuoi elementi `.card` dovrebbe avere una classe utility per impostare un bordo arrotondato.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);

const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

cards.forEach(card => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => card.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
})
```

Ciascuno dei tuoi elementi `.card` dovrebbe avere una classe utility per impostare la spaziatura interna.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);
const hasPadding = [...cards].every(card => {
  return [...card.classList].some(c => c.startsWith("p-"));
});
assert.isTrue(hasPadding);
```

Ciascuno dei tuoi elementi `.card` dovrebbe contenere un elemento immagine.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
```

Ciascuno dei tuoi elementi `img` dovrebbe avere un attributo `src`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("src")));
images.forEach(img => assert.isTrue(img.getAttribute("src").length > 0));
```

Ciascuno dei tuoi elementi `img` dovrebbe avere un attributo `alt`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("alt")));
images.forEach(img => assert.isTrue(img.getAttribute("alt").length > 0));
```

Ciascuno dei tuoi elementi `.card` dovrebbe contenere un elemento con la classe chiamata `subheading`.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isAtLeast(subheadings.length, 3);
```

Ciascuno dei tuoi elementi `.subheading` dovrebbe avere del contenuto testuale.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(subheadings);

subheadings.forEach(subheading => {
  assert.isNotEmpty(subheading.textContent.trim());
});
```

Ciascuno dei tuoi elementi `.card` dovrebbe contenere un elemento con la classe chiamata `description`.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
assert.isAtLeast(descriptions.length, 3);
```

Ciascuno dei tuoi elementi `.description` dovrebbe avere del contenuto testuale.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(descriptions);

descriptions.forEach(description => {
  assert.isNotEmpty(description.textContent.trim());
});
```

Ciascuno dei tuoi elementi `img` dovrebbe avere una classe utility per impostare un bordo arrotondato.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isNotEmpty(images);
const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

images.forEach(img => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => img.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
});
```

Ciascuno dei tuoi elementi `.subheading` dovrebbe usare classi utility per impostare il peso del carattere tipografico.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedFontWeights = [
  "font-thin",
  "font-extralight",
  "font-light",
  "font-normal",
  "font-medium",
  "font-semibold",
  "font-bold",
  "font-extrabold",
  "font-black"
];

subheadings.forEach(subheading => {
  const hasFontWeight = predefinedFontWeights.some(weight => subheading.classList.contains(weight));
  assert.isTrue(hasFontWeight);
});
```

Ciascuno dei tuoi elementi `.subheading` dovrebbe usare classi utility per impostare la dimensione del carattere tipografico.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

subheadings.forEach(subheading => {
  const hasFontSize = predefinedSizes.some(size => subheading.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

Ciascuno dei tuoi elementi `.description` dovrebbe usare una classe utility per impostare la dimensione del carattere tipografico.

```js
const descriptions = document.querySelectorAll(".description");
assert.isNotEmpty(descriptions);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

descriptions.forEach(description => {
  const hasFontSize = predefinedSizes.some(size => description.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
 
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-purple-800 text-yellow-400 min-h-screen">
    <h1 class="text-center text-4xl font-bold my-12">Photography Exhibit</h1>

    <main class="main-container grid grid-cols-1 md:grid-cols-3 gap-8 px-4">
      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="Ancient Colosseum"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Colosseum, Rome</h2>
          <p class="description text-sm mb-4">
            There is so much history and beauty. The light hits it just right.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="Alpine Mountains"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">The Alps</h2>
          <p class="description text-sm mb-4">
            The mountains go on forever and the view is breathtaking.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="Calm Sea"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Seascape</h2>
          <p class="description text-sm mb-4">
            The water was still, the air was cool and the feeling is pure peace.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>
    </main>
  </body>
</html>
```
