---
id: 694a9ff65d4472dd451d8255
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次工作坊中，你将通过创建一个电影票预订计算器来练习如何在 Python 中使用布尔值和条件式语句。

首先声明保存有关用户和电影放映信息的变量。

正如你在之前的课程中学到的，变量是使用赋值操作符（`=`）这样赋值的：

```py
# Assigns the value 10 to variable x
x = 10
```

创建一个名为 `base_price` 的变量来保存电影票的基础价格，并将其值设置为 `15`。创建另一个名为 `age` 的变量来保存用户的年龄，并将其值设置为 `21`。

# --hints--

你应该声明一个名为 `base_price` 的变量。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

你应该将整数值 `15` 赋值给 `base_price`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

你应该声明一个名为 `age` 的变量。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

你应该将整数值 `21` 赋值给 `age`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
