---
id: 694cc1b515976889afd4cded
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

现在你需要保存一些关于电影票的字串值。你可以用这种方式在 Python 中保存字串：

```py
name = 'Alice'
```

创建一个名为 `seat_type` 的变量来保存用户选择的座位类型，并将其值设置为字串 `Gold`。创建另一个名为 `show_time` 的变量来保存电影的放映时间，并将其值设置为字串 `Evening`。

记得用单引号或双引号将两个值都括起来。

# --hints--

你应该声明一个名为 `seat_type` 的变量。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

你应该将字串 `Gold` 赋值给 `seat_type`。记住要用单引号或双引号将该值括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

你应该声明一个名为 `show_time` 的变量。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

你应该将字串 `Evening` 赋值给 `show_time`。记住要用单引号或双引号将值括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
