---
id: 6958e2377cacbc46a7f7383a
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

现在，为你的 `if` 语句添加一个 `else` 子句，并在 `else` 体内打印 `No extra charges will be applied`。只有当额外收费条件为假时，这才会打印该消息。

然后，在 `else` 子句下面，使用 `print()` 调用显示一条消息，内容为 `Extra charges:`，后跟更新后的 `extra_charges` 值，并在终端查看输出。

# --hints--

你应该有一个 `else` 子句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

你应该在你的 `else` 子句中打印 `No extra charges will be applied`。记得用单引号或双引号将消息括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

你应该在你的 `if...else` 语句下面有 `print('Extra charges:', extra_charges)`。

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
