---
id: 6958e0ecc121b0ae3dd9c5e0
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

你还应该处理用户不符合会员折扣资格的分支。

为你的 `if` 语句添加一个 `else` 子句，并在 `else` 体内打印 `User does not qualify for membership discount`。

你可以像这样打印一条消息，后面跟着一个变量：

```py
print('Total:', total)
```

你还想显示更新后的 `discount` 值。在 `if...else` 语句下面，使用 `print()` 调用来显示一条消息，内容是 `Discount:`，后面跟着更新后的 `discount` 值。然后，在终端查看输出。

# --hints--

你应该有一个 `else` 子句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

你应该在你的 `else` 子句中打印 `User does not qualify for membership discount`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

你应该在你的 `if...else` 语句下面有 `print('Discount:', discount)`。

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
