---
id: 6958e1875e592b58e8852757
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

现在创建一个名为 `extra_charges` 的变量并将其设置为 `0`。这将表现为在周末应用于电影票的额外费用。

然后，创建一个 `if` 语句来查看 `is_weekend` 是否为真。在 `if` 语句的体内，将 `extra_charges` 的值更新为 `2` 并在终端打印 `Extra charges will be applied`。

# --hints--

你应该声明一个名为 `extra_charges` 的变量。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

你应该将整数值 `0` 赋值给 `extra_charges`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

你应该有第四个 `if` 语句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

你的第四个 `if` 语句条件应为 `is_weekend`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

你应该在你的新建 `if` 语句中为 `extra_charges` 变量赋值整数值 `2`。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

你应该在你的新建 `if` 语句中打印 `Extra charges will be applied`。记得用单引号或双引号将消息括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
