---
id: 6958de0d773289a89515a6b8
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

大于或等于操作符（`>=`）用于查看其左侧的值是否大于或等于右侧的值。与大于操作符（`>`）不同，当两个值相等时，`>=` 也会返回 `True`：

```py
score = 50
print(score > 50) # False
print(score >= 50) # True
```

现在你将根据用户的年龄查看他们是否被允许预订晚场演出。

创建一个 `if` 语句来查看 `age` 是否大于或相等于 `21`。在 `if` 语句的体内，打印 `User is eligible for Evening shows`。

# --hints--

你的代码中应该有两个 `if` 语句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

你的第二个 `if` 语句条件应为 `age >= 21`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

你应该在你的新建 `if` 语句中打印 `User is eligible for Evening shows`。记得用单引号或双引号将消息括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
