---
id: 6958dd85587e10c99162dd81
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

正如你在之前的课程中学到的，在 Python 中，可以使用 `if` 语句根据条件是否为真来运行代码。

一个 `if` 语句由 `if` 关键字组成，后跟一个条件和一个冒号。当条件为真时要运行的代码，必须缩排，称为 `if` 语句的体。

```py
if condition:
    # Code to execute if condition is True
```

在此步骤中，你将根据用户的年龄查看其是否有资格预订电影票。

创建一个 `if` 语句来查看 `age` 是否大于 `17`。在 `if` 语句的体内，打印 `User is eligible to book a ticket`。只有当用户的年龄大于 `17` 时，才会打印该消息。

记得缩排 `if` 语句的体，并在 `print()` 调用内用单引号或双引号将消息括起来。

# --hints--

你应该有一个 `if` 语句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

`if` 语句条件应为 `age > 17`。

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

你应该在你的 `if` 语句内有一个 `print()` 调用。记得缩排 `if` 语句的体。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

你应该在你的 `if` 语句内打印 `User is eligible to book a ticket`。记得用单引号或双引号将消息括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
