---
id: 6958dfa0c74b237e1b56ce04
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

如果用户是成员，则他们有资格获得会员折扣。

创建一个名为 `discount` 的变量并将其值设置为 `0`。这将保存用户在电影票上获得的折扣。

# --hints--

你应该声明一个名为 `discount` 的变量。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

你应该将整数值 `0` 赋值给 `discount`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
