---
id: 6958e59845777391b86d7331
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

在 Python 中，`if` 语句也可以放置在另一个 `if` 语句的体内。这称为嵌套的 `if` 语句。

嵌套的 `if` 语句允许你仅在第一个条件已经满足后查看额外的条件。只有当外层 `if` 条件为真时，内层 `if` 语句才会运行。

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

现在你将根据用户选择的座位类型计算服务费。

仍然在你最后一个 `if` 语句体内，创建一个名为 `service_charges` 的变量并将其设置为 `0`。

然后，创建一个嵌套的 `if` 语句来查看 `seat_type` 是否相等于 `Premium`。在 `if` 语句的体内，将 `service_charges` 的值更新为 `5`。

# --hints--

你应该在最后一个 `if` 语句体内声明一个名为 `service_charges` 的变量。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

你应该将整数值 `0` 赋值给 `service_charges`。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

你应该在你最后一个 `if` 语句体内有一个 `if` 语句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

新的 `if` 语句条件应为 `seat_type == 'Premium'`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

你应该在你的新建 `if` 语句中为 `service_charges` 变量赋值整数值 `5`。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
