---
id: 696099c1ac72a4e585de1d5b
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

有些信息只能是真或假。正如你在之前的课程中学到的，这可以用布尔值来表现。

创建一个名为 `is_member` 的变量来表示用户是否是成员，并将其值设置为 `True`。

在 `is_member` 变量下创建另一个名为 `is_weekend` 的变量，用于指示电影放映是否在周末，并将其值设置为 `False`。不要用引号括起该值。

# --hints--

你应该声明一个名为 `is_member` 的变量。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

你应该将布尔值 `True` 赋值给 `is_member`。不要用引用将该值括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

你应该声明一个名为 `is_weekend` 的变量。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

你应该将布尔值 `False` 赋值给 `is_weekend`。不要用引用包围该值。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
