---
id: 6960ba6e6c95ed4ad896c6f2
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

预订规则还有一个例外。年龄在 `18` 到 `21` 之间的用户如果是成员，可以预订晚场演出。

当在 `if` 语句中使用多个逻辑操作符时，使用 `and` 连接的条件会在使用 `or` 连接的条件之前被计算。Python 中使用括号 `()` 来分组条件并控制它们的计算顺序。

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

现在，使用 `or` 操作符为你现有的 `if` 语句添加另一个条件，以查看 `is_member` 是否为真。使用括号 `()` 将 `show_time != 'Evening'` 和 `is_member` 条件组合在一起，如上例所示。

# --hints--

你最后的 `if` 语句条件应为 `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
