---
id: 696099abb6acc539b27055a1
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

在上一步中，你使用 `if` 语句查看了一个条件。`else` 分支用于处理条件为假的分支。以下是 `if...else` 语句的语法：

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

现在，为你的 `if` 语句添加一个 `else` 子句，并在 `else` 体内打印 `User is not eligible for Evening shows`。只有当 `if` 语句中的条件为假时，这才会打印。

# --hints--

你应该有一个 `else` 子句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

你应该在你的 `else` 子句中打印 `User is not eligible for Evening shows`。记得用单引号或双引号将消息括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
