---
id: 6960a848aacf859f7bb15447
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

现在将 `is_member` 变量的值更改为 `False`，因为该用户不是成员。

之后，你会看到 `discount` 值现在保持为 `0`，因为必须满足两个条件才能应用折扣。

# --hints--

你应该将 `is_member` 变量的值更新为 `False`。不要用引号括住该值。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--
is_member = True
--fcc-editable-region--
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
