---
id: 6960b6a6ea06e27dcc7eef37
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

现在你将查看用户是否满足预订电影票的条件。年龄在 `21` 岁或以上的用户可以始终无任何限制地预订票。

创建一个 `if` 语句来查看 `age` 是否大于或相等于 `21`。在 `if` 语句的体内，打印 `Ticket booking condition satisfied` 到终端。

然后，为你的 `if` 语句添加一个 `else` 子句，并在 `else` 体内打印 `Ticket booking failed due to restrictions`。

# --hints--

你应该有第五个 `if` 语句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

你的第五个 `if` 语句条件应为 `age >= 21`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

你应该在你的新建 `if` 语句中打印 `Ticket booking condition satisfied`。记得用单引号或双引号将消息括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

你的新建 `if` 语句应包含一个 `else` 子句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

你应该在你的 `else` 子句中打印 `Ticket booking failed due to restrictions`。记得用单引号或双引号将消息括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
