---
id: 6960a5165f7efe52a0c0aa39
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

在 Python 中，`and` 操作符用于查看多个条件是否为真。以下是如何在 `if` 语句中将两个条件组合使用的方法：

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

会员折扣应仅在其 `age` 大于或相等于 `21` 时应用于成员。

使用 `and` 操作符将你 `if` 语句的现有条件与另一个检查 `age` 是否大于或相等于 `21` 的条件结合起来。

# --hints--

你的第三个 `if` 语句条件应为 `is_member and age >= 21`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
