---
id: 6960b8632b9544d1f57295f4
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

年龄在 `18` 和 `21` 之间的用户可以预订票，但仅当 `show_time` 不是 `Evening` 时。

使用 `and` 操作符构建一个表达式，查看 `age` 是否大于或相等于 `18` 且 `show_time` 不等于 `Evening`。使用 `or` 操作符将其与最后一个 `if` 语句的现有条件组合。不要创建新的变量。 

# --hints--

你最后的 `if` 语句条件应为 `age >= 21 or age >= 18 and show_time != 'Evening'`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
_first = _cond.is_equivalent("age >= 21 or age >= 18 and show_time != 'Evening'")
_second = _cond.is_equivalent("age >= 21 or 18 <= age and show_time != 'Evening'")
assert _first or _second
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21:
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
