---
id: 696cb889fc4fdeb337e52864
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

在 Python 中，`if` 条件不必显式地将值与 `True` 或 `False` 进行比较。相反，它们依赖于值的真值性。真值性值在布尔上下文中（例如 `if` 语句的条件）会被评估为 `True`。这些包括非零数字、非空字串和布尔值 `True`。

```py
if value:
   print('value is truthy')
```

创建一个 `if` 语句来查看 `is_member` 是否为真。在 `if` 语句的体内，将 `discount` 的值更新为 `3` 并打印 `User qualifies for membership discount` 到终端。

# --hints--

你的代码中应该有三个 `if` 语句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

你的第三个 `if` 语句条件应为 `is_member`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

你应该在你的新建 `if` 语句中为 `discount` 变量赋值整数值 `3`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

你应该在你的新建 `if` 语句中打印 `User qualifies for membership discount`。记得用单引号或双引号将消息括起来。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
