---
id: 587d778e367417b2b2512aab
title: покращення читабельності за допомогою тексту з високим контрастом
challengeType: 0
videoUrl: 'https://scrimba.com/c/cKb3nCq'
forumTopicId: 301017
dashedName: improve-readability-with-high-contrast-text
---

# --description--

Низький контраст між кольорами переднього плану та фону може ускладнювати читання тексту. Достатній контраст покращує читабельність вашого вмісту, але що саме означає «достатній»?

Рекомендації з доступності веб-вмісту (WCAG) радять мати коефіцієнт контрастності не менше 4,5:1 для звичайного тексту. Цей коефіцієнт обчислюється шляхом порівняння відносних значень яскравості двох кольорів. Він варіюється від 1:1 для однакових кольорів, тобто без контрасту, до 21:1 для білого на чорному, що є найбільшим контрастом. В Інтернеті доступно багато інструментів для перевірки коефіцієнта контрастності, які обчислять його за вас.

# --instructions--

Вибір котом Кемпером світло-сірого тексту на білому фоні для його останнього блогу має коефіцієнт контрастності 1,5:1, що ускладнює читання. Змініть `color` тексту з поточного сірого (`#D3D3D3`) на темніший сірий (`#636363`), щоб покращити коефіцієнт контрастності до 6:1.

# --hints--

Ваш код має змінити `color` тексту для `body` на темніший сірий.

```js
const body = document.querySelector('body');
const bodyColor = window.getComputedStyle(body).color; 
assert(bodyColor == 'rgb(99, 99, 99)');
```

Ваш код не повинен змінювати `background-color` для `body`.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body).backgroundColor; 
assert.equal(backgroundColor , 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: #D3D3D3;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: #636363;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
