---
id: 587d7789367417b2b2512aa4
title: покращення доступності аудіоконтенту за допомогою елемента audio
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVkcZ'
forumTopicId: 301014
dashedName: improve-accessibility-of-audio-content-with-the-audio-element
---

# --description--

Елемент `audio` у HTML5 надає семантичне значення, коли він обгортає звуковий або аудіопотоковий вміст у вашій розмітці. Аудіоконтент також потребує текстової альтернативи, щоб бути доступним для людей з вадами слуху або глухих. Це можна зробити за допомогою тексту поруч на сторінці або посилання на транскрипт.

Тег `audio` підтримує атрибут `controls`. Це показує браузерні елементи керування за замовчуванням, такі як відтворення, пауза та інші, а також підтримує функціональність клавіатури. Це булевий атрибут, тобто йому не потрібно значення, його наявність у тегу вмикає цю опцію.

Ось приклад:

```html
<audio id="meowClip" controls>
  <source src="audio/meow.mp3" type="audio/mpeg">
  <source src="audio/meow.ogg" type="audio/ogg">
</audio>
```

**Примітка:** Мультимедійний вміст зазвичай має як візуальні, так і аудіальні компоненти. Для користувачів із порушеннями зору та/або слуху потрібні синхронізовані підписи та транскрипт. Зазвичай веброзробник не відповідає за створення підписів або транскрипту, але повинен знати, що їх потрібно включати.

# --instructions--

Час зробити перерву від кота Кемпера і познайомитися з іншим кемпером Zersiax (@zersiax), захисником доступності та користувачем читача екрана. Щоб почути фрагмент його читача екрана в дії, додайте елемент `audio` після елемента `p`. Включіть атрибут `controls`. Потім помістіть елемент `source` всередину тегів `audio` з атрибутом `src`, встановленим у `https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3`, і атрибутом `type`, встановленим у `"audio/mpeg"`.

**Примітка:** Аудіокліп може звучати швидко і бути складним для розуміння, але це нормальна швидкість для користувачів читачів екрана.

# --hints--

Ваш код повинен містити один тег `audio`.

```js
assert.lengthOf(document.querySelectorAll('audio'),1);
```

Ваш елемент `audio` повинен мати кінцевий тег.

```js
assert.match(code,/<audio.*>[\s\S]*<\/audio>/g);
assert.lengthOf(code.match(/<\/audio>/g),1);
```

Тег `audio` повинен мати атрибут `controls`.

```js
assert.exists(document.querySelector('audio')?.getAttribute('controls'));
```

Ваш код повинен містити один тег `source`.

```js
assert.lengthOf(document.querySelectorAll('source'), 1);
```

Ваш тег `source` повинен бути всередині тегів `audio`.

```js
const audio = document.querySelector('audio');
const children = audio.querySelectorAll(`:scope ${'source'}`);
assert.lengthOf(children,1);
```

Значення атрибута `src` у тегу `source` має точно відповідати посиланню в інструкціях.

```js
assert.equal(
  document.querySelector('source')?.getAttribute('src'),
  'https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3'
);
```

Ваш код повинен містити атрибут `type` у тегу `source` зі значенням audio/mpeg.

```js
assert.equal(document.querySelector('source')?.getAttribute('type'), 'audio/mpeg');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>



  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>
    <audio controls>
      <source src="https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3" type="audio/mpeg"/>
    </audio>
  </main>
</body>
```
