---
id: 587d778b367417b2b2512aa7
title: Обгорніть радіокнопки в елемент fieldset для кращої доступності
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVefw'
forumTopicId: 301030
dashedName: wrap-radio-buttons-in-a-fieldset-element-for-better-accessibility
---

# --description--

Наступна тема форми охоплює доступність радіокнопок. Кожному вибору присвоєно `label` з атрибутом `for`, який пов’язаний з `id` відповідного елемента, як було розглянуто в останньому завданні. Оскільки радіокнопки часто йдуть у групі, де користувач повинен вибрати один варіант, існує спосіб семантично показати, що вибори є частиною множини.

Тег `fieldset` охоплює всю групу радіокнопок, щоб досягти цього. Зазвичай він використовує тег `legend` для надання опису групи, який читачі екрана озвучують для кожного вибору в елементі `fieldset`.

Обгортка `fieldset` і тег `legend` не є необхідними, коли вибори зрозумілі самі по собі, наприклад, вибір статі. Використання `label` з атрибутом `for` для кожної радіокнопки є достатнім.

Ось приклад:

```html
<form>
  <fieldset>
    <legend>Choose one of these three items:</legend>
    <input id="one" type="radio" name="items" value="one">
    <label for="one">Choice One</label><br>
    <input id="two" type="radio" name="items" value="two">
    <label for="two">Choice Two</label><br>
    <input id="three" type="radio" name="items" value="three">
    <label for="three">Choice Three</label>
  </fieldset>
</form>
```

# --instructions--

Кіт Кемпер хоче отримати інформацію про рівень ніндзя своїх користувачів, коли вони підписуються на його список е-пошти. Він додав множину радіокнопок і дізнався з нашого останнього уроку, що потрібно використовувати теги `label` з атрибутами `for` для кожного вибору. Молодець, кіт Кемпер! Проте його код ще потребує деяких змін. Замініть тег `div`, який обгортає радіокнопки, на тег `fieldset`, а тег `p` всередині нього — на тег `legend`.

# --hints--

Ваш код має містити тег `fieldset`, який обгортає множину радіокнопок.

```js
assert.lengthOf(document.querySelectorAll('fieldset'),1);
```

Елемент `fieldset` повинен мати кінцевий тег.

```js
assert.isTrue(code.match(/<\/fieldset>/g)?.length === code.match(/<fieldset>/g)?.length);
```

Ваш код має містити тег `legend`, який обгортає текст із запитанням про рівень ніндзя користувача.

```js
assert.lengthOf(document.querySelectorAll('legend') ,1);
```

У вашому коді не повинно бути тегів `div`.

```js
assert.lengthOf(document.querySelectorAll('div') ,0);
```

У вашому коді більше не має бути тегу `p`, який обгортає текст із запитанням про рівень ніндзя користувача.

```js
assert.lengthOf(document.querySelectorAll('p') ,4);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <!-- Only change code below this line -->
      <div>
        <p>What level ninja are you?</p>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </div>
      <!-- Only change code above this line -->


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">

      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </fieldset>

      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
