---
id: 587d7790367417b2b2512ab1
title: Використання tabindex для визначення порядку фокусу клавіатури для кількох елементів
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzRRcb'
forumTopicId: 301028
dashedName: use-tabindex-to-specify-the-order-of-keyboard-focus-for-several-elements
---

# --description--

Атрибут `tabindex` також визначає точний порядок переходу табуляції між елементами. Це досягається, коли значення атрибута встановлено на додатне число 1 або більше.

Встановлення `tabindex="1"` призведе до того, що фокус клавіатури спочатку перейде на цей елемент. Потім він проходить послідовність заданих значень `tabindex` (2, 3 тощо), перш ніж перейти до елементів за замовчуванням і `tabindex="0"`.

Важливо зазначити, що коли порядок табуляції встановлено таким чином, він заміщує порядок за замовчуванням (який базується на HTML-коді). Це може збити з пантелику користувачів, які очікують почати навігацію з верхньої частини сторінки. Цю техніку іноді потрібно застосовувати, але з огляду на доступність, будьте обережні перед її використанням.

Ось приклад:

```html
<div tabindex="1">I get keyboard focus, and I get it first!</div>
```

```html
<div tabindex="2">I get keyboard focus, and I get it second!</div>
```

# --instructions--

Кіт Кемпер має поле пошуку на сторінці з надихаючими цитатами, яке він планує розмістити у верхньому правому куті за допомогою CSS. Він хоче, щоб елементи керування формою пошуку `input` і відправлення `input` були першими двома елементами в порядку табуляції. Додайте атрибут `tabindex` зі значенням `1` до `search` `input`, а атрибут `tabindex` зі значенням `2` до `submit` `input`.

Ще варто зазначити, що деякі браузери можуть розміщувати фокус посередині вашого порядку табуляції, коли ви натискаєте на елемент. На сторінку додано елемент, який гарантує, що ви завжди починатимете з початку вашого порядку табуляції.

# --hints--

Ваш код має додати атрибут `tabindex` до тегу `search` `input`.

```js
assert.isNotNull(document.querySelector('#search')?.getAttribute('tabindex'));
```

Ваш код має додати атрибут `tabindex` до тегу `submit` `input`.

```js
assert.isNotNull(document.querySelector('#submit')?.getAttribute('tabindex'));
```

Ваш код має встановити атрибут `tabindex` на тегу `search` `input` зі значенням 1.

```js
assert.equal(document.querySelector('#search')?.getAttribute('tabindex'), '1');
```

Ваш код має встановити атрибут `tabindex` на тегу `submit` `input` зі значенням 2.

```js
assert.equal(document.querySelector('#submit')?.getAttribute('tabindex'), '2');
```

# --seed--

## --seed-contents--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input type="search" name="search" id="search">
    <input type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```

# --solutions--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input tabindex="1" type="search" name="search" id="search">
    <input tabindex="2" type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```
