---
id: 587d774e367417b2b2512aa0
title: Обгорніть вміст у елемент article
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp79S3'
forumTopicId: 301029
dashedName: wrap-content-in-the-article-element
---

# --description--

`article` — це ще один із нових елементів HTML5, які додають семантичне значення вашій розмітці. `article` — це секційний елемент, який використовується для обгортання незалежного, самодостатнього вмісту. Цей тег добре підходить для блогових записів, повідомлень на форумах або новинних статей.

Визначити, чи може вміст існувати сам по собі, зазвичай — це справа оцінки, але ви можете скористатися кількома простими тестами. Запитайте себе: якщо ви видалите весь навколишній контекст, чи матиме цей вміст сенс? Аналогічно, для тексту — чи збережеться його зміст, якщо він буде у RSS-стрічці?

Пам’ятайте, що люди, які користуються допоміжними технологіями, покладаються на організовану, семантично значущу розмітку, щоб краще розуміти вашу роботу.

**Примітка:** Елемент `section` також новий у HTML5 і має трохи інше семантичне значення, ніж `article`. `article` призначений для самостійного вмісту, а `section` — для групування тематично пов’язаного вмісту. Їх можна використовувати один в одному за потреби. Наприклад, якщо книга — це `article`, тоді кожен розділ — це `section`. Коли між групами вмісту немає зв’язку, використовуйте `div`.

`<div>` — групує вміст  
`<section>` — групує пов’язаний вміст  
`<article>` — групує незалежний, самодостатній вміст

# --instructions--

Кіт Кемпер використав теги `article` для обгортання записів на своїй сторінці блогу, але забув застосувати їх навколо верхнього запису. Змініть тег `div` на тег `article`.

# --hints--

Ваш код має містити три теги `article`.

```js
assert.lengthOf(document.querySelectorAll('article') ,3);
```

У вашому коді не повинно бути жодних тегів `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <div>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </div>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```
