---
id: 587d7790367417b2b2512aaf
title: зробіть посилання навігаційними за допомогою HTML-атрибутів доступу
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQvmaTp'
forumTopicId: 301021
dashedName: make-links-navigable-with-html-access-keys
---

# --description--

HTML пропонує атрибут `accesskey` для вказівки клавіші швидкого доступу, щоб активувати або встановити фокус на елементі. Додавання атрибута `accesskey` може зробити навігацію більш ефективною для користувачів, які працюють лише з клавіатурою.

HTML5 дозволяє використовувати цей атрибут на будь-якому елементі, але він особливо корисний, коли застосовується до інтерактивних елементів. Це включає посилання, кнопки та елементи керування формою.

Ось приклад:

```html
<button accesskey="b">Important Button</button>
```

# --instructions--

Кіт Кемпер хоче, щоб посилання навколо двох заголовків статей у блозі мали клавіші швидкого доступу, щоб користувачі його сайту могли швидко переходити до повного тексту. Додайте атрибут `accesskey` до обох посилань і встановіть перший на `g` (для Гарфілда), а другий — на `c` (для Чака Норріса).

# --hints--

Ваш код має додати атрибут `accesskey` до тегу `a` зі значенням `id` `first`.

```js
assert.isTrue(document.querySelector('#first')?.hasAttribute('accesskey'));
```

Ваш код має додати атрибут `accesskey` до тегу `a` зі значенням `id` `second`.

```js
assert.isTrue(document.querySelector('#second')?.hasAttribute('accesskey'));
```

Ваш код має встановити атрибут `accesskey` на тег `a` зі значенням `id` `first` на `g`. Зверніть увагу, що регістр має значення.

```js
assert.equal(document.querySelector('#first')?.getAttribute('accesskey'), 'g');
```

Ваш код має встановити атрибут `accesskey` на тег `a` зі значенням `id` `second` на `c`. Зверніть увагу, що регістр має значення.

```js
assert.equal(document.querySelector('#second')?.getAttribute('accesskey'), 'c');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" accesskey="g" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" accesskey="c" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
