---
id: 587d778f367417b2b2512aae
title: Надати посиланням значення за допомогою описового тексту посилання
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437DcV'
forumTopicId: 301013
dashedName: give-links-meaning-by-using-descriptive-link-text
---

# --description--

Користувачі читачів екрана мають різні варіанти того, який тип вмісту їхній пристрій читає. Ці варіанти включають пропуск орієнтирних елементів, перехід до основного вмісту або отримання зведення сторінки з заголовків. Ще один варіант — чути лише посилання, доступні на сторінці.

Читачі екрана роблять це, читаючи текст посилання, або те, що знаходиться між тегами якоря (`a`). Список посилань на кшталт «натисніть тут» або «читати далі» не є корисним. Замість цього використовуйте короткий, але описовий текст у тегах `a`, щоб надати більше значення цим користувачам.

# --instructions--

Текст посилання, який використовує кіт Кемпер, не дуже описовий без навколишнього контексту. Перемістіть теги якоря (`a`), щоб вони охоплювали текст `information about batteries` замість `Click here`.

# --hints--

Ваш код має перемістити теги якоря `a` з навколо слів `Click here` так, щоб вони охоплювали слова `information about batteries`.

```js
assert.match(document.querySelector('a')?.textContent, /^(information about batteries)$/g);
```

Елемент `a` має мати атрибут `href` зі значенням порожнього рядка `""`.

```js
assert.isEmpty(document.querySelector('a')?.getAttribute('href'));
```

Елемент `a` має мати кінцевий тег.

```js
assert.isTrue(code.match(/<\/a>/g)?.length === code.match(/<a href=(''|"")>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. <a href="">Click here</a> for information about batteries</p>
  </article>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. Click here for <a href="">information about batteries</a></p>
  </article>
</body>
```
