---
id: 587d778b367417b2b2512aa8
title: додати доступний вибір дати
challengeType: 0
videoUrl: 'https://scrimba.com/c/cR3bRbCV'
forumTopicId: 301008
dashedName: add-an-accessible-date-picker
---

# --description--

Форми часто містять поле `input`, яке можна використати для створення кількох різних елементів керування формою. Атрибут `type` цього елемента вказує, який тип елемента `input` буде створено.

Ви могли помітити типи вхідних даних `text` і `submit` у попередніх завданнях, а HTML5 додав можливість вказати поле `date`. Залежно від підтримки браузера, у полі `input` з’являється вибір дати, коли воно у фокусі, що полегшує заповнення форми для всіх користувачів.

Для старіших браузерів тип за замовчуванням буде `text`, тому корисно показати користувачам очікуваний формат дати у тексті `label` або `placeholder` на всяк випадок.

Ось приклад:

```html
<label for="input1">Enter a date:</label>
<input type="date" id="input1" name="input1">
```

# --instructions--

Кіт Кемпер організовує турнір Mortal Kombat і хоче запитати учасників, яка дата їм підходить найкраще. Додайте тег `input` з атрибутом `type` зі значенням `date`, атрибутом `id` зі значенням `pickdate` і атрибутом `name` зі значенням `date`.

# --hints--

Ваш код має додати один тег `input` для поля вибору дати.

```js
assert.lengthOf(document.querySelectorAll('input'), 2);
```

Ваш тег `input` має мати атрибут `type` зі значенням `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('type'), 'date');
```

Ваш тег `input` має мати атрибут `id` зі значенням `pickdate`.

```js
assert.equal(document.querySelector('input')?.getAttribute('id'),'pickdate');
```

Ваш тег `input` має мати атрибут `name` зі значенням `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('name'), 'date');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>

        <!-- Only change code below this line -->



        <!-- Only change code above this line -->

        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>
        <input type="date" id="pickdate" name="date">
        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
