---
id: 587d774c367417b2b2512a9d
title: "Знати, коли alt-текст слід залишати порожнім"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

У попередньому завданні ви дізналися, що додавання атрибута `alt` при використанні тегів `img` є обов’язковим. Однак іноді зображення згруповані з підписом, який їх уже описує, або використовуються лише для декоративних цілей. У таких випадках `alt` текст може здаватися зайвим або непотрібним.

Коли зображення вже пояснене текстовим вмістом або не додає значення сторінці, `img` все одно потребує атрибута `alt`, але його можна встановити як порожній рядок. Ось приклад:

```html
<img src="visualDecoration.jpeg" alt="">
```

Фонові зображення зазвичай також відносять до категорії «декоративних». Проте вони зазвичай застосовуються за допомогою CSS-правил і тому не є частиною розмітки, яку обробляють читачі екрана.

**Примітка:** Для зображень із підписом ви все одно можете додати `alt` текст, оскільки це допомагає пошуковим системам каталогізувати вміст зображення.

# --instructions--

Кіт Кемпер створив каркас сторінки для блогу свого сайту. Він планує додати візуальний розрив між двома своїми статтями за допомогою декоративного зображення самурайського меча. Додайте атрибут `alt` до тегу `img` і встановіть його як порожній рядок. (Зверніть увагу, що `src` зображення не посилається на реальний файл — не хвилюйтеся, що на екрані не видно мечів.)

# --hints--

Ваш тег `img` має містити атрибут `alt`.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

Атрибут `alt` слід встановити як порожній рядок.

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```
