---
id: 587d774c367417b2b2512a9c
title: Додайте текстову альтернативу до зображень для доступності при порушеннях зору
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

Ви, ймовірно, бачили атрибут `alt` у тегу `img` в інших завданнях. Текст `alt` описує вміст зображення і надає для нього текстову альтернативу. Атрибут `alt` допомагає у випадках, коли зображення не завантажується або його не може побачити користувач. Пошукові системи також використовують його, щоб зрозуміти, що містить зображення, і включити його у результати пошуку. Ось приклад:

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

Люди з порушеннями зору покладаються на читачі екрана, які перетворюють веб-вміст у аудіоінтерфейс. Вони не отримають інформацію, якщо вона подана лише візуально. Для зображень читачі екрана можуть отримати доступ до атрибута `alt` і прочитати його вміст, щоб передати ключову інформацію.

Добрий текст `alt` надає читачеві короткий опис зображення. Ви завжди повинні додавати атрибут `alt` до свого зображення. Згідно зі специфікацією HTML5, це тепер вважається обов’язковим.

# --instructions--

Кіт Кемпер — це і кодер-ніндзя, і справжній ніндзя, який створює вебсайт, щоб ділитися своїми знаннями. Фото профілю, яке він хоче використати, показує його навички і має бути оцінене всіма відвідувачами сайту. Додайте атрибут `alt` у тег `img`, який пояснює, що кіт Кемпер займається карате. (Посилання на файл зображення `src` не веде до реального файлу, тому ви повинні побачити текст `alt` на екрані.)

# --hints--

Ваш тег `img` повинен мати атрибут `alt`, і він не має бути порожнім.

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">
```
