---
id: 587d7787367417b2b2512aa1
title: Спрощення навігації для читачів екрана за допомогою орієнтира header
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB76vtv'
forumTopicId: 301023
dashedName: make-screen-reader-navigation-easier-with-the-header-landmark
---

# --description--

Наступний HTML5 елемент, який додає семантичне значення і покращує доступність, — це тег `header`. Його використовують для обгортання вступної інформації або навігаційних посилань для батьківського елемента, і він добре підходить для вмісту, який повторюється у верхній частині на кількох сторінках.

`header` має функцію вставленого орієнтира, яку ви бачили у `main`, що дозволяє допоміжним технологіям швидко переходити до цього вмісту.

**Примітка:** `header` призначений для використання у `body` тегу вашого HTML-документа. Він відрізняється від елемента `head`, який містить заголовок сторінки, метаінформацію тощо.

# --instructions--

Кіт Кемпер пише чудові статті про тренування ніндзя і хоче додати сторінку для них на свій сайт. Змініть верхній `div`, який зараз містить `h1`, на тег `header`.

# --hints--

Ваш код повинен містити один тег `header`.

```js
assert.lengthOf(document.querySelectorAll('header'),1);
```

Ваші теги `header` повинні обгортати `h1`.

```js
const header = document.querySelector('header');
const children = header?.querySelectorAll(`:scope ${'h1'}`);
assert.lengthOf(children , 1);
```

У вашому коді не повинно бути тегів `div`.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

Ваш елемент `header` повинен мати кінцевий тег.

```js
assert.isTrue(code.match(/<\/header>/g)?.length === code.match(/<header>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>

  <div>
    <h1>Training with Camper Cat</h1>
  </div>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>

  <header>
    <h1>Training with Camper Cat</h1>
  </header>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
