---
id: 587d778f367417b2b2512aac
title: "Уникайте проблем із порушенням сприйняття кольору, використовуючи достатній контраст"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMEUw'
forumTopicId: 301012
dashedName: avoid-colorblindness-issues-by-using-sufficient-contrast
---

# --description--

Колір є важливою частиною візуального дизайну, але його використання створює дві проблеми з доступністю. По-перше, колір не повинен бути єдиним способом передачі важливої інформації, оскільки користувачі читачів екрана його не побачать. По-друге, кольори переднього плану та фону повинні мати достатній контраст, щоб користувачі з порушенням сприйняття кольору могли їх розрізняти.

Попередні завдання охоплювали наявність текстових альтернатив для вирішення першої проблеми. Останнє завдання познайомило з інструментами перевірки контрасту, які допомагають із другою. Рекомендований WCAG коефіцієнт контрастності 4.5:1 застосовується як для використання кольору, так і для комбінацій у відтінках сірого.

Користувачам з порушенням сприйняття кольору важко розрізняти деякі кольори — зазвичай за відтінком, але іноді й за світлістю. Ви, можливо, пам’ятаєте, що коефіцієнт контрастності обчислюється за відносною світністю (або світлістю) кольорів переднього плану та фону.

На практиці коефіцієнт контрастності 4.5:1 можна досягти, затемнюючи (додаючи чорний) темніший колір і освітлюючи (додаючи білий) світліший колір. Темніші відтінки на колірному колі вважаються відтінками синього, фіолетового, мадженти та червоного, тоді як світліші освітлені кольори — це помаранчевий, жовтий, зелений і синьо-зелений.

# --instructions--

Кіт Кемпер експериментує з використанням кольору для тексту та фону свого блогу, але його поточне поєднання зеленуватого `background-color` з маруновим текстом `color` має коефіцієнт контрастності 2.5:1. Ви можете легко відрегулювати світлість кольорів, оскільки він оголосив їх за допомогою CSS-властивості `hsl()` (яка означає відтінок, насиченість, світлість), змінивши третій аргумент. Збільшіть значення світлості `background-color` з 35% до 55%, а значення світлості `color` зменшіть з 20% до 15%. Це покращить контраст до 5.9:1.

# --hints--

Ваш код повинен змінити лише значення світлості для властивості тексту `color` на 15%.

```js
assert.match(code ,/color:\s*?hsl\(0,\s*?55%,\s*?15%\)/gi);
```

Ваш код повинен змінити лише значення світлості для властивості `background-color` на 55%.

```js
assert.match(code ,/background-color:\s*?hsl\(120,\s*?25%,\s*?55%\)/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 20%);
    background-color: hsl(120, 25%, 35%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 15%);
    background-color: hsl(120, 25%, 55%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
