---
id: 587d7790367417b2b2512ab0
title: Використання tabindex для додавання клавіатурного фокусу на елемент
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMDHW'
forumTopicId: 301027
dashedName: use-tabindex-to-add-keyboard-focus-to-an-element
---

# --description--

HTML-атрибут `tabindex` має три різні функції, пов’язані з клавіатурним фокусом елемента. Коли він застосований до тегу, це означає, що елемент може отримувати фокус. Значення (ціле число, яке може бути додатним, від’ємним або нулем) визначає поведінку.

Деякі елементи, такі як посилання та елементи керування формою, автоматично отримують клавіатурний фокус, коли користувач переходить між елементами сторінки за допомогою клавіші Tab. Порядок фокусу відповідає порядку елементів у HTML-розмітці. Таку саму функціональність можна надати іншим елементам, наприклад `div`, `span` та `p`, додавши до них атрибут `tabindex="0"`. Ось приклад:

```html
<div tabindex="0">I need keyboard focus!</div>
```

**Примітка:** Від’ємне значення `tabindex` (зазвичай -1) означає, що елемент може отримувати фокус, але до нього неможливо дістатися за допомогою клавіатури. Цей метод зазвичай використовується для програмного встановлення фокусу на вміст (наприклад, коли активується `div` для спливаючого вікна) і виходить за межі цих завдань.

# --instructions--

Кіт Кемпер створив нове опитування для збору інформації про своїх користувачів. Він знає, що поля введення автоматично отримують клавіатурний фокус, але хоче переконатися, що користувачі клавіатури зупиняються на інструкціях під час переходу між елементами за допомогою Tab. Додайте атрибут `tabindex` до тегу `p` і встановіть його значення `0`. Бонус — використання `tabindex` також дозволяє працювати CSS-псевдокласу `:focus` на тегу `p`.

# --hints--

Ваш код має додати атрибут `tabindex` до тегу `p`, який містить інструкції форми.

```js
assert.isNotNull(document.querySelector('p')?.getAttribute('tabindex'));
```

Ваш код має встановити атрибут `tabindex` на тегу `p` зі значенням 0.

```js
assert.equal(document.querySelector('p')?.getAttribute('tabindex'), '0');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p>Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p tabindex="0">Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
