---
id: 587d778a367417b2b2512aa6
title: покращення доступності поля форми за допомогою елемента label
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMMAN'
forumTopicId: 301016
dashedName: improve-form-field-accessibility-with-the-label-element
---

# --description--

Покращення доступності за допомогою семантичної HTML-розмітки стосується використання як відповідних назв тегів, так і атрибутів. Наступні кілька завдань охоплюють важливі сценарії використання атрибутів у формах.

Тег `label` обгортає текст для конкретного елемента керування формою, зазвичай це ім’я або мітка для вибору. Це надає значення елементу і робить форму більш зрозумілою. Атрибут `for` на тегу `label` явно пов’язує цей `label` з елементом керування формою і використовується читачами екрана.

Ви вивчали радіокнопки та їхні мітки в уроці з розділу Basic HTML. У тому уроці ми обгорнули елемент вводу радіокнопки всередині елемента `label` разом із текстом мітки, щоб зробити текст активним. Інший спосіб досягти цього — використати атрибут `for`, як пояснено в цьому уроці.

Значення атрибута `for` має співпадати зі значенням атрибута `id` елемента керування формою. Ось приклад:

```html
<form>
  <label for="name">Name:</label>
  <input type="text" id="name" name="name">
</form>
```

# --instructions--

Кіт Кемпер очікує багато зацікавленості до своїх продуманих блог-постів і хоче додати форму підписки на е-пошту. Додайте атрибут `for` до елемента е-пошти `label`, який співпадає з `id` на його полі `input`.

# --hints--

Ваш код має містити атрибут `for` на тегу `label`, який не є порожнім.

```js
assert.isNotEmpty(document.querySelector('label')?.getAttribute('for'));
```

Значення вашого атрибута `for` має співпадати зі значенням `id` на елементі е-пошти `input`.

```js
assert.equal(document.querySelector('label')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label>Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
