---
id: 660ef31a5be625914a0102cd
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Коли змінну оголошено за допомогою ключового слова `let`, ви можете <dfn>перепризначити</dfn> (або змінити значення) цю змінну пізніше. У цьому прикладі значення `programmer` змінено з `"Naomi"` на `"CamperChan"`.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Зауважте, що під час перепризначення змінної ви **не** використаєте знову ключове слово `let`.

Після `console.log` призначте значення `"World"` до змінної `character`.

# --hints--

Код має використовувати `character` двічі.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

Використайте оператор присвоєння, щоб перепризначити `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

Код має використовувати рядок `"World"`.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Змінна `character` повинна мати значення `"World"`.

```js
assert.equal(character, "World");
```

Перепризначення не повинно використовувати `let`.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
