---
id: 660ef0f7c4b8e68ccd1f0786
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

У JavaScript є сім примітивних типів даних, одним із яких є String. У JavaScript <dfn>рядок</dfn> представляє послідовність символів і може бути оточений в одинарні (`'`) або подвійні (`"`) лапки.

Зауважте, що рядки є <dfn>незмінними</dfn>, а це означає, що після створення вони не можуть бути змінені. Змінній все ще можна перепризначити інше значення.

Змініть рядок `"Hello"` використати одинарні лапки.

# --hints--

Код не повинен містити подвійних лапок.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

Використайте одинарні лапки для рядка `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

Код все ще повинен використати `let`.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Все ще оголошуйте змінну `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Змінна `character` все ще повинна мати рядок `"Hello"` як значення.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
