---
id: 660eebd83100d37862268781
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Змінна `character` зараз не має значення. Призначте значення за допомогою оператора <dfn>присвоєння</dfn> `=`. Наприклад:

```js
let hello = "Hello";
```

Призначення значення змінній у момент її оголошення відоме як <dfn>ініціалізація</dfn>.

Ініціалізуйте змінну `character`, призначивши до неї значення `"Hello"` під час її оголошення.

# --hints--

Ви повинні використати оператор присвоювання `=`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Використайте рядок `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Використайте подвійні лапки для рядка `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Призначте `"Hello"` до змінної `character`.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
