---
id: 660ef55dd468079679ee0092
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Тепер ви готові оголосити свою наступну змінну. Вилучіть інструкції `console.log` та перепризначення `character`.

Також видаліть змінну `secondCharacter`, але залиште ініціалізацію `character` незміненою.

# --hints--

Код не повинен містити жодних інструкцій `console.log()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Не перепризначайте значення `character`.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

Змінна `character` повинна мати значення `"Hello"`.

```js
assert.equal(character, "Hello");
```

Ви не повинні мати змінну `secondCharacter`.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
