---
id: 660ef5c1904955978a986a5c
title: Крок 15
challengeType: 1
dashedName: step-15
---

# --description--

Використайте `let`, щоб оголосити змінну `count`. Призначте до неї <dfn>число</dfn> `8`. Коли використається числове значення, лапки не використовуються. Наприклад:

```js
let money = 100;
```

# --hints--

Код повинен мати друге ключове слово `let`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Використайте `let`, щоб оголосити змінну `count`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

Призначте число `8` до змінної `count`.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
