---
id: 660ef857f2806aa626d29d17
title: Крок 17
challengeType: 1
dashedName: step-17
---

# --description--

Ви також можете виконувати віднімання (`-`), множення (`*`) та ділення (`/`). Не соромтеся експериментувати з операторами та числами у `console.log`. Коли ви будете готові рухатися далі, видаліть `console.log`.

# --hints--

Код не повинен містити `console.log()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Змінна `character` все ще повинна мати значення `"Hello"`.

```js
assert.equal(character, "Hello");
```

Змінна `count` все ще повинна мати значення `8`.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
