---
id: 660ee6e3a242da6bd579de69
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

One of the most important concepts in programming is variables. <dfn>Змінна</dfn> вказує на конкретну адресу пам’яті, яка зберігає значення. Variables are given a name which can be used throughout your code to access that value.

Declaring a variable means giving it a name. In JavaScript, this is often done with the `let` keyword. Наприклад, ось так можна оголосити змінну `hello`:

```js
let hello;
```

Variable naming follows specific rules: names can include letters, numbers, dollar signs, and underscores, but cannot contain spaces and must not begin with a number.

Використайте ключове слово `let`, щоб оголосити змінну під назвою `character`.

_Note_: It is common practice to end statements in JavaScript with a semicolon. `;`

# --hints--

Використайте `let` в коді.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Використайте `character` в коді.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

Використайте `let`, щоб оголосити змінну `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Ваше оголошення має закінчуватися крапкою з комою.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
