---
id: 660f4de78f775e480ba2e451
title: Крок 107
challengeType: 1
dashedName: step-107
---

# --description--

Метод масиву `.unshift()` дозволяє додавати значення до **початку** масиву, на відміну від `.push()`, який додає значення в кінець масиву. `.unshift()` повертає нову довжину масиву, на якому було викликано.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

Використовуйте `const`, щоб оголосити змінну `unshifted` та призначте до неї результат виклику `.unshift()` на масиві `numbers`. Pass `5` as the argument. Then print your `unshifted` variable.

# --hints--

You should declare your `unshifted` variable.

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

You should call the `.unshift()` method on your `numbers` array.

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

You should pass `5` as the argument to your `.unshift()` call.

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

Призначте результат виклику `.unshift()` до змінної `unshifted`.

```js
assert.equal(unshifted, 4);
```

Зареєструйте змінну `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
