---
id: 660f23b53db70af0f2620e78
title: Крок 43
challengeType: 1
dashedName: step-43
---

# --description--

Тепер усі числа відображаються в одному рядку. Це не спрацює для створення піраміди.

Потрібно буде додати новий рядок до кожного рядка. Однак натискання клавіші повернення для вставлення розриву рядка між лапками у JavaScript призведе до помилки розбору. Замість, потрібно використовувати спеціальну <dfn>послідовність що керує</dfn> `\n`, яка інтерпретується як новий рядок, коли рядок записується в журнал. Наприклад:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Використовуйте другий оператор додавання, щоб об’єднати новий рядок між наявним значенням `result` та доданим значенням `row`.

# --hints--

Використовуйте послідовність що керує `\n`. Пам’ятайте, що це має бути рядок, тому він узятий у лапки.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

Об’єднайте послідовність що керує `\n` до змінної `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*('|")\\n\1/);
```

Об’єднайте змінну `row` до послідовності що керує `\n`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*('|")\\n\1\s*\+\s*row/);
```

Призначте усю конкатенацію назад до змінної `result`. Не забудьте про крапку з комою.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*result\s*\+\s*('|")\\n\1\s*\+\s*row;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
