---
id: 660f0a55847d6cc485f29ba5
title: Крок 26
challengeType: 1
dashedName: step-26
---

# --description--

<dfn>Метод</dfn> у JavaScript — це функція, пов’язана з певними значеннями чи об’єктами. Прикладом, з яким ви вже стикалися, є метод `.log()`, який є частиною об’єкта `console`.

Масиви мають власні методи, і перше, що ви дослідите, це метод `.push()`. Це дозволяє вам "проштовхнути" значення в кінець масиву. Ось приклад додавання числа `12` у кінець масиву:

```js
array.push(12);
```

Використайте `.push()`, щоб додати рядок `"freeCodeCamp"` у кінець масиву `rows`. Додайте цей код перед `console.log`, щоб ви могли побачити зміни, які ви внесли у масив.

# --hints--

Код повинен використати `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

Використайте метод `.push()` на масиві `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Передайте рядок `"freeCodeCamp"` у метод `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
