---
id: 660f17294346b7d69e79db3d
title: Крок 32
challengeType: 1
dashedName: step-32
---

# --description--

Тепер ви готові почати будувати генератор піраміди. Змінна `character` слугуватиме будівельним блоком для піраміди.

`"Hello"` може не працювати дуже добре для цього. Змініть значення `character` на символ ґрат `"#"`.

# --hints--

Змінна `character` має бути символом ґрат.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
