---
id: 660f07d231941bc11719f664
title: Крок 21
challengeType: 1
dashedName: step-21
---

# --description--

Особливість масивів полягає в тому, що вони вважаються <dfn>змінюваними</dfn>. Це означає, що ви безпосередньо можете змінити значення за індексом.

Наприклад, цей код призначить число `25` другому елементу в масиві:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

Оновіть **третій** елемент масиву `rows` до числа `10`. Потім надрукуйте масив `rows` на консоль.

# --hints--

Знову використайте дужкову нотацію на масиві `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

Використайте третій елемент масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

Використайте інструкцію присвоєння на третьому елементі масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

Призначте значення `10` третьому елементу масиву `rows`.

```js
assert.equal(rows[2], 10);
```

Код повинен мати другу інструкцію `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Зареєструйте масив `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
