---
id: 660f3ba3cceef11b6ba08b59
title: Крок 75
challengeType: 1
dashedName: step-75
---

# --description--

На жаль, нині дно піраміди зникло. Це тому, що ви створили ще одну <dfn>помилку на одиницю</dfn>.

Ваш оригінальний цикл йшов по значеннях `i` від `0` до `7`, оскільки `count` дорівнює `8` і ваша умова вимагає, щоб `i` був меншим за `count`. Зараз цикл виконується для значень `i` від `1` до `7`.

Цикл також потрібно оновити, щоб виконуватися, коли `i` є `8`. Дивлячись на вашу логіку, тобто цикл має запускатися, коли `i` <dfn>менше або дорівнює</dfn> `count`. Для цього можна використовувати оператор менше або дорівнює `<=`.

Оновіть умову циклу, щоб виконуватися, коли `i` менше або дорівнює `count`.

# --hints--

Цикл `for` не повинен перевіряти, чи `i` менше ніж `count`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Цикл `for` має перевіряти, чи `i` менше або дорівнює `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
