---
id: 660f45ccf4ca5c31f253005a
title: Крок 86
challengeType: 1
dashedName: step-86
---

# --description--

Прямо зараз, якщо ви зміните `continueLoop` на true, цикл `while` працюватиме вічно. Це називається <dfn>нескінченним циклом</dfn>, і ви мусите бути обережним з цим. Нескінченний цикл може заблокувати вашу систему, вимагаючи повного перезавантаження, щоб вийти.

Щоб уникнути цього, почніть із використання оператора збільшення, щоб збільшити значення змінної `done` у циклі.

# --hints--

Цикл `while` має збільшити змінну `done`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
