---
id: 660f447efc0e722f016c1be0
title: Крок 81
challengeType: 1
dashedName: step-81
---

# --description--

Текст знову з’явився! Це пояснюється тим, що `"false"` є рядком, який коли оцінюється до булевого стає `true`. Це означає, що `"false"` є правдивим значенням.

<dfn>Істинне значення</dfn> — це значення, яке вважається істинним, коли його оцінюють як булеве. Більшість значень, які ви зустрічаєте в JavaScript, будуть правдивими.

<dfn>Хибне значення</dfn> є протилежним — значення вважається хибним коли оцінюється як булеве. JavaScript має визначений список хибних значень. Деякі з них включають `false`, `0`, `""`, `null`, `undefined`, та `NaN`.

Спробуйте змінити умову `if` на порожній рядок `""`, який є хибним значенням.

# --hints--

Інструкція `if` має містити `""` як умову.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
