---
id: 660f039ff313dbb696b007ca
title: Крок 19
challengeType: 1
dashedName: step-19
---

# --description--

Коли масив тримає значення або <dfn>елементи</dfn>, ці значення розділяються комами. Ось масив, який тримає два рядки:

```js
let array = ["first", "second"];
```

Змініть оголошення `rows` на масив із рядками `"Naomi"`, `"Quincy"` та `"CamperChan"`. Порядок значень у масиві важливий, тому дотримуйтеся цього порядку. Пам’ятайте, що рядки чутливі до регістру.

# --hints--

Першим елементом у масиві має бути рядок `"Naomi"`.

```js
assert.equal(rows[0], "Naomi");
```

Другим елементом у масиві має бути рядок `"Quincy"`.

```js
assert.equal(rows[1], "Quincy");
```

Третім елементом у масиві має бути рядок `"CamperChan"`.

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
