---
id: 660f0da9bf1035c9097af20a
title: Крок 28
challengeType: 1
dashedName: step-28
---

# --description--

Ви повинні були побачити `"freeCodeCamp"`, надруковане на консолі. Це тому, що `.pop()` повертає значення, яке було вилучено з масиву, а ви раніше вставили `"freeCodeCamp"` у кінець масиву.

Але що повертає `.push()`? Призначте наявний `rows.push()` до нової змінної `pushed` та зареєструйте це.

# --hints--

Оголосіть змінну `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

Використайте `let`, щоб оголосити змінну `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

Призначте `rows.push("freeCodeCamp")` до змінної `pushed`.

```js
assert.equal(pushed, 4);
```

Зареєструйте змінну `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
