---
id: 660f18f059fe0fda192ce394
title: Крок 34
challengeType: 1
dashedName: step-34
---

# --description--

Циклу тепер потрібен належний ітератор. <dfn>Ітератор</dfn> — це змінна, яку можна оголосити спеціально у циклі `for`, щоб контролювати, як цикл ітерує чи іде через вашу логіку.

Загальною конвенцією є використання `i` як змінної ітератора в циклі. Цикл `for` дозволяє оголосити це в круглих дужках `()`. Наприклад, ось цикл `for`, який оголошує змінну `index` та призначає до неї значення `100`.

```js
for (let index = 100; "second"; "third") {

}
```

Замініть рядок `"iterator"` на оголошення `let` для змінної `i`. Для початку призначте до неї значення `0`. Це надасть змінній `i` значення `0` під час **першого** запуску циклу.

# --hints--

Використайте `let`, щоб оголосити змінну `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

Призначте `0` до змінної `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

Цикл `for` має починати ітератор `i` з `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
