---
id: 660f415b76859a2736771607
title: Крок 78
challengeType: 1
dashedName: step-78
---

# --description--

Піраміда знову зникла. Це нормально — цього й слід було очікувати.

Перш ніж створити новий цикл, вам потрібно дізнатися про інструкції `if`. <dfn>Інструкція `if`</dfn> дозволяє запускати блок коду лише тоді, коли виконується умова. Вона використає наступний синтаксис:

```js
if (condition) {
  logic
}
```

Створіть інструкцію `if` із логічним `true` як умовою. У тілі надрукуйте рядок `"Condition is true"`.

# --hints--

Створіть інструкцію `if`.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

Інструкція `if` має містити `true` як умову.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

Тіло `if` має реєструвати `"Condition is true"`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
