---
id: 660f38c34a4de6141c0c369f
title: Крок 71
challengeType: 1
dashedName: step-71
---

# --description--

Ваш генератор пірамід тепер працює як очікувано. Але це чудова можливість глибше дослідити код, який ви написали.

Оператор додавання — не єдиний спосіб додати значення до змінної. Оператор <dfn>присвоєння додавання</dfn> можна використовувати як скорочення, щоб означати «взяти оригінальне значення змінної, додати це значення та призначити результат назад змінній.» Наприклад, ці дві інструкції дадуть однаковий результат:

```js
test = test + 1;
test += 1;
```

Update your iteration statement in the `for` loop to use addition assignment.

# --hints--

Цикл `for` не повинен використати `i = i + 1`;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Цикл `for` має використовувати призначення додавання з `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
