---
id: 660f2a70ad6225fa503e71c3
title: Крок 46
challengeType: 1
dashedName: step-46
---

# --description--

Ви наближаєтеся! У цей момент ви зіткнулися з так званою <dfn>помилкою на одиницю</dfn>, частою проблемою в мовах з індексуванням на базі нуля, як-от JavaScript.

Перший індекс масиву `rows` є `0`, тому ви починаєте цикл `for` із `i = 0`. Але повторення рядка нуль разів призводить до того, що нічого не треба друкувати.

Щоб виправити це, додайте `1` до значення `i` у виклику `.repeat()`. Не призначайте його назад до `i`, як це було в умовах циклу.

# --hints--

Додайте `1` `i` у методі `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
