---
id: 660f255022991ef34ed0ee88
title: Крок 44
challengeType: 1
dashedName: step-44
---

# --description--

Друк чисел не призведе до візуально привабливої ​​піраміди. Тепер, коли ви виводите відформатований вміст масиву `rows`, настав час оновити оригінальний цикл.

Замість того, щоб штовхати `i` до масиву, штовхніть значення змінної `character`.

# --hints--

Більш не штовхайте змінну `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Штовхніть змінну `character`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
