---
id: 660f23b53db70af0f2620e78
title: Крок 43
challengeType: 1
dashedName: step-43
---

# --description--

Тепер усі числа відображаються в одному рядку. Це не спрацює для створення піраміди.

Потрібно буде додати новий рядок до кожного рядка. Однак натискання клавіші повернення для вставлення розриву рядка між лапками у JavaScript призведе до помилки розбору. Замість, потрібно використовувати спеціальну <dfn>послідовність що керує</dfn> `\n`, яка інтерпретується як новий рядок, коли рядок записується в журнал. Наприклад:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Use a second addition operator to append a new line after the `result` and `row` values.

# --hints--

Використовуйте послідовність що керує `\n`. Пам’ятайте, що це має бути рядок, тому він узятий у лапки.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

You should concatenate your `row` variable to your `result` variable.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

You should concatenate the `\n` escape sequence to your `row` variable.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

You should assign the entire concatenation back to your `result` variable.

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
