---
id: 660f2fbd45b520046cac68e8
title: Крок 65
challengeType: 1
dashedName: step-65
---

# --description--

Пам’ятайте, що на попередньому кроці ви дізналися про значення, що повертаються. Функція може <dfn>повернути</dfn> значення, для вашої програми, щоб вона споживала його окремо.

У функції ключове слово `return` використається для визначення значення, яке повертається. Наприклад, ця функція буде повертати значення, надане першому параметру:

```js
function name(parameter) {
  return parameter;
}
```

Використайте ключове слово `return`, щоб повернути значення змінної `character`, повторене `rowNumber` разів.

# --hints--

Використайте метод `.repeat()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

Використайте метод `.repeat()` на змінній `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

Передайте `rowNumber` до виклику `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

Використайте ключове слово `return`.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

Поверніть результат виклику `.repeat()`.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
