---
id: 660f0f980e98e8cf77f1ce31
title: Крок 30
challengeType: 1
dashedName: step-30
---

# --description--

Змініть декларацію `rows`, щоб знову призначити порожній масив.

Також змініть рядок `'Hello'`, щоб знову використати подвійні лапки. Взагалі, неважливо, якому з двох ви віддаєте перевагу, але ви захочете бути послідовним у цьому виборі протягом усього проєкту.

# --hints--

Масив `rows` має бути порожнім.

```js
assert.empty(rows);
```

Рядок `"Hello"` має використовувати подвійні лапки.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
