---
id: 660f3ce51f70571e1c5227c8
title: Крок 76
challengeType: 1
dashedName: step-76
---

# --description--

Коментарі можуть бути корисними, щоб пояснити, чому код використає певний підхід, або залишити нотатки щодо справ яки було б непогано зробити для вас майбутнього.

У JavaScript ви можете використати `//`, щоб залишити однорядковий коментар у коді.

Add a single-line comment above your `for` loop to remind yourself to change the code to a different kind of loop.

# --hints--

Почніть однорядковий коментар з `//`.

```js
assert.match(code, /\/\//);
```

Ваш однорядковий коментар має містити принаймні п’ять символів.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
