---
id: 660f4f79e2a82a4e92290f44
title: Крок 110
challengeType: 1
dashedName: step-110
---

# --description--

Іноді вам може знадобитися повернути попередній код, який ви закоментували. Можна це зробити, видаливши `/*` та `*/` навколо цього коду. Це називається <dfn>розкоментуванням</dfn>.

Розкоментуйте лише перший цикл `for`. Залиште однорядковий коментар та два інші багаторядкові коментарі на місці.

# --hints--

Розкоментуйте перший цикл `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

Не видаляйте однорядковий коментар.

```js
assert.match(code, /\/\/\sTODO:/);
```

Не розкоментовуйте цикл `while`.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Не розкоментовуйте другий цикл `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
