---
id: 660f3dd626be3a1ffe27e5d1
title: Крок 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript також підтримує багаторядкові коментарі. Багаторядковий коментар починається з `/*` та закінчується на `*/`.

На відміну від однорядкового коментаря, багаторядковий коментар інкапсулює кілька рядків.

Використайте `/*` та `*/`, щоб перетворити поточний цикл `for`, включно з тілом, на багаторядковий коментар.

# --hints--

Почніть багаторядковий коментар з `/*`.

```js
assert.match(code, /\/\*/);
```

Закінчіть багаторядковий коментар з `*/`.

```js
assert.match(code, /\*\//);
```

Весь цикл `for` потрібно закоментувати.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
