---
id: 660f1e3f047bf4e403268713
title: Крок 39
challengeType: 1
dashedName: step-39
---

# --description--

На жаль, зараз ви не можете бачити, що робить цикл.

Використайте `let`, щоб оголосити змінну `result` та призначте до неї порожній рядок. Порожній рядок представляється лапками без нічого між ними, як от `""`.

# --hints--

Оголосіть змінну `result`.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

Використайте `let`, щоб оголосити змінну `result`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

Змінна `result` має бути порожнім рядком.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
