---
id: 660f359af3e32e0f1a6880b7
title: Крок 68
challengeType: 1
dashedName: step-68
---

# --description--

Тепер ви повинні побачити ту ж саму групу символів у консолі. Функція `padRow` робить те саме, що ви робили раніше, але тепер вона знаходиться у власній секції для багаторазового використання.

Використайте оператор додавання, щоб об’єднати один пробіл `" "` до початку та кінця повторюваного рядка `character`.

Пам’ятайте, що ви можете використати оператор `+` для об’єднання рядків таким чином:

```js
" " + "string"
```

# --hints--

З’єднайте один пробіл до початку повернутого значення.

```js
assert.match(padRow(1, 1), /^\s/);
```

З’єднайте один пробіл до кінця повернутого значення.

```js
assert.match(padRow(1, 1), /\s$/);
```

Функція `padRow()` має повернути повторювану серію `character` із пробілом перед та після серії.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
