---
id: 660f48a419b40238e2b8b4d5
title: Крок 93
challengeType: 1
dashedName: step-93
---

# --description--

Цикл більше не покладається на змінну `continueLoop`. Це робить змінну <dfn>невикористаною декларацією</dfn>. Як правило, потрібно уникати невикористаних декларацій, щоб запобігти плутанині в майбутньому.

Видаліть змінну `continueLoop`.

# --hints--

Вам більше не потрібна змінна `continueLoop`.

```js
assert.notMatch(__helpers.removeJSComments(code), /continueLoop/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
