---
id: 660f540c2176ea5dec01306d
title: Крок 118
challengeType: 1
dashedName: step-118
---

# --description--

Гарна праця! Експериментуйте з різними значеннями для змінних `character`, `count` та `inverted`.

Коли ви будете готові перейти до наступного проєкту, встановіть `character` на `"!"`, `count` на `10`, та `inverted` на `false`, щоб продовжити.

Вітаємо із завершенням вашого першого проєкту JavaScript!

# --hints--

Встановіть `character` на `"!"`.

```js
assert.equal(character, "!");
```

Встановіть `count` на `10`.

```js
assert.equal(count, 10);
```

Встановіть `inverted` на `false`.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
