---
id: 660f455b044d3230ed971e98
title: Крок 85
challengeType: 1
dashedName: step-85
---

# --description--

Цикл <dfn>`while`</dfn> виконуватиметься знову та знову, доки вказана `condition` перестане бути правдивою. Він має наступний синтаксис:

```js
while (condition) {
  logic;
}
```

Використайте цей синтаксис, щоб оголосити цикл `while` з умовою `continueLoop`. Тіло повинно бути порожнім.

# --hints--

Ви повинні використати цикл `while`.

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

Цикл `while` має використовувати `continueLoop` як умову.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
