---
id: 660f34626216270c682e2f7b
title: Крок 66
challengeType: 1
dashedName: step-66
---

# --description--

<dfn>Виклик функції</dfn> фактично дозволяє використовувати функцію. Можливо, ви не знали про це, але такі методи, як `.push()`, які ви використовували, були викликами функцій.

Функція викликається пославшись на назву функції та додавши `()`. Ось як викликати функцію `test`:

```js
test();
```

Замініть `character.repeat(i + 1)` у виклику `.push()` на виклик функції для функції `padRow`.

# --hints--

Не використовуйте `i + 1` у виклику `push`.

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

Не використовуйте `character.repeat` у виклику `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

Викличте `padRow` у виклику `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*\)\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
