---
id: 660f4455f457ef2e3ec6920f
title: Крок 80
challengeType: 1
dashedName: step-80
---

# --description--

Тепер рядок більше не друкується, оскільки `false` не є `true`(правдивим). Але як щодо інших значень?

Спробуйте змінити умову на рядок `"false"`.

# --hints--

Інструкція `if` має містити рядок `"false"` як умову.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*('|")false\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (false) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
