---
id: 660f207334fabaeac3269c38
title: Крок 41
challengeType: 1
dashedName: step-41
---

# --description--

Щоб маніпулювати рядком `result`, ви будете використовувати інший тип циклу. Зокрема, цикл `for...of`, який ітерує над кожним елементом ітерованого об’єкта та тимчасово призначає його до змінної.

Синтаксис циклу `for...of` виглядає так:

```js
for (const value of iterable) {

}
```

Зауважте, що можна використати `const`, оскільки змінна існує лише для окремої ітерації, а не протягом усього циклу.

Створіть цикл `for...of` для ітерації над масивом `rows`, призначаючи кожне значення до змінної `row`.

# --hints--

Використайте інше ключове слово `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

Оголосіть змінну `row`.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

Використайте `const`, щоб оголосити змінну `row`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Цикл `for...of` має оголосити змінну `row`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

Видобудьте змінну `row` з `rows` за допомогою ключового слова `of`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

Тіло циклу `for...of` повинно бути порожнім.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
