---
id: 660f2eccfe3f820304af1b39
title: Крок 64
challengeType: 1
dashedName: step-64
---

# --description--

Для того, щоб знати, як відформатувати рядок, функція `padRow` має знати, на якому номері рядка ви перебуваєте та скільки рядків загалом генерується.

Найкращий спосіб зробити це, створивши для них параметри функції. Надайте функції `padRow` параметри `rowNumber` та `rowCount`. Кілька параметрів розділяються комою:

```js
function name(first, second) {

}
```

# --hints--

Функція `padRow` повинна мати параметр `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

Додайте кому після параметра `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

Функція `padRow` повинна мати параметр `rowCount`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
