---
id: 660f0c34aad72dc712b97624
title: Крок 27
challengeType: 1
dashedName: step-27
---

# --description--

Інший метод, суттєвий для цього проєкту, це метод `.pop()`. Він видаляє останній елемент із масиву та <dfn>повертає</dfn> цей елемент.

Коли метод повертає значення, ви можете думати про нього, що він дає вам назад значення, роблячи його доступним для використання в інших частинах коду.

Створіть нову змінну під назвою `popped` та призначте до неї результат `rows.pop()`. Потім зареєструйте `popped` на консоль.

# --hints--

Оголосіть змінну під назвою `popped`.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

Використайте `let`, щоб оголосити змінну під назвою `popped`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

Викличте метод `.pop()`.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

Викличте метод `.pop()` на масиві `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

Зареєструйте змінну `popped`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
