---
id: 660f0ee51d7460ce88cd248d
title: Крок 29
challengeType: 1
dashedName: step-29
---

# --description--

Ви очікували побачити `4` на консолі? `.push()` повертає нову довжину масиву після додавання наданого значення.

Важливо бути в курсі, які значення повертає метод. Приділіть трохи часу, щоб поекспериментувати з `.push()` та `.pop()`. Коли ви будете готові, видаліть усі виклики `.push()` та `.pop()`, а також інструкції `console.log`.

# --hints--

Вам не потрібен виклик `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

Вам не потрібен виклик `.pop()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

Ви не повинні мати жодних інструкцій `log`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

Ви не повинні мати змінну `popped`.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

Ви не повинні мати змінну `pushed`.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
