---
id: 660f48e1d3682f39e81843c4
title: Крок 94
challengeType: 1
dashedName: step-94
---

# --description--

Генератор пірамід все ще працює. Однак можна знову отримати нескінченний цикл.

Оскільки ви лише перевіряєте, чи `done` не дорівнює `count`, якщо `done` був би **більшим** ніж `count` цикл би тривав вічно.

Оновіть умову циклу, щоб перевіряти чи `done` менше або дорівнює `count`.

# --hints--

Цикл `while` має перевіряти, чи `done` менше або дорівнює `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
