---
id: 660f17d4e9f227d86e834abd
title: Крок 33
challengeType: 1
dashedName: step-33
---

# --description--

Щоб створити піраміду, вам потрібно буде створити кілька рядків. Коли вам потрібно виконати завдання повторно, доки не буде виконано умову, ви скористаєтеся <dfn>циклом</dfn>. Є багато шляхів написати цикл.

Почніть з базового циклу `for`. Цикли `for` використовують такий синтаксис:

```js
for (iterator; condition; iteration) {
  logic;
}
```

У наступних кроках ви детально дослідите кожен компонент циклу. Наразі побудуйте цикл `for`, який містить терміни `"iterator"`, `"condition"` та `"iteration"` для трьох компонентів. Тримайте <dfn>body</dfn>(тіло) циклу, розділ у фігурних дужках `{}`, порожнім.


# --hints--

Ви повинні мати цикл `for`.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

Першим компонентом циклу `for` має бути рядок `"iterator"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

Другим компонентом циклу `for` має бути рядок `"condition"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

Третім компонентом циклу `for` має бути рядок `"iteration"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

Тіло інструкції `for` має бути порожнім.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
