---
id: 660f1cedf3676fe26122ebf6
title: Крок 38
challengeType: 1
dashedName: step-38
---

# --description--

Ви повинні побачити на консолі цифри від нуля до семи, надрукованими по одній на рядок. Це послужить основою для генерації піраміди.

Замініть інструкцію журналювання на інструкцію, яка штовхає `i` до масиву `rows`.

# --hints--

Вам не потрібен виклик `console.log`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Викличте `.push()` на масиві `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Штовхніть `i` до масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

`.push()` має коїтися у циклі `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
