---
id: 660f4a83373de83ea101685f
title: Крок 100
challengeType: 1
dashedName: step-100
---

# --description--

Оскільки ви збираєтеся виконувати цикл у протилежному напрямку, цикл має виконуватися, поки `i` більше за `0`. Для цього можна використати оператор <dfn>більше ніж</dfn> `>`.

Встановіть умову циклу для виконання, коли `i` більше за `0`.

# --hints--

Цикл `for` має запускатися коли `i` більше за `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
