---
id: 660f033cf051ebb50ea3bf48
title: Крок 18
challengeType: 1
dashedName: step-18
---

# --description--

У програмуванні вам часто доведеться працювати з великою кількістю даних. Існує багато структур даних, які можуть допомогти вам впорядкувати дані та керувати ними. Однією з базових структур даних є <dfn>масив</dfn>.

<dfn>Масив</dfn> — це непримітивний тип даних, який може тримати серію значень. Непримітивні типи даних відрізняються від примітивних тим, що вони можуть тримати складніші дані. Примітивні типи даних, такі як рядки та числа, можуть тримати лише одне значення за раз.

Масиви позначаються квадратними дужками (`[]`). Ось приклад змінної зі значенням порожнього масиву:

```js
let array = [];
```

Оголосіть змінну `rows` та призначте до неї порожній масив.

# --hints--

Ви повинні мати змінну `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

Використайте `let`, щоб оголосити змінну `rows`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

Призначте порожній масив до змінної `rows`.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
