---
id: 660f1a00ac619ddc1e259a66
title: Крок 35
challengeType: 1
dashedName: step-35
---

# --description--

<dfn>Умова</dfn> циклу `for` каже циклу, скільки разів він має повторюватися. Коли `condition`(умова) стане хибною, цикл зупиниться.

У JavaScript булеве значення може бути `true` або `false`. Це не рядки — ви дізнаєтеся більше про різницю пізніше.

Наразі ви використовуватимете оператор <dfn>менше</dfn> (`<`). Це дозволяє перевірити, чи значення ліворуч менше за значення праворуч. Наприклад, `count < 3` буде оцінено як `true`, якщо `count` дорівнює `2`, та `false`, якщо `count` дорівнює `4`.

Замініть рядок `"condition"` на умову, щоб перевірити, чи значення `i` менше за `count`.

# --hints--

Ви повинні використати оператор менше ніж.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

Використайте оператор «менше», щоб перевірити, чи `i` менше ніж `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Цикл `for` має використовувати `i < count` як умову.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
