---
id: 660f09a2694b59c3a10ee304
title: Крок 23
challengeType: 1
dashedName: step-23
---

# --description--

Наразі видаліть перший `console log` та призначення `rows[rows.length - 1]`. Залиште другу інструкцію журналювання `rows` на потім.

# --hints--

Видаліть інструкцію `console.log(rows[0])`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

Видаліть перепризначення `rows[rows.length - 1]`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

Не видаляйте інструкцію `console.log(rows)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
