---
id: 660f34e99571070d56d2f231
title: Крок 67
challengeType: 1
dashedName: step-67
---

# --description--

Функція `padRow` має два параметри, які ви визначили. Під час виклику функції цим параметрам надаються значення.

Значення, які ви надаєте виклику функції, називаються <dfn>аргументами</dfn>, та ви <dfn>передаєте</dfn> аргументи виклику функції. Ось виклик функції з `"Hello"`, переданим як аргумент:

```js
test("Hello");
```

Передайте `i + 1` та `count` як аргументи виклику `padRow`. Як і параметри, аргументи відокремлюються комою.

# --hints--

Передайте `i + 1` у виклик `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

У вас має бути кома після аргументу `i + 1`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

Передайте `count` як другий аргумент.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
