---
id: 660f3b664421471aa595170f
title: Крок 74
challengeType: 1
dashedName: step-74
---

# --description--

Зараз піраміда виглядає трохи смішно. Оскільки ви починаєте цикл із `1` замість `0`, вам не потрібно додавати один до `i`, коли ви передаєте його до `padRow`.

Оновіть перший аргумент виклику `padRow` бути `i`.

# --hints--

Не передавайте `i + 1` до `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

Передайте `i` до `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
