---
id: 660f1b6e60bd9edf902c81fd
title: Крок 36
challengeType: 1
dashedName: step-36
---

# --description--

Інструкція <dfn>ітерації</dfn> вкаже циклу, що робити з ітератором після кожного запуску.

Коли ви перепризначаєте змінну, ви можете використовувати змінну для посилання на попереднє значення перед перепризначенням. Це дозволяє вам робити такі дії, як додавати три до наявного числа. For example, `bees = bees + 3` would increase the value of `bees` by three.

Використайте цей синтаксис, щоб замінити рядок `"iteration"` інструкцією перепризначення, яка збільшує `i` на одиницю.

# --hints--

Додайте одиницю до змінної `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

Призначте назад `i + 1` до змінної `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Цикл `for` має збільшити `i` на `1` на кожній ітерації.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
