---
id: 660f280dda5040f707c76b4a
title: Крок 45
challengeType: 1
dashedName: step-45
---

# --description--

Тепер у вас є серія символів `#`, але форма піраміди все ще відсутня. На щастя, змінна `i` представляє поточний номер «рядка» у циклі, що дозволяє використати її для майстрування пірамідо-подібної структури.

Щоб досягти цього, ви скористаєтеся методом `.repeat()`, доступним для рядків. Цей метод приймає число як аргумент, визначаючи скільки разів повторювати цільовий рядок. Наприклад, використання `.repeat()` для генерації рядка `"Code! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

Використайте метод `.repeat()` на `character` та надайте йому `i` для числа.

# --hints--

Використайте метод `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

Використайте метод `.repeat()` на змінній `character`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

Передайте `i` методу `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
