---
id: 660f46460f9c36330ebc07d8
title: Крок 88
challengeType: 1
dashedName: step-88
---

# --description--

Оператор рівності може призвести до дивної поведінки в JavaScript. Наприклад, `"0" == 0` є істинним, навіть якщо одне є рядком, а друге є числом.

Оператор <dfn>суворої рівності</dfn> `===` використається, щоб перевірити, чи два значення є рівними та мають однаковий тип. Як правило, це оператор рівності, який ви завжди повинні використовувати. З оператором суворої рівності `"0" === 0` стає хибним, тому що хоча вони можуть мати однакове значення нуль, вони не одного типу.

Оновіть умову `done == count`, щоб використати оператор суворої рівності.

# --hints--

Умова `if` має використовуватисувору рівність.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
