---
id: 660f39b444fd6f16d1e49c1f
title: Крок 73
challengeType: 1
dashedName: step-73
---

# --description--

Замість того, щоб передавати `i + 1` у виклик `padRow`, ви можете натомість розпочати цикл із `1`. Це дозволить вам створити цикл індексований з одиниці.

Оновіть ітератор, щоб він починався з `1` замість `0`.

# --hints--

Цикл `for` має ініціалізувати `i` в `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
