---
id: 660f47afe4c98536715d5fa4
title: Крок 91
challengeType: 1
dashedName: step-91
---

# --description--

Оператор <dfn>суворої нерівності</dfn> `!==` дозволяє перевірити, чи два значення не є рівними або не мають однаковий тип. Синтаксис подібний до оператора нерівності: `value !== 4`.

Оновіть умову циклу `while`, щоб перевірити, чи `done` не дорівнює `count`.

# --hints--

Цикл `while` має перевірити, чи не рівні `done` та `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
