---
id: 6610b9f7619764fad5fd516d
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Ви також можете призначити значення однієї змінної до іншої змінної. Наприклад:

```js
let first = "One";
let second = "Two";
second = first;
```

Змінна `second` тепер матиме значення `"One"`.

Щоб побачити це в дії, змініть призначення `secondCharacter` з `"Test"` на змінну `character`.

Потім відкрийте консоль, щоб побачити, що реєструється.

# --hints--

Не присвоюйте значення `"Test"` до змінної `secondCharacter`.

```js
assert.notEqual(secondCharacter, "Test");
```

Призначте значення змінної `character` до змінної `secondCharacter`. Не забудьте про крапку з комою.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

Тепер змінна `secondCharacter` повинна мати значення `"World"`.


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
