---
id: 6610c538372aa61cc0f5b122
title: Крок 57
challengeType: 1
dashedName: step-57
---

# --description--

Змінні в JavaScript доступні в певній <dfn>області</dfn>. Іншими словами, місце де змінна оголошена визначає, де в коді її можна використовувати.

Перша область — це глобальна область. Змінні яки оголошені поза будь-яким «блоком», як-от функція або цикл `for`, знаходяться в <dfn>глобальній області</dfn>. Змінні `character`, `count` та `rows` знаходяться в глобальній області видимості.

Коли змінна знаходиться в глобальній області видимості, функція може отримати до неї доступ у своєму визначенні. Ось приклад функції, яка використає глобальну змінну `title`:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

Цей приклад поверне `"Professor Naomi"`. Оновіть функцію `padRow`, щоб повернути значення конкатенації змінної `character` до початку параметра `name`.

# --hints--

Функція `padRow` має конкатенувати `character` до початку `name`.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

Функція `padRow` має повернути результат `character + name`.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
