---
id: 6610c83b52583e245a079217
title: Крок 61
challengeType: 1
dashedName: step-61
---

# --description--

Тепер змінна `call` має значення `"Testing"`. Але функція більше не використає параметр `name`.

Видаліть параметр `name` з оголошення функції, а потім видаліть рядок `"CamperChan"` із виклику `padRow`.

Також видаліть обидві `console.log` із функції `padRow`.
# --hints--

Your `padRow` function should not have a `name` parameter.

```js
assert.notMatch(padRow.toString(), /name/);
```

Не передавайте `"CamperChan"` у виклик `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

You should still call your `padRow` function.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Ви не повинні мати `console.log` перед ключовим словом `return`.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

Ви не повинні мати `console.log` після ключового слова `return`.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
