---
id: 6610c71600966a2191d3a64a
title: Крок 59
challengeType: 1
dashedName: step-59
---

# --description--

Значення, повернуті з функції, використовуються викликнувши функцію. Ви можете використати виклик функції безпосередньо як значення, яке вона повертає, або захопити повернуте значення у змінній. Таким чином, ви можете використати значення, призначене змінній у локальної області видимості, поза функцією, у якій її було створено.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

Щоб використати значення `"Testing"`, поверніть його з функції `padRow`, оновивши інструкцію `return`, щоб повернути лише змінну `test`.


# --hints--

Функція `padRow` має повернути змінну `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
