---
id: 6610b741b54b90f0c0fb3d58
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Коли імена змінних складаються з кількох слів, існують спеціальні конвенції іменування про те, як капіталізувати слова. У JavaScript за конвенцією використовується <dfn>верблюдячий регістр</dfn>.

Верблюдячий регістр означає, що перше слово в назві повністю з малої літери, а наступні слова всі з заголовком. Ось кілька прикладів верблюдячого регістру:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

Використайте верблюдячий регістр, щоб оголосити нову змінну `secondCharacter`.

# --hints--

Оголосіть змінну `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

Використайте `let`, щоб оголосити змінну `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

Не присвоюйте значення змінній `secondCharacter`. Не забудьте про крапку з комою.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
