---
id: 6610c6541c82551f95e765ab
title: Крок 58
challengeType: 1
dashedName: step-58
---

# --description--

Змінні також можна оголошувати в межах функції. Вважається, що ці змінні належать до <dfn>локальної області</dfn> або <dfn>області блоку</dfn>. Змінна, оголошена в межах функції, може бути використана лише у цієї функції. Якщо ви спробуєте отримати до неї доступ поза функцією, ви отримаєте помилку посилання.

Щоб побачити це в дії, використайте `const`, щоб оголосити змінну `test` у функції `padRow`. Ініціалізуйте її значенням `"Testing"`.

Потім під функцією спробуйте зареєструвати `test` на консолі. Ви побачите помилку, оскільки вона не визначена за межами локальної області видимості функції. Видаліть цей `console.log`, щоб пройти тести та продовжити.

# --hints--

Функція має оголосити змінну `test`.

```js
assert.match(padRow.toString(), /var\s+test/);
```

Ініціалізуйте `test` значенням `"Testing"`. Не забудьте про крапку з комою.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

Змінна `test` має йти перед ключовим словом `return`.

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
