---
id: 6610b8f6a98d25f4d485a94d
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Значенням за замовчуванням неініціалізованої змінної є `undefined`. Це спеціальний тип даних, який представляє значення, яке ще не має визначення.

Ви все ще можете призначити значення неініціалізованій змінній. Ось приклад:

```js
let uninitialized;
uninitialized = "assigned";
```

Призначте рядок `"Test"` до змінної `secondCharacter` під оголошенням. Відкрийте консоль, щоб побачити, як змінився ваш реєстр.

# --hints--

Не ініціалізуйте `secondCharacter`. Пам’ятайте, що ініціалізація означає присвоєння значення під час оголошення змінної.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

Використайте оператор присвоєння на `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

Призначте рядок `"Test"` до змінної `secondCharacter`.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
