---
id: 6610c424b7119919b62932f4
title: Крок 53
challengeType: 1
dashedName: step-53
---

# --description--

Функція може і не повертати жорстко закодоване значення. Вона може повертати значення, що зберігається в змінній. Параметри є спеціальними змінними для функції, тому їх також можна повертати.

Змініть функцію `padRow`, щоб `return`(повернути) безпосередньо параметр `name`.

# --hints--

Функція `padRow` має повернути значення параметра `name`.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
