---
id: 6610c77d50636722e5b6be17
title: Крок 60
challengeType: 1
dashedName: step-60
---

# --description--

Під інструкцією `return` зареєструйте рядок `"This works!"` на консоль.

Зробивши це, ви побачите, що рядок `"This works!"` не відображається в консолі, а рядок `console.log("This works!")` відображається сірим кольором.

Скопіюйте `console log` та вставте його над інструкцією `return`. Тепер у консолі має з’явитися рядок `"This works!"`.

Треба знати важливу річ про ключове слово `return`, це то що воно не тільки визначає значення, яке буде повернуто з функції, воно також зупиняє виконання коду в межах функції чи інструкції блоку. Це означає, що будь-який код після інструкції `return` не виконуватиметься.

# --hints--

Your `padRow` function should return the `test` variable.

```js
assert.equal(padRow("Naomi"), "Testing");
```

Перший `console.log` має іти після ключового слова `return`.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

Другий `console.log` має іти перед ключовим словом `return`.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";

  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
