---
id: 6610c48c4ea0891afa7c4696
title: Крок 54
challengeType: 1
dashedName: step-54
---

# --description--

Якщо ви знову відкриєте консоль, ви побачите, що функція `padRow` повертає `undefined`, навіть якщо ви визначили значення, яке повертається! Це пояснюється тим, що параметрам потрібно надати значення під час **виклику** функції.

Коли ви передаєте значення виклику функції, це значення називається <dfn>аргументом</dfn>. Ось приклад виклику функції `demo` та передачі `"Naomi"` як аргументу для параметра `name`.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

Передайте своє ім’я як аргумент для параметра `name` у виклику `padRow`. Пам’ятайте, що ваше ім’я є рядком, тож потрібно буде використовувати лапки.

# --hints--

Передайте рядок у виклик `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
