---
id: 6610c8cfe4cf4d278e35c156
title: Крок 63
challengeType: 1
dashedName: step-63
---

# --description--

Як і очікувалося, зараз функція знову повертає `undefined`. Змінна `call` більше не потрібна, тому видаліть оголошення `call` та `console.log` для змінної `call`.

# --hints--

Ви не повинні мати оголошення `call`.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s*call/);
```

Не реєструйте змінну `call`.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
