---
id: 6610b8017d1671f2814e8c77
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Коли ви оголошуєте змінну без її ініціалізації, вона вважається <dfn>неініціалізованою</dfn>. Наразі змінна `secondCharacter` неініціалізована.

Додайте `console.log()`, щоб побачити яке значення змінної `secondCharacter`.

# --hints--

Додайте другий `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Зареєструйте змінну `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
