---
id: 6610c87eac0f0b256d7b037e
title: Крок 62
challengeType: 1
dashedName: step-62
---

# --description--

Оскільки функція більше не використовує параметр, зміна аргументу, не вплинула на неї.

Ідіть уперед і видаліть оголошення `test` та інструкцію `return` із функції `padRow`, щоб функція знову була порожньою.

# --hints--

Функція `padRow` не повинна мати змінну `test`.

```js
assert.notMatch(padRow.toString(), /test/);
```

Функція `padRow` не повинна повертати значення.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
