---
id: 6610c105bbdacc114d6cdc44
title: Крок 48
challengeType: 1
dashedName: step-48
---

# --description--

Щоб скористатися функцією, її потрібно викликати. <dfn>Виклик функції</dfn> повідомляє програмі запускати код із функції, де б ви не вибрали її викликати. Синтаксис виклику функції — це ім’я функції, за яким слідують дужки. Наприклад, цей код визначає та викликає функцію `test`.

```js
function test() {

}
test();
```

Викличте функцію `padRow`.

# --hints--

Викличте функцію `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Тіло функції `padRow` має бути порожнім.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
