---
id: 6610c16c4fa0df12c0e30675
title: Крок 49
challengeType: 1
dashedName: step-49
---

# --description--

Ви викликаєте функцію `padRow`, але нічого не робите з цим викликом функції. Усі функції в JavaScript <dfn>повертають</dfn> значення, тобто вони вам надають визначений результат їх виклику для використання в іншому місці.

Щоб побачити результат виклику функції `padRow`, оголосіть змінну `call` та призначте наявний виклик `padRow` цій змінній.

# --hints--

Оголосіть змінну `call`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let)\s+call/);
```

Використайте `const`, щоб оголосити змінну `call`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

Призначте `call` результат виклику `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
