---
id: 664599653fcd6e97104f9261
title: Крок 83
challengeType: 1
dashedName: step-83
---

# --description--

Іноді вам потрібно запустити інший код, коли всі умови `if...else if` є `false`. Ви можете зробити це, додавши блок `else`.

Блок `else` оцінюється, лише тоді, коли блоки `if` та `else if` не відповідають умовам.

Тут блок `else` додається до блоку `else if`.

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

Додайте блок `else` до блоку `else if`. Усередині блоку `else`, зареєструйте рядок `"This is the else block"` на консоль.

Щоб побачити результати на консолі, можна вручну змінити `<` в інструкції `else if` на `>`. Це зробить умову `false` та запуститься блок `else`.

# --hints--

Ви повинні мати блок `else`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

Блок `else` має зареєструвати рядок `"This is the else block"` на консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

