---
id: 664ed97c55a99e5ffab759dc
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Перш ніж рухатися далі, вам слід взяти хвилинку, щоб переглянути вивчені концепції.

Використайте ключове слово `let`, щоб оголосити змінні `profession` та `age`. Ініціалізуйте `profession` рядком `"teacher"`, але не ініціалізуйте `age` жодним значенням.

Зареєструйте обидві змінні на консоль, щоб побачити результати.

# --hints--

Оголосіть змінну `profession`.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

Використайте `let`, щоб оголосити змінну `profession`.

```js
assert.match(code, /let\s+profession/);
```

Призначте рядок `"teacher"` до змінної `profession`.

```js
assert.strictEqual(profession, "teacher");
```

Оголосіть змінну `age`.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

Використайте `let`, щоб оголосити змінну `age`.

```js
assert.match(code, /let\s+age/);
```

Не присвоюйте значення до змінної `age`.

```js
assert.isUndefined(age);
```

Виведіть `age` на консолі.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

Виведіть `profession` на консолі.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
