---
id: 664ed9344ba3495e1f67f8ea
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Тепер, коли ви переглянули оголошення та ініціалізацію, видаліть код, який ви написали для перегляду.

Не видаляйте змінну `character`.

# --hints--

Ви не повинні мати змінну `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

Ви не повинні мати змінну `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

Код не повинен містити інструкцію `console.log` для `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

Код не повинен містити інструкцію `console.log` для `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

Ви все ще повинні мати змінну `character`.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
