---
id: 665732da4815b70bb083915e
title: Крок 24
challengeType: 1
dashedName: step-24
---


# --description--

У кількох останніх кроках ви дізналися все про роботу з масивами. Візьміть момент, щоб переглянути те, що ви дізналися.

Почніть з оголошення змінної `cities` та ініціалізації її як масив рядків `"London"`, `"New York"` та `"Mumbai"`. Потім зареєструйте цю змінну на консолі.

Після реєстрації змініть останній елемент з `cities` на рядок `"Mexico City"`, а потім знову зареєструйте змінну `cities`.

Якщо все зроблено правильно, ви повинні побачити цей вивід у консолі.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

Використайте `let`, щоб оголосити змінну `cities`.

```js
assert.match(code, /let\s+cities/);
```

Призначте масив рядків `"London"`, `"New York"` та `"Mumbai"` до змінної `cities`.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*\]/);
```

Використайте `console.log()` для реєстрації всього масиву `cities` на консолі.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

Оновіть останній елемент масиву `cities` до рядка `"Mexico City"`. Пам’ятайте, що ви можете отримати доступ до останнього елемента масиву за допомогою `array[array.length - 1]`.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

Код повинен містити дві інструкції `console.log(cities)`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
