---
id: 66643c93e05093c728abdbe9
title: Крок 55
challengeType: 1
dashedName: step-55
---

# --description--

Перш ніж рухатися далі, візьміть момент переглянути, як працюють функції.

Оголосіть функцію під назвою `addTwoNumbers`. Ця функція має приймати два аргументи та повернути суму цих двох аргументів.

Функція не повинна використати жорстко закодовані значення. Прикладом жорстко закодованої функції може бути:

```js
function sayName(firstName, lastName) {
  return "John Doe";
}

sayName("Camper", "Cat");
```

Ця функція повертатиме `"John Doe"` незалежно від аргументів, переданих до параметрів `firstName` та `lastName`, тому `"John Doe"` вважається жорстко закодованим значенням.

Оголосіть змінну `sum` і призначте до неї значення виклику функції `addTwoNumbers` з `5` та `10` як аргументами. Log the `sum` variable to the console.

# --hints--

Ви повинні мати функцію під назвою `addTwoNumbers`.

```js
assert.isFunction(addTwoNumbers);
```

Your function `addTwoNumbers` should have two parameters.

```js
assert.lengthOf(addTwoNumbers, 2);
```

Your function should return the sum of the two parameters.

```js
assert.strictEqual(addTwoNumbers(5,10), 15);
```

Your function should not return a hard-coded value. That is, it should work with any two number arguments.

```js
assert.strictEqual(addTwoNumbers(3, 5), 8);
```

You should declare a `sum` variable.

```js
assert.isDefined(sum);
```

Your `sum` variable should have the value `15`.

```js
assert.strictEqual(sum, 15);
```

You should assign `sum` the value from calling the `addTwoNumbers` function with `5` and `10` for the arguments.

```js
assert.match(code, /sum\s*=\s*addTwoNumbers\s*\(/);
```

You should log your `sum` variable.

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
