---
id: 616d55bd160a95e22453a081
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Agora que você tem um marcador centrado com cor, é hora de adicionar os outros marcadores.

No elemento `.container`, adicione mais dois elementos `div` e dê a cada um a classe `marker`.

# --hints--

O primeiro novo elemento `div` deve ter uma tag de abertura.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

O primeiro novo elemento `div` deve ter uma tag de fechamento.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

O segundo novo elemento `div` deve ter uma tag de abertura.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

O segundo novo elemento `div` deve ter uma tag de fechamento.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Os novos elementos `div` devem estar dentro de um `div` com a classe `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

Os dois novos elementos `div` devem ter a classe `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
