---
id: 616d525007b8c5d8b3308b1c
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Seu marcador ficaria melhor se estivesse centralizado na página. Uma maneira fácil de fazer isso é usando a propriedade <dfn>shorthand</dfn> `margin`.

Você pode definir a área de margem dos elementos separadamente com propriedades como `margin-top` e `margin-left` e a propriedade abreviada `margin` facilita definir múltiplas áreas de margem ao mesmo tempo.

Para centralizar seu marcador na página, defina a sua propriedade `margin` como `auto`. Isso define `margin-top`, `margin-right`, `margin-bottom` e `margin-left` como `auto`.

# --hints--

A regra do CSS `.marker` deve ter uma propriedade `margin` definida como `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
