---
id: 616d3a67ccf800ad94ec89ae
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Agora, você vai adicionar alguns elementos que mais pra frente serão estilizados em marcadores de cor.

Primeiro, dentro do elemento `body`, adicione um elemento `div` com um atributo `class` definido como `container`. Certifique-se de que o elemento `div` esteja colocado abaixo do elemento `h1`.

# --hints--

Você deve acrescentar uma tag de abertura para `div`.

```js
assert.match(code, /<div.*>/i);
```

Você deve acrescentar uma tag de fechamento para `div`.

```js
assert.match(code, /<\/div\s*>/i);
```

Seu elemento `div` deve estar dentro do seu elemento `body`.

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

O elemento `div` deve estar abaixo do elemento `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

O elemento `div` deve ter o atributo `class` definido como `container`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
