---
id: 61764c602bee6974e7790f35
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Para dar cores diferentes aos marcadores, você precisará adicionar uma classe única a cada um.Você pode adicionar múltiplas classes a um elemento listando-as no atributo `class` e separando-as com um espaço.Por exemplo, o código a seguir adiciona as classes `animal` e `dog` a um elemento `div`:

```html
<div class="animal dog">
```

Se você adicionar múltiplas classes a um elemento HTML, os estilos das primeiras classes que você listar na folha de estilos podem ser sobrescritos por classes posteriores.

Para começar, adicione a classe `one` ao primeiro elemento do marcador `div`.


# --hints--

Adicione a class `one` ao primeiro elemento do marcador `div`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

O primeiro marcador `div` deve ter os marcadores de classe `marker` e `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
