---
id: a105e963526e7de52b219be9
title: Implemente uma União Ordenada Única
challengeType: 26
dashedName: implement-a-unique-sorted-union
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função chamada `uniteUnique`.
1. A função `uniteUnique` deve aceitar dois ou mais arrays como argumentos.
1. A função deve retornar um novo array que contenha valores únicos dos arrays de argumento, na ordem em que são encontrados pela primeira vez nos argumentos. Por exemplo, uma entrada como `[1, 2, 4], [2, 3, 5]` teria uma saída de `[1, 2, 4, 3, 5]`.

# --hints--

Você deve ter uma função `uniteUnique`.

```js
assert.isFunction(uniteUnique);
```

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` deve retornar `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` deve retornar `[1, 2, 3, 5]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` deve retornar `[1, 2, 3, 5, 4, 6, 7, 8]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` deve retornar `[1, 3, 2, 5, 4, 6]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` deve retornar `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const uniteUnique = (...arrays) => {
  return arrays.reduce((a, b) => {
    return [
      ...a,
      ...b.filter((e, currentIndex) => b.indexOf(e) === currentIndex && a.indexOf(e) === -1)
    ];
  }, []);
};
```
