---
id: 68f1f8c248306da204655117
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Das Attribut `muted` schaltet den Ton bei der ersten Wiedergabe stumm. Wenn Sie `controls` aktiviert haben, kann der Benutzer den Ton wieder einschalten. Wenn Sie das Attribut `muted` weglassen, wird der Ton bei der ersten Wiedergabe abgespielt.

Das Attribut `muted` ist ein Boolean-Attribut und benötigt keinen Wert.

Fügen Sie das Attribut `muted` dem Element `video` hinzu.

# --hints--

Ihr Element `video` sollte das Attribut `muted` haben.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
