---
id: 68f1fac259a4afabb478b960
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Sie haben vielleicht bemerkt, dass Sie noch nicht auf das eigentliche Video verlinkt haben. Das werden Sie in der nächsten Phase tun. Bei Videodateitypen gibt es Unterschiede in der Browserunterstützung. Um dem gerecht zu werden, können Sie `source`-Elemente innerhalb des `video`-Elements verwenden, und der Browser wählt das erste kompatible `source` aus.

Hier ist ein Beispiel für ein `source`-Element:

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

Das `source`-Element ist ein sogenanntes Void-Element, daher hat es kein schließendes Tag.

Fügen Sie ein `source`-Element innerhalb Ihres `video`-Elements hinzu.


# --hints--

Sie sollten ein `source`-Element innerhalb Ihres `video`-Elements haben.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
