---
id: 68f1f57723debf9383c75947
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Das Attribut `loop` startet das Video neu, sobald die Wiedergabe abgeschlossen ist. Denken Sie an ein Internet-Meme, das die Wiedergabe wiederholt. Wenn Sie das Attribut `loop` weglassen, wird das Video nur einmal abgespielt.

Das Attribut `loop` ist ein Boolean-Attribut und benötigt keinen Wert.

Fügen Sie das Attribut `loop` zum Element `video` hinzu.

# --hints--

Ihr Element `video` sollte das Attribut `loop` haben.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
