---
id: 68f1f6f97af9749a42260a43
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Das Attribut `controls` stellt Wiedergabesteuerungen bereit, darunter Wiedergabe, Zurückspulen und Lautstärkeregelung für das Element `video`.

Das Attribut `controls` ist ein Boolean-Attribut und benötigt keinen Wert.

Fügen Sie das Attribut `controls` zum Element `video` hinzu.

Jetzt sollten Sie das Element `video` auf der Seite sehen.

# --hints--

Ihr Element `video` sollte das Attribut `controls` haben.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
