---
id: 68f1f371fefd448b4f70b898
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

In einer vorherigen Lektion haben Sie verschiedene Attribute kennengelernt, die für das `video`-Element verfügbar sind. Das `width`-Attribut bestimmt die Breite des Videos in Pixeln.

Fügen Sie dem `video`-Element das `width`-Attribut mit dem Wert `640` hinzu.

# --hints--

Ihr `video`-Element sollte das Attribut `width` mit dem Wert `640` haben.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
