---
id: 68f18724d71db1dda8d6fec7
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop bauen Sie einen HTML-Videoplayer. Der HTML-Boilerplate-Code wurde für Sie bereitgestellt.

Erstellen Sie ein `h1`-Element und geben Sie ihm den Text `Working with the HTML Video Element`.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
assert.exists(document.querySelector('h1'));
```

Ihr `h1`-Element sollte den Text `Working with the HTML Video Element` enthalten.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Working with the HTML Video Element');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
