---
id: 68f1f993460f06a64c033061
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Das `poster`-Attribut ist ein Vorschaubild des Videos. Denken Sie an die Videos, die Sie auf YouTube ansehen. Es wird angezeigt, während das Video heruntergeladen wird. Wenn das Attribut weggelassen wird, wird während der Downloadphase der erste Videobildschirm angezeigt.

Fügen Sie nun das `poster`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` zu Ihrem `video`-Element hinzu.

# --hints--

Ihr `video`-Element sollte das `poster`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` haben.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
