---
id: 68f1eb9ad0a7575e4dbc31d5
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Erstellen Sie als Nächstes ein `video`-Element unterhalb des `h1`. In den nächsten Schritten fügen Sie die notwendigen Attribute hinzu, damit der Videoplayer funktioniert.

# --hints--

Sie sollten ein `video`-Element haben.

```js
assert.exists(document.querySelector('video'));
```

Das neue `video`-Element sollte unterhalb des `h1` sein.

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
