---
id: 68f26d52b13e0f5b6e2d2e09
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Ein weiterer häufiger MIME-Typ ist der `video/webm`-MIME-Typ.

WebM ist ein quelloffenes audiovisuelle Medienformat, das von Google entwickelt wurde und hauptsächlich für webbasierte Medieninhalte konzipiert ist. Es unterstützt Videocodecs wie VP8, VP9 und AV1 sowie Audiocodecs wie Vorbis und Opus, was es zu einer beliebten Wahl für HTML5-Video- und Audioelemente macht.

Fügen Sie unter Ihrem ersten `source`-Element ein weiteres `source`-Element hinzu und geben Sie ihm ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` sowie ein `type`-Attribut mit dem Wert `video/webm`.

# --hints--

Sie sollten ein zweites `source`-Element haben.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

Ihr zweites `source`-Element sollte ein `src`-Attribut besitzen.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

Ihr zweites `source`-Element sollte ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` besitzen.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

Ihr zweites `source`-Element sollte ein `type`-Attribut besitzen.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

Ihr zweites `source`-Element sollte ein `type`-Attribut mit dem Wert `video/webm` besitzen.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
