---
id: 68f2701b2160c765f23b60f8
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Ein weiterer gängiger MIME-Typ ist der `video/ogg`-MIME-Typ.

Ogg ist ein digitales Multimedia-Containerformat, das für effizientes Streaming und die Bearbeitung digitaler Multimediainhalte entwickelt wurde. Es wird von der Xiph.Org Foundation gepflegt und ist frei sowie offen, ohne Einschränkungen durch Softwarepatente. Sein Name stammt von „ogging“, einem Fachbegriff aus dem Computerspiel Netrek.

Fügen Sie unter Ihrem zweiten `source`-Element ein drittes `source`-Element hinzu und geben Sie ihm ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` sowie ein `type`-Attribut mit dem Wert `video/ogg`.

# --hints--

Sie sollten ein drittes `source`-Element haben.

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

Ihr drittes `source`-Element sollte ein `src`-Attribut besitzen.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

Ihr drittes `source`-Element sollte ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` haben.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

Ihr drittes `source`-Element sollte ein `type`-Attribut besitzen.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

Ihr drittes `source`-Element sollte ein `type`-Attribut mit dem Wert `video/ogg` haben.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
