---
id: 68f26a5f79e97c5032c10642
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Sie haben eine Videodatei mit der `mp4` Dateiendung verwendet und müssen dem Browser mitteilen, wie er die Datei lesen soll.

Sie verwenden das `type`-Attribut, um den `video/mp4`-MIME-Typ anzugeben.

MIME (Multipurpose Internet Mail Extensions) ist ein Standard, um Dokumente in anderen Formen als ASCII-Text zu beschreiben, zum Beispiel Audio, Video und Bilder.

MP4, offiziell bekannt als MPEG-4 Part 14, ist ein digitales Multimedia-Containerformat. Es wird häufig zum Speichern von Video und Audio verwendet, kann aber auch andere Datentypen wie Untertitel und Standbilder enthalten. MP4-Dateien sind für das Streaming im Internet ausgelegt und mit vielen Geräten und Plattformen kompatibel.

Fügen Sie nun das `type`-Attribut mit dem Wert `video/mp4` hinzu.

# --hints--

Ihr `source`-Element sollte ein `type`-Attribut haben.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

Ihr `source`-Element sollte ein `type`-Attribut mit dem Wert `video/mp4` haben.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
