---
id: 68f27139bddf9f6bf2821acf
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Das letzte `source`-Element, das Sie hinzufügen, wird für den `video/quicktime`-MIME-Typ sein.

QuickTime ist eine erweiterbare Multimedia-Architektur von Apple, die das Abspielen, Streamen, Kodieren und Transkodieren verschiedener digitaler Medienformate unterstützt. Nicht so populär wie das MP4-Format, benötigen Sie es möglicherweise für die Unterstützung veralteter Anwendungen.

Fügen Sie unter Ihrem dritten `source`-Element ein viertes `source`-Element hinzu und geben Sie ihm ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` und ein `type`-Attribut mit dem Wert `video/quicktime`.

Herzlichen Glückwunsch! Sie haben den HTML-Video-Player-Workshop abgeschlossen.

# --hints--

Sie sollten ein viertes `source`-Element haben.

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

Ihr viertes `source`-Element sollte ein `src`-Attribut haben.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

Ihr viertes `source`-Element sollte ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` haben.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

Ihr viertes `source`-Element sollte ein `type`-Attribut haben.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

Ihr viertes `source`-Element sollte ein `type`-Attribut mit dem Wert `video/quicktime` haben.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
