---
id: 587d7fb7367417b2b2512c0d
title: 使用 model.findById() 以 _id 搜尋你的資料庫
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

當儲存文件時，MongoDB 會自動添加欄位 `_id`，並將其設為唯一的字母數字鍵。以 `_id` 進行搜尋是非常頻繁的操作行為，因此 Mongoose 提供了專用的方法。

# --instructions--

修改 `findPersonById` 以使用 `Model.findById() -> Person` 找出唯一擁有指定 `_id` 的人。使用函式引數 `personId` 作為搜尋鍵。

# --hints--

透過 Id 找到一個項目應該成功

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```

