---
id: 587d7fb6367417b2b2512c09
title: 創建並儲存模型的記錄
challengeType: 2
forumTopicId: 301536
dashedName: create-and-save-a-record-of-a-model
---

# --description--

在這個挑戰中，你必須創建並儲存一個模型的記錄。

# --instructions--

在 `createAndSavePerson` 函式中，使用你之前建立的 `Person` 模型建構式來建立一個文件實體。將一個具有欄位 `name`、`age` 和 `favoriteFoods` 的物件傳遞給建構式。它們的型別必須符合 `personSchema` 中的型別。然後，對傳回的文件實體呼叫方法 `document.save()`。使用 Node 慣例將回呼傳遞給它。這是一個常見的樣式；所有以下的 CRUD 方法都將這樣的回呼函式作為最後一個引數。

```js
/* Example */

// ...
person.save(function(err, data) {
  //   ...do your stuff here...
});
```

# --hints--

建立並儲存一個 `db item` 應該成功

```js
  const response = await fetch(code + '/_api/create-and-save-person');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isString(data.name, '"item.name" should be a String');
  assert.isNumber(data.age, '28', '"item.age" should be a Number');
  assert.isArray(
    data.favoriteFoods,
    '"item.favoriteFoods" should be an Array'
  );
  assert.equal(data.__v, 0, 'The db item should be not previously edited');
```

