---
id: 587d7fb7367417b2b2512c0a
title: 使用 model.create() 創建多筆記錄
challengeType: 2
forumTopicId: 301537
dashedName: create-many-records-with-model-create
---

# --description--

有時你需要創建許多模型實體，例如在使用初始資料為資料庫進行種子資料時。`Model.create()` 將像 `[{name: 'John', ...}, {...}, ...]` 這樣的物件陣列作為第一個引數，並將它們全部存檔到資料庫中。

# --instructions--

修改 `createManyPeople` 函式（程式），使用帶有引數 `arrayOfPeople` 的 `Model.create()` 來創建多個人。

**注意：** 你可以重複使用你在前一個練習中具現化的模型。

# --hints--

一次建立多個 db 項目應該會成功

```js
  const response = await fetch(code + '/_api/create-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'John', age: 24, favoriteFoods: ['pizza', 'salad'] },
      { name: 'Mary', age: 21, favoriteFoods: ['onions', 'chicken'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an array');
  assert.equal(
    data.length,
    2,
    'the response does not contain the expected number of items'
  );
  assert.equal(data[0].name, 'John', 'The first item is not correct');
  assert.equal(
    data[0].__v,
    0,
    'The first item should be not previously edited'
  );
  assert.equal(data[1].name, 'Mary', 'The second item is not correct');
  assert.equal(
    data[1].__v,
    0,
    'The second item should be not previously edited'
  );
```

