---
id: 587d7fb7367417b2b2512c0c
title: 使用 model.findOne() 從你的資料庫傳回單一符合的文件
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()` 的行為類似於 `Model.find()`，但它只傳回一個文件（不是陣列），即使有多個條款。當你依據已宣告為唯一的屬性進行搜尋時，它特別有用。

# --instructions--

修改 `findOneByFood` 函式（程式），以使用 `Model.findOne() -> Person` 找出在該人的最愛中有特定食物的唯一一個人。使用函式（程式）引數 `food` 作為搜尋鍵。

# --hints--

尋找一個項目應該成功

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```

