---
id: 587d7fb8367417b2b2512c0f
title: 使用 model.findOneAndUpdate() 對文件執行新更新
challengeType: 2
forumTopicId: 301542
dashedName: perform-new-updates-on-a-document-using-model-findoneandupdate
---

# --description--

較新版本的 Mongoose 提供了簡化文件更新的方法。一些進階特性（例如 pre/post 掛鉤、驗證）在此方法下的行為有所不同，因此經典方法在許多情況下仍然有用。當以 id 進行搜尋時，可以使用 `findByIdAndUpdate()`。

# --instructions--

修改 `findAndUpdate` 函式（程式），以透過 `Name` 找到一個人，並將該人的年齡設為 `20`。使用函式（程式）參數 `personName` 作為搜尋鍵。

**注意：**你應該傳回更新後的文件。為此，你需要將選項文件 `{ new: true }` 作為第 3 個引數傳遞給 `findOneAndUpdate()`。預設情況下，這些方法會傳回未修改的物件。

# --hints--

findOneAndUpdate 一個項目應該成功

```js
  const response = await fetch(code + '/_api/find-one-update', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Dorian Gray',
      age: 35,
      favoriteFoods: ['unknown']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Dorian Gray', 'item.name is not what expected');
  assert.equal(data.age, 20, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['unknown'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(
    data.__v,
    0,
    'findOneAndUpdate does not increment version by design!'
  );
```

