---
id: 587d7fb7367417b2b2512c0b
title: 使用 model.find() 來搜尋你的資料庫
challengeType: 2
forumTopicId: 301543
dashedName: use-model-find-to-search-your-database
---

# --description--

在最簡單的用法中，`Model.find()` 接受一個查詢文件（一個 JSON 物件）作為第一個引數，然後是一個回呼。它傳回一個符合條件的陣列。它支援非常廣泛的搜尋選項。請參閱文件以了解更多資訊。

# --instructions--

修改 `findPeopleByName` 函式（程式）以使用 <code>Model.find() -\> [Person]</code> 找出所有具有指定名稱的人。

將函式（程式）引數 `personName` 作為搜尋鍵。

# --hints--

尋找符合條件的所有項目應該成功

```js
  const response = await fetch(code + '/_api/find-all-by-name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'r@nd0mN4m3',
      age: 24,
      favoriteFoods: ['pizza']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data[0].name,
    'r@nd0mN4m3',
    'item.name is not what expected'
  );
  assert.equal(data[0].__v, 0, 'The item should be not previously edited');
```

