---
id: 587d7fb6367417b2b2512c07
title: 創建一個 Model
challengeType: 2
forumTopicId: 301535
dashedName: create-a-model
---

# --description--

**C**RUD 零件 I - 創建

首先，我們需要一個 Schema。每個 schema 對應到一個 MongoDB 群集。它定義該群集中文件的結構。Schemas 是 Models 的建造區塊。它們可以巢狀以創建複雜的模型，但在這個案例中，我們將保持簡單。Model 允許你創建你的物件實體，稱為文件。

在伺端中，與資料庫的互動發生在處理常式函式中。當某些事件發生時（例如有人呼叫你的 API 上的端點），這些函式會被執行。我們將在這些練習中採用相同的方法。`done()` 函式是一個回呼，用來告訴我們在完成非同步操作（例如插入、搜尋、更新或刪除）後可以繼續。它遵循 Node 的慣例，成功時應以 `done(null, data)` 呼叫，錯誤時則以 `done(err)` 呼叫。

警告－與遠端服務互動時，可能會發生錯誤！

```js
/* Example */

const someFunc = function(done) {
  //... do something (risky) ...
  if (error) return done(error);
  done(null, result);
};
```

# --instructions--

創建一個名為 `personSchema` 的人物結構，具有以下形狀：

* 一個必填的 `name` 欄位，型別為 `String`
* 一個型別為 `Number` 的 `age` 欄
* 一個型別為 `[String]` 的 `favoriteFoods` 欄

使用 Mongoose 基本的 schema 型別。如果你想，你也可以為你的欄添加更多欄位，使用簡單的驗證器如 required 或 unique，並設定預設值。請參考我們的 <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">Mongoose 文章</a>。

現在，從 `personSchema` 創建一個模型並將其指定到現有的變數 `Person`。

# --hints--

從 mongoose schema 建立實體應該會成功

```js
  const response = await fetch(code + '/_api/mongoose-model', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Mike',
      age: 28,
      favoriteFoods: ['pizza', 'cheese']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Mike', '"model.name" is not what expected');
  assert.equal(data.age, '28', '"model.age" is not what expected');
  assert.isArray(
    data.favoriteFoods,
    '"model.favoriteFoods" is not an Array'
  );
  assert.include(
    data.favoriteFoods,
    'pizza',
    '"model.favoriteFoods" does not include the expected items'
  );
  assert.include(
    data.favoriteFoods,
    'cheese',
    '"model.favoriteFoods" does not include the expected items'
  );
```

