---
id: 587d7fb8367417b2b2512c11
title: 使用 model.remove() 刪除多個文件
challengeType: 2
forumTopicId: 301538
dashedName: delete-many-documents-with-model-remove
---

# --description--

`Model.remove()` 對刪除所有符合指定條件的文件很有用。

# --instructions--

修改 `removeManyPeople` 函式（程式），使用 `Model.remove()` 刪除所有名稱在變數 `nameToRemove` 內的人。將其傳遞給一個查詢文件，並設定 `name` 欄位，以及一個回呼。

**注意：** `Model.remove()` 不會傳回被刪除的文件，而是包含操作結果及受影響條款數量的 JSON 物件。別忘了將它傳給 `done()` 回呼，因為我們在測試中會使用它。

# --hints--

一次刪除多個項目應該成功

```js
  const response = await fetch(code + '/_api/remove-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Mary', age: 16, favoriteFoods: ['lollipop'] },
      { name: 'Mary', age: 21, favoriteFoods: ['steak'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(!!data.ok, 'The mongo stats are not what expected');
  assert.equal(
    data.n,
    2,
    'The number of items affected is not what expected'
  );
  assert.equal(data.count, 0, 'the db items count is not what expected');
```

