---
id: 587d7fb8367417b2b2512c10
title: 使用 model.findByIdAndRemove 刪除一個文件
challengeType: 2
forumTopicId: 301539
dashedName: delete-one-document-using-model-findbyidandremove
---

# --description--

`findByIdAndRemove` 和 `findOneAndRemove` 類似於先前的更新方法。它們會將被移除的文件傳遞給資料庫。照常，使用函式引數 `personId` 作為搜尋鍵。

# --instructions--

修改 `removeById` 函式（程式），以透過該人的 `_id` 刪除一個人。你應該使用 `findByIdAndRemove()` 或 `findOneAndRemove()` 其中一個方法。

# --hints--

刪除一個項目應該成功

```js
  const response = await fetch(code + '/_api/remove-one-person', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Jason Bourne',
      age: 36,
      favoriteFoods: ['apples']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Jason Bourne', 'item.name is not what expected');
  assert.equal(data.age, 36, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['apples'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0);
  assert.equal(data.count, 0, 'the db items count is not what expected');
```

