---
id: 587d7fb9367417b2b2512c12
title: 串鏈搜尋查詢輔助工具以縮小搜尋結果
challengeType: 2
forumTopicId: 301533
dashedName: chain-search-query-helpers-to-narrow-search-results
---

# --description--

如果你沒有將回呼作為最後一個引數傳遞給 `Model.find()`（或其他搜尋方法），查詢將不會被執行。你可以將查詢儲存在變數中以供稍後使用。這種類型的物件讓你能夠使用串鏈語法來建置查詢。實際的資料庫搜尋會在你最終串鏈 `.exec()` 方法時執行。你必須將回呼傳遞給這個最後的方法。有許多查詢輔助工具，這裡我們將使用最常用的。

# --instructions--

修改 `queryChain` 函式（程式）以尋找喜歡由名為 `foodToSearch` 的變數指定的食物的人。依照 `name` 排序，將結果限制為兩份文件，並隱藏他們的年齡。串鏈 `.find()`、`.sort()`、`.limit()`、`.select()`，然後是 `.exec()`。將 `done(err, data)` 回呼傳遞給 `exec()`。

# --hints--

串鏈查詢輔助工具應該成功

```js
  const response = await fetch(code + '/_api/query-tools', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Pablo', age: 26, favoriteFoods: ['burrito', 'hot-dog'] },
      { name: 'Bob', age: 23, favoriteFoods: ['pizza', 'nachos'] },
      { name: 'Ashley', age: 32, favoriteFoods: ['steak', 'burrito'] },
      { name: 'Mario', age: 51, favoriteFoods: ['burrito', 'prosciutto'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data.length,
    2,
    'the data array length is not what expected'
  );
  assert.notProperty(
    data[0],
    'age',
    'The returned first item has too many properties'
  );
  assert.equal(
    data[0].name,
    'Ashley',
    'The returned first item name is not what expected'
  );
  assert.notProperty(
    data[1],
    'age',
    'The returned second item has too many properties'
  );
  assert.equal(
    data[1].name,
    'Mario',
    'The returned second item name is not what expected'
  );
```

