---
id: 587d7fb6367417b2b2512c06
title: 安裝並設定 Mongoose
challengeType: 2
forumTopicId: 301540
dashedName: install-and-set-up-mongoose
---

# --description--

處理這些挑戰將涉及你使用以下其中一種方法來撰寫你的程式碼：

- 複製<a href="https://github.com/freeCodeCamp/boilerplate-mongomongoose/" target="_blank" rel="noopener noreferrer nofollow">這個 GitHub 倉庫</a>並在區域端完成這些挑戰。
- 使用你選擇的網站建置工具來完成專案。務必將我們 GitHub 倉庫中的所有檔案納入其中。

在這個挑戰中，你將設定一個 MongoDB Atlas 資料庫並匯入連接它所需的軟體包。

請依照<a href='https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/' target="_blank" rel="noopener noreferrer nofollow">此教學</a>設定 MongoDB Atlas 上的託管資料庫。

# --instructions--

`mongoose@^5.11.15` 已經被添加到你的專案的 `package.json` 檔案中。首先，在 `myApp.js` 中以 `mongoose` 的名稱引入 mongoose。接著，創建一個 `.env` 檔案並為它添加一個 `MONGO_URI` 變數。它的值應該是你的 MongoDB Atlas 資料庫 URI。請務必用單引號或雙引號包圍該 URI，並且記住在環境變數中的 `=` 兩邊不能有空格。例如，`MONGO_URI='VALUE'`。

完成後，請在你的 `myApp.js` 檔案中使用以下語法，透過呼叫 `connect` 方法來連接資料庫：

```js
mongoose.connect(<Your URI>, { useNewUrlParser: true, useUnifiedTopology: true });
```

# --hints--

`mongoose version ^5.11.15` 相依性應該在 `package.json` 中。

```js
  const response = await fetch(code + '/_api/file/package.json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  var packJson = JSON.parse(data);
  assert.property(packJson.dependencies, 'mongoose');
  assert.match(
    packJson.dependencies.mongoose,
    /^\^5\.11\.15/,
    'Wrong version of "mongoose". It should be ^5.11.15'
  );
```

`mongoose` 應該連接到資料庫

```js
  const response = await fetch(code + '/_api/is-mongoose-ok');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.isMongooseOk, 'mongoose is not connected');
```

