---
id: 698c50535b54b2d75fa985fb
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`loop` 속성은 재생이 완료되면 오디오를 다시 시작합니다. 반복 재생되는 인터넷 밈을 생각하세요. `loop` 속성을 생략하면 오디오가 한 번만 재생됩니다.

`loop` 속성은 불리언 속성이며 값이 필요 없습니다.

`loop` 속성을 `audio` 요소에 추가하세요.

# --hints--

`audio` 요소에 `loop` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
