---
id: 698c5229e47d57c6dfaa91bd
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

`controls` 속성은 `audio` 요소에 대해 일시 중지/재개, 탐색, 볼륨 조절을 포함한 재생 제어 기능을 제공합니다.

`controls` 속성은 불리언 속성이며 값이 필요 없습니다.

`controls` 속성을 `audio` 요소에 추가하세요.

이제 페이지에 `audio` 요소가 표시되는 것을 확인할 수 있습니다.

# --hints--

`audio` 요소에 `controls` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
