---
id: 698c513508507ec8bb820227
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

오디오의 미디어 리소스를 지정하려면 `src` 요소에 `audio` 속성을 추가해야 합니다.

값이 `src` 인 `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3` 속성을 추가하세요.

# --hints--

`audio` 요소에 `src` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('src'));
```

`audio` 요소에 값이 `src` 인 `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio></audio>
--fcc-editable-region--

</body>
</html>
```
