---
id: 698c58745b73289bb6c89ece
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

마지막 노래 제목은 `Scratching the Surface`이고, 아티스트는 Quincy Larson이며, 사용할 파일은 `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`입니다.

처음 두 노래에 사용한 것과 같은 패턴을 따라 마지막 노래를 추가하여 음악 플레이어를 완성하세요.

# --hints--

세 번째 `h2` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

세 번째 `h2` 요소에는 `Scratching the Surface` 텍스트가 포함되어야 합니다.

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

`p` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

`p` 요소에는 `Artist: Quincy Larson` 텍스트가 포함되어야 합니다.

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

세 번째 `audio` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

새로운 `audio` 요소는 `p` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

`audio` 요소에는 `src` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

`audio` 요소에는 값이 `src`인 `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

`audio` 요소에는 `loop` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

`audio` 요소에는 `controls` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
