---
id: 698c4f4e965bd853ad5084e1
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 아래에 `h2` 요소를 추가하고 텍스트 `Can't Stay Down`를 넣으세요. 이것은 첫 번째 노래의 제목입니다.

`h2` 요소 아래에 `p` 요소를 추가하고 텍스트 `Artist: Quincy Larson`를 넣으세요.

# --hints--

`h2` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h2'));
```

`h2` 요소는 `h1` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('h1 + h2'));
```

`h2` 요소에는 텍스트 `Can't Stay Down`가 포함되어야 합니다.

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

`p` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('p'));
```

`p` 요소에는 텍스트 `Artist: Quincy Larson`가 포함되어야 합니다.

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
