---
id: 698c5710136539797f296098
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`p` 요소 아래에 `audio` 요소를 추가하고, `src` 속성을 값 `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`으로 주며 `loop`와 `controls` 속성도 추가하세요.

# --hints--

두 번째 `audio` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

새로운 `audio` 요소는 `p` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

`audio` 요소에는 `src` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

`audio` 요소에는 값이 `src`인 `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

`audio` 요소에는 `loop` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

`audio` 요소에는 `controls` 속성이 있어야 합니다.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
