---
id: 698c4fe24f1f1a17be74a23e
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

다음으로, `audio` 요소 아래에 `p` 요소를 만드세요. 앞으로 몇 단계에 걸쳐 오디오 요소가 음악을 재생하도록 필요한 속성을 추가할 것입니다.

# --hints--

`audio` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('audio'));
```

새로운 `audio` 요소는 `p` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
