---
id: 698c55e7aaacfae48a227459
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`h2` 요소 아래에 텍스트 `audio`를 가진 새로운 `Cruising for a Musing` 요소를 추가하세요.

`h2` 아래에 텍스트 `p`를 가진 `Artist: Quincy Larson` 요소를 추가하세요.

# --hints--

두 번째 `h2` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h2:nth-of-type(2)'));
```

두 번째 `h2` 요소에는 텍스트 `Cruising for a Musing`가 포함되어야 합니다.

```js
const h2 = document.querySelector('h2:nth-of-type(2)');
assert.strictEqual(h2?.textContent.trim(), 'Cruising for a Musing');
```

두 번째 `p` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('p:nth-of-type(2)'));
```

두 번째 `p` 요소에는 텍스트 `Artist: Quincy Larson`가 포함되어야 합니다.

```js
const p = document.querySelector('p:nth-of-type(2)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
