---
id: 68af276c60e009597dd861e7
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Erstellen Sie schließlich eine letzte `if`-Anweisung, um zu überprüfen, ob `duration` nicht größer als `0` ist, und lösen Sie eine `ValueError` mit der Nachricht `Duration must be positive` aus.

# --hints--

Sie sollten eine `ValueError` mit der Nachricht `Duration must be positive` auslösen, wenn `duration` keine positive Zahl ist.

```js
({
  test: () => runPython(`
  _msg = 'Duration must be positive'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1999, 'The Wachowskis', 0)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 1999, 'The Wachowskis', -1)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Sie sollten keine `ValueError` auslösen, wenn `duration` eine gültige Zahl ist.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid duration'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
--fcc-editable-region--
        
--fcc-editable-region--
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
