---
id: 68af274560e009597dd861e3
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Erstellen Sie nun eine Instanz von `Movie`, indem Sie Argumente Ihrer Wahl übergeben, und weisen Sie sie einer Variablen namens `movie1` zu. Geben Sie anschließend `movie1` in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `movie1` haben.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

Sie sollten eine Instanz von `Movie` der Variablen `movie1` zuweisen.

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

Sie sollten `movie1` in der Konsole ausgeben.

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

--fcc-editable-region--

--fcc-editable-region--
```
