---
id: 68af274e60e009597dd861e4
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

In den vorherigen Schritten haben Sie ein einfaches Filmobjekt erstellt. Jetzt fügen Sie eine Validierung hinzu, um sicherzustellen, dass Filme nicht mit ungültigen Attributen instanziiert werden können.

Erstellen Sie am Anfang der `__init__`-Methode eine `if`-Anweisung, die überprüft, ob der `title` entweder leer ist oder nur Leerzeichen enthält (Sie können die `strip`-Methode verwenden, um die Bedingung zu formulieren).

Falls ja, lösen Sie eine `ValueError` mit der Nachricht `Title cannot be empty` aus.

# --hints--

Sie sollten eine `ValueError` mit der Nachricht `Title cannot be empty` auslösen, wenn `title` entweder leer ist oder nur Leerzeichen enthält.

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Sie sollten keine `ValueError` auslösen, wenn `title` ein gültiger String ist.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
