---
id: 68af277e60e009597dd861e9
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Erstellen Sie nun eine neue Klasse namens `MediaCatalogue` und geben Sie ihr eine `__init__`-Methode, die ein `items`-Attribut auf eine leere Liste setzt. Dies wird die Liste der Elemente im Katalog darstellen.

# --hints--

Sie sollten eine Klasse namens `MediaCatalogue` haben.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("MediaCatalogue")`)
})
```

Ihre `MediaCatalogue`-Klasse sollte eine `__init__`-Methode haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_class("MediaCatalogue").has_function("__init__")
    `)
})
```

Ihre `__init__`-Methode sollte das `items`-Attribut auf eine leere Liste setzen.

```js
({
    test: () => runPython(`
    _catalogue = MediaCatalogue()
    assert _catalogue.items == []
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
