---
id: 68af272f60e009597dd861e1
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop erstellen Sie einen Medienkatalog. Beginnen Sie damit, eine Klasse mit dem Namen `Movie` zu definieren.

Geben Sie Ihrer Klasse eine `__init__`-Methode mit den folgenden Parametern: `self`, `title`, `year`, `director` und `duration`. Weisen Sie innerhalb der `__init__`-Methode die Parameter `title`, `year`, `director` und `duration` den Instanzattributen mit demselben Namen zu.

# --hints--

Sie sollten eine Klasse mit dem Namen `Movie` erstellen.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

Ihre `Movie`-Klasse sollte eine `__init__`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

Ihre `__init__`-Methode sollte die folgenden Parameter haben: `self`, `title`, `year`, `director` und `duration`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

Sie sollten `title` in Ihrer `self.title`-Methode `__init__` zuweisen.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

Sie sollten `year` in Ihrer `self.year`-Methode `__init__` zuweisen.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

Sie sollten `director` in Ihrer `self.director`-Methode `__init__` zuweisen.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

Sie sollten `duration` in Ihrer `self.duration`-Methode `__init__` zuweisen.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
