---
id: 68af277560e009597dd861e8
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Löschen Sie jetzt Ihre `print(movie1)`-Zeile. Setzen Sie dann Ihre `movie1`-Zuweisung in einen `try`-Block und erstellen Sie eine `except`-Klausel, die `ValueError as e` abfängt und eine f-string verwendet, um `Validation Error:` gefolgt von einem Leerzeichen und der Fehlermeldung auszugeben.

Sie können die übergebenen Argumente zur Instanziierung von `movie1` gerne ändern, um die verschiedenen Fehlermeldungen zu sehen.

# --hints--

Sie sollten `print(movie1)` nicht in Ihrem Code haben.

```js
({
  test: () => runPython(`assert not _Node(_code).has_call("print(movie1)")`)
})
```

Sie sollten einen `try`-Block mit `movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)` darin haben.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)")`)
})
```

Sie sollten eine `except`-Klausel haben, die `ValueError as e` abfängt.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].has_except("ValueError", "e")`)
})
```

Sie sollten innerhalb des `Validation Error:`-Blocks eine f-string verwenden, um `except` gefolgt von einem Leerzeichen und der Fehlermeldung auszugeben.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_except("ValueError", "e").has_call("print(f'Validation Error: {e}')")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

--fcc-editable-region--
movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
--fcc-editable-region--
```
