---
id: 68af275860e009597dd861e5
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Der erste Film wurde 1895 hergestellt. Erstellen Sie eine weitere `if`-Anweisung, die überprüft, ob `year` kleiner als `1895` ist, und lösen Sie eine `ValueError` mit der Nachricht `Year must be 1895 or later` aus.

# --hints--

Sie sollten eine `ValueError` mit der Nachricht `Year must be 1895 or later` auslösen, wenn `year` kleiner als `1895` ist.

```js
({
  test: () => runPython(`
  _msg = 'Year must be 1895 or later'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1894, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 0, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Sie sollten keine `ValueError` auslösen, wenn `year` eine gültige Zahl ist.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1895, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid year'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
--fcc-editable-region--
        
--fcc-editable-region--
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
