---
id: 68af273960e009597dd861e2
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Sie möchten eine gut lesbare Ausgabe zum Debuggen und Ausgeben. Erstellen Sie eine `__str__`-Methode, die einen String im Format `title (year) - duration min, director` zurückgibt (wobei `title`, `year`, `duration` und `director` durch die Werte der entsprechenden Attribute ersetzt werden sollten), sodass das Ausgeben eines Films eine schöne Zusammenfassung ergibt.

# --hints--

Ihre `Movie`-Klasse sollte eine `__str__`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

Ihre `__str__`-Methode sollte einen String im Format `title (year) - duration min, director` zurückgeben, wobei `title`, `year`, `duration` und `director` durch die Werte der entsprechenden Attribute ersetzt werden sollten.

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    
--fcc-editable-region--
```
