---
id: 68b704862dd6559e890c2a86
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Erstellen Sie nach der `if`-Anweisung eine Variable namens `result`. Dies wird der String sein, der von der Methode zurückgegeben wird.

Weisen Sie ihr also einen f-String im Format `Media Catalogue (N items):` zu. Ersetzen Sie `N` durch die Anzahl der Elemente in der `items`-Liste und lassen Sie den String mit zwei Zeilenumbrüchen enden.

# --hints--

Sie sollten eine Variable namens `result` in Ihrer `__str__`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_variable("result")`)
})
```

Sie sollten `f'Media Catalogue ({len(self.items)} items):\n\n'` der Variablen `result` zuweisen.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_stmt("result = f'Media Catalogue ({len(self.items)} items):\\\\n\\\\n'")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)
--fcc-editable-region--
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'
        
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
