---
id: 68b804f3be853a355f055fef
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Die `super`-Funktion wird verwendet, um auf die Elternklasse zu verweisen, ohne sie explizit zu benennen. Zum Beispiel ruft `super().__init__(arg1, arg2, ...)` die in der Elternklasse definierte `__init__`-Methode auf.

Im Moment haben Sie die von `__init__` geerbte `Movie`-Methode effektiv überschrieben. Das bedeutet, dass ein `TVSeries`-Objekt nur ein `seasons`-Attribut und ein `total_episodes`-Attribut besitzt. Das ist nicht das, was Sie wollen.

Um den gesamten in der `Movie`-Klasse definierten Code in der `__init__` beizubehalten, verwenden Sie `super`, um die `__init__` der Elternklasse aufzurufen. Beachten Sie, dass Sie die Argumente übergeben müssen, die die Methode benötigt. In diesem Fall `title`, `year`, `director` und `duration`.

# --hints--

Sie sollten `super().__init__(title, year, director, duration)` innerhalb Ihrer `__init__`-Methode aufrufen.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("super().__init__(title, year, director, duration)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
--fcc-editable-region--
        
--fcc-editable-region--
        
        self.seasons = seasons
        self.total_episodes = total_episodes

class MediaCatalogue:
    def __init__(self):
        self.items = []
    
    def add(self, media_item):
        self.items.append(media_item)
    
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'
        
        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
