---
id: 66c644c08ebcd2ef90c750bd
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

前のレッスンでは、`push` メソッドを使って配列の末尾にアイテムを追加する方法を学びました。

```js
fruits.push('pear');
```

`push` メソッドを使って、文字列 `"Apples"` を `shoppingList` 配列に追加してください。

# --hints--

`shoppingList` 配列には1つのアイテムが入っているはずです。

```js
assert.lengthOf(shoppingList, 1);
```

`shoppingList` 配列には文字列 `"Apples"` が含まれているはずです。スペルや大文字・小文字の間違いがないかもう一度確認してください。

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
