---
id: 66c64095f61166ede6e39a84
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

この買い物リストでは、購入する必要があるアイテムを表すために配列を使います。

前のレッスンでは、このように配列を作成する方法を学びました。

```js
const fruits = ['apple', 'banana', 'orange'];
```

このステップでは、`shoppingList` という変数を作成し、空の配列を代入してください。

# --hints--

`shoppingList` という変数があるはずです。

```js
assert.isNotNull(shoppingList);
```

`shoppingList` 変数に配列を代入してください。

```js
assert.isArray(shoppingList);
```

配列は空である必要があります。

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
