---
id: 66c6491fe4c8e0f16845425f
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

リストにもう一つ果物を追加する時間です。

先ほどと同じ配列メソッドを使って、文字列 `"Grapes"` を `shoppingList` 配列の末尾に追加してください。

次に、`console.log` を追加し、`getShoppingListMsg` の中で `shoppingList` 配列を引数として `console.log` 関数を呼び出し、更新されたリストがコンソールに表示されるようにしてください。

# --hints--

あなたの `shoppingList` 配列には今、2つのアイテムが含まれているはずです。

```js
assert.lengthOf(shoppingList, 2);
```

あなたの `shoppingList` 配列には文字列 `"Grapes"` が含まれているはずです。

```js
assert.equal(shoppingList[1], "Grapes");
```

`getShoppingListMsg(shoppingList)` の中で `console.log` 関数を呼び出す必要があります。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
