---
id: 66c64210413532ee9d3bd342
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

このワークショップの次の部分では、食料品リストに食品アイテムを追加する練習をします。

まず、`console.log`を使ってメッセージ`"It will be nice to have some fruit to eat."`をログに出力してください。

# --hints--

メッセージ`"It will be nice to have some fruit to eat."`をコンソールにログ出力する必要があります。

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
