---
id: 66c63bfa4685e4e3843fa0dc
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、配列について学びながら、食料品の買い物リストを作成していきます。

まず、コンソールに文字列 `console.log` を出力する `"Grocery shopping list"` を追加してください。

# --hints--

コード内に `console.log` があるはずです。

```js
assert.match(code, /console\.log(.*)/);
```

`console.log` はメッセージ `"Grocery shopping list"` を出力する必要があります。スペルや大文字・小文字の間違いがないかもう一度確認してください。

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
