---
id: 66c64ad16796c7f2419b45c5
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

いよいよ買い物リストの先頭にアイテムを追加する時間です。

まず、コンソールにメッセージ `console.log` を記録する `"It looks like we need to get some cooking oil."` を追加してください。

# --hints--

コンソールにメッセージ `"It looks like we need to get some cooking oil."` を記録してください。

```js
assert.match(code, /console\.log\((['"])(It\s+looks\s+like\s+we\s+need\s+to\s+get\s+some\s+cooking\s+oil\.)\1\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
