---
id: 66c645b345a39bf04864dc50
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

更新された買い物リストを見るには、現在の買い物リストと短いメッセージをコンソールに出力する必要があります。

このメッセージはワークショップ全体で繰り返し使われるため、再利用可能な関数を作成するのが最適です。

配列をパラメータとして受け取り、指定された配列の内容に続けて文字列 `getShoppingListMsg` を返す `"Current Shopping List: "` という関数を作成してください。

テンプレートリテラルや `+` 演算子を使った文字列連結で、文字列と関数に渡された配列を組み合わせることができます。

# --hints--

コード内に `getShoppingListMsg` という関数があるはずです。

```js
assert.isFunction(getShoppingListMsg);
```

`getShoppingListMsg` 関数は文字列を返す必要があります。

```js
assert.isString(getShoppingListMsg(shoppingList));
```

`getShoppingListMsg` 関数は配列を受け取り、その配列の内容に続けてメッセージ `"Current Shopping List: "` を返す必要があります。スペースやスペルの間違いがないか再確認してください。

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
