---
id: 66c73a7798f6f62b2ae58f22
title: ステップ16
challengeType: 1
dashedName: step-16
---

# --description--

今度は食料品リストの先頭にアイテムを追加する時間です。

まず、メッセージ`console.log`をログに出力する`"It might be nice to get a dessert."`文を追加してください。

その`console`文の下に、正しい配列メソッドを使って文字列`"Chocolate Cake"`を`shoppingList`配列の先頭に追加してください。

最後に、`console.log`を追加し、`getShoppingListMsg`の中で`shoppingList`配列を引数にして`console.log`関数を呼び出し、更新されたリストがコンソールにログされるのを確認してください。

# --hints--

メッセージ`console.log`をログに出力する`"It might be nice to get a dessert."`があるはずです。

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

`shoppingList`配列は合計で`6`個のアイテムを持っているはずです。

```js
assert.lengthOf(shoppingList, 6);
```

文字列`unshift`を`"Chocolate Cake"`配列の先頭に追加するには`shoppingList`メソッドを使うべきです。

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

`getShoppingListMsg(shoppingList)`の中で`console.log`を呼び出すべきです。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
