---
id: 66c73a0c5b264f2a75164d94
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

更新された `shoppingList` 配列をコンソールにログ出力する時です。

`console.log` の中に `getShoppingListMsg` を追加し、`shoppingList` 配列を引数として `console.log` 関数を呼び出して、更新されたリストがコンソールにログ出力されるようにしてください。

# --hints--

`getShoppingListMsg(shoppingList)` の中で `console.log` を呼び出す必要があります。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 5);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();

--fcc-editable-region--

--fcc-editable-region--
```
