---
id: 66c726c34ecf1e238aa9d7d5
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

前のレッスンでは、`unshift()` メソッドを使って配列の先頭に要素を追加する方法を学びました。

ここで `unshift()` メソッドの使い方をもう一度確認しましょう。

```js
array.unshift(item1, item2, ..., itemX);
```

`unshift()` メソッドを使って、文字列 `"Vegetable Oil"` を `shoppingList` 配列の先頭に追加してください。

# --hints--

`shoppingList` 配列には3つの要素があるはずです。

```js
assert.lengthOf(shoppingList, 3);
```

`unshift()` メソッドを使って、`"Vegetable Oil"` を `shoppingList` 配列の先頭に追加してください。

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
