---
id: 66c72f4d0528bd268a82107b
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

更新された `shoppingList` 配列をコンソールにログ出力する時です。

もう一つ `console.log` を追加し、`getShoppingListMsg` の中で `shoppingList` 配列を引数にして `console.log` 関数を呼び出し、更新されたリストがコンソールにログ出力されるようにしてください。

# --hints--

`getShoppingListMsg(shoppingList)` の中で `console.log` を呼び出す必要があります。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 4);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");

--fcc-editable-region--

--fcc-editable-region--
```
