---
id: 66c730183f4020275cbf0611
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

このワークショップの次の部分では、配列の末尾からアイテムを削除する方法をレビューします。

まず、メッセージ `console.log` をログに出力する `"This looks like too much junk food."` を追加してください。

# --hints--

メッセージ `"This looks like too much junk food."` をログに出力してください。

```js
assert.match(code, /console\.log\((['"])(This\s+looks\s+like\s+too\s+much\s+junk\s+food.)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
