---
id: 66c72b0ffbc5522525768558
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

前のステップでは、`push` メソッドを使って配列の末尾にアイテムを追加する方法をレビューしました。

しかし、`push` メソッドは複数の引数を受け取れるので、このように配列の末尾に複数のアイテムを追加できます。

```js
array.push(item1, item2, item3);
```

このステップでは、`push` メソッドを使って `"Popcorn", "Beef Jerky", "Potato Chips"` という文字列を `shoppingList` 配列に追加してください。

順番が重要なので、リストにある順番でアイテムを追加することを忘れないでください。

# --hints--

あなたの `shoppingList` 配列には合計で `6` 個の文字列が含まれているはずです。

```js
assert.lengthOf(shoppingList, 6);
```

あなたの `shoppingList` には文字列 `"Popcorn"` が含まれているはずです。

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

あなたの `shoppingList` には文字列 `"Beef Jerky"` が含まれているはずです。

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

あなたの `shoppingList` には文字列 `"Potato Chips"` が含まれているはずです。

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
