---
id: 66c74079c30b1c2d166cb9a4
title: ステップ 18
challengeType: 1
dashedName: step-18
---

# --description--

前のレッスンでは、`shift` メソッドを使って配列の先頭からアイテムを削除する方法を学びました。

ここで `shift` メソッドの使い方をおさらいしましょう。

```js
const array = [1, 2, 3, 4, 5];
array.shift();

// Result: [2, 3, 4, 5]
console.log(array); 
```

`shift` メソッドを使って、`shoppingList` 配列の最初のアイテムを削除してください。

# --hints--

`shoppingList` 配列には合計で `5` 個のアイテムがあるはずです。

```js
assert.lengthOf(shoppingList, 5);
```

`shift` 配列の最初のアイテムを削除するには、`shoppingList` メソッドを使うべきです。

```js
assert.strictEqual(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

--fcc-editable-region--

--fcc-editable-region--
```
