---
id: 66c730ee6ae076281721d0b9
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

前のレッスンでは、`pop` メソッドを使って配列の末尾からアイテムを削除する方法を学びました。

ここで `pop` メソッドの使い方をおさらいしましょう。

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

`pop` メソッドを使って、`shoppingList` 配列の最後のアイテムを削除してください。

# --hints--

あなたの `shoppingList` 配列には、今 `5` 個のアイテムが入っているはずです。

```js
assert.lengthOf(shoppingList, 5);
```

あなたの `shoppingList` 配列には、文字列 `"Potato Chips"` がもう含まれていないはずです。

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
