---
id: 66c72a55418cc9247b710827
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

次に、`console.log` を追加し、`getShoppingListMsg` の中で `shoppingList` 配列を引数として `console.log` 関数を呼び出して、更新されたリストがコンソールに表示されるようにしてください。

# --hints--

`getShoppingListMsg(shoppingList)` の中で `console.log` を呼び出す必要があります。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
