---
id: 66cbe2319d3845545a293a0b
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

コンソールに記録されたメッセージを確認する時間です。

`console.log`の中で`getShoppingListMsg`を追加し、`shoppingList`配列を引数として`console.log`関数を呼び出して、コンソールにメッセージが記録されるのを確認してください。

# --hints--

`getShoppingListMsg(shoppingList)`の中で`console.log`関数を呼び出す必要があります。

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
