---
id: 67fccc2463253b213a000142
title: Створення мультимедійного плеєра
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

У попередніх уроках ви ознайомилися з роботою з елементами `audio` та `video`. У цій лабораторній роботі ви створите мультимедійний плеєр, який відтворюватиме аудіодоріжку `audio` та відео `video` з транскриптом.

Для елемента `audio` потрібно додати елемент `source`, який використовується для вказання медіафайлу.

Ось приклад:

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

Елемент `source` також можна використовувати всередині елемента `video` ось так:

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати елемент `h1` для основного заголовка сторінки.
2. Ви повинні мати три елементи `section`.
3. Всередині першого елемента `section` має бути елемент `h2` для назви пісні, що відтворюється.
4. Під елементом `h2` має бути елемент `audio` з атрибутом `controls` та атрибутом `aria-label`.
5. Всередині елемента `audio` має бути елемент `source` з атрибутом `src`, що вказує на аудіофайл, та атрибутом `type`. Ви можете використати цей URL аудіо, якщо хочете: `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. Всередині другого елемента `section` має бути елемент `h2` для назви відео, що відтворюється.
7. Під елементом `h2` має бути елемент `video` з атрибутами `controls`, `width` та атрибутом `aria-label`.
8. Всередині елемента `video` має бути елемент `source` з атрибутом `src`, що вказує на відеофайл, та атрибутом `type`. Ви можете використати цей URL відео, якщо хочете: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. Під елементом `source` має бути елемент `track` з атрибутом `src`, що вказує на файл субтитрів, а також атрибутами `kind`, `srclang` і `label`. Ви можете використати цей URL субтитрів, якщо хочете: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. Всередині третього елемента `section` має бути елемент `h2` для заголовка розділу, наприклад, "Transcript".
11. Під елементом `h2` має бути елемент `p` з транскриптом відео.

# --hints--

Ви повинні мати елемент `h1`.

```js
assert.isNotNull(document.querySelector("h1"));
```

Ваш `h1` має містити текст, який представляє основний заголовок сторінки.

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

Ви повинні мати три елементи `section`.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

Перший `section` має містити елемент `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

Ваш `h2` має містити текст, який представляє назву пісні, що відтворюється.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

Перший `section` має містити елемент `audio`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

Ваш елемент `audio` має мати атрибут `controls`.

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

Ваш елемент `audio` має мати атрибут `aria-label` з текстом, що описує аудіо.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

Ваш елемент `audio` має містити елемент `source`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

Ваш елемент `source` має мати атрибут `src`, що вказує на аудіофайл.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

Ваш елемент `source` має мати атрибут `type` зі значенням, що вказує тип аудіо.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

Другий `section` має містити елемент `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

Ваш `h2` має містити текст, який представляє назву відео, що відтворюється.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

Другий `section` має містити елемент `video`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

Ваш елемент `video` має мати атрибут `controls`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

Ваш елемент `video` має мати атрибут `width` зі значенням, що вказує ширину відео.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

Елемент `video` має містити елемент `source`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

Елемент `source` має мати атрибут `src`, що вказує на відеофайл.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

Елемент `source` має мати атрибут `type` зі значенням, що вказує тип відео.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

Ваш елемент `video` має містити елемент `track`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

Ваш елемент `track` має мати атрибут `src`, що вказує на файл субтитрів.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

Ваш елемент `track` має мати атрибут `kind` з текстом, що описує тип файлу.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

Ваш елемент `track` має мати атрибут `srclang` з текстом, що описує мову субтитрів.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

Ваш елемент `track` має мати атрибут `label` з текстом, що описує мову субтитрів.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

Третій `section` має містити елемент `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

Ваш `h2` має містити текст, який представляє заголовок розділу.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

Третій `section` має містити елемент `p` для транскрипту.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
