---
id: 66629f407d679d3105e8317f
title: "Crie um jogo de pedra, papel e tesoura"
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
O jogo serĂĄ jogado contra o computador. VocĂȘ escreverĂĄ uma funĂ§ĂŁo que retorna aleatoriamente `"rock"`, `"paper"` ou `"scissors"`.

VocĂȘ nĂŁo precisa se preocupar com a parte de front-end do jogo. VocĂȘ apenas escreverĂĄ a lĂłgica para o jogo. Abra `script.js` e comece a seguir as instruĂ§Ă”es.

**HistĂłrias de usuĂĄrio:**

1. VocĂȘ deve ter uma funĂ§ĂŁo chamada `getComputerChoice`.

1. A funĂ§ĂŁo `getComputerChoice` deve retornar `"rock"`, `"paper"` ou `"scissors"` aleatoriamente.

**Dica:** O mĂ©todo [Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random) retorna um nĂșmero aleatĂłrio que Ă© maior ou igual a 0 e menor que 1. Pense em como vocĂȘ pode usar isso para retornar condicionalmente uma das mĂșltiplas opĂ§Ă”es.

O jogo serĂĄ jogado por um jogador humano. VocĂȘ escreverĂĄ uma funĂ§ĂŁo que recebe a escolha do usuĂĄrio e a retorna.

1. Crie uma funĂ§ĂŁo chamada `getHumanChoice`.

1. Escreva o cĂłdigo para que `getHumanChoice` retorne uma das escolhas vĂĄlidas, dependendo da entrada do usuĂĄrio.

**Dica:** Use o mĂ©todo [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt) para obter a entrada do usuĂĄrio.

O jogo acompanharĂĄ a pontuaĂ§ĂŁo do jogador. VocĂȘ escreverĂĄ variĂĄveis para acompanhar a pontuaĂ§ĂŁo do jogador.

1. Crie duas variĂĄveis chamadas `humanScore` e `computerScore` no escopo global.

1. Inicialize essas variĂĄveis com o valor de `0`.

O jogo serĂĄ jogado rodada por rodada. VocĂȘ escreverĂĄ uma funĂ§ĂŁo que recebe as escolhas do jogador humano e do computador como argumentos, joga uma Ășnica rodada, incrementa a pontuaĂ§ĂŁo do vencedor da rodada e registra o anĂșncio do vencedor.

1. Crie uma funĂ§ĂŁo chamada `playRound`.

1. Defina dois parĂąmetros para `playRound`. ParĂąmetro um `humanChoice` e parĂąmetro dois `computerChoice`. Use esses dois parĂąmetros para receber as escolhas do humano e do computador como argumentos.

1. Certifique-se de que o parĂąmetro `humanChoice` da funĂ§ĂŁo nĂŁo faĂ§a a distinĂ§ĂŁo entre maiĂșsculas e minĂșsculas para que os jogadores possam digitar `"rock"`, `"ROCK"`, `"RocK"` ou outras variaĂ§Ă”es.

1. Escreva o cĂłdigo para a funĂ§ĂŁo `playRound`, que retorna um valor de string representando o vencedor da rodada.

- Se for empate, ela deve retornar `"It's a tie!"`.

- Se o jogador vencer, deve retornar `"You win! [player choice] beats [computer choice]"`.

- Se o computador vencer, ele deve retornar `"You lose! [computer choice] beats [player choice]"`.

1. Incremente a variĂĄvel `humanScore` ou `computerScore` com base no vencedor da rodada.

O jogo terĂĄ 3 rodadas. VocĂȘ escreverĂĄ uma funĂ§ĂŁo chamada `playGame` que chama `playRound` para jogar 3 rodadas, mantĂ©m o controle das pontuaĂ§Ă”es e declara um vencedor no final.

1. Crie uma funĂ§ĂŁo chamada `playGame`.

1. Crie um laĂ§o que Ă© executado por 3 rodadas e chame a funĂ§ĂŁo `playRound` a cada vez, com as funĂ§Ă”es de escolha do humano e do computador como argumentos.

1. No final do jogo, retorne o vencedor do jogo com base em quem ganhou mais rodadas.

- Se o jogador humano ganhar mais rodadas que o computador, retorne uma mensagem que diga `"You win the game!"`.

- Se o computador ganhar mais rodadas que o jogador humano, retorne uma mensagem que diga `"You lose the game!"`.

# --hints--

VocĂȘ deve ter uma funĂ§ĂŁo chamada `getComputerChoice`

```js
assert.isFunction(getComputerChoice);
```

A funĂ§ĂŁo `getComputerChoice` deve retornar `"rock"`, `"paper"` ou `"scissors"` aleatoriamente.

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

VocĂȘ deve ter uma funĂ§ĂŁo chamada `getHumanChoice`

```js
assert.isFunction(getHumanChoice);
```

VocĂȘ deve ter duas variĂĄveis chamadas `humanScore` e `computerScore` no escopo global.

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

VocĂȘ deve ter uma funĂ§ĂŁo chamada `playRound`

```js
assert.isFunction(playRound);
```

A funĂ§ĂŁo `playRound` deve receber as escolhas do jogador humano e do computador como argumentos com os parĂąmetros `humanChoice` e `computerChoice`.

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

A funĂ§ĂŁo `playRound` nĂŁo deve distinguir maiĂșsculas de minĂșsculas para que os jogadores possam digitar `"rock"`, `"ROCK"`, `"RocK"` ou outras variaĂ§Ă”es.

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

A funĂ§ĂŁo `playRound` deve retornar a string `"It's a tie!"` se as escolhas do humano e do computador forem as mesmas.

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

Sua funĂ§ĂŁo `playRound` deve retornar a string `"You win! [player choice] beats [computer choice]"` se o jogador vencer.

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

Sua funĂ§ĂŁo `playRound` deve retornar a string `"You lose! [computer choice] beats [player choice]"` se o computador vencer.

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

A funĂ§ĂŁo `playRound` deve incrementar a variĂĄvel `humanScore` ou `computerScore` com base no vencedor da rodada.

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

VocĂȘ deve ter uma funĂ§ĂŁo chamada `playGame`.

```js
assert.isFunction(playGame);
```

VocĂȘ deve usar um laĂ§o para jogar 3 rodadas.

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

VocĂȘ deve retornar o vencedor do jogo com base em quem ganhou mais rodadas.

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```

