---
id: a2f1d72d9b908d0bd72bb9f6
title: 建立一個 Person
challengeType: 1
forumTopicId: 16020
dashedName: make-a-person
---

# --description--

使用以下方法填寫物件建構式：

```js
getFirstName()
getLastName()
getFullName()
setFirstName(first)
setLastName(last)
setFullName(first, last)
```

執行測試以查看每個方法的預期輸出。這些方法必須是與該物件互動的唯一可用方式。每個測試都會宣告一個新的 `Person` 實體為 `new Person('Bob', 'Ross')`。

# --hints--

你不應該更改函式標記。

```js
assert.match(code, /const\s+Person\s*=\s*function\s*\(\s*first\s*,\s*last\s*\)\s*{/);
```

你不應該重新指定 `first` 參數。

```js
assert.notMatch(code, /\bfirst\s*=\s*/);
```

你不應該重新指定 `last` 參數。

```js
assert.notMatch(code, /\blast\s*=\s*/);
```

不應該添加任何屬性。`Object.keys(Person).length` 應該始終傳回 6。

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
_person.setLastName('Curry');
_person.setFullName('John', 'Smith');
assert.lengthOf(Object.keys(_person), 6);
```

你應該能夠實例化你的 `Person` 物件。

```js
const _person = new Person('Bob', 'Ross');
assert.instanceOf(_person, Person);
```

你的 `Person` 物件不應該有 `firstName` 屬性。

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'firstName');
```

你的 `Person` 物件不應該有 `lastName` 屬性。

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'lastName');
```

`.getFirstName()` 方法應該傳回字串 `Bob`。

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFirstName(), 'Bob');
```

`.getLastName()` 應該傳回字串 `Ross`。

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getLastName(), 'Ross');
```

`.getFullName()` 方法應該傳回字串 `Bob Ross`。

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFullName(), 'Bob Ross');
```

`.setFirstName('Haskell')` 呼叫後，`.getFullName()` 方法應該傳回字串 `Haskell Ross`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
assert.strictEqual(_person.getFullName(), 'Haskell Ross');
```

`.getFullName()` 方法在呼叫 `.setLastName('Curry')` 後，應該傳回字串 `Bob Curry`。

```js
const _person = new Person('Bob', 'Ross');
_person.setLastName('Curry');
assert.strictEqual(_person.getFullName(), 'Bob Curry');
```

`.getFullName()` 方法在呼叫 `.setFullName('Haskell', 'Curry')` 之後，應該傳回字串 `Haskell Curry`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFullName(), 'Haskell Curry');
```

在呼叫 `.setFullName('Haskell', 'Curry')` 之後，`.getFirstName()` 方法應該傳回字串 `Haskell`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFirstName(), 'Haskell');
```

`.setFullName('Haskell', 'Curry')` 呼叫後，`.getLastName()` 方法應該傳回字串 `Curry`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getLastName(), 'Curry');
```

`.getFullName()` 方法在呼叫 `.setFullName('Emily Martinez', 'de la Rosa')` 後，應該傳回字串 `Emily Martinez de la Rosa`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFullName(), 'Emily Martinez de la Rosa');
```

`.setFullName('Emily Martinez', 'de la Rosa')` 呼叫後，`.getFirstName()` 屬性應該傳回字串 `Emily Martinez`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFirstName(), 'Emily Martinez');
```

`.getLastName()` 屬性在呼叫 `.setFullName('Emily Martinez', 'de la Rosa')` 後，應該傳回字串 `de la Rosa`。

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getLastName(), 'de la Rosa');
```

# --seed--

## --seed-contents--

```js
const Person = function(first, last) {
  this.getFullName = function() {
    return "";
  };
  return "";
};
```

# --solutions--

```js
const Person = function(first, last) {
  let firstName = first;
  let lastName = last;

  this.getFirstName = function(){
    return firstName;
  };

  this.getLastName = function(){
    return lastName;
  };

  this.getFullName = function(){
    return firstName + " " + lastName;
  };

  this.setFirstName = function(str){
    firstName = str;
  };


  this.setLastName = function(str){
    lastName = str;
  };

  this.setFullName = function(first, last){
    firstName = first;
    lastName = last;
  };
};
```
