---
id: a3566b1109230028080c9345
title: 加總範圍內的所有數字
challengeType: 1
forumTopicId: 16083
dashedName: sum-all-numbers-in-a-range
---

# --description--

我們會傳給你一個包含兩個數字的陣列。傳回這兩個數字的總和加上它們之間所有數字的總和。較小的數字不一定會先出現。

例如，`sumAll([4,1])` 應該傳回 `10`，因為 1 和 4（含兩端）之間所有數字的總和是 `10`。

# --hints--

`sumAll([1, 4])` 應該傳回一個數字。

```js
assert(typeof sumAll([1, 4]) === 'number');
```

`sumAll([1, 4])` 應該傳回 10。

```js
assert.deepEqual(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` 應該傳回 10。

```js
assert.deepEqual(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` 應該傳回 45。

```js
assert.deepEqual(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` 應該傳回 45。

```js
assert.deepEqual(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
function sumAll(arr) {
  return 1;
}

sumAll([1, 4]);
```

# --solutions--

```js
function sumAll(arr) {
  var sum = 0;
  arr.sort(function(a,b) {return a-b;});
  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }
  return sum;
}
```
