---
id: a5229172f011153519423690
title: 將所有奇數 `Fibonacci` 數字相加
challengeType: 1
forumTopicId: 16084
dashedName: sum-all-odd-fibonacci-numbers
---

# --description--

給定一個正整數 `num`，傳回所有小於或等於 `num` 的奇數 Fibonacci 數字的總和。

費波那契數列的前兩個數字是 0 和 1。數列中的每個額外數字都是前兩個數字的總和。費波那契數列的前七個數字是 0、1、1、2、3、5 和 8。

例如，`sumFibs(10)` 應該傳回 `10`，因為所有小於或相等於 `10` 的奇數費波那契數字是 1、1、3 和 5。

# --hints--

`sumFibs(1)` 應該傳回一個數字。

```js
assert(typeof sumFibs(1) === 'number');
```

`sumFibs(1000)` 應該傳回 1785。

```js
assert(sumFibs(1000) === 1785);
```

`sumFibs(4000000)` 應該傳回 4613732。

```js
assert(sumFibs(4000000) === 4613732);
```

`sumFibs(4)` 應該傳回 5。

```js
assert(sumFibs(4) === 5);
```

`sumFibs(75024)` 應該傳回 60696。

```js
assert(sumFibs(75024) === 60696);
```

`sumFibs(75025)` 應該傳回 135721。

```js
assert(sumFibs(75025) === 135721);
```

# --seed--

## --seed-contents--

```js
function sumFibs(num) {
  return num;
}

sumFibs(4);
```

# --solutions--

```js
function sumFibs(num) {
  var a = 1;
  var b = 1;
  var s = 0;
  while (a <= num) {
    if (a % 2 !== 0) {
      s += a;
    }
    a = [b, b=b+a][0];
  }
  return s;
}
```
