---
id: a97fd23d9b809dac9921074f
title: 引數 選用
challengeType: 1
forumTopicId: 14271
dashedName: arguments-optional
---

# --description--

創建一個函式（程式），將兩個引數相加。如果只提供一個引數，則傳回一個期望一個引數並傳回總和的函式（程式）。

例如，`addTogether(2, 3)` 應該傳回 `5`，而 `addTogether(2)` 應該傳回一個函式（程式）。

使用單一引數呼叫此傳回的函式（程式）將會傳回總和：

```js
var sumTwoAnd = addTogether(2);
```

`sumTwoAnd(3)` 傳回 `5`。

如果任一引數不是有效的數字，傳回 undefined。

# --hints--

`addTogether(2, 3)` 應該傳回 5。

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` 應該傳回 53.4。

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` 應該傳回一個函式（程式）。

```js
assert.deepEqual(typeof(addTogether(5)), 'function');
```

`addTogether(5)(7)` 應該傳回 12。

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js
function addTogether() {
  return false;
}

addTogether(2,3);
```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
