---
id: aa7697ea2477d1316795783b
title: 豬拉丁語
challengeType: 1
forumTopicId: 16039
dashedName: pig-latin
---

# --description--

Pig Latin 是一種改變英文字的方式。規則如下：

- 如果一個字以子音開頭，取出第一個子音或子音群，將它移到字尾，並加上 `ay`。

- 如果一個字以母音開頭，只需在結尾加上 `way`。

# --instructions--

將提供的字串翻譯成 Pig Latin。輸入的字串保證是全部小寫的英文字。

# --hints--

`translatePigLatin("california")` 應該傳回字串 `aliforniacay`。

```js
assert.deepEqual(translatePigLatin('california'), 'aliforniacay');
```

`translatePigLatin("paragraphs")` 應該傳回字串 `aragraphspay`。

```js
assert.deepEqual(translatePigLatin('paragraphs'), 'aragraphspay');
```

`translatePigLatin("glove")` 應該傳回字串 `oveglay`。

```js
assert.deepEqual(translatePigLatin('glove'), 'oveglay');
```

`translatePigLatin("algorithm")` 應該傳回字串 `algorithmway`。

```js
assert.deepEqual(translatePigLatin('algorithm'), 'algorithmway');
```

`translatePigLatin("eight")` 應該傳回字串 `eightway`。

```js
assert.deepEqual(translatePigLatin('eight'), 'eightway');
```

應該控制代碼處理第一個母音出現在字中間的字。`translatePigLatin("schwartz")` 應該傳回字串 `artzschway`。

```js
assert.deepEqual(translatePigLatin('schwartz'), 'artzschway');
```

應該能處理沒有母音的字。`translatePigLatin("rhythm")` 應該會傳回字串 `rhythmay`。

```js
assert.deepEqual(translatePigLatin('rhythm'), 'rhythmay');
```

# --seed--

## --seed-contents--

```js
function translatePigLatin(str) {
  return str;
}

translatePigLatin("consonant");
```

# --solutions--

```js
function translatePigLatin(str) {
  if (isVowel(str.charAt(0))) return str + "way";
  var front = [];
  str = str.split('');
  while (str.length && !isVowel(str[0])) {
    front.push(str.shift());
  }
  return [].concat(str, front).join('') + 'ay';
}

function isVowel(c) {
  return ['a', 'e', 'i', 'o', 'u'].indexOf(c.toLowerCase()) !== -1;
}
```
