---
id: af4afb223120f7348cdfc9fd
title: 繪製碎片地圖
challengeType: 1
forumTopicId: 16021
dashedName: map-the-debris
---

# --description--

根據開普勒第三定律，兩個點質量在圓形或橢圓軌道上互相繞行的軌道週期 $T$ 為：

$$
T = 2 \pi \sqrt{\frac{a^{3}}{\mu}}
$$

- $a$ 是軌道的半長軸
- $μ = GM$ 是標準重力參數
- $G$ 是萬有引力常數，
- $M$ 是質量較大的天體的質量。

傳回一個新的陣列，將元素的平均高度轉換為它們的軌道週期（以秒為單位）。

該陣列將包含格式為 `{name: 'name', avgAlt: avgAlt}` 的物件。

數值應該四捨五入到最接近的整數。被繞行的天體是地球。

地球的半徑是 6367.4447 公里，地球的 GM 值是 398600.4418 km<sup>3</sup>s<sup>-2</sup>。

# --hints--

`orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}])` 應該傳回 `[{name: "sputnik", orbitalPeriod: 86400}]`。

```js
assert.deepEqual(orbitalPeriod([{ name: 'sputnik', avgAlt: 35873.5553 }]), [
  { name: 'sputnik', orbitalPeriod: 86400 }
]);
```

`orbitalPeriod([{name: "iss", avgAlt: 413.6}, {name: "hubble", avgAlt: 556.7}, {name: "moon", avgAlt: 378632.553}])` 應該傳回 `[{name : "iss", orbitalPeriod: 5557}, {name: "hubble", orbitalPeriod: 5734}, {name: "moon", orbitalPeriod: 2377399}]`。

```js
assert.deepEqual(
  orbitalPeriod([
    { name: 'iss', avgAlt: 413.6 },
    { name: 'hubble', avgAlt: 556.7 },
    { name: 'moon', avgAlt: 378632.553 }
  ]),
  [
    { name: 'iss', orbitalPeriod: 5557 },
    { name: 'hubble', orbitalPeriod: 5734 },
    { name: 'moon', orbitalPeriod: 2377399 }
  ]
);
```

# --seed--

## --seed-contents--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  return arr;
}

orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}]);
```

# --solutions--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  const TAU = 2 * Math.PI;
  return arr.map(function(obj) {
    return {
      name: obj.name,
      orbitalPeriod: Math.round(TAU * Math.sqrt(Math.pow(obj.avgAlt+earthRadius, 3)/GM))
    };
  });
}
```
