---
id: 5d5a813321b9e3db6c106a46
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript 是一種強大的語言，允許你建造具有互動性的網站。 

*注意*：對於本課程中所有剩餘的專案，你將需要具備基本的 HTML 和 CSS 知識層級。如果你是 HTML 和 CSS 新手，請完成 [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/)。

要開始，請使用 `DOCTYPE`、`html`、`head` 和 `body` 創建你的標準 HTML 範本，然後為 `charset` 添加一個 `meta` 頁籤。為 `title` 元素添加文字 `RPG - Dragon Repeller`。包含一個 `link` 頁籤以連結你的樣式表 `styles.css` 檔案。

最後，在你的 `body` 中創建一個 `div` 元素，並將 `id` 設定為 `game`。

# --hints--

你的程式碼應該包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應該在 `DOCTYPE` 引用後面加上一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該定義文件型別為 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你應該在 `DOCTYPE` 宣告式的型別後面以 `>` 關閉。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個開啟的標記。不要忘記 `lang` 屬性。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 宣告應該放在你的 HTML 開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你應該有一個開啟的 `head` 頁籤。

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個結束的 `head` 頁籤。

```js
assert(code.match(/<\/head\s*>/i));
```

你應該有一個開啟的 `body` 頁籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個結束的 `body` 頁籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是同層元素。

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

`head` 元素應該位於 `html` 元素內。

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

`body` 元素應該位於 `html` 元素內。

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

你的程式碼應該有一個 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

你的 `meta` 元素應該有一個值為 `UTF-8` 的 `charset` 屬性。

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

你的程式碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你應該有一個結束的 `title` 頁籤。

```js
assert(code.match(/<\/title\s*>/i));
```

你的程式碼應該有一個 `link` 元素。

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

你應該有一個 `div` 元素。

```js
const div = document.querySelector('div');
assert.exists(div);
```

你的 `div` 元素應該有一個值為 `game` 的 `id` 屬性。

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

你的 `div` 元素應該位於 `body` 元素內。

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
