---
id: 62a1166ed9a56d439c0770e7
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

在你的 `#stats` 元素中創建三個 `span` 元素。為它們每個添加 `stat` 類別。然後為第一個設定文字為 `XP: 0`，第二個設定文字為 `Health: 100`，第三個設定文字為 `Gold: 50`。

# --hints--

你應該在 `#stats` 元素內有三個 `span` 元素。

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

你應該為新的三個 `span` 元素指定一個 `stat` 類別。

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

你的第一個 `.stat` 元素應該有提供的文字 `XP: 0`。

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

你的第二個 `.stat` 元素應該有提供的文字 `Health: 100`。

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

你的第三個 `.stat` 元素應該有提供的文字 `Gold: 50`。

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
