---
id: 62a115879a6d51422652cbfc
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

在你的角色扮演遊戲中，使用者將能夠追蹤他們的屬性、購買武器，並與怪物戰鬥。在你能繼續進行遊戲中互動式 `JavaScript` 部分之前，你需要先創建將顯示遊戲資訊的 `HTML` 元素。

在你的 `#game` 元素中創建四個 `div` 元素。依序為它們指定以下各自的 `id` 值：`stats`、`controls`、`monsterStats` 和 `text`。

# --hints--

你應該創建四個新的 `div` 元素。

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

你應該為其中一個新的 `div` 元素指定一個 `id` 為 `stats`。

```js
assert.exists(document.querySelector('div#stats'));
```

你應該為其中一個新的 `div` 元素指定一個 `id` 為 `controls`。

```js
assert.exists(document.querySelector('div#controls'));
```

你應該為其中一個新的 `div` 元素指定一個 `id` 為 `monsterStats`。

```js
assert.exists(document.querySelector('div#monsterStats'));
```

你應該為其中一個新的 `div` 元素指定一個 `id` 為 `text`。

```js
assert.exists(document.querySelector('div#text'));
```

你應該將新的 `div` 元素放在正確的順序中。

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

你應該將新的 `div` 元素放置在 `#game` 元素內。

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
