---
id: 62a240c67f3dbb1a1e6d95ee
title: 第 23 步
challengeType: 0
dashedName: step-23
---

# --description--

為你的 `#game` 設定最大寬度為 `500px` 和最大高度為 `400px`。將 `background-color` 設定為 `#ffffff`，並將 `color` 設定為 `#ffffff`。 

透過將上邊距設為 `30px`、下邊距設為 `0px`，以及左右邊距設為 `auto`，使用邊距來置中它。 

最後，為四個邊都添加 `10px` 的內距。

# --hints--

你應該有一個 `#game` 選擇器。

```js
const game = new __helpers.CSSHelp(document).getStyle('#game');
assert.exists(game);
```

你的 `#game` 選擇器應該有 `max-width` 為 `500px`。

```js
const maxWidth = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-width');
assert.equal(maxWidth, '500px');
```

你的 `#game` 選擇器應該有 `max-height` 為 `400px`。

```js
const maxHeight = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-height');
assert.equal(maxHeight, '400px');
```

你的 `#game` 選擇器應該有 `background-color` 為 `#ffffff`。

```js
const background = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('background-color');
assert.equal(background, 'rgb(255, 255, 255)');
```

你的 `#game` 選擇器應該有 `color` 為 `#ffffff`。

```js
const color = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('color');
assert.equal(color, 'rgb(255, 255, 255)');
```

你的 `#game` 選擇器應該設定 `margin-top` 為 `30px`。

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-top');
assert.equal(margin, '30px');
```

你的 `#game` 選擇器應該將 `margin-left` 設為 `auto`。

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-left');
assert.equal(margin, 'auto');
```

你的 `#game` 選擇器應該將 `margin-right` 設為 `auto`。

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-right');
assert.equal(margin, 'auto');
```


你的 `#game` 選擇器應該將 `margin-bottom` 設為 `0px`。

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-bottom');
assert.equal(margin, '0px');
```

你的 `#game` 選擇器四邊應該有 `10px` 的內距。

```js
const padding = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('padding');
assert.equal(padding, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
