---
id: 62a255dae245b52317da824a
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

最強大的工具之一是你的開發者主控台。根據你的瀏覽器，這可能可以透過按下 `F12` 或 `Ctrl+Shift+I` 來開啟。在 Mac 上，你可以按下 `Option + ⌘ + C` 並選擇「主控台」。你也可以點擊預覽視窗上方的「主控台」按鈕來查看我們內建的主控台。 

開發者主控台會包含你的程式碼所產生的錯誤，但你也可以用它來查看程式碼中變數的值，這對偵錯很有幫助。

在你的 `script` 頁籤之間添加一行 `console.log("Hello World");`。然後點擊「Console」按鈕以開啟主控台。你應該會看到文字 `"Hello World"`。


# --hints--

你應該為你的程式碼添加一行 `console.log("Hello World");`。 

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

你的 `console.log("Hello World");` 行應該位於你的 `script` 頁籤之間。

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
