---
id: 62a23c1d505bfa13747c8a9b
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

將數字 `0`、`100` 和 `50` 包裹在 `span` 元素中，然後將這些新的 `span` 元素包裹在 `strong` 元素中。接著分別為你的新 `span` 元素指定 `id` 值為 `xpText`、`healthText` 和 `goldText`。

你的答案應該遵循這個基本結構：

```html
<span class="stat">TEXT <strong><span id="VALUE">TEXT</span></strong></span>
```

# --hints--

你應該在你的第一個 `.stat` 元素中添加一個 `strong` 元素。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

你的第一個新的 `strong` 元素應該有一個 `span` 元素。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

你的第一個巢狀 `span` 元素應該有 `id` 為 `xpText`。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'xpText');
```

你的第一個 `span` 元素應該包裹在文字 `0` 周圍。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '0');
```

你第一個 `.stat` 元素的文字仍應為 `XP: 0`。

```js
const stat = document.querySelectorAll('.stat')[0];
assert.equal(stat.innerText.trimEnd(), 'XP: 0');
```

你應該在你的第二個 `.stat` 元素中添加一個 `strong` 元素。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

你的第二個新的 `strong` 元素應該有一個 `span` 元素。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

你的第二個巢狀 `span` 元素應該有 `id` 為 `healthText`。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'healthText');
```

你的第二個 `span` 元素應該包裹住文字 `100`。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '100');
```

你的第二個 `.stat` 元素的文字仍應該是 `Health: 100`。

```js
const stat = document.querySelectorAll('.stat')[1];
assert.equal(stat.innerText.trimEnd(), 'Health: 100');
```

你應該在你的第三個 `.stat` 元素中添加一個 `strong` 元素。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

你的第三個新的 `strong` 元素應該有一個 `span` 元素。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

你的第三個巢狀 `span` 元素應該有 `id` 為 `goldText`。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'goldText');
```

你的第三個 `span` 元素應該包裹住文字 `50`。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '50');
```

第三個 `.stat` 元素的文字仍應為 `Gold: 50`。

```js
const stat = document.querySelectorAll('.stat')[2];
assert.equal(stat.innerText.trimEnd(), 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
--fcc-editable-region--
        <span class="stat">XP: 0</span>
        <span class="stat">Health: 100</span>
        <span class="stat">Gold: 50</span>
--fcc-editable-region--
      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
