---
id: 62a23cb9bc467a147516b500
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

為你的 `#controls` 元素創建三個 `button` 元素。第一個應該將 `id` 設定為 `button1`，並且文字為 `Go to store`。第二個應該將 `id` 設定為 `button2`，並且文字為 `Go to cave`。第三個應該將 `id` 設定為 `button3`，並且文字為 `Fight dragon`。

# --hints--

你應該為你的 `#controls` 元素添加三個 `button` 元素。

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

你的第一個按鈕應該將 `id` 設定為 `button1`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

你的第一個按鈕應該有文字 `Go to store`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

你的第二個按鈕應該將 `id` 設定為 `button2`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

你的第二個按鈕應該有文字 `Go to cave`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

你的第三個按鈕應該將 `id` 設定為 `button3`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

你的第三個按鈕應該有文字 `Fight dragon`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
--fcc-editable-region--
      <div id="controls">

      </div>
--fcc-editable-region--
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
