---
id: 62a23d1c5f1c93161f3582ae
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

類似於你的 `#stats` 元素，你的 `#monsterStats` 元素需要兩個 `span` 元素。為它們添加 `stat` 類別，並且第一個元素的文字為 `Monster Name: `，第二個元素的文字為 `Health: `。在每個文字後面，添加一個 `strong` 元素，裡面巢狀一個空的 `span` 元素。為第一個內部的 `span` 元素設定 `id` 為 `monsterName`，為第二個內部的 `span` 元素設定 `id` 為 `monsterHealth`。

# --hints--

你的 `#monsterStats` 元素應該有兩個 `span` 元素。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

你的新 `span` 元素都應該有一個 `class` 值為 `stat`。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

你的第一個 `span` 元素應該有一個包含空的巢狀 `span` 元素的 `strong` 元素。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

你的第二個 `span` 元素應該有一個 `strong` 元素，裡面包含一個空的巢狀 `span` 元素。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

你的第一個 `span` 元素應該有文字 `Monster Name: `。請確保間距正確。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+/);
```

你的第二個 `span` 元素應該有文字 `Health: `。請確保間距正確。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[1].innerHTML, /Health:\s+/);
```

你的文字後面應該接著 `strong` 和 `span` 元素。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+<strong>/);
assert.match(spans[1].innerHTML, /Health:\s+<strong>/);
```

你的第一個內部 `span` 元素應該有一個 `id` 為 `monsterName`。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[0]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterName');
```

你的第二個內部 `span` 元素應該有一個 `id` 為 `monsterHealth`。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[1]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterHealth');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
--fcc-editable-region--
      <div id="monsterStats">

      </div>
--fcc-editable-region--
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
