---
id: 62a2509ba163e020bb9d84ea
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

現在你可以開始撰寫你的 JavaScript。首先建立一個 `script` 元素。此元素用於將 JavaScript 載入你的 HTML 檔案中。 

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

你應該有一個 `script` 元素。

```js
assert.isAtLeast(document.querySelectorAll('script').length, 1);
```

你的 `script` 元素應該有一個開啟的頁籤。

```js
assert.match(code, /<script\s*>/i);
```

你的 `script` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
