---
id: 62a257659d0d1e2456f24ba2
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在你開始撰寫你的專案程式碼之前，你應該將它移動到自己的檔案中以保持整齊。 

移除你的 `console.log("Hello World");` 行。然後為你現在空的 `script` 元素添加一個 `src` 屬性，設定為 `./script.js`。

# --hints--

你的程式碼中不應該有 `console.log("Hello World");` 這一行。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

你的 `script` 元素應該有一個 `src` 屬性設定為 `./script.js`。

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
