---
id: 62a3a7e4f1060e2fc5ffb34b
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

創建另一個名為 `currentWeaponIndex` 的變數並將其設為 `0`。

# --hints--

你應該使用 `let` 宣告一個名為 `currentWeaponIndex` 的變數。

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

你應該使用 camelCase 來命名你的變數。

```js
assert.match(code, /currentWeaponIndex/);
```

你的 `currentWeaponIndex` 變數應該設為 `0`。

```js
assert.equal(currentWeaponIndex, 0);
```

你應該將你的變數初始化為 `0`。

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
