---
id: 62a3b3eab50e193608c19fc6
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

宣告一個名為 `fighting` 的變數，但不要用值初始化它。 

# --hints--

你應該使用 `let` 來宣告變數 `fighting`。

```js
assert.match(code, /let\s+fighting/);
```

你的 `fighting` 變數不應該有值。

```js
assert.isUndefined(fighting);
```

你不應該為你的 `fighting` 變數指定值。 

```js
assert.match(code, /let\s+fighting\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
--fcc-editable-region--
```
