---
id: 62a3b79d520a7f3d0e25afd6
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript 使用 <dfn>文件物件模型</dfn>（DOM）與 HTML 交談。DOM 是一個表述 HTML 的物件樹。你可以使用表述你整個 HTML 文件的 `document` 物件來存取 HTML。

尋找 HTML 中特定元素的一種方法是使用 `querySelector()` 方法。`querySelector()` 方法會將 CSS 選擇器作為引數，並傳回符合該選擇器的第一個元素。 例如，要尋找 HTML 中的 `<h1>` 元素，你可以寫：

```js
let h1 = document.querySelector("h1");
```

請注意，`h1` 是一個字串，並且符合你會使用的 CSS 選擇器。

創建一個 `button1` 變數，並使用 `querySelector()` 將你的 `id` 為 `button1` 的元素指定給它。請記得 CSS `id` 選擇器前置式為 `#`。

# --hints--

你應該使用 `let` 來宣告一個 `button1` 變數。

```js
assert.match(code, /let\s+button1/);
```

你應該使用 `document.querySelector()`。

```js
assert.match(code, /document\.querySelector/);
```

你應該使用 `"#button1"` 選擇器。

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
