---
id: 62a3b41c9494f937560640ab
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

宣告另外兩個名為 `monsterHealth` 和 `inventory` 的變數。

為你的 `inventory` 變數指定一個包含字串 `"stick"` 的陣列值。

請記得你在前一個專案中這樣使用過陣列：

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

你應該使用 `let` 來宣告你的 `monsterHealth` 變數。

```js
assert.match(code, /let\s+monsterHealth/i);
```

你應該使用 camelCase 來命名你的 `monsterHealth` 變數。

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` 不應該有值。

```js
assert.isUndefined(monsterHealth);
```

你不應該為你的 `monsterHealth` 變數指定值。

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

你應該使用 `let` 來宣告你的 `inventory` 變數。

```js
assert.match(code, /let\s+inventory/i);
```

你的 `inventory` 變數仍然應該是陣列。

```js
assert.isArray(inventory);
```

你的 `inventory` 變數應該只有一個值。

```js
assert.lengthOf(inventory, 1);
```

你的 `inventory` 變數應該包含字串 `"stick"`。

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
