---
id: 62a39f5a5790eb27c1e5d4bf
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

你的檢視已切換到你的新 `script.js` 檔案。請記得你可以使用上方的分頁在檔案之間切換。

將你的 `console.log("Hello World");` 行添加到這個檔案，並在你的主控台中看到它出現。

# --hints--

你的程式碼中應該有一行 `console.log("Hello World");`。

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
