---
id: 62a3c2fccf186146b59c6e96
title: 第 36 步
challengeType: 0
dashedName: step-36
---

# --description--

`button1` 表現你的第一個 `button` 元素。這些元素有一個特殊的屬性叫做 `onclick`，你可以用它來決定當有人點擊該按鈕時會發生什麼。

你可以用幾種不同的方式在 JavaScript 中存取屬性。第一種是使用<dfn>點記法</dfn>。以下是一個使用點記法將按鈕的 `onclick` 屬性設定為函式引用的範例。

```js
button.onclick = myFunction;
```

在此範例中，`button` 是按鈕元素，而 `myFunction` 是對函式的引用。當按鈕被點擊時，將會呼叫 `myFunction`。

使用點記法將你的 `button1` 的 `onclick` 屬性設定為 `goStore` 的函式引用。請注意，`button1` 已經宣告過，所以你不需要使用 `let` 或 `const`。

# --hints--

你應該使用點記法來存取 `button1` 的 `onclick` 屬性。

```js
assert.match(code, /button1\.onclick/);
```

你不應該使用 `let` 或 `const`。

```js
assert.notMatch(code, /(let|const)\s+button1\.onclick/);
```

你應該將 `goStore` 函式（程式）引用指定給 `button1.onclick`。請確保不要呼叫該函式（程式）。

```js
assert.match(code, /button1\.onclick\s*=\s*goStore\s*;?\s*$/m);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");

--fcc-editable-region--
// initialize buttons

--fcc-editable-region--
function goStore() {
  console.log("Going to store.");
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
