---
id: 62a3a0a3c0a4b32915d26a6e
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

移除你的 `console.log("Hello World");` 行以開始撰寫你的專案程式碼。

在先前的專案中，你學會了如何像這樣使用變數和 `let` 關鍵字：

```js
let camperbot;
```

你也學會了如何用這樣的方式初始化一個變數並賦值：

```js
let age = 32;
```

使用 `let` 關鍵字宣告一個名為 `xp` 的變數，並將其指定為數字 `0`。

# --hints--

你的程式碼中不應該有 `console.log("Hello World");` 這一行。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

你應該使用 `let` 關鍵字來宣告你的變數。

```js
assert.match(code, /let/);
```

你應該宣告一個名為 `xp` 的變數。

```js
assert.match(code, /let\s+xp/);
```

`xp` 應該有一個值為 `0`。

```js
assert.equal(xp, 0);
```

你應該將 `xp` 變數初始化為 `0`。 

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
