---
id: 62a3a75d8466a12e009eff76
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

初始化另一個名為 `health` 的變數，值為 `100`，以及一個名為 `gold` 的變數，值為 `50`。

# --hints--

你應該使用 `let` 來宣告一個名為 `health` 的變數。

```js
assert.match(code, /let\s+health/);
```

你應該初始化一個名為 `health` 的變數，值為 `100`。

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

你應該使用 `let` 來宣告一個名為 `gold` 的變數。

```js
assert.match(code, /let\s+gold/);
```

你應該初始化一個名為 `gold` 的變數，值為 `50`。

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` 應該有一個值為 `100`。

```js
assert.equal(health, 100);
```

`gold` 應該有一個值為 `50`。

```js
assert.equal(gold, 50);
```

`xp` 應該仍然具有 `0` 的值。

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
