---
id: 62a7bf06d2ad9d1c5024e833
title: 步驟 40
challengeType: 0
dashedName: step-40
---

# --description--

你還需要更新在按鈕被點擊時再次執行的函式（程式）。

在你的 `goStore()` 函式中，更新每個按鈕的 `onclick` 屬性，分別執行 `buyHealth`、`buyWeapon` 和 `goTown`。

# --hints--

你應該使用點記法來存取 `button1` 的 `onclick` 屬性。

```js
assert.match(goStore.toString(), /button1\.onclick/);
```

你不應該使用 `let` 或 `const` 來存取 `button1` 的 `onclick` 屬性。

```js
assert.notMatch(code, /(let|const)\s+button1\.onclick/);
```

你應該將 `button1` 的 `onclick` 屬性設為 `buyHealth`。

```js
assert.match(code, /button1\.onclick\s*=\s*buyHealth/);
```

你應該在你的 `goStore` 函式中設定 `button1` 的 `onclick` 屬性。

```js
assert.match(goStore.toString(), /button1\.onclick\s*=\s*buyHealth/);
```

你應該使用點記法來存取 `button2` 的 `onclick` 屬性。

```js
assert.match(goStore.toString(), /button2\.onclick/);
```

你不應該使用 `let` 或 `const` 來存取 `button2` 的 `onclick` 屬性。

```js
assert.notMatch(code, /(let|const)\s+button2\.onclick/);
```

你應該將 `button2` 的 `onclick` 屬性設為 `buyWeapon`。

```js
assert.match(code, /button2\.onclick\s*=\s*buyWeapon/);
```

你應該在你的 `goStore` 函式中設定 `button2` 的 `onclick` 屬性。

```js
assert.match(goStore.toString(), /button2\.onclick\s*=\s*buyWeapon/);
```

你應該使用點記法來存取 `button3` 的 `onclick` 屬性。

```js
assert.match(goStore.toString(), /button3\.onclick/);
```

你不應該使用 `let` 或 `const` 來存取 `button3` 的 `onclick` 屬性。

```js
assert.notMatch(code, /(let|const)\s+button3\.onclick/);
```

你應該將 `button3` 的 `onclick` 屬性設為 `goTown`。

```js
assert.match(code, /button3\.onclick\s*=\s*goTown/);
```

你應該在你的 `goStore` 函式中設定 `button3` 的 `onclick` 屬性。

```js
assert.match(goStore.toString(), /button3\.onclick\s*=\s*goTown/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function goStore() {
  button1.innerText = "Buy 10 health (10 gold)";
  button2.innerText = "Buy weapon (30 gold)";
  button3.innerText = "Go to town square";
}
--fcc-editable-region--

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
