---
id: 62b46e3a8d4be31be5af793d
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

我們遇到了一點小問題。你正嘗試查詢你的頁面中的按鈕元素，但你的 `script` 頁籤位於 HTML 的 `head` 中。這表示你的程式碼會在瀏覽器完成讀取 HTML 之前執行，而你的 `document.querySelector()` 不會看到該按鈕－因為瀏覽器尚未處理它。

為了解決此問題，請將你的 `script` 元素從 `head` 元素中移出，並將它放置在你的 `body` 元素的結尾（就在關閉的 `</body>` 標籤之前）。

# --hints--

你的 `script` 元素不應該放在你的 `head` 元素中。

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

你的 `script` 元素應該放在你的 `body` 元素的最後。

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
