---
id: 6140c9d35015ae0ba0c250e8
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Agregua un `div` dentro de  tu elemento `body` y dale una `class` de `wheel`.

Dentro de tu nuevo `div`, agrega seis elementos `span` con una `class` establecida a `line`, y seis elementos `div` con una `class` establecida a `cabin`.

# --hints--

Debes crear un nuevo `div` dentro de tu elemento `body`.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert.lengthOf(divs, 1);
```

Tu elemento `div` dentro de `body` debe tener una `class` de `wheel`.

```js
assert.isTrue(document.querySelector('body div')?.classList?.contains('wheel'));
```

Tu elemento `.wheel` debe tener seis elementos `span` dentro.

```js
assert.lengthOf(document.querySelectorAll('.wheel span'), 6);
```

Tus seis elementos `span` dentro del elemento `.wheel` deben tener una `class` de `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert.isTrue(spans?.every(span => span?.classList?.contains('line')));
assert.lengthOf(document.querySelectorAll('.line'), 6);
```

Tu elemento `.wheel` debe tener seis elementos `div` dentro.

```js
assert.lengthOf(document.querySelectorAll('.wheel div'), 6);
```

Tus seis elementos `div` dentro del elemento `.wheel` deben tener una `class` de `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert.isTrue(divs?.every(div => div?.classList?.contains('cabin')));
assert.lengthOf(document.querySelectorAll('.cabin'), 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
