---
id: 6140e0d875ec16262f26432b
title: Paso 21
challengeType: 0
dashedName: step-21
---

# --description--

La propiedad `animation-iteration-count` define cuántas veces debe repetirse tu animación. Esto se puede configurar con un número, o con `infinite` para repetir indefinidamente la animación. Tu rueda de la fortuna nunca debe detenerse, así que configura el selector `.wheel` para que tenga un `animation-iteration-count` de `infinite`.

La propiedad `animation-timing-function` establece cómo debe progresar la animación con el tiempo. Hay algunos valores diferentes para esta propiedad, pero deseas que la animación de la rueda de la fortuna se ejecute al mismo ritmo de principio a fin. Establece la `animation-timing-function` a `linear` en tu selector `.wheel`.

# --hints--

Tu selector `.wheel` debe tener una propiedad `animation-iteration-count` establecida en `infinite`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationIterationCount, 'infinite');
```

Tu selector `.wheel` debe tener una propiedad `animation-timing-function` establecida en `linear`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationTimingFunction, 'linear');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
  animation-name: wheel;
  animation-duration: 10s;
}
--fcc-editable-region--

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

@keyframes wheel {
   0% {
     transform: rotate(0deg);
   }
   100% {
     transform: rotate(360deg);
   }
}
```
