---
id: 6140cdebd39d6a101e747529
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Crea un selector para tus elementos `.line`. Empieza configurando el `background-color` a `black`, el `width` a `50%` y el `height` a `2px`.

# --hints--

Debes tener un selector `.line`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line'));
```

Tu selector `.line` debe tener una propiedad `background-color` configurada a `black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor, "black");
```

Tu selector `.line` debe tener una propiedad `width` configurada a `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.width, "50%");
```

Tu selector `.line` debe tener una propiedad `height` configurada a `2px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.height, "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
