---
id: 68773ee26f332a80bc0295db
title: Den Tower-of-Hanoi-Algorithmus implementieren
challengeType: 27
saveSubmissionToDB: true
dashedName: implement-the-tower-of-hanoi-algorithm
---

# --description--

In diesem Labor lösen Sie das mathematische Rätsel, das als Tower of Hanoi bekannt ist. Das Rätsel besteht aus drei Stäben und einer Anzahl von Scheiben mit unterschiedlichen Durchmessern.

<img alt="Folge der Züge, die erforderlich sind, um ein Tower-of-Hanoi-Rätsel mit 3 Scheiben zu lösen" src="https://cdn.freecodecamp.org/curriculum/python/tower-of-hanoi.gif" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Das Rätsel beginnt damit, dass die Scheiben auf dem ersten Stab gestapelt sind, in abnehmender Größe, wobei die kleinste Scheibe oben und die größte Scheibe unten liegt.

Das Ziel des Tower-of-Hanoi-Rätsels ist es, alle Scheiben auf den letzten Stab zu bewegen. Dafür müssen Sie drei einfache Regeln befolgen:

- Sie dürfen nur die obersten Scheiben bewegen.
- Sie dürfen jeweils nur eine Scheibe bewegen.
- Sie dürfen keine größere Scheibe auf eine kleinere legen.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `hanoi_solver` haben, die eine ganze Zahl als Argument annimmt, welche die Gesamtzahl der Scheiben des Rätsels darstellt.
1. Die `hanoi_solver`-Funktion sollte das Rätsel unter Einhaltung der angegebenen Regeln in 2<sup>n</sup> - 1 Zügen lösen, wobei `n` die Gesamtzahl der Scheiben ist.
1. Die `hanoi_solver`-Funktion sollte einen String zurückgeben, der alle Züge zur Lösung des Rätsels enthält, einschließlich der Anfangsanordnung, wobei jeder Zug in einer neuen Zeile steht. Die Stäbe sollten als Listen von ganzen Zahlen dargestellt werden (die kleinste Scheibe wird durch die Zahl `1` dargestellt), wobei jeder Stab durch ein Leerzeichen getrennt ist. Zum Beispiel sollte `hanoi_solver(3)` Folgendes zurückgeben:

```md
[3, 2, 1] [] []
[3, 2] [] [1]
[3] [2] [1]
[3] [2, 1] []
[] [2, 1] [3]
[1] [2] [3]
[1] [] [3, 2]
[] [] [3, 2, 1]
```

# --hints--

Sie sollten eine Funktion namens `hanoi_solver` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_function('hanoi_solver')`)) })
```

Ihre `hanoi_solver`-Funktion sollte ein einzelnes Argument annehmen.

```js
({ test: () => runPython(`
import inspect
sig = inspect.signature(hanoi_solver)
assert len(sig.parameters) == 1
`) })
```

Ihre `hanoi_solver`-Funktion sollte einen String zurückgeben.

```js
({ test: () => runPython(`
assert isinstance(hanoi_solver(2), str)
assert isinstance(hanoi_solver(6), str)
assert isinstance(hanoi_solver(11), str)
`) })
```

`hanoi_solver(2)` sollte `[2, 1] [] []\n[2] [1] []\n[] [1] [2]\n[] [] [2, 1]` zurückgeben.

```js
({ test: () => runPython(`
expected = """[2, 1] [] []
[2] [1] []
[] [1] [2]
[] [] [2, 1]"""
assert hanoi_solver(2) == expected
`) })
```

`hanoi_solver(3)` sollte `[3, 2, 1] [] []\n[3, 2] [] [1]\n[3] [2] [1]\n[3] [2, 1] []\n[] [2, 1] [3]\n[1] [2] [3]\n[1] [] [3, 2]\n[] [] [3, 2, 1]` zurückgeben.

```js
({ test: () => runPython(`
expected = """[3, 2, 1] [] []
[3, 2] [] [1]
[3] [2] [1]
[3] [2, 1] []
[] [2, 1] [3]
[1] [2] [3]
[1] [] [3, 2]
[] [] [3, 2, 1]"""
assert hanoi_solver(3) == expected
`) })
```

`hanoi_solver(4)` sollte `[4, 3, 2, 1] [] []\n[4, 3, 2] [1] []\n[4, 3] [1] [2]\n[4, 3] [] [2, 1]\n[4] [3] [2, 1]\n[4, 1] [3] [2]\n[4, 1] [3, 2] []\n[4] [3, 2, 1] []\n[] [3, 2, 1] [4]\n[] [3, 2] [4, 1]\n[2] [3] [4, 1]\n[2, 1] [3] [4]\n[2, 1] [] [4, 3]\n[2] [1] [4, 3]\n[] [1] [4, 3, 2]\n[] [] [4, 3, 2, 1]` zurückgeben.

```js
({ test: () => runPython(`
expected = """[4, 3, 2, 1] [] []
[4, 3, 2] [1] []
[4, 3] [1] [2]
[4, 3] [] [2, 1]
[4] [3] [2, 1]
[4, 1] [3] [2]
[4, 1] [3, 2] []
[4] [3, 2, 1] []
[] [3, 2, 1] [4]
[] [3, 2] [4, 1]
[2] [3] [4, 1]
[2, 1] [3] [4]
[2, 1] [] [4, 3]
[2] [1] [4, 3]
[] [1] [4, 3, 2]
[] [] [4, 3, 2, 1]"""
assert hanoi_solver(4) == expected
`) })
```

`hanoi_solver(5)` sollte `[5, 4, 3, 2, 1] [] []\n[5, 4, 3, 2] [] [1]\n[5, 4, 3] [2] [1]\n[5, 4, 3] [2, 1] []\n[5, 4] [2, 1] [3]\n[5, 4, 1] [2] [3]\n[5, 4, 1] [] [3, 2]\n[5, 4] [] [3, 2, 1]\n[5] [4] [3, 2, 1]\n[5] [4, 1] [3, 2]\n[5, 2] [4, 1] [3]\n[5, 2, 1] [4] [3]\n[5, 2, 1] [4, 3] []\n[5, 2] [4, 3] [1]\n[5] [4, 3, 2] [1]\n[5] [4, 3, 2, 1] []\n[] [4, 3, 2, 1] [5]\n[1] [4, 3, 2] [5]\n[1] [4, 3] [5, 2]\n[] [4, 3] [5, 2, 1]\n[3] [4] [5, 2, 1]\n[3] [4, 1] [5, 2]\n[3, 2] [4, 1] [5]\n[3, 2, 1] [4] [5]\n[3, 2, 1] [] [5, 4]\n[3, 2] [] [5, 4, 1]\n[3] [2] [5, 4, 1]\n[3] [2, 1] [5, 4]\n[] [2, 1] [5, 4, 3]\n[1] [2] [5, 4, 3]\n[1] [] [5, 4, 3, 2]\n[] [] [5, 4, 3, 2, 1]` zurückgeben.

```js
({ test: () => runPython(`
expected = """[5, 4, 3, 2, 1] [] []
[5, 4, 3, 2] [] [1]
[5, 4, 3] [2] [1]
[5, 4, 3] [2, 1] []
[5, 4] [2, 1] [3]
[5, 4, 1] [2] [3]
[5, 4, 1] [] [3, 2]
[5, 4] [] [3, 2, 1]
[5] [4] [3, 2, 1]
[5] [4, 1] [3, 2]
[5, 2] [4, 1] [3]
[5, 2, 1] [4] [3]
[5, 2, 1] [4, 3] []
[5, 2] [4, 3] [1]
[5] [4, 3, 2] [1]
[5] [4, 3, 2, 1] []
[] [4, 3, 2, 1] [5]
[1] [4, 3, 2] [5]
[1] [4, 3] [5, 2]
[] [4, 3] [5, 2, 1]
[3] [4] [5, 2, 1]
[3] [4, 1] [5, 2]
[3, 2] [4, 1] [5]
[3, 2, 1] [4] [5]
[3, 2, 1] [] [5, 4]
[3, 2] [] [5, 4, 1]
[3] [2] [5, 4, 1]
[3] [2, 1] [5, 4]
[] [2, 1] [5, 4, 3]
[1] [2] [5, 4, 3]
[1] [] [5, 4, 3, 2]
[] [] [5, 4, 3, 2, 1]"""
assert hanoi_solver(5) == expected
`) })
```

`hanoi_solver(n)` sollte das Tower-of-Hanoi-Rätsel für jeden positiven Wert von `n` lösen.

```js
({ test: () => runPython(`
expected = """[6, 5, 4, 3, 2, 1] [] []
[6, 5, 4, 3, 2] [1] []
[6, 5, 4, 3] [1] [2]
[6, 5, 4, 3] [] [2, 1]
[6, 5, 4] [3] [2, 1]
[6, 5, 4, 1] [3] [2]
[6, 5, 4, 1] [3, 2] []
[6, 5, 4] [3, 2, 1] []
[6, 5] [3, 2, 1] [4]
[6, 5] [3, 2] [4, 1]
[6, 5, 2] [3] [4, 1]
[6, 5, 2, 1] [3] [4]
[6, 5, 2, 1] [] [4, 3]
[6, 5, 2] [1] [4, 3]
[6, 5] [1] [4, 3, 2]
[6, 5] [] [4, 3, 2, 1]
[6] [5] [4, 3, 2, 1]
[6, 1] [5] [4, 3, 2]
[6, 1] [5, 2] [4, 3]
[6] [5, 2, 1] [4, 3]
[6, 3] [5, 2, 1] [4]
[6, 3] [5, 2] [4, 1]
[6, 3, 2] [5] [4, 1]
[6, 3, 2, 1] [5] [4]
[6, 3, 2, 1] [5, 4] []
[6, 3, 2] [5, 4, 1] []
[6, 3] [5, 4, 1] [2]
[6, 3] [5, 4] [2, 1]
[6] [5, 4, 3] [2, 1]
[6, 1] [5, 4, 3] [2]
[6, 1] [5, 4, 3, 2] []
[6] [5, 4, 3, 2, 1] []
[] [5, 4, 3, 2, 1] [6]
[] [5, 4, 3, 2] [6, 1]
[2] [5, 4, 3] [6, 1]
[2, 1] [5, 4, 3] [6]
[2, 1] [5, 4] [6, 3]
[2] [5, 4, 1] [6, 3]
[] [5, 4, 1] [6, 3, 2]
[] [5, 4] [6, 3, 2, 1]
[4] [5] [6, 3, 2, 1]
[4, 1] [5] [6, 3, 2]
[4, 1] [5, 2] [6, 3]
[4] [5, 2, 1] [6, 3]
[4, 3] [5, 2, 1] [6]
[4, 3] [5, 2] [6, 1]
[4, 3, 2] [5] [6, 1]
[4, 3, 2, 1] [5] [6]
[4, 3, 2, 1] [] [6, 5]
[4, 3, 2] [1] [6, 5]
[4, 3] [1] [6, 5, 2]
[4, 3] [] [6, 5, 2, 1]
[4] [3] [6, 5, 2, 1]
[4, 1] [3] [6, 5, 2]
[4, 1] [3, 2] [6, 5]
[4] [3, 2, 1] [6, 5]
[] [3, 2, 1] [6, 5, 4]
[] [3, 2] [6, 5, 4, 1]
[2] [3] [6, 5, 4, 1]
[2, 1] [3] [6, 5, 4]
[2, 1] [] [6, 5, 4, 3]
[2] [1] [6, 5, 4, 3]
[] [1] [6, 5, 4, 3, 2]
[] [] [6, 5, 4, 3, 2, 1]"""
assert hanoi_solver(6) == expected

expected = """[7, 6, 5, 4, 3, 2, 1] [] []
[7, 6, 5, 4, 3, 2] [] [1]
[7, 6, 5, 4, 3] [2] [1]
[7, 6, 5, 4, 3] [2, 1] []
[7, 6, 5, 4] [2, 1] [3]
[7, 6, 5, 4, 1] [2] [3]
[7, 6, 5, 4, 1] [] [3, 2]
[7, 6, 5, 4] [] [3, 2, 1]
[7, 6, 5] [4] [3, 2, 1]
[7, 6, 5] [4, 1] [3, 2]
[7, 6, 5, 2] [4, 1] [3]
[7, 6, 5, 2, 1] [4] [3]
[7, 6, 5, 2, 1] [4, 3] []
[7, 6, 5, 2] [4, 3] [1]
[7, 6, 5] [4, 3, 2] [1]
[7, 6, 5] [4, 3, 2, 1] []
[7, 6] [4, 3, 2, 1] [5]
[7, 6, 1] [4, 3, 2] [5]
[7, 6, 1] [4, 3] [5, 2]
[7, 6] [4, 3] [5, 2, 1]
[7, 6, 3] [4] [5, 2, 1]
[7, 6, 3] [4, 1] [5, 2]
[7, 6, 3, 2] [4, 1] [5]
[7, 6, 3, 2, 1] [4] [5]
[7, 6, 3, 2, 1] [] [5, 4]
[7, 6, 3, 2] [] [5, 4, 1]
[7, 6, 3] [2] [5, 4, 1]
[7, 6, 3] [2, 1] [5, 4]
[7, 6] [2, 1] [5, 4, 3]
[7, 6, 1] [2] [5, 4, 3]
[7, 6, 1] [] [5, 4, 3, 2]
[7, 6] [] [5, 4, 3, 2, 1]
[7] [6] [5, 4, 3, 2, 1]
[7] [6, 1] [5, 4, 3, 2]
[7, 2] [6, 1] [5, 4, 3]
[7, 2, 1] [6] [5, 4, 3]
[7, 2, 1] [6, 3] [5, 4]
[7, 2] [6, 3] [5, 4, 1]
[7] [6, 3, 2] [5, 4, 1]
[7] [6, 3, 2, 1] [5, 4]
[7, 4] [6, 3, 2, 1] [5]
[7, 4, 1] [6, 3, 2] [5]
[7, 4, 1] [6, 3] [5, 2]
[7, 4] [6, 3] [5, 2, 1]
[7, 4, 3] [6] [5, 2, 1]
[7, 4, 3] [6, 1] [5, 2]
[7, 4, 3, 2] [6, 1] [5]
[7, 4, 3, 2, 1] [6] [5]
[7, 4, 3, 2, 1] [6, 5] []
[7, 4, 3, 2] [6, 5] [1]
[7, 4, 3] [6, 5, 2] [1]
[7, 4, 3] [6, 5, 2, 1] []
[7, 4] [6, 5, 2, 1] [3]
[7, 4, 1] [6, 5, 2] [3]
[7, 4, 1] [6, 5] [3, 2]
[7, 4] [6, 5] [3, 2, 1]
[7] [6, 5, 4] [3, 2, 1]
[7] [6, 5, 4, 1] [3, 2]
[7, 2] [6, 5, 4, 1] [3]
[7, 2, 1] [6, 5, 4] [3]
[7, 2, 1] [6, 5, 4, 3] []
[7, 2] [6, 5, 4, 3] [1]
[7] [6, 5, 4, 3, 2] [1]
[7] [6, 5, 4, 3, 2, 1] []
[] [6, 5, 4, 3, 2, 1] [7]
[1] [6, 5, 4, 3, 2] [7]
[1] [6, 5, 4, 3] [7, 2]
[] [6, 5, 4, 3] [7, 2, 1]
[3] [6, 5, 4] [7, 2, 1]
[3] [6, 5, 4, 1] [7, 2]
[3, 2] [6, 5, 4, 1] [7]
[3, 2, 1] [6, 5, 4] [7]
[3, 2, 1] [6, 5] [7, 4]
[3, 2] [6, 5] [7, 4, 1]
[3] [6, 5, 2] [7, 4, 1]
[3] [6, 5, 2, 1] [7, 4]
[] [6, 5, 2, 1] [7, 4, 3]
[1] [6, 5, 2] [7, 4, 3]
[1] [6, 5] [7, 4, 3, 2]
[] [6, 5] [7, 4, 3, 2, 1]
[5] [6] [7, 4, 3, 2, 1]
[5] [6, 1] [7, 4, 3, 2]
[5, 2] [6, 1] [7, 4, 3]
[5, 2, 1] [6] [7, 4, 3]
[5, 2, 1] [6, 3] [7, 4]
[5, 2] [6, 3] [7, 4, 1]
[5] [6, 3, 2] [7, 4, 1]
[5] [6, 3, 2, 1] [7, 4]
[5, 4] [6, 3, 2, 1] [7]
[5, 4, 1] [6, 3, 2] [7]
[5, 4, 1] [6, 3] [7, 2]
[5, 4] [6, 3] [7, 2, 1]
[5, 4, 3] [6] [7, 2, 1]
[5, 4, 3] [6, 1] [7, 2]
[5, 4, 3, 2] [6, 1] [7]
[5, 4, 3, 2, 1] [6] [7]
[5, 4, 3, 2, 1] [] [7, 6]
[5, 4, 3, 2] [] [7, 6, 1]
[5, 4, 3] [2] [7, 6, 1]
[5, 4, 3] [2, 1] [7, 6]
[5, 4] [2, 1] [7, 6, 3]
[5, 4, 1] [2] [7, 6, 3]
[5, 4, 1] [] [7, 6, 3, 2]
[5, 4] [] [7, 6, 3, 2, 1]
[5] [4] [7, 6, 3, 2, 1]
[5] [4, 1] [7, 6, 3, 2]
[5, 2] [4, 1] [7, 6, 3]
[5, 2, 1] [4] [7, 6, 3]
[5, 2, 1] [4, 3] [7, 6]
[5, 2] [4, 3] [7, 6, 1]
[5] [4, 3, 2] [7, 6, 1]
[5] [4, 3, 2, 1] [7, 6]
[] [4, 3, 2, 1] [7, 6, 5]
[1] [4, 3, 2] [7, 6, 5]
[1] [4, 3] [7, 6, 5, 2]
[] [4, 3] [7, 6, 5, 2, 1]
[3] [4] [7, 6, 5, 2, 1]
[3] [4, 1] [7, 6, 5, 2]
[3, 2] [4, 1] [7, 6, 5]
[3, 2, 1] [4] [7, 6, 5]
[3, 2, 1] [] [7, 6, 5, 4]
[3, 2] [] [7, 6, 5, 4, 1]
[3] [2] [7, 6, 5, 4, 1]
[3] [2, 1] [7, 6, 5, 4]
[] [2, 1] [7, 6, 5, 4, 3]
[1] [2] [7, 6, 5, 4, 3]
[1] [] [7, 6, 5, 4, 3, 2]
[] [] [7, 6, 5, 4, 3, 2, 1]"""
assert hanoi_solver(7) == expected

expected = """[8, 7, 6, 5, 4, 3, 2, 1] [] []
[8, 7, 6, 5, 4, 3, 2] [1] []
[8, 7, 6, 5, 4, 3] [1] [2]
[8, 7, 6, 5, 4, 3] [] [2, 1]
[8, 7, 6, 5, 4] [3] [2, 1]
[8, 7, 6, 5, 4, 1] [3] [2]
[8, 7, 6, 5, 4, 1] [3, 2] []
[8, 7, 6, 5, 4] [3, 2, 1] []
[8, 7, 6, 5] [3, 2, 1] [4]
[8, 7, 6, 5] [3, 2] [4, 1]
[8, 7, 6, 5, 2] [3] [4, 1]
[8, 7, 6, 5, 2, 1] [3] [4]
[8, 7, 6, 5, 2, 1] [] [4, 3]
[8, 7, 6, 5, 2] [1] [4, 3]
[8, 7, 6, 5] [1] [4, 3, 2]
[8, 7, 6, 5] [] [4, 3, 2, 1]
[8, 7, 6] [5] [4, 3, 2, 1]
[8, 7, 6, 1] [5] [4, 3, 2]
[8, 7, 6, 1] [5, 2] [4, 3]
[8, 7, 6] [5, 2, 1] [4, 3]
[8, 7, 6, 3] [5, 2, 1] [4]
[8, 7, 6, 3] [5, 2] [4, 1]
[8, 7, 6, 3, 2] [5] [4, 1]
[8, 7, 6, 3, 2, 1] [5] [4]
[8, 7, 6, 3, 2, 1] [5, 4] []
[8, 7, 6, 3, 2] [5, 4, 1] []
[8, 7, 6, 3] [5, 4, 1] [2]
[8, 7, 6, 3] [5, 4] [2, 1]
[8, 7, 6] [5, 4, 3] [2, 1]
[8, 7, 6, 1] [5, 4, 3] [2]
[8, 7, 6, 1] [5, 4, 3, 2] []
[8, 7, 6] [5, 4, 3, 2, 1] []
[8, 7] [5, 4, 3, 2, 1] [6]
[8, 7] [5, 4, 3, 2] [6, 1]
[8, 7, 2] [5, 4, 3] [6, 1]
[8, 7, 2, 1] [5, 4, 3] [6]
[8, 7, 2, 1] [5, 4] [6, 3]
[8, 7, 2] [5, 4, 1] [6, 3]
[8, 7] [5, 4, 1] [6, 3, 2]
[8, 7] [5, 4] [6, 3, 2, 1]
[8, 7, 4] [5] [6, 3, 2, 1]
[8, 7, 4, 1] [5] [6, 3, 2]
[8, 7, 4, 1] [5, 2] [6, 3]
[8, 7, 4] [5, 2, 1] [6, 3]
[8, 7, 4, 3] [5, 2, 1] [6]
[8, 7, 4, 3] [5, 2] [6, 1]
[8, 7, 4, 3, 2] [5] [6, 1]
[8, 7, 4, 3, 2, 1] [5] [6]
[8, 7, 4, 3, 2, 1] [] [6, 5]
[8, 7, 4, 3, 2] [1] [6, 5]
[8, 7, 4, 3] [1] [6, 5, 2]
[8, 7, 4, 3] [] [6, 5, 2, 1]
[8, 7, 4] [3] [6, 5, 2, 1]
[8, 7, 4, 1] [3] [6, 5, 2]
[8, 7, 4, 1] [3, 2] [6, 5]
[8, 7, 4] [3, 2, 1] [6, 5]
[8, 7] [3, 2, 1] [6, 5, 4]
[8, 7] [3, 2] [6, 5, 4, 1]
[8, 7, 2] [3] [6, 5, 4, 1]
[8, 7, 2, 1] [3] [6, 5, 4]
[8, 7, 2, 1] [] [6, 5, 4, 3]
[8, 7, 2] [1] [6, 5, 4, 3]
[8, 7] [1] [6, 5, 4, 3, 2]
[8, 7] [] [6, 5, 4, 3, 2, 1]
[8] [7] [6, 5, 4, 3, 2, 1]
[8, 1] [7] [6, 5, 4, 3, 2]
[8, 1] [7, 2] [6, 5, 4, 3]
[8] [7, 2, 1] [6, 5, 4, 3]
[8, 3] [7, 2, 1] [6, 5, 4]
[8, 3] [7, 2] [6, 5, 4, 1]
[8, 3, 2] [7] [6, 5, 4, 1]
[8, 3, 2, 1] [7] [6, 5, 4]
[8, 3, 2, 1] [7, 4] [6, 5]
[8, 3, 2] [7, 4, 1] [6, 5]
[8, 3] [7, 4, 1] [6, 5, 2]
[8, 3] [7, 4] [6, 5, 2, 1]
[8] [7, 4, 3] [6, 5, 2, 1]
[8, 1] [7, 4, 3] [6, 5, 2]
[8, 1] [7, 4, 3, 2] [6, 5]
[8] [7, 4, 3, 2, 1] [6, 5]
[8, 5] [7, 4, 3, 2, 1] [6]
[8, 5] [7, 4, 3, 2] [6, 1]
[8, 5, 2] [7, 4, 3] [6, 1]
[8, 5, 2, 1] [7, 4, 3] [6]
[8, 5, 2, 1] [7, 4] [6, 3]
[8, 5, 2] [7, 4, 1] [6, 3]
[8, 5] [7, 4, 1] [6, 3, 2]
[8, 5] [7, 4] [6, 3, 2, 1]
[8, 5, 4] [7] [6, 3, 2, 1]
[8, 5, 4, 1] [7] [6, 3, 2]
[8, 5, 4, 1] [7, 2] [6, 3]
[8, 5, 4] [7, 2, 1] [6, 3]
[8, 5, 4, 3] [7, 2, 1] [6]
[8, 5, 4, 3] [7, 2] [6, 1]
[8, 5, 4, 3, 2] [7] [6, 1]
[8, 5, 4, 3, 2, 1] [7] [6]
[8, 5, 4, 3, 2, 1] [7, 6] []
[8, 5, 4, 3, 2] [7, 6, 1] []
[8, 5, 4, 3] [7, 6, 1] [2]
[8, 5, 4, 3] [7, 6] [2, 1]
[8, 5, 4] [7, 6, 3] [2, 1]
[8, 5, 4, 1] [7, 6, 3] [2]
[8, 5, 4, 1] [7, 6, 3, 2] []
[8, 5, 4] [7, 6, 3, 2, 1] []
[8, 5] [7, 6, 3, 2, 1] [4]
[8, 5] [7, 6, 3, 2] [4, 1]
[8, 5, 2] [7, 6, 3] [4, 1]
[8, 5, 2, 1] [7, 6, 3] [4]
[8, 5, 2, 1] [7, 6] [4, 3]
[8, 5, 2] [7, 6, 1] [4, 3]
[8, 5] [7, 6, 1] [4, 3, 2]
[8, 5] [7, 6] [4, 3, 2, 1]
[8] [7, 6, 5] [4, 3, 2, 1]
[8, 1] [7, 6, 5] [4, 3, 2]
[8, 1] [7, 6, 5, 2] [4, 3]
[8] [7, 6, 5, 2, 1] [4, 3]
[8, 3] [7, 6, 5, 2, 1] [4]
[8, 3] [7, 6, 5, 2] [4, 1]
[8, 3, 2] [7, 6, 5] [4, 1]
[8, 3, 2, 1] [7, 6, 5] [4]
[8, 3, 2, 1] [7, 6, 5, 4] []
[8, 3, 2] [7, 6, 5, 4, 1] []
[8, 3] [7, 6, 5, 4, 1] [2]
[8, 3] [7, 6, 5, 4] [2, 1]
[8] [7, 6, 5, 4, 3] [2, 1]
[8, 1] [7, 6, 5, 4, 3] [2]
[8, 1] [7, 6, 5, 4, 3, 2] []
[8] [7, 6, 5, 4, 3, 2, 1] []
[] [7, 6, 5, 4, 3, 2, 1] [8]
[] [7, 6, 5, 4, 3, 2] [8, 1]
[2] [7, 6, 5, 4, 3] [8, 1]
[2, 1] [7, 6, 5, 4, 3] [8]
[2, 1] [7, 6, 5, 4] [8, 3]
[2] [7, 6, 5, 4, 1] [8, 3]
[] [7, 6, 5, 4, 1] [8, 3, 2]
[] [7, 6, 5, 4] [8, 3, 2, 1]
[4] [7, 6, 5] [8, 3, 2, 1]
[4, 1] [7, 6, 5] [8, 3, 2]
[4, 1] [7, 6, 5, 2] [8, 3]
[4] [7, 6, 5, 2, 1] [8, 3]
[4, 3] [7, 6, 5, 2, 1] [8]
[4, 3] [7, 6, 5, 2] [8, 1]
[4, 3, 2] [7, 6, 5] [8, 1]
[4, 3, 2, 1] [7, 6, 5] [8]
[4, 3, 2, 1] [7, 6] [8, 5]
[4, 3, 2] [7, 6, 1] [8, 5]
[4, 3] [7, 6, 1] [8, 5, 2]
[4, 3] [7, 6] [8, 5, 2, 1]
[4] [7, 6, 3] [8, 5, 2, 1]
[4, 1] [7, 6, 3] [8, 5, 2]
[4, 1] [7, 6, 3, 2] [8, 5]
[4] [7, 6, 3, 2, 1] [8, 5]
[] [7, 6, 3, 2, 1] [8, 5, 4]
[] [7, 6, 3, 2] [8, 5, 4, 1]
[2] [7, 6, 3] [8, 5, 4, 1]
[2, 1] [7, 6, 3] [8, 5, 4]
[2, 1] [7, 6] [8, 5, 4, 3]
[2] [7, 6, 1] [8, 5, 4, 3]
[] [7, 6, 1] [8, 5, 4, 3, 2]
[] [7, 6] [8, 5, 4, 3, 2, 1]
[6] [7] [8, 5, 4, 3, 2, 1]
[6, 1] [7] [8, 5, 4, 3, 2]
[6, 1] [7, 2] [8, 5, 4, 3]
[6] [7, 2, 1] [8, 5, 4, 3]
[6, 3] [7, 2, 1] [8, 5, 4]
[6, 3] [7, 2] [8, 5, 4, 1]
[6, 3, 2] [7] [8, 5, 4, 1]
[6, 3, 2, 1] [7] [8, 5, 4]
[6, 3, 2, 1] [7, 4] [8, 5]
[6, 3, 2] [7, 4, 1] [8, 5]
[6, 3] [7, 4, 1] [8, 5, 2]
[6, 3] [7, 4] [8, 5, 2, 1]
[6] [7, 4, 3] [8, 5, 2, 1]
[6, 1] [7, 4, 3] [8, 5, 2]
[6, 1] [7, 4, 3, 2] [8, 5]
[6] [7, 4, 3, 2, 1] [8, 5]
[6, 5] [7, 4, 3, 2, 1] [8]
[6, 5] [7, 4, 3, 2] [8, 1]
[6, 5, 2] [7, 4, 3] [8, 1]
[6, 5, 2, 1] [7, 4, 3] [8]
[6, 5, 2, 1] [7, 4] [8, 3]
[6, 5, 2] [7, 4, 1] [8, 3]
[6, 5] [7, 4, 1] [8, 3, 2]
[6, 5] [7, 4] [8, 3, 2, 1]
[6, 5, 4] [7] [8, 3, 2, 1]
[6, 5, 4, 1] [7] [8, 3, 2]
[6, 5, 4, 1] [7, 2] [8, 3]
[6, 5, 4] [7, 2, 1] [8, 3]
[6, 5, 4, 3] [7, 2, 1] [8]
[6, 5, 4, 3] [7, 2] [8, 1]
[6, 5, 4, 3, 2] [7] [8, 1]
[6, 5, 4, 3, 2, 1] [7] [8]
[6, 5, 4, 3, 2, 1] [] [8, 7]
[6, 5, 4, 3, 2] [1] [8, 7]
[6, 5, 4, 3] [1] [8, 7, 2]
[6, 5, 4, 3] [] [8, 7, 2, 1]
[6, 5, 4] [3] [8, 7, 2, 1]
[6, 5, 4, 1] [3] [8, 7, 2]
[6, 5, 4, 1] [3, 2] [8, 7]
[6, 5, 4] [3, 2, 1] [8, 7]
[6, 5] [3, 2, 1] [8, 7, 4]
[6, 5] [3, 2] [8, 7, 4, 1]
[6, 5, 2] [3] [8, 7, 4, 1]
[6, 5, 2, 1] [3] [8, 7, 4]
[6, 5, 2, 1] [] [8, 7, 4, 3]
[6, 5, 2] [1] [8, 7, 4, 3]
[6, 5] [1] [8, 7, 4, 3, 2]
[6, 5] [] [8, 7, 4, 3, 2, 1]
[6] [5] [8, 7, 4, 3, 2, 1]
[6, 1] [5] [8, 7, 4, 3, 2]
[6, 1] [5, 2] [8, 7, 4, 3]
[6] [5, 2, 1] [8, 7, 4, 3]
[6, 3] [5, 2, 1] [8, 7, 4]
[6, 3] [5, 2] [8, 7, 4, 1]
[6, 3, 2] [5] [8, 7, 4, 1]
[6, 3, 2, 1] [5] [8, 7, 4]
[6, 3, 2, 1] [5, 4] [8, 7]
[6, 3, 2] [5, 4, 1] [8, 7]
[6, 3] [5, 4, 1] [8, 7, 2]
[6, 3] [5, 4] [8, 7, 2, 1]
[6] [5, 4, 3] [8, 7, 2, 1]
[6, 1] [5, 4, 3] [8, 7, 2]
[6, 1] [5, 4, 3, 2] [8, 7]
[6] [5, 4, 3, 2, 1] [8, 7]
[] [5, 4, 3, 2, 1] [8, 7, 6]
[] [5, 4, 3, 2] [8, 7, 6, 1]
[2] [5, 4, 3] [8, 7, 6, 1]
[2, 1] [5, 4, 3] [8, 7, 6]
[2, 1] [5, 4] [8, 7, 6, 3]
[2] [5, 4, 1] [8, 7, 6, 3]
[] [5, 4, 1] [8, 7, 6, 3, 2]
[] [5, 4] [8, 7, 6, 3, 2, 1]
[4] [5] [8, 7, 6, 3, 2, 1]
[4, 1] [5] [8, 7, 6, 3, 2]
[4, 1] [5, 2] [8, 7, 6, 3]
[4] [5, 2, 1] [8, 7, 6, 3]
[4, 3] [5, 2, 1] [8, 7, 6]
[4, 3] [5, 2] [8, 7, 6, 1]
[4, 3, 2] [5] [8, 7, 6, 1]
[4, 3, 2, 1] [5] [8, 7, 6]
[4, 3, 2, 1] [] [8, 7, 6, 5]
[4, 3, 2] [1] [8, 7, 6, 5]
[4, 3] [1] [8, 7, 6, 5, 2]
[4, 3] [] [8, 7, 6, 5, 2, 1]
[4] [3] [8, 7, 6, 5, 2, 1]
[4, 1] [3] [8, 7, 6, 5, 2]
[4, 1] [3, 2] [8, 7, 6, 5]
[4] [3, 2, 1] [8, 7, 6, 5]
[] [3, 2, 1] [8, 7, 6, 5, 4]
[] [3, 2] [8, 7, 6, 5, 4, 1]
[2] [3] [8, 7, 6, 5, 4, 1]
[2, 1] [3] [8, 7, 6, 5, 4]
[2, 1] [] [8, 7, 6, 5, 4, 3]
[2] [1] [8, 7, 6, 5, 4, 3]
[] [1] [8, 7, 6, 5, 4, 3, 2]
[] [] [8, 7, 6, 5, 4, 3, 2, 1]"""
assert hanoi_solver(8) == expected

expected = """[9, 8, 7, 6, 5, 4, 3, 2, 1] [] []
[9, 8, 7, 6, 5, 4, 3, 2] [] [1]
[9, 8, 7, 6, 5, 4, 3] [2] [1]
[9, 8, 7, 6, 5, 4, 3] [2, 1] []
[9, 8, 7, 6, 5, 4] [2, 1] [3]
[9, 8, 7, 6, 5, 4, 1] [2] [3]
[9, 8, 7, 6, 5, 4, 1] [] [3, 2]
[9, 8, 7, 6, 5, 4] [] [3, 2, 1]
[9, 8, 7, 6, 5] [4] [3, 2, 1]
[9, 8, 7, 6, 5] [4, 1] [3, 2]
[9, 8, 7, 6, 5, 2] [4, 1] [3]
[9, 8, 7, 6, 5, 2, 1] [4] [3]
[9, 8, 7, 6, 5, 2, 1] [4, 3] []
[9, 8, 7, 6, 5, 2] [4, 3] [1]
[9, 8, 7, 6, 5] [4, 3, 2] [1]
[9, 8, 7, 6, 5] [4, 3, 2, 1] []
[9, 8, 7, 6] [4, 3, 2, 1] [5]
[9, 8, 7, 6, 1] [4, 3, 2] [5]
[9, 8, 7, 6, 1] [4, 3] [5, 2]
[9, 8, 7, 6] [4, 3] [5, 2, 1]
[9, 8, 7, 6, 3] [4] [5, 2, 1]
[9, 8, 7, 6, 3] [4, 1] [5, 2]
[9, 8, 7, 6, 3, 2] [4, 1] [5]
[9, 8, 7, 6, 3, 2, 1] [4] [5]
[9, 8, 7, 6, 3, 2, 1] [] [5, 4]
[9, 8, 7, 6, 3, 2] [] [5, 4, 1]
[9, 8, 7, 6, 3] [2] [5, 4, 1]
[9, 8, 7, 6, 3] [2, 1] [5, 4]
[9, 8, 7, 6] [2, 1] [5, 4, 3]
[9, 8, 7, 6, 1] [2] [5, 4, 3]
[9, 8, 7, 6, 1] [] [5, 4, 3, 2]
[9, 8, 7, 6] [] [5, 4, 3, 2, 1]
[9, 8, 7] [6] [5, 4, 3, 2, 1]
[9, 8, 7] [6, 1] [5, 4, 3, 2]
[9, 8, 7, 2] [6, 1] [5, 4, 3]
[9, 8, 7, 2, 1] [6] [5, 4, 3]
[9, 8, 7, 2, 1] [6, 3] [5, 4]
[9, 8, 7, 2] [6, 3] [5, 4, 1]
[9, 8, 7] [6, 3, 2] [5, 4, 1]
[9, 8, 7] [6, 3, 2, 1] [5, 4]
[9, 8, 7, 4] [6, 3, 2, 1] [5]
[9, 8, 7, 4, 1] [6, 3, 2] [5]
[9, 8, 7, 4, 1] [6, 3] [5, 2]
[9, 8, 7, 4] [6, 3] [5, 2, 1]
[9, 8, 7, 4, 3] [6] [5, 2, 1]
[9, 8, 7, 4, 3] [6, 1] [5, 2]
[9, 8, 7, 4, 3, 2] [6, 1] [5]
[9, 8, 7, 4, 3, 2, 1] [6] [5]
[9, 8, 7, 4, 3, 2, 1] [6, 5] []
[9, 8, 7, 4, 3, 2] [6, 5] [1]
[9, 8, 7, 4, 3] [6, 5, 2] [1]
[9, 8, 7, 4, 3] [6, 5, 2, 1] []
[9, 8, 7, 4] [6, 5, 2, 1] [3]
[9, 8, 7, 4, 1] [6, 5, 2] [3]
[9, 8, 7, 4, 1] [6, 5] [3, 2]
[9, 8, 7, 4] [6, 5] [3, 2, 1]
[9, 8, 7] [6, 5, 4] [3, 2, 1]
[9, 8, 7] [6, 5, 4, 1] [3, 2]
[9, 8, 7, 2] [6, 5, 4, 1] [3]
[9, 8, 7, 2, 1] [6, 5, 4] [3]
[9, 8, 7, 2, 1] [6, 5, 4, 3] []
[9, 8, 7, 2] [6, 5, 4, 3] [1]
[9, 8, 7] [6, 5, 4, 3, 2] [1]
[9, 8, 7] [6, 5, 4, 3, 2, 1] []
[9, 8] [6, 5, 4, 3, 2, 1] [7]
[9, 8, 1] [6, 5, 4, 3, 2] [7]
[9, 8, 1] [6, 5, 4, 3] [7, 2]
[9, 8] [6, 5, 4, 3] [7, 2, 1]
[9, 8, 3] [6, 5, 4] [7, 2, 1]
[9, 8, 3] [6, 5, 4, 1] [7, 2]
[9, 8, 3, 2] [6, 5, 4, 1] [7]
[9, 8, 3, 2, 1] [6, 5, 4] [7]
[9, 8, 3, 2, 1] [6, 5] [7, 4]
[9, 8, 3, 2] [6, 5] [7, 4, 1]
[9, 8, 3] [6, 5, 2] [7, 4, 1]
[9, 8, 3] [6, 5, 2, 1] [7, 4]
[9, 8] [6, 5, 2, 1] [7, 4, 3]
[9, 8, 1] [6, 5, 2] [7, 4, 3]
[9, 8, 1] [6, 5] [7, 4, 3, 2]
[9, 8] [6, 5] [7, 4, 3, 2, 1]
[9, 8, 5] [6] [7, 4, 3, 2, 1]
[9, 8, 5] [6, 1] [7, 4, 3, 2]
[9, 8, 5, 2] [6, 1] [7, 4, 3]
[9, 8, 5, 2, 1] [6] [7, 4, 3]
[9, 8, 5, 2, 1] [6, 3] [7, 4]
[9, 8, 5, 2] [6, 3] [7, 4, 1]
[9, 8, 5] [6, 3, 2] [7, 4, 1]
[9, 8, 5] [6, 3, 2, 1] [7, 4]
[9, 8, 5, 4] [6, 3, 2, 1] [7]
[9, 8, 5, 4, 1] [6, 3, 2] [7]
[9, 8, 5, 4, 1] [6, 3] [7, 2]
[9, 8, 5, 4] [6, 3] [7, 2, 1]
[9, 8, 5, 4, 3] [6] [7, 2, 1]
[9, 8, 5, 4, 3] [6, 1] [7, 2]
[9, 8, 5, 4, 3, 2] [6, 1] [7]
[9, 8, 5, 4, 3, 2, 1] [6] [7]
[9, 8, 5, 4, 3, 2, 1] [] [7, 6]
[9, 8, 5, 4, 3, 2] [] [7, 6, 1]
[9, 8, 5, 4, 3] [2] [7, 6, 1]
[9, 8, 5, 4, 3] [2, 1] [7, 6]
[9, 8, 5, 4] [2, 1] [7, 6, 3]
[9, 8, 5, 4, 1] [2] [7, 6, 3]
[9, 8, 5, 4, 1] [] [7, 6, 3, 2]
[9, 8, 5, 4] [] [7, 6, 3, 2, 1]
[9, 8, 5] [4] [7, 6, 3, 2, 1]
[9, 8, 5] [4, 1] [7, 6, 3, 2]
[9, 8, 5, 2] [4, 1] [7, 6, 3]
[9, 8, 5, 2, 1] [4] [7, 6, 3]
[9, 8, 5, 2, 1] [4, 3] [7, 6]
[9, 8, 5, 2] [4, 3] [7, 6, 1]
[9, 8, 5] [4, 3, 2] [7, 6, 1]
[9, 8, 5] [4, 3, 2, 1] [7, 6]
[9, 8] [4, 3, 2, 1] [7, 6, 5]
[9, 8, 1] [4, 3, 2] [7, 6, 5]
[9, 8, 1] [4, 3] [7, 6, 5, 2]
[9, 8] [4, 3] [7, 6, 5, 2, 1]
[9, 8, 3] [4] [7, 6, 5, 2, 1]
[9, 8, 3] [4, 1] [7, 6, 5, 2]
[9, 8, 3, 2] [4, 1] [7, 6, 5]
[9, 8, 3, 2, 1] [4] [7, 6, 5]
[9, 8, 3, 2, 1] [] [7, 6, 5, 4]
[9, 8, 3, 2] [] [7, 6, 5, 4, 1]
[9, 8, 3] [2] [7, 6, 5, 4, 1]
[9, 8, 3] [2, 1] [7, 6, 5, 4]
[9, 8] [2, 1] [7, 6, 5, 4, 3]
[9, 8, 1] [2] [7, 6, 5, 4, 3]
[9, 8, 1] [] [7, 6, 5, 4, 3, 2]
[9, 8] [] [7, 6, 5, 4, 3, 2, 1]
[9] [8] [7, 6, 5, 4, 3, 2, 1]
[9] [8, 1] [7, 6, 5, 4, 3, 2]
[9, 2] [8, 1] [7, 6, 5, 4, 3]
[9, 2, 1] [8] [7, 6, 5, 4, 3]
[9, 2, 1] [8, 3] [7, 6, 5, 4]
[9, 2] [8, 3] [7, 6, 5, 4, 1]
[9] [8, 3, 2] [7, 6, 5, 4, 1]
[9] [8, 3, 2, 1] [7, 6, 5, 4]
[9, 4] [8, 3, 2, 1] [7, 6, 5]
[9, 4, 1] [8, 3, 2] [7, 6, 5]
[9, 4, 1] [8, 3] [7, 6, 5, 2]
[9, 4] [8, 3] [7, 6, 5, 2, 1]
[9, 4, 3] [8] [7, 6, 5, 2, 1]
[9, 4, 3] [8, 1] [7, 6, 5, 2]
[9, 4, 3, 2] [8, 1] [7, 6, 5]
[9, 4, 3, 2, 1] [8] [7, 6, 5]
[9, 4, 3, 2, 1] [8, 5] [7, 6]
[9, 4, 3, 2] [8, 5] [7, 6, 1]
[9, 4, 3] [8, 5, 2] [7, 6, 1]
[9, 4, 3] [8, 5, 2, 1] [7, 6]
[9, 4] [8, 5, 2, 1] [7, 6, 3]
[9, 4, 1] [8, 5, 2] [7, 6, 3]
[9, 4, 1] [8, 5] [7, 6, 3, 2]
[9, 4] [8, 5] [7, 6, 3, 2, 1]
[9] [8, 5, 4] [7, 6, 3, 2, 1]
[9] [8, 5, 4, 1] [7, 6, 3, 2]
[9, 2] [8, 5, 4, 1] [7, 6, 3]
[9, 2, 1] [8, 5, 4] [7, 6, 3]
[9, 2, 1] [8, 5, 4, 3] [7, 6]
[9, 2] [8, 5, 4, 3] [7, 6, 1]
[9] [8, 5, 4, 3, 2] [7, 6, 1]
[9] [8, 5, 4, 3, 2, 1] [7, 6]
[9, 6] [8, 5, 4, 3, 2, 1] [7]
[9, 6, 1] [8, 5, 4, 3, 2] [7]
[9, 6, 1] [8, 5, 4, 3] [7, 2]
[9, 6] [8, 5, 4, 3] [7, 2, 1]
[9, 6, 3] [8, 5, 4] [7, 2, 1]
[9, 6, 3] [8, 5, 4, 1] [7, 2]
[9, 6, 3, 2] [8, 5, 4, 1] [7]
[9, 6, 3, 2, 1] [8, 5, 4] [7]
[9, 6, 3, 2, 1] [8, 5] [7, 4]
[9, 6, 3, 2] [8, 5] [7, 4, 1]
[9, 6, 3] [8, 5, 2] [7, 4, 1]
[9, 6, 3] [8, 5, 2, 1] [7, 4]
[9, 6] [8, 5, 2, 1] [7, 4, 3]
[9, 6, 1] [8, 5, 2] [7, 4, 3]
[9, 6, 1] [8, 5] [7, 4, 3, 2]
[9, 6] [8, 5] [7, 4, 3, 2, 1]
[9, 6, 5] [8] [7, 4, 3, 2, 1]
[9, 6, 5] [8, 1] [7, 4, 3, 2]
[9, 6, 5, 2] [8, 1] [7, 4, 3]
[9, 6, 5, 2, 1] [8] [7, 4, 3]
[9, 6, 5, 2, 1] [8, 3] [7, 4]
[9, 6, 5, 2] [8, 3] [7, 4, 1]
[9, 6, 5] [8, 3, 2] [7, 4, 1]
[9, 6, 5] [8, 3, 2, 1] [7, 4]
[9, 6, 5, 4] [8, 3, 2, 1] [7]
[9, 6, 5, 4, 1] [8, 3, 2] [7]
[9, 6, 5, 4, 1] [8, 3] [7, 2]
[9, 6, 5, 4] [8, 3] [7, 2, 1]
[9, 6, 5, 4, 3] [8] [7, 2, 1]
[9, 6, 5, 4, 3] [8, 1] [7, 2]
[9, 6, 5, 4, 3, 2] [8, 1] [7]
[9, 6, 5, 4, 3, 2, 1] [8] [7]
[9, 6, 5, 4, 3, 2, 1] [8, 7] []
[9, 6, 5, 4, 3, 2] [8, 7] [1]
[9, 6, 5, 4, 3] [8, 7, 2] [1]
[9, 6, 5, 4, 3] [8, 7, 2, 1] []
[9, 6, 5, 4] [8, 7, 2, 1] [3]
[9, 6, 5, 4, 1] [8, 7, 2] [3]
[9, 6, 5, 4, 1] [8, 7] [3, 2]
[9, 6, 5, 4] [8, 7] [3, 2, 1]
[9, 6, 5] [8, 7, 4] [3, 2, 1]
[9, 6, 5] [8, 7, 4, 1] [3, 2]
[9, 6, 5, 2] [8, 7, 4, 1] [3]
[9, 6, 5, 2, 1] [8, 7, 4] [3]
[9, 6, 5, 2, 1] [8, 7, 4, 3] []
[9, 6, 5, 2] [8, 7, 4, 3] [1]
[9, 6, 5] [8, 7, 4, 3, 2] [1]
[9, 6, 5] [8, 7, 4, 3, 2, 1] []
[9, 6] [8, 7, 4, 3, 2, 1] [5]
[9, 6, 1] [8, 7, 4, 3, 2] [5]
[9, 6, 1] [8, 7, 4, 3] [5, 2]
[9, 6] [8, 7, 4, 3] [5, 2, 1]
[9, 6, 3] [8, 7, 4] [5, 2, 1]
[9, 6, 3] [8, 7, 4, 1] [5, 2]
[9, 6, 3, 2] [8, 7, 4, 1] [5]
[9, 6, 3, 2, 1] [8, 7, 4] [5]
[9, 6, 3, 2, 1] [8, 7] [5, 4]
[9, 6, 3, 2] [8, 7] [5, 4, 1]
[9, 6, 3] [8, 7, 2] [5, 4, 1]
[9, 6, 3] [8, 7, 2, 1] [5, 4]
[9, 6] [8, 7, 2, 1] [5, 4, 3]
[9, 6, 1] [8, 7, 2] [5, 4, 3]
[9, 6, 1] [8, 7] [5, 4, 3, 2]
[9, 6] [8, 7] [5, 4, 3, 2, 1]
[9] [8, 7, 6] [5, 4, 3, 2, 1]
[9] [8, 7, 6, 1] [5, 4, 3, 2]
[9, 2] [8, 7, 6, 1] [5, 4, 3]
[9, 2, 1] [8, 7, 6] [5, 4, 3]
[9, 2, 1] [8, 7, 6, 3] [5, 4]
[9, 2] [8, 7, 6, 3] [5, 4, 1]
[9] [8, 7, 6, 3, 2] [5, 4, 1]
[9] [8, 7, 6, 3, 2, 1] [5, 4]
[9, 4] [8, 7, 6, 3, 2, 1] [5]
[9, 4, 1] [8, 7, 6, 3, 2] [5]
[9, 4, 1] [8, 7, 6, 3] [5, 2]
[9, 4] [8, 7, 6, 3] [5, 2, 1]
[9, 4, 3] [8, 7, 6] [5, 2, 1]
[9, 4, 3] [8, 7, 6, 1] [5, 2]
[9, 4, 3, 2] [8, 7, 6, 1] [5]
[9, 4, 3, 2, 1] [8, 7, 6] [5]
[9, 4, 3, 2, 1] [8, 7, 6, 5] []
[9, 4, 3, 2] [8, 7, 6, 5] [1]
[9, 4, 3] [8, 7, 6, 5, 2] [1]
[9, 4, 3] [8, 7, 6, 5, 2, 1] []
[9, 4] [8, 7, 6, 5, 2, 1] [3]
[9, 4, 1] [8, 7, 6, 5, 2] [3]
[9, 4, 1] [8, 7, 6, 5] [3, 2]
[9, 4] [8, 7, 6, 5] [3, 2, 1]
[9] [8, 7, 6, 5, 4] [3, 2, 1]
[9] [8, 7, 6, 5, 4, 1] [3, 2]
[9, 2] [8, 7, 6, 5, 4, 1] [3]
[9, 2, 1] [8, 7, 6, 5, 4] [3]
[9, 2, 1] [8, 7, 6, 5, 4, 3] []
[9, 2] [8, 7, 6, 5, 4, 3] [1]
[9] [8, 7, 6, 5, 4, 3, 2] [1]
[9] [8, 7, 6, 5, 4, 3, 2, 1] []
[] [8, 7, 6, 5, 4, 3, 2, 1] [9]
[1] [8, 7, 6, 5, 4, 3, 2] [9]
[1] [8, 7, 6, 5, 4, 3] [9, 2]
[] [8, 7, 6, 5, 4, 3] [9, 2, 1]
[3] [8, 7, 6, 5, 4] [9, 2, 1]
[3] [8, 7, 6, 5, 4, 1] [9, 2]
[3, 2] [8, 7, 6, 5, 4, 1] [9]
[3, 2, 1] [8, 7, 6, 5, 4] [9]
[3, 2, 1] [8, 7, 6, 5] [9, 4]
[3, 2] [8, 7, 6, 5] [9, 4, 1]
[3] [8, 7, 6, 5, 2] [9, 4, 1]
[3] [8, 7, 6, 5, 2, 1] [9, 4]
[] [8, 7, 6, 5, 2, 1] [9, 4, 3]
[1] [8, 7, 6, 5, 2] [9, 4, 3]
[1] [8, 7, 6, 5] [9, 4, 3, 2]
[] [8, 7, 6, 5] [9, 4, 3, 2, 1]
[5] [8, 7, 6] [9, 4, 3, 2, 1]
[5] [8, 7, 6, 1] [9, 4, 3, 2]
[5, 2] [8, 7, 6, 1] [9, 4, 3]
[5, 2, 1] [8, 7, 6] [9, 4, 3]
[5, 2, 1] [8, 7, 6, 3] [9, 4]
[5, 2] [8, 7, 6, 3] [9, 4, 1]
[5] [8, 7, 6, 3, 2] [9, 4, 1]
[5] [8, 7, 6, 3, 2, 1] [9, 4]
[5, 4] [8, 7, 6, 3, 2, 1] [9]
[5, 4, 1] [8, 7, 6, 3, 2] [9]
[5, 4, 1] [8, 7, 6, 3] [9, 2]
[5, 4] [8, 7, 6, 3] [9, 2, 1]
[5, 4, 3] [8, 7, 6] [9, 2, 1]
[5, 4, 3] [8, 7, 6, 1] [9, 2]
[5, 4, 3, 2] [8, 7, 6, 1] [9]
[5, 4, 3, 2, 1] [8, 7, 6] [9]
[5, 4, 3, 2, 1] [8, 7] [9, 6]
[5, 4, 3, 2] [8, 7] [9, 6, 1]
[5, 4, 3] [8, 7, 2] [9, 6, 1]
[5, 4, 3] [8, 7, 2, 1] [9, 6]
[5, 4] [8, 7, 2, 1] [9, 6, 3]
[5, 4, 1] [8, 7, 2] [9, 6, 3]
[5, 4, 1] [8, 7] [9, 6, 3, 2]
[5, 4] [8, 7] [9, 6, 3, 2, 1]
[5] [8, 7, 4] [9, 6, 3, 2, 1]
[5] [8, 7, 4, 1] [9, 6, 3, 2]
[5, 2] [8, 7, 4, 1] [9, 6, 3]
[5, 2, 1] [8, 7, 4] [9, 6, 3]
[5, 2, 1] [8, 7, 4, 3] [9, 6]
[5, 2] [8, 7, 4, 3] [9, 6, 1]
[5] [8, 7, 4, 3, 2] [9, 6, 1]
[5] [8, 7, 4, 3, 2, 1] [9, 6]
[] [8, 7, 4, 3, 2, 1] [9, 6, 5]
[1] [8, 7, 4, 3, 2] [9, 6, 5]
[1] [8, 7, 4, 3] [9, 6, 5, 2]
[] [8, 7, 4, 3] [9, 6, 5, 2, 1]
[3] [8, 7, 4] [9, 6, 5, 2, 1]
[3] [8, 7, 4, 1] [9, 6, 5, 2]
[3, 2] [8, 7, 4, 1] [9, 6, 5]
[3, 2, 1] [8, 7, 4] [9, 6, 5]
[3, 2, 1] [8, 7] [9, 6, 5, 4]
[3, 2] [8, 7] [9, 6, 5, 4, 1]
[3] [8, 7, 2] [9, 6, 5, 4, 1]
[3] [8, 7, 2, 1] [9, 6, 5, 4]
[] [8, 7, 2, 1] [9, 6, 5, 4, 3]
[1] [8, 7, 2] [9, 6, 5, 4, 3]
[1] [8, 7] [9, 6, 5, 4, 3, 2]
[] [8, 7] [9, 6, 5, 4, 3, 2, 1]
[7] [8] [9, 6, 5, 4, 3, 2, 1]
[7] [8, 1] [9, 6, 5, 4, 3, 2]
[7, 2] [8, 1] [9, 6, 5, 4, 3]
[7, 2, 1] [8] [9, 6, 5, 4, 3]
[7, 2, 1] [8, 3] [9, 6, 5, 4]
[7, 2] [8, 3] [9, 6, 5, 4, 1]
[7] [8, 3, 2] [9, 6, 5, 4, 1]
[7] [8, 3, 2, 1] [9, 6, 5, 4]
[7, 4] [8, 3, 2, 1] [9, 6, 5]
[7, 4, 1] [8, 3, 2] [9, 6, 5]
[7, 4, 1] [8, 3] [9, 6, 5, 2]
[7, 4] [8, 3] [9, 6, 5, 2, 1]
[7, 4, 3] [8] [9, 6, 5, 2, 1]
[7, 4, 3] [8, 1] [9, 6, 5, 2]
[7, 4, 3, 2] [8, 1] [9, 6, 5]
[7, 4, 3, 2, 1] [8] [9, 6, 5]
[7, 4, 3, 2, 1] [8, 5] [9, 6]
[7, 4, 3, 2] [8, 5] [9, 6, 1]
[7, 4, 3] [8, 5, 2] [9, 6, 1]
[7, 4, 3] [8, 5, 2, 1] [9, 6]
[7, 4] [8, 5, 2, 1] [9, 6, 3]
[7, 4, 1] [8, 5, 2] [9, 6, 3]
[7, 4, 1] [8, 5] [9, 6, 3, 2]
[7, 4] [8, 5] [9, 6, 3, 2, 1]
[7] [8, 5, 4] [9, 6, 3, 2, 1]
[7] [8, 5, 4, 1] [9, 6, 3, 2]
[7, 2] [8, 5, 4, 1] [9, 6, 3]
[7, 2, 1] [8, 5, 4] [9, 6, 3]
[7, 2, 1] [8, 5, 4, 3] [9, 6]
[7, 2] [8, 5, 4, 3] [9, 6, 1]
[7] [8, 5, 4, 3, 2] [9, 6, 1]
[7] [8, 5, 4, 3, 2, 1] [9, 6]
[7, 6] [8, 5, 4, 3, 2, 1] [9]
[7, 6, 1] [8, 5, 4, 3, 2] [9]
[7, 6, 1] [8, 5, 4, 3] [9, 2]
[7, 6] [8, 5, 4, 3] [9, 2, 1]
[7, 6, 3] [8, 5, 4] [9, 2, 1]
[7, 6, 3] [8, 5, 4, 1] [9, 2]
[7, 6, 3, 2] [8, 5, 4, 1] [9]
[7, 6, 3, 2, 1] [8, 5, 4] [9]
[7, 6, 3, 2, 1] [8, 5] [9, 4]
[7, 6, 3, 2] [8, 5] [9, 4, 1]
[7, 6, 3] [8, 5, 2] [9, 4, 1]
[7, 6, 3] [8, 5, 2, 1] [9, 4]
[7, 6] [8, 5, 2, 1] [9, 4, 3]
[7, 6, 1] [8, 5, 2] [9, 4, 3]
[7, 6, 1] [8, 5] [9, 4, 3, 2]
[7, 6] [8, 5] [9, 4, 3, 2, 1]
[7, 6, 5] [8] [9, 4, 3, 2, 1]
[7, 6, 5] [8, 1] [9, 4, 3, 2]
[7, 6, 5, 2] [8, 1] [9, 4, 3]
[7, 6, 5, 2, 1] [8] [9, 4, 3]
[7, 6, 5, 2, 1] [8, 3] [9, 4]
[7, 6, 5, 2] [8, 3] [9, 4, 1]
[7, 6, 5] [8, 3, 2] [9, 4, 1]
[7, 6, 5] [8, 3, 2, 1] [9, 4]
[7, 6, 5, 4] [8, 3, 2, 1] [9]
[7, 6, 5, 4, 1] [8, 3, 2] [9]
[7, 6, 5, 4, 1] [8, 3] [9, 2]
[7, 6, 5, 4] [8, 3] [9, 2, 1]
[7, 6, 5, 4, 3] [8] [9, 2, 1]
[7, 6, 5, 4, 3] [8, 1] [9, 2]
[7, 6, 5, 4, 3, 2] [8, 1] [9]
[7, 6, 5, 4, 3, 2, 1] [8] [9]
[7, 6, 5, 4, 3, 2, 1] [] [9, 8]
[7, 6, 5, 4, 3, 2] [] [9, 8, 1]
[7, 6, 5, 4, 3] [2] [9, 8, 1]
[7, 6, 5, 4, 3] [2, 1] [9, 8]
[7, 6, 5, 4] [2, 1] [9, 8, 3]
[7, 6, 5, 4, 1] [2] [9, 8, 3]
[7, 6, 5, 4, 1] [] [9, 8, 3, 2]
[7, 6, 5, 4] [] [9, 8, 3, 2, 1]
[7, 6, 5] [4] [9, 8, 3, 2, 1]
[7, 6, 5] [4, 1] [9, 8, 3, 2]
[7, 6, 5, 2] [4, 1] [9, 8, 3]
[7, 6, 5, 2, 1] [4] [9, 8, 3]
[7, 6, 5, 2, 1] [4, 3] [9, 8]
[7, 6, 5, 2] [4, 3] [9, 8, 1]
[7, 6, 5] [4, 3, 2] [9, 8, 1]
[7, 6, 5] [4, 3, 2, 1] [9, 8]
[7, 6] [4, 3, 2, 1] [9, 8, 5]
[7, 6, 1] [4, 3, 2] [9, 8, 5]
[7, 6, 1] [4, 3] [9, 8, 5, 2]
[7, 6] [4, 3] [9, 8, 5, 2, 1]
[7, 6, 3] [4] [9, 8, 5, 2, 1]
[7, 6, 3] [4, 1] [9, 8, 5, 2]
[7, 6, 3, 2] [4, 1] [9, 8, 5]
[7, 6, 3, 2, 1] [4] [9, 8, 5]
[7, 6, 3, 2, 1] [] [9, 8, 5, 4]
[7, 6, 3, 2] [] [9, 8, 5, 4, 1]
[7, 6, 3] [2] [9, 8, 5, 4, 1]
[7, 6, 3] [2, 1] [9, 8, 5, 4]
[7, 6] [2, 1] [9, 8, 5, 4, 3]
[7, 6, 1] [2] [9, 8, 5, 4, 3]
[7, 6, 1] [] [9, 8, 5, 4, 3, 2]
[7, 6] [] [9, 8, 5, 4, 3, 2, 1]
[7] [6] [9, 8, 5, 4, 3, 2, 1]
[7] [6, 1] [9, 8, 5, 4, 3, 2]
[7, 2] [6, 1] [9, 8, 5, 4, 3]
[7, 2, 1] [6] [9, 8, 5, 4, 3]
[7, 2, 1] [6, 3] [9, 8, 5, 4]
[7, 2] [6, 3] [9, 8, 5, 4, 1]
[7] [6, 3, 2] [9, 8, 5, 4, 1]
[7] [6, 3, 2, 1] [9, 8, 5, 4]
[7, 4] [6, 3, 2, 1] [9, 8, 5]
[7, 4, 1] [6, 3, 2] [9, 8, 5]
[7, 4, 1] [6, 3] [9, 8, 5, 2]
[7, 4] [6, 3] [9, 8, 5, 2, 1]
[7, 4, 3] [6] [9, 8, 5, 2, 1]
[7, 4, 3] [6, 1] [9, 8, 5, 2]
[7, 4, 3, 2] [6, 1] [9, 8, 5]
[7, 4, 3, 2, 1] [6] [9, 8, 5]
[7, 4, 3, 2, 1] [6, 5] [9, 8]
[7, 4, 3, 2] [6, 5] [9, 8, 1]
[7, 4, 3] [6, 5, 2] [9, 8, 1]
[7, 4, 3] [6, 5, 2, 1] [9, 8]
[7, 4] [6, 5, 2, 1] [9, 8, 3]
[7, 4, 1] [6, 5, 2] [9, 8, 3]
[7, 4, 1] [6, 5] [9, 8, 3, 2]
[7, 4] [6, 5] [9, 8, 3, 2, 1]
[7] [6, 5, 4] [9, 8, 3, 2, 1]
[7] [6, 5, 4, 1] [9, 8, 3, 2]
[7, 2] [6, 5, 4, 1] [9, 8, 3]
[7, 2, 1] [6, 5, 4] [9, 8, 3]
[7, 2, 1] [6, 5, 4, 3] [9, 8]
[7, 2] [6, 5, 4, 3] [9, 8, 1]
[7] [6, 5, 4, 3, 2] [9, 8, 1]
[7] [6, 5, 4, 3, 2, 1] [9, 8]
[] [6, 5, 4, 3, 2, 1] [9, 8, 7]
[1] [6, 5, 4, 3, 2] [9, 8, 7]
[1] [6, 5, 4, 3] [9, 8, 7, 2]
[] [6, 5, 4, 3] [9, 8, 7, 2, 1]
[3] [6, 5, 4] [9, 8, 7, 2, 1]
[3] [6, 5, 4, 1] [9, 8, 7, 2]
[3, 2] [6, 5, 4, 1] [9, 8, 7]
[3, 2, 1] [6, 5, 4] [9, 8, 7]
[3, 2, 1] [6, 5] [9, 8, 7, 4]
[3, 2] [6, 5] [9, 8, 7, 4, 1]
[3] [6, 5, 2] [9, 8, 7, 4, 1]
[3] [6, 5, 2, 1] [9, 8, 7, 4]
[] [6, 5, 2, 1] [9, 8, 7, 4, 3]
[1] [6, 5, 2] [9, 8, 7, 4, 3]
[1] [6, 5] [9, 8, 7, 4, 3, 2]
[] [6, 5] [9, 8, 7, 4, 3, 2, 1]
[5] [6] [9, 8, 7, 4, 3, 2, 1]
[5] [6, 1] [9, 8, 7, 4, 3, 2]
[5, 2] [6, 1] [9, 8, 7, 4, 3]
[5, 2, 1] [6] [9, 8, 7, 4, 3]
[5, 2, 1] [6, 3] [9, 8, 7, 4]
[5, 2] [6, 3] [9, 8, 7, 4, 1]
[5] [6, 3, 2] [9, 8, 7, 4, 1]
[5] [6, 3, 2, 1] [9, 8, 7, 4]
[5, 4] [6, 3, 2, 1] [9, 8, 7]
[5, 4, 1] [6, 3, 2] [9, 8, 7]
[5, 4, 1] [6, 3] [9, 8, 7, 2]
[5, 4] [6, 3] [9, 8, 7, 2, 1]
[5, 4, 3] [6] [9, 8, 7, 2, 1]
[5, 4, 3] [6, 1] [9, 8, 7, 2]
[5, 4, 3, 2] [6, 1] [9, 8, 7]
[5, 4, 3, 2, 1] [6] [9, 8, 7]
[5, 4, 3, 2, 1] [] [9, 8, 7, 6]
[5, 4, 3, 2] [] [9, 8, 7, 6, 1]
[5, 4, 3] [2] [9, 8, 7, 6, 1]
[5, 4, 3] [2, 1] [9, 8, 7, 6]
[5, 4] [2, 1] [9, 8, 7, 6, 3]
[5, 4, 1] [2] [9, 8, 7, 6, 3]
[5, 4, 1] [] [9, 8, 7, 6, 3, 2]
[5, 4] [] [9, 8, 7, 6, 3, 2, 1]
[5] [4] [9, 8, 7, 6, 3, 2, 1]
[5] [4, 1] [9, 8, 7, 6, 3, 2]
[5, 2] [4, 1] [9, 8, 7, 6, 3]
[5, 2, 1] [4] [9, 8, 7, 6, 3]
[5, 2, 1] [4, 3] [9, 8, 7, 6]
[5, 2] [4, 3] [9, 8, 7, 6, 1]
[5] [4, 3, 2] [9, 8, 7, 6, 1]
[5] [4, 3, 2, 1] [9, 8, 7, 6]
[] [4, 3, 2, 1] [9, 8, 7, 6, 5]
[1] [4, 3, 2] [9, 8, 7, 6, 5]
[1] [4, 3] [9, 8, 7, 6, 5, 2]
[] [4, 3] [9, 8, 7, 6, 5, 2, 1]
[3] [4] [9, 8, 7, 6, 5, 2, 1]
[3] [4, 1] [9, 8, 7, 6, 5, 2]
[3, 2] [4, 1] [9, 8, 7, 6, 5]
[3, 2, 1] [4] [9, 8, 7, 6, 5]
[3, 2, 1] [] [9, 8, 7, 6, 5, 4]
[3, 2] [] [9, 8, 7, 6, 5, 4, 1]
[3] [2] [9, 8, 7, 6, 5, 4, 1]
[3] [2, 1] [9, 8, 7, 6, 5, 4]
[] [2, 1] [9, 8, 7, 6, 5, 4, 3]
[1] [2] [9, 8, 7, 6, 5, 4, 3]
[1] [] [9, 8, 7, 6, 5, 4, 3, 2]
[] [] [9, 8, 7, 6, 5, 4, 3, 2, 1]"""
assert hanoi_solver(9) == expected
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def hanoi_solver(n):
    # Initialize the three rods
    A = list(range(n, 0, -1))
    B = []
    C = []
    
    # List to store all moves as strings
    moves = [f"{A} {B} {C}"]
    
    def move(num_disks, source, auxiliary, target):
        if num_disks <= 0:
            return
        
        # Move n-1 disks from source to auxiliary, so they are out of the way
        move(num_disks - 1, source, target, auxiliary)
        
        # Move the nth disk from source to target
        target.append(source.pop())
        
        # Record the move
        moves.append(f"{A} {B} {C}")
        
        # Move the n-1 disks that we left on auxiliary onto target
        move(num_disks - 1, auxiliary, source, target)
    
    # Start the recursive process
    move(n, A, B, C)
    
    return '\n'.join(moves)
```
