---
id: 60f02e7361b68405e27b62a5
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, aprenderás cómo estilizar formularios diseñando un formulario de registro.

Se te ha proporcionado todo el código base HTML (`DOCTYPE`, `html`, `head`, y `body`).

Dentro del `body`, añade un encabezado para contextualizar el formulario usando un elemento `h1` con el texto `Registration Form`.

# --hints--

Debes añadir el `h1` dentro del `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Debes dar al `h1` el texto `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
