---
id: 60f030d388cb74067cf291c3
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Debajo del encabezado, utilice el siguiente texto dentro de un elemento de párrafo para animar a los usuarios a registrarse:

```md
Please fill out this form with the required information
```

# --hints--

Debe agregar un elemento `p` dentro del `body`.

```js
assert.exists(document.querySelector('body > p'));
```

Debe agregar el elemento `p` debajo del `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

Debe darle al elemento `p` un texto de `Please fill out this form with the required information`.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  --fcc-editable-region--
  
  --fcc-editable-region--
</body>
</html>
```
