---
id: 60f1922fcbd2410527b3bd89
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

En lecciones anteriores, aprendiste cómo trabajar con unidades de viewport.

Recuerda que la unidad `vh` significa altura del viewport, y es igual al 1% del `height` del viewport. Esto la hace relativa a la altura de la pantalla/ventana gráfica.

Es hora de engalanar el proyecto con un poco de CSS. Empieza dando al `body` un `width` de `100%`, y un `height` de `100vh`.

# --hints--

Deberías usar el selector de elemento `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Debes dar al `body` un `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Debes dar al `body` un `height` de `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
