---
id: 60f1922fcbd2410527b3bd89
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

In the previous lectures, you learned how to work with viewport units.

Recuerda que la unidad `vh` significa altura del viewport, y es igual al 1% del `height` del viewport. Esto lo hace relativo a la altura del viewport.

Es hora de mejorar el proyecto con algo de CSS. Comienza dando al `body` un `width` de `100%`, y un `height` de `100vh`.

# --hints--

Debes utilizar el elemento selector `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Debes darle al `body` un `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Debes darle al `body` un `height` de `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

