---
id: 68e7e729f6d14f69a23786ef
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Successivamente, all'interno del tuo elemento `main`, aggiungi un elemento `h1` con il testo `Parent Teacher Conference Form` e le classi `title` e `center`. 

# --hints--

Dovresti avere un elemento `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Il tuo elemento `h1` dovrebbe avere il testo `Parent Teacher Conference Form`.

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

Il tuo elemento `h1` dovrebbe avere le classi `title` e `center`.

```js
assert.exists(document.querySelector("h1.title.center"));
```

Il tuo elemento `h1` dovrebbe essere all'interno dell'elemento `main`.

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>

</html>
```
