---
id: 68e707257ea91c2e12d489fb
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop, eserciterai come aggiungere definizioni di stile personalizzate ai pulsanti di opzione costruendo un modulo per la conferenza genitori-insegnanti. Il codice boilerplate HTML ti è stato fornito.

Inizia aggiungendo un elemento `main` con una classe chiamata `container`.

# --hints--

Dovresti avere un elemento `main`.

```js
assert.exists(document.querySelector("main"));
```

Il tuo elemento `main` dovrebbe avere una classe chiamata `container`.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
