---
id: 68ea8a5db73b155708992687
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Il passaggio successivo è aggiungere l'elemento `input` associato alle informazioni dello studente.

Inizia aggiungendo un elemento `input` con un attributo `type` impostato su `"text"`. Poi aggiungi un attributo `name` e un attributo `id`, entrambi impostati su `student-name`.

Successivamente, aggiungi un attributo `placeholder` impostato su `E.g., Jane Doe`. E infine, aggiungi un attributo `required`.

# --hints--

Dovresti avere un elemento `input`.

```js
assert.exists(document.querySelector("fieldset input"));
```

Il tuo elemento `input` dovrebbe avere un attributo `type` impostato su `"text"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Il tuo elemento `input` dovrebbe avere un attributo `name` impostato su `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

Il tuo elemento `input` dovrebbe avere un attributo `id` impostato su `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

Il tuo elemento `input` dovrebbe avere un attributo `placeholder` impostato su `E.g., Jane Doe`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

Il tuo elemento `input` dovrebbe avere un attributo `required`.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
