---
id: 68ea8666b8b473455f5b33ef
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

La prima sezione del modulo si concentrerà sulle informazioni dello studente, come il suo nome e la sua classe.

All'interno dell'elemento `form`, aggiungi un elemento `fieldset`. All'interno di quell'elemento `fieldset`, aggiungi un elemento `legend` con il testo `Student Information`.

Sotto l'elemento `legend`, aggiungi un elemento `label` con il testo `Full Name: ` e un attributo `for` con il valore `student-name`.


# --hints--

Dovresti avere un elemento `fieldset`.

```js
assert.exists(document.querySelector("fieldset"));
```

Il tuo elemento `fieldset` dovrebbe essere all'interno dell'elemento `form`.

```js
assert.exists(document.querySelector("form fieldset"));
```

Dovresti avere un elemento `legend`.

```js
assert.exists(document.querySelector("legend"));
```

Il tuo elemento `legend` dovrebbe essere all'interno dell'elemento `fieldset`.

```js
assert.exists(document.querySelector("fieldset legend"));
```

Il tuo elemento `legend` dovrebbe avere il testo `Student Information`.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

Dovresti avere un elemento `label`.

```js
assert.exists(document.querySelector("label"));
```

Il tuo elemento `label` dovrebbe avere il testo `Full Name: `. Non dimenticare lo spazio dopo i due punti.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

Il tuo elemento `label` dovrebbe avere un attributo `for`.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

L'attributo `label` del tuo elemento `for` dovrebbe avere il valore `student-name`.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

Il tuo elemento `label` dovrebbe essere sotto l'elemento `legend`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
