---
id: 68ea8c330ed07b60879f3d38
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Ora è il momento di aggiungere gli elementi del modulo per raccogliere le informazioni sul voto dello studente.

Inizia aggiungendo un altro elemento `label` con il testo `Student Grade: ` e l'attributo `for` impostato su `"grade"`.

Poi, sotto il tuo elemento `label`, aggiungi un elemento `input` con l'attributo `type` impostato su `"number"`. Gli attributi `name` e `id` dovrebbero essere impostati su `"grade"`. L'attributo `placeholder` dovrebbe essere impostato su `"E.g., 4"`. Infine, il tuo elemento numero `input` dovrebbe essere obbligatorio.

# --hints--

Dovresti avere un totale di due elementi `label` nella pagina.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

Il tuo secondo elemento `label` dovrebbe avere il testo `Student Grade: `. Non dimenticare lo spazio dopo i due punti.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

Il tuo secondo elemento `label` dovrebbe avere un attributo `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

L'attributo `label` del tuo secondo elemento `for` dovrebbe avere il valore `"grade"`.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

Dovresti avere due elementi `input` nella pagina.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

Il tuo secondo elemento `input` dovrebbe essere sotto il tuo secondo elemento `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Il tuo secondo elemento `input` dovrebbe avere un attributo `type` impostato su `"number"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

Il tuo secondo elemento `input` dovrebbe avere un attributo `name` impostato su `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

Il tuo secondo elemento `input` dovrebbe avere un attributo `id` impostato su `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

Il tuo secondo elemento `input` dovrebbe avere un attributo `placeholder` impostato su `"E.g., 4"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

Il tuo secondo elemento `input` dovrebbe avere un attributo `required`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
