---
id: 68ea9f01c51b66b937793611
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

Successivamente, aggiungi un'altra etichetta con una classe `"contact-method"` e un attributo `for` impostato su `"phone"`. Il testo dell'etichetta dovrebbe essere `Phone: `.

Sotto il tuo elemento `label`, aggiungi un pulsante di opzione con attributi `id` e `value` impostati su `"phone"`. L'attributo `name` dovrebbe essere impostato su `"contact-method"` e la classe dovrebbe essere impostata su `"contact-method-radio-btn"`.

# --hints--

Dovresti avere due elementi `label` nell'ultimo `fieldset`.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type label"), 2);
```

L'elemento `label` dovrebbe avere una classe `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

L'attributo `label` dell'elemento `for` dovrebbe essere impostato su `"phone"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label[for='phone']"));
```

L'elemento `label` dovrebbe avere il testo `Phone: `.

```js
const label = document.querySelector("fieldset:last-of-type label[for='phone']");
assert.match(label?.innerText, /Phone:\s?/);
```

Dovresti avere due elementi `input` all'interno dell'ultimo `fieldset`.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type input"), 2);
```

Il tuo elemento `input` dovrebbe avere un attributo `type` impostato su `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

L'elemento `input` dovrebbe avere un attributo `id` impostato su `"phone"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("id"), "phone");
```

L'elemento `input` dovrebbe avere un attributo `value` impostato su `"phone"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("value"), "phone");
```

L'elemento `input` dovrebbe avere un attributo `name` impostato su `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

L'elemento `input` dovrebbe avere una classe `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input[id='phone'].contact-method-radio-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
