---
id: 68ea8fdf2c9b167307e7f671
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Ora è il momento di aggiungere gli elementi del modulo per raccogliere le informazioni del genitore.

Inizia aggiungendo un elemento `label` con il testo `Parent/Guardian Name: ` e l'attributo `for` impostato su `"parent-name"`.

Poi, sotto il tuo elemento `label`, aggiungi un elemento `input` con l'attributo `type` impostato su `"text"`. Gli attributi `name` e `id` dovrebbero essere impostati su `"parent-name"`. L'attributo `placeholder` dovrebbe essere impostato su `"E.g., Nancy Doe"`. Infine, il tuo `input` dovrebbe essere obbligatorio.

# --hints--

Dovresti avere un totale di tre elementi `label` nella pagina.

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

Il tuo terzo elemento `label` dovrebbe avere il testo `Parent/Guardian Name: `. Non dimenticare lo spazio dopo i due punti.

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

Il tuo terzo elemento `label` dovrebbe avere un attributo `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

L'attributo `label` del tuo terzo elemento `for` dovrebbe avere il valore `"parent-name"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

Dovresti avere tre elementi `input` nella pagina.

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

Il tuo terzo elemento `input` dovrebbe essere sotto il tuo terzo elemento `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Il tuo terzo elemento `input` dovrebbe avere un attributo `type` impostato su `"text"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Il tuo terzo elemento `input` dovrebbe avere un attributo `name` impostato su `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

Il tuo terzo elemento `input` dovrebbe avere un attributo `id` impostato su `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

Il tuo terzo elemento `input` dovrebbe avere un attributo `placeholder` impostato su `"E.g., Nancy Doe"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

Il tuo terzo elemento `input` dovrebbe avere un attributo `required`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
