---
id: 68ea95c833f0518ae6e366cd
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Nei prossimi passaggi aggiungerai gli elementi del modulo responsabili di raccogliere il metodo di contatto preferito dall'utente.

Inizia aggiungendo un altro elemento `fieldset` con un elemento `legend` annidato all'interno. Il tuo elemento `legend` dovrebbe contenere il testo `Preferred Contact Method`.

# --hints--

Hai un totale di tre elementi `fieldset` nella pagina.

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 3);
```

Dovresti avere un elemento `legend` all'interno del tuo terzo elemento `fieldset`.

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

Il tuo elemento `legend` dovrebbe contenere il testo `Preferred Contact Method`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Preferred Contact Method");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
