---
id: 68ea975c7cbbb894a386e402
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Il passaggio successivo è aggiungere gli elementi `label` e `input` per il metodo di contatto via email.

Inizia aggiungendo un'etichetta con una classe `"contact-method"` e un attributo `for` impostato su `"email"`. Il testo dell'etichetta dovrebbe essere `Email: `.

Sotto il tuo elemento `label`, aggiungi un pulsante di opzione con gli attributi `id` e `value` impostati su `"email"`. L'attributo `name` dovrebbe essere impostato su `"contact-method"` e la classe dovrebbe essere impostata su `"contact-method-radio-btn"`.

Infine, assicurati che questo pulsante di opzione sia selezionato per impostazione predefinita.

# --hints--


Il tuo ultimo fieldset dovrebbe contenere un elemento `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

Questo nuovo elemento `label` (per email) dovrebbe avere una classe `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

L'attributo `label` del nuovo elemento `for` dovrebbe essere impostato su `"email"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

L'elemento `label` dovrebbe avere il testo `Email: `. Assicurati dello spazio dopo i due punti.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

Dovresti avere un elemento `input` all'interno dell'ultimo `fieldset`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

Il tuo elemento `input` dovrebbe avere un attributo `type` impostato su `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

L'elemento `input` dovrebbe avere un attributo `id` impostato su `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

L'elemento `input` dovrebbe avere un attributo `value` impostato su `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

L'elemento `input` dovrebbe avere un attributo `name` impostato su `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

L'elemento `input` dovrebbe avere una classe `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

Il tuo pulsante `radio` dovrebbe essere `checked` per impostazione predefinita.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
