---
id: 68eb058081504963447bfa60
title: Passaggio 17
challengeType: 0
dashedName: step-17
---

# --description--

Ora dovresti definire lo stile del contenitore principale. Seleziona l'elemento con la classe `container` e imposta la sua `background-color` sul codice esadecimale `#ffffff1a`.

Un codice esadecimale è una combinazione di sei cifre tra numeri e lettere usata in HTML e CSS per rappresentare i colori. Inizia con un `#` seguito da tre coppie:

- Le prime due cifre rappresentano il rosso
- Le due successive rappresentano il verde
- Le ultime due rappresentano il blu

Per esempio, `#ffffff` è bianco puro perché ha il valore massimo per rosso, verde e blu.

Puoi anche aggiungere due cifre extra alla fine per controllare l'opacità (chiamata canale alfa). In `#ffffff1a`, la `1a` rende il colore bianco semi-trasparente. Più basso è il valore alfa, più trasparente appare il colore. Imparerai di più sui codici esadecimali più avanti!

# --hints--

Dovresti mirare all'elemento con la classe `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

Imposta la proprietà `background-color` su `#ffffff1a`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.container')
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```


```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

--fcc-editable-region--

--fcc-editable-region--
```
