---
id: 587d78b0367417b2b2512b08
title: створити медіазапит
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cqwKrtm'
forumTopicId: 301139
dashedName: create-a-media-query
---

# --description--

Медіазапити — це нова техніка, введена в CSS3, яка змінює подання вмісту залежно від різних розмірів вікна перегляду. Вікно перегляду — це видима користувачу область вебсторінки, яка відрізняється залежно від пристрою, що використовується для доступу до сайту.

Медіазапити складаються з типу медіа, і якщо цей тип медіа відповідає типу пристрою, на якому відображається документ, стилі застосовуються. Ви можете мати стільки селекторів і стилів у своєму медіазапиті, скільки хочете.

Ось приклад медіазапиту, який повертає вміст, коли ширина пристрою менша або дорівнює `100px`:

```css
@media (max-width: 100px) {
  /* CSS Rules */
}
```

а наступний медіазапит повертає вміст, коли висота пристрою більша або дорівнює `350px`:

```css
@media (min-height: 350px) {
  /* CSS Rules */
}
```

Пам’ятайте, що CSS всередині медіазапиту застосовується лише якщо тип медіа відповідає типу пристрою, який використовується.

# --instructions--

Додайте медіазапит, щоб тег `p` мав `font-size` зі значенням `10px`, коли висота пристрою менша або дорівнює `800px`.

# --hints--

Ви повинні оголосити `@media` запит для пристроїв з `height`, меншою або рівною `800px`.

```js
const media = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(
  media.some(x => x.media?.mediaText?.includes('(max-height: 800px)'))
);
```

Ваш елемент `p` повинен мати `font-size` зі значенням `10px`, коли висота пристрою `height` менша або дорівнює `800px`.

```js
const rules = new __helpers.CSSHelp(document).getRuleListsWithinMedia(
  '(max-height: 800px)'
);
assert.strictEqual(
  rules?.find(x => x.selectorText === 'p')?.style?.fontSize,
  '10px'
);
```

Ваш елемент `p` повинен мати ініціальний `font-size` зі значенням `20px`, коли висота пристрою `height` більша за `800px`.

```js
const ifPFirst = new __helpers.CSSHelp(document)
  .getCSSRules()
  ?.find(x => x?.selectorText === 'p' || x?.media);
assert.strictEqual(ifPFirst?.style?.fontSize, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  p {
    font-size: 20px;
  }

  /* Only change code below this line */

  /* Only change code above this line */
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  p {
    font-size: 20px;
  }

  @media (max-height: 800px) {
    p {
      font-size: 10px;
    }
  }
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
