---
id: 587d78b1367417b2b2512b0c
title: Зробіть типографіку адаптивною
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/crzN7T8'
forumTopicId: 301141
dashedName: make-typography-responsive
---

# --description--

Замість того, щоб використовувати `em` або `px` для задання розміру тексту, ви можете використовувати одиниці вікна перегляду для адаптивної типографіки. Одиниці вікна перегляду, як і відсотки, є відносними одиницями, але вони базуються на різних параметрах. Одиниці вікна перегляду відносні до розмірів вікна перегляду (ширини або висоти) пристрою, а відсотки відносні до розміру батьківського контейнера.

Існує чотири різні одиниці вікна перегляду:

<ul><li><code>vw</code> (ширина вікна перегляду): <code>10vw</code> означає 10% ширини вікна перегляду.</li><li><code>vh</code> (висота вікна перегляду): <code>3vh</code> означає 3% висоти вікна перегляду.</li><li><code>vmin</code> (мінімум вікна перегляду): <code>70vmin</code> означає 70% меншої розмірності вікна перегляду (висоти або ширини).</li><li><code>vmax</code> (максимум вікна перегляду): <code>100vmax</code> означає 100% більшої розмірності вікна перегляду (висоти або ширини).</li></ul>

Ось приклад, який задає тегу `body` розмір 30% ширини вікна перегляду.

```css
body {
  width: 30vw;
}
```

# --instructions--

Задайте `width` тегу `h2` розмір 80% ширини вікна перегляду, а `width` абзацу — 75% меншої розмірності вікна перегляду.

# --hints--

Ваш тег `h2` повинен мати `width` зі значенням 80vw.

```js
assert.match(
  __helpers.removeCssComments(code),
  /h2\s*?{\s*?width:\s*?80vw;\s*?}/g
);
```

Ваш тег `p` повинен мати `width` зі значенням 75vmin.

```js
assert.match(
  __helpers.removeCssComments(code),
  /p\s*?{\s*?width:\s*?75vmin;\s*?}/g
);
```

# --seed--

## --seed-contents--

```html
<style>

</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  h2 {
    width: 80vw;
  }
  p {
    width: 75vmin;
  }
</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
