---
id: 587d78b1367417b2b2512b0a
title: Використання зображення Retina для дисплеїв з вищою роздільністю
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cVZ4Rfp'
forumTopicId: 301142
dashedName: use-a-retina-image-for-higher-resolution-displays
---

# --description--

Зі збільшенням кількості пристроїв, підключених до інтернету, їхні розміри та характеристики різняться, і дисплеї, які вони використовують, можуть відрізнятися зовні та внутрішньо. Щільність пікселів — це аспект, який може відрізнятися на одному пристрої від інших, і ця щільність відома як Pixel Per Inch (PPI) або Dots Per Inch (DPI). Найвідомішим таким дисплеєм є той, що відомий як "Retina Display" на останніх ноутбуках Apple MacBook Pro і нещодавно на комп’ютерах iMac. Через різницю в щільності пікселів між дисплеями "Retina" та "Non-Retina" деякі зображення, які не були створені з урахуванням дисплеїв з високою роздільністю, можуть виглядати "пікселізованими" при відтворенні на дисплеї з високою роздільністю.

Найпростіший спосіб зробити так, щоб ваші зображення правильно відтворювалися на дисплеях з високою роздільністю, таких як "retina display" MacBook Pro, — це визначити їхні `width` і `height` значення як лише половину від оригінального файлу. Ось приклад зображення, яке використовує лише половину оригінальної висоти та ширини:

```html
<style>
  img {
    height: 250px;
    width: 250px;
  }
</style>
<img src="coolPic500x500" alt="A most excellent picture" />
```

# --instructions--

Встановіть `width` і `height` тегу `img` на половину їхніх оригінальних значень. У цьому випадку і оригінальний `height`, і оригінальний `width` дорівнюють `200px`.

# --hints--

Ваш тег `img` повинен мати `width` 100 пікселів.

```js
assert.strictEqual(document.querySelector('img').width, 100);
```

Ваш тег `img` повинен мати `height` 100 пікселів.

```js
assert.strictEqual(document.querySelector('img').height, 100);
```

# --seed--

## --seed-contents--

```html
<style></style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```

# --solutions--

```html
<style>
  img {
    height: 100px;
    width: 100px;
  }
</style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```
