---
id: 587d78b1367417b2b2512b09
title: зробити зображення адаптивним
challengeType: 0
forumTopicId: 301140
dashedName: make-an-image-responsive
---

# --description--

Зробити зображення адаптивним за допомогою CSS насправді дуже просто. Вам потрібно лише додати ці властивості до зображення:

```css
img {
  max-width: 100%;
  height: auto;
}
```

`max-width` `100%` гарантуватиме, що зображення ніколи не буде ширшим за контейнер, у якому воно знаходиться, а `height` `auto` дозволить зображенню зберігати свій початковий співвідношення сторін.

# --instructions--

Додайте правила стилю до класу `responsive-img`, щоб зробити його адаптивним. Воно ніколи не повинно бути ширшим за свій контейнер (у цьому випадку це вікно попереднього перегляду) і має зберігати своє початкове співвідношення сторін. Після того, як ви додасте свій код, змініть розмір вікна попереднього перегляду, щоб побачити, як поводяться ваші зображення.

# --hints--

Ваш клас `responsive-img` має мати множину `max-width`, встановлену в `100%`.

```js
assert.strictEqual(getComputedStyle(document.querySelector('.responsive-img')).maxWidth, '100%');
```

Ваш клас `responsive-img` має мати множину `height`, встановлену в `auto`.

```js
assert.match(code, /height:\s*?auto;/g);
```

# --seed--

## --seed-contents--

```html
<style>
.responsive-img {


}

img {
  width: 600px;
}
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```

# --solutions--

```html
<style>
  .responsive-img {
    max-width: 100%;
    height: auto;
  }

  img {
    width: 600px;
  }
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```
