---
id: 6735a7370e0ae93a4577c689
title: 設計報紙版面
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

在本實驗中，你將使用 CSS Grid 設計報紙版面配置，包括像是 grid 列 和 grid 行 等概念。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。 

**使用者故事：**

1. 你的頁面應該包含一個帶有類別 `newspaper-layout` 的 `main` 元素。
1. `.newspaper-layout` 應該包含一個帶有類別 `title` 的 `header`，其中包含一個 `h1` 用來顯示報紙的名稱。
1. `.newspaper-layout` 應該包含一個帶有 `feature-article` 類別的 `article`，用於主要新聞文章。
1. `.feature-article` 應包含一個用於文章標題的 `h2` 元素，後面接著一個用於文章內容的 `p` 元素。
1. 你應該為較小的文章添加三個更多的 `article` 元素，並帶有類別 `small-article1`、`small-article2` 和 `small-article3`。
1. 每個較小的文章應包含一個用於文章標題的 `h3` 元素，後面接著一個用於文章內容的 `p` 元素。
1. `.newspaper-layout` 應該包含一個帶有類別 `secondary-article` 的 `article` 元素，用於額外的新聞區塊。
1. `.secondary-article` 應包含一個用於文章標題的 `h2` 元素，後面接著一個用於文章內容的 `p` 元素。
1. `.newspaper-layout` 應該包含一個帶有類別 `cover-image` 的 `figure`，用來顯示表述報紙內容的影像。
1. 具有 `title`、`feature-article`、`secondary-article`、`cover-image`、`small-article1`、`small-article2` 和 `small-article3` 類別的元素，其 `grid-area` 屬性應設定為相同的類別名稱。
1. 你的 `.newspaper-layout` 應該使用 CSS Grid 並搭配 `grid-template-areas` 屬性來定義區塊的排列方式：
   - `.title` 應該跨越頂部列。
   - `.feature-article` 應該在第二列跨越兩個行，並且 `.cover-image` 位於第三行。
   - `.secondary-article` 應該在第三列跨越兩個行，並且 `.cover-image` 位於第三行。
   - 這三篇小文章應該出現在第四列。
1. `.newspaper-layout` 應該有一個 `grid-template-columns` 屬性，將空間分成三個相等的行。
1. 你應該將 `.newspaper-layout` 的 `grid-template-rows` 屬性設定為第一列為 `auto`，並將剩餘空間平均分配給其他列。
1. 你應該在網格項目之間添加間隙。
1. 確保 `.cover-image` 內的影像符合指定空間，方法是將其 `max-width` 設為 `100%`。

**注意：** 請務必在你的 HTML 中聯結你的樣式表並套用你的 CSS。

# --hints--

你應該有一個 `main` 元素。

```js
assert.exists(document.querySelector('main'));
```

你的 `main` 元素應該有一個 `class` 為 `newspaper-layout`。

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

你應該在你的 `.newspaper-layout` 元素內有一個 `header` 元素，其 `class` 為 `title`。

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

你的 `header` 元素應該包含一個 `h1` 元素。

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

你的 `.title` 元素應該包含一些文字。

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

你應該為每篇文章使用一個 `article` 元素，並分別帶有類別 `feature-article`、`small-article1`、`small-article2`、`small-article3` 和 `secondary-article`。

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

`.feature-article` 應包含一個 `h2` 元素，後面接著一個 `p` 元素。

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`、`.small-article2` 和 `.small-article3` 應該包含一個 `h3` 元素，後面跟著一個 `p` 元素。

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

`.secondary-article` 應該包含一個 `h2` 元素，後面接著一個 `p` 元素。

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

你應該在你的 `.newspaper-layout` 中有一個帶有 `cover-image` 類別的 `figure` 元素。

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

你的 `.cover-image` 應該包含一個 `img` 元素。

```js
assert.exists(document.querySelector(".cover-image img"));
```

`.newspaper-layout` 元素的 `display` 屬性應該設為 `grid`。

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

你的 `.newspaper-layout` 元素應該有一個 `grid-template-areas` 屬性，並且 `title` 應該跨越第一列的所有三個行。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

網格模板的第二列應包含一個跨越兩個行的 `feature-article` 和一個 `cover-image`。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

網格模板的第三列應包含一個跨越兩個行的 `secondary-article` 和一個 `cover-image`。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

網格模板的第四列應包含 `small-article1`、`small-article2` 和 `small-article3`。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

你的 `.newspaper-layout` 選擇器應該使用 `grid-template-columns` 屬性將空間分成三個相等的部分。請記得你可以使用 `repeat` 函式（程式）或手動定義每個分數單位。

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

你的 `.newspaper-layout` 選擇器應該有一個 `grid-template-rows` 屬性，設定為 `auto 1fr 1fr 1fr`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

你應該在網格項目之間添加間隙。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` 應該有一個 `grid-area` 為 `title`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` 應該有一個 `grid-area` 為 `feature-article`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` 應該有一個 `grid-area` 為 `cover-image`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` 應該有一個 `grid-area` 為 `secondary-article`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` 應該有一個 `grid-area` 為 `small-article1`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` 應該有一個 `grid-area` 為 `small-article2`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` 應該有一個 `grid-area` 為 `small-article3`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

`.cover-image` 內的 `img` 應該有 `max-width` 為 `100%`。

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
