---
id: 61329d52e5010e08d9b9d66b
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Another important `meta` element for accessibility and SEO is the `description` definition. The value of the `content` attribute is used by search engines to provide a description of your page.

Adicione um elemento `meta` com o atributo `name` definido como `description`, e dê a ele um atributo `content`.

# --hints--

Você deve adicionar um novo elemento `meta` à `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Você deve dar ao elemento `meta` um atributo `name` de `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Você deve dar ao elemento `meta` um atributo `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

O valor do atributo `content` não deve ter mais de 165 caracteres. _Esse é o tamanho máximo da descrição para o Google._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
