---
id: 687d8c6f066877768f7ee441
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

داخل صنف `LinkedList` الخاص بك، استبدل الكلمة المفتاحية `pass` بصنف داخلي يُسمى `Node`. هذا سيمثل العقد الفردية التي تخزن البيانات في القائمة.

# --hints--

يجب أن تنشئ صنف `Node` داخل صنف `LinkedList`. تذكر إضافة الكلمة المفتاحية `pass` داخل صنف `Node`.

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    pass
--fcc-editable-region--
```
