---
id: 687d7deea6a2a22fa8816a73
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في الدروس السابقة، تعلمت عن بعض هياكل البيانات الشائعة. في هذه الورشة، ستقوم بتنفيذ قائمة مرتبطة.

أولًا، أنشئ صنفًا باسم `LinkedList` لتمثيل الهيكل الذي سيحتوي على عقد القائمة. أضف الكلمة المفتاحية `pass` داخل الصنف لكي ينجح الاختبار.

# --hints--

يجب أن تنشئ صنفًا باسم `LinkedList`. تذكر إضافة الكلمة المفتاحية `pass`.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
