---
id: 687d8e05a529a4838f1c961c
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

بعد ذلك، احذف الكلمة المفتاحية `pass` داخل صنف `Node`، ثم أضف طريقة `__init__` إلى صنف `Node` بمعلمتين، `self` و `element`.

داخل الطريقة، عيّن `element` إلى `self.element`. هذا يخزن القيمة التي تُمرر عند إنشاء العقدة.

# --hints--

يجب أن تعرف طريقة `__init__` داخل صنف `Node` الخاص بك.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

يجب أن تحتوي طريقة `__init__` على معلمتين، `self` و `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

يجب أن تعيّن `element` إلى `self.element` داخل طريقة `__init__`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        pass
--fcc-editable-region--
```
