---
id: 687e366b3004a9831e4cbd41
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

أضف طريقة تسمى `is_empty` تُرجع `self.length == 0`.

ستستخدم هذا لتتبع عدد العقد في القائمة. إذًا، عندما تكون `self.length == 0`، تُرجع `True`، مما يعني أن القائمة لا تحتوي على عقد.

إذا كانت قيمة `self.length` أكبر من `0`، تُرجع القيمة `False`، مما يعني أن القائمة تحتوي على عقدة واحدة على الأقل.

# --hints--

يجب أن تنشئ طريقة `is_empty` تُرجع `self.length == 0`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
