---
id: 687e344cc1993f7316af030c
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

داخل صنف `LinkedList`، أضف طريقة `__init__` التي تُعيّن `self.length` إلى `0`. هذا يُهيئ عدادًا لتتبع عدد العقد (العناصر) الموجودة حاليًا في القائمة.

# --hints--

يجب أن تُعيّن `self.length` إلى `0` داخل طريقة `__init__` في صنف `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
