---
id: 68800a64b06e80d9a31c131b
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

لا تزال في طريقة `__init__` من صنف `LinkedList`، عرّف `self.head` على أنه `None`. هذا يعمل كمرجع للعقدة الأولى في القائمة. القائمة فارغة في الوقت الحالي، لذا يجب تعيينها إلى `None`.


# --hints--

يجب تعيين `self.head` إلى `None` داخل طريقة `__init__` من صنف `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

    def __init__(self):
        self.length = 0
--fcc-editable-region--
        
--fcc-editable-region--
```
