---
id: 688003f589b875af60273dd9
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

لا تزال داخل الطريقة `__init__`، عيّن القيمة `None` إلى `self.next`. هذا يعيّن مؤشر `next` إلى `None`، مما يعني أن العقدة الجديدة لا تشير بعد إلى أي عقدة أخرى في القائمة.

لاحقًا، يمكن تحديث `next` ليشير إلى عقدة أخرى لتشكيل سلسلة من قائمة مرتبطة.

# --hints--

داخل الطريقة `__init__`، يجب تهيئة الخاصية `next` إلى `None`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
            
--fcc-editable-region--
```
