---
id: 68857e0a49baf4751f737a2a
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

الآن، بعد طريقة `is_empty`، أنشئ طريقة تسمى `add` داخل صنف `LinkedList`. يجب أن تحتوي الطريقة الجديدة على معلمتين، `self` و `element`.

داخل طريقتك الجديدة، أنشئ متغيرًا اسمه `node` وقم بتعيين نسخة جديدة من `Node` له باستخدام `self.Node(element)`. هذا سينشئ عقدة جديدة تُضاف إلى القائمة المرتبطة.

# --hints--

يجب أن تنشئ طريقة `add` داخل صنف `LinkedList`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

يجب أن تحتوي طريقة `add` على معلمتين، `self` و `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

يجب أن تعيّن `self.Node(element)` إلى `node` داخل طريقة `add`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
