---
id: 688649c8673f842f49228a5e
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

إذا لم تكن القائمة فارغة، يجب أن تجد الطريقة العقدة الأخيرة وتربط العقدة الجديدة بمؤشر `next` الخاص بها.

للقيام بذلك، أنشئ كتلة `else` بعد كتلة `if`. داخل كتلة `else`، أنشئ متغيرًا باسم `current_node` وقم بتعيين `self.head` له.

# --hints--

يجب تعيين `self.head` إلى `current_node` داخل كتلة `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
