---
id: 68867c3898b09115d8c5140f
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

تُمكّنك حلقة `while` التي أنشأتها للتو من المرور بكل عقدة في القائمة حتى تصل إلى العقدة الأخيرة. بمجرد حدوث ذلك، يمكنك إضافة العقدة إلى القائمة.

في كتلة `else`، مباشرة بعد كتلة `while`، عيّن `node` إلى `current_node.next`.

# --hints--

يجب تعيين `node` إلى `current_node.next` في كتلة `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].has_stmt("current_node.next = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
--fcc-editable-region--
            
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
