---
id: 6888fe6c745b6852ed240229
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

بعد ذلك، ستتكرر عبر القائمة المرتبطة لتجد العقدة التي تحتوي على العنصر المحدد لإزالته.

أنشئ حلقة `while` تتحقق من أن `current_node` ليست `None` وأن `current_node.element` لا يساوي `element`. داخل حلقة الـ `while`، قم بتعيين `current_node` إلى `previous_node`.

# --hints--

يجب أن تحتوي طريقة `remove` على حلقة `while`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
`)})
```

يجب أن تتحقق حلقة `while` من أن `current_node` ليست `None` وأن `current_node.element` لا يساوي `element`.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0].find_conditions()[0]
  assert _cond.is_equivalent("current_node is not None and current_node.element != element")
`)})
```

يجب تعيين `current_node` إلى `previous_node` داخل حلقة الـ `while`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
