---
id: 6888f9827637322d4f309192
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

داخل الطريقة `remove`، أنشئ متغيرًا باسم `current_node` وقم بتعيين قيمة `self.head` له.

هذا يبدأ التنقل في القائمة المرتبطة من العقدة الأولى. سيُستخدم `current_node` لتتبع العقدة التي تُفحص حاليًا أثناء البحث عن العقدة التي تحتوي على العنصر المحدد لإزالته.

# --hints--

يجب تعيين `self.head` إلى `current_node` داخل الطريقة `remove`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").has_stmt("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
