---
id: 68b84df2648b4c96b25022f3
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

حان الوقت لاختبار قائمة مرتبطة حتى الآن! في أسفل شفرتك، أنشئ مثيلًا من `LinkedList` باسم `my_list`. ثم، باستخدام الطريقة `is_empty` التي أنشأتها في الخطوة السابقة، اطبع `my_list.is_empty()`.

# --hints--

يجب أن تنشئ مثيلًا من `LinkedList` باسم `my_list`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

يجب أن تطبع `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
