---
id: 68b9e951e4892721632acc26
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

رائع! حان الوقت الآن لاختبار إضافة عناصر إلى القائمة المرتبطة. في أسفل شفرتك، بعد إنشاء نسخة من `LinkedList` باسم `my_list` وطباعة `my_list.is_empty()`، أضف `1` و `2` إلى القائمة باستخدام الطريقة `add`.

ثم، اطبع `my_list.is_empty()` مرة أخرى لتتأكد أن القائمة لم تعد فارغة. وأخيرًا، اطبع `my_list.length` لرؤية الطول الحالي للقائمة المرتبطة.

# --hints--

يجب أن تضيف `1` إلى القائمة باستخدام الطريقة `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

يجب أن تضيف `2` إلى القائمة باستخدام الطريقة `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

يجب أن تطبع `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

يجب أن تطبع `my_list.length`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
