---
id: 68c0da2df8ce865dcf7769a7
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

لا تزال داخل كتلة `else`، أنشئ حلقة `while` تتحقق من أن `current_node.next` ليست `None`.

داخل حلقة `while`، قم بتعيين `current_node.next` إلى `current_node`.

# --hints--

يجب أن تنشئ حلقة `while` تتحقق مما إذا كانت `current_node.next` ليست `None`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("current_node.next is not None")
`) })
```

يجب أن تعيّن `current_node.next` إلى `current_node` داخل حلقة `while`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0]
assert _cond.is_equivalent("while current_node.next is not None:\\n current_node = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
--fcc-editable-region--
            
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
