---
id: 579e2a2c335b9d72dd32e05c
title: スライスとスプライスのアルゴリズムを実装する
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

このラボでは、2つの配列をマージするアルゴリズムを作成する必要があります。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. 2つの配列とインデックスを受け取る`frankenSplice`関数を作成してください。
1. 最初の配列の各要素を、指定されたインデックスから始めて順番に2番目の配列にコピーし、結果の配列を返してください。
1. 関数実行後も入力の配列は変わらないようにしてください。

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)`は`[4, 1, 2, 3, 5]`を返すべきです。

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)`は`["a", 1, 2, "b"]`を返すべきです。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)`は`["head", "shoulders", "claw", "tentacle", "knees", "toes"]`を返すべきです。

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

最初の配列のすべての要素は元の順序のまま2番目の配列に追加されるべきです。`frankenSplice([1, 2, 3, 4], [], 0)`は`[1, 2, 3, 4]`を返すべきです。

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

関数実行後、最初の配列は変わらないようにしてください。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

関数実行後、2番目の配列は変わらないようにしてください。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
