---
id: 587d8259367417b2b2512c86
title: Insertion Sort implementieren
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

In diesem Labor implementieren Sie den Insertion-Sort-Algorithmus. Diese Methode funktioniert, indem sie am Anfang der Liste ein sortiertes Array aufbaut. Sie beginnt das sortierte Array mit dem ersten Element. Dann betrachtet sie das nächste Element und tauscht es rückwärts in das sortierte Array, bis es an der richtigen Stelle einsortiert ist. Sie fährt fort, die Liste zu durchlaufen und neue Elemente rückwärts in den sortierten Abschnitt zu tauschen, bis sie das Ende erreicht. Dieser Algorithmus hat im Durchschnitts- und im schlimmsten Fall eine quadratische Zeitkomplexität.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `insertionSort`-Funktion erstellen.
1. Die `insertionSort`-Funktion sollte ein Array von ganzen Zahlen entgegennehmen und ein Array mit denselben Zahlen zurückgeben, das von klein nach groß sortiert ist.

# --hints--

`insertionSort` sollte eine Funktion sein.

```js
assert.isFunction(insertionSort);
```

`insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` sollte ein Array zurückgeben, das bis auf die Reihenfolge unverändert ist.

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort` sollte ein sortiertes Array zurückgeben (von klein nach groß).

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`insertionSort([5, 4, 33, 2, 8])` sollte `[2, 4, 5, 8, 33]` zurückgeben.

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` sollte die eingebaute `.sort()`-Methode nicht verwenden.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
