---
id: 6140c7e645d8e905819f1dd4
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Щоб почати, додайте елемент `link` для файлу `./styles.css`.

# --hints--

Ваш код має містити елемент `link`.

```js
assert.match(code, /<link/);
```

У вас має бути один елемент `link`.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Ваш елемент `link` має бути всередині елемента `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Ваш елемент `link` має мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ваш елемент `link` має мати атрибут `href` зі значенням `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
