---
id: 5d822fd413a79914d39e98cd
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

CSS では、アスタリスク (*) ですべての要素を選択できます。 すべての要素に境界線を追加するため、`*` セレクターを使用して `border` を `1px solid black` に設定してください。 これは要素の位置とサイズを視覚化するのに役立つ技です。 この設定は後ほど削除します。

# --hints--

You should use the `*` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

You should use the `border` property to style all the elements.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

All elements should have a `1px solid black` border.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
