---
id: 5d822fd413a79914d39e98ce
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

更に、すべての要素の `box-sizing` を `border-box` に設定してください。 この設定により、境界線を追加しても要素のサイズが増加しなくなります。

# --hints--

You should use the `box-sizing` property.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

You should make use of the existing `*` selector.

```js
// Two selectors create two CSSStyleRule objects
assert.lengthOf(new __helpers.CSSHelp(document).getStyleDeclarations('*'), 1);
```

All elements should have a `box-sizing` of `border-box`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```

