---
id: 5d822fd413a79914d39e98d1
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`.background-buildings` の要素について `width` と `height` を `100%` に設定し、親要素である `body` の幅および高さ全体に広がるようにしましょう。

# --hints--

You should use the `background-buildings` class to select the correct element.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

Your `.background-buildings` element should have a `width` of `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

Your `.background-buildings` element should have a `height` of `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--

```

