---
id: a5deed1811a43193f9f1c841
title: تنفيذ متخطٍ للعناصر
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

في هذا المختبر ستنشئ دالة تتخطى العناصر في مصفوفة حتى تجد عنصرًا مقبولًا بناءً على دالة اختبار محددة.

على سبيل المثال، لمصفوفة مثل `[1, 1, 1, 2, 1, 1, 1]` ودالة اختبار `function(n) {return n === 2}`، أول عنصر مقبول هو العنصر عند الفهرس 3، لذا يجب تجاهل كل العناصر التي قبله، ويجب أن يكون الناتج هو العناصر المتبقية `[2, 1, 1, 1]`.

**الهدف**: حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة `dropElements` تقبل مصفوفة (`arr`) ودالة (`func`) كمعلمات.
1. يجب أن تقوم دالة `dropElements` بالتكرار خلال المصفوفة وإزالة العناصر بدءًا من الأول حتى تُرجع `func` القيمة `true` لعنصر ما.
1. يجب أن تُرجع دالة `dropElements` العناصر المتبقية في المصفوفة إذا تحقق الشرط.
1. إذا لم يتحقق الشرط أبدًا، يجب أن تُرجع مصفوفة فارغة.

# --hints--

يجب أن تملك دالة `dropElements`.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` يجب أن تُرجع `[3, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` يجب أن تُرجع `[1, 0, 1]`.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` يجب أن تُرجع `[1, 2, 3]`.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` يجب أن تُرجع `[]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` يجب أن تُرجع `[7, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` يجب أن تُرجع `[3, 9, 2]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
