---
id: 579e2a2c335b9d72dd32e05c
title: تنفيذ خوارزمية القطع والدمج
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

في هذا المختبر، ستحتاج إلى إنشاء خوارزمية لدمج مصفوفتين.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. أنشئ دالة `frankenSplice` تقبل مصفوفتين وفهرسًا.
1. انسخ كل عنصر من المصفوفة الأولى إلى المصفوفة الثانية، بترتيبها، بدءًا من الفهرس المعطى، وأرجع المصفوفة الناتجة.
1. يجب أن تبقى المصفوفات المدخلة كما هي بعد تشغيل الدالة.

# --hints--

يجب أن تُرجع `frankenSplice([1, 2, 3], [4, 5], 1)` القيمة `[4, 1, 2, 3, 5]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

يجب أن تُرجع `frankenSplice([1, 2], ["a", "b"], 1)` القيمة `["a", 1, 2, "b"]`.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

يجب أن تُرجع `frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` القيمة `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

يجب إضافة كل العناصر من المصفوفة الأولى إلى المصفوفة الثانية بترتيبها الأصلي. يجب أن تُرجع `frankenSplice([1, 2, 3, 4], [], 0)` القيمة `[1, 2, 3, 4]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

يجب أن تبقى المصفوفة الأولى كما هي بعد تشغيل الدالة.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

يجب أن تبقى المصفوفة الثانية كما هي بعد تشغيل الدالة.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
