---
id: 6421f98f4999d1179ce37cb4
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

تحتوي المصفوفات على طريقة `some()`. مثل طريقة `filter()`، تقبل `some()` دالة استدعاء راجع يجب أن تأخذ عنصرًا من المصفوفة كمعلمة. سترجع طريقة `some()` القيمة `true` إذا أعادت دالة الاستدعاء الراجع القيمة `true` لعنصر واحد على الأقل في المصفوفة.

إليك مثال على استدعاء طريقة `.some()` للتحقق مما إذا كان أي عنصر في المصفوفة حرفًا كبيرًا.

```js
const arr = ["A", "b", "C"];
arr.some(letter => letter === letter.toUpperCase());
```

حدّث دالة `isSpam()` لاستخدام طريقة `some()`، التي ستتحقق مما إذا كان اختبار `msg` مقابل أي من التعبيرات النمطية في `denyList` يرجع `true`.

# --hints--

يجب أن تستخدم دالة `isSpam` طريقة `.some()`.

```js
let flag = false;
const temp = Array.prototype.some;
Array.prototype.some = (arg) => {flag = true};
try {
  isSpam("");
  assert.isTrue(flag);
} finally {
  Array.prototype.some = temp;
}
```

يجب أن تستدعي دالة `isSpam` `denyList.some()`.

```js
let flag = false;
const temp = denyList.some;
denyList.some = (arg) => {flag = true};
try {
  isSpam("");
  assert.isTrue(flag);
} finally {
  denyList.some = temp;
} 
```

يجب أن تُرجع دالة `isSpam` نتيجة `denyList.some()`، حيث يجب أن تُقيّم دالة الاستدعاء الراجع الخاصة بها `.test(msg)` على كل عنصر من عناصر `denyList`.

```js
const help = "plEase hElp";
const assist = "AssIst ME";
const notSpam = "not spam";
assert.match(
  isSpam.toString(),
  /denyList\.some\(\s*\(?\s*(\w+)\s*\)?\s*=>\s*\1\.test\(\s*msg\s*\)\s*\)/
);
assert.isTrue(isSpam(help));
assert.isTrue(isSpam(assist));
assert.isFalse(isSpam(notSpam));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result-message {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;

const denyList = [helpRegex];

--fcc-editable-region--
const isSpam = (msg) => helpRegex.test(msg);
--fcc-editable-region--

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
