---
id: 63b607af6fcdb119aae9b16a
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

En el formulario, los usuarios deben tener la capacidad para agregar varios tipos de comidas junto con sus conteos de calorías.

Crea un elemento `fieldset` con el `id` establecido a `breakfast`.

Dentro de ese elemento, crea un `legend` con el texto `Breakfast`, y un `div` vacío con la `clase` establecida a `input-container`.

# --hints--

Debes crear un elemento `fieldset` en tu `form`.

```js
assert.exists(document.querySelector('form fieldset'));
```

Tu elemento `fieldset` debe venir después de tu elemento `input`.

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

Tu elemento `fieldset` debe tener un `id` establecido a `breakfast`.

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

Debes crear un elemento `legend` dentro de tu `fieldset`.

```js
assert.exists(document.querySelector('form fieldset legend'));
```

Tu elemento `legend` debe tener el texto `Breakfast`.

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

Debes crear un elemento `div` dentro de tu `fieldset`.

```js
assert.exists(document.querySelector('form fieldset div'));
```

Tu elemento `div` debe venir después de tu elemento `legend`.

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

Tu elemento `div` debe tener una `clase` establecida a `input-container`.

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
