---
id: 63b60aaaa65f8922bfce6b7e
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

En tu elemento `span`, crea un elemento `label` para un `entry-dropdown` y dale el texto `Agregar comida o ejercicio:`. Luego crea un elemento `select` con el `id` establecido en `entry-dropdown` y un `name` establecido en `options`. Debajo de eso, agrega un elemento `button` con el `id` establecido en `add-entry` y el texto `Agregar Entrada`.

Dale a tu elemento `button` un atributo `type` configurado como `button` para evitar el envío automático del formulario.

# --hints--

Deberías agregar un elemento `label` a tu elemento `span`.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

Tu nuevo elemento `label` debería tener un atributo `for` establecido en `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

Tu nuevo elemento `label` debería tener el texto `Agregar comida o ejercicio:`.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

Deberías agregar un elemento `select` a tu elemento `span`.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

Tu elemento `select` debería venir después de tu elemento `label`.

```js
assert.equal(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName, 'LABEL');
```

Tu nuevo elemento `select` debería tener un atributo `id` establecido en `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

Tu nuevo elemento `select` debería tener un atributo `name` establecido en `options`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

Tu elemento `select` debería estar vacío y tener una etiqueta de cierre `</select>`.

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

Deberías agregar un elemento `button` a tu elemento `span`.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

Tu elemento `button` debería venir después de tu elemento `select`.

```js
assert.equal(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName, 'SELECT');
```

Tu nuevo elemento `button` debería tener un atributo `id` establecido en `add-entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

Tu nuevo elemento `button` debería tener el texto `Agregar Entrada`.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
