---
id: 68e52994bc7ea2e3dec530a0
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

Agora você precisa de uma maneira de gerenciar múltiplas estratégias de desconto e encontrar o melhor preço. No topo do seu arquivo, após a importação existente, adicione uma importação para o tipo `List` do módulo `typing`:

```py
from typing import List
```

Isso permite que você especifique que um parâmetro de função é uma lista contendo tipos específicos de objetos.

# --hints--

Você deve importar `List` do módulo `typing` conforme mencionado nas instruções após as importações existentes.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[1].is_equivalent("from typing import List")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
