---
id: 68e5293bd00d2fe134f58981
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Crie uma classe chamada `DiscountStrategy` que herda de `ABC`. Esta será uma classe base abstrata que define a interface para todas as estratégias de desconto.

Por enquanto, apenas use `pass` no corpo da classe para criar uma classe vazia.

# --hints--

Você deve criar uma classe chamada `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

Sua classe `DiscountStrategy` deve herdar de `ABC`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
