---
id: 68e529226b0defdff26c4a89
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Neste momento, se você criar uma instância de `Product` e imprimi-la, verá uma representação padrão do objeto que não é muito informativa. Por exemplo, pode ser algo como `<__main__.Product object at 0x7f9b8c2d1d30>`.

Para converter a representação padrão do objeto em uma string legível, crie um método `__str__` na sua classe `Product` que retorne uma string com o formato `name - $price` (onde `name` e `price` devem ser substituídos pelos valores dos atributos correspondentes).

# --hints--

Sua classe `Product` deve ter um método `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

Seu método `__str__` deve retornar uma string no formato `name - $price`.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
