---
id: 68e529226b0defdff26c4a8b
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Agora você começará a construir a interface do sistema de desconto. Python fornece uma maneira de definir classes abstratas que servem como modelos para outras classes. Isso é feito usando a classe `ABC` (Abstract Base Class) do módulo `abc`.

No início do seu código, adicione a seguinte declaração de importação:

```py
from abc import ABC, abstractmethod
```

Isto importa a classe `ABC` e o decorador `abstractmethod`, que você usará para definir métodos que devem ser implementados por qualquer classe que herde da sua classe abstrata.

# --hints--

Você deve ter a declaração de importação mencionada no topo do seu código.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
