---
id: 68e5293bd00d2fe134f58983
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Agora adicione um segundo método abstrato chamado `apply_discount` à sua classe `DiscountStrategy`. Use o decorador `@abstractmethod` e dê a ele os parâmetros `self` e `product`.

Adicione dicas de tipo: `product` deve ser do tipo `Product` e o método deve retornar um `float`. Use `pass` no corpo do método.

# --hints--

Sua classe `DiscountStrategy` deve ter um método chamado `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

Seu método `apply_discount` deve ter o decorador `@abstractmethod`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

Seu método `apply_discount` deve ter os parâmetros `self` e `product: Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

Seu método `apply_discount` deve ter uma dica de tipo de retorno `-> float`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
