---
id: 68e5293bd00d2fe134f58985
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Como `PercentageDiscount` herda de uma classe abstrata, você deve implementar todos os métodos abstratos. Comece implementando o método `is_applicable`. O método `is_applicable` deve receber `self`, `product: Product` e `user_tier: str` como parâmetros. Ele deve ter uma dica de tipo de retorno `bool`.

Este método deve determinar se um desconto percentual é válido. Retorne `True` se `self.percent` for menor ou igual a `70`, caso contrário retorne `False`. Isso previne descontos irreais como 90% de desconto.

# --hints--

Sua classe `PercentageDiscount` deve ter um método `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

Seu método `is_applicable` deve ter os parâmetros `self`, `product: Product` e `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

Seu método `is_applicable` deve retornar `True` quando `self.percent` for menor ou igual a `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

Seu método `is_applicable` deve retornar `False` quando `self.percent` for maior que `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

Seu método `is_applicable` deve ter uma dica de tipo de retorno `bool`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
