---
id: 68e529226b0defdff26c4a8a
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Agora, crie uma instância de `Product` com o nome `Wireless Mouse` e um preço de `50.0`. Atribua-a a uma variável chamada `product`. Em seguida, imprima `product` no console para ver como ele se parece.

# --hints--

Você deve ter uma variável chamada `product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

Você deve atribuir uma instância de `Product` à variável `product`.

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

Sua instância `Product` deve ter o nome `Wireless Mouse` e o preço `50.0`.

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

Você deve imprimir `product` no console.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

Quando você imprime um `Product` com nome `'Wireless Mouse'` e preço `50.0`, ele deve exibir `Wireless Mouse - $50.0`.

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
