---
id: 68e5293bd00d2fe134f58984
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Agora é hora de criar uma implementação concreta da sua classe abstrata. Crie uma classe chamada `PercentageDiscount` que herda de `DiscountStrategy`.

Dê a ele um método `__init__` que receba `self` e `percent` como parâmetros. Dê uma dica de tipo `int` para `percent` e defina o tipo de retorno do método como `None`. Dentro do método `__init__`, armazene `percent` como um atributo da instância.

# --hints--

Você deve criar uma classe chamada `PercentageDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

Sua classe `PercentageDiscount` deve herdar de `DiscountStrategy`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

Sua classe `PercentageDiscount` deve ter um método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

Seu método `__init__` deve ter os parâmetros `self` e `percent: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

Seu método `__init__` deve ter uma dica de tipo de retorno `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

Você deve atribuir `percent` a `self.percent` dentro do seu método `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
