---
id: 68e524c94e2e4ab7b4bfc7e2
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você vai construir um calculador de desconto que pode aplicar diferentes estratégias de desconto a produtos. O sistema vai determinar o melhor preço para um cliente com base em múltiplas regras de desconto.

Comece definindo uma classe chamada `Product`. Dê à sua classe um método `__init__` com os seguintes parâmetros: `self`, `name` e `price`. Dentro do método `__init__`, atribua os parâmetros `name` e `price` aos atributos da instância com o mesmo nome.

# --hints--

Você deve criar uma classe chamada `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

Sua classe `Product` deve ter um método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

Seu método `__init__` deve ter os seguintes parâmetros: `self`, `name` e `price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

Você deve atribuir `name` a `self.name` dentro do seu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

Você deve atribuir `price` a `self.price` dentro do seu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
