---
id: 68e79d91daf8fb4808f93c47
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Após implementar o método `apply_discount`, teste sua estratégia de desconto. Crie uma instância de `PercentageDiscount` com `10` por cento e atribua-a a uma variável chamada `discount`. Em seguida, aplique o desconto usando `discount.apply_discount(product)` e imprima o resultado para ver o desconto em ação.

# --hints--

Você deve criar uma variável chamada `discount` com uma instância de `PercentageDiscount`.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

Sua instância `PercentageDiscount` deve ser inicializada com `10` por cento.

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

Você deve chamar `discount.apply_discount(product)` e imprimir o resultado.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
