---
id: 68e9207ddcba053c5cb24413
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Em Python, as dicas de tipo são sinais opcionais que indicam a outros desenvolvedores qual é o tipo de dado esperado de uma variável ou função. Aqui está um exemplo que adiciona a dica de tipo `str` ao parâmetro `name`.

```py
def greet(name: str):
    return 'Hello, ' + name
```

Para os parâmetros do método `__init__`, atualize o parâmetro `name` para ter a dica de tipo `str` e o parâmetro `price` para ter a dica de tipo `float`.

# --hints--

Você deve adicionar a dica de tipo `str` ao parâmetro `name` e a dica de tipo `float` ao parâmetro `price` no método `__init__` da classe `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
