---
id: 68f119ad905b29ff59d858c8
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Assim como na etapa anterior, no seu método `__str__`, adicione uma dica de tipo de retorno `str` já que este método retorna uma representação em string do objeto.

# --hints--

Seu método `__str__` deve ter uma dica de tipo de retorno `str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
