---
id: 68ce56d54d3cad1149a1c4cf
title: 设计一个特征选择页面
challengeType: 25
dashedName: design-a-feature-selection-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你的页面应有一个 `h1` 元素，文本为 `Feature Selection`。
2. 你的页面应有一个带有类名为 `feature-card-container` 的 `div` 元素。
3. 你的页面应至少有两个 `label` 元素，每个元素的类为 `feature-card`，且都在 `div` 内。
4. 你的每个 `label` 元素都应包含标签文本和一个类型为 `checkbox` 的 `input` 元素。
5. 创建一个选择器以定位复选框，并应用以下样式：
    - 你所有的复选框元素都应在你的 CSS 中设置为 `appearance: none;`。
    - 你所有的复选框元素都应该有你选择的边框宽度、颜色和样式。
6. 当复选框被选中时，应在复选框中心显示一个勾选标记 `✓`。
7. 当复选框被选中时，复选框的背景颜色应更改为你选择的颜色。
8. 当复选框被选中时，复选框的边框颜色应更改为你选择的颜色。

**注意：** 请确保在你的超文本标记语言中链接你的 CSS 文件。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 应该有 `Feature Selection` 的文本。

```js
assert.equal(document.querySelector('h1').innerText.trim(), `Feature Selection`);
```

你应该有一个带有类 `feature-card-container` 的 `div` 元素。

```js
assert.exists(document.querySelector('div.feature-card-container'));
```

你的 `.feature-card-container` 内应至少有两个 `label` 元素。

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isAtLeast(labels.length, 2);
```

`.feature-card-container` 内的每个 `label` 元素都应具有 `feature-card` 类。

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isTrue(label.classList.contains('feature-card')));
```

在每个 `label` 元素内你应该有标签文本。

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isNotEmpty(label.innerText.trim()));
```

每个 label 元素内应包含一个类型为 `checkbox` 的 `input` 元素。

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.exists(label.querySelector('input[type="checkbox"]')));
```

你应该有一个复选框的选择器。

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你的复选框应设置为 `appearance: none`。

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.equal(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('appearance'), 'none');
```

你的复选框应具有你选择的边框宽度。

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-width'));
```

你的复选框应具有你选择的边框颜色。

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-color'));
```

你的复选框应具有你选择的边框样式。

```js
const selectors = [  
  '.feature-card [type="checkbox"]',
  '.feature-card input[type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input', // there aren't other input elements, so this is a valid approach
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-style'));
```

当复选框被选中时，复选框中央应有一个勾选标记 `✓`。

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;

    checkbox.checked = true;
    const after = window.getComputedStyle(checkbox, ':after');
    assert.equal(after.content, '"✓"');
});
```

当复选框被选中时，复选框的背景颜色应更改为你选择的颜色。

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBG = window.getComputedStyle(checkbox).backgroundColor;
    checkbox.checked = true;
    const checkedBG = window.getComputedStyle(checkbox).backgroundColor;

    assert.notEqual(checkedBG, uncheckedBG);
});
```

当复选框被选中时，复选框的边框颜色应更改为你选择的颜色。

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBC = window.getComputedStyle(checkbox).borderColor;
    checkbox.checked = true;
    const checkedBC = window.getComputedStyle(checkbox).borderColor;

    assert.notEqual(checkedBC, uncheckedBC);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Selection Feature Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Feature Selection</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Feature Selection</h1>
    <div class="feature-card-container">
      <label class="feature-card">
        <input type="checkbox" />
        <h4>Cloud Storage</h4>
        <p>100 Gigabyte secure storage</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Dedicated Support</h4>
        <p>24/7 customer help</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Advanced Analytics</h4>
        <p>Insights & reports</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Custom User Themes</h4>
        <p>Personalized dashboard design</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Multi-User Collab</h4>
        <p>Team access and sharing</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>API Access</h4>
        <p>Integrate with your custom built tools</p>
      </label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  flex-direction: column;
  height: 100vh;
}

.feature-card-container {
  display: flex;
  flex-wrap: wrap;
  gap: 20px;
  max-width: 600px;
  justify-content: center;
}

.feature-card-container > * {
  flex: 1 1 180px;
}

.feature-card {
  border: 2px solid #ccc;
  border-radius: 12px;
  padding: 20px;
  text-align: center;
  position: relative;
  cursor: pointer;
}

.feature-card:hover {
  border-color: #f1be32;
  box-shadow: 0 4px 12px rgba(255, 201, 53, 0.6);
}

.feature-card input[type='checkbox'] {
  position: absolute;
  top: 5px;
  right: 5px;
  width: 24px;
  height: 24px;
  cursor: pointer;
  appearance: none;
  -webkit-appearance: none;
  border: 2px solid #ccc;
  border-radius: 4px;
  background-color: white;
}

.feature-card input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #f1be32;
}

.feature-card input[type='checkbox']:checked::after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  line-height: 20px;
}

.feature-card h4 {
  margin: 10px 0 5px;
  font-size: 1.2rem;
}

@media (max-width: 600px) {
  body {
    height: auto;
  }
}
```
