---
id: 5d5a813321b9e3db6c106a46
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

O JavaScript é uma linguagem poderosa, que permite que você construa sites que são interativos. 

*Note*: For all remaining projects in this curriculum, you will need a basic level of knowledge in HTML and CSS. If you are new to HTML and CSS, please go through the [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

Para começar, crie seu boilerplate HTML padrão com um `DOCTYPE`, `html`, `head` e `body`. Em seguida, adicione uma tag `meta` para o `charset`.Adicione um elemento `title` e use o texto `RPG - Dragon Repeller` a ele.Inclua uma tag `link` para sua folha de estilos e vincule-a ao arquivo `styles.css`.

Por fim, crie um elemento `div` com o `id` definido como `game` dentro de `body`.

# --hints--

O código deve conter a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Você deve incluir um espaço após a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Você deve definir o tipo de documento para ser `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Você deve fechar a declaração `DOCTYPE` com um `>` após o tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

O elemento `html` deve ter uma tag de abertura.Não se esqueça do atributo `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert(code.match(/<\/html\s*>/));
```

A declaração `DOCTYPE` deve estar no ínicio do HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Você deve acrescentar uma tag de abertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Você deve ter uma tag de fechamento para `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Você deve ter uma tag de abertura para `body`.

```js
assert(code.match(/<body\s*>/i));
```

Você deve ter uma tag de fechamento para `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Os elementos `head` e `body` devem ser irmãos.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

O elemento `head` deve estar dentro do elemento `html`.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

O elemento `body` deve estar dentro do elemento `html`.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

Seu código deve ter um elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

O elemento `meta` deve ter o atributo `charset` com o valor `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

O código deve ter um elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Você deve acrescentar uma tag de fechamento para `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

O código deve ter um elemento `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Você deve ter um elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

O elemento `div` deve ter um atributo `id` com o valor `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Seu elemento `div` deve estar dentro do elemento `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
