---
id: 62a3a0a3c0a4b32915d26a6e
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Remova a linha `console.log("Hello World");` para começar a escrever o código do projeto.

No projeto anterior, você aprendeu como trabalhar com variáveis e a palavra-chave `let` assim:

```js
let camperbot;
```

Você também aprendeu como inicializar uma variável com um valor assim:

```js
let age = 32;
```

Use a palavra-chave `let` para declarar uma variável chamada `xp` e atribuir a ela o valor `0`, um número.

# --hints--

Você não deve ter uma linha `console.log("Hello World");` no código.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Você deve usar a palavra-chave `let` para declarar a variável.

```js
assert.match(code, /let/);
```

Você deve declarar uma variável chamada `xp`.

```js
assert.match(code, /let\s+xp/);
```

`xp` deve ter um valor de `0`.

```js
assert.equal(xp, 0);
```

Você deve inicializar a variável `xp` como `0`. 

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
