---
id: 62a3a7e4f1060e2fc5ffb34b
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Crie outra variável chamada `currentWeaponIndex` e defina-a como `0`.

# --hints--

Você deve usar `let` para declarar uma variável chamada `currentWeaponIndex`.

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

Você deve usar CamelCase para nomear a variável.

```js
assert.match(code, /currentWeaponIndex/);
```

Sua variável `currentWeaponIndex` deve ser definida como `0`.

```js
assert.equal(currentWeaponIndex, 0);
```

Você deve inicializar a variável como `0`.

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
