---
id: 62a3b41c9494f937560640ab
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Declare mais duas variáveis chamadas `monsterHealth` e `inventory`.

Para sua variável `inventory`, atribua a ela o valor de um array contendo a string `"stick"`.

Lembre-se de que você trabalhou com arrays no projeto anterior assim:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

Você deve usar `let` para declarar a variável `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Você deve usar camelCase para nomear a variável `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` não deve ter um valor.

```js
assert.isUndefined(monsterHealth);
```

Você não deve atribuir um valor à variável `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

Você deve usar `let` para declarar a variável `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

A variável `inventory` ainda deve ser um array.

```js
assert.isArray(inventory);
```

A variável `inventory` deve ter somente um valor.

```js
assert.lengthOf(inventory, 1);
```

Sua variável `inventário` deve incluir o texto `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
