---
id: 62fc20387ef88d1d1998aac5
title: Passo 18
challengeType: 0
dashedName: step-18
---

# --description--

Use `querySelector()` para obter os outros dois elementos `button`, usando os `id`s `button2` e `button3`. Armazene-os em variáveis chamadas `button2` e `button3`. Lembre-se de usar `const`.

# --hints--

Você deve declarar uma variável `button2` com `const`.

```js
assert.match(code, /const\s+button2/);
```

Sua variável `button2` deve ter o valor do seu elemento `"#button2"`.

```js
assert.match(code, /const\s+button2\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#button2\1\s*\)/);
```

Você deve declarar uma variável `button3` com `const`.

```js
assert.match(code, /const\s+button3/);
```

Sua variável `button3` deve ter o valor do seu elemento `"#button3"`.

```js
assert.match(code, /const\s+button3\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#button3\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");

--fcc-editable-region--
```
