---
id: 5e44412c903586ffb414c94c
title: Jenga mradi wa mpangaji hesabu
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462359
dashedName: build-an-arithmetic-formatter-project
---

# --description--

Wanafunzi wa shule ya msingi mara nyingi huandaa matatizo ya hesabu wima ili kuyafanya kuwa rahisi kuyatatua. Kwa mfano, "235 + 52" inakuwa:

```py
  235
+  52
-----
```

Maliza kitendakazi cha `arithmetic_arranger` kinachopokea orodha ya mfuatano wa herufi ambao ni matatizo ya hesabu, na kurudisha matatizo yakiwa yamepangwa wima na kando-kando. Kitendakazi kinapaswa kuchukua hoja ya pili hiari. Wakati hoja ya pili imewekwa kuwa `True`, majibu yanapaswa kuonyeshwa.

## Mfano

Wito wa kitendakazi:

```py
arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])
```

Matokeo:

```py
   32      3801      45      123
+ 698    -    2    + 43    +  49
-----    ------    ----    -----
```

Wito wa kitendakazi:

```py
arithmetic_arranger(["32 + 8", "1 - 3801", "9999 + 9999", "523 - 49"], True)
```

Matokeo:

```py
  32         1      9999      523
+  8    - 3801    + 9999    -  49
----    ------    ------    -----
  40     -3800     19998      474
```

## Sheria

Kitendakazi kitarejesha uongofu sahihi ikiwa matatizo yaliyotolewa yamepangwa ipasavyo, vinginevyo, kitarejesha **mfuatano wa herufi** unaoelezea kosa lenye maana kwa mtumizi.

- Hali zitakazorejesha kosa:
  - Ikiwa matatizo yaliyotolewa kwa kitendakazi ni mengi sana. Kiwango ni **matatizo matano**, zaidi ya hayo kitarudisha:
    `'Error: Too many problems.'`
  - Waendeshaji wanaokubalika kwa kitendakazi ni **jumlisha** na **toa**. Kuzidisha na kugawanya kitarudisha kosa. Waendeshaji wengine wasiotajwa hapa hawatahitaji kujaribiwa. Kosa litakalorejeshwa ni:
    `"Error: Operator must be '+' or '-'."`
  - Kila nambari (operandi) inapaswa kuwa na tarakimu tu. Vinginevyo, kitendakazi kitarudisha:
    `'Error: Numbers must only contain digits.'`
  - Kila operandi (yaani nambari upande wowote wa mwezeshaji) inaruhusiwa kuwa na tarakimu nne kwa upana. Vinginevyo, mfuatano wa kosa utarejeshwa:
    `'Error: Numbers cannot be more than four digits.'`
- Ikiwa mtumizi atatoa matatizo yenye muundo sahihi, uongofu utakaoletwa utazingatia sheria hizi:
  - Inapaswa kuwepo na nafasi moja kati ya mwezeshaji na operandi ndefu zaidi kati ya mbili, mwezeshaji atakuwa kwenye mstari mmoja na operandi wa pili, operand zote zitawekwa kwa mpangilio ule ule kama ulivyotolewa (ya kwanza itakuwa juu na ya pili chini).
  - Nambari zitawekwa zikiwa zimepangwa kulia.
  - Kutakuwa na nafasi nne kati ya kila tatizo.
  - Kutakuwa na mistari ya dashes chini ya kila tatizo. Mistari hii itakuwa inakimbia kwa urefu mzima wa tatizo kila moja kwa pekee. (Mfano ulio juu unaonyesha jinsi inavyopaswa kuonekana.)

Kumbuka: fungua konsoli ya kivinjari kwa F12 kuona matokeo ya majaribio kwa undani zaidi.

# --hints--

`arithmetic_arranger(["3801 - 2", "123 + 49"])` inapaswa kurudisha `  3801      123\n-    2    +  49\n------    -----`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3801 - 2", "123 + 49"]), '  3801      123\\n-    2    +  49\\n------    -----')`);
  }
})
```

`arithmetic_arranger(["1 + 2", "1 - 9380"])` inapaswa kurudisha `  1         1\n+ 2    - 9380\n---    ------`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["1 + 2", "1 - 9380"]), '  1         1\\n+ 2    - 9380\\n---    ------')`);
  }
})
```

`arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"])` inapaswa kurudisha `    3      3801      45      123\n+ 855    -    2    + 43    +  49\n-----    ------    ----    -----`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"]), '    3      3801      45      123\\n+ 855    -    2    + 43    +  49\\n-----    ------    ----    -----')`);
  }
})
```

`arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"])` inapaswa kurudisha `  11      3801      1      123         1\n+  4    - 2999    + 2    +  49    - 9380\n----    ------    ---    -----    ------`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"]), '  11      3801      1      123         1\\n+  4    - 2999    + 2    +  49    - 9380\\n----    ------    ---    -----    ------')`);
  }
})
```

`arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"])` inapaswa kurudisha `'Error: Too many problems.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"]), 'Error: Too many problems.')`);
  }
})
```

`arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"])` inapaswa kurudisha `"Error: Operator must be '+' or '-'."`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Operator must be '+' or '-'.")`);
  }
})
```

`arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"])` inapaswa kurudisha `'Error: Numbers cannot be more than four digits.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers cannot be more than four digits.")`);
  }
})
```

`arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"])` inapaswa kurudisha `'Error: Numbers must only contain digits.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers must only contain digits.")`);
  }
})
```

`arithmetic_arranger(["3 + 855", "988 + 40"], True)` inapaswa kurudisha `    3      988\n+ 855    +  40\n-----    -----\n  858     1028`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "988 + 40"], True), "    3      988\\n+ 855    +  40\\n-----    -----\\n  858     1028")`);
  }
})
```

`arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True)` inapaswa kurudisha `   32         1      45      123      988\n- 698    - 3801    + 43    +  49    +  40\n-----    ------    ----    -----    -----\n -666     -3800      88      172     1028`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True), "   32         1      45      123      988\\n- 698    - 3801    + 43    +  49    +  40\\n-----    ------    ----    -----    -----\\n -666     -3800      88      172     1028")`);
  }
})
```

# --seed--

## --seed-contents--

```py
def arithmetic_arranger(problems, show_answers=False):

    return problems

print(f'\n{arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])}')
```

# --solutions--

```py
def arithmetic_arranger(problems, result=False):

	lin1 = ""
	lin2 = ""
	lin3 = ""
	lin4 = ""

	if len(problems) > 5:
		return 'Error: Too many problems.'
	for problem in problems:
		[num1, sym, num2] = problem.split()
		sign = ['+', '-']
		if sym not in sign:
			return ("Error: Operator must be '+' or '-'.")
		if len(num1) > 4 or len(num2) > 4:
			return ("Error: Numbers cannot be more than four digits.")
		if not num1.isnumeric() or not num2.isnumeric():
			return ("Error: Numbers must only contain digits.")

		lin1 += "  " + num1 + "    " if len(num1) >= len(
		    num2) else " " * (len(num2) + 2 - len(num1)) + num1 + "    "
		lin2 += sym + " " + num2 + "    " if len(num2) >= len(
		    num1) else sym + " " * (len(num1) - len(num2) + 1) + num2 + "    "
		nmax = (len(num1) + 2) if len(num1) >= len(num2) else (len(num2) + 2)
		lin3 += "-" * nmax + "    "
		op = int(num1) + int(num2) if sym == "+" else int(num1) - int(num2)
		lin4 += (" " * (nmax - len(str(op)))) + str(op) + "    "

	arranged_problems = lin1.rstrip() + "\n" + lin2.rstrip(
	) + "\n" + lin3.rstrip()
	if result:
		arranged_problems += "\n" + lin4.rstrip()

	return arranged_problems
```
