---
id: 61fd6343fb42013d99bcd7f3
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `body` 元素內，將一個 `section` 元素巢狀於一個 `main` 元素中。

# --hints--

`body` 元素內的第一個元素應該是你的 `main` 元素。 

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

你應該在你的 `main` 元素中有一個 `section` 元素。

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
