---
id: 61fd5a93fd62bb35968adeab
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

設定你的 HTML，包含 `DOCTYPE`、表示此文件為英文的 `html`、`head` 和 `body` 元素。

為你的 `head` 元素添加適當的 `meta` 元素以設定 `charset` 和 `viewport`，一個具有適當標題的 `title` 元素，以及一個用於你的樣式表的 `link` 元素。

# --hints--

你的程式碼應該有一個 `<!DOCTYPE html>` 宣告式。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你的程式碼應該包含一個 `html` 元素。

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

你的 `html` 元素應該有一個開啟標籤，並且帶有 `lang` 屬性且值為 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的程式碼應該在 `html` 元素內有一個 `head` 元素。

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

你的程式碼應該在 `html` 元素內有一個 `body` 元素。

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

你的 `head` 元素應該位於你的 `body` 元素之前。

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

你應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

一個 `meta` 元素應該設定 `name` 為 `viewport`，並將 `content` 設定為 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一個 `meta` 元素應該將 `charset` 屬性設為 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

你的程式碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的 `title` 應該包含一些文字。

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

你的程式碼應該有一個 `link` 元素。

```js
assert(/<link/.test(code))
```

你的 `link` 元素應該位於你的 `head` 元素內。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
