---
id: 6391d1a4f7ac71efd0621380
title: بناء صفحة وصفة
challengeType: 25
dashedName: top-build-a-recipe-project
---

# --description--

سيتكون الموقع من صفحة فهرس رئيسية تحتوي على روابط لعدد من الوصفات. لن يبدو الموقع جميلاً عند الانتهاء منه.

**قصص المستخدم:**

1. يجب أن تحتوي صفحة الوصفة على وسم `DOCTYPE`.
1. يجب أن تتضمن صفحة الوصفة عنصر `html` يحتوي على عنصرَي `head` و`body` كأبناء.
1. يجب أن يكون لديك عنصر `title` داخل عنصر `head` يحتوي على النص `The Odin Recipes`.
1. يجب أن ترى عنصر `h1` يحتوي على النص `Creamy Chocolate Fudge`.
1. يجب أن ترى صورة ذات صلة تحتوي على خاصية `alt`.
1. يجب أن يكون هناك عنصر `h2` يحتوي على النص `Description` تحت الصورة.
1. يجب أن ترى فقرتين تحت `Description` تصفان الوصفة.
1. يجب أن يكون هناك عنصر `h2` يحتوي على النص `Ingredients`.
1. تحت عنوان `Ingredients` يجب أن تكون هناك قائمة غير مرتبة تحتوي على المكونات اللازمة للوصفة.
1. تحت قائمة المكونات أضف عنوانًا آخر يسمى `Steps`.
1. يجب أن ترى قائمة مرتبة تحتوي على بعض الخطوات اللازمة لإكمال الوصفة.
1. تحت الخطوات يجب أن يكون هناك عنصر `h2` يحتوي على النص `More Recipes`.
1. يجب أن ترى بعض الروابط لوصفات أخرى داخل قائمة غير مرتبة تحتوي على بعض عناصر القائمة التي بداخلها عناصر وسم الرابط.
1. يجب أن تحتوي عناصر وسم الرابط هذه على خاصية `href` بقيمة `#`.

# --hints--

يجب أن تحتوي على وسم `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+?html\s*?>/gi));
```

يجب أن تحتوي على عنصر `html` مع عنصرَي `head` و`body`.

```js
const html = document.querySelectorAll('html')[0];
const head = document.querySelectorAll('html > head')[0];
const body = document.querySelectorAll('html > body')[0];

assert(html && head && body);
```

يجب أن يحتوي عنصر `title` داخل عنصر `head` على النص 
`The Odin Recipes`.

```js
assert(document.querySelectorAll('HEAD > TITLE')[0]?.innerText == 'The Odin Recipes');
```

يجب أن يحتوي عنصر `h1` داخل عنصر `body` على النص
`Creamy Chocolate Fudge`.

```js
assert(document.querySelectorAll('BODY > H1')[0]?.innerText == 'Creamy Chocolate Fudge');
```

يجب أن تحتوي على صورة مع خاصية `alt`.

```js
const img = document.querySelectorAll('IMG')[0];

assert(img && img.alt !='' && img.src != '')
```

يجب أن يحتوي على عنصر `h2` يحتوي على النص `Description`.

```js
const h2 = document.querySelectorAll('H2')[0];

assert(h2?.innerText == 'Description');
```

يجب أن يحتوي على فقرتين على الأقل `p` تصفان الوصفة.

```js
const paragraphs = document.querySelectorAll('P');

assert(paragraphs?.length > 1);
```

يجب أن يحتوي على عنصر `h2` يحتوي على النص `Ingredients`.

```js
const h2 = document.querySelectorAll('H2')[1];

assert(h2?.innerText == 'Ingredients');
```

يجب أن يحتوي على عنصر `ul` مسكن مع عناصر `li` تحتوي على المكونات.

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "Ingredients") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);
        } else {
            assert(false);
        }
    }
});
```

يجب أن يحتوي على عنصر `h2` يحتوي على النص `Steps`.

```js
const h2 = document.querySelectorAll('H2')[2];

assert(h2?.innerText == 'Steps');
```

يجب أن يحتوي على عنصر `ol` يحتوي على عناصر `li` مع خطوات الوصفة.

```js
const orderedList = document.querySelectorAll('OL')[0];
const listItems = document.querySelectorAll('OL > LI');

assert(orderedList && listItems && listItems.length > 1);
```

يجب أن يحتوي على عنصر `h2` يحتوي على النص `More Recipes`.

```js
const h2 = document.querySelectorAll('H2')[3];

assert(h2?.innerText == 'More Recipes');
```

يجب أن يحتوي على عنصر `ul` مسكن مع عناصر `li` تحتوي على عناصر `a` تربط بوصفات أخرى.

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "More Recipes") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);

            listItems.forEach(listItem => {
                const anchor = listItem.querySelector("A");

                assert(anchor);
            });
        } else {
            assert(false);
        }
    }
});
```

يجب أن تحتوي عناصر وسم الرابط التي تربط بالوصفات على خاصية `href` بقيمة `#`.

```js
const anchorTags = document.querySelectorAll("a");

const allAnchorsHaveHrefHash = [...anchorTags].every(function(anchorTag) {
  return anchorTag.hasAttribute("href") && anchorTag.getAttribute("href") === "#";
});

assert(allAnchorsHaveHrefHash && anchorTags.length > 0); 
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>The Odin Recipes</title>
  </head>
  <body>
    <h1>Creamy Chocolate Fudge</h1>
    <img src="https://cdn.freecodecamp.org/curriculum/odin-project/build-a-recipe-page/build-a-recipe-page-01.jpg" alt="A delicious chocolate fudge dessert">
    <h2>Description</h2>
    <p>This recipe is for a rich and creamy chocolate fudge that is sure to satisfy your sweet tooth. It's perfect for a special occasion or as a tasty treat for any time of the year.</p>
    <p>This recipe is easy to follow and only requires a few simple ingredients. With just a few steps, you'll be able to create a delicious dessert that everyone will love.</p>
    <h2>Ingredients</h2>
    <ul>
      <li>1 cup sugar</li>
      <li>1/2 cup unsalted butter</li>
      <li>1/4 cup milk</li>
      <li>1/4 cup cocoa powder</li>
      <li>1/4 cup chocolate chips</li>
      <li>1/4 tsp salt</li>
      <li>1 tsp vanilla extract</li>
    </ul>
    <h2>Steps</h2>
    <ol>
      <li>In a medium saucepan, melt the butter over medium heat.</li>
      <li>Add the sugar, milk, cocoa powder, and salt to the saucepan and stir until well combined.</li>
      <li>Bring the mixture to a boil, stirring constantly, and then reduce the heat to low and simmer for 5 minutes.</li>
      <li>Remove the saucepan from the heat and stir in the chocolate chips and vanilla extract until the chocolate is melted and the mixture is smooth.</li>
      <li>Pour the fudge into a greased 8-inch square pan and let it cool completely before cutting into squares.</li>
    </ol>
    <h2>More Recipes</h2>
    <ul>
      <li><a href="#">Peanut Butter Cookies</a></li>
      <li><a href="#">Lemon Bars</a></li>
      <li><a href="#">Chocolate Chip Pancakes</a></li>
    </ul>
  </body>
</html>
```

```css

```
