---
id: 6690e10ebe2181212abc9652
title: Tengeneza kadi ya biashara
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na `div` yenye sifa ya kipengele ya `class` yenye thamani ya `business-card` ambayo itakuwa na vipengele vyote vingine ndani yake. 
1. Ndani ya kipengele cha `.business-card`, kunapaswa kuwepo na kipengele cha `img` chenye sifa ya kipengele ya `class` yenye thamani ya `profile-image`. Unaweza kuweka chanzo cha picha kuwa `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg` ikiwa unataka. Kunapaswa kuwepo na `alt` yenye maelezo yenye maana.
1. Ndani ya kipengele cha `.business-card`, baada ya kipengele cha `img`, unapaswa kuwa na vipengele vitatu vya `p` vyenye sifa ya kipengele ya `class` zenye thamani za `full-name`, `designation`, na `company` kwa mpangilio huo.
1. Kipengele cha kwanza cha `p` kinapaswa kuwa na jina lako.
1. Kipengele cha pili cha `p` kinapaswa kuwa na cheo chako.
1. Kipengele cha tatu cha `p` kinapaswa kuwa na jina la kampuni yako. 
1. Kunapaswa kuwepo na kipengele cha `hr` chini ya kipengele cha tatu cha `p`.
1. Baada ya kipengele cha `hr`, kunapaswa kuwepo na vipengele viwili vya `p`. Katika kipengele cha kwanza cha `p`, unapaswa kuwa na anwani ya barua pepe kama maandishi. Katika kipengele cha pili cha `p`, maandishi yanapaswa kuwa nambari ya simu.
1. Baada ya vipengele viwili vya `p`, kunapaswa kuwepo na kipengele cha `a` chenye darasa la vitu la `portfolio-link`, maandishi ya `Portfolio`, na kinapaswa kuunganisha na URL halali.
1. Kunapaswa kuwepo na kipengele cha `hr` baada ya kipengele cha `a` chenye kiungo cha portfolio.
1. Unapaswa kuwa na kipengele kingine cha `div` chenye sifa ya kipengele ya `class` yenye thamani ya `social-media`. Ndani ya kipengele hiki, kunapaswa kuwepo na kipengele cha `h2` chenye maandishi ya `Connect with me`. 
1. Ndani ya vipengele vya `.social-media`, kunapaswa kuwepo na vipengele vitatu vya `a`. Vipengele vya `a` vinapaswa kuwa na maandishi ya `Twitter`, `LinkedIn` na `GitHub` kwa mpangilio huo na viungo vya tovuti halali.
1. Unapaswa kuunganisha karatasi ya mtindo `styles.css` kwa kutumia lebo ya `link` katika sehemu ya `head` ya faili ya HTML.
1. Mandhari ya ukurasa wako inapaswa kuwa `rosybrown`. Fonti kuu inapaswa kuwa `Arial`, na njia mbadala ya fonti `sans-serif`.   
1. Kichaguzi cha `.business-card` kinapaswa kuwa na sifa za kuweka kipengele kuwa na upana wa `300px` na rangi ya mandhari unayochagua. Pia unapaswa kuweka ukingo wa `padding` wa `20px` pande zote na ukingo wa juu wa `margin` wa `100px`. Maandishi yanapaswa kuwa katikati na ukubwa wa fonti uwe `16px`. Ukingo wa kushoto na kulia unapaswa kuwekwa kuwa `auto`.
1. Kichaguzi cha `.profile-image` kinapaswa kuwa na sifa ya `max-width` yenye thamani ya `100%`.
1. Vipengele vyote vya `p` vinapaswa kuwa na ukingo wa juu na chini wa `margin` wa `5px`.
1. Viungo vyote kwenye ukurasa havipaswi kuwa na mstari wa chini.

# --hints--

Unapaswa kuwa na `div` kuu yenye sifa ya kipengele ya `class` yenye thamani ya `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

Ndani ya kipengele cha `.business-card`, kunapaswa kuwepo na kipengele cha `img` chenye sifa ya kipengele ya `class` yenye thamani ya `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

Chanzo cha picha kinapaswa kuwekwa kuwa picha halali.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

Sifa ya kipengele ya `alt` ya picha inapaswa kuwekwa kuwa maelezo yenye maana.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

Ndani ya kipengele cha `.business-card`, baada ya kipengele cha `img`, kunapaswa kuwepo na kipengele cha `p` chenye sifa ya kipengele ya `class` yenye thamani ya `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Baada ya kipengele cha `.full-name`, kunapaswa kuwepo na kipengele cha `p` chenye sifa ya kipengele ya `class` yenye thamani ya `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Baada ya kipengele cha `.designation`, kunapaswa kuwepo na kipengele cha `p` chenye sifa ya kipengele ya `class` yenye thamani ya `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

Kipengele cha kwanza cha `p`, kile chenye darasa la vitu la `full-name`, kinapaswa kuwa na jina lako.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

Kipengele cha pili cha `p`, kile chenye sifa ya kipengele ya `class` yenye thamani ya `designation`, kinapaswa kuwa na cheo chako.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

Kipengele cha tatu cha `p`, kile chenye sifa ya kipengele ya `class` yenye thamani ya `company`, kinapaswa kuwa na jina la kampuni yako.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

Baada ya kipengele cha tatu cha `p`, kile chenye sifa ya kipengele ya `class` yenye thamani ya `company`, kunapaswa kuwepo na kipengele cha `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Baada ya kipengele cha kwanza cha `hr`, kunapaswa kuwepo na kipengele cha `p` chenye anwani ya barua pepe kama maandishi yako.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

Baada ya kipengele cha barua pepe cha `p`, kunapaswa kuwepo na kipengele kingine cha `p` chenye nambari ya simu kama maandishi yako.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Baada ya kipengele cha nambari ya simu cha `p`, kunapaswa kuwepo na kipengele cha `a` chenye darasa la vitu la `portfolio-link`. 

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

Kipengele cha `.portfolio-link` kinapaswa kuwa na maandishi ya `Portfolio`.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

Kipengele chako cha `.portfolio-link` kinapaswa kuwa na `href` halali.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

Baada ya kipengele cha `a` chenye portfolio yako, kunapaswa kuwepo na kipengele kingine cha `hr` kama kigawanyo.

```js
assert.exists(document.querySelector('a + hr'));
```

Baada ya kipengele cha pili cha `hr`, kunapaswa kuwepo na kipengele cha `div` chenye sifa ya kipengele ya `class` yenye thamani ya `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

Ndani ya kipengele cha `.social-media`, kunapaswa kuwepo na kipengele cha `h2` chenye maandishi ya `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

Ndani ya kipengele cha `.social-media` kunapaswa kuwepo na vipengele vitatu vya `a`.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

Kipengele cha kwanza cha `a` kinapaswa kuwa na maandishi ya `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

Sifa ya kipengele ya `href` ya kipengele cha kwanza cha `a` inapaswa kuunganisha na kiungo halali.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

Kipengele cha pili cha `a` kinapaswa kuwa na maandishi ya `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

Sifa ya kipengele ya `href` ya kipengele cha pili cha `a` inapaswa kuunganisha na kiungo halali.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

Kipengele cha tatu cha `a` kinapaswa kuwa na maandishi ya `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

Sifa ya kipengele ya `href` ya kipengele cha tatu cha `a` inapaswa kuunganisha na kiungo halali.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `rel` yenye thamani ya `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `href` yenye thamani ya `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Unapaswa kuweka rangi ya mandhari ya ukurasa katika kichaguzi cha `body` kuwa `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

Unapaswa kuweka fonti ya ukurasa kuwa `Arial` katika kipengele cha `body` na njia mbadala ya `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

Kichaguzi chako cha `.business-card` kinapaswa kuwa na sifa ya `width` yenye thamani ya `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

Kichaguzi chako cha `.business-card` kinapaswa kuwa na sifa ya `background-color` yenye thamani halali ya rangi.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Kichaguzi chako cha `.business-card` kinapaswa kuwa na sifa ya `padding` yenye thamani ya `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

Kichaguzi chako cha `.business-card` kinapaswa kuwa na ukingo wa juu wa `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

Maandishi ndani ya kipengele cha `.business-card` yanapaswa kuwa yamewekwa katikati.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

Ukubwa wa fonti wa maandishi ndani ya kipengele cha `.business-card` unapaswa kuwa `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

Ukingo wa kushoto na kulia wa kipengele cha `.business-card` unapaswa kuwekwa kuwa `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

Kichaguzi chako cha `.profile-image` kinapaswa kuwa na sifa ya `max-width` yenye thamani ya `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

Vipengele vyako vya `p` vinapaswa kuwa na ukingo wa juu na chini wa `margin` wa `5px`.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

Viungo vya ukurasa havipaswi kuwa na mstari wa chini.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
