---
id: 66f17db06803d11a1bd19a20
title: Construir un Juego de Cuestionario
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historia de usuario:**

1. Deberías crear un array llamado `questions`.
1. El array `questions` debe contener al menos cinco objetos, cada uno con las claves `category`, `question`, `choices` y `answer`.
1. La clave `category` debe tener el valor de una cadena que represente una categoría de pregunta.
1. La clave `question` debe tener el valor de una cadena que represente una pregunta.
1. La clave `choices` debe tener el valor de un array que contenga tres cadenas, que son respuestas alternativas a la pregunta.
1. La clave `answer` debe tener el valor de una cadena, que representa la respuesta correcta a la pregunta. Además, el valor de `answer` debe estar incluido en el array `choices`.
1. Debes tener una función llamada `getRandomQuestion` que tome un array de preguntas como parámetro y devuelva un objeto de pregunta aleatoria del array.
1. Debes tener una función llamada `getRandomComputerChoice` que tome el array de las opciones disponibles como parámetro y devuelva una respuesta aleatoria a la pregunta seleccionada.
1. Debes tener una función llamada `getResults` que tome el objeto de pregunta como primer parámetro y la elección de la computadora como segundo parámetro. La función debe devolver `¡La elección de la computadora es correcta!` si la respuesta es correcta. De lo contrario, devuelve `La elección de la computadora es incorrecta. La respuesta correcta es: <code><correct-answer>`, donde `<correct-answer>` es el valor de la respuesta correcta a la pregunta elegida.

# --hints--

Deberías crear un array llamado `questions`.

```js
assert.isArray(questions);
```

El array `questions` debe contener al menos cinco objetos, cada uno con las claves `category`, `question`, `choices` y `answer`.

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

La clave `category` debe tener el valor de una cadena que represente una categoría de pregunta.

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => assert.isString(category));
```

La clave `question` debe tener el valor de una cadena que represente una pregunta.

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
})
```

La clave `choices` debe tener el valor de un array que contenga tres cadenas.

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    choices.forEach(choice => assert.isString(choice));
})
```

La clave `answer` debe tener el valor de una cadena.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => assert.isString(answer));
```

El valor de `answer` debe estar incluido en el array `choices`.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

Debes tener una función llamada `getRandomQuestion` que tome un array de preguntas como parámetro y devuelva un objeto de pregunta aleatoria del array.

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

Debes tener una función llamada `getRandomComputerChoice` que tome el array de las opciones disponibles como parámetro y devuelva una respuesta aleatoria a la pregunta seleccionada.

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

Debes tener una función llamada `getResults`.

```js
assert.isFunction(getResults);
```

Tu función `getResults` debe tomar el objeto de pregunta como primer parámetro y la elección de la computadora como segundo parámetro.

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
assert.notEqual(getResults(questions[0].answer, questions[0]), `The computer's choice is correct!`);
```

Si la elección de la computadora coincide con la respuesta, `getResults` debería devolver `¡La elección de la computadora es correcta!`

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

Si la elección de la computadora no coincide con la respuesta, `getResults` debería devolver `La elección de la computadora es incorrecta. La respuesta correcta es: <code><correct-answer>`, donde `<correct-answer>` es el valor de la respuesta correcta a la pregunta elegida.

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
