---
id: 66ea8239b9eeb76f12299030
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Наступна частина практичного заняття — повторення того, як згенерувати випадкове число між двома значеннями.

Почніть із додавання ще одного `console.log()`, який виводить повідомлення `"Now, generate a random number between two values."` у консоль.

Потім створіть змінну з назвою `min` і призначте їй значення `1`, а також створіть змінну з назвою `max` і призначте їй значення `100`.

У наступному кроці ви згенеруєте випадкове число між цими двома значеннями.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

У вас має бути `console.log()`, який виводить повідомлення `"Now, generate a random number between two values."`.

```js
assert.equal(getLogs()[3], "Now, generate a random number between two values.");
```

У вас має бути змінна з назвою `min`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.min);
```

Ваша змінна `min` має бути числом.

```js
assert.isNumber(min);
```

Ви маєте призначити число `1` змінній `min`.

```js
assert.equal(min, 1);
```

У вас має бути змінна з назвою `max`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.max);
```

Ваша змінна `max` має бути числом.

```js
assert.isNumber(max);
```

Ви маєте призначити число `100` змінній `max`.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
