---
id: 66ea7fc048a6e26e33c598fe
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Тепер настав час згенерувати випадкове число за допомогою методу `Math.random()`.

Створіть змінну з назвою `randomNum` і призначте їй значення результату виклику методу `Math.random()`.

Потім виведіть змінну `randomNum` у консоль.

Спробуйте додати пробіл у код, щоб знову запустити бота і побачити різні випадкові числа, які генеруються.

# --hints--

У вас має бути змінна з назвою `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomNum);
```

Ви повинні призначити результат виклику методу `Math.random()` змінній `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
const randomNum = explorer.variables.randomNum;
assert.equal(randomNum?.value?.toString(), "Math.random()");
```

Виведіть змінну `randomNum` у консоль.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
