---
id: 66ea7adae8053065a64f9002
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви повторите роботу з різними методами об’єкта `Math`, створюючи Mathbot. Цей Mathbot виводитиме в консоль деякі математичні операції та повідомлення, щоб навчити вас про об’єкт `Math` у JavaScript.

Для цього першого кроку почніть зі створення змінної з назвою `botName` і призначте їй рядкове значення `"MathBot"`.

Потім створіть іншу змінну з назвою `greeting` і призначте їй речення `"Hi there! My name is [botName goes here] and I am here to teach you about the Math object!"`. Замість заповнювача `[botName goes here]` використайте змінну `botName`.

Ви можете використовувати шаблонні літерали або конкатенацію рядків з оператором `+`, щоб досягти цього.

Нарешті, виведіть у консоль змінну `greeting`.

# --hints--

У вас має бути змінна з назвою `botName`.

```js
assert.isNotNull(botName);
```

Змінна `botName` має бути рядком.

```js
assert.isString(botName);
```

Ви маєте призначити рядкове значення `"MathBot"` змінній `botName`.

```js
assert.strictEqual(botName, "MathBot");
```

У вас має бути змінна з назвою `greeting`.

```js
assert.isNotNull(greeting);
```

Змінна `greeting` має бути рядком.

```js
assert.isString(greeting);
```

Ви маєте призначити наступне речення змінній `greeting`: `"Hi there! My name is [MathBot goes here] and I am here to teach you about the Math object!"`. Переконайтеся, що замінили `[MathBot goes here]` на змінну `botName`.

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

Ви маєте вивести у консоль змінну `greeting`.

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
