---
id: 66ea845f54195c7020cc521b
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Формула для генерації випадкового числа між двома значеннями виглядає так:

```js
Math.random() * (maximum - minimum) + minimum;
```

Це дасть результат у вигляді числа з плаваючою комою між двома значеннями.

Створіть змінну з назвою `randomNum2` і призначте їй результат генерації значення між `min` та `max`.

Потім виведіть у консоль змінну `randomNum2`, щоб побачити результат. Спробуйте оновити сторінку, щоб побачити різні результати.

# --hints--

У вас має бути змінна з назвою `randomNum2`.

```js
assert.isNotNull(randomNum2);
```

Змінна `randomNum2` має бути числом.

```js
assert.isNumber(randomNum2);
```

Ви маєте призначити результат генерації значення між `min` та `max`. Не використовуйте жорстко закодовані значення. Використайте формулу, наведену вище.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

Виведіть у консоль змінну `randomNum2`.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
