---
id: 66eb488c0129e192542f0865
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

У попередніх уроках ви навчились працювати з методом `Math.ceil()` так:

```js
const price = 10.01;
Math.ceil(price); // 11
```

Метод `Math.ceil()` округлює значення вгору до найближчого цілого числа.

Створіть змінну з назвою `numRoundedUp` і призначте їй результат округлення числа з плаваючою крапкою `3.2` вгору до найближчого цілого числа.

Потім виведіть змінну `numRoundedUp` у консоль, щоб побачити результат.

# --hints--

У вас має бути змінна `numRoundedUp`.

```js
assert.isNotNull(numRoundedUp);
```

Змінна `numRoundedUp` має бути числом.

```js
assert.isNumber(numRoundedUp);
```

Ваша змінна `numRoundedUp` має дорівнювати результату округлення числа з плаваючою крапкою `3.2` вниз до найближчого цілого числа.

```js
assert.equal(numRoundedUp, 4);
```

Ви не повинні жорстко задавати значення `4` для змінної `numRoundedUp`.

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

Ви маєте вивести змінну `numRoundedUp` у консоль.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
