---
id: 66eb3ec3b07dfe8f898646c0
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Як ви дізналися в попередніх уроках, метод `Math.floor()` округлює значення вниз до найближчого цілого числа.

```js
const price = 10.99;
Math.floor(price); // 10
```

Створіть змінну з назвою `numRoundedDown` і призначте їй результат округлення числа з плаваючою крапкою `6.7` вниз до найближчого цілого числа.

Потім виведіть змінну `numRoundedDown` у консоль, щоб побачити результат.

# --hints--

У вас має бути змінна `numRoundedDown`.

```js
assert.isNotNull(numRoundedDown);
```

Змінна `numRoundedDown` має бути числом.

```js
assert.isNumber(numRoundedDown);
```

Ваша змінна `numRoundedDown` має дорівнювати результату округлення числа з плаваючою крапкою `6.7` вниз до найближчого цілого числа.

```js
assert.equal(numRoundedDown, 6);
```

Ви не повинні жорстко кодувати значення `6` для змінної `numRoundedDown`.

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

Ви маєте вивести змінну `numRoundedDown` у консоль.

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
