---
id: 66ece147323c91f654bf8dd9
title: Крок 15
challengeType: 1
dashedName: step-15
---

# --description--

У попередніх уроках ви навчились отримувати максимальні та мінімальні значення з діапазону чисел ось так:

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

Створіть змінну з назвою `maxNum` і призначте їй результат знаходження максимального числа між числами `3, 125, 55, 24`. Потім виведіть значення `maxNum` у консоль.

Нижче створіть змінну з назвою `minNum` і призначте їй результат знаходження мінімального числа між числами `6, 90, 14, 90, 2`. Потім виведіть значення `minNum` у консоль.

# --hints--

Ви повинні мати змінну з назвою `maxNum`.

```js
assert.isNotNull(maxNum);
```

Вашій змінній `maxNum` має бути призначений результат знаходження максимального числа між числами `3, 125, 55, 24`.

```js
assert.equal(maxNum, 125);
```

Ви не повинні жорстко кодувати значення `125` у змінну `maxNum`. Переконайтеся, що замість цього використали метод `Math.max()`.

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

Ви повинні вивести значення `maxNum` у консоль.

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

Ви повинні мати змінну з назвою `minNum`.

```js
assert.isNotNull(minNum);
```

Вашій змінній `minNum` має бути призначений результат знаходження мінімального числа між числами `6, 90, 14, 90, 2`.

```js
assert.equal(minNum, 2);
```

Ви не повинні жорстко кодувати значення `2` у змінну `minNum`. Переконайтеся, що замість цього використали метод `Math.min()`.

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

Ви повинні вивести значення `minNum` у консоль.

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
