---
id: 69218e46f4490b33f2481a1a
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Формула для генерації випадкового цілого числа між двома значеннями виглядає так:

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

Це дасть результат у вигляді цілого числа між двома значеннями.

Створіть змінну з назвою `randomInt` і призначте їй результат генерації значення між `min` та `max`.

Потім виведіть у консоль змінну `randomInt`, щоб побачити результат. Спробуйте оновити сторінку, щоб побачити різні результати.

# --hints--

У вас має бути змінна з назвою `randomInt`.

```js
assert.isNotNull(randomInt);
```

Змінна `randomInt` має бути числом.

```js
assert.isNumber(randomInt);
```

Ви маєте призначити результат генерації значення між `min` та `max`. Переконайтеся, що не використовуєте жорстко закодоване значення. Використовуйте формулу, наведену вище.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

Виведіть у консоль змінну `randomInt`.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
