---
id: 69218a15fe536ee1be38aefb
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Наступна частина практичного заняття — повторення того, як згенерувати випадкове ціле число між двома значеннями.

Почніть із додавання ще одного `console.log()`, який виводить повідомлення `"Now, generate a random integer between two values."` у консоль.

У наступному кроці ви згенеруєте випадкове ціле число між значеннями змінних `min` і `max`, які ви вже створили.

# --hints--

У вас має бути `console.log()`, який виводить повідомлення `"Now, generate a random integer between two values."`.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+integer\s+between\s+two\s+values\.\s*\1\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

--fcc-editable-region--

--fcc-editable-region--
```
