---
id: 6690e10ebe2181212abc9652
title: Створіть дизайн візитівки
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Прив’яжіть таблицю стилів `styles.css` за допомогою тегу `link` у розділі `head` HTML-файлу.
1. Кольором фону сторінки має бути `rosybrown`. Основним шрифтом має бути `Arial`, а резервним — `sans-serif`.
1. Всі елементи `p` повинні мати верхнє і нижнє поля (`margin`) зі значенням `5px`.
1. Посилання на сторінці не мають бути підкресленими.
1. Ви повинні мати елемент `div` з атрибутом `class` зі значенням `business-card`, який буде містити всі інші елементи.
1. Селектор `.business-card` повинен мати властивості, щоб встановити ширину елемента на `300px` та колір фону на ваш вибір. Також потрібно встановити `padding` зі значенням `20px` зі всіх сторін і верхнє поле (`margin`) зі значенням `100px`. Текст потрібно вирівняти за центром, а розміром шрифту має бути `16px`. Ліве та праве поля повинні мати значення `auto`.
1. У межах елемента `.business-card` має бути елемент `img` з атрибутом `class` зі значенням `profile-image`. Джерелом зображення можна вказати `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`. Не забудьте про атрибут `alt` зі змістовним описом.
1. Селектор `.profile-image` повинен мати властивість `max-width` зі значенням `100%`.
1. У межах елемента `.business-card` після елемента `img` має бути три елементи `p` з атрибутом `class` з відповідними значеннями `full-name`, `designation` та `company`.
1. Перший елемент `p` має містити ім’я.
1. Другий елемент `p` має містити посаду.
1. Третій елемент `p` має містити назву компанії.
1. Під третім елементом `p` має бути елемент `hr`.
1. Після елемента `hr` має бути два елементи `p`. Текстом першого елемента `p` має бути електронна пошта. Текстом другого елемента `p` має бути номер телефону. Після двох елементів `p` має бути елемент `a` з текстом `Portfolio`, що містить дійсне посилання. Після елемента `a` додайте елемент `hr`, що містить посилання на портфоліо.
1. Ви повинні мати ще один елемент `div` з атрибутом `class` зі значенням `social-media`. У межах цього елемента має бути елемент `h2` з текстом `Connect with me`.
1. У межах елементів `.social-media` має бути три елементи `a`. Елементи `a` мають містити текст `Twitter`, `LinkedIn` та `GitHub` з дійсними посиланнями на сторінки.

# --hints--

Елемент `link` має бути в межах елемента `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

У селекторі `body` встановіть колір фону сторінки (`rosybrown`).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

В елементі `body` встановіть головний шрифт `Arial` і резервний варіант `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```


Всі елементи `p` повинні мати верхнє і нижнє поля (`margin`) зі значенням `5px`.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

Посилання на сторінці не мають бути підкресленими.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

Ви повинні мати основний елемент `div` з атрибутом `class` зі значенням `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

Селектор `.business-card` повинен мати властивість `width` зі значенням `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

Селектор `.business-card` повинен мати властивість `background-color` з дійсним значенням кольору.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Селектор `.business-card` повинен мати властивість `padding` зі значенням `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

Селектор `.business-card` повинен мати верхнє поле зі значенням `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

Текст в елементі `.business-card` потрібно вирівняти за центром.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

Розміром тексту в елементі `.business-card` має бути `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

Ліве та праве поля елемента `.business-card` повинні мати значення `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

У межах елемента `.business-card` має бути елемент `img` з атрибутом `class` зі значенням `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

Джерелом зображення потрібно вказати дійсне зображення.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

Зображення повинне мати атрибут `alt` зі змістовним описом.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

Селектор `.profile-image` повинен мати властивість `max-width` зі значенням `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

У межах елемента `.business-card` після елемента `img` має бути елемент `p` з атрибутом `class` зі значенням `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Після елемента `.full-name` має бути елемент `p` з атрибутом `class` зі значенням `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Після елемента `.designation` має бути елемент `p` з атрибутом `class` зі значенням `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

Перший елемент `p` (той, що має `class` зі значенням `full-name`) має містити ім’я.

```js
assert.isAbove(document.querySelector(".full-name")?.innerText.length, 0);
```

Другий елемент `p` (той, що має `class` зі значенням `designation`) має містити посаду.

```js
assert.isAbove(document.querySelector(".designation")?.innerText.length, 0);
```

Третій елемент `p` (той, що має `class` зі значенням `company`) має містити назву компанії.

```js
assert.isAbove(document.querySelector(".company")?.innerText.length, 0);
```

Після третього елемента `p` з атрибутом `class` зі значенням `company` має бути елемент `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Після першого елемента `hr` має бути елемент `p` з електронною поштою.

```js
assert.isAbove(document.querySelectorAll('.business-card hr + p')[0].innerText.length, 0);
```

Після елемента `p` з електронною поштою має бути ще один елемент `p` з номером телефону.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Після елемента `p` з номером телефону має бути елемент `a` з текстом `Portfolio`.

```js
assert.equal(document.querySelector('.business-card > a')?.textContent, 'Portfolio');
```

Елемент `a` з портфоліо повинен мати дійсний атрибут `href`.

```js
assert.isAbove(document.querySelectorAll('.business-card a')[0].href.length, 0);
```

Після елемента `a` з портфоліо має бути ще один елемент `hr`, який діє як розмежовувач.

```js
assert.exists(document.querySelector('a + hr'));
```

Після другого елемента `hr` має бути елемент `div` з атрибутом `class` зі значенням `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

У межах елемента `.social-media` має бути елемент `h2`, що містить текст `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent, 'Connect with me');
```

У межах елемента `.social-media` має бути три елементи `a`.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

Перший елемент `a` має містити текст `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent, 'Twitter');
```

`href` першого елемента `a` має вказувати на дійсне посилання.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

Другий елемент `a` має містити текст `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent, 'LinkedIn');
```

`href` другого елемента `a` має вказувати на дійсне посилання.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

Третій елемент `a` має містити текст `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent, 'GitHub');
```

`href` третього елемента `a` має вказувати на дійсне посилання.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
