---
id: 63b60821c855d01b1eda3c0b
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

次に、`fieldset` 要素を作成し、`id` を `lunch` に設定してください。

その要素の中に、テキストが `legend` の `Lunch` 要素と、`div` を `class` に設定した空の `input-container` を作成してください。

# --hints--

`fieldset` の中に2つ目の `form` 要素を作成してください。

```js
assert.exists(document.querySelectorAll('form fieldset')[1]);
```

2つ目の `fieldset` 要素は、1つ目の `fieldset` 要素の後に配置してください。

```js
assert.equal(document.querySelectorAll('form fieldset')[1]?.previousElementSibling?.tagName, "FIELDSET");
```

2つ目の `fieldset` 要素には、`id` を `lunch` に設定してください。

```js
assert.equal(document.querySelectorAll('form fieldset')[1]?.id, "lunch");
```

2つ目の `fieldset` 要素の中に `legend` 要素を含めてください。

```js
assert.exists(document.querySelectorAll('form fieldset')[1]?.querySelector('legend'));
```

新しい `legend` 要素には、テキスト `Lunch` を設定してください。

```js
assert.equal(document.querySelectorAll('form fieldset')[1]?.querySelector('legend')?.innerText, "Lunch");
```

2つ目の `fieldset` 要素の中に `div` 要素を含めてください。

```js
assert.exists(document.querySelectorAll('form fieldset')[1]?.querySelector('div'));
```

新しい`div`要素には`class`が`input-container`に設定されている必要があります。

```js
assert.equal(document.querySelectorAll('form fieldset')[1]?.querySelector('div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
