---
id: 63b60aaaa65f8922bfce6b7e
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

あなたの`span`要素の中に、`label`用の`entry-dropdown`要素を作成し、テキストに`Add food or exercise:`を設定してください。次に、`select`が`id`に設定され、`entry-dropdown`が`name`に設定された`options`要素を作成してください。その下に、`button`が`id`に設定され、テキストに`add-entry`を持つ`Add Entry`要素を追加してください。

自動フォーム送信を防ぐために、`button`要素に`type`属性を`button`に設定してください。

# --hints--

`label`要素に`span`要素を追加してください。

```js
assert.exists(document.querySelector('.controls > span > label'));
```

新しい`label`要素には、`for`属性を`entry-dropdown`に設定してください。

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

新しい`label`要素には、テキスト`Add food or exercise:`を設定してください。

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

`select`要素に`span`要素を追加してください。

```js
assert.exists(document.querySelector('.controls > span > select'));
```

`select`要素は`label`要素の後に配置してください。

```js
assert.equal(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName, 'LABEL');
```

新しい`select`要素には、`id`属性を`entry-dropdown`に設定してください。

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

新しい`select`要素には、`name`属性を`options`に設定してください。

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

`select`要素は空で、閉じる`</select>`タグを持つ必要があります。

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

`button`要素に`span`要素を追加してください。

```js
assert.exists(document.querySelector('.controls > span > button'));
```

`button`要素は`select`要素の後に配置してください。

```js
assert.equal(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName, 'SELECT');
```

新しい`button`要素には、`id`属性を`add-entry`に設定してください。

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

新しい`button`要素には、テキスト`Add Entry`を設定してください。

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
