---
id: 660404511dbf1b90eb23b617
title: Step 35
challengeType: 0
dashedName: step-35
---

# --description--

عôعéîعد`isInvalidInput`لûتµـ░عéْعâعé╣عâêعآعéïµآéعîµإحع╛عùعاعéعôع«عâعé╣عâêعدع»ع`1e3`عé`10e2`ع«عéêععزµîçµـ░كةذكذءعéْلûتµـ░عîµج£فç║عدععéïعïعرععïعéْقت║كزعùع╛عآعéعôعéîع»µـ░فجعéْكةذعآµ£ëفè╣عزµû╣µ│ـعدعآعîععéسعâصعâزعâ╝عéسعéخعâ│عé┐عâ╝عâùعâصعé╕عéدعé»عâêع«µ£ëفè╣عزفàحفèؤعدع»عéعéèع╛عؤعéôعé

`isInvalidInput`لûتµـ░ع«غ╕ïعسعé│عâ│عé╜عâ╝عâسµûçعéْك┐╜فèبعùعخععبعـععéعإع«عé│عâ│عé╜عâ╝عâسµûçع«غ╕صعدعف╝ـµـ░عس`isInvalidInput`عéْµ╕ةعùعخ`"1e3"`لûتµـ░عéْفّ╝ع│فç║عùعخععبعـععé

ق╡µئ£عéْكخïعéïعاعéعسعé│عâ│عé╜عâ╝عâسعéْلûïععخععبعـععéµشةع«عé╣عâعââعâùعدععإع«ق╡µئ£عîغ╜ـعéْµفّ│عآعéïع«عïعéْكر│عùعفصخع│ع╛عآعé

# --hints--

`console.log` لûتµـ░ع«غ╕ïعس `isInvalidInput` عîعéعéïع»عأعدعآعé

```js
assert.match(code, /console\.log\(/);
```

عé│عâ│عé╜عâ╝عâسµûçعسع»`isInvalidInput("1e3")`ع«فجعîعéعéïع╣ععدعآعé

```js
assert.match(code, /console\.log\(isInvalidInput\(('|")1e3\s*\1\);?\)/);

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>Breakfast</option>
                <option value="lunch">Lunch</option>
                <option value="dinner">Dinner</option>
                <option value="snacks">Snacks</option>
                <option value="exercise">Exercise</option>
              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
          <div>
            <button type="submit">
              Calculate Remaining Calories
            </button>
            <button type="button" id="clear">Clear</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');
const addEntryButton = document.getElementById('add-entry');
const clearButton = document.getElementById('clear');
const output = document.getElementById('output');
let isError = false;

function cleanInputString(str) {
  const regex = /[+-\s]/g;
  return str.replace(regex, '');
}

function isInvalidInput(str) {
  const regex = /\d+e\d+/i;
  return str.match(regex);
}
--fcc-editable-region--

--fcc-editable-region--
```
