---
id: 5d822fd413a79914d39e98dd
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

المباني متباعدة جدًا. اجمعها معًا بإضافة عنصرين فارغين من نوع `div` في أعلى عنصر `.background-buildings`، وعنصرين آخرين في أسفله، وعنصر واحد بين `.bb3` و `.bb4`. ستُضاف هذه العناصر بشكل متساوٍ عبر الحاوية، مما يحرك المباني أقرب إلى المركز.

# --hints--

يجب إضافة عنصرين جديدين من نوع `div` قبل عنصر `.bb1`.

```js
const bBuildings = document.querySelector('.background-buildings')?.children;
assert.isTrue(![...bBuildings?.[0]?.classList]?.includes('bb1'));
assert.isTrue(![...bBuildings?.[1]?.classList]?.includes('bb1'));
```

يجب إضافة عنصر جديد من نوع `div` بين عنصري `.bb3` و `.bb4`.

```js
assert.equal(document.querySelector('.bb3')?.nextElementSibling, document.querySelector('.bb4')?.previousElementSibling);
```

يجب إضافة عنصرين جديدين من نوع `div` بعد عنصر `.bb4`.

```js
const bb4 = document.querySelector('.bb4');
assert.exists(bb4?.nextElementSibling);
assert.exists(bb4?.nextElementSibling?.nextElementSibling);
```

يجب إضافة 5 عناصر جديدة من نوع `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 14);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      
      <div class="bb4"></div>
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
}

.bb3 {
  width: 10%;
  height: 55%;
}

.bb4 {
  width: 11%;
  height: 58%;
}
    
```

