---
id: 5d822fd413a79914d39e98ce
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

أضف أيضًا خاصية `box-sizing` بقيمة `border-box` لكل شيء. هذا يجعل الحدود التي أضفتها لا تضيف أي حجم إلى العناصر.

# --hints--

يجب استخدام الخاصية `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

يجب استخدام المحدِّد `*` الموجود.

```js
// Two selectors create two CSSStyleRule objects
assert.lengthOf(new __helpers.CSSHelp(document).getStyleDeclarations('*'), 1);
```

يجب أن تحتوي كل العناصر على خاصية `box-sizing` بقيمة `border-box`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
--fcc-editable-region--
  
--fcc-editable-region--
}

```

