---
id: 5d822fd413a79914d39e98cc
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

مرحبًا بك في مشروع CSS Variables Skyline! ابدأ بإضافة عنصر `link` يربط ملف `styles.css` داخل عنصر `head`.

# --hints--

يجب أن يحتوي شفرتك على عنصر `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

لا تغير وسم `head` الموجود لديك. تأكد من عدم حذف وسم الإغلاق.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

يجب أن يكون لديك عنصر `link` فارغ واحد.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
--fcc-editable-region--
    
--fcc-editable-region--
  </head>

  <body>
  </body>
</html>
```
