---
id: 5d822fd413a79914d39e98cd
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

في CSS، يمكنك استهداف كل شيء باستخدام النجمة *. أضف حدودًا لكل العناصر باستخدام المحدِّد `*`، ومنحها خاصية `border` بقيمة `1px solid black`. هذه خدعة تساعدك على تصور مكان العناصر وحجمها. ستزيل هذا لاحقًا.

# --hints--

يجب استخدام المحدِّد `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

يجب استخدام خاصية `border` لتنسيق كل العناصر.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

يجب أن تحتوي كل العناصر على حدود `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
