---
id: 5d822fd413a79914d39e98e6
title: الخطوة 27
challengeType: 0
dashedName: step-27
---

# --description--

قم بتسكين ستة عناصر `div` داخل `.foreground-buildings` وأعطها الأصناف من `fb1` إلى `fb6` بهذا الترتيب. "fb" تعني "مبنى المقدمة". ستكون هذه ستة مبانٍ إضافية للمقدمة.

# --hints--

يجب إنشاء عنصر `div` جديد بصنف `fb1`.

```js
assert.exists(document.querySelector('div.fb1'));
```

يجب إنشاء عنصر `div` جديد بصنف `fb2`.

```js
assert.exists(document.querySelector('div.fb2'));
```

يجب إنشاء عنصر `div` جديد بصنف `fb3`.

```js
assert.exists(document.querySelector('div.fb3'));
```

يجب إنشاء عنصر `div` جديد بصنف `fb4`.

```js
assert.exists(document.querySelector('div.fb4'));
```

يجب إنشاء عنصر `div` جديد بصنف `fb5`.

```js
assert.exists(document.querySelector('div.fb5'));
```

يجب إنشاء عنصر `div` جديد بصنف `fb6`.

```js
assert.exists(document.querySelector('div.fb6'));
```

يجب وضع هذه العناصر الجديدة `div` داخل عنصر `.foreground-buildings`.

```js
assert.exists(document.querySelector('div.foreground-buildings > div.fb1'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb2'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb3'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb4'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb5'));
assert.exists(document.querySelector('div.foreground-buildings > div.fb6'));
```

يجب وضع عناصر `div` الجديدة بالترتيب الصحيح.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.fb1','fb2') && __t('div.fb2','fb3') && __t('div.fb3','fb4') && __t('div.fb4', 'fb5') && __t('div.fb5', 'fb6'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
}

```

