---
id: 5d822fd413a79914d39e98d0
title: خطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أنشئ عنصر `div` داخل الـ `body` مع صنف `background-buildings`. ستكون هذه حاوية لمجموعة من المباني.

# --hints--

يجب أن تنشئ عنصر `div`.

```js
assert.exists(document.querySelector('div'));
```

يجب أن يكون عنصر `div` داخل الـ `body`.

```js
assert.equal(document.querySelector('div')?.parentElement?.localName, 'body');
```

يجب أن يحتوي عنصر `div` على صنف `background-buildings`.

```js
assert.isTrue([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

