---
id: 5d822fd413a79914d39e98d3
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

قم بتسكين أربعة عناصر `div` داخل الحاوية `.bb1`. امنحهم الأصناف `bb1a`، `bb1b`، `bb1c`، و `bb1d` بهذا الترتيب. سيكون لهذا المبنى أربعة أقسام.

# --hints--

يجب إنشاء أربعة عناصر `div` جديدة.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

يجب إعطاء أحد عناصر `div` الجديدة الصنف `bb1a`.

```js
assert.exists(document.querySelector('div.bb1a'));
```

يجب إعطاء أحد عناصر `div` الجديدة الصنف `bb1b`.

```js
assert.exists(document.querySelector('div.bb1b'));
```

يجب إعطاء أحد عناصر `div` الجديدة الصنف `bb1c`.

```js
assert.exists(document.querySelector('div.bb1c'));
```

يجب إعطاء أحد عناصر `div` الجديدة الصنف `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

يجب وضع عناصر `div` الجديدة بالترتيب الصحيح.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

يجب وضع عناصر `div` الجديدة داخل العنصر `.bb1`.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
    
```

