---
id: af7588ade1100bde429baf20
title: Einen fehlenden Buchstaben erkennen
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `fearNotLetter` haben.
1. Die `fearNotLetter`-Funktion sollte ein Argument akzeptieren: einen String, der einen Buchstabenbereich in alphabetischer Reihenfolge darstellt, in dem ein Buchstabe fehlen kann.
1. Die Funktion sollte den fehlenden Buchstaben im übergebenen Buchstabenbereich finden und zurückgeben.
1. Wenn alle Buchstaben im Bereich vorhanden sind, sollte die Funktion `undefined` zurückgeben.

# --hints--

Sie sollten eine `fearNotLetter`-Funktion haben.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` sollte den String `d` zurückgeben.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` sollte den String `i` zurückgeben.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` sollte den String `u` zurückgeben.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` sollte den String `e` zurückgeben.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` sollte `undefined` zurückgeben.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
