---
id: 63507b5ac8a55134223374c7
title: 步驟2
challengeType: 0
dashedName: step-2
---

# --description--

首先，`calculate` 函數需要找到在 `#numbers` 輸入字段中輸入的數字。 要做到這一點，使用 `.querySelector` 來定位輸入字段，然後使用 `.value` 屬性以獲得輸入的數字。

將其存儲在一個 `value` 變量中。

# --hints--

你的 `calculate` 函數應該有一個 `value` 變量。

```js
assert.match(calculate.toString(), /value/);
```

你的 `calculate` 函數應該調用 `document.querySelector()`，以 `"#numbers"` 作爲參數。

```js
assert.match(calculate.toString(), /document\.querySelector\(\s*("|')#numbers\1\s*\)/)
```

你應該訪問你的 `document.querySelector()` 調用的 `value` 屬性。

```js
assert.match(calculate.toString(), /document\.querySelector\(\s*("|')#numbers\1\s*\)\.value/);
```

你應該將 `value` 屬性賦值給你的 `value` 變量。

```js
assert.match(calculate.toString(), /value\s*=\s*document\.querySelector\(\s*("|')#numbers\1\s*\)\.value/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {

}
--fcc-editable-region--
```
