---
id: 63507bcbfe4ede356e624395
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

現在你有了輸入的值，你需要將其拆分成一個數字的數組。 使用 `.split()` 方法來實現這一點。

`.split()` 方法接收一個字符串，並將其拆分成一個字符串的數組。 你可以給它傳入一個字母的字符串或一個正則表達式作爲分隔符。 例如，`string.split(",")` 會在每個逗號處分割字符串並返回一個字符串的數組。

使用 `/,\s*/g` 正則通過逗號分割 `value` 字符串。 你可以根據值之間的空格數進行調整。 將該數組存儲在一個 `array` 變量中。

# --hints--

你的 `calculate` 函數應該有一個 `array` 變量。

```js
assert.match(code.toString(), /array/);
```

你的 `calculate` 函數應該調用 `.split()` 方法，處理 `value` 變量。

```js
assert.match(code.toString(), /value\.split()/);
```

你的 `calculate` 函數應該將 `.split()` 方法的結果賦值給 `array` 變量。

```js
assert.match(code.toString(),  /array\s*=\s*value\.split()/);
```

你應該使用 `/,\s*/g` 作爲 `split()` 方法的分隔符。

```js
assert.match(code.toString(), /value\.split\(\s*\/,\s*\\s*\*\s*\/g\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
}
--fcc-editable-region--
```
