---
id: a0b5010f579e69b815e7c5d6
title: 构建一个智能单词替换函数
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

在本实验中，你将创建一个函数，使用提供的参数对句子执行查找和替换操作，然后返回新的句子。

**注意：** 在更换原始单词时保留原始单词中第一个字符的大小写。 即如果传入的第二个参数为 `Book`，第三个参数为 `dog`，那么替换后的结果应为 `Dog`

**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `myReplace` 的函数。
1. `myReplace` 函数应接受三个参数：一个字串、一个要被替换的单词，以及用来替换的单词。
1. `myReplace` 函数应返回带有替换的新字串。
1. 你应该保留原始单词中第一个字符的分支。

# --hints--

你应该创建一个名为 `myReplace` 的函数。

```js
assert.isFunction(myReplace);
```

`myReplace` 应该接受三个参数。

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` 应返回 `Let us go to the mall`。

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` 应返回 `He is Sitting on the couch`。

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` 应返回 `I think we should look down there`。

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` 应返回 `This has a spelling error`。

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` 应返回 `His name is John`。

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` 应返回 `Let us get back to more Algorithms`。

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
