---
id: 6842a6cd9836f0114a5b7a8a
title: Створіть генератор числових шаблонів
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

У цій лабораторній роботі ви попрактикуєтеся у базових навичках Python, створюючи невеликий застосунок, який генерує числовий шаблон.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні визначити функцію з назвою `number_pattern`, яка приймає один параметр `n` (що представляє додатнє ціле число).
1. `number_pattern` має використовувати цикл `for`.
1. `number_pattern(n)` має повертати рядок зі всіма цілими числами від 1 до `n` (включно), розділеними пробілом. Наприклад, `number_pattern(4)` має повертати рядок `1 2 3 4`.
1. Якщо аргумент, переданий функції, не є цілим числом, функція має повертати `Argument must be an integer value.`.
1. Якщо аргумент, переданий функції, менший за 1, функція має повертати `Argument must be an integer greater than 0.`.

# --hints--

Ви повинні мати функцію `number_pattern`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

Функція `number_pattern` має мати один параметр з ім’ям `n`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` має повертати `1 2 3 4`.

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` має повертати `1 2 3 4 5 6 7 8 9 10 11 12`.

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` має повертати список чисел, розділених пробілом, для будь-якого додатнього цілого числа.

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

`number_pattern` має повертати `Argument must be an integer value.`, якщо передане значення не є цілим числом.

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` має повертати `Argument must be an integer greater than 0.`, якщо передане значення не є додатнім цілим числом.

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
