---
id: 6690b9626a66f51257ae5a91
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Die `h1`- bis `h6`-Überschriftenelemente werden verwendet, um die Wichtigkeit der Inhalte zu beschreiben. Je kleiner die Zahl desto wichtiger. `h2`-Elemente sind dementsprechend weniger wichtig als `h1`-Elemente.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Only use one `h1` element per page and place lower-importance headings below higher-importance headings.

Füge unter dem `h1`-Element ein `h2`-Element mit folgendem Text hinzu:

`Cat Photos`

# --hints--

Dein `h1`-Element sollte einen öffnenden Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Dein `h1`-Element sollte ein schließenden Tag haben. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Du solltest nur ein `h1`-Element haben. Entferne die überflüssigen.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Der Text deines `h1`-Elements sollte "CatPhotoApp" sein. Du hast entweder den Text vergessen oder hast einen Tippfehler gemacht.

```js
assert.equal(document.querySelector('h1').innerText.toLowerCase(), 'catphotoapp');
```

Dein `h2`-Element sollte ein öffnenden Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Dein `h2`-Element sollte einen schließenden Tag haben. Schließende Tags haben ein `/` direkt nach dem `<` Zeichen haben.

```js
assert.match(code, /<\/h2\>/);
```

Your `h2` element's text should be `Cat Photos`. Platziere den Text `Cat Photos` zwischen dem öffnenden und dem schließenden `h2` Tag.

```js
assert.equal(document.querySelector('h2').innerText.toLowerCase(), 'cat photos');
```

Dein `h2`-Element sollte unter dem `h1`-Element sein. Das `h1`-Element hat eine höhere Wichtigkeit und sollte über dem `h2`-Element sein.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```
