---
id: 6690bb0b4062b114997f708a
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 has some elements that identify different content areas. These elements make your HTML easier to read and help with Search Engine Optimization (SEO) and accessibility.

Das `main`-Element wird verwendet, um den Hauptinhalt des Hauptteils eines HTML-Dokuments darzustellen. Der Inhalt des `main`-Elements sollte für das Dokument einzigartig sein und nicht in anderen Teilen des Dokuments wiederholt werden.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identify the main section of this page by adding a `<main>` opening tag before the `h1` element, and a `</main>` closing tag after the `p` element.

# --hints--

Your `main` element should have an opening tag. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

Your `main` element should have a closing tag. Schließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert.match(code, /<\/main\>/);
```

Your `main` element's opening tag should be below the `body` element's opening tag. Sie sind in falscher Reihenfolge.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Your `main` element's opening tag should be above the `h1` element. Sie sind in falscher Reihenfolge.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

Your `main` element's closing tag should be below the `p` element. Sie stehen in falscher Reihenfolge.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode.contains(pNode));
assert.match(pNode.textContent.toLowerCase(), /see more cat photos in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```
