---
id: 6690ba16cf76f613b4e36197
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Kommentare sind dazu da, dass du Nachrichten im Code hinterlassen kannst, ohne dass sie auf der Webseite angezeigt werden. Kommentare können außerdem Code deaktivieren. Ein Kommentar in HTML beginnt mit `<!--`, enthält eine beliebige Anzahl von Textzeilen und endet mit `-->`.

Hier ist ein Beispiel für einen Kommentar mit dem `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Füge einen Kommentar über das `p`-Element mit diesem Text hinzu:

`TODO: Add link to cat photos`

# --hints--

Dein Kommentar sollte mit `<!--` beginnen. Dir fehlt ein oder mehrere Zeichen, die den Anfang eines Kommentars definieren.

```js
assert.match(code, /<!--/);
```

Dein Kommentar sollte mit `-->` aufhören.  Dir fehlt ein oder mehrere Zeichen, die das Ende eines Kommentars definieren.

```js
assert.match(code, /-->/);
```

Dein Code sollte keinen zusätzlichen öffnende oder schließende Kommentarzeichen enthalten. Dir wird ein zusätziches `<!--` oder `-->` im Browser angezeigt.

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g).length, 2)
assert.isBelow(noSpaces.match(/-->/g).length, 2);
```

Dein Kommentar sollte den Text `TODO: Add link to cat photos` enthalten.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

Dein Kommentar sollte über dem `p`-Element sein. Sie sind in falscher Reihenfolge.

```js
assert.match(
  code
    .replace(/\s/g, ''),
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```


