---
id: 6690be42942f421648b119eb
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Du kannst Bilder zu deiner Webseite mit dem `img`-Element hinzufügen. `img`-Elemente haben einen öffnenden Tag aber keinen schließenden. Ein Element ohne ein schließenden Tag wird auch <dfn>void element</dfn> genannt.

Füge ein `img`-Element unter dem `p`-Element hinzu. An diesem Punkt wird kein Bild im Browser angezeigt.

# --hints--

Dein `img`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Dein `img`-Element sollte kein schließendes Tag haben.

```js
assert.notMatch(code, /<\/img\>/);
```

Du solltest nur ein `img`-Element haben. Entferne alles Zusätzliche.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

Dein `img`-Element sollte unter dem `p`-Element stehen. Sie sind in falscher Reihenfolge.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```


