---
id: 6690b9a24f5b0f1300040c76
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Das `p`-Element wird verwendet, um einen Textabsatz auf den Webseiten zu erstellen. Erstelle ein `p`-Element unter deinem `h2`-Element und gib ihm den folgenden Text:

`See more cat photos in our gallery.`

# --hints--

Your `p` element should have an opening tag. Opening tags have the following syntax: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

Your `p` element should have a closing tag. Abschließende Tags enthalten ein `/` direkt nach dem `<`-Zeichen.

```js
assert.match(code, /<\/p\>/);
```

Dein `p` Elementtext sollte `See more cat photos in our gallery.` Du hast entweder den Text weggelassen oder hast einen Tippfehler im Text.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /see more cat photos in our gallery\.?$/i);
```

Your `p` element should be below the `h2` element. Du hast die Reihenfolge vertauscht.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```
