---
id: 66984e0cb4febb50c3794edf
title: Schritt 33
challengeType: 0
dashedName: step-33
---

# --description--

Das `footer`-Element wird verwendet, um eine Fußzeile für ein Dokument oder eine Sektion zu definieren. Eine Fußzeile enthält in der Regel Informationen über den Autor des Dokuments, Urheberinformationen, Links zu Nutzungsbedingungen, Kontaktinformationen und mehr.

Füge nach dem `main`-Element ein `footer`-Element hinzu.

# --hints--

Du hast entweder das `main`-Element gelöscht oder es fehlt ein öffnendes oder schließendes Tag.

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

Dein `footer`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben die folgende Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('footer'));
```

Dein `footer`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert.match(code, /<\/footer\>/);
```

Dein `footer`-Element sollte sich unterm schließenden `main`-Element-Tag befinden. Du hast es woanders hingetan.

```js
assert.equal(document.querySelector('main').nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
--fcc-editable-region--
    </main>

  </body>
--fcc-editable-region--
</html>
```


