---
id: 66984ec64a3652522e81c678
title: Schritt 35
challengeType: 0
dashedName: step-35
---

# --description--

Turn the existing `freeCodeCamp.org` text into a link by enclosing it in an anchor (`a`) element. The `href` attribute should be set to `https://www.freecodecamp.org`.

# --hints--

Dein (`a`)-Ankerelement sollte innerhalb des `p`-Elements im `footer` eingebettet sein. Stelle sicher, dass du ein öffnendes Tag und ein schließendes Tag für die (`a`)-Ankerelemente hinzufügst.

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

Dein (`a`)-Element sollte ein schließendes Tag haben. Abschließende Tags haben einen `/` direkt nach dem `<`-Zeichen haben.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 3);
```

Dein (`a`)-Ankerelement sollte ein `href`-Attribut mit dem Wert `https://www.freecodecamp.org` haben. Möglicherweise hast du das Attribut/den Wert weggelassen oder einen Tippfehler gemacht.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

Der Linktext sollte `freeCodeCamp.org` sein. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' '),
    'freecodecamp.org'
);
```

Nach der Einbettung des (`a`)-Ankerelements sollte der einzige Inhalt des `p`-Elements, der im Browser sichtbar ist, `No Copyright - freeCodeCamp.org` sein. Überprüfe den Text, die Abstände und die Zeichensetzung des `p` und des eingebetteten Ankerelements.

```js
const pText = document.querySelectorAll('footer > p')[0].innerText.toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```


