---
id: 669842cee55b683aa13ee758
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Füge innerhalb des zweiten `section`-Elements ein neues `h2`-Element mit dem Text `Cat Lists` hinzu.

# --hints--

Dein `section`Element sollte ein öffnendes Tag haben. Öffnende Tags besitzen die folgende Syntax: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Dein `h2`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

Dein `h2`-Element sollte ein schließendes Tag haben. Abschließende Tags enthalten ein `/` direkt nach dem `<`-Zeichen.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

Dein zweites `h2`-Element sollte sich genau über dem abschließenden Tag des zweiten `section`-Elements befinden. Es ist nicht in der richtigen Position.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection.lastElementChild.nodeName, 'H2');
```

Das zweite `h2`-Element sollte den Text `Cat Lists` haben. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```


