---
id: 6698419d8f52cf3891f34853
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

In previous steps, you used an anchor element to turn text into a link. Auch andere Arten von Inhalten können in einen Link umgewandelt werden, indem sie mit Anker-Tags umschlossen werden.

Hier ist ein Beispiel für die Umwandlung eines Bildes in einen Link:

```html
<a href="example-link">
  <img src="image-link.jpg" alt="A photo of a cat.">
</a>
```

Verwandle das Bild in einen Link, indem du es mit den notwendigen Tags umschließt. Verwende `https://freecatphotoapp.com` als den `href`-Attributwert des Ankers.

# --hints--

Du solltest ein `img`-Element mit einem `src` von `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg` haben. Möglicherweise hast du es versehentlich gelöscht.

```js
assert.equal(
  document.querySelector('img')?.getAttribute('src'),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg'
);
```

Dein (`a`)-Ankerelement sollte einen öffnenden Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('a').length, 2);
```

Dir fehlt ein schließender (`a`)-Tag nach dem Bild.

```js
assert.lengthOf(document.querySelectorAll('a'), 2);
```

Dein (`a`)-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert.isAtLeast(code.match(/<\/a>/g).length, 2);
```

Du solltest nur ein schließendes (`a`)-Ankertag hinzufügen. Entferne alle zusätzlichen.

```js
assert.lengthOf(code.match(/<\/a>/g), 2);
```

Dein (`a`)-Element hat kein `href`-Attribut. Überprüfe, ob ein Leerzeichen nach dem Namen des öffnenden Tags und/oder Leerzeichen vor allen Attributnamen vorhanden sind.

```js
assert.isTrue(document.querySelector('a').hasAttribute('href'));
```

Dein (`a`)-Ankerelement sollte auf `https://freecatphotoapp.com` verlinken. Du hast die URL entweder vergessen oder einen Tippfehler gemacht.

```js
assert.equal(
  document.querySelectorAll('a')[1].getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

Dein `img`-Element sollte innerhalb deines (`a`)-Ankerelements verschachtelt sein. Das gesamte `img`-Element sollte sich innerhalb der öffnenden und schließenden Tags des (`a`)-Ankerelements befinden.

```js
assert.equal(document.querySelector('img').parentNode.nodeName, 'A');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
--fcc-editable-region--
    </main>
  </body>
</html>
```


