---
id: 66985068aeb06e54562df39e
title: Schritt 38
challengeType: 0
dashedName: step-38
---

# --description--

Vielleicht ist dir aufgefallen, dass der gesamte Inhalt der Seite in einem `html` Element ist. Das `html`-Element ist der Grundbaustein einer HTML-Seite und umhüllt den gesamten Inhalt der Seite.

Du kannst auch die Sprache deiner Seite angeben, indem du das `lang`-Attribut zum `html`-Element hinzufügst.

Füge das `lang`-Attribut mit dem Wert `en` zum öffnenden `html`-Tag hinzu, um anzugeben, dass die Sprache der Seite Englisch ist.

# --hints--

Du hast entweder das `html`-Element gelöscht oder es fehlt ein öffnendes oder schließendes Tag.

```js
assert.match(code, /\<html.*?\>/);
assert.match(code, /\<\/html\>/);
```

Dein `html`-Element sollte ein `lang`-Attribut mit dem Wert `en` haben. Möglicherweise hast du das Attribut/den Wert weggelassen oder einen Tippfehler gemacht.

```js
const extraSpacesRemoved = code.replace(/\s+/g, " ");
assert.match(extraSpacesRemoved, /\<html\s+lang\s*\=\s*("?|'?)en\1\s*\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```


