---
id: 66984330a510e63bf0787898
title: Schritt 20
challengeType: 0
dashedName: step-20
---

# --description--

Um eine ungeordnete Liste von Elementen zu erstellen, kannst du das `ul`-Element verwenden.

Füge nach dem `h3`-Element mit dem `Things cats love:` Text ein ungeordnetes (`ul`)-Listenelement hinzu. Beachte, dass an diesem Punkt nichts angezeigt wird.

# --hints--

Dein `ul`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('ul'));
```

Dein `ul`-Element sollte ein schließendes Tag haben. Abschließende Tags enthalten ein `/` direkt nach dem `<`-Zeichen.

```js
assert.match(code, /<\/ul>/);
```

Das `ul`-Element sollte über dem schließenden Tag des zweiten `section`-Elements stehen.

```js
const secondSectionLastElemNode = $('main > section')[1].lastElementChild;
assert.equal(secondSectionLastElemNode.nodeName, 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```


