---
id: 6698457295f6403fa49050b8
title: Schritt 25
challengeType: 0
dashedName: step-25
---

# --description--

To place emphasis on a specific word or phrase, you can use the `em` element.

Emphasize the word `love` in the `figcaption` element by wrapping it in an emphasis `em` element.

# --hints--

Your emphasis `em` element should have an opening tag. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('em'));
```

Your emphasis `em` element should have a closing tag. Abschließende Tags enthalten ein `/` direkt nach dem `<`-Zeichen.

```js
assert.match(code, /<\/em\>/);
```

You have either deleted the `figcaption` element or it is missing an opening or closing tag.

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

Your emphasis `em` element should surround the text `love`. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.toLowerCase(), 'love'
);
```

The `figcaption`'s text should be `Cats love lasagna`. Check for typos and that the necessary spaces are present around the `em` element's opening and closing tags.

```js
assert.match(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' '),
    /cats love lasagna\.?/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```


