---
id: 6698496716e821448ba724a6
title: Schritt 32
challengeType: 0
dashedName: step-32
---

# --description--

Das `strong`-Element wird verwendet, um anzuzeigen, dass bestimmter Text wichtig, bzw. von großer Bedeutung ist.

In the `figcaption` you just added, indicate that `hate` is of strong importance by wrapping it in a `strong` element.

# --hints--

Your `strong` element should have an opening tag. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('strong'));
```

Your `strong` element should have a closing tag. Abschließende Tags enthalten ein `/` direkt nach dem `<`-Zeichen.

```js
assert.match(code, /<\/strong\>/);
```

Dein `strong`-Element sollte die Wörter `hate` in dem Text `Cats hate other cats.`. Du hast den Text entweder gelöscht oder einen Tippfehler gemacht.

```js
assert.equal(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.toLowerCase(), 'hate'
);
```

Der Text des `figcaption` Elements sollte `Cats hate other cats.` sein. Überprüfe ihn auf Tippfehler und schau, ob um das `strong` Element Leerzeichen vorhanden sind.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert.match(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim(),
      /cats hate other cats\.?/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
--fcc-editable-region--

          <figcaption>Cats hate other cats.</figcaption>  

--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```


