---
id: 66984535d474083eec6bb8da
title: Schritt 24
challengeType: 0
dashedName: step-24
---

# --description--

Ein Abbildungsbeschriftungselement (`figcaption`) wird verwendet, um eine Beschriftung hinzuzufügen, die das im `figure` enthaltene Bild beschreibt.

Hier ist ein Beispiel für ein `figcaption`-Element mit der Beschriftung von `A cute cat`:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

Füge nach dem im `figure`-Element eingebetteten Bild ein `figcaption`-Element mit dem folgenden Text hinzu:

`Cats love lasagna.`

# --hints--

Your `figcaption` element should have an opening tag. Opening tags have the following syntax: `<elementName>`.

```js
assert.exists(document.querySelector('figcaption'));
```

Your `figcaption` element should have a closing tag. Abschließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert.match(code, /<\/figcaption\>/);
```

The `figcaption` element should be nested in the `figure` element.

```js
assert.exists(document.querySelector('figure > figcaption'));
```

The lasagna `img` element should be nested in the `figure` element.

```js
assert.exists(document.querySelector('figure > img'))
assert.equal(
    document.querySelector('figure > img')?.getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

The `figcaption` element nested in the `figure` element should be below the `img` element. Sie sind in falscher Reihenfolge.

```js
assert.equal(
  document.querySelector('figcaption').previousElementSibling.nodeName, 'IMG'
);
```

Der Text deines `figcaption`-Elements sollte `Cats love lasagna.` sein. Du hast ihn entweder vergessen oder einen Tippfehler gemacht.

```js
assert.match(
  document.querySelector('figcaption').innerText, /Cats love lasagna.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```


