---
id: 669850ce0ccf395511ff4254
title: Step 39
challengeType: 0
dashedName: step-39
---

# --description--

Alle Seiten sollten mit `<!DOCTYPE html>` beginnen. Dieser spezielle String wird als eine <dfn>declaration</dfn> (Deklaration) bezeichnet und stellt sicher, dass der Browser versucht die branchenweiten Vorgaben zu erfüllen.

`<!DOCTYPE html>` teilt Browsern mit, dass das Dokument ein HTML5-Dokument ist, welches die neueste Version von HTML ist.

Füge diese Deklaration als erste Zeile des Codes hinzu.

# --hints--

Dein Code sollte mit der Deklaration `<!DOCTYPE html>` beginnen. Möglicherweise hast du die Deklaration weggelassen, einen Tippfehler gemacht oder die Deklaration ist nicht die erste Zeile des Codes.

```js
assert.match(code, /\<\s*!DOCTYPE\s+html\s*\>/);
```

Dein `<!DOCTYPE html>` sollte sich am Anfang des Dokuments befinden.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /^\<\!DOCTYPEhtml\>\<html/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

<html lang="en">
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```


