---
id: 669845af34899240d9196258
title: Schritt 27
challengeType: 0
dashedName: step-27
---

# --description--

Der Code für eine geordnete Liste (`ol`) ist einer ungeordneten Liste ähnlich, nur dass die Listenelemente in einer geordneten Liste nummeriert werden, wenn sie auf der Webseite angezeigt werden.

After the second `section` element's last `h3` element, add an ordered list with these three list items:

`flea treatment` `thunder` `other cats`

# --hints--

Your `ol` element should have an opening tag. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('ol'));
```

Your `ol` element should have a closing tag. Schließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert.match(code, /<\/ol>/);
```

The `ol` element should be above the second `section` element's closing tag. Du hast die Reihenfolge vertauscht.

```js
assert.equal(document.querySelectorAll('main > section')[1].lastElementChild.nodeName, 'OL');
```

The three `li` elements should be nested inside the `ol` element.

```js
assert.lengthOf(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ), 3
);
```

You should have three `li` elements with the text `flea treatment`, `thunder` and `other cats` in any order.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

You should only have one `ol` element.

```js 
assert.lengthOf([...document.querySelectorAll('ol')], 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```


