---
id: 61695d1fbc003856628bf561
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

El atributo `charset` especifica la codificación de caracteres que usa el documento. `utf-8` (Unicode Transformation Format – 8-bit) es un estándar de codificación de caracteres usado para la comunicación electrónica.

Dentro del elemento `head`, anida un elemento `meta` con un atributo `charset` con el valor `"utf-8"`.

# --hints--

Debes tener un elemento `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

Tu elemento `meta` debe ser un elemento vacío, no debe tener una etiqueta de cierre `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Tu elemento `meta` debe tener un atributo `charset` con el valor `"utf-8"`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Tu elemento `meta` debe estar dentro del elemento `head`.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
