---
id: 64dc92a9718fb53ca3d1bad7
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

يمكن حل لغز برج هانوي في 2<sup>n</sup> - 1 خطوة، حيث n هو عدد الأقراص. أعلن متغيرًا باسم `number_of_moves` وقم بتعيين العدد الإجمالي للخطوات لهذا المتغير.

عامل القوة في Python هو `**`.

# --hints--

يجب إعلان متغير باسم `number_of_moves`.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

يجب أن تكون قيمة `number_of_moves` هي التعبير لحساب عدد الخطوات.

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
