---
id: 64dc985e6720a23edac63b51
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

تُظهر الحركات المسموح بها للأقراص بين الأعمدة نمطًا متكررًا يحدث كل ثلاث حركات. على سبيل المثال، الحركات بين العمود `A` والعمود `C` مسموح بها في الحركة الأولى والرابعة والسابعة، حيث يكون الباقي من قسمة رقم الحركة على 3 مساويًا لـ 1.

داخل حلقة `for` التي أنشأتها سابقًا، استبدل استدعاء `print()` الحالي بعبارة `if` تُنفذ عندما يكون `(i + 1) % 3 == 1`. داخل هذه العبارة `if`، اطبع `f'Move {i + 1} allowed between {source} and {target}'` باستخدام سلسلة نصية من نوع f-string. لاحظ أن `i + 1` هو رقم الحركة لأن `i` يبدأ من صفر في التكرار الأول.

# --hints--

يجب أن تحتوي على عبارة `if` تُنفذ عندما يكون `(i + 1) % 3 == 1`.

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

يجب أن تطبع السلسلة النصية `f'Move {i + 1} allowed between {source} and {target}'` باستخدام f-string.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
