---
id: 64dc90949ff85e3b37be40d0
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

ترجع الدالة `range()` متتالية أرقام غير قابلة للتغيير. كما ترى، نوع البيانات لـ `rods['A']` هو `range`، لكنك تريد أن تكون قائمة.

مرر استدعاء `range()` إلى الدالة `list()` لتحقيق ذلك.

# --hints--

يجب أن تمرر استدعاء `range()` إلى الدالة `list()`.

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
