---
id: 64dc8b1faf9c60304ca8a9b6
title: خطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذا المشروع، ستحل اللغز الرياضي المعروف باسم برج هانوي. يتكون اللغز من ثلاثة أعمدة وعدد من الأقراص ذات أقطار مختلفة.

هدف هذا اللغز هو نقل الأقراص من العمود الأول إلى العمود الثالث، مع اتباع قواعد محددة تمنع وضع قرص أكبر فوق قرص أصغر.

ابدأ بإنشاء قاموس فارغ باسم `rods` لتمثيل الأعمدة.

# --hints--

يجب أن يكون لديك متغير باسم `rods`.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

يجب أن يكون المتغير `rods` قاموسًا فارغًا.

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
