---
id: 64dcdee07a14f0299ec30559
title: الخطوة 22
challengeType: 20
dashedName: step-22
---

# --description--

كتبت الشفرة لإيجاد الحركة المسموح بها بين القضبان، لكن الحركة الفعلية قد تكون في كلا الاتجاهين.
بعد استدعاء `print()`، أعلن متغيرًا باسم `forward` وضعه على `False`. ستستخدم هذا المتغير للتحقق من الاتجاه الذي تحتاج لتحريك القرص بين القضبان فيه.

# --hints--

يجب إعلان متغير باسم `forward`.

```js
({ test: () => assert.match(code, /forward\s*=\s*./) })
```

يجب تعيين متغير `forward` إلى `False`.

```js
({ test: () => assert.match(code, /forward\s*=\s*False/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
