---
id: 64dc9004f658183afa85f4f1
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

تحقق الآن من نوع البيانات للمفتاح `'A'` بتمريره إلى الدالة `type()` واطبع النتيجة على الطرفية.

# --hints--

يجب تمرير المفتاح `'A'` إلى الدالة `type()`.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

يجب طباعة نوع `rods['A']` باستخدام الدالة `print()`.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
