---
id: 64dc922df2919e3c38cead61
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

في أعلى شفرتك، أعلن عن متغير باسم `NUMBER_OF_DISKS` ليخزن عدد الأقراص وأعطه القيمة `3`. ثم، استبدل المعلمة الأولى الممررة إلى دالة `range()` بالمتغير الجديد.

# --hints--

يجب إعلان متغير باسم `NUMBER_OF_DISKS`.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

يجب أن يحتوي المتغير `NUMBER_OF_DISKS` على القيمة `3`.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

يجب تغيير قيمة المعلمة الأولى في دالة `range()` إلى المتغير `NUMBER_OF_DISKS`.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
