---
id: 64dc916138dbe03bdfe6cec6
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

الهدف من برج هانوي هو نقل جميع الأقراص إلى القضيب الأخير. لتحقيق ذلك، يجب اتباع ثلاث قواعد بسيطة:

1. يمكنك تحريك الأقراص العلوية فقط
2. يمكنك تحريك قرص واحد فقط في كل مرة
3. لا يمكنك وضع الأقراص الأكبر فوق الأقراص الأصغر

أسفل شفرتك الحالية، أعلن دالة فارغة باسم `move`. ستستخدم هذه الدالة لاحقًا لنقل الأقراص بين القضبان. الآن، لتجنب الأخطاء، استخدم الكلمة المفتاحية `pass` داخل جسم الدالة.

# --hints--

يجب إعلان دالة فارغة باسم `move`. تذكر استخدام الكلمة المفتاحية `pass` داخل جسم الدالة مع التنسيق الصحيح.

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
