---
id: 64dcd9bbc2268127e7898d77
title: الخطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

أخيرًا، عندما يكون رقم الحركة مقسومًا على 3 بدون باقي، يُسمح بالحركة بين `'B'` و `'C'`.

أضف جملة `elif` لذلك. ثم اطبع السلسلة النصية المناسبة إذا تحقق الشرط.

# --hints--

يجب أن يكون لديك جملة `elif` تُنفذ عندما يكون `remainder == 0`.

```js
({test: () => assert.match(code, /elif\s+remainder\s*==\s*0\s*:/)})
```

يجب أن تطبع السلسلة النصية `f'Move {i + 1} allowed between {auxiliary} and {target}'`.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*auxiliary\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
