---
id: 64dc8ea01436383a88256d7f
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

تبدأ اللغز بتكديس الأقراص على القضيب الأول، بترتيب تنازلي في الحجم، بحيث يكون أصغر قرص في الأعلى وأكبر قرص في الأسفل. عليك ملء القائمة الأولى بأرقام تمثل أحجام الأقراص المختلفة.

بدلاً من إضافة العناصر يدويًا إلى القائمة الأولى، أنشئ متتالية أرقام تنازلية من `3` إلى `1` باستخدام دالة `range()` وقم بتعيينها إلى `rods['A']`. هنا، يمثل الرقم `3` أكبر قرص في أسفل التكديس ويمثل الرقم `1` أصغر قرص في الأعلى.

الصيغة النحوية هي `range(x, y, h)`، حيث `x` هو العدد الصحيح الابتدائي (شامل)، و`y` هو العدد الصحيح الأخير (غير شامل)، و`h` هو الفرق بين رقم والرقم التالي في المتتالية.

# --hints--

يجب استخدام دالة `range()` لتعيين متتالية أرقام إلى `rods['A']`. صيغة استدعاء دالة `range()` هي `range(x, y, h)`.

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

يجب استخدام دالة `range()` لتعيين متتالية الأرقام من `3` إلى `1` إلى `rods['A']`.

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
