---
id: 64dc97005bc7943e2978df0a
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

يمكنك الآن إزالة استدعاء `print()` الخاص بك. ثم، داخل دالة `move()`، احذف الكلمة المفتاحية `pass` واطبع محتوى قاموس `rods`.

# --hints--

يجب ألا يحتوي شفرتك على `print(number_of_moves)`.

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

يجب إزالة الكلمة المفتاحية `pass` واطبع قاموس `rods`.

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
