---
id: 64de1469fccce61940f498c4
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

سيُمثل قاموس `rods` الأعمدة الثلاثة مع الأقراص الخاصة بها. أعطه السلاسل النصية `'A'` و `'B'` و `'C'` كمفاتيح واضبط كلًا منها على قائمة فارغة.

# --hints--

يجب أن يحتوي قاموس `rods` على المفتاح `'A'`.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

يجب أن تكون قيمة `rods['A']` قائمة فارغة.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

يجب أن يحتوي قاموس `rods` على المفتاح `'B'`.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

يجب أن تكون قيمة `rods['B']` قائمة فارغة.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

يجب أن يحتوي قاموس `rods` على المفتاح `'C'`.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

يجب أن تكون قيمة `rods['C']` قائمة فارغة.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
