---
id: 64df3f1011888113fbd3d81b
title: الخطوة 44
challengeType: 20
dashedName: step-44
---

# --description--

لحل اللغز باستخدام العودية، أول شيء يجب فعله هو تقسيم المشكلة الأصلية إلى مشكلات فرعية أصغر.

يمكن الحصول على التكوين النهائي مع `n` أقراص مكدسة على القضيب الثالث بترتيب تنازلي عن طريق تحريك:

- `n - 1` قرصًا من المصدر إلى القضيب المساعد
- أكبر قرص من المصدر إلى الهدف
- ثم `n - 1` قرصًا من القضيب المساعد إلى الهدف.

لذا، أول شيء يجب أن تفعله دالة `move` هو استدعاء نفسها مع `n - 1` كأول معلمة. لكن إذا حاولت فعل ذلك دون تعريف حالة أساسية، ستحصل على `RecursionError`. يحدث هذا لأن الدالة تستدعي نفسها بلا نهاية.

قبل تعليقك واستدعاء `print()`، أضف استدعاء الدالة العودية مع `n - 1` كأول معلمة وتأكد من تنفيذ جسم الدالة فقط عندما تكون قيمة `n` أكبر من صفر. اترك المعلمات الأخرى في نفس الترتيب مؤقتًا.

# --hints--

يجب أن يبدأ جسم دالة `move` بعبارة `if` تُفعّل عندما تكون قيمة `n` أكبر من صفر.

```js
({ test: () => assert.match(code, /def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:\s+if\s+n\s*>\s*0\s*:/) })
```

يجب أن تستدعي `move(n - 1, source, auxiliary, target)` داخل عبارة `if`. تذكر أن تزيح استدعاء `print()` بمقدار مسافة بادئة.

```js
const transformedCode = "\n" + code.replace(/\r/g, '');
const move = __helpers.python.getDef(transformedCode, "move");
const { function_parameters, function_body } = move;

assert.match(function_parameters, /n\s*,\s*source\s*,\s*auxiliary\s*,\s*target/);
const functionBodyIndent = function_body.match(/ +/)[0];
const re = new RegExp(`${functionBodyIndent}if\\s+n\\s*>\\s*0\\s*:\\s+^${functionBodyIndent}( +)move\\(\\s*n\\s*-\\s*1\\s*,\\s*source\\s*,\\s*auxiliary\\s*,\\s*target\\s*\\).+?^${functionBodyIndent}\\1print\\s*\\(\\s*rods\\s*,\\s*("|')\\\\n\\2\\s*\\)`, "ms");
assert.match(function_body, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
