---
id: 64df496c6a8ddf1b38db1ed6
title: الخطوة 50
challengeType: 20
dashedName: step-50
---

# --description--

في خطوة سابقة، كتبت الشفرة لنقل أكبر قرص من المشكلة الجزئية إلى العمود الهدف.

الآن، كل ما عليك فعله هو إضافة استدعاء تكراري آخر لنقل الأقراص `n - 1` التي قمت بتحريكها بالفعل. انسخ الاستدعاء التكراري الأول والصقه في نهاية كتلة `if`.

لاحظ أن معلمات الدالة ليست بالترتيب الصحيح. حاول معرفة الترتيب الصحيح.

# --hints--

يجب أن تستدعي `move(n - 1,  auxiliary, source, target)` في نهاية دالة `move` الخاصة بك.

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
