---
id: 64e34146860065146733883b
title: الخطوة 53
challengeType: 20
dashedName: step-53
---

# --description--

على الرغم من أن العودية قد تكون أحيانًا أقل سهولة في الفهم، إلا أنها تمنحك القدرة على إنشاء شفرة أكثر إيجازًا. في هذه الحالة، لا تحتاج حتى إلى التمييز بين عدد الأقراص الزوجي والفردي.

عيّن `NUMBER_OF_DISKS` إلى `5` وتحقق من الناتج.

# --hints--

يجب أن تعيّن المتغير `NUMBER_OF_DISKS` إلى 5.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 5) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
A = list(range(NUMBER_OF_DISKS, 0, -1))
B = []
C = []

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        target.append(source.pop())
        
        # display our progress
        print(A, B, C, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, A, B, C)
--fcc-editable-region--
```
