---
id: 650c6082e5586f9e3acfcd3b
title: الخطوة 52
challengeType: 20
dashedName: step-52
---

# --description--

احذف الآن قاموس `rods` وحول مفاتيحه إلى المتغيرات `A` و `B` و `C`، مع الاحتفاظ بقيمها. أعد هيكلة شفرتك لتعكس هذه التغييرات. إذا رأيت الناتج على الطرفية، فقد نفذت هذه الخطوة بشكل صحيح.

# --hints--

حوّل `rods['A']` إلى متغير اسمه `A` مع الاحتفاظ بقيمته.

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

حوّل `rods['B']` إلى متغير اسمه `B` مع الاحتفاظ بقيمته.

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

حوّل `rods['C']` إلى متغير اسمه `C` مع الاحتفاظ بقيمته.

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

عدّل استدعاء `print()` ليطبع `A` و `B` و `C` بدلاً من كائن `rods`. احتفظ بحرف السطر الجديد في استدعاء `print()`.

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

يجب أن يحتوي شفرتك على `target.append(source.pop())`.

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```

