---
id: 65782342850feb3b8d62f936
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

بما أنك ستستخدم التعبير `(i + 1) % 3` عدة مرات، من المناسب تخزينه في متغير.

فوق جملة `if` مباشرة، أعلن عن متغير `remainder` وقم بتعيين القيمة `(i + 1) % 3` لهذا المتغير.

# --hints--

يجب إعلان متغير `remainder` فوق كتلة `if` مباشرة.

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=/m) })
```

يجب أن تكون قيمة متغير `remainder` هي التعبير `(i + 1) % 3`.

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=\s*\(\s*i\s*\+\s*1\s*\)\s*%\s*3/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
