---
id: 657b667a772ed53e82962c81
title: الخطوة 54
challengeType: 20
dashedName: step-54
---

# --description--

لا يزال بإمكانك فعل شيء واحد لتحسين قابلية قراءة شفرتك.

عدّل شرط `if` ليُنفذ عندما تكون قيمة `n` أقل من أو تساوي صفرًا، وأضف عبارة `return` لإيقاف تنفيذ الدالة.

# --hints--

يجب تعديل شرط `if` الحالي للتحقق مما إذا كانت قيمة `n` أقل من أو تساوي `0`.

```js
({ test: () =>
  {
    const hanoi = __helpers.python.getDef(code, "move");
    const {function_body} = hanoi;    
    assert(function_body.match(/^\s{4}if\s+n\s*<=\s*0\s*:/m));
  }
})
```

يجب إضافة عبارة `return` في السطر الأول من كتلة `if`.

```js
({ test: () =>
  {
    const hanoi = __helpers.python.getDef(code, "move");
    const {function_body} = hanoi;    
    assert(function_body.match(/^(\s{4})if\s+n\s*<=\s*0\s*:\s*\1\1return/m));
  }
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 5
A = list(range(NUMBER_OF_DISKS, 0, -1))
B = []
C = []
--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:

--fcc-editable-region--        
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        target.append(source.pop())
        
        # display our progress
        print(A, B, C, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, A, B, C)

```
