---
id: 615f3cafd794015aa9547a6d
title: Hatua ya 12
challengeType: 0
dashedName: step-12
---

# --description--

Remember that the use of `h1`, `h2`, and similar tags determine the semantic structure of your HTML. However, you can adjust the CSS of these elements to control the visual flow and hierarchy.

Unda sheria ya `h1` na uweke sifa ya `font-weight` kuwa `800`. Hii itafanya maandishi yako ya `h1` kuwa mazito zaidi (bolder).

# --hints--

Unapaswa kuunda kichaguzi cha `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

Kichaguzi chako cha `h1` kinafaa kuwa na sifa ya `font-weight` iliyowekwa kuwa `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
