---
id: 6852f06fdde1b4634a263f03
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora mettiamo alla prova la classe `Email` creando un oggetto email e verificando che archivi correttamente le informazioni. Stamperai un paio di attributi per controllare che tutto funzioni.

Crea un oggetto email chiamato `email_obj` con `alice@example.com` come mittente, `bob@example.com` come destinatario, `Hello` come oggetto e `Hi Bob!` come corpo. Poi stampa gli attributi `sender` e `subject` come istruzioni di stampa separate per verificare che siano archiviati correttamente.

# --hints--

Dovresti creare un oggetto email chiamato `email_obj` con i parametri specificati.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

Dovresti stampare `email_obj.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

Dovresti stampare `email_obj.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
