---
id: 6852ea0c3cadb54e897acdac
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, costruirai un *simulatore di email* che simula l'invio, la ricezione e la gestione delle email tra utenti diversi. Imparerai a conoscere le classi, gli oggetti e come organizzare il codice in modo orientato agli oggetti.

Inizia creando una classe chiamata `Email` usando la parola chiave `class`.

# --hints--

Devi creare una classe chiamata `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
