---
id: 6852f11d412c29665985a915
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Ora che gli utenti possono inviare email, hanno bisogno di un modo per archiviare le email che ricevono. Ogni utente dovrebbe avere una casella di posta per raccogliere le proprie email.

Aggiungi un attributo `inbox` al metodo `User` della classe `__init__` e inizializzalo come una lista vuota `[]`. Questo archivio conterrà tutte le email che l’utente riceve.

# --hints--

Dovresti aggiungere `self.inbox = []` al metodo `User` della classe `__init__`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
