---
id: 6852f0c3fcedb964ab959921
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora verifichiamo che l'attributo `read` sia stato aggiunto correttamente alla tua classe `Email`. Poiché hai già un oggetto email dai passaggi precedenti, stampa l'attributo `read` per vedere che ora è `False` per impostazione predefinita.

# --hints--

Dovresti stampare `email_obj.read` per testare l'attributo di lettura.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
