---
id: 6852ef95ccc27660275692cd
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

La tua classe `Email` deve archiviare informazioni su ogni email quando viene creata.

All'interno della tua classe `Email`, rimuovi la parola chiave `pass` esistente e crea il metodo `__init__` che prende `self`, `sender`, `receiver`, `subject` e `body` come parametri.

# --hints--

Dovresti definire un metodo `__init__` all'interno della classe `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Il tuo metodo `__init__` dovrebbe avere i parametri: `self`, `sender`, `receiver`, `subject` e `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
