---
id: 6852f01116912861c7d1b0e6
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

All'interno del metodo `__init__` devi assegnare i valori dei parametri agli attributi dell'oggetto in modo che ogni email possa archiviare le sue informazioni.

Nel tuo metodo `__init__`, rimuovi la parola chiave `pass` e assegna i parametri a `self.sender`, `self.receiver`, `self.subject` e `self.body`.

# --hints--

Dovresti assegnare il parametro `sender` a `self.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

Dovresti assegnare il parametro `receiver` a `self.receiver`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

Dovresti assegnare il parametro `subject` a `self.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

Dovresti assegnare il parametro `body` a `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
