---
id: 6853f4ded9f84896ff182dd8
title: Passaggio 31
challengeType: 20
dashedName: step-31
---

# --description--

Dopo il controllo della casella di posta vuota, stampa il messaggio `\nYour Emails:`. Dopodiché, itera su tutte le email nella casella di posta usando un loop `for` con enumerazione. Questa è la sintassi per l'enumerazione con un numero di partenza: `enumerate(iterable, start=start_number)`.

Usa l'enumerazione per numerare le email a partire da `1`. Usa `i` (l'indice) e `email` (l'oggetto email) come variabili di iterazione.

All'interno del loop, stampa una f-string con l'indice di iterazione seguito da un `.`, poi uno spazio e la rappresentazione in stringa dell'oggetto email in questo formato:

```md
i. string_representation
```

Dove `i` è l'indice e `string_representation` è la rappresentazione dell'oggetto email.

# --hints--

Dovresti stampare l'intestazione `\nYour Emails:` esternamente all'istruzione `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

Dovresti avere un loop `for` che itera su `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

Il tuo loop `for` dovrebbe usare `i` e `email` per iterare su `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

All'interno del loop, dovresti stampare ogni riepilogo email con un prefisso numerato seguito dalla stringa formattata dal metodo `__str__`. Non dimenticare di aggiungere un `.` dopo il numero.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        
--fcc-editable-region--
```
