---
id: 6853f4584f508e95cc7a9f4d
title: Passaggio 30
challengeType: 20
dashedName: step-30
---

# --description--

Ora creerai un metodo per elencare tutte le email nella posta in arrivo. Questo metodo dovrebbe gestire il caso in cui la posta in arrivo è vuota e mostrare un elenco numerato delle email quando sono presenti.

Aggiungi un metodo chiamato `list_emails` alla tua classe `Inbox` che prende `self` come parametro. Per prima cosa, crea un'istruzione `if` per verificare se la posta in arrivo è vuota controllando la lista `emails`. Se è vuota, stampa il messaggio `Your inbox is empty.\n`. Aggiungi un'istruzione `return` per uscire dal metodo.

# --hints--

Dovresti definire un metodo chiamato `list_emails` nella classe `Inbox` che prende `self` come parametro.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

Dovresti verificare se `self.emails` è vuota usando `if not self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

Dovresti stampare il messaggio `Your inbox is empty.\n` se non ci sono email.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

Dovresti avere un'istruzione `return` alla fine dell'istruzione `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
