---
id: 6859400f8f7f872e50e1238c
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

La classe `User` ha bisogno di un metodo per inviare email ad altri utenti. Quando invii un'email, creerai un nuovo oggetto `Email` e lo aggiungerai alla casella di posta del destinatario.

Crea un metodo chiamato `send_email` nella tua classe `User`. Questo metodo dovrebbe prendere come parametri `receiver`, `subject` e `body`.

All'interno del metodo, crea un nuovo oggetto `Email` con i seguenti valori di parametro e assegnalo a una variabile chiamata `email`:

- `sender`: `self` (l'utente corrente)
- `receiver`: `receiver` (l'utente che riceve l'email)
- `subject`: `subject` (l'oggetto dell'email)
- `body`: `body` (il corpo dell'email)

# --hints--

Dovresti definire un metodo `send_email` nella classe `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

Il metodo `send_email` dovrebbe avere come parametri: `self`, `receiver`, `subject` e `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

Il metodo dovrebbe creare un oggetto `Email` chiamato `email` con i parametri corretti.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
