---
id: 68594df7779a164e6a61e8ae
title: Passaggio 32
challengeType: 20
dashedName: step-32
---

# --description--

La casella di posta ha bisogno di un metodo per leggere una email specifica. Quando un utente vuole leggere una email, indicherà quale numero di email (a partire dall’indice 0) vuole vedere, e il metodo mostrerà il contenuto completo della email.

Aggiungi un metodo chiamato `read_email` alla tua classe `Inbox` che prende un parametro `index`. Per prima cosa, verifica se la casella di posta è vuota e stampa il messaggio `Inbox is empty.\n` se lo è. Aggiungi un’istruzione `return` subito dopo per uscire dal metodo.

# --hints--

Dovresti definire un metodo chiamato `read_email` nella classe `Inbox` che prende `self` e `index` come parametri.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index")`));
    }
})
```

Dovresti controllare se `self.emails` è vuota usando `if not self.emails`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`));
    }
})
```

Dovresti stampare `Inbox is empty.\n` se non ci sono email e uscire dal metodo.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\\\n')")`));
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[1].is_equivalent("return")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    
--fcc-editable-region--
```
