---
id: 68593ea22eac7e228ed585ce
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Le email dovrebbero tenere traccia se sono state lette o meno. Aggiungi un attributo `read` al metodo `__init__` e impostalo su `False` per impostazione predefinita, dato che le nuove email iniziano come non lette.

# --hints--

Dovresti assegnare `False` a `self.read` nel metodo `__init__`.

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
