---
id: 685c1e32254e5993a5c8823e
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Ora che hai una configurazione iniziale della classe `Email`, è il momento di creare utenti che possano inviare e ricevere email. Ogni utente avrà un nome e potrà eseguire operazioni di posta elettronica.

Crea una classe `User` con un metodo `__init__` che prende `self` e `name` come parametri. All'interno della classe, assegna il parametro `name` a `self.name`.

# --hints--

Dovresti creare una classe `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

La classe `User` dovrebbe avere un metodo `__init__` che prende `self` e `name` come parametri.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

Dovresti assegnare il parametro `name` a `self.name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
