---
id: 685c1eba966c6396850736ea
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Anche se gli utenti possono inviare email e avere caselle di posta, abbiamo bisogno di una classe dedicata per gestire le operazioni sulla casella in modo efficiente.

La classe `Inbox` archivierà una lista di email e fornirà metodi per aggiungere nuove email, elencare tutte le email e gestirle.

Crea una nuova classe chiamata `Inbox` con un metodo `__init__` che inizializza una lista vuota chiamata `emails`.

# --hints--

Dovresti creare una classe `Inbox`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

La classe `Inbox` dovrebbe avere un metodo `__init__` che prende `self` come parametro.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

Dovresti inizializzare `self.emails = []` nel metodo `Inbox` della classe `__init__`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
