---
id: 685d2fd1f309b9191f257c55
title: Passaggio 40
challengeType: 20
dashedName: step-40
---

# --description--

Prima di integrare i timestamp nel nostro sistema email, esercitiamoci a lavorare con la formattazione delle date e degli orari. La funzione `datetime.datetime.now()` ci fornisce la data e l'ora correnti, e possiamo usare il metodo `strftime()` per formattarli in modi diversi.

Ecco come `strftime()` funziona con i codici di formato:

```python
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # Output: 2024-03-15 (year-month-day with - separator)
```

I codici di formato come `%Y` (anno), `%m` (mese), `%d` (giorno) indicano a `strftime()` cosa includere, e puoi aggiungere separatori come `-` tra di essi.

Alla fine del tuo codice, crea una variabile chiamata `current_time` e assegnale `datetime.datetime.now()`. Poi usa `strftime()` per stampare l'ora nel formato ore:minuti:secondi usando `:` come separatore.

Usa questi codici di formato: `%H` per le ore (formato 24 ore), `%M` per i minuti e `%S` per i secondi.

# --hints--

Dovresti creare una variabile chiamata `current_time` e assegnarle `datetime.datetime.now()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("current_time = datetime.datetime.now()")`));
    }
})
```

Dovresti usare `strftime()` per stampare l'ora nel formato `%H:%M:%S`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(current_time.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

--fcc-editable-region--

--fcc-editable-region--
```
