---
id: 68879d6de6398893be8c4b7b
title: Passaggio 25
challengeType: 20
dashedName: step-25
---

# --description--

Ora aggiungi il corpo dell'email per completare il contenuto principale. Nel metodo `display_full_email`, aggiungi un'altra istruzione di stampa nel formato `Body: body` (dove `body` è il contenuto dell'email) per mostrare il contenuto dell'email.

# --hints--

Dovresti stampare il corpo dell'email usando `f'Body: {self.body}'`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Body: {self.body}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

```
