---
id: 68875554c6c32de788afc64e
title: Passaggio 21
challengeType: 20
dashedName: step-21
---

# --description--

Gli utenti hanno bisogno di un modo per leggere correttamente le loro email.

Aggiungi un metodo `display_full_email` alla classe `Email` che prenda solo `self` come parametro. All'interno di questo metodo, chiama il metodo `mark_as_read()` per segnare l'email come letta quando qualcuno la visualizza.

Questo metodo verrà usato per mostrare il contenuto completo dell'email in modo ben formattato.

# --hints--

Dovresti definire un metodo `display_full_email` nella classe `Email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

Il metodo `display_full_email` dovrebbe prendere solo `self` come parametro.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

Il metodo dovrebbe chiamare `self.mark_as_read()` per segnare l'email come letta.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
