---
id: 688754d6d59578df1313cb3c
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Ora è il momento di collegare le classi User e Inbox così che le email possano essere effettivamente consegnate!

Gli utenti devono avere veri oggetti Inbox invece di semplici liste, e il metodo `send_email` dovrebbe consegnare le email nella inbox del destinatario.

Aggiorna la classe `User` `__init__` in modo che ogni utente abbia un vero oggetto Inbox invece della lista vuota creata in precedenza.

# --hints--

Dovresti aggiornare la classe User per inizializzare `self.inbox = Inbox()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
