---
id: 68b9350183ff05bd106c79e4
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Ora, testa il metodo `mark_as_read` che hai creato nella classe `Email`. Usa il metodo sull'oggetto `Email` esistente `email_obj` per cambiare lo stato. Dopo aver segnato l'email come letta, stampa l'attributo `read` di `email_obj` per confermare la modifica.

# --hints--

Dovresti chiamare il metodo `mark_as_read` su `email_obj` per segnare l'email come letta.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

Dovresti stampare l'attributo `read` di `email_obj` dopo aver chiamato il metodo `mark_as_read`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
