---
id: 68b9366fae8b13c1278761b1
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

La tua casella di posta ha bisogno di un modo per ricevere nuove email. Quando qualcuno invia un’email a un utente, questa dovrebbe essere aggiunta alla sua casella di posta.

Aggiungi un metodo chiamato `receive_email` alla tua classe `Inbox` che prende `self` e un oggetto email `email` come parametri. All’interno del corpo del metodo, aggiungi `email` alla lista `emails` usando il metodo `append`.

# --hints--

Dovresti definire un metodo chiamato `receive_email` nella classe `Inbox` che prende `self` e `email` come parametri.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

Dovresti aggiungere `email` a `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
