---
id: 68b96d6a290d52d5e57a7485
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Rimuovi `email_obj` e le istruzioni di stampa successive.

# --hints--

Non dovresti avere la riga di creazione `email_obj` nel tuo codice.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`));
    }
})
```

Non dovresti avere `print(email_obj.sender)` nel tuo codice.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.sender)")`));
    }
})
```

Non dovresti avere `print(email_obj.subject)` nel tuo codice.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.subject)")`));
    }
})
```

Non dovresti avere `print(email_obj.read)` nel tuo codice.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

Non dovresti avere `email_obj.mark_as_read()` nel tuo codice.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
email_obj.mark_as_read()
print(email_obj.read)
--fcc-editable-region--
```
