---
id: 68b930ee0e7b8fb81895e59a
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Ora aggiungerai un metodo per segnare un'email come letta. Quando qualcuno apre un'email, vuoi cambiare il suo stato da non letta a letta. Questo metodo sarà semplice: deve solo impostare l'attributo `read` su `True`.

Aggiungi un metodo chiamato `mark_as_read` alla tua classe `Email`. Questo metodo dovrebbe prendere come parametro solo `self` e impostare l'attributo `read` su `True`.

# --hints--

Dovresti definire un metodo chiamato `mark_as_read` nella classe `Email`.
che dovrebbe prendere come parametro solo `self`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

Il metodo dovrebbe impostare `self.read` su `True`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
