---
id: 68be6af1f43cc44541debe41
title: Passaggio 37
challengeType: 20
dashedName: step-37
---

# --description--

Dovresti gestire il caso in cui l'utente tenta di eliminare un'email usando un indice non valido proprio come hai fatto nel metodo `read_email`.

All'interno del metodo `delete_email`:

- Converti il parametro `index`, che parte da 1, in un indice che parte da 0 sottraendo 1 e archivialo in una variabile chiamata `actual_index`.
- Crea un'istruzione if per verificare se `actual_index` è minore di 0 oppure maggiore o uguale alla lunghezza della lista `self.emails`. Se è così, stampa il messaggio `Invalid email number.\n` e usa `return` per uscire dal metodo.

# --hints--


Dovresti sottrarre 1 dal parametro `index` e archiviarlo in una variabile chiamata `actual_index`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_body().find_variable("actual_index").is_equivalent("actual_index = index - 1")`))
    }
})
```

Dovresti creare un'istruzione if per verificare se `actual_index` è minore di 0 oppure maggiore o uguale alla lunghezza della lista `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_conditions()[0].is_equivalent("actual_index < 0 or actual_index >= len(self.emails)")`))
    }
})
```

All'interno dell'istruzione `if`, dovresti stampare il messaggio `Invalid email number.\n` e usare `return` per uscire dal metodo.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[0].is_equivalent("print('Invalid email number.\\\\n')")`))
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()
            
--fcc-editable-region--
    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        
--fcc-editable-region--

```
