---
id: 594810f028c0303b75339ace
title: Fábrica de acumuladores
challengeType: 1
forumTopicId: 302222
dashedName: accumulator-factory
---

# --description--

Un problema planteado por Paul Graham es el de crear una función que tome un solo argumento (numérico) y que devuelva otra función que sea un acumulador. La función acumuladora devuelta, a su vez, también toma un solo argumento numérico y devuelve la suma de todos los valores numéricos pasados hasta ahora a ese acumulador (incluyendo el valor inicial pasado cuando se creó el acumulador).

# --instructions--

Crea una función que tome un número $n$ y genere funciones acumuladoras que devuelvan la suma de todos los números que se les hayan pasado.

**Reglas:**

No uses variables globales.

**Consejo:**

Los closures guardan el estado externo.

# --hints--

`accumulator` debe ser una función.

```js
assert(typeof accumulator === 'function');
```

`accumulator(0)` debería devolver una función.

```js
assert(typeof accumulator(0) === 'function');
```

`accumulator(0)(2)` debería devolver un número.

```js
assert(typeof accumulator(0)(2) === 'number');
```

Pasar los valores 3, -4, 1.5 y 5 debería devolver 5.5.

```js
assert(testFn(5) === 5.5);
```

# --seed--

## --after-user-code--

```js
const testFn = typeof accumulator(3) === 'function' && accumulator(3);
if (testFn) {
  testFn(-4);
  testFn(1.5);
}
```

## --seed-contents--

```js
function accumulator(sum) {

}
```

# --solutions--

```js
function accumulator(sum) {
  return function(n) {
    return sum += n;
  };
}
```
