---
id: 594810f028c0303b75339acf
title: Función de Ackermann
challengeType: 1
forumTopicId: 302223
dashedName: ackermann-function
---

# --description--

La función Ackermann es un ejemplo clásico de una función recursiva, notable especialmente porque no es una función recursiva primitiva. Crece muy rápido en valor, al igual que el tamaño de su árbol de llamadas.

La función de Ackermann generalmente se define de la siguiente manera:

$A(m, n) = \\begin{cases} n+1 & \\mbox{if } m = 0 \\\\ A(m-1, 1) & \\mbox{if } m > 0 \\mbox{ and } n = 0 \\\\ A(m-1, A(m, n-1)) & \\mbox{if } m > 0 \\mbox{ and } n > 0. \\end{cases}$

Sus argumentos nunca son negativos y siempre termina.

# --instructions--

Escribe una función que devuelva el valor de $A(m, n)$. Se prefiere precisión arbitraria (ya que la función crece muy rápido), pero no es obligatorio.

# --hints--

`ack` debería ser una función.

```js
assert(typeof ack === 'function');
```

`ack(0, 0)` debería devolver 1.

```js
assert(ack(0, 0) === 1);
```

`ack(1, 1)` debería devolver 3.

```js
assert(ack(1, 1) === 3);
```

`ack(2, 5)` debería devolver 13.

```js
assert(ack(2, 5) === 13);
```

`ack(3, 3)` debería devolver 61.

```js
assert(ack(3, 3) === 61);
```

# --seed--

## --seed-contents--

```js
function ack(m, n) {

}
```

# --solutions--

```js
function ack(m, n) {
  return m === 0 ? n + 1 : ack(m - 1, n === 0 ? 1 : ack(m, n - 1));
}
```
