---
id: 594810f028c0303b75339ad8
title: Matriz zig-zag
challengeType: 1
forumTopicId: 302348
dashedName: zig-zag-matrix
---

# --description--

Un arreglo 'zig-zag' es una estructura cuadrada de los primeros $N^2$ enteros, donde los números aumentan secuencialmente mientras haces zig-zag a lo largo de las anti-diagonales del arreglo.

Por ejemplo, para la entrada `5`, se debe producir el siguiente resultado:

<pre>
 0  1  5  6 14
 2  4  7 13 15
 3  8 12 16 21
 9 11 17 20 22
10 18 19 23 24
</pre>

# --instructions--

Escribe una función que reciba el tamaño de la matriz zig-zag y devuelva la matriz correspondiente como un arreglo bidimensional.

# --hints--

ZigZagMatrix debería ser una función.

```js
assert.equal(typeof ZigZagMatrix, 'function');
```

ZigZagMatrix debería devolver un array.

```js
assert.equal(typeof ZigZagMatrix(1), 'object');
```

ZigZagMatrix debería devolver un arreglo de arreglos anidados.

```js
assert.equal(typeof ZigZagMatrix(1)[0], 'object');
```

ZigZagMatrix(1) debería devolver \[[0]].

```js
assert.deepEqual(ZigZagMatrix(1), zm1);
```

ZigZagMatrix(2) debe devolver \[[0, 1], [2, 3]].

```js
assert.deepEqual(ZigZagMatrix(2), zm2);
```

ZigZagMatrix(5) debería devolver la matriz especificada.

```js
assert.deepEqual(ZigZagMatrix(5), zm5);
```

# --seed--

## --after-user-code--

```js
const zm1 = [[0]];
const zm2 = [[0, 1], [2, 3]];
const zm5 = [
  [0, 1, 5, 6, 14],
  [2, 4, 7, 13, 15],
  [3, 8, 12, 16, 21],
  [9, 11, 17, 20, 22],
  [10, 18, 19, 23, 24]
];
```

## --seed-contents--

```js
function ZigZagMatrix(n) {

  return [[], []];
}
```

# --solutions--

```js
function ZigZagMatrix(n) {
  const mtx = [];
  for (let i = 0; i < n; i++) {
    mtx[i] = [];
  }

  let i = 1;
  let j = 1;
  for (let e = 0; e < n * n; e++) {
    mtx[i - 1][j - 1] = e;
    if ((i + j) % 2 === 0) {
      // Even stripes
      if (j < n) j++;
      else i += 2;
      if (i > 1) i--;
    } else {
      // Odd stripes
      if (i < n) i++;
      else j += 2;
      if (j > 1) j--;
    }
  }
  return mtx;
}
```
