---
id: 594dc6c729e5700999302b45
title: Corchetes balanceados
challengeType: 1
forumTopicId: 302230
dashedName: balanced-brackets
---

# --description--

Determina si una cadena generada de corchetes está balanceada; es decir, si consiste completamente en pares de corchetes de apertura/cierre (en ese orden), ninguno de los cuales está mal anidado.

**Ejemplos:**
| Entrada                     | Resultado |
| ------------------------- | ------ |
| <code>\[]</code>          | true   |
| <code>]\[</code>          | false  |
| <code>[][]</code>         | true   |
| <code>]\[]</code>         | false  |
| <code>\[]]\[\[]</code>    | false  |
| <code>\[\[\[\[]]]]</code> | true   |

# --hints--

`isBalanced` debería ser una función.

```js
assert(typeof isBalanced === 'function');
```

`isBalanced("[]")` debería devolver true.

```js
assert(isBalanced(testCases[0]));
```

`isBalanced("]][[[][][][]][")` debería devolver false.

```js
assert(!isBalanced(testCases[1]));
```

`isBalanced("[][[[[][][[[]]]]]]")` debería devolver true.

```js
assert(isBalanced(testCases[2]));
```

`isBalanced("][")` debería devolver false.

```js
assert(!isBalanced(testCases[3]));
```

`isBalanced("[[[]]]][[]")` debería devolver false.

```js
assert(!isBalanced(testCases[4]));
```

`isBalanced("][[]")` debería devolver false.

```js
assert(!isBalanced(testCases[5]));
```

`isBalanced("][[][]][[[]]")` debería devolver false.

```js
assert(!isBalanced(testCases[6]));
```

`isBalanced("[[][]]][")` debería devolver false.

```js
assert(!isBalanced(testCases[7]));
```

`isBalanced("[[[]]][[]]]][][[")` debería devolver false.

```js
assert(!isBalanced(testCases[8]));
```

`isBalanced("[]][[]]][[[[][]]")` debería devolver false.

```js
assert(!isBalanced(testCases[9]));
```

`isBalanced("][]][[][")` debería devolver false.

```js
assert(!isBalanced(testCases[10]));
```

`isBalanced("[[]][[][]]")` debería devolver true.

```js
assert(isBalanced(testCases[11]));
```

`isBalanced("[[]]")` debería devolver true.

```js
assert(isBalanced(testCases[12]));
```

`isBalanced("]][]][[]][[[")` debería devolver false.

```js
assert(!isBalanced(testCases[13]));
```

`isBalanced("][]][][[")` debería devolver false.

```js
assert(!isBalanced(testCases[14]));
```

`isBalanced("][][")` debería devolver false.

```js
assert(!isBalanced(testCases[15]));
```

`isBalanced("[]]]")` debería devolver false.

```js
assert(!isBalanced(testCases[16]));
```

`isBalanced("")` debería devolver true.

```js
assert(isBalanced(testCases[17]));
```

# --seed--

## --after-user-code--

```js
const testCases = [
  '[]',
  ']][[[][][][]][',
  '[][[[[][][[[]]]]]]',
  '][',
  '[[[]]]][[]',
  '][[]',
  '][[][]][[[]]',
  '[[][]]][',
  '[[[]]][[]]]][][[',
  '[]][[]]][[[[][]]',
  '][]][[][',
  '[[]][[][]]',
  '[[]]',
  ']][]][[]][[[',
  '][]][][[',
  '][][',
  '[]]]',
  ''
];
```

## --seed-contents--

```js
function isBalanced(str) {

  return true;
}
```

# --solutions--

```js
function isBalanced(str) {
  if (str === '') return true;
  let a = str;
  let b;
  do {
    b = a;
    a = a.replace(/\[\]/g, '');
  } while (a !== b);
  return !a;
}
```
