---
id: 594d8d0ab97724821379b1e6
title: Promedios/Moda
challengeType: 1
forumTopicId: 302226
dashedName: averagesmode
---

# --description--

Escribe una función `mode` para encontrar el valor que aparece más en un arreglo.

Se puede ignorar el caso en que la colección esté vacía. Debes tener cuidado al manejar el caso en que el modo no sea único.

Si no es apropiado o posible soportar una colección general, usa un vector (array), si es posible. Si no es apropiado o posible soportar un tipo de valor no especificado, usa enteros.

# --hints--

`mode` debe ser una función.

```js
assert(typeof mode === 'function');
```

`mode([1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17])` debería ser igual a `[6]`

```js
assert.deepEqual(mode(arr1), [6]);
```

`mode([1, 2, 4, 4, 1])` debería ser igual a `[1, 4]`.

```js
assert.deepEqual(mode(arr2).sort(), [1, 4]);
```

# --seed--

## --after-user-code--

```js
const arr1 = [1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17];
const arr2 = [1, 2, 4, 4, 1];
```

## --seed-contents--

```js
function mode(arr) {

  return true;
}
```

# --solutions--

```js
function mode(arr) {
  const counter = {};
  let result = [];
  let max = 0;
  // for (const i in arr) {
  arr.forEach(el => {
    if (!(el in counter)) {
      counter[el] = 0;
    }
    counter[el]++;

    if (counter[el] === max) {
      result.push(el);
    }
    else if (counter[el] > max) {
      max = counter[el];
      result = [el];
    }
  });
  return result;
}
```
