---
id: 594da033de4190850b893874
title: Promedios/Raíz cuadrada media
challengeType: 1
forumTopicId: 302228
dashedName: averagesroot-mean-square
---

# --description--

Calcula la raíz cuadrada media (RMS) de los números del 1 al 10 inclusive.

El <abbr title="Root mean square">RMS</abbr> se calcula tomando la raíz cuadrada de la media de los cuadrados de los números, dada por la ecuación:

$$x\_{\\mathrm{rms}} = \\sqrt {{{x_1}^2 + {x_2}^2 + \\cdots + {x_n}^2} \\over n}. $$

# --hints--

`rms` debería ser una función.

```js
assert(typeof rms === 'function');
```

`rms([1, 2, 3, 4, 5, 6, 7, 8, 9, 10])` debería ser igual a `6.2048368229954285`.

```js
assert.equal(rms(arr1), answer1);
```

# --seed--

## --after-user-code--

```js
const arr1 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const answer1 = 6.2048368229954285;
```

## --seed-contents--

```js
function rms(arr) {

}
```

# --solutions--

```js
function rms(arr) {
  const sumOfSquares = arr.reduce((s, x) => s + x * x, 0);
  return Math.sqrt(sumOfSquares / arr.length);
}
```
