---
id: 5951ed8945deab770972ae56
title: Torres de Hanoi
challengeType: 1
forumTopicId: 302341
dashedName: towers-of-hanoi
---

# --description--

Resuelve el problema de las Torres de Hanoi. El número de objetos se dará como el primer parámetro, seguido de las cadenas usadas para identificar cada pila de objetos. Crea un arreglo anidado que contenga la lista de movimientos, `["source", "destination"]`.
 
Por ejemplo, los parámetros `(4, 'A', 'B', 'C')` resultarán en un arreglo anidado de movimientos `[['A', 'C'], ['B', 'A']]`, indicando que el primer movimiento fue mover un objeto de la pila `A` a `C` y el segundo movimiento fue mover un objeto de la pila `B` a `A`.

# --instructions--

Escribe una función que devuelva los movimientos para apilar los objetos en un array anidado.

# --hints--

`towerOfHanoi` debe ser una función.

```js
assert(typeof towerOfHanoi === 'function');
```

`towerOfHanoi(3, ...)` debe devolver 7 movimientos.

```js
assert(res3.length === 7);
```

`towerOfHanoi(3, 'A', 'B', 'C')` debería devolver `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B']]`.

```js
assert.deepEqual(towerOfHanoi(3, 'A', 'B', 'C'), res3Moves);
```

El movimiento 10 de `towerOfHanoi(5, "X", "Y", "Z")` debe ser de Y a X.

```js
assert.deepEqual(res5[9], ['Y', 'X']);
```

`towerOfHanoi(7, 'A', 'B', 'C')` los primeros diez movimientos deben ser `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B'], ['A','C'], ['B','C'], ['B','A']]`

```js
assert.deepEqual(towerOfHanoi(7, 'A', 'B', 'C').slice(0, 10), res7First10Moves);
```

# --seed--

## --after-user-code--

```js
const res3 = towerOfHanoi(3, 'A', 'B', 'C');
const res3Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B']];
const res5 = towerOfHanoi(5, 'X', 'Y', 'Z');
const res7First10Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B'], ['A', 'C'], ['B', 'C'], ['B', 'A']];
```

## --seed-contents--

```js
function towerOfHanoi(n, a, b, c) {

  return [[]];
}
```

# --solutions--

```js
function towerOfHanoi(n, a, b, c) {
  const res = [];
  towerOfHanoiHelper(n, a, c, b, res);
  return res;
}

function towerOfHanoiHelper(n, a, b, c, res) {
  if (n > 0) {
    towerOfHanoiHelper(n - 1, a, c, b, res);
    res.push([a, c]);
    towerOfHanoiHelper(n - 1, b, a, c, res);
  }
}
```
