---
id: 595608ff8bcd7a50bd490181
title: Secuencia Hailstone
challengeType: 1
forumTopicId: 302279
dashedName: hailstone-sequence
---

# --description--

La secuencia Hailstone de números puede generarse a partir de un entero positivo inicial, `n`, mediante:

- Si `n` es `1` entonces la secuencia termina
- Si `n` es `even` entonces el siguiente `n` de la secuencia `= n/2`
- Si `n` es `odd` entonces el siguiente `n` de la secuencia `= (3 * n) + 1`

La conjetura (no demostrada) de Collatz es que la secuencia hailstone para cualquier número inicial siempre termina.

La secuencia hailstone también es conocida como números hailstone (porque los valores suelen estar sujetos a múltiples descensos y ascensos como los granizos en una nube), o como la secuencia Collatz.

# --instructions--

1. Crea una rutina para generar la secuencia hailstone para un número
2. Tu función debe devolver un arreglo con el número menor que `limit` que tenga la secuencia de hailstone más larga y la longitud de esa secuencia. (¡Pero no muestres la secuencia real!)

# --hints--

`hailstoneSequence` debería ser una función.

```js
assert(typeof hailstoneSequence === 'function');
```

`hailstoneSequence(30)` debería devolver un arreglo.

```js
assert(Array.isArray(hailstoneSequence(30)));
```

`hailstoneSequence(30)` debería devolver `[27, 112]`.

```js
assert.deepEqual(hailstoneSequence(30), [27, 112]);
```

`hailstoneSequence(50000)` debería devolver `[35655, 324]`.

```js
assert.deepEqual(hailstoneSequence(50000), [35655, 324]);
```

`hailstoneSequence(100000)` debería devolver `[77031, 351]`.

```js
assert.deepEqual(hailstoneSequence(100000), [77031, 351]);
```

# --seed--

## --seed-contents--

```js
function hailstoneSequence(limit) {
  const res = [];


  return res;
}
```

# --solutions--

```js
function hailstoneSequence (limit) {
  function hailstone(n) {
    const seq = [n];
    while (n > 1) {
      n = n % 2 ? 3 * n + 1 : n / 2;
      seq.push(n);
    }
    return seq;
  }

  let n = 0;
  let max = 0;
  for (let i = limit; --i;) {
    const seq = hailstone(i);
    const sLen = seq.length;

    if (sLen > max) {
      n = i;
      max = sLen;
    }
  }

  return [n, max];
}
```
