---
id: 595668ca4cfe1af2fb9818d4
title: Serie de Harshad o Niven
challengeType: 1
forumTopicId: 302281
dashedName: harshad-or-niven-series
---

# --description--

Los números Harshad o Niven son enteros positivos ≥ 1 que son divisibles por la suma de sus dígitos.

Por ejemplo, `42` es un número Harshad ya que `42` es divisible por `(4 + 2)` sin residuo.

Asume que la serie está definida como los números en orden ascendente.

# --instructions--

Implementa una función para generar miembros sucesivos de la secuencia Harshad.

Úsalo para devolver un arreglo con diez miembros de la secuencia, comenzando con el primer número Harshad mayor que `n`.

# --hints--

`isHarshadOrNiven` debería ser una función.

```js
assert(typeof isHarshadOrNiven === 'function');
```

`isHarshadOrNiven(10)` debería devolver `[12, 18, 20, 21, 24, 27, 30, 36, 40, 42]`

```js
assert.deepEqual(isHarshadOrNiven(10), [12, 18, 20, 21, 24, 27, 30, 36, 40, 42]);
```

`isHarshadOrNiven(400)` debería devolver `[402, 405, 407, 408, 410, 414, 420, 423, 432, 440]`

```js
assert.deepEqual(isHarshadOrNiven(400), [402, 405, 407, 408, 410, 414, 420, 423, 432, 440]);
```

`isHarshadOrNiven(1000)` debería devolver `[1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]`

```js
assert.deepEqual(isHarshadOrNiven(1000), [1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]);
```

# --seed--

## --seed-contents--

```js
function isHarshadOrNiven(n) {
  const res = [];

  return res;
}
```

# --solutions--

```js
function isHarshadOrNiven(n) {
  function isHarshad(num) {
    let s = 0;
    const nStr = num.toString();
    for (let i = 0; i < nStr.length; ++i) {
      s += parseInt(nStr.charAt(i), 10);
    }
    return n % s === 0;
  }

  const res = [];
  let count = 0;

  while (count < 10) {
    n++;
    if (isHarshad(n)) {
      count++;
      res.push(n);
    }
  }

  return res;
}
```
