---
id: 59622f89e4e137560018a40e
title: Secuencias Figura-Figura de Hofstadter
challengeType: 1
forumTopicId: 302286
dashedName: hofstadter-figure-figure-sequences
---

# --description--

Las secuencias Hofstadter Figure-Figure $R_n$ y $S_n$ están dadas por

$R_1 = 1\\ ;\\ S_1 = 2 \\\\R_n = R_{n-1} + S_{n-1}, \\quad n>1.$

Específicamente, la secuencia $R_n$ contiene los valores

<pre>1, 3, 7, 12, 18, ...</pre>

y la secuencia $S_n$ contiene los valores

<pre>2, 4, 5, 6, 8, ...</pre>

La secuencia $R_n$ se define por la relación de recurrencia $R_n = R_{n-1} + S_{n-1}$, mientras que $S_n$ se define como una secuencia de enteros positivos que no están incluidos en la secuencia $R_n$.

# --instructions--

Crea dos funciones llamadas `ffr` y `ffs` que devuelvan `R(n)` o `S(n)`, respectivamente, para cualquier índice `n`. Ten en cuenta que las secuencias Hofstadter Figure-Figure están indexadas desde 1, con $R_1 = 1$ y $S_1 = 2$.

No se debe asumir un valor máximo para `n`.

**Referencias**

 <p>Rosetta: <a href='https://rosettacode.org/wiki/Hofstadter_Figure-Figure_sequences' target='_blank'>Hofstadter Figure-Figure sequences</a></p>.


# --hints--

`ffr` debería ser una función.

```js
assert(typeof ffr === 'function');
```

`ffs` debería ser una función.

```js
assert(typeof ffs === 'function');
```

`ffr` debería devolver un entero.

```js
assert(Number.isInteger(ffr(1)));
```

`ffs` debería devolver un entero.

```js
assert(Number.isInteger(ffs(1)));
```

`ffr(10)` debería devolver `69`

```js
assert.equal(ffr(ffrParamRes[0][0]), ffrParamRes[0][1]);
```

`ffr(50)` debería devolver `1509`

```js
assert.equal(ffr(ffrParamRes[1][0]), ffrParamRes[1][1]);
```

`ffr(100)` debería devolver `5764`

```js
assert.equal(ffr(ffrParamRes[2][0]), ffrParamRes[2][1]);
```

`ffr(1000)` debería devolver `526334`

```js
assert.equal(ffr(ffrParamRes[3][0]), ffrParamRes[3][1]);
```

`ffs(10)` debería devolver `14`

```js
assert.equal(ffs(ffsParamRes[0][0]), ffsParamRes[0][1]);
```

`ffs(50)` debería devolver `59`

```js
assert.equal(ffs(ffsParamRes[1][0]), ffsParamRes[1][1]);
```

`ffs(100)` debería devolver `112`

```js
assert.equal(ffs(ffsParamRes[2][0]), ffsParamRes[2][1]);
```

`ffs(1000)` debería devolver `1041`

```js
assert.equal(ffs(ffsParamRes[3][0]), ffsParamRes[3][1]);
```

# --seed--

## --after-user-code--

```js
const ffrParamRes = [[10, 69], [50, 1509], [100, 5764], [1000, 526334]];
const ffsParamRes = [[10, 14], [50, 59], [100, 112], [1000, 1041]];
```

## --seed-contents--

```js
function ffr(n) {
  return n;
}

function ffs(n) {
  return n;
}
```

# --solutions--

```js
const R = [null, 1];
const S = [null, 2];

function extendSequences (n) {
  let current = Math.max(R[R.length - 1], S[S.length - 1]);
  let i;
  while (R.length <= n || S.length <= n) {
    i = Math.min(R.length, S.length) - 1;
    current += 1;
    if (current === R[i] + S[i]) {
      R.push(current);
    } else {
      S.push(current);
    }
  }
}

function ffr (n) {
  extendSequences(n);
  return R[n];
}

function ffs (n) {
  extendSequences(n);
  return S[n];
}
```
