---
id: 5966f99c45e8976909a85575
title: Día de la semana
challengeType: 1
forumTopicId: 302245
dashedName: day-of-the-week
---

# --description--

Una empresa decide que cada vez que la Navidad caiga en domingo, dará a sus trabajadores todos los días extra pagados de vacaciones para que, junto con cualquier día festivo, los trabajadores no tengan que trabajar la semana siguiente (entre el 25 de diciembre y el primero de enero).

# --instructions--

Escribe una función que tome un año de inicio y un año de fin y devuelva un arreglo con todos los años en los que el 25 de diciembre será domingo.

# --hints--

`findXmasSunday` debería ser una función.

```js
assert(typeof findXmasSunday === 'function');
```

`findXmasSunday(2000, 2100)` debería devolver un arreglo.

```js
assert(typeof findXmasSunday(2000, 2100) === 'object');
```

`findXmasSunday(1970, 2017)` debería devolver `[1977, 1983, 1988, 1994, 2005, 2011, 2016]`

```js
assert.deepEqual(findXmasSunday(1970, 2017), firstSolution);
```

`findXmasSunday(2008, 2121)` debería devolver `[2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118]`

```js
assert.deepEqual(findXmasSunday(2008, 2121), secondSolution);
```

# --seed--

## --after-user-code--

```js
const firstSolution = [1977, 1983, 1988, 1994, 2005, 2011, 2016];
const secondSolution = [2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118];
```

## --seed-contents--

```js
function findXmasSunday(start, end) {

  return true;
}
```

# --solutions--

```js
function findXmasSunday(start, end) {
  const xmasSunday = [];
  for (let year = start; year <= end; year++) {
    const xmas = new Date(year, 11, 25);
    if (xmas.getDay() === 0) {
      xmasSunday.push(year);
    }
  }
  return xmasSunday;
}
```
