---
id: 59669d08d75b60482359409f
title: Formato de fecha
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

Devuelve un arreglo con dos cadenas de fecha de la fecha actual con las siguientes especificaciones:

- El orden de la fecha en la primera cadena debe ser el número de año, número de mes y número de día separados por guiones (`-`).
- El año de la primera cadena debe tener una longitud de cuatro dígitos.
- El mes y el día de la primera cadena no deben contener ceros a la izquierda.
- Los nombres de los días de la semana y los meses en la segunda cadena no deben abreviarse.
- El día de la segunda cadena no debe contener ceros a la izquierda.

Ejemplos de resultados:

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

`getDateFormats` debería ser una función.

```js
assert(typeof getDateFormats === 'function');
```

`getDateFormats` debería devolver un objeto.

```js
assert(typeof getDateFormats() === 'object');
```

`getDateFormats` debería devolver un arreglo con 2 elementos.

```js
assert(getDateFormats().length === 2);
```

`getDateFormats` debería devolver la fecha correcta en el formato adecuado

```js
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --after-user-code--

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
```

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
