---
id: 59667989bf71cf555dd5d2ff
title: S-Expresiones
challengeType: 1
forumTopicId: 302303
dashedName: s-expressions
---

# --description--

<a href="https://rosettacode.org/wiki/S-expressions" target="_blank" rel="noopener noreferrer nofollow">S-Expressions</a> son una forma conveniente de analizar y almacenar datos.

# --instructions--

Escribe un lector/analizador simple para S-Expressions que maneje cadenas entre comillas y sin comillas, enteros y flotantes.

La función debe leer una sola pero anidada S-Expression desde una cadena y devolverla como un arreglo (anidado).

Los saltos de línea y otros espacios en blanco pueden ser ignorados a menos que estén contenidos dentro de una cadena entre comillas.

"`()`" dentro de cadenas entre comillas no se interpretan, sino que se tratan como parte de la cadena.

Manejar comillas escapadas dentro de una cadena es opcional; por lo tanto, "`(foo"bar)`" puede ser tratado como una cadena "`foo"bar`", o como un error.

Para esto, el lector no necesita reconocer `\` para escapar, pero debe, además, reconocer números si el lenguaje tiene tipos de datos apropiados.

Ten en cuenta que, con la excepción de `()"` (`\` si se admite escape) y los espacios en blanco, no hay caracteres especiales. Cualquier otra cosa está permitida sin comillas.

El lector debería poder leer la siguiente entrada

<pre>((data "quoted data" 123 4.5)
(data (!@# (4.5) "(more" "data)")))
</pre>

y conviértelo en una estructura de datos nativa.

# --hints--

`parseSexpr` debería ser una función.

```js
assert(typeof parseSexpr === 'function');
```

`parseSexpr('(data1 data2 data3)')` debería devolver `['data1', 'data2', 'data3']`

```js
assert.deepEqual(parseSexpr(simpleSExpr), simpleSolution);
```

`parseSexpr('((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))')` debería devolver `[['data', '"quoted data"', 123, 4.5], ['data', ['!@#', [4.5], '"(more"', '"data)"']]]`.

```js
assert.deepEqual(parseSexpr(basicSExpr), basicSolution);
```

# --seed--

## --after-user-code--

```js
const simpleSExpr = '(data1 data2 data3)';
const simpleSolution = ['data1', 'data2', 'data3'];

const basicSExpr = '((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))';
const basicSolution = [["data","\"quoted data\"",123,4.5],["data",["!@#",[4.5],"\"(more\"","\"data)\""]]];
```

## --seed-contents--

```js
function parseSexpr(str) {

  return true;
}
```

# --solutions--

```js
function parseSexpr(str) {
  const t = str.match(/\s*("[^"]*"|\(|\)|"|[^\s()"]+)/g);
  for (var o, c = 0, i = t.length - 1; i >= 0; i--) {
    var n,
      ti = t[i].trim();
    if (ti == '"') return;
    else if (ti == '(') t[i] = '[', c += 1;
    else if (ti == ')') t[i] = ']', c -= 1;
    else if ((n = +ti) == ti) t[i] = n;
    else t[i] = `'${ti.replace('\'', '\\\'')}'`;
    if (i > 0 && ti != ']' && t[i - 1].trim() != '(') t.splice(i, 0, ',');
    if (!c) if (!o) o = true; else return;
  }
  return c ? undefined : eval(t.join(''));
}
```
