---
id: 596a8888ab7c01048de257d5
title: Copia profunda
challengeType: 1
forumTopicId: 302247
dashedName: deepcopy
---

# --description--

Escribe una función que devuelva una copia profunda de un objeto dado. La copia no debe ser el mismo objeto que fue dado.

Esta tarea no evaluará:

<ul>
  <li>Objects with properties that are functions</li>
  <li>Date objects or object with properties that are Date objects</li>
  <li>RegEx or object with properties that are RegEx objects</li>
  <li>Prototype copying</li>
</ul>

# --hints--

`deepcopy` debería ser una función.

```js
assert(typeof deepcopy === 'function');
```

`deepcopy({test: "test"})` debería devolver un objeto.

```js
assert(typeof deepcopy(obj1) === 'object');
```

`deepcopy` no debería devolver el mismo objeto que se proporcionó.

```js
assert(deepcopy(obj2) != obj2);
```

Cuando se pasa un objeto que contiene un arreglo, `deepcopy` debería devolver una copia profunda del objeto.

```js
assert.deepEqual(deepcopy(obj2), obj2);
```

Cuando se pasa un objeto que contiene otro objeto, `deepcopy` debería devolver una copia profunda del objeto.

```js
assert.deepEqual(deepcopy(obj3), obj3);
```

# --seed--

## --after-user-code--

```js
const obj1 = { test: 'test' };
const obj2 = {
  t: 'test',
  a: ['an', 'array']
};
const obj3 = {
  t: 'try',
  o: obj2
};
```

## --seed-contents--

```js
function deepcopy(obj) {

  return true;
}
```

# --solutions--

```js
function deepcopy(obj) {
  return JSON.parse(JSON.stringify(obj));
}
```
