---
id: 596e457071c35c882915b3e4
title: Compara una lista de cadenas
challengeType: 1
forumTopicId: 302235
dashedName: compare-a-list-of-strings
---

# --description--

Una lista es un conjunto ordenado de valores que puede contener duplicados. Aquí tienes un ejemplo:

```js
const list = [['AA',  'BB', 'CC'], ['AA', 'ACB', 'AA'], [], ['AA']];
````

Dada una lista de una cantidad arbitraria de cadenas, implementa una función para cada una de las siguientes condiciones:

<ul>
  <li>test if they are all lexically equal</li>
  <li>test if every string is lexically less than the one after it  (i.e. whether the list is in strict ascending order)</li>
</ul>

# --hints--

`allEqual` debería ser una función.

```js
assert(typeof allEqual === 'function');
```

`azSorted` debe ser una función.

```js
assert(typeof azSorted === 'function');
```

`allEqual(["AA", "AA", "AA", "AA"])` debería devolver true.

```js
assert(allEqual(testCases[0]));
```

`azSorted(["AA", "AA", "AA", "AA"])` debería devolver false.

```js
assert(!azSorted(testCases[0]));
```

`allEqual(["AA", "ACB", "BB", "CC"])` debería devolver false.

```js
assert(!allEqual(testCases[1]));
```

`azSorted(["AA", "ACB", "BB", "CC"])` debería devolver true.

```js
assert(azSorted(testCases[1]));
```

`allEqual([])` debería devolver true.

```js
assert(allEqual(testCases[2]));
```

`azSorted([])` debería devolver true.

```js
assert(azSorted(testCases[2]));
```

`allEqual(["AA"])` debería devolver true.

```js
assert(allEqual(testCases[3]));
```

`azSorted(["AA"])` debería devolver true.

```js
assert(azSorted(testCases[3]));
```

`allEqual(["BB", "AA"])` debería devolver false.

```js
assert(!allEqual(testCases[4]));
```

`azSorted(["BB", "AA"])` debería devolver false.

```js
assert(!azSorted(testCases[4]));
```

# --seed--

## --after-user-code--

```js
const testCases = [['AA', 'AA', 'AA', 'AA'], ['AA', 'ACB', 'BB', 'CC'], [], ['AA'], ['BB', 'AA']];
```

## --seed-contents--

```js
function allEqual(arr) {

  return true;
}

function azSorted(arr) {

  return true;
}
```

# --solutions--

```js
function allEqual(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] === a[i]);
  } return out;
}

function azSorted(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] < a[i]);
  } return out;
}
```
