---
id: 596fda99c69f779975a1b67d
title: Cuenta las ocurrencias de una subcadena
challengeType: 1
forumTopicId: 302237
dashedName: count-occurrences-of-a-substring
---

# --description--

Crea una función, o muestra una función incorporada, para contar el número de ocurrencias no superpuestas de una subcadena dentro de una cadena.

La función debe tomar dos argumentos:

<ul>
  <li>the first argument being the string to search, and</li>
  <li>the second a substring to be searched for.</li>
</ul>

Debe devolver un conteo entero.

La coincidencia debe producir el mayor número de coincidencias no superpuestas.

En general, esto significa esencialmente hacer coincidir de izquierda a derecha o de derecha a izquierda.

# --hints--

`countSubstring` debería ser una función.

```js
assert(typeof countSubstring === 'function');
```

`countSubstring("the three truths", "th")` debería devolver `3`.

```js
assert.equal(countSubstring(testCases[0], searchString[0]), results[0]);
```

`countSubstring("ababababab", "abab")` debería devolver `2`.

```js
assert.equal(countSubstring(testCases[1], searchString[1]), results[1]);
```

`countSubstring("abaabba*bbaba*bbab", "a*b")` debería devolver `2`.

```js
assert.equal(countSubstring(testCases[2], searchString[2]), results[2]);
```

# --seed--

## --after-user-code--

```js
const testCases = ['the three truths', 'ababababab', 'abaabba*bbaba*bbab'];
const searchString = ['th', 'abab', 'a*b'];
const results = [3, 2, 2];
```

## --seed-contents--

```js
function countSubstring(str, subStr) {

  return true;
}
```

# --solutions--

```js
function countSubstring(str, subStr) {
  const escapedSubStr = subStr.replace(/[.+*?^$[\]{}()|/]/g, '\\$&');
  const matches = str.match(new RegExp(escapedSubStr, 'g'));
  return matches ? matches.length : 0;
}
```
