---
id: 597089c87eec450c68aa1643
title: Define un tipo de dato primitivo
challengeType: 1
forumTopicId: 302248
dashedName: define-a-primitive-data-type
---

# --description--

Define un tipo que se comporte como un entero pero que tenga un valor válido mínimo de 1 y un valor válido máximo de 10.

Manejo de errores:

<ul>
  <li>If you try to instantiate a <code>Num</code> with a value outside of 1 - 10, it should throw a <code>TypeError</code> with an error message of <code>'Out of range'</code>.</li>
  <li>If you try to instantiate a <code>Num</code> with a value that is not a number, it should throw a <code>TypeError</code> with an error message of <code>'Not a Number'</code>.</li>
</ul>

# --hints--

`Num` debe ser una función.

```js
assert(typeof Num === 'function');
```

`new Num(4)` debería devolver un objeto.

```js
assert(typeof new Num(4) === 'object');
```

`new Num('test')` debería lanzar un TypeError con el mensaje 'Not a Number'.

```js
assert.throws(() => new Num('test'), TypeError);
```

`new Num(0)` debería lanzar un TypeError con el mensaje 'Out of range'.

```js
assert.throws(() => new Num(0), TypeError);
```

`new Num(-5)` debería lanzar un TypeError con el mensaje 'Out of range'.

```js
assert.throws(() => new Num(-5), TypeError);
```

`new Num(10)` debería lanzar un TypeError con el mensaje 'Out of range'.

```js
assert.throws(() => new Num(11), TypeError);
```

`new Num(20)` debería lanzar un TypeError con el mensaje 'Out of range'.

```js
assert.throws(() => new Num(20), TypeError);
```

`new Num(3) + new Num(4)` debería ser igual a 7.

```js
assert.equal(new Num(3) + new Num(4), 7);
```

`new Num(3) - new Num(4)` debería ser igual a -1.

```js
assert.equal(new Num(3) - new Num(4), -1);
```

`new Num(3) * new Num(4)` debería ser igual a 12.

```js
assert.equal(new Num(3) * new Num(4), 12);
```

`new Num(3) / new Num(4)` debería ser igual a 0.75.

```js
assert.equal(new Num(3) / new Num(4), 0.75);
```

`new Num(3) < new Num(4)` debería ser true.

```js
assert(new Num(3) < new Num(4));
```

`new Num(3) > new Num(4)` debería ser false.

```js
assert(!(new Num(3) > new Num(4)));
```

`(new Num(5)).toString()` debería devolver '5'

```js
assert.equal(new Num(5).toString(), '5');
```

# --seed--

## --seed-contents--

```js
function Num(n) {

  return n;
}
```

# --solutions--

```js
function Num(n) {
  if (isNaN(n)) {
    throw new TypeError('Not a Number');
  }
  if (n < 1 || n > 10) {
    throw new TypeError('Out of range');
  }

  this._value = +n;
}
Num.prototype.valueOf = function() { return this._value; };
Num.prototype.toString = function () { return this._value.toString(); };
```
