---
id: 598de241872ef8353c58a7a2
title: Evalúa coeficientes binomiales
challengeType: 1
forumTopicId: 302259
dashedName: evaluate-binomial-coefficients
---

# --description--

Escribe una función para calcular el coeficiente binomial para el valor dado de n y k.

Esta fórmula es recomendada:

$\\binom{n}{k} = \\frac{n!}{(n-k)!k!} = \\frac{n(n-1)(n-2)\\ldots(n-k+1)}{k(k-1)(k-2)\\ldots 1}$

# --hints--

`binom` debería ser una función.

```js
assert(typeof binom === 'function');
```

`binom(5,3)` debería devolver 10.

```js
assert.equal(binom(5, 3), 10);
```

`binom(7,2)` debería devolver 21.

```js
assert.equal(binom(7, 2), 21);
```

`binom(10,4)` debería devolver 210.

```js
assert.equal(binom(10, 4), 210);
```

`binom(6,1)` debería devolver 6.

```js
assert.equal(binom(6, 1), 6);
```

`binom(12,8)` debería devolver 495.

```js
assert.equal(binom(12, 8), 495);
```

# --seed--

## --seed-contents--

```js
function binom(n, k) {

}
```

# --solutions--

```js
function binom(n, k) {
  let coeff = 1;
  for (let i = n - k + 1; i <= n; i++) coeff *= i;
  for (let i = 1; i <= k; i++) coeff /= i;
  return coeff;
}
```
