---
id: 598eef80ba501f1268170e1e
title: Secuencias numéricas de Fibonacci n-pasos
challengeType: 1
forumTopicId: 302267
dashedName: fibonacci-n-step-number-sequences
---

# --description--

Estas series numéricas son una expansión de la secuencia ordinaria de Fibonacci donde:

<ol>
  <li>Para $n = 2$ tenemos la secuencia de Fibonacci; con valores iniciales $[1, 1]$ y $F_k^2 = F_{k-1}^2 + F_{k-2}^2$</li>
  <li>Para $n = 3$ tenemos la secuencia tribonacci; con valores iniciales $[1, 1, 2]$ y $F_k^3 = F_{k-1}^3 + F_{k-2}^3 + F_{k-3}^3$</li>
  <li>Para $n = 4$ tenemos la secuencia tetranacci; con valores iniciales $[1, 1, 2, 4]$ y $F_k^4 = F_{k-1}^4 + F_{k-2}^4 + F_{k-3}^4 + F_{k-4}^4$...</li>
  <li>Para un $n>2$ general tenemos la secuencia Fibonacci de $n$ pasos - $F_k^n$; con valores iniciales que son los primeros $n$ valores de la secuencia Fibonacci de $n$ pasos $(n-1)$'sima $F_k^{n-1}$; y el valor $k$-ésimo de esta $n$-ésima secuencia siendo $F_k^n = \sum_{i=1}^{(n)} {F_{k-i}^{(n)}}$</li>
</ol>

Para valores pequeños de $n$, a veces se usan prefijos numéricos griegos para nombrar individualmente cada serie.

Secuencias Fibonacci de $n$ pasos:

| $n$ | Nombre de la serie | Valores                                                 |
| --- | ------------------ | ------------------------------------------------------ |
| 2   | fibonacci          | 1 1 2 3 5 8 13 21 34 55 89 144 233 377 610 ...         |
| 3   | tribonacci         | 1 1 2 4 7 13 24 44 81 149 274 504 927 1705 3136 ...    |
| 4   | tetranacci         | 1 1 2 4 8 15 29 56 108 208 401 773 1490 2872 5536 ...  |
| 5   | pentanacci         | 1 1 2 4 8 16 31 61 120 236 464 912 1793 3525 6930 ...  |
| 6   | hexanacci          | 1 1 2 4 8 16 32 63 125 248 492 976 1936 3840 7617 ...  |
| 7   | heptanacci         | 1 1 2 4 8 16 32 64 127 253 504 1004 2000 3984 7936 ... |
| 8   | octonacci          | 1 1 2 4 8 16 32 64 128 255 509 1016 2028 4048 8080 ... |
| 9   | nonanacci          | 1 1 2 4 8 16 32 64 128 256 511 1021 2040 4076 8144 ... |
| 10  | decanacci          | 1 1 2 4 8 16 32 64 128 256 512 1023 2045 4088 8172 ... |

Se pueden generar secuencias relacionadas donde se cambian los valores iniciales: La serie Lucas suma los dos valores precedentes como la serie Fibonacci para $n=2$ pero usa $\[2, 1]$ como sus valores iniciales.

# --instructions--

Escribe una función para generar secuencias numéricas Fibonacci de $n$ pasos y secuencias Lucas. El primer parámetro será $n$. El segundo parámetro será la cantidad de elementos que se deben devolver. El tercer parámetro especificará si se debe devolver la secuencia Fibonacci o la secuencia Lucas. Si el parámetro es `"f"` entonces devuelve la secuencia Fibonacci y si es `"l"`, entonces devuelve la secuencia Lucas. Las secuencias deben devolverse como un arreglo.

# --hints--

`fib_luc` debe ser una función.

```js
assert(typeof fib_luc === 'function');
```

`fib_luc(2,10,"f")` debe devolver `[1,1,2,3,5,8,13,21,34,55]`.

```js
assert.deepEqual(fib_luc(2, 10, 'f'), globalThis.ans[0]);
```

`fib_luc(3,15,"f")` debe devolver `[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136]`.

```js
assert.deepEqual(fib_luc(3, 15, 'f'), globalThis.ans[1]);
```

`fib_luc(4,15,"f")` debe devolver `[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536]`.

```js
assert.deepEqual(fib_luc(4, 15, 'f'), globalThis.ans[2]);
```

`fib_luc(2,10,"l")` debe devolver `[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76]`.

```js
assert.deepEqual(fib_luc(2, 10, 'l'), globalThis.ans[3]);
```

`fib_luc(3,15,"l")` debe devolver `[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ]`.

```js
assert.deepEqual(fib_luc(3, 15, 'l'), globalThis.ans[4]);
```

`fib_luc(4,15,"l")` debe devolver `[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ]`.

```js
assert.deepEqual(fib_luc(4, 15, 'l'), globalThis.ans[5]);
```

`fib_luc(5,15,"l")` debe devolver `[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]`.

```js
assert.deepEqual(fib_luc(5, 15, 'l'), globalThis.ans[6]);
```

# --before-all--

```js
globalThis.ans = [[1,1,2,3,5,8,13,21,34,55],
[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136],
[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536],
[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76],
[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ],
[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ],
[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]];
```

# --seed--

## --seed-contents--

```js
function fib_luc(n, len, w) {

}
```

# --solutions--

```js
function fib_luc(n, len, w) {
    function nacci(a, n, len) {
        while (a.length < len) {
            let sum = 0;
            for (let i = Math.max(0, a.length - n); i < a.length; i++)
                sum += a[i];
            a.push(sum);
        }
        return a;
    }
    if(w=="f"){
        return nacci(nacci([1,1], n, n), n, len);
    }else{
        return nacci(nacci([2,1], n, n), n, len);
    }
}
```
