---
id: 598eea87e5cf4b116c3ff81a
title: Factores De Un Numero Mersenne
challengeType: 1
forumTopicId: 302264
dashedName: factors-of-a-mersenne-number
---

# --description--

Un número de Mersenne es un número en la forma de <code>2<sup>P</sup>-1</code>.

Si `P` es primo, el número de Mersenne puede ser un primo de Mersenne. (Si `P` no es primo, el número de Mersenne tampoco es primo.)

En la búsqueda de números primos de Mersenne es ventajoso eliminar exponentes encontrando un factor pequeño antes de comenzar una, potencialmente larga, <a href="https://rosettacode.org/wiki/Lucas-Lehmer test" target="_blank" rel="noopener noreferrer nofollow">prueba de Lucas-Lehmer</a>.

Existen algoritmos muy eficientes para determinar si un número divide <code>2<sup>P</sup>-1</code> (o, de manera equivalente, si <code>2<sup>P</sup> mod (the number) = 1</code>).

Algunos lenguajes ya tienen implementaciones integradas de esta operación exponent-and-mod (denominada modPow o similar).

Lo siguiente es cómo implementar este modPow por ti mismo:

Por ejemplo, calculemos <code>2<sup>23</sup> mod 47</code>.

Convierte el exponente 23 a binario, obtienes 10111. Comenzando con <code><tt>square</tt> = 1</code>, elévalo al cuadrado repetidamente.

Quita la parte superior del exponente, y si es 1 multiplica `square` por la base de la exponenciación (2), luego calcula <code><tt>square</tt> módulo 47</code>.

Usa el resultado del modulo desde el ultimo paso como el valor inicial de `square` en el proximo paso:

<pre>Remove   Optional
square        top bit  multiply by 2  mod 47
------------  -------  -------------  ------
1*1 = 1       1  0111  1*2 = 2           2
2*2 = 4       0   111     no             4
4*4 = 16      1    11  16*2 = 32        32
32*32 = 1024  1     1  1024*2 = 2048    27
27*27 = 729   1        729*2 = 1458      1
</pre>

Dado que <code>2<sup>23</sup> mod 47 = 1</code>, 47 es un factor de <code>2<sup>P</sup>-1</code>.

(Para ver esto, resta 1 a ambos lados: <code>2<sup>23</sup>-1 = 0 mod 47</code>.)

Como hemos demostrado que 47 es un factor, <code>2<sup>23</sup>-1</code> no es primo.

Propiedades adicionales de los números de Mersenne nos permiten refinar el proceso aún más.

Cualquier factor `q` de <code>2<sup>P</sup>-1</code> debe ser de la forma `2kP+1`, siendo `k` un entero positivo o cero. Además, `q` debe ser `1` o `7 mod 8`.

Finalmente, cualquier factor potencial `q` debe ser <a href="https://rosettacode.org/wiki/Primality_by_trial_division" target="_blank" rel="noopener noreferrer nofollow">primo</a>.

Como en otros algoritmos de división por prueba, el algoritmo se detiene cuando `2kP+1 > sqrt(N)`. Estas pruebas principales solo funcionan en números de Mersenne donde `P` es primo. Por ejemplo, <code>M<sub>4</sub>=15</code> no produce factores usando estas técnicas, pero se factoriza en 3 y 5, ninguno de los cuales cumple con `2kP+1`.

# --instructions--

Usando el método anterior encuentra un factor de <code>2<sup>p</sup>-1</code>.

# --hints--

`check_mersenne` debería ser una función.

```js
assert(typeof check_mersenne === 'function');
```

`check_mersenne(3)` debería devolver una cadena.

```js
assert(typeof check_mersenne(3) == 'string');
```

`check_mersenne(3)` debería devolver la cadena `M3 = 2^3-1 is prime`.

```js
assert.equal(check_mersenne(3), 'M3 = 2^3-1 is prime');
```

`check_mersenne(23)` debería devolver la cadena `M23 = 2^23-1 is composite with factor 47`.

```js
assert.equal(check_mersenne(23), 'M23 = 2^23-1 is composite with factor 47');
```

`check_mersenne(929)` debería devolver la cadena `M929 = 2^929-1 is composite with factor 13007`.

```js
assert.equal(
  check_mersenne(929),
  'M929 = 2^929-1 is composite with factor 13007'
);
```

# --seed--

## --seed-contents--

```js
function check_mersenne(p) {

}
```

# --solutions--

```js
function check_mersenne(p){
    function isPrime(value){
      for (let i=2; i < value; i++){
        if (value % i == 0){
          return false;
        }
        if (value % i != 0){
          return true;
         }
      }
    }

    function trial_factor(base, exp, mod){
      let square, bits;
      square = 1;
      bits = exp.toString(2).split('');
      for (let i=0,ln=bits.length; i<ln; i++){
        square = Math.pow(square, 2) * (bits[i] == 1 ? base : 1) % mod;
      }
      return (square == 1);
    }

    function mersenne_factor(p){
      let limit, k, q;
      limit = Math.sqrt(Math.pow(2,p) - 1);
      k = 1;
      while ((2*k*p - 1) < limit){
        q = 2*k*p + 1;
        if (isPrime(q) && (q % 8 == 1 || q % 8 == 7) && trial_factor(2,p,q)){
          return q; // q is a factor of 2**p-1
        }
        k++;
      }
      return null;
    }
  let f, result;
  result="M"+p+" = 2^"+p+"-1 is ";
  f = mersenne_factor(p);
  result+=f == null ? "prime" : "composite with factor "+f;
  return result;
}
```
