---
id: 5992e222d397f00d21122931
title: Palabra fibonacci
challengeType: 1
forumTopicId: 302269
dashedName: fibonacci-word
---

# --description--

La Secuencia de Palabras de Fibonacci puede crearse de una manera análoga a la Secuencia de Fibonacci, pero se centra en iterar la concatenación.

<pre>Define  F_Word<sub>1</sub>  as  <strong>1</strong>
Define  F_Word<sub>2</sub>  as  <strong>0</strong>
Form   F_Word<sub>3</sub>  as  F_Word<sub>2</sub>   concatenated with  F_Word<sub>1</sub>   i.e.:  <strong>01</strong>
Form   F_Word<sub>n</sub>  as  F_Word<sub>n-1</sub>  concatenated with  F_word<sub>n-2</sub>
</pre>

El cálculo de entropía es necesario en este desafío, <a href="https://www.freecodecamp.org/learn/coding-interview-prep/rosetta-code/entropy" target="_blank" rel="noopener noreferrer nofollow">como se muestra en este desafío de Rosetta Code</a>

# --instructions--

Escribe una función para devolver las primeras `n` Fibonacci Words. El número de `n` se proporciona como un parámetro para la función. La función debe devolver un arreglo de objetos. Los objetos deben tener la forma: `{ N: 1, Length: 1, Entropy: 0, Word: '1' }`. `Entropy` se calcula para la cadena `Word` y se redondea a 8 dígitos decimales de precisión. Ten en cuenta que los índices de esta secuencia comienzan en `1`.

# --hints--

`fibWord` debería ser una función.

```js
assert(typeof fibWord === 'function');
```

`fibWord(5)` debe devolver una matriz.

```js
assert(Array.isArray(fibWord(5)));
```

`fibWord(5)` debería volver `[{ N:1, Longitud:1, Entropía:0, Palabra:"1" },{ N:2, Longitud:1, Entropía:0, Palabra:"0" },{ N:3, Longitud:2, Entropía :1, Palabra:"01" },{ N:4, Longitud:3, Entropía:0,91829583, Palabra:"010" },{ N:5, Longitud:5, Entropía:0,97095059, Palabra:"01001" }].

```js
assert.deepEqual(fibWord(5), words5);
```

`fibWord(7)` debería devolver `[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }, { N:6, Length:8, Entropy:0.954434, Word:'01001010' }, { N:7, Length:13, Entropy:0.9612366, Word:'0100101001001' }]`.

```js
assert.deepEqual(fibWord(7), words7);
```

# --seed--

## --after-user-code--

```js
const words5 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' }
];

const words7 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' },
  { N: 6, Length: 8, Entropy: 0.954434, Word: '01001010' },
  { N: 7, Length: 13, Entropy: 0.9612366, Word: '0100101001001' }
];
```

## --seed-contents--

```js
function fibWord(n) {

}
```

# --solutions--

```js
// Round to digits
function roundFloat(num, digits) {
  return Math.round(num * 10.0**digits) / (10.0**digits);
}

// Entropy calculation for string with only 0 and 1
function entropy(word) {
  function digitEntropy(count) {
    return count < 1 ? 0
      : - count / word.length * Math.log2(count / word.length);
  }
  const numZeros = word.split('').filter(e => e === '0').length;
  const numOnes  = word.length - numZeros;
  return roundFloat(digitEntropy(numZeros) + digitEntropy(numOnes), 8);
}

// Compute array of Fibonacci words
function fibWord(n) {
  return [...Array(n).keys()]
    .reduce((words, i) => {
      const word = i === 0 ? "1"
                 : i === 1 ? "0"
                 : words[i - 1].Word + words[i - 2].Word;
      words.push(
        { N: i + 1, Length: word.length, Entropy: entropy(word), Word: word }
      );
      return words;
    }, []);
}
```
