---
id: 599d0ba974141b0f508b37d5
title: Primos Emirp
challengeType: 1
forumTopicId: 302253
dashedName: emirp-primes
---

# --description--

Un emirp (**prime** escrito al revés) son números primos que, al invertirse (en su representación decimal), son un primo diferente.

# --instructions--

Escribe una función que:

<ul>
  <li>Shows the first <code>n</code> emirp numbers.</li>
  <li>Shows the emirp numbers in a range.</li>
  <li>Shows the number of emirps in a range.</li>
  <li>Shows the <code>n<sup>th</sup></code> emirp number.</li>
</ul>

La función debe aceptar dos parámetros. El primero recibirá `n` o el rango como un arreglo. El segundo recibirá un booleano, que especifica si la función devuelve los emirps como un arreglo o un solo número (el número de primos en el rango o el <code>n<sup>th</sup></code> primo). Según los parámetros, la función debe devolver un arreglo o un número.

# --hints--

`emirps` debería ser una función.

```js
assert(typeof emirps === 'function');
```

`emirps(20,true)` debería devolver `[13,17,31,37,71,73,79,97,107,113,149,157,167,179,199,311,337,347,359,389]`

```js
assert.deepEqual(emirps(20, true), [
  13,
  17,
  31,
  37,
  71,
  73,
  79,
  97,
  107,
  113,
  149,
  157,
  167,
  179,
  199,
  311,
  337,
  347,
  359,
  389
]);
```

`emirps(1000)` debería devolver `70529`

```js
assert.deepEqual(emirps(1000), 70529);
```

`emirps([7700,8000],true)` debería devolver `[7717,7757,7817,7841,7867,7879,7901,7927,7949,7951,7963]`

```js
assert.deepEqual(emirps([7700, 8000], true), [
  7717,
  7757,
  7817,
  7841,
  7867,
  7879,
  7901,
  7927,
  7949,
  7951,
  7963
]);
```

`emirps([7700,8000],false)` debería devolver `11`

```js
assert.deepEqual(emirps([7700, 8000], false), 11);
```

# --seed--

## --seed-contents--

```js
function emirps(n) {

}
```

# --solutions--

```js
function emirps(num, showEmirps)
{
  const is_prime = function(n)
    {
    if (!(n % 2) || !(n % 3)) return false;
    let p = 1;
    while (p * p < n)
                    { if (n % (p += 4) == 0 || n % (p += 2) == 0)
                            { return false; } }
    return true;
  };
  const is_emirp = function(n) {
    const r = parseInt(n.toString().split('').reverse().join(''));
    return r != n && is_prime(n) && is_prime(r);
  };

  let i,
    arr = [];
  if (typeof num === 'number') {
    for (i = 0; arr.length < num; i++) if (is_emirp(i)) arr.push(i);
    // first x emirps
    if (showEmirps) return arr;
    // xth emirp
    return arr.pop();
  }

  if (Array.isArray(num)) {
    for (i = num[0]; i <= num[1]; i++) if (is_emirp(i)) arr.push(i);
    // emirps between x .. y
    if (showEmirps) return arr;
    // number of emirps between x .. y
    return arr.length;
  }
}
```
