---
id: 59e0a8df964e4540d5abe599
title: Ejecutar brain****
challengeType: 1
forumTopicId: 302261
dashedName: execute-brain
---

# --description--

Escribe una función para implementar un intérprete de Brain\*\*\*\*. La función recibirá una cadena como parámetro y debe devolver una cadena como resultado. Más detalles se dan a continuación:

RCBF es un conjunto de <a href="https://rosettacode.org/wiki/Brainf***" target="_blank" rel="noopener noreferrer nofollow">compiladores e intérpretes de Brainf\*\*\*</a> escritos para Rosetta Code en varios lenguajes.

A continuación hay enlaces a cada una de las versiones de RCBF.

Una implementación solo necesita implementar correctamente las siguientes instrucciones:

| Comando           | Descripción                                                                       |
| ----------------- | --------------------------------------------------------------------------------- |
| <code>></code>    | Mover el puntero a la derecha                                                     |
| <code>&lt;</code> | Mover el puntero a la izquierda                                                   |
| <code>+</code>    | Incrementar la celda de memoria bajo el puntero                                   |
| <code>-</code>    | Decrementar la celda de memoria bajo el puntero                                   |
| <code>.</code>    | Mostrar el carácter indicado por la celda bajo el puntero                         |
| <code>,</code>    | Ingresar un carácter y almacenarlo en la celda bajo el puntero                    |
| <code>\[</code>   | Saltar más allá del <code>]</code> correspondiente si la celda bajo el puntero es 0 |
| <code>]</code>    | Saltar de regreso al <code>\[</code> correspondiente si la celda bajo el puntero no es 0 |

Se permite cualquier tamaño de celda, el soporte para EOF (*E*nd-*O*-*F*ile) es opcional, al igual que si la memoria es limitada o ilimitada.

# --hints--

`brain(globalThis.hello)` debe devolver una cadena

```js
assert(typeof brain(globalThis.hello) === 'string');
```

`brain("++++++[>++++++++++<-]>+++++.")` debe devolver "A"

```js
assert.equal(brain('++++++[>++++++++++<-]>+++++.'), 'A');
```

`brain(globalThis.hello)` debe devolver `Hello World!\n`

```js
assert.equal(brain(globalThis.hello), 'Hello World!\n');
```

`brain(globalThis.fib)` debe devolver `1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89`

```js
assert.equal(brain(globalThis.fib), '1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89');
```

# --before-all--

```js
globalThis.fib=`+

++

+++

++++

+>+>>

>>++++

+++++++

++++++++

+++++++++

++++++++++

++++++>++++

++++++++++++

+++++++++++++

+++<<<<<<[>[>>

>>>>+>+<<<<<<<-

]>>>>>>>[<<<<<<<

+>>>>>>>-]<[>++++

++++++[-<-[>>+>+<<

<-]>>>[<<<+>>>-]+<[

>[-]<[-]]>[<<[>>>+<<

<-]>>[-]]<<]>>>[>>+>+

<<<-]>>>[<<<+>>>-]+<[>

[-]<[-]]>[<<+>>[-]]<<<<

<<<]>>>>>[++++++++++++++

+++++++++++++++++++++++++

+++++++++.[-]]++++++++++<[

->-<]>+++++++++++++++++++++

+++++++++++++++++++++++++++.

[-]<<<<<<<<<<<<[>>>+>+<<<<-]>

>>>[<<<<+>>>>-]<-[>>.>.<<<[-]]

<<[>>+>+<<<-]>>>[<<<+>>>-]<<[<+

>-]>[<+>-]<<<-]`;
globalThis.hello='++++++++[>++++[>++>+++>+++>+<<<<-]>+>+>->>+[<]<-]>>.>---.+++++++..+++.>>.<-.<.+++.------.--------.>>+.>++.'
```

# --seed--

## --seed-contents--

```js
function brain(prog) {

}
```

# --solutions--

```js
function brain(prog){
  var output="";
    var code; // formatted code
  var ip = 0; // current instruction within code
  var nest = 0; // current bracket nesting (for Out button)
  var ahead = []; // locations of matching brackets

  var data = [0]; // data array (mod by +, -)
  var dp = 0; // index into data (mod by <, >)

  var inp = 0; // current input character (fetch with ,)
  var quit = 0;
    var commands = {
    '>':function() { if (++dp >= data.length) data[dp]=0 },
    '<':function() { if (--dp < 0) quit++ },
    '+':function() { ++data[dp] },
    '-':function() { --data[dp] },
    '[':function() { if (!data[dp]) ip = ahead[ip]; else ++nest },
    ']':function() { if ( data[dp]) ip = ahead[ip]; else --nest },
    ',':function() {
        var c = document.getElementById("input").value.charCodeAt(inp++);
        data[dp] = isNaN(c) ? 0 : c; // EOF: other options are -1 or no change
    },
    '.':function() {
            output+=String.fromCharCode(data[dp]);
            /*var s = document.getElementById("output").innerHTML)
             + String.fromCharCode(data[dp]);
            s = s.replace(/\n/g,"<br>").replace(/ /g,"&amp;nbsp;");
            document.getElementById("output").innerHTML = s;*/
        },
    };

    let ar=prog.split('');
    var st = [], back, error = -1;
    for (ip=0; ip<ar.length; ip++) {
        switch(ar[ip]) {
        case '[':
            st.push(ip);
            break;
        case ']':
            if (st.length == 0) error = ip;
            back = st.pop();
            ahead[ip] = back;
            ahead[back] = ip;
            break;
        }
    }

    for(ip=0;ip<ar.length;ip++){
    if(commands.hasOwnProperty(ar[ip]))
          commands[ar[ip]]();
    }

    return output;
}
```
