---
id: 5a23c84252665b21eecc8014
title: Estabilidad de ordenamiento
challengeType: 1
forumTopicId: 302308
dashedName: sort-stability
---

# --description--

Al ordenar registros en una tabla por una columna o campo en particular, un <a href="https://www.freecodecamp.org/news/stability-in-sorting-algorithms-a-treatment-of-equality-fa3140a5a539/" target="_blank" rel="noopener noreferrer nofollow">ordenamiento estable</a> siempre conservará el orden relativo de los registros que tienen la misma clave.

Por ejemplo, en esta tabla de países y ciudades, un ordenamiento estable en la **segunda** columna, las ciudades, mantendría a Birmingham de EE. UU. por encima de Birmingham del Reino Unido. (Aunque un ordenamiento inestable *podría*, en este caso, colocar a Birmingham de EE. UU. por encima de Birmingham del Reino Unido, una rutina de ordenamiento estable lo *garantizaría*).

<pre>UK  London
US  New York
US  Birmingham
UK  Birmingham
</pre>

De manera similar, ordenar de forma estable solo por la primera columna generaría "UK London" como el primer elemento y "US Birmingham" como el último elemento (ya que se mantendría el orden de los elementos que tienen la misma primera palabra – "UK" o "US").

# --instructions--

Escribe una función que tome un arreglo 2D como parámetro. Cada elemento tiene 2 elementos similares al ejemplo anterior. La función debe ordenar el arreglo como se mencionó previamente y devolver el arreglo ordenado.

# --hints--

`stableSort` debería ser una función.

```js
assert(typeof stableSort == 'function');
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` debería devolver un arreglo.

```js
assert(
  Array.isArray(
    stableSort([
      ['UK', 'London'],
      ['US', 'New York'],
      ['US', 'Birmingham'],
      ['UK', 'Birmingham']
    ])
  )
);
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` debería devolver `[["US", "Birmingham"], ["UK", "Birmingham"], ["UK", "London"], ["US", "New York"]]`.

```js
assert.deepEqual(
  stableSort([
    ['UK', 'London'],
    ['US', 'New York'],
    ['US', 'Birmingham'],
    ['UK', 'Birmingham']
  ]),
  [
    ['US', 'Birmingham'],
    ['UK', 'Birmingham'],
    ['UK', 'London'],
    ['US', 'New York']
  ]
);
```

`stableSort([[2, 2], [1, 2], [1, 4], [1, 5]])` debería devolver `[[2, 2], [1, 2], [1, 4], [1, 5]]`.

```js
assert.deepEqual(
  stableSort([
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]),
  [
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]
);
```

`stableSort([[11, 55], [12, 45], [11, 45], [32, 45]])` debería devolver `[[12, 45], [11, 45], [32, 45], [11, 55]]`.

```js
assert.deepEqual(
  stableSort([
    [11, 55],
    [12, 45],
    [11, 45],
    [32, 45]
  ]),
  [
    [12, 45],
    [11, 45],
    [32, 45],
    [11, 55]
  ]
);
```

`stableSort([[10, 22], [1, 2], [1, 4], [1, 5], [10, 9]])` debería devolver `[[1, 2], [1, 4], [1, 5], [10, 9], [10, 22]]`.

```js
assert.deepEqual(
  stableSort([
    [10, 22],
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9]
  ]),
  [
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9],
    [10, 22]
  ]
);
```

`stableSort([[55, 54], [12, 22], [31, 43], [31, 54], [10, 49]])` debería devolver `[[12, 22], [31, 43], [10, 49], [55, 54], [31, 54]]`.

```js
assert.deepEqual(
  stableSort([
    [55, 54],
    [12, 22],
    [31, 43],
    [31, 54],
    [10, 49]
  ]),
  [
    [12, 22],
    [31, 43],
    [10, 49],
    [55, 54],
    [31, 54]
  ]
);
```

# --seed--

## --seed-contents--

```js
function stableSort(arr) {

}
```

# --solutions--

```js
function stableSort(arr) {
  arr.sort(function(a, b) {
    return a[1] < b[1] ? -1 : a[1] > b[1] ? 1 : 0;
  });
  return arr;
}
```
