---
id: 5a23c84252665b21eecc8036
title: Desmantela códigos de control y caracteres extendidos desde una cadena
challengeType: 1
forumTopicId: 302327
dashedName: strip-control-codes-and-extended-characters-from-a-string
---

# --description--

La tarea es desmantelar códigos de control y caracteres extendidos desde una cadena. La solución debería demostrar cómo lograr cada uno de los siguientes resultados: Una cadena con códigos de control y caracteres extendidos desmantelados. En ASCII, los códigos de control tienen códigos decimales de 0 hasta 31 y 127. En un sistema basado en ASCII, si los códigos de control son desmantelados, la cadena resultante tendría todos sus caracteres dentro del rango de 32 a 126 decimales en la tabla ASCII. En un sistema no basado en ASCII, consideramos caracteres que no tienen un correspondiente glifo en la tabla ASCII (dentro del rango ASCII de 32 a 126 decimales) a ser caracteres extendidos para el propósito de esta tarea.

# --hints--

`strip` debería ser una función.

```js
assert(typeof strip == 'function');
```

`strip("abc")` debería devolver una cadena.

```js
assert(typeof strip('abc') == 'string');
```

`strip("\ba\x00b\n\rc\fd\xc3")` debería devolver `"abcd"`.

```js
assert.equal(strip('\ba\x00b\n\rc\fd\xc3'), 'abcd');
```

`strip("\u0000\n abc\u00E9def\u007F")` debería devolver `" abcdef"`.

```js
assert.equal(strip('\u0000\n abc\u00E9def\u007F'), ' abcdef');
```

`strip("a\n\tb\u2102d\u2147f")` debería devolver `"abdf"`.

```js
assert.equal(strip('a\n\tb\u2102d\u2147f'), 'abdf');
```

`strip("Français.")` debería devolver `"Franais."`.

```js
assert.equal(strip('Français.'), 'Franais.');
```

`strip("123\tabc\u0007DEF\u007F+-*/€æŧðłþ")` debería devolver `"123abcDEF+-*/"`.

```js
assert.equal(strip('123\tabc\u0007DEF\u007F+-*/€æŧðłþ'), '123abcDEF+-*/');
```

# --seed--

## --seed-contents--

```js
function strip(s) {

}
```

# --solutions--

```js
function strip(s) {
  return s
    .split('')
    .filter(function(x) {
      var n = x.charCodeAt(0);

      return 31 < n && 127 > n;
    })
    .join('');
}
```
