---
id: 5a23c84252665b21eecc7ec4
title: JortSort
challengeType: 1
forumTopicId: 302293
dashedName: jortsort
---

# --description--

jortSort es un conjunto de herramientas de ordenamiento que hace que tú hagas el trabajo y garantiza eficiencia porque no tendrás que ordenar nunca más. Fue presentado originalmente por Jenn "Moneydollars" Schiffer en la prestigiosa JSConf2014.

jortSort debe ser una función que tome un solo arreglo de objetos comparables como argumento. Luego ordena el arreglo en orden ascendente y compara el arreglo ordenado con el arreglo proporcionado originalmente. Si los arreglos coinciden (es decir, el arreglo original ya estaba ordenado), la función devuelve true. Si los arreglos no coinciden (es decir, el arreglo original no estaba ordenado), la función devuelve false.

# --hints--

`jortsort` debería ser una función.

```js
assert(typeof jortsort == 'function');
```

`jortsort([1,2,3,4,5])` debería devolver un booleano.

```js
assert(typeof jortsort([1, 2, 3, 4, 5]) == 'boolean');
```

`jortsort([1,2,3,4,5])` debería devolver `true`.

```js
assert.equal(jortsort([1, 2, 3, 4, 5]), true);
```

`jortsort([1,2,13,4,5])` debería devolver `false`.

```js
assert.equal(jortsort([1, 2, 13, 4, 5]), false);
```

`jortsort([12,4,51,2,4])` debería devolver `false`.

```js
assert.equal(jortsort([12, 4, 51, 2, 4]), false);
```

`jortsort([1,2])` debería devolver `true`.

```js
assert.equal(jortsort([1, 2]), true);
```

`jortsort([5,4,3,2,1])` debería devolver `false`.

```js
assert.equal(jortsort([5, 4, 3, 2, 1]), false);
```

`jortsort([1,1,1,1,1])` debería devolver `true`.

```js
assert.equal(jortsort([1, 1, 1, 1, 1]), true);
```

# --seed--

## --seed-contents--

```js
function jortsort(array) {

}
```

# --solutions--

```js
function jortsort(array) {
  // sort the array
  var originalArray = array.slice(0);
  array.sort( function(a,b){return a - b} );

  // compare to see if it was originally sorted
  for (var i = 0; i < originalArray.length; ++i) {
    if (originalArray[i] !== array[i]) return false;
  }

  return true;
};
```
