---
id: 5a23c84252665b21eecc7ee0
title: Factoriales izquierdos
challengeType: 1
forumTopicId: 302302
dashedName: left-factorials
---

# --description--

**Factoriales izquierdos**, $ !n $, pueden referirse tanto a *subfactoriales* como a *sumas factoriales*. La misma notación puede verse confusamente usada para las dos definiciones diferentes. A veces, los *subfactoriales* (también conocidos como *desarreglos*) pueden usar cualquiera de las notaciones:

<ul>
  <li>$!n`$</li>
  <li>$!n$</li>
  <li>$n¡$</li>
</ul>

(Eso puede no ser visualmente obvio, pero el último ejemplo usa un signo de exclamación invertido.) Esta tarea usará esta fórmula para el **factorial izquierdo**:

$ !n = \\sum\_{k=0}^{n-1} k! $

donde $!0 = 0$

# --instructions--

Escribe una función para calcular el factorial izquierdo de un número dado.

# --hints--

`leftFactorial` debería ser una función.

```js
assert(typeof leftFactorial == 'function');
```

`leftFactorial(0)` debería devolver un número.

```js
assert(typeof leftFactorial(0) == 'number');
```

`leftFactorial(0)` debería devolver `0`.

```js
assert.equal(leftFactorial(0), 0);
```

`leftFactorial(1)` debería devolver `1`.

```js
assert.equal(leftFactorial(1), 1);
```

`leftFactorial(2)` debería devolver `2`.

```js
assert.equal(leftFactorial(2), 2);
```

`leftFactorial(3)` debería devolver `4`.

```js
assert.equal(leftFactorial(3), 4);
```

`leftFactorial(10)` debería devolver `409114`.

```js
assert.equal(leftFactorial(10), 409114);
```

`leftFactorial(17)` debería devolver `22324392524314`.

```js
assert.equal(leftFactorial(17), 22324392524314);
```

`leftFactorial(19)` debería devolver `6780385526348314`.

```js
assert.equal(leftFactorial(19), 6780385526348314);
```

# --seed--

## --seed-contents--

```js
function leftFactorial(n) {

}
```

# --solutions--

```js
function leftFactorial(n) {
  if (n == 0) return 0;
  if (n == 1) return 1;

  // Note: for n>=20, the result may not be correct.
  // This is because JavaScript uses 53 bit integers and
  // for n>=20 result becomes too large.

  let res = 2,
    fact = 2;
  for (var i = 2; i < n; i++) {
    res += fact;
    fact *= i + 1;
  }

  return res;
}
```
