---
id: 5a23c84252665b21eecc8010
title: Algoritmos de ordenamiento/Ordenamiento Shell
challengeType: 1
forumTopicId: 302317
dashedName: sorting-algorithmsshell-sort
---

# --description--

Escribe una función para ordenar un arreglo de elementos usando el algoritmo Shell sort, un ordenamiento de incremento decreciente. La función debe devolver el arreglo ordenado.

El ordenamiento Shell (también conocido como Shellsort o método de Shell) lleva el nombre de su inventor, Donald Shell, quien publicó el algoritmo en 1959.

Shell sort es una secuencia de ordenamientos por inserción entrelazados basada en una secuencia de incrementos. El tamaño del incremento se reduce después de cada pasada hasta que el tamaño del incremento es 1.

Con un tamaño de incremento de 1, la ordenación es una estructura básica de inserción, pero para ese momento los datos están garantizados a estar casi ordenados, que es el "mejor caso" de la ordenación por inserción.

Cualquier secuencia ordenará los datos siempre que termine en 1, pero algunas funcionan mejor que otras.

Los estudios empíricos han demostrado que una secuencia de incremento geométrico con una razón de aproximadamente 2.2 funciona bien en la práctica.

# --hints--

`shellSort` debería ser una función.

```js
assert(typeof shellSort == 'function');
```

`shellSort([25, 32, 12, 7, 20])` debería devolver un arreglo.

```js
assert(Array.isArray(shellSort([25, 32, 12, 7, 20])));
```

`shellSort([25, 32, 12, 7, 20])` debería devolver `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(shellSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`shellSort([38, 45, 35, 8, 13])` debería devolver `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(shellSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`shellSort([43, 36, 20, 34, 24])` debería devolver `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(shellSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`shellSort([12, 33, 26, 18, 1, 16, 38])` debería devolver `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(shellSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`shellSort([3, 39, 48, 16, 1, 4, 29])` debería devolver `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(shellSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function shellSort(a) {

}
```

# --solutions--

```js
function shellSort(a) {
  for (var h = a.length; h > 0; h = parseInt(h / 2)) {
    for (var i = h; i < a.length; i++) {
      var k = a[i];
      for (var j = i; j >= h && k < a[j - h]; j -= h) a[j] = a[j - h];
      a[j] = k;
    }
  }
  return a;
}
```
