---
id: 5a23c84252665b21eecc8041
title: Suma de una serie
challengeType: 1
forumTopicId: 302333
dashedName: sum-of-a-series
---

# --description--

Calcula el **n**<sup>th</sup> término de una <em>serie</em>, es decir, la suma de los primeros **n** términos de la <em>secuencia</em> correspondiente. De manera informal, este valor, o su límite cuando **n** tiende a infinito, también se llama *suma de la serie*, de ahí el título de esta tarea. Para esta tarea, usa: $S_n = \displaystyle\sum_{k=1}^n \frac{1}{k^2}$.

# --instructions--

Escribe una función que tome $a$ y $b$ como parámetros y devuelva la suma de los miembros $a^{th}$ a $b^{th}$ de la secuencia.

# --hints--

`sum` debería ser una función.

```js
assert(typeof sum == 'function');
```

`sum(1, 100)` debería devolver un número.

```js
assert(typeof sum(1, 100) == 'number');
```

`sum(1, 100)` debería devolver `1.6349839001848923`.

```js
assert.equal(sum(1, 100), 1.6349839001848923);
```

`sum(33, 46)` debería devolver `0.009262256361481223`.

```js
assert.equal(sum(33, 46), 0.009262256361481223);
```

`sum(21, 213)` debería devolver `0.044086990748706555`.

```js
assert.equal(sum(21, 213), 0.044086990748706555);
```

`sum(11, 111)` debería devolver `0.08619778593108679`.

```js
assert.equal(sum(11, 111), 0.08619778593108679);
```

`sum(1, 10)` debería devolver `1.5497677311665408`.

```js
assert.equal(sum(1, 10), 1.5497677311665408);
```

# --seed--

## --seed-contents--

```js
function sum(a, b) {

}
```

# --solutions--

```js
function sum(a, b) {
  function fn(x) {
    return 1 / (x * x);
  }
  var s = 0;
  for (; a <= b; a++) s += fn(a);
  return s;
}
```
