---
id: 5a23c84252665b21eecc7ec1
title: Elevación al cuadrado de dígitos iterados
challengeType: 1
forumTopicId: 302291
dashedName: iterated-digits-squaring
---

# --description--

Si sumas el cuadrado de los dígitos de un número Natural (un entero mayor que cero), siempre terminas con 1 o 89:

<pre>15 -> 26 -> 40 -> 16 -> 37 -> 58 -> 89
7 -> 49 -> 97 -> 130 -> 10 -> 1
</pre>

# --instructions--

Escribe una función que tome un número como parámetro y devuelva 1 o 89 después de realizar el proceso mencionado.

# --hints--

`iteratedSquare` debería ser una función.

```js
assert(typeof iteratedSquare == 'function');
```

`iteratedSquare(4)` debería devolver un número.

```js
assert(typeof iteratedSquare(4) == 'number');
```

`iteratedSquare(4)` debería devolver `89`.

```js
assert.equal(iteratedSquare(4), 89);
```

`iteratedSquare(7)` debería devolver `1`.

```js
assert.equal(iteratedSquare(7), 1);
```

`iteratedSquare(15)` debería devolver `89`.

```js
assert.equal(iteratedSquare(15), 89);
```

`iteratedSquare(20)` debería devolver `89`.

```js
assert.equal(iteratedSquare(20), 89);
```

`iteratedSquare(70)` debería devolver `1`.

```js
assert.equal(iteratedSquare(70), 1);
```

`iteratedSquare(100)` debería devolver `1`.

```js
assert.equal(iteratedSquare(100), 1);
```

# --seed--

## --seed-contents--

```js
function iteratedSquare(n) {

}
```

# --solutions--

```js
function iteratedSquare(n) {
    var total;
    while (n != 89 && n != 1) {
        total = 0;
        while (n > 0) {
            total += Math.pow(n % 10, 2);
            n = Math.floor(n/10);
        }
        n = total;
    }
    return n;
}
```
