---
id: 5a23c84252665b21eecc7edf
title: Mínimo común múltiplo
challengeType: 1
forumTopicId: 302301
dashedName: least-common-multiple
---

# --description--

El mínimo común múltiplo de 12 y 18 es 36, porque 12 es un factor (12 × 3 = 36), y 18 es un factor (18 × 2 = 36), y no existe ningún entero positivo menor que 36 que tenga ambos factores. Como caso especial, si $m$ o $n$ es cero, entonces el mínimo común múltiplo es cero. Una forma de calcular el mínimo común múltiplo es iterar todos los múltiplos de $m$, hasta encontrar uno que también sea múltiplo de $n$. Si ya tienes $gcd$ para el <a href="https://rosettacode.org/wiki/Greatest_common_divisor" target="_blank" rel="noopener noreferrer nofollow">máximo común divisor</a>, entonces esta fórmula calcula $lcm$. 

$$
\operatorname{lcm}(m, n) = \frac{|m \times n|}{\operatorname{gcd}(m, n)}
$$

# --instructions--

Calcula el mínimo común múltiplo de un arreglo de enteros. Dado *m* y *n*, el mínimo común múltiplo es el entero positivo más pequeño que tiene tanto a *m* como a *n* como factores.

# --hints--

`LCM` debería ser una función.

```js
assert(typeof LCM == 'function');
```

`LCM([2, 4, 8])` debería devolver un número.

```js
assert(typeof LCM([2, 4, 8]) == 'number');
```

`LCM([2, 4, 8])` debería devolver `8`.

```js
assert.equal(LCM([2, 4, 8]), 8);
```

`LCM([4, 8, 12])` debería devolver `24`.

```js
assert.equal(LCM([4, 8, 12]), 24);
```

`LCM([3, 4, 5, 12, 40])` debería devolver `120`.

```js
assert.equal(LCM([3, 4, 5, 12, 40]), 120);
```

`LCM([11, 33, 90])` debería devolver `990`.

```js
assert.equal(LCM([11, 33, 90]), 990);
```

`LCM([-50, 25, -45, -18, 90, 447])` debería devolver `67050`.

```js
assert.equal(LCM([-50, 25, -45, -18, 90, 447]), 67050);
```

# --seed--

## --seed-contents--

```js
function LCM(A) {

}
```

# --solutions--

```js
function LCM(A) {
  var n = A.length,
    a = Math.abs(A[0]);
  for (var i = 1; i < n; i++) {
    var b = Math.abs(A[i]),
      c = a;
    while (a && b) {
      a > b ? (a %= b) : (b %= a);
    }
    a = Math.abs(c * A[i]) / (a + b);
  }
  return a;
}
```
