---
id: 5a23c84252665b21eecc8013
title: Algoritmos de ordenamiento/Strand sort
challengeType: 1
forumTopicId: 302319
dashedName: sorting-algorithmsstrand-sort
---

# --description--

El **Strand sort** crea subconjuntos ordenados que se combinan para crear el resultado final. 

Considera un `unsortedArray = [3, 1, 4, 2]`. Elige el primer elemento `3` y cópialo en un arreglo separado. Busca cualquier elemento mayor que siga a este elemento. Cuando encuentres un elemento más grande, en este caso `4`, cópialo al arreglo separado, `[3, 4]`, y compara los elementos siguientes con este nuevo valor, `4`.

Después de que hayas llegado al final del array, elimina los elementos que copiaste, `[3, 4]`, y comienza de nuevo con el primer elemento que queda en el `unsortedArray`, en este caso `1`.

Siguiendo este proceso se obtienen dos arreglos ordenados, `[3, 4]` y `[1, 2]`. Fusiona estos dos arreglos para crear el `strandSortedArray`.

```js
const unsortedArray = [3, 1, 4, 2];
const strandsortedArray = [1, 2, 3, 4];
```

Escribe una función para ordenar un arreglo usando el **Strand sort**. La función debe devolver el arreglo ordenado.


# --hints--

`strandSort` debería ser una función.

```js
assert(typeof strandSort == 'function');
```

`strandSort([25, 32, 12, 7, 20])` debería devolver un arreglo.

```js
assert(Array.isArray(strandSort([25, 32, 12, 7, 20])));
```

`strandSort([25, 32, 12, 7, 20])` debería devolver `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(strandSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`strandSort([38, 45, 35, 8, 13])` debería devolver `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(strandSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`strandSort([43, 36, 20, 34, 24])` debería devolver `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(strandSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`strandSort([12, 33, 26, 18, 1, 16, 38])` debería devolver `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(strandSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`strandSort([3, 39, 48, 16, 1, 4, 29])` debería devolver `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(strandSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function strandSort(list) {

}
```

# --solutions--

```js
function strandSort(list) {
  function merge(left, right) {
    var result = [];
    while (left.length != 0 && right.length != 0) {
      if (left[0] <= right[0]) result.push(left.shift());
      else result.push(right.shift());
    }
    result.push.apply(result, left);
    result.push.apply(result, right);
    return result;
  }

  if (list.length <= 1) return list;
  var result = [];
  while (list.length > 0) {
    var sorted = [];
    sorted.push(list.shift());
    var len = list.length;
    for (var i = 1; i < len; i++) {
      var elem = list[i];
      if (sorted[i - 1] <= elem) {
        sorted.push(elem);
        sorted.splice(i, 1);
      }
    }

    result = merge(result, sorted);
  }
  return result;
}
```
