---
id: 5a23c84252665b21eecc7ede
title: Año bisiesto
challengeType: 1
forumTopicId: 302300
dashedName: leap-year
---

# --description--

Determina si un año dado es un año bisiesto en el calendario gregoriano.

# --hints--

`isLeapYear` debería ser una función.

```js
assert(typeof isLeapYear == 'function');
```

`isLeapYear()` debería devolver un booleano.

```js
assert(typeof isLeapYear(2018) == 'boolean');
```

`isLeapYear(2018)` debería devolver `false`.

```js
assert.equal(isLeapYear(2018), false);
```

`isLeapYear(2016)` debería devolver `true`.

```js
assert.equal(isLeapYear(2016), true);
```

`isLeapYear(2000)` debería devolver `true`.

```js
assert.equal(isLeapYear(2000), true);
```

`isLeapYear(1900)` debería devolver `false`.

```js
assert.equal(isLeapYear(1900), false);
```

`isLeapYear(1996)` debería devolver `true`.

```js
assert.equal(isLeapYear(1996), true);
```

`isLeapYear(1800)` debería devolver `false`.

```js
assert.equal(isLeapYear(1800), false);
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

}
```

# --solutions--

```js
function isLeapYear(year) {
  return year % 100 === 0 ? year % 400 === 0 : year % 4 === 0;
}
```
