---
id: 5eaf48389ee512d4d103684b
title: Números que se describen a sí mismos
challengeType: 1
forumTopicId: 385289
dashedName: self-describing-numbers
---

# --description--

Hay varios enteros llamados "auto-descriptivos" o "auto-descriptivos".

Un entero se dice que es "auto-descriptivo" si tiene la propiedad de que, cuando las posiciones de los dígitos se etiquetan de 0 a N-1, el dígito en cada posición es igual al número de veces que ese dígito aparece en el número.

Por ejemplo, **2020** es un número autodescriptivo de cuatro dígitos:

<ul>
    <li> position 0 has value 2 and there are two 0s in the number; </li>
    <li> position 1 has value 0 and there are no 1s in the number; </li>
    <li> position 2 has value 2 and there are two 2s; </li>
    <li> position 3 has value 0 and there are zero 3s; </li>
</ul>

Los números autodescriptivos &lt; 100,000,000 son: 1210, 2020, 21200, 3211000, 42101000.

# --instructions--

Escribe una función que tome un entero positivo como parámetro. Si es auto-descriptivo, devuelve true. De lo contrario, devuelve false.

# --hints--

`isSelfDescribing` debería ser una función.

```js
assert(typeof isSelfDescribing == 'function');
```

`isSelfDescribing()` debería devolver un booleano.

```js
assert(typeof isSelfDescribing(2020) == 'boolean');
```

`isSelfDescribing(2020)` debería devolver `true`.

```js
assert.equal(isSelfDescribing(2020), true);
```

`isSelfDescribing(3021)` debería devolver `false`.

```js
assert.equal(isSelfDescribing(3021), false);
```

`isSelfDescribing(3211000)` debería devolver `true`.

```js
assert.equal(isSelfDescribing(3211000), true);
```

# --seed--

## --seed-contents--

```js
function isSelfDescribing(n) {

}
```

# --solutions--

```js
function isSelfDescribing(n) {
    let digits = String(n).split("");
    digits = digits.map(function(e) {return parseInt(e)});
    let count = digits.map((x) => {return 0})
    digits.forEach((d) =>{
        if (d >= count.length) {
            return false
        }
        count[d] += 1;
    });

     if (digits === count) {
        return true;
    }
    if (digits.length != count.length) {
        return false;
    }
    
    for (let i=0; i< digits.length; i++){
      if (digits[i] !== count[i]) {
        return false;
      }
    }
    return true;
}
```
