---
id: 5eb3e497b8d6d7f63c5517ea
title: Buscar en una lista de registros
challengeType: 1
forumTopicId: 385315
dashedName: search-a-list-of-records
---

# --description--

Un registro consiste en atributos que describen una entidad. Cada atributo tiene un nombre y un valor. Por ejemplo, una persona puede tener un atributo `age` con un valor de 25. Una operación importante en una lista de registros es encontrar un registro con un valor de atributo particular.

# --instructions--

Escribe una función que tome una cadena como parámetro. La función debe devolver el índice del elemento en `list` para el cual el valor del atributo `name` coincide con la cadena dada.

# --hints--

`searchCity` debería ser una función.

```js
assert(typeof searchCity === 'function');
```

`searchCity("Dar Es Salaam")` debería devolver un número.

```js
assert(typeof searchCity('Dar Es Salaam') === 'number');
```

`searchCity("Dar Es Salaam")` debería devolver `6`.

```js
assert.equal(searchCity('Dar Es Salaam'), 6);
```

`searchCity("Casablanca")` debería devolver `9`.

```js
assert.equal(searchCity('Casablanca'), 9);
```

`searchCity("Cairo")` debería devolver `1`.

```js
assert.equal(searchCity('Cairo'), 1);
```

`searchCity("Mogadishu")` debería devolver `4`.

```js
assert.equal(searchCity('Mogadishu'), 4);
```

`searchCity("Lagos")` debería devolver `0`.

```js
assert.equal(searchCity('Lagos'), 0);
```

# --seed--

## --seed-contents--

```js
function searchCity(name) {

}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```

# --solutions--

```js
function searchCity(name) {
  return list.findIndex(item => item.name === name);
}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```
