---
id: 6762ec275cef87635acc4fe3
title: Crea una batería electrónica
challengeType: 25
dashedName: build-drum-machine
demoType: onClick
---

# --description--

**Objetivo:** Cumple con las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

1. Deberías tener un elemento `div` con un `id` de `drum-machine` que contenga todos los demás elementos.

2. Dentro del elemento `#drum-machine` deberías tener otro `div` con un `id` de `pad-bank`.

3. Dentro del elemento `#drum-machine` deberías tener un elemento `p` con un `id` de `display`.

4. Dentro de tu elemento `#pad-bank` deberías tener nueve elementos de drum pad clickables, cada uno con una clase de `drum-pad`, un `id` único que describa el clip de audio que el drum pad activará, y un texto interior que corresponda a una de las siguientes teclas del teclado: `Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`. Los pads de la batería DEBEN estar en este orden.

5. Cada `.drum-pad` debería tener un elemento `audio` que tenga una clase de `clip`, un atributo `src` que apunte a un clip de audio, y un `id` correspondiente al texto interior de su elemento padre `.drum-pad` (por ejemplo, `id="Q"`, `id="W"`, `id="E"`, etc.).

6. Cuando haces clic en un elemento `.drum-pad`, el clip de audio contenido en su elemento hijo de audio debería activarse.

7. Cuando presionas la tecla de disparo asociada a cada `.drum-pad`, el clip de audio contenido en su elemento hijo de audio debería activarse (por ejemplo, al presionar la tecla `Q` se debe activar el drum pad que contiene el string `Q`, al presionar la tecla `W` se debe activar el drum pad que contiene el string `W`, etc.).

8. Cuando se activa un `.drum-pad`, debes mostrar una cadena describiendo el clip de audio asociado como texto interior del elemento `#display` (cada cadena debe ser única).

Algunas muestras de audio que puedes usar para tu drum machine se pueden encontrar en `https://cdn.freecodecamp.org/curriculum/drum/<fileName>`, donde `<fileName>` es el siguiente:

| nombre del tambor | `fileName`       |
| ----------------- | ---------------- |
| Titulo 1          | `Heater-1.mp3`   |
| Titulo 2          | `Heater-2.mp3`   |
| Titulo 3          | `Heater-3.mp3`   |
| Titulo 4          | `Heater-4_1.mp3` |
| Aplauso           | `Heater-6.mp3`   |
| HH Abierto        | `Dsc_Oh.mp3`     |
| Golpe y platillo  | `Kick_n_Hat.mp3` |
| Golpe             | `RP4_KICK_1.mp3` |
| HH Cerrado        | `Cev_H2.mp3`     |

# --hints--

Deberías tener un elemento `div` con un `id` de `drum-machine` que contenga todos los demás elementos.

```js
assert.isNotNull(document.querySelector('div#drum-machine'));
```

Dentro del elemento `#drum-machine` deberías tener otro `div` con un `id` de `pad-bank`.

```js
assert.isNotNull(document.querySelector('#drum-machine div#pad-bank'));
```

Dentro del elemento `#drum-machine` deberías tener un elemento `p` con un `id` de `display`.

```js
assert.isNotNull(document.querySelector('#drum-machine p#display'));
```

Dentro de tu elemento `#pad-bank` deberías tener nueve elementos de drum pad clickables, cada uno con una clase de `drum-pad`.

```js
assert.lengthOf(document.querySelectorAll('#pad-bank .drum-pad'), 9);
```

Cada `.drum-pad` debería tener una de las siguientes letras como `innerText`, en orden: `Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`.

```js
const pads = document.querySelectorAll('.drum-pad');
const letters = [`Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`];
for (let i = 0; i < 9; i++) {
  assert.equal(pads[i].innerText.trim(), letters[i]);
}
```

Cada `.drum-pad` debería tener un elemento `audio` que tenga una clase de `clip`, un atributo `src` que apunte a un clip de audio, y un `id` correspondiente al texto interior de su elemento padre `.drum-pad` (por ejemplo, `id="Q"`, `id="W"`, `id="E"`, etc.).

```js
const letters = [`Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`];
const audios = document.querySelectorAll('#drum-machine .drum-pad audio');

for (let i = 0; i < 9; i++) {
  assert.equal(audios[i].id, letters[i]);
  assert.isTrue(audios[i].classList.contains('clip'));
  assert.isTrue(audios[i].hasAttribute('src'));
}
```

Cuando haces clic en un elemento `.drum-pad`, el clip de audio contenido en su elemento hijo de audio debería activarse.

```js
const pads = document.querySelectorAll('.drum-pad');
assert.isNotEmpty(pads);
pads.forEach(el => {
  const audio = el.querySelector('audio');
  el.click();
  assert.isFalse(audio.paused);
  audio.pause();
});
```

Cuando presionas una de las teclas Q, W, E, A, S, D, Z, X, C en tu teclado, el elemento `audio` correspondiente debería reproducir el sonido correspondiente.

```js
const letters = [`Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`];
const audios = document.querySelectorAll('#drum-machine .drum-pad audio');

for (let i = 0; i < 9; i++) {
  const audio = audios[i];
  document.dispatchEvent(
    new KeyboardEvent('keydown', { key: letters[i].toLowerCase() })
  );
  assert.isFalse(audio.paused);
  audio.pause();
}
```

Cuando se activa un `.drum-pad`, debes mostrar una cadena describiendo el clip de audio asociado como texto interior del elemento `#display` (cada cadena debe ser única).

```js
const collection = new Set();
const disp = document.querySelector('#display');
const pads = document.querySelectorAll('.drum-pad');
assert.lengthOf(pads, 9);
pads.forEach(el => {
  const audio = el.querySelector('audio');
  el.click();
  collection.add(disp.innerText.trim());
  audio.pause();
});
assert.lengthOf(collection, 9);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Drum Machine</title>
  </head>
  <body></body>
</html>
```

```css

```

```js

```

# --solutions--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Drum Machine</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div class="inner-container" id="drum-machine">
      <div id="pad-bank"></div>

      <div class="controls-container">
        <div class="control power-control">
          <p>Power: <span id="power-status">On</span></p>
          <div class="select">
            <div class="inner"></div>
          </div>
        </div>

        <!-- Display for Drum Pad Name and Volume -->
        <p id="display"></p>

        <div class="volume-slider">
          <input type="range" min="0" max="1" step="0.01" value="0.3" />
        </div>

        <div class="control bank-control">
          <p>Bank: <span id="current-bank">Heater Kit</span></p>
          <div class="select">
            <div class="inner"></div>
          </div>
        </div>
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
/* General Reset */
* {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

body {
  font-family: 'Arial', sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
  background-color: #282c34;
  color: white;
}

.inner-container {
  display: flex;
  flex-direction: column;
  align-items: center;
  background: #3b3b5f;
  border-radius: 8px;
  padding: 20px;
  width: 300px;
}

/* Drum Pad Styles */
#pad-bank {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 10px;
  margin-bottom: 15px;
}

.drum-pad {
  display: flex;
  align-items: center;
  justify-content: center;
  width: 70px;
  height: 70px;
  background: #444;
  border-radius: 5px;
  font-size: 1.2em;
  cursor: pointer;
}

.drum-pad.active {
  background: #ffaa33;
}

.controls-container {
  display: flex;
  flex-direction: column;
  align-items: center;
  width: 100%;
}

.control {
  display: flex;
  justify-content: space-between;
  width: 100%;
  margin: 10px 0;
}

.select {
  width: 50px;
  height: 20px;
  background: #555;
  border-radius: 10px;
  cursor: pointer;
  position: relative;
}

.select .inner {
  position: absolute;
  width: 18px;
  height: 18px;
  background: #ffaa33;
  border-radius: 50%;
  top: 1px;
  left: calc(100% - 18px);
}

#display {
  margin: 15px 0;
  text-align: center;
  font-size: 1em;
  padding: 20px;
  background: #444;
  border-radius: 5px;
  width: 100%;
}

.volume-slider input[type='range'] {
  width: 100%;
  -webkit-appearance: none;
  appearance: none;
  height: 4px;
  background: #ffaa33;
  border-radius: 2px;
  margin-bottom: 10px;
}

.volume-slider input[type='range']::-webkit-slider-thumb {
  -webkit-appearance: none;
  width: 12px;
  height: 12px;
  background: #fff;
  border-radius: 50%;
  cursor: pointer;
}
```

```js
const projectName = 'drum-machine';

const bankOne = [
  {
    keyCode: 81,
    keyTrigger: 'Q',
    id: 'Heater-1',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-1.mp3'
  },
  {
    keyCode: 87,
    keyTrigger: 'W',
    id: 'Heater-2',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-2.mp3'
  },
  {
    keyCode: 69,
    keyTrigger: 'E',
    id: 'Heater-3',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-3.mp3'
  },
  {
    keyCode: 65,
    keyTrigger: 'A',
    id: 'Heater-4',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-4_1.mp3'
  },
  {
    keyCode: 83,
    keyTrigger: 'S',
    id: 'Clap',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-6.mp3'
  },
  {
    keyCode: 68,
    keyTrigger: 'D',
    id: 'Open-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Dsc_Oh.mp3'
  },
  {
    keyCode: 90,
    keyTrigger: 'Z',
    id: "Kick-n'-Hat",
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Kick_n_Hat.mp3'
  },
  {
    keyCode: 88,
    keyTrigger: 'X',
    id: 'Kick',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/RP4_KICK_1.mp3'
  },
  {
    keyCode: 67,
    keyTrigger: 'C',
    id: 'Closed-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Cev_H2.mp3'
  }
];

const bankTwo = [
  {
    keyCode: 81,
    keyTrigger: 'Q',
    id: 'Chord-1',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Chord_1.mp3'
  },
  {
    keyCode: 87,
    keyTrigger: 'W',
    id: 'Chord-2',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Chord_2.mp3'
  },
  {
    keyCode: 69,
    keyTrigger: 'E',
    id: 'Chord-3',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Chord_3.mp3'
  },
  {
    keyCode: 65,
    keyTrigger: 'A',
    id: 'Shaker',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Give_us_a_light.mp3'
  },
  {
    keyCode: 83,
    keyTrigger: 'S',
    id: 'Open-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Dry_Ohh.mp3'
  },
  {
    keyCode: 68,
    keyTrigger: 'D',
    id: 'Closed-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Bld_H1.mp3'
  },
  {
    keyCode: 90,
    keyTrigger: 'Z',
    id: 'Punchy-Kick',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/punchy_kick_1.mp3'
  },
  {
    keyCode: 88,
    keyTrigger: 'X',
    id: 'Side-Stick',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/side_stick_1.mp3'
  },
  {
    keyCode: 67,
    keyTrigger: 'C',
    id: 'Snare',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Brk_Snr.mp3'
  }
];

const state = {
  power: true,
  display: ' ',
  currentPadBank: bankOne,
  currentPadBankId: 'Heater Kit',
  sliderVal: 0.3
};

const display = document.getElementById('display');

const setDisplay = text => {
  state.display = text;
  display.innerText = text;
};

const adjustVolume = e => {
  if (state.power) {
    state.sliderVal = e.target.value;
    setDisplay(`Volume: ${Math.round(e.target.value * 100)}`);
    setTimeout(() => setDisplay(' '), 1000);
    document.querySelectorAll('.clip').forEach(sound => {
      sound.volume = state.sliderVal;
    });
  }
};

const activatePad = pad => {
  pad.classList.add('active');
  setTimeout(() => pad.classList.remove('active'), 100);
};

const playSound = (id, name) => {
  if (!state.power) return;

  const sound = document.getElementById(id);
  sound.currentTime = 0;
  sound.play();
  activatePad(sound.parentNode);
  setDisplay(name);
};

const renderPadBank = () => {
  const padBank = document.querySelector('#pad-bank');
  padBank.innerHTML = '';

  state.currentPadBank.forEach(drum => {
    const pad = document.createElement('div');
    pad.classList.add('drum-pad');
    pad.id = drum.id;
    pad.innerText = drum.keyTrigger;

    const audio = document.createElement('audio');
    audio.classList.add('clip');
    audio.id = drum.keyTrigger;
    audio.src = drum.url;
    pad.appendChild(audio);

    pad.addEventListener('click', () =>
      playSound(drum.keyTrigger, drum.id.replace(/-/g, ' '))
    );
    document.addEventListener('keydown', e => {
      if (e.key === drum.keyTrigger.toLowerCase())
        playSound(drum.keyTrigger, drum.id.replace(/-/g, ' '));
    });

    padBank.appendChild(pad);
  });
};

const powerControl = () => {
  state.power = !state.power;
  document.getElementById('power-status').textContent = state.power
    ? 'On'
    : 'Off';

  setDisplay('');

  const powerSlider = document.querySelector('.power-control .select .inner');
  powerSlider.style.left = state.power ? 'calc(100% - 18px)' : '0';
};

const selectBank = () => {
  if (!state.power) return;
  setDisplay(' ');
  state.currentPadBank =
    state.currentPadBankId === 'Heater Kit' ? bankTwo : bankOne;
  state.currentPadBankId =
    state.currentPadBankId === 'Heater Kit' ? 'Smooth Piano Kit' : 'Heater Kit';

  renderPadBank();
  document.getElementById('current-bank').textContent =
    state.currentPadBankId === 'Heater Kit' ? 'Heater Kit' : 'Smooth Piano Kit';
  const bankControl = document.querySelector('.bank-control .select .inner');
  bankControl.style.left =
    state.currentPadBankId === 'Heater Kit' ? 'calc(100% - 18px)' : '0';
};

document
  .querySelector('.power-control')
  .addEventListener('click', powerControl);
document.querySelector('.bank-control').addEventListener('click', selectBank);
document
  .querySelector('.volume-slider input')
  .addEventListener('input', adjustVolume);

renderPadBank();
```
