---
id: 68cabc534d2e741f33a3fe1d
title: بناء متتبع إحصائيات شخصية لعبة
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

في هذا المختبر، ستبني متتبع إحصائيات شخصية لعبة. سيسمح البرنامج بإنشاء شخصية بخصائص محددة، تحديث تلك الخصائص، واسترجاع الإحصائيات الحالية للشخصية.

**الهدف:** أتمم قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. أنشئ صنفًا باسم `GameCharacter` يمثل شخصية لعبة ويدير إحصائيات الشخصية.
2. عند إنشاء كائن جديد من `GameCharacter` يجب أن يحتوي على الخصائص التالية:
    - `_name` يعين إلى السلسلة النصية المعطاة عند لحظة الإنشاء.
    - `_health` يعين إلى `100`.
    - `_mana` يعين إلى `50`.
    - `_level` يعين إلى `1`.
3. أنشئ خاصية `name` للوصول للقراءة فقط إلى اسم الشخصية.
4. بالنسبة لخاصية `health`:
     - عرّف دالة getter تُرجع الصحة الحالية.
     - عرّف دالة setter تمنع تعيين الصحة إلى أقل من `0`، وتحد أقصى الصحة عند `100`.
5. بالنسبة لخاصية `mana`:
     - عرّف دالة getter تُرجع المانا الحالية.
     - عرّف دالة setter تمنع تعيين المانا إلى أقل من `0`، وتحد أقصى المانا عند `50`.
6. أنشئ دالة getter لـ `level` تُرجع المستوى الحالي للشخصية.
7. عرّف طريقة باسم `level_up` التي:
    - تزيد مستوى الشخصية بمقدار 1.
    - تعيد تعيين الصحة إلى `100` والمانا إلى `50` باستخدام دوال setter الخاصة بهما.
    - تطبع رسالة بالشكل `<name> leveled up to <level>!` (حيث يجب استبدال `<name>` و `<level>` باسم الشخصية والمستوى الجديد، على التوالي).
8. عرّف دالة `__str__` تُرجع سلسلة نصية منسقة تشمل:
    - اسم الشخصية.
    - مستوى الشخصية.
    - الصحة الحالية للشخصية.
    - المانا الحالية للشخصية.
   على سبيل المثال، شخصية باسم `Kratos`، مباشرة بعد الإنشاء، يجب تمثيلها كما يلي:

   ```md
   Name: Kratos
   Level: 1
   Health: 100
   Mana: 50
   ```

## مثال على الاستخدام

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

يجب أن يكون لديك صنف `GameCharacter`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

يجب أن يكون لديك دالة `__init__` في صنف `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

يجب أن تحتوي دالة `__init__` على معلمتين، الأولى هي `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

يجب تعيين `self._name` إلى السلسلة النصية المعطاة عند لحظة الإنشاء داخل دالة `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

يجب تعيين `100` إلى `self._health` داخل دالة `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

يجب تعيين `50` إلى `self._mana` داخل دالة `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

يجب تعيين `1` إلى `self._level` داخل دالة `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

يجب أن يكون لديك دالة `name` في صنف `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

يجب أن تحتوي دالة `name` على معلمة واحدة فقط، وهي `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

يجب إرجاع `self._name` في دالة `name`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

يجب أن تحتوي دالة `name` على مزين `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

يجب أن يكون لديك دالة `health` في صنف `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

يجب أن تحتوي دالة `health` على معلمة واحدة فقط، وهي `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

يجب إرجاع `self._health` في دالة `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

يجب أن تحتوي دالة `health` على مزين `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

يجب إنشاء دالة setter لـ `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

يجب أن تحتوي دالة setter لـ `health` على مزين `@health.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

يجب أن تحتوي دالة setter لـ `health` على معلمتين، الأولى هي `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

يجب تعيين `self._health` إلى `0` إذا كانت القيمة المعطاة إلى setter الخاص بـ `health` أقل من `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

يجب تعيين `self._health` إلى `100` إذا كانت القيمة المعطاة إلى setter الخاص بـ `health` أكبر من `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

يجب تعيين `self._health` إلى القيمة المعطاة إلى setter الخاص بـ `health` عندما تكون القيمة بين `0` و `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
    char.health = 0
    assert char.health == 0
    char.health = 100
    assert char.health == 100
`)})
```

يجب أن يكون لديك دالة `mana` في صنف `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

يجب أن تحتوي دالة `mana` على معلمة واحدة فقط، وهي `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

يجب إرجاع `self._mana` في دالة `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

يجب أن تحتوي دالة `mana` على مزين `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

يجب إنشاء دالة setter لـ `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

يجب أن تحتوي دالة setter لـ `mana` على مزين `@mana.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

يجب أن تحتوي دالة setter لـ `mana` على معلمتين، الأولى هي `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

يجب تعيين `self._mana` إلى `0` إذا كانت القيمة المعطاة إلى setter الخاص بـ `mana` أقل من `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

يجب تعيين `self._mana` إلى `50` إذا كانت القيمة المعطاة إلى setter الخاص بـ `mana` أكبر من `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

يجب تعيين `self._mana` إلى القيمة المعطاة إلى setter الخاص بـ `mana` إذا كانت القيمة بين `0` و `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

يجب أن يكون لديك دالة `level` في صنف `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

يجب أن تحتوي دالة `level` على معلمة واحدة فقط، وهي `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

يجب إرجاع `self._level` في دالة `level`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

يجب أن تحتوي دالة `level` على مزين `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

يجب أن يكون لديك دالة `level_up` في صنف `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

يجب أن تحتوي دالة `level_up` على معلمة واحدة فقط، وهي `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

يجب زيادة `self._level` بمقدار `1` في دالة `level_up`.

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

يجب تعيين `self.health` إلى `100` في دالة `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

يجب تعيين `self.mana` إلى `50` في دالة `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

يجب أن تطبع دالة `level_up` الرسالة `<name> leveled up to <level>!` (حيث يجب استبدال `<name>` و `<level>` باسم الشخصية والمستوى الجديد، على التوالي).

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

يجب أن يكون لديك دالة `__str__` في صنف `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

يجب أن تحتوي دالة `__str__` على معلمة واحدة فقط، وهي `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

يجب أن تُرجع دالة `__str__` سلسلة نصية تحتوي على إحصائيات الشخصية باستخدام التنسيق المقدم.

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
