---
id: 66a83e5e491625454b6f62c3
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Форми складаються з `inputs`, де користувачі можуть вводити свої вхідні дані. Ви можете згрупувати пов’язані вхідні дані разом, використовуючи елемент `fieldset`.

Ось приклад використання елемента `fieldset`:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

Всередині вашого елемента `form` додайте елемент `fieldset`.

# --hints--

У вас має бути відкриваючий тег `fieldset`.

```js
assert.match(code, /<fieldset>/i);
```

У вас має бути закриваючий тег `fieldset`.

```js
assert.match(code, /<\/fieldset>/i);
```

Ваш елемент `fieldset` має бути всередині елемента `form`.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
