---
id: 66a84111965a0c46df6bbd0a
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

У попередніх уроках ви навчились пов’язувати елемент `label` з `input` таким чином:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

Атрибут `for` на елементі `label` має відповідати атрибуту `id` на елементі `input`. Це називається явним зв’язком.

Під вашим елементом `legend` додайте елемент `label` з текстом `Name (required):`. Встановіть його атрибут `for` зі значенням `"full-name"`.

Потім під вашим елементом `label` додайте елемент `input` без атрибутів. У наступних кроках ви додасте необхідні атрибути.

# --hints--

У вас має бути відкриваючий тег `label`.

```js
assert.match(code, /<label.*>/);
```

У вас має бути закриваючий тег `label`.

```js
assert.match(code, /<\/label>/);
```

Ваш елемент `label` має бути розміщений під вашим `legend` всередині елемента `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

Ваш елемент `label` має мати атрибут `for` зі значенням `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

Ваш елемент `label` має містити текст `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

У вас має бути елемент `input` під вашим елементом `label`.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
