---
id: 66a84111965a0c46df6bbd0a
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

In previous lessons, you learned how to associate a `label` element with an `input` like this:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

Атрибут `for` елемента `label` має відповідати атрибуту `id` елемента `input`. Такий спосіб називають явним зв’язуванням.

Додайте елемент `label` з текстом `Name (required):` під елементом `legend`. Встановіть значення його атрибута `for` на `"full-name"`.

Потім під елементом `label` додайте елемент `input` без атрибутів. Ви додасте потрібні атрибути в наступних кроках.

# --hints--

Ви повинні мати початковий тег `label`.

```js
assert.match(code, /<label.*>/);
```

Ви повинні мати кінцевий тег `label`.

```js
assert.match(code, /<\/label>/);
```

Елемент `label` має бути під `legend` в межах елемента `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

Елемент `label` повинен мати атрибут `for` зі значенням `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

Елемент `label` повинен мати текст `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

Ви повинні мати елемент `input` під елементом `label`.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--

--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
