---
id: 66a83601cd819e37f0dccd14
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з HTML-формами, створюючи форму зворотного зв’язку для готелю.

Для вступного тексту потрібно відобразити головний заголовок, а потім коротку нотатку про залишення відгуку.

Всередині вашого `body` елемента додайте `header` елемент.

Всередині `header` елемента додайте `h1` елемент. Текст `h1` елемента має бути `Hotel Feedback Form`.

Під вашим `h1` елементом додайте `p` елемент. Текст `p` елемента має бути `Thank you for staying with us. Please provide feedback on your recent stay.`.

# --hints--

У вас має бути відкриваючий тег `header`.

```js
assert.match(code, /<header>/i);
```

У вас має бути закриваючий тег `header`.

```js
assert.match(code, /<\/header>/i);
```

У вас має бути відкриваючий тег `h1`.

```js
assert.match(code, /<h1>/i);
```

У вас має бути закриваючий тег `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Ваш `h1` елемент має містити текст `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

Ваш `h1` елемент має бути всередині `header` елемента.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

У вас має бути відкриваючий тег `p`.

```js
assert.match(code, /<p>/i);
```

У вас має бути закриваючий тег `p`.

```js
assert.match(code, /<\/p>/i);
```

Ваш абзац має містити текст `Thank you for staying with us. Please provide feedback on your recent stay.`.

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Ваш абзац має бути всередині `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

Ваш абзац має йти після `h1`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
