---
id: 66a833e88d08593618f22285
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

В межах елемента `head` вкладіть елемент `meta` з атрибутом `charset` зі значенням `"UTF-8"`.

Додайте елемент `title` під цим елементом `meta`. Текстом елемента `title` має бути `Hotel Feedback Form`.

# --hints--

Ви повинні мати елемент `meta`.

```js
assert.isNotNull(document.querySelector("meta"));
```

Елемент `meta` є порожнім елементом, він не повинен мати кінцевий тег `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Тег `meta` повинен мати атрибут `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Атрибут `charset` повинен мати значення `"UTF-8"`.

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

Елемент `meta` має бути вкладеним в елементі `head`.

```js
assert.match(code, /<head>.*\s*<meta.*charset.*>.*\s*<\/head>/si);
```

Ви повинні мати початковий тег `title`.

```js
assert.match(code, /<title>/i);
```

Ви повинні мати кінцевий тег `title`.

```js
assert.match(code, /<\/title>/i);
```

Елемент `title` має бути вкладеним в елементі `head`.

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

Елемент `title` повинен містити текст `Hotel Feedback Form`. Можливо, доведеться перевірити написання.

```js
const titleText = document.querySelector('title')?.innerText.trim();
assert.strictEqual(titleText, "Hotel Feedback Form");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    --fcc-editable-region--
    <head>

    </head>
    --fcc-editable-region--
</html>
```
