---
id: 66a83fec026a7a4631e084d2
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Працюючи з елементами `fieldset`, часто використовують підпис, щоб описати групу вхідних даних. Для цього можна використати елемент `legend`.

Ось приклад використання елемента `legend`:

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

Всередині вашого елемента `fieldset` додайте елемент `legend` з текстом `Personal Information`.

# --hints--

У вас має бути відкриваючий тег `legend`.

```js
assert.match(code, /<legend>/);
```

У вас має бути закриваючий тег `legend`.

```js
assert.match(code, /<\/legend>/);
```

Ваш елемент `legend` має бути всередині елемента `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

Ваш елемент `legend` має містити текст `Personal Information`.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
