---
id: 66a83bdcf425e7446900b7c4
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У попередніх уроках ви навчились працювати з елементом `form` ось так:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

Атрибут `action` використовується, щоб вказати, куди слід відправляти дані форми при її відправленні.

Атрибут `method` використовується, щоб вказати HTTP-метод, який слід застосувати при відправленні даних форми. Найпоширеніші методи — `GET` та `POST`.

**ПРИМІТКА**: Ви дізнаєтесь, як працюють HTTP-методи, у наступних модулях.

Всередині вашого елемента `main` додайте елемент `form` з атрибутом `action`, встановленим у `"https://hotel-feedback.freecodecamp.org"`, та атрибутом `method`, встановленим у `"POST"`.

# --hints--

У вас має бути відкриваючий тег `form`.

```js
assert.match(code, /<form.*>/i);
```

У вас має бути закриваючий тег `form`.

```js
assert.match(code, /<\/form>/i);
```

Ваш елемент `form` має мати атрибут `action`, встановлений у `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Ваш елемент `form` має мати атрибут `method`, встановлений у `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
