---
id: 66a937e74920ba68ebe5e86d
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

У попередніх уроках ви навчились працювати з атрибутами `placeholder` та `required` ось так:

```html
<input type="text" placeholder="e.g., John Doe" required>
```

Для вашого існуючого елемента `input` додайте атрибут `placeholder` зі значенням `"e.g., John Doe"`.

Також додайте атрибут `required` до елемента `input`.

# --hints--

Ви не повинні видаляти існуючий елемент `input`.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Ваш елемент `input` має мати атрибут `placeholder` зі значенням `"e.g., John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

Ваш елемент `input` має мати атрибут `required`.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
