---
id: 66a93730719e1f68410cce54
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Атрибут `name` використовується для ідентифікації даних форми після їх відправлення на сервер.

Ось приклад, як використовувати атрибут `name`:

```html
<input type="email" name="email">
```

Додайте атрибут `name` до елемента `input` зі значенням `"name"`.

# --hints--

Не слід видаляти існуючий елемент `input`.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Не слід створювати новий елемент `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Ваш елемент `input` має мати атрибут `name` зі значенням `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
