---
id: 66a962954f4e0d76223b37ed
title: Крок 21
challengeType: 0
dashedName: step-21
---

# --description--

Далі додайте ще один прапорець `input` з атрибутами `id` і `value`, встановленими в `"location"`, а атрибут `name` — у `"choice"`.

Для елемента `label` текст `Location` і атрибут `for` мають бути встановлені в `"location"`.

Під елементом `label` додайте ще один прапорець `input` з атрибутами `id` і `value`, встановленими в `"reputation"`, а атрибут `name` — у `"choice"`.

Для елемента `label` текст `Reputation` і атрибут `for` мають бути встановлені в `"reputation"`.

# --hints--

У вас має бути третій прапорець `input`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

Ваш третій прапорець `input` має мати атрибут `id`, встановлений у `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

Ваш третій прапорець `input` має мати атрибут `name`, встановлений у `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'choice');
```

Ваш третій прапорець `input` має мати атрибут `value`, встановлений у `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

Після прапорця `label` має бути елемент `input` з текстом `Location`.

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

Ваш `label` має мати атрибут `for`, встановлений у `"location"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

Під `input` має бути четвертий прапорець `label`.

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

Ваш четвертий прапорець `input` має мати атрибут `id`, встановлений у `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

Ваш четвертий прапорець `input` має мати атрибут `name`, встановлений у `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'choice');
```

Ваш четвертий прапорець `input` має мати атрибут `value`, встановлений у `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

Під прапорцем репутації має бути елемент `label` з текстом `Reputation`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

Ваш `label` має мати атрибут `for`, встановлений у `"reputation"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
