---
id: 66a9419e2d18476c645ce693
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Число `input` використовується для створення числового поля введення.

Ось приклад числового поля введення:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

Атрибути `min` і `max` використовуються, щоб встановити мінімальне та максимальне значення, які можна ввести у поле введення.

Під вашим елементом `label` додайте поле введення з атрибутом `type`, встановленим у `"number"`, і ідентифікатором `id`, рівним `"age"`.

Атрибут `name` має бути встановлений у `"age"`, атрибут `min` — у `"3"`, а атрибут `max` — у `"100"`.

# --hints--

У вас має бути числове поле введення `input` під вашим елементом `label`.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

Ваше числове поле введення `input` має мати ідентифікатор `id`, рівний `age`.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

Ваше числове поле введення `input` має мати атрибут `name`, встановлений у `age`.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

Ваше числове поле введення `input` має мати атрибут `min`, встановлений у `"3"`.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

Ваше числове поле введення `input` має мати атрибут `max`, встановлений у `"100"`.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
