---
id: 66a9521bc70162712caf118d
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Якщо ви хочете, щоб користувачі вибирали одну опцію зі списку варіантів, ви можете використати множину радіокнопок.

Ось приклад двох радіокнопок:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

У цьому прикладі радіокнопки згруповані разом за допомогою однакового значення атрибута `name`. Це означає, що одночасно можна вибрати лише одну радіокнопку.

Під вашим елементом `legend` додайте кнопку `radio` з атрибутом `id`, встановленим у `"yes-option"`, атрибутом `name`, встановленим у `"hotel-stay"`, і атрибутом `value`, встановленим у `"yes"`.

Під вашою кнопкою `radio` додайте елемент `label` з текстом `Yes` і атрибутом `for`, встановленим у `"yes-option"`.

# --hints--

У вас має бути кнопка `radio` під елементом `legend`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

Ваша кнопка `radio` має мати атрибут `id`, встановлений у `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

Ваша кнопка `radio` має мати атрибут `name`, встановлений у `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

Ваша радіокнопка має мати атрибут `value`, встановлений у `"yes"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

Під вашою кнопкою `label` має бути елемент `radio`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

Ваш елемент `label` має мати атрибут `for`, встановлений у `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

Ваш елемент `label` має містити текст `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
