---
id: 66a954b2bcddba72076c1857
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Під вашим елементом `label` додайте кнопку `radio` з множиною `id`, встановленою на `"no-option"`, атрибутом `name`, встановленим на `"hotel-stay"`, і атрибутом `value`, встановленим на `"no"`.

Під вашою новою кнопкою `radio` додайте ще один елемент `label` з атрибутом `for`, встановленим на `"no-option"`. Текст для `label` має бути `No`.

Коли закінчите, ви зможете спробувати радіокнопки, вибираючи по одній опції за раз.

# --hints--

Під вашим першим елементом `input` має бути елемент `label`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

Ваш елемент `input` має мати атрибут `type` зі значенням `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

Ваш елемент `input` має мати атрибут `id` зі значенням `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

Ваш елемент `input` має мати атрибут `name` зі значенням `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

Ваш елемент input має мати атрибут `value`, встановлений на `"no"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) label + input')?.value,
  'no'
);
```

Під вашим другим елементом `label` має бути другий елемент `input`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

Ваш елемент `label` має мати атрибут `for`, встановлений на `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

Ваш елемент `label` має містити текст `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
