---
id: 66a9401c9d660d6bb15993e2
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

У вашій формі відгуку про готель має бути опція для користувачів, щоб додати свій вік.

Почніть із додавання елемента `label` з текстом `Age (optional):` до форми.

Атрибут `for` має бути встановлений у `"age"`.

# --hints--

У вас має бути елемент `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

Ваш елемент `label` має мати атрибут `for`, встановлений у `"age"`.

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

Ваш елемент `label` має містити текст `Age (optional):`.

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
