---
id: 66a94ea5df66236ebed260e8
title: Крок 18
challengeType: 0
dashedName: step-18
---

# --description--

У наступному розділі форми потрібно запитати користувачів, чи вони зупинялись в цьому готелі раніше.

Для початку додайте елемент `fieldset`.

Додайте елемент `legend` з текстом `Was this your first time at our hotel?` до елемента `fieldset`.

# --hints--

Ви повинні мати елемент `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 2);
```

Ви повинні мати елемент `legend` в межах елемента `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset legend'), 2);
```

Елемент `legend` повинен містити текст `Was this your first time at our hotel?`.

```js
assert.strictEqual(document.querySelectorAll('fieldset legend')[1]?.textContent.trim(), 'Was this your first time at our hotel?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
      --fcc-editable-region--

      --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
