---
id: 66a93c95bc58e26a8fe95818
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Далі додайте `input` з типом `"email"` під вашим елементом е-пошти `label`. Цей `input` має мати атрибут `id` зі значенням `"email"`. Також встановіть атрибут `name` зі значенням `"email"`.

Цей `input` також є обов’язковим, тому обов’язково додайте атрибут `required`.

Нарешті, додайте атрибут `placeholder` зі значенням `"example@email.com"`.

# --hints--

У вас має бути елемент `input` під вашим елементом е-пошти `label`.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

Ваш елемент `input` має мати атрибут `type` зі значенням `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

Ваш елемент `input` має мати атрибут `id` зі значенням `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

Ваш елемент `input` має мати атрибут `name` зі значенням `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

Ваш елемент `input` має мати атрибут `required`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

Ваш елемент `input` має мати атрибут `placeholder` зі значенням `"example@email.com"`.
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
