---
id: 66a95d0eff8168747805f1f3
title: Крок 19
challengeType: 0
dashedName: step-19
---

# --description--

Коли ви хочете, щоб користувач вибрав кілька варіантів зі списку, можна використовувати прапорці.

Ось приклад, як працювати з прапорцями на прикладі варіантів їжі:

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

Атрибут `value` використовується, щоб вказати значення, яке буде надіслано на сервер при відправленні форми.

Під вашим елементом `legend` додайте прапорець `input` з атрибутами `id` і `value`, встановленими в `"ads"`, а атрибут `name` встановіть у `"choice"`.

Під вашим прапорцем `input` додайте елемент `label` з текстом `Social Media Ads`. Атрибут `for` має бути встановлений у `"ads"`.

# --hints--

У вас має бути прапорець `input` під вашим елементом `legend`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

Ваш прапорець `input` має мати атрибут `id`, встановлений у `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

Ваш прапорець `input` має мати атрибут `name`, встановлений у `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "choice");
```

Ваш прапорець `input` має мати атрибут `value`, встановлений у `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

Під вашим прапорцем `label` має бути елемент `input`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

Ваш елемент `label` має містити текст `Social Media Ads`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

Ваш елемент `label` має мати атрибут `for`, встановлений у `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
