---
id: 66a930b20f589b6664c51cb0
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Коли користувач вводить своє повне ім’я, `input` прийматиме звичайний текст.

У попередніх уроках ви навчились працювати з атрибутом `type` таким чином:

```html
<input type="text">
```

Для вашого існуючого елемента `input` додайте атрибут `type` зі значенням `"text"`.

Також додайте елементу `input` атрибут `id` зі значенням `"full-name"`.

# --hints--

Ви не повинні видаляти існуючий елемент `input`.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Ви не повинні створювати новий елемент `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Ваш елемент `input` має мати атрибут `type` зі значенням `"text"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

Ваш елемент `input` має мати атрибут `id` зі значенням `"full-name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
