---
id: 66a93bbe65a26169dbf3bc39
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Ваша форма зворотного зв’язку для готелю також має збирати адресу е-пошти користувача.

Почніть із додавання нового `label` елемента з текстом `Email address (required):` до форми. Ваш `label` елемент має мати атрибут `for`, встановлений у значення `"email"`.

# --hints--

У вас має бути `label` елемент.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

Ваш `label` елемент має містити текст `Email address (required):`.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

Ваш `label` елемент має мати атрибут `for`, встановлений у значення `"email"`.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
