---
id: 67a51d3a8a6fe123b77b0c6e
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Елементи `input` можуть мати атрибут `size`. Цей атрибут визначає кількість символів, які користувач буде бачити під час введення даних. Значенням атрибута `size` має бути невід’ємне число, більше за нуль. Якщо значення атрибута `size` не вказано або воно не дійсне, то елемент матиме ширину за замовчуванням браузера.

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

Надайте атрибут `size` зі значенням `"20"` тим елементам `input`, які стосуються імені та електронної адреси.

# --hints--

Надайте елементу `input`, що стосується імені, атрибут `size` зі значенням `"20"`.

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

Надайте елементу `input`, що стосується електронної адреси, атрибут `size` зі значенням `"20"`.

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
