---
id: 67a51d3a8a6fe123b77b0c6e
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Елементи `input` можуть мати атрибут `size`. Цей атрибут визначає кількість символів, які мають бути видимими під час введення користувачем у поле вводу. Значення `size` має бути невід’ємним цілим числом більшим за нуль. Якщо `size` не вказано або вказано з некоректним значенням, ширина поля вводу буде встановлена за замовчуванням браузером.

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

Додайте атрибут `size` з значенням `"20"` для полів введення імені та е-пошти.

# --hints--

Ви маєте додати атрибут `input` зі значенням `size` для `"20"` імені.

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

Ви маєте додати атрибут `input` зі значенням `size` для `"20"` е-пошти.

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
