---
id: 657a0ea50da0c8d9d6d7950a
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Projekt lernen Sie algorithmisches Denken, indem Sie ein Würfelspiel erstellen. Es gibt insgesamt 6 Runden, und in jeder Runde kann der Spieler bis zu 3 Mal würfeln und Punkte sammeln.

Das HTML und CSS wurden für Sie bereitgestellt. Erkunden Sie diese gerne. Wenn Sie bereit sind, müssen Sie Ihre HTML-Variablen einrichten. Holen Sie alle Ihre `.die`-Elemente und weisen Sie sie einer `listOfAllDice`-Variable zu. Holen Sie Ihre Punkte-Eingaben (die `input`-Elemente in Ihrem `#score-options`) und Punkte-Spans und weisen Sie sie `scoreInputs` und `scoreSpans` zu. Weisen Sie das `#current-round`-Element `roundElement` zu und das `#current-round-rolls`-Element `rollsElement`, dann machen Sie dasselbe für Ihre `#total-score`- und `#score-history`-Elemente. Weisen Sie Ihre `#roll-dice-btn`, `#keep-score-btn`, `#rules-btn` und `.rules-container` Variablen mit korrekt formatierten Namen zu.

Wenn der Benutzer auf die `Show rules`-Schaltfläche klickt, sollte er zwischen dem Anzeigen und Verbergen der Spielregeln umschalten können. Erstellen Sie eine Variable `isModalShowing`, um den Zustand der Spielregeln zu verfolgen.

Jedes Mal, wenn der Benutzer würfelt, müssen Sie alle Würfelwerte verfolgen. Erstellen Sie eine Variable `diceValuesArr`, um dies zu verfolgen.

Während des Spiels müssen Sie den aktuellen Punktestand, die Gesamtpunktzahl, die Anzahl der Würfe und die Runde, in der sich der Spieler befindet, verfolgen. Deklarieren Sie `rolls`, `score` und `round`-Variablen zu diesem Zweck.

Überlegen Sie, welchen Anfangswert jede dieser Variablen haben sollte. Alle diese Variablen sollten neu zuweisbar sein.

# --hints--

Sie sollten Ihre `.die`-Elemente `listOfAllDice` zuweisen.

```js
assert.deepEqual(listOfAllDice, document.querySelectorAll(".die"));
```

Sie sollten Ihre `#score-options input`-Elemente `scoreInputs` zuweisen.

```js
assert.deepEqual(scoreInputs, document.querySelectorAll("#score-options input"));
```

Sie sollten Ihre `#score-options span`-Elemente `scoreSpans` zuweisen.

```js
assert.deepEqual(scoreSpans, document.querySelectorAll("#score-options span"));
```

Sie sollten Ihr `#current-round`-Element `roundElement` zuweisen.

```js
assert.deepInclude([document.getElementById("current-round"), document.querySelector("#current-round")], roundElement);
```

Sie sollten Ihr `#current-round-rolls`-Element `rollsElement` zuweisen.

```js
assert.deepInclude([document.getElementById("current-round-rolls"), document.querySelector("#current-round-rolls")], rollsElement);
```

Sie sollten Ihr `#total-score`-Element `totalScoreElement` zuweisen.

```js
assert.deepInclude([document.getElementById("total-score"), document.querySelector("#total-score")], totalScoreElement);
```

Sie sollten Ihr `#score-history`-Element `scoreHistory` zuweisen.

```js
assert.deepInclude([document.getElementById("score-history"), document.querySelector("#score-history")], scoreHistory);
```

Sie sollten Ihr `#roll-dice-btn`-Element `rollDiceBtn` zuweisen.

```js
assert.deepInclude([document.getElementById("roll-dice-btn"), document.querySelector("#roll-dice-btn")], rollDiceBtn);
```

Sie sollten Ihr `#keep-score-btn`-Element `keepScoreBtn` zuweisen.

```js
assert.deepInclude([document.getElementById("keep-score-btn"), document.querySelector("#keep-score-btn")], keepScoreBtn);
```

Sie sollten Ihr `#rules-btn`-Element `rulesBtn` zuweisen.

```js
assert.deepInclude([document.getElementById("rules-btn"), document.querySelector("#rules-btn")], rulesBtn);
```

Sie sollten Ihr `.rules-container`-Element `rulesContainer` zuweisen.

```js
assert.deepEqual(rulesContainer, document.querySelector(".rules-container"));
```

Ihre `isModalShowing`-Variable sollte den Wert `false` haben.

```js
assert.isFalse(isModalShowing);
```

Ihre `diceValuesArr` sollte ein leeres Array als Wert haben.

```js
assert.isArray(diceValuesArr);
assert.isEmpty(diceValuesArr);
```

Ihre `rolls`-Variable sollte den Wert `0` haben.

```js
assert.strictEqual(rolls, 0);
```

Ihre `score`-Variable sollte den Wert `0` haben.

```js
assert.strictEqual(score, 0);
```

Ihre `round`-Variable sollte den Wert `1` haben.

```js
assert.strictEqual(round, 1);
```

Alle Ihre neuen Variablen sollten mit `let` deklariert werden.

```js
assert.match(code, /let isModalShowing/);
assert.match(code, /let diceValuesArr/);
assert.match(code, /let rolls/);
assert.match(code, /let score/);
assert.match(code, /let round/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Advanced Dice Game</title>
  <link rel="stylesheet" href="styles.css" />
</head>

<body>
  <header>
    <h1>Advanced Dice Game</h1>
    <button class="btn" id="rules-btn" type="button">Show rules</button>
    <div class="rules-container">
      <h2>Rules</h2>
      <ul>
        <li>There are total of six rounds</li>
        <li>You can only roll the dice three times per round</li>
        <li>To start the game, roll the dice</li>
        <li>
          Then, choose from one of the selected scores or roll the dice again
        </li>
        <li>
          If you choose a selected score, then you will move to the next round
        </li>
        <li>
          If you decline to choose a selected score, then you can roll the
          dice again two more times
        </li>
      </ul>
      <h2 class="points">Points</h2>
      <ul>
        <li>Three of a kind: Sum of all five dice</li>
        <li>Four of a kind: Sum of all five dice</li>
        <li>Full house: Three of a kind and a pair - 25 points</li>
        <li>
          Small straight: Four of the dice have consecutive values - 30 points
        </li>
        <li>
          Large straight: All five dice have consecutive values - 40 points
        </li>
      </ul>
    </div>
  </header>

  <main>
    <form id="game">
      <div id="dice">
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
      </div>

      <p class="rounds-text">
        <strong>Rolls:</strong> <span id="current-round-rolls">0</span> |
        <strong>Round:</strong> <span id="current-round">1</span>
      </p>

      <div id="score-options">
        <div>
          <input type="radio" name="score-options" id="three-of-a-kind" value="three-of-a-kind" disabled />
          <label for="three-of-a-kind">Three of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="four-of-a-kind" value="four-of-a-kind" disabled />
          <label for="four-of-a-kind">Four of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="full-house" value="full-house" disabled />
          <label for="full-house">Full house<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="small-straight" value="small-straight" disabled />
          <label for="small-straight">Small straight<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="large-straight" value="large-straight" disabled />
          <label for="large-straight">Large straight<span></span></label>
        </div>

        <div>
          <input type="radio" name="score-options" id="none" value="none" disabled />
          <label for="none">None of the above<span></span></label>
        </div>
      </div>

      <button class="btn" id="keep-score-btn" type="button">
        Keep the above selected score
      </button>
      <button class="btn" id="roll-dice-btn" type="button">
        Roll the dice
      </button>
    </form>

    <div id="scores">
      <h3>Score history (Total score: <span id="total-score">0</span>)</h3>
      <ol id="score-history"></ol>
    </div>
  </main>
  <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

header {
  color: var(--light-grey);
  text-align: center;
}

h1 {
  font-size: 2.5rem;
  margin: 25px 0;
}

.rules-container {
  display: none;
  background-color: var(--light-grey);
  color: var(--black);
  width: 50%;
  margin: 20px auto;
  height: 300px;
  border-radius: 10px;
  overflow-y: scroll;
}

.rules-container ul {
  list-style-type: none;
}

.points {
  margin-top: 15px;
}

main {
  background-color: var(--light-grey);
  padding: 20px;
  display: grid;
  grid-template-columns: repeat(auto-fit, minmax(500px, 1fr));
  gap: 1rem;
  margin: auto;
  justify-items: center;
  width: 50%;
  border-radius: 10px;
}

#dice {
  display: flex;
  justify-content: space-around;
  margin-bottom: 15px;
}

.die {
  width: 40px;
  height: 40px;
  text-align: center;
  margin-right: 15px;
  border: 4px solid var(--black);
  padding: 10px;
}

.rounds-text {
  text-align: center;
}

input[type="radio"]:disabled + label {
  color: var(--grey);
}

#score-history {
  margin-top: 15px;
  text-align: center;
  list-style-position: inside;
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px 0 10px 0.5rem;
  color: var(--black);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover:enabled {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

@media (max-width: 992px) {
  main {
    width: 100%;
  }
}

@media (max-width: 500px) {
  .btn {
    width: 120px;
  }
}

```

```js
--fcc-editable-region--

--fcc-editable-region--
```
