---
id: 6a89c69cbe1bb69845992d8b
title: 숫자 범위 생성기 만들기
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

이 실습에서는 재귀를 사용해 지정된 범위 내 숫자 목록을 생성하는 `range_of_numbers` 함수를 만드세요.

예를 들어:

- `range_of_numbers(3, 9)`는 `[3, 4, 5, 6, 7, 8, 9]`를 반환해야 합니다.
- `range_of_numbers(5, 5)`는 `[5]`를 반환해야 합니다.

아래 스토리를 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `range_of_numbers`라는 이름의 함수를 만들고 `start_num`과 `end_num` 두 매개변수를 사용하세요.
1. 함수는 `start_num`부터 `end_num`까지 연속된 정수 목록을 반환해야 하며, 두 값 모두 포함해야 합니다.
1. `start_num` 값은 항상 `end_num`보다 작거나 같습니다.
1. 함수는 자기 자신을 호출하는 재귀를 사용해야 합니다.
1. 함수는 `for` 또는 `while` 루프, 리스트 내포, `range()` 함수를 사용하지 않아야 합니다.
1. `start_num`이 `end_num`과 같을 때는 `start_num`을 포함하는 리스트를 반환해야 하며, 이것이 기본 사례입니다.
1. 재귀 사례에서는 기본 사례로 가는 인수를 사용해 자기 자신을 호출한 뒤, 현재 숫자를 반환된 리스트에 추가해야 합니다.
1. 반환된 리스트를 저장하기 위해 전역 변수를 사용하지 않아야 합니다.

# --hints--

`start_num`과 `end_num` 매개변수를 가진 `range_of_numbers` 함수를 만들어야 합니다.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)`는 `[1, 2, 3, 4, 5]`를 반환해야 합니다.

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)`는 `[6, 7, 8, 9]`를 반환해야 합니다.

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)`는 `[4]`를 반환해야 합니다.

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)`는 `[10, 11, 12, 13, 14, 15]`를 반환해야 합니다.

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

`range_of_numbers` 함수는 자기 자신을 호출해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

`range_of_numbers` 함수는 루프, 리스트 내포, `range()` 함수를 사용하지 않아야 합니다.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

`range_of_numbers(1, 3)` 호출이 이후 호출에 영향을 주지 않아야 합니다. 그 호출 후에도 `range_of_numbers(6, 9)`는 여전히 `[6, 7, 8, 9]`를 반환해야 합니다.

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
