---
id: 0095d0567b6817114989a1dd
title: passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai implementar o algoritmo do caminho mais curto em JavaScript. Você vai escrever uma função que calcula o caminho mais curto entre nós em um grafo ponderado e retorna tanto as distâncias quanto os caminhos percorridos.

Por exemplo, dado um grafo onde cidades estão conectadas por estradas com diferentes distâncias, o algoritmo vai encontrar a rota mais curta de uma cidade para outra. Se você quiser viajar da Cidade A para a Cidade D, o algoritmo pode descobrir que ir A ⇨ B ⇨ C ⇨ D (total: 15km) é mais curto do que ir diretamente A ⇨ D (20km).

Para começar, defina uma variável chamada `INF` e atribua a ela o valor `Infinity`. Depois, você vai usá-la para indicar que não há conexão direta entre dois nós.

# --hints--

Você deve definir uma variável chamada `INF`.

```js
assert.isDefined(INF);
```

Sua variável `INF` deve ter o valor `Infinity`.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
