---
id: 69c918841987bf16c1d1babd
title: passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Você precisa acompanhar a menor distância conhecida do nó inicial até cada outro nó. Para começar, você vai assumir que todos os nós estão a uma distância infinita.

Em JavaScript, você pode criar um array pré-preenchido com um valor usando `new Array(n).fill(value)`. Por exemplo, `new Array(3).fill(0)` cria `[0, 0, 0]`.

Dentro da função `shortestPath`, crie uma variável chamada `distances` e atribua a ela `new Array(n).fill(INF)`.

# --hints--

Você deve criar uma variável chamada `distances` dentro de `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

Sua variável `distances` deve ser atribuída a `new Array(n).fill(INF)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
