---
id: 69c918841987bf16c1d1bac5
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Após escanear todos os nós para buscar o nó não visitado mais próximo, você precisa tratar o caso em que nenhum nó válido é encontrado. Isso acontece quando todos os nós restantes estão inacessíveis.

Depois do loop interno `for`, mas ainda dentro do loop externo, adicione uma declaração `if` que verifica se `current` é estritamente igual a `-1`. Se isso for verdade, use `break` dentro do bloco `if` para interromper o loop externo antecipadamente.

# --hints--

Você deve adicionar uma declaração `if` que verifica se `current === -1`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

Você deve usar `break` dentro do bloco `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
