---
id: 69c918841987bf16c1d1bac4
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Se a condição for verdadeira, o nó atual é o melhor candidato não visitado encontrado até agora. Você precisa atualizar suas variáveis de rastreamento para refletir isso.

Dentro da declaração `if`, atualize `minDistance` para `distances[nodeNo]` e `current` para `nodeNo`.

# --hints--

Você deve definir `minDistance` como `distances[nodeNo]` dentro do bloco `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /minDistance\s*=\s*distances\s*\[\s*nodeNo\s*\]/
);
```

Você deve definir `current` como `nodeNo` dentro do bloco `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /current\s*=\s*nodeNo/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
