---
id: 69c918841987bf16c1d1bac8
title: Passo 17
challengeType: 1
dashedName: step-17
---

# --description--

Antes de atualizar as distâncias, você precisa verificar se o vizinho vale a pena ser considerado. Deve haver uma aresta real para ele (`distance !== INF`) e ele não pode ter sido visitado ainda.

Dentro do loop `for` dos vizinhos, adicione uma declaração `if` que verifique se `distance` não é estritamente igual a `INF` **e** se `nodeNo` não foi visitado.

Dentro do bloco `if`, declare uma variável `newDistance` atribuída a `distances[current] + distance`. Esse é o custo total de alcançar o vizinho através do nó atual.

# --hints--

Você deve adicionar uma declaração `if` com a condição `distance !== INF && !visited[nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

Você deve declarar uma `newDistance` com o valor `distances[current] + distance` dentro do bloco `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
