---
id: 69c937c73f23b1140bcd297a
title: passo 4
challengeType: 1
dashedName: step-4
---

# --description--

O parâmetro `targetNode` é opcional. Quando não fornecido, a função calculará os caminhos mais curtos de `startNode` para todos os outros nós no grafo.

Dê ao parâmetro `targetNode` um valor padrão de `null`.

# --hints--

O parâmetro `targetNode` da sua função `shortestPath` deve ter como padrão `null`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {

}
--fcc-editable-region--
```
