---
id: 69c918841987bf16c1d1babb
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Agora você vai criar a função principal que faz todo o trabalho.

Crie uma função chamada `shortestPath` que receba três parâmetros: `matrix`, `startNode` e `targetNode`.

Nos próximos passos, você vai desenvolver a lógica dessa função.

# --hints--

Você deve criar uma função chamada `shortestPath`.

```js
assert.isFunction(shortestPath);
```

Sua função `shortestPath` deve ter os parâmetros `matrix`, `startNode` e `targetNode`.

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
