---
id: 69c918841987bf16c1d1bac2
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Dentro do loop, você precisa escanear cada nó para buscar aquele com a menor distância conhecida que ainda não foi visitado.

Adicione um loop `for` dentro do loop principal que itere por cada nó. Use `nodeNo` como variável do loop indo de `0` até `n`.

Deixe o corpo vazio por enquanto.

# --hints--

Você deve ter um segundo loop `for` (interno) dentro do loop principal com `nodeNo` como variável do loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+nodeNo\s*=\s*0\s*;\s*nodeNo\s*<\s*n\s*;\s*nodeNo\+\+\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

--fcc-editable-region--
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;

  }
--fcc-editable-region--
}
```
