---
id: 69c918841987bf16c1d1babf
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Você também precisa acompanhar o caminho real percorrido para alcançar cada nó. Você vai armazenar uma lista de caminhos onde cada entrada é um array de índices de nós representando a rota tomada.

Inicialmente, o caminho de cada nó conterá apenas ele mesmo. Você pode criar essa estrutura usando `Array.from()`:

```js
Array.from({ length: n }, (_, i) => [i]);
```

Isso cria um array de `n` elementos, onde o elemento no índice `i` é `[i]`.

Ainda dentro da função `shortestPath`, crie uma variável chamada `paths` e atribua a ela `Array.from({ length: n }, (_, i) => [i])`.

# --hints--

Você deve criar uma variável `paths` dentro da função `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

Sua variável `paths` deve ser inicializada usando `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

Sua variável `paths` deve usar `{ length: n }` como o primeiro argumento para `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

O segundo argumento para `Array.from` deve ser um callback `(_, i) => [i]`, onde `_` é um marcador de posição para o elemento não usado e `i` é o índice. Ele deve retornar `[i]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
