---
id: 69c918841987bf16c1d1bac1
title: passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Agora você vai adicionar o loop principal que conduz o algoritmo. Ele é executado uma vez para cada nó no grafo, selecionando o nó não visitado mais próximo a cada vez.

Dentro da função `shortestPath`, adicione um loop `for` que execute `n` vezes usando `i` como variável do loop.

No início do corpo do loop, usando `let`, declare duas variáveis `minDistance` e `current` definidas como `INF` e `-1`, respectivamente. Elas vão acompanhar o nó não visitado mais próximo encontrado em cada iteração.

# --hints--

Você deve ter um loop `for` dentro de `shortestPath` que execute `n` vezes. Não esqueça de definir `i` com `let` dentro dele.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

Você deve ter uma variável `minDistance` definida como `INF` dentro do loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

Você deve ter uma variável `current` definida como `-1` dentro do loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
