---
id: 69c918841987bf16c1d1bac7
title: passo 16
challengeType: 1
dashedName: step-16
---

# --description--

Agora que o nó atual está marcado como visitado, você precisa olhar para todos os seus vizinhos para ver se consegue alcançá-los de forma mais eficiente através do nó atual.

Após `visited[current] = true`, adicione um loop de vizinhos `for` que itere por cada `nodeNo` de `0` até `n`. Note que este é outro loop interno no mesmo nível do primeiro.

Dentro do loop, declare uma variável chamada `distance` e atribua a ela `matrix[current][nodeNo]`. Isso fornece o peso da aresta entre o nó atual e seu vizinho.

# --hints--

Você deve adicionar um loop de vizinhos `for` após marcar o nó atual como visitado.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

O loop de vizinhos `for` deve iterar de `0` até `n`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

Você deve ter uma variável `distance` atribuída a `matrix[current][nodeNo]` dentro do loop de vizinhos.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
