---
id: 69c918841987bf16c1d1bac0
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

À medida que o algoritmo executa, você precisa acompanhar quais nós já foram processados para não visitá-los novamente.

Crie uma variável chamada `visited` e atribua a ela `new Array(n).fill(false)`.

# --hints--

Você deve criar uma variável chamada `visited`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+visited\s*=/
);
```

Sua variável `visited` deve ser atribuída a `new Array(n).fill(false)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*false\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
--fcc-editable-region--

--fcc-editable-region--
}
```
