---
id: 69c918841987bf16c1d1baba
title: passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Você vai precisar de um array 2D para representar a matriz de adjacência do grafo. Cada valor representa o peso (distância) da aresta entre dois nós. Um valor `INF` significa que não há uma aresta direta.

Crie outra variável chamada `adjMatrix` e atribua a ela um array 2D com os seguintes valores:

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Você deve criar uma variável chamada `adjMatrix`.

```js
assert.isDefined(adjMatrix);
```

`adjMatrix` deve ser um array 2D com 6 linhas e 6 colunas.

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix` deve conter os valores corretos.

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
