---
id: 587d7b7d367417b2b2512b1f
title: تعديل مصفوفة مخزنة في كائن
challengeType: 1
forumTopicId: 301163
dashedName: modify-an-array-stored-in-an-object
---

# --description--

الآن رأيت كل العمليات الأساسية لكائنات JavaScript. يمكنك إضافة، تعديل، وحذف أزواج المفتاح-القيمة، التحقق من وجود المفاتيح، والتكرار على كل المفاتيح في الكائن. مع استمرار تعلمك لـ JavaScript سترى تطبيقات أكثر تنوعًا للكائنات. بالإضافة إلى ذلك، دروس هياكل البيانات الموجودة في قسم التحضير لمقابلات البرمجة تغطي أيضًا كائنات ES6 <dfn>Map</dfn> و <dfn>Set</dfn>، وكلاهما مشابه للكائنات العادية لكنه يوفر ميزات إضافية. بعد أن تعلمت أساسيات المصفوفات والكائنات، أنت مستعد تمامًا لبدء حل مشكلات أكثر تعقيدًا باستخدام JavaScript!

# --instructions--

ألق نظرة على الكائن الذي وفرناه في محرر الشفرة. يحتوي كائن `user` على ثلاثة مفاتيح. يحتوي مفتاح `data` على خمسة مفاتيح، أحدها يحتوي على مصفوفة من `friends`. من هذا، يمكنك أن ترى مدى مرونة الكائنات كـ هياكل بيانات. بدأنا بكتابة دالة `addFriend`. أكمل كتابتها بحيث تأخذ كائن `user` وتضيف اسم المعلمة `friend` إلى المصفوفة المخزنة في `user.data.friends` وتُرجع تلك المصفوفة.

# --hints--

يجب أن يحتوي كائن `user` على مفاتيح `name` و `age` و `data`.

```js
assert('name' in user && 'age' in user && 'data' in user);
```

يجب أن تقبل دالة `addFriend` كائن `user` وسلسلة نصية `friend` كمعلمات وتضيف الصديق إلى مصفوفة `friends` في كائن `user`.

```js
assert(
  (function () {
    let L1 = user.data.friends.length;
    addFriend(user, 'Sean');
    let L2 = user.data.friends.length;
    return L2 === L1 + 1;
  })()
);
```

يجب أن تُرجع `addFriend(user, "Pete")` المصفوفة `["Sam", "Kira", "Tomo", "Pete"]`.

```js
assert.deepEqual(
  (function () {
    delete user.data.friends;
    user.data.friends = ['Sam', 'Kira', 'Tomo'];
    return addFriend(user, 'Pete');
  })(),
  ['Sam', 'Kira', 'Tomo', 'Pete']
);
```

# --seed--

## --seed-contents--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  // Only change code below this line

  // Only change code above this line
}

console.log(addFriend(user, 'Pete'));
```

# --solutions--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  userObj.data.friends.push(friend);
  return userObj.data.friends;
}
```
