---
id: 587d7b7d367417b2b2512b1c
title: تحقق مما إذا كان الكائن يحتوي على خاصية
challengeType: 1
forumTopicId: 301155
dashedName: check-if-an-object-has-a-property
---

# --description--

يمكننا الآن إضافة المفاتيح وتعديلها وحذفها من الكائنات. لكن ماذا لو أردنا فقط معرفة ما إذا كان الكائن يحتوي على خاصية معينة؟ توفر JavaScript طريقتين مختلفتين للقيام بذلك. تستخدم إحداهما الطريقة `hasOwnProperty()` والأخرى الكلمة المفتاحية `in`. إذا كان لدينا كائن `users` يحتوي على خاصية `Alan`، يمكننا التحقق من وجودها بإحدى الطريقتين التاليتين:

```js
users.hasOwnProperty('Alan');
'Alan' in users;
```

كلا الطريقتين سترجعان القيمة `true`.

# --instructions--

أكمل كتابة الدالة بحيث تُرجع `true` إذا كان الكائن الممرر إليها يحتوي على الأسماء الأربعة كلها: `Alan` و`Jeff` و`Sarah` و`Ryan`، وتُرجع `false` خلاف ذلك.

# --hints--

يجب عدم الوصول إلى كائن `users` مباشرة

```js 

assert(__helpers.removeJSComments(code).match(/users/gm).length <= 2)

```

يجب أن يحتوي كائن `users` فقط على المفاتيح `Alan` و`Jeff` و`Sarah` و`Ryan`

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

يجب أن تُرجع الدالة `isEveryoneHere` القيمة `true` إذا كانت `Alan` و`Jeff` و`Sarah` و`Ryan` خصائص في الكائن الممرر إليها.

```js
assert(isEveryoneHere(users) === true);
```

يجب أن تُرجع الدالة `isEveryoneHere` القيمة `false` إذا لم تكن `Alan` خاصية في الكائن الممرر إليها.

```js
assert(
  (function () {
    delete users.Alan;
    return isEveryoneHere(users);
  })() === false
);
```

يجب أن تُرجع الدالة `isEveryoneHere` القيمة `false` إذا لم تكن `Jeff` خاصية في الكائن الممرر إليها.

```js
assert(
  (function () {
    delete users.Jeff;
    return isEveryoneHere(users);
  })() === false
);
```

يجب أن تُرجع الدالة `isEveryoneHere` القيمة `false` إذا لم تكن `Sarah` خاصية في الكائن الممرر إليها.

```js
assert(
  (function () {
    delete users.Sarah;
    return isEveryoneHere(users);
  })() === false
);
```

يجب أن تُرجع الدالة `isEveryoneHere` القيمة `false` إذا لم تكن `Ryan` خاصية في الكائن الممرر إليها.

```js
assert(
  (function () {
    delete users.Ryan;
    return isEveryoneHere(users);
  })() === false
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  // Only change code below this line
  
  // Only change code above this line
}

console.log(isEveryoneHere(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  return [
    'Alan',
    'Jeff',
    'Sarah',
    'Ryan'
  ].every(user => userObj.hasOwnProperty(user));
}

console.log(isEveryoneHere(users));
```
