---
id: 587d7b7d367417b2b2512b1d
title: التكرار خلال مفاتيح كائن باستخدام جملة for...in
challengeType: 1
forumTopicId: 301162
dashedName: iterate-through-the-keys-of-an-object-with-a-for---in-statement
---

# --description--

أحيانًا تحتاج إلى التكرار خلال جميع المفاتيح داخل كائن. يمكنك استخدام حلقة <dfn>for...in</dfn> للقيام بذلك. تبدو حلقة for...in كما يلي:

```javascript
const refrigerator = {
  'milk': 1,
  'eggs': 12,
};

for (const food in refrigerator) {
  console.log(food, refrigerator[food]);
}
```

تُسجل هذه الشفرة `milk 1` و `eggs 12`، مع ظهور كل زوج مفتاح-قيمة في سطر منفصل.

عرّفنا المتغير `food` في رأس الحلقة، وتم تعيين هذا المتغير لكل مفتاح من مفاتيح الكائن في كل تكرار، مما أدى إلى طباعة اسم كل طعام في وحدة التحكم.

**ملاحظة:** الكائنات لا تحافظ على ترتيب المفاتيح المخزنة كما تفعل المصفوفات؛ لذلك فإن موقع المفتاح في الكائن، أو الترتيب النسبي لظهوره، لا يؤثر عند الرجوع إلى ذلك المفتاح أو الوصول إليه.

# --instructions--

عرّفنا دالة `countOnline` التي تقبل معلمة واحدة، `allUsers`. استخدم جملة <dfn>for...in</dfn> داخل هذه الدالة للتكرار خلال كائن `allUsers` وأرجع عدد المستخدمين الذين خاصية `online` لديهم مضبوطة على `true`. مثال على كائن يمكن تمريره إلى `countOnline` موضح أدناه. سيكون لكل مستخدم خاصية `online` مضبوطة إما على `true` أو `false`.

```js
{
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
```

# --hints--

يجب أن تستخدم الدالة `countOnline` جملة `for in` للتكرار خلال مفاتيح الكائن الذي يُمرر إليها.

```js
assert.match(
  __helpers.removeJSComments(code),
    /for\s*\(\s*(var|let|const)\s+[a-zA-Z_$]\w*\s+in\s+[a-zA-Z_$]\w*\s*\)/
);
```

يجب أن تُرجع الدالة `countOnline` القيمة `1` عندما يُمرر إليها الكائن `{ Alan: { online: false }, Jeff: { online: true }, Sarah: { online: false } }`

```js
const usersObj1 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj1), 1);
```

يجب أن تُرجع الدالة `countOnline` القيمة `2` عندما يُمرر إليها الكائن `{ Alan: { online: true }, Jeff: { online: false }, Sarah: { online: true } }`

```js
const usersObj2 = {
  Alan: {
    online: true
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: true
  }
}
assert.equal(countOnline(usersObj2), 2);
```

يجب أن تُرجع الدالة `countOnline` القيمة `0` عندما يُمرر إليها الكائن `{ Alan: { online: false }, Jeff: { online: false }, Sarah: { online: false } }`

```js
const usersObj3 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj3), 0);
```

# --seed--

## --seed-contents--

```js
const users = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}

function countOnline(allUsers) {
  // Only change code below this line

  // Only change code above this line
}

console.log(countOnline(users));
```

# --solutions--

```js
function countOnline(allUsers) {
  let numOnline = 0;
  for(const user in allUsers){
    if(allUsers[user].online) {
      numOnline++;
    }
  }
  return numOnline;
}
```
