---
id: 587d78b2367417b2b2512b0e
title: إضافة عناصر إلى مصفوفة باستخدام push() و unshift()
challengeType: 1
forumTopicId: 301151
dashedName: add-items-to-an-array-with-push-and-unshift
---

# --description--

طول المصفوفة، مثل أنواع البيانات التي يمكن أن تحتويها، ليس ثابتًا. يمكن تعريف المصفوفات بطول أي عدد من العناصر، ويمكن إضافة أو إزالة العناصر مع مرور الوقت؛ بعبارة أخرى، المصفوفات <dfn>قابلة للتغيير</dfn>. في هذا التحدٍّ، سننظر في طريقتين يمكننا من خلالهما تعديل المصفوفة برمجيًا: `Array.push()` و `Array.unshift()`.

تأخذ كلتا الطريقتين عنصرًا أو أكثر كمعلمات وتضيف تلك العناصر إلى المصفوفة التي تُستدعى عليها الطريقة؛ تضيف طريقة `push()` العناصر إلى نهاية المصفوفة، بينما تضيف `unshift()` العناصر إلى البداية. انظر المثال التالي:

```js
let twentyThree = 'XXIII';
let romanNumerals = ['XXI', 'XXII'];

romanNumerals.unshift('XIX', 'XX');
```

سيكون لـ `romanNumerals` القيمة `['XIX', 'XX', 'XXI', 'XXII']`.

```js
romanNumerals.push(twentyThree);
```

سيكون لـ `romanNumerals` القيمة `['XIX', 'XX', 'XXI', 'XXII', 'XXIII']`. لاحظ أننا يمكننا أيضًا تمرير متغيرات، مما يمنحنا مرونة أكبر في تعديل بيانات المصفوفة ديناميكيًا.

# --instructions--

لقد عرّفنا دالة `mixedNumbers` التي نمرر إليها مصفوفة كمعلمة. عدّل الدالة باستخدام `push()` و `unshift()` لإضافة `'I', 2, 'three'` إلى بداية المصفوفة و `7, 'VIII', 9` إلى نهايتها بحيث تحتوي المصفوفة المعادة على تمثيلات الأرقام من 1 إلى 9 بالترتيب.

# --hints--

يجب أن تُعيد `mixedNumbers(["IV", 5, "six"])` الآن `["I", 2, "three", "IV", 5, "six", 7, "VIII", 9]`

```js
assert.deepEqual(mixedNumbers(['IV', 5, 'six']), [
  'I',
  2,
  'three',
  'IV',
  5,
  'six',
  7,
  'VIII',
  9
]);
```

يجب أن تستخدم دالة `mixedNumbers` طريقة `push()`

```js
assert(mixedNumbers.toString().match(/\.push/));
```

يجب أن تستخدم دالة `mixedNumbers` طريقة `unshift()`

```js
assert(mixedNumbers.toString().match(/\.unshift/));
```

# --seed--

## --seed-contents--

```js
function mixedNumbers(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(mixedNumbers(['IV', 5, 'six']));
```

# --solutions--

```js
function mixedNumbers(arr) {
  arr.push(7,'VIII',9);
  arr.unshift('I',2,'three');
  return arr;
}
```
