---
id: 587d7b7b367417b2b2512b16
title: إنشاء مصفوفات متعددة الأبعاد معقدة
challengeType: 1
forumTopicId: 301159
dashedName: create-complex-multi-dimensional-arrays
---

# --description--

رائع! لقد تعلمت للتو الكثير عن المصفوفات! كان هذا نظرة عامة على مستوى عالٍ إلى حد ما، وهناك الكثير لتتعلمه عن العمل مع المصفوفات، سترى الكثير منه في الأقسام التالية. لكن قبل الانتقال إلى النظر في <dfn>الكائنات</dfn>، دعنا نلقي نظرة أخرى، ونرى كيف يمكن أن تصبح المصفوفات أكثر تعقيدًا مما رأينا في التحديات السابقة.

واحدة من أقوى الميزات عند التفكير في المصفوفات كـ هياكل بيانات، هي أن المصفوفات يمكن أن تحتوي، أو حتى تتكون بالكامل من مصفوفات أخرى. لقد رأينا مصفوفات تحتوي على مصفوفات في التحديات السابقة، لكنها كانت بسيطة إلى حد ما. ومع ذلك، يمكن للمصفوفات أن تحتوي على عمق لا نهائي من المصفوفات التي تحتوي على مصفوفات أخرى، كل منها بمستويات عمق عشوائية خاصة بها، وهكذا. بهذه الطريقة، يمكن أن تصبح المصفوفة بسرعة هيكل بيانات معقد جدًا، يُعرف باسم مصفوفة <dfn>متعددة الأبعاد</dfn>، أو مصفوفة مسكنة. انظر المثال التالي:

```js
let nestedArray = [
  ['deep'],
  [
    ['deeper'], ['deeper'] 
  ],
  [
    [
      ['deepest'], ['deepest']
    ],
    [
      [
        ['deepest-est?']
      ]
    ]
  ]
];
```

المصفوفة `deep` مسكنة على عمق مستويين. المصفوفات `deeper` مسكنة على عمق 3 مستويات. المصفوفات `deepest` على 4 مستويات، و`deepest-est?` على 5 مستويات.

رغم أن هذا المثال قد يبدو معقدًا، إلا أن هذا المستوى من التعقيد ليس غريبًا أو غير معتاد عند التعامل مع كميات كبيرة من البيانات. ومع ذلك، يمكننا بسهولة كبيرة الوصول إلى أعمق مستويات المصفوفة المعقدة هذه باستخدام ترميز الأقواس:

```js
console.log(nestedArray[2][1][0][0][0]);
```

هذا يسجل السلسلة النصية `deepest-est?`. والآن بعد أن عرفنا مكان هذه البيانات، يمكننا إعادة تعيينها إذا احتجنا إلى ذلك:

```js
nestedArray[2][1][0][0][0] = 'deeper still';

console.log(nestedArray[2][1][0][0][0]);
```

الآن يسجل `deeper still`.

# --instructions--

لقد عرّفنا متغيرًا باسم `myNestedArray`، مساويًا لمصفوفة. عدّل `myNestedArray` باستخدام أي تركيبة من <dfn>السلاسل النصية</dfn>، و<dfn>الأرقام</dfn>، و<dfn>القيم المنطقية</dfn> كعناصر بيانات، بحيث يحتوي على خمسة مستويات من العمق بالضبط (تذكر، المصفوفة الخارجية هي المستوى 1). في مكان ما في المستوى الثالث، أدرج السلسلة النصية `deep`، وفي المستوى الرابع، أدرج السلسلة النصية `deeper`، وفي المستوى الخامس، أدرج السلسلة النصية `deepest`.

# --hints--

يجب أن يحتوي `myNestedArray` على أرقام وقيم منطقية وسلاسل نصية فقط كعناصر بيانات

```js
assert.strictEqual(
  (function (arr) {
    let flattened = (function flatten(arr) {
      const flat = [].concat(...arr);
      return flat.some(Array.isArray) ? flatten(flat) : flat;
    })(arr);
    for (let i = 0; i < flattened.length; i++) {
      if (
        typeof flattened[i] !== 'number' &&
        typeof flattened[i] !== 'string' &&
        typeof flattened[i] !== 'boolean'
      ) {
        return false;
      }
    }
    return true;
  })(myNestedArray),
  true
);
```

يجب أن يحتوي `myNestedArray` على خمسة مستويات من العمق بالضبط

```js
assert.strictEqual(
  (function (arr) {
    let depth = 0;
    function arrayDepth(array, i, d) {
      if (Array.isArray(array[i])) {
        arrayDepth(array[i], 0, d + 1);
      } else {
        depth = d > depth ? d : depth;
      }
      if (i < array.length) {
        arrayDepth(array, i + 1, d);
      }
    }
    arrayDepth(arr, 0, 0);
    return depth;
  })(myNestedArray),
  4
);
```

يجب أن يحتوي `myNestedArray` على ظهور واحد فقط للسلسلة النصية `deep` في مصفوفة مسكنة على عمق 3 مستويات

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deep').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deep')[0] === 2
);
```

يجب أن يحتوي `myNestedArray` على ظهور واحد فقط للسلسلة النصية `deeper` في مصفوفة مسكنة على عمق 4 مستويات

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deeper').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deeper')[0] === 3
);
```

يجب أن يحتوي `myNestedArray` على ظهور واحد فقط للسلسلة النصية `deepest` في مصفوفة مسكنة على عمق 5 مستويات

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deepest').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deepest')[0] === 4
);
```

# --seed--

## --seed-contents--

```js
let myNestedArray = [
  // Only change code below this line
  ['unshift', false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
  // Only change code above this line
];
```

# --solutions--

```js
let myNestedArray = [
  ['unshift', ['deep', ['deeper', ['deepest']]],false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
];
```
