---
id: 587d7b7c367417b2b2512b18
title: إضافة أزواج المفتاح-القيمة إلى كائنات JavaScript
challengeType: 1
forumTopicId: 301153
dashedName: add-key-value-pairs-to-javascript-objects
---

# --description--

في أبسط صورها، الكائنات هي مجرد مجموعات من أزواج <dfn>المفتاح-القيمة</dfn>. بعبارة أخرى، هي قطع من البيانات (<dfn>القيم</dfn>) مرتبطة بمعرّفات فريدة تُسمى <dfn>خصائص</dfn> (<dfn>مفاتيح</dfn>). ألقِ نظرة على المثال التالي:

```js
const tekkenCharacter = {
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true
};
```

تعرف الشفرة أعلاه كائن شخصية من لعبة Tekken يُسمى `tekkenCharacter`. يحتوي على ثلاث خصائص، كل منها مرتبط بقيمة محددة. إذا أردت إضافة خاصية إضافية، مثل "origin"، يمكن فعل ذلك بتعيين `origin` إلى الكائن:

```js
tekkenCharacter.origin = 'South Korea';
```

هذا يستخدم ترميز النقطة. إذا راقبت كائن `tekkenCharacter`، ستجد أنه الآن يتضمن الخاصية `origin`. كان لدى Hwoarang أيضًا شعر برتقالي مميز. يمكنك إضافة هذه الخاصية باستخدام ترميز الأقواس كما يلي:

```js
tekkenCharacter['hair color'] = 'dyed orange';
```

يُطلب استخدام ترميز الأقواس إذا كانت الخاصية تحتوي على فراغ أو إذا أردت استخدام متغير لتسمية الخاصية. في الحالة أعلاه، تُحاط الخاصية بعلامات اقتباس للدلالة على أنها سلسلة نصية وستُضاف كما هي. بدون علامات الاقتباس، ستُقيّم كمتغير واسم الخاصية سيكون القيمة التي يحملها المتغير. إليك مثالًا مع متغير:

```js
const eyes = 'eye color';

tekkenCharacter[eyes] = 'brown';
```

بعد إضافة كل الأمثلة، سيبدو الكائن هكذا:

```js
{
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true,
  origin: 'South Korea',
  'hair color': 'dyed orange',
  'eye color': 'brown'
};
```

# --instructions--

تم إنشاء كائن `foods` بثلاث مدخلات. باستخدام الصيغة التي تختارها، أضف ثلاث مدخلات أخرى إليه: `bananas` بقيمة `13`، و`grapes` بقيمة `35`، و`strawberries` بقيمة `27`.

# --hints--

يجب أن يكون `foods` كائنًا.

```js
assert(typeof foods === 'object');
```

يجب أن يحتوي كائن `foods` على مفتاح `bananas` بقيمة `13`.

```js
assert(foods.bananas === 13);
```

يجب أن يحتوي كائن `foods` على مفتاح `grapes` بقيمة `35`.

```js
assert(foods.grapes === 35);
```

يجب أن يحتوي كائن `foods` على مفتاح `strawberries` بقيمة `27`.

```js
assert(foods.strawberries === 27);
```

يجب عدم تغيير تعريف كائن `foods`.

```js
assert(
  __helpers.removeJSComments(code).search(/let foods/) === -1 &&
  __helpers.removeJSComments(code).search(/const\s+foods\s*=\s*{\s*apples:\s*25,\s*oranges:\s*32,\s*plums:\s*28\s*};/
) !== -1
);
```

# --seed--

## --seed-contents--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

foods['bananas'] = 13;
foods['grapes']  = 35;
foods['strawberries'] = 27;
```
