---
id: 587d7b7b367417b2b2512b13
title: نسخ مصفوفة باستخدام معامل النشر
challengeType: 1
forumTopicId: 301157
dashedName: copy-an-array-with-the-spread-operator
---

# --description--

بينما تسمح لنا الدالة `slice()` بالاختيار الدقيق للعناصر التي نريد نسخها من المصفوفة، من بين عدة مهام مفيدة أخرى، يتيح لنا <dfn>معامل النشر</dfn> الجديد في ES6 نسخ *جميع* عناصر المصفوفة بسهولة، وبترتيبها، باستخدام صيغة نحوية بسيطة وسهلة القراءة. صيغة النشر تبدو هكذا ببساطة: `...`

عمليًا، يمكننا استخدام معامل النشر لنسخ مصفوفة كما يلي:

```js
let thisArray = [true, true, undefined, false, null];
let thatArray = [...thisArray];
```

تساوي `thatArray` القيمة `[true, true, undefined, false, null]`. تبقى `thisArray` دون تغيير وتحتوي `thatArray` على نفس عناصر `thisArray`.

# --instructions--

لقد عرّفنا دالة باسم `copyMachine` تأخذ `arr` (مصفوفة) و `num` (رقم) كمعلمات. من المفترض أن تُرجع الدالة مصفوفة جديدة مكونة من `num` نسخ من `arr`. لقد أنجزنا معظم العمل من أجلك، لكنها لا تعمل بشكل صحيح بعد. عدّل الدالة باستخدام صيغة النشر لتعمل بشكل صحيح (تلميح: قد تفيدك طريقة أخرى سبق أن تعلمناها!).

# --hints--

يجب أن تُرجع `copyMachine([true, false, true], 2)` القيمة `[[true, false, true], [true, false, true]]`

```js
assert.deepEqual(copyMachine([true, false, true], 2), [
  [true, false, true],
  [true, false, true]
]);
```

يجب أن تُرجع `copyMachine([1, 2, 3], 5)` القيمة `[[1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3]]`

```js
assert.deepEqual(copyMachine([1, 2, 3], 5), [
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3]
]);
```

يجب أن تُرجع `copyMachine([true, true, null], 1)` القيمة `[[true, true, null]]`

```js
assert.deepEqual(copyMachine([true, true, null], 1), [[true, true, null]]);
```

يجب أن تُعيد الدالة `copyMachine(["it works"], 3)` القيمة `[["it works"], ["it works"], ["it works"]]`

```js
assert.deepEqual(copyMachine(['it works'], 3), [
  ['it works'],
  ['it works'],
  ['it works']
]);
```

يجب أن تستخدم الدالة `copyMachine` معامل النشر `spread operator` مع المصفوفة `arr`

```js
assert(__helpers.removeJSComments(__helpers.removeJSComments(code)).match(/\.\.\.\s*arr/));
```

# --seed--

## --seed-contents--

```js
function copyMachine(arr, num) {
  let newArr = [];
  while (num >= 1) {
    // Only change code below this line

    // Only change code above this line
    num--;
  }
  return newArr;
}

console.log(copyMachine([true, false, true], 2));
```

# --solutions--

```js
function copyMachine(arr,num){
    let newArr=[];
    while(num >=1){
    newArr.push([...arr]);
    num--;
    }
    return newArr;
}
console.log(copyMachine([true, false, true], 2));
```
