---
id: 587d7b7b367417b2b2512b14
title: التحقق من وجود عنصر باستخدام indexOf()
challengeType: 1
forumTopicId: 301154
dashedName: check-for-the-presence-of-an-element-with-indexof
---

# --description--

بما أن المصفوفات يمكن تغييرها أو *تعديلها* في أي وقت، فلا يوجد ضمان لمكان وجود قطعة معينة من البيانات في مصفوفة معينة، أو حتى إذا كان ذلك العنصر لا يزال موجودًا. لحسن الحظ، توفر JavaScript طريقة مدمجة أخرى، `indexOf()`، تتيح لك التحقق بسرعة وسهولة من وجود عنصر في مصفوفة. تأخذ `indexOf()` عنصرًا كمعلمة، وعند استدعائها، تُرجع موقع ذلك العنصر، أو الفهرس، أو `-1` إذا لم يكن العنصر موجودًا في المصفوفة.

على سبيل المثال:

```js
let fruits = ['apples', 'pears', 'oranges', 'peaches', 'pears'];

fruits.indexOf('dates');
fruits.indexOf('oranges');
fruits.indexOf('pears');
```

`indexOf('dates')` تُرجع `-1`، و`indexOf('oranges')` تُرجع `2`، و`indexOf('pears')` تُرجع `1` (أول فهرس يوجد فيه كل عنصر).

# --instructions--

يمكن أن تكون `indexOf()` مفيدة جدًا للتحقق بسرعة من وجود عنصر في مصفوفة. لقد عرّفنا دالة، `quickCheck`، تأخذ مصفوفة وعنصر كمعلمات. عدّل الدالة باستخدام `indexOf()` بحيث تُرجع `true` إذا كان العنصر الممرر موجودًا في المصفوفة، و`false` إذا لم يكن موجودًا.

# --hints--

يجب أن تُرجع دالة `quickCheck` قيمة منطقية (`true` أو `false`)، وليس سلسلة نصية (`"true"` أو `"false"`).

```js
assert.isBoolean(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

`quickCheck(["squash", "onions", "shallots"], "mushrooms")` يجب أن تُرجع `false`

```js
assert.strictEqual(
  quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'),
  false
);
```

`quickCheck(["onions", "squash", "shallots"], "onions")` يجب أن تُرجع `true`

```js
assert.strictEqual(
  quickCheck(['onions', 'squash', 'shallots'], 'onions'),
  true
);
```

`quickCheck([3, 5, 9, 125, 45, 2], 125)` يجب أن تُرجع `true`

```js
assert.strictEqual(quickCheck([3, 5, 9, 125, 45, 2], 125), true);
```

`quickCheck([true, false, false], undefined)` يجب أن تُرجع `false`

```js
assert.strictEqual(quickCheck([true, false, false], undefined), false);
```

يجب أن تستخدم دالة `quickCheck` طريقة `indexOf()`

```js
assert.notStrictEqual(quickCheck.toString().search(/\.indexOf\(/), -1);
```

# --seed--

## --seed-contents--

```js
function quickCheck(arr, elem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

# --solutions--

```js
function quickCheck(arr, elem) {
  return arr.indexOf(elem) >= 0; 
}
```
