---
id: 587d7b7e367417b2b2512b20
title: استخدام مصفوفة لتخزين مجموعة من البيانات
challengeType: 1
forumTopicId: 301167
dashedName: use-an-array-to-store-a-collection-of-data
---

# --description--

فيما يلي مثال على أبسط تنفيذ لهيكل بيانات المصفوفة. يُعرف هذا باسم <dfn>مصفوفة أحادية البعد</dfn>، أي أنها تحتوي على مستوى واحد فقط، أو أنها لا تحتوي على أي مصفوفات أخرى مسكنة بداخلها. لاحظ أنها تحتوي على <dfn>قيم منطقية</dfn>، و<dfn>سلاسل نصية</dfn>، و<dfn>أرقام</dfn>، إلى جانب أنواع بيانات جافاسكريبت صالحة أخرى:

```js
let simpleArray = ['one', 2, 'three', true, false, undefined, null];
console.log(simpleArray.length);
```

يعرض استدعاء `console.log` القيمة `7`.

كل المصفوفات لها خاصية الطول، والتي كما هو موضح أعلاه، يمكن الوصول إليها بسهولة باستخدام الصيغة النحوية `Array.length`. يمكن رؤية تنفيذ أكثر تعقيدًا للمصفوفة أدناه. يُعرف هذا باسم <dfn>مصفوفة متعددة الأبعاد</dfn>، أو مصفوفة تحتوي على مصفوفات أخرى. لاحظ أن هذه المصفوفة تحتوي أيضًا على <dfn>كائنات</dfn> جافاسكريبت، والتي سنفحصها بدقة في القسم التالي، ولكن كل ما تحتاج لمعرفته الآن هو أن المصفوفات قادرة أيضًا على تخزين كائنات معقدة.

```js
let complexArray = [
  [
    {
      one: 1,
      two: 2
    },
    {
      three: 3,
      four: 4
    }
  ],
  [
    {
      a: "a",
      b: "b"
    },
    {
      c: "c",
      d: "d"
    }
  ]
];
```

# --instructions--

لقد عرّفنا متغيرًا يسمى `yourArray`. أكمل العبارة بتعيين مصفوفة لا يقل طولها عن 5 عناصر إلى المتغير `yourArray`. يجب أن تحتوي مصفوفاتك على الأقل على <dfn>سلسلة نصية</dfn> واحدة، و<dfn>رقم</dfn> واحد، و<dfn>قيمة منطقية</dfn> واحدة.

# --hints--

يجب أن يكون `yourArray` مصفوفة.

```js
assert.strictEqual(Array.isArray(yourArray), true);
```

يجب أن يكون طول `yourArray` على الأقل 5 عناصر.

```js
assert.isAtLeast(yourArray.length, 5);
```

يجب أن تحتوي `yourArray` على عنصر واحد على الأقل من نوع `boolean`.

```js
assert(yourArray.filter((el) => typeof el === 'boolean').length >= 1);
```

يجب أن تحتوي `yourArray` على عنصر واحد على الأقل من نوع `number`.

```js
assert(yourArray.filter((el) => typeof el === 'number').length >= 1);
```

يجب أن تحتوي `yourArray` على عنصر واحد على الأقل من نوع `string`.

```js
assert(yourArray.filter((el) => typeof el === 'string').length >= 1);
```

# --seed--

## --seed-contents--

```js
let yourArray; // Change this line
```

# --solutions--

```js
let yourArray = ['a string', 100, true, ['one', 2], 'another string'];
```
