---
id: 587d7b7c367417b2b2512b1b
title: استخدام الكلمة المفتاحية delete لحذف خصائص الكائن
challengeType: 1
forumTopicId: 301168
dashedName: use-the-delete-keyword-to-remove-object-properties
---

# --description--

الآن تعرف ما هي الكائنات وميزاتها الأساسية وفوائدها. باختصار، هي مخازن لأزواج المفتاح-القيمة توفر طريقة مرنة وبديهية لتنظيم هيكل البيانات، ***وأيضًا***، توفر وقت بحث سريع جدًا. خلال بقية هذه التحديات، سنشرح عدة عمليات شائعة يمكنك تنفيذها على الكائنات لتصبح مرتاحًا في تطبيق هذه الهياكل المفيدة في برامجك.

في التحديات السابقة، أضفنا وعدلنا أزواج المفتاح-القيمة في الكائن. هنا سنرى كيف يمكننا *حذف* زوج مفتاح-قيمة من كائن.

لنعد إلى مثال كائن `foods` مرة أخيرة. إذا أردنا حذف المفتاح `apples`، يمكننا حذفه باستخدام الكلمة المفتاحية `delete` هكذا:

```js
delete foods.apples;
```

# --instructions--

استخدم الكلمة المفتاحية delete لحذف المفاتيح `oranges`، `plums`، و`strawberries` من كائن `foods`.

# --hints--

يجب أن يحتوي كائن `foods` على ثلاثة مفاتيح فقط: `apples`، `grapes`، و`bananas`.

```js
assert(
  !foods.hasOwnProperty('oranges') &&
    !foods.hasOwnProperty('plums') &&
    !foods.hasOwnProperty('strawberries') &&
    Object.keys(foods).length === 3
);
```

يجب حذف المفاتيح `oranges`، `plums`، و`strawberries` باستخدام `delete`.

```js
assert(
  __helpers.removeJSComments(code).search(/oranges:/) !== -1 &&
    __helpers.removeJSComments(code).search(/plums:/) !== -1 &&
    __helpers.removeJSComments(code).search(/strawberries:/) !== -1
);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

delete foods.oranges;
delete foods.plums;
delete foods.strawberries;

console.log(foods);
```
