---
id: 587d7b7c367417b2b2512b1a
title: الوصول إلى أسماء الخصائص باستخدام ترميز الأقواس
challengeType: 1
forumTopicId: 301150
dashedName: access-property-names-with-bracket-notation
---

# --description--

في أول تحدٍّ للكائنات ذكرنا استخدام ترميز الأقواس كطريقة للوصول إلى قيم الخصائص باستخدام تقييم متغير. على سبيل المثال، تخيل أن كائن `foods` يُستخدم في برنامج لصندوق تسجيل في سوبرماركت. لدينا دالة تُعيّن `selectedFood` ونريد التحقق من وجود هذا الطعام في كائن `foods`. قد يبدو ذلك كما يلي:

```js
let selectedFood = getCurrentFood(scannedItem);
let inventory = foods[selectedFood];
```

تُقيّم هذه الشفرة القيمة المخزنة في متغير `selectedFood` وتُرجع قيمة ذلك المفتاح في كائن `foods`، أو `undefined` إذا لم يكن موجودًا. ترميز الأقواس مفيد جدًا لأن خصائص الكائنات أحيانًا لا تُعرف قبل وقت التشغيل أو نحتاج إلى الوصول إليها بطريقة أكثر ديناميكية.

# --instructions--

عرّفنا دالة `checkInventory` التي تستقبل عنصرًا ممسوحًا كمعلمة. أرجع القيمة الحالية للمفتاح `scannedItem` في كائن `foods`. يمكنك افتراض أن المعلمات المقدمة إلى `checkInventory` ستكون فقط مفاتيح صحيحة.

# --hints--

يجب أن تكون `checkInventory` دالة.

```js
assert.strictEqual(typeof checkInventory, 'function');
```

يجب أن يحتوي كائن `foods` فقط على أزواج المفتاح-القيمة التالية: `apples: 25`، `oranges: 32`، `plums: 28`، `bananas: 13`، `grapes: 35`، `strawberries: 27`.

```js
assert.deepEqual(foods, {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
});
```

يجب أن تُرجع `checkInventory("apples")` القيمة `25`.

```js
assert.strictEqual(checkInventory('apples'), 25);
```

يجب أن تُرجع `checkInventory("bananas")` القيمة `13`.

```js
assert.strictEqual(checkInventory('bananas'), 13);
```

يجب أن تُرجع `checkInventory("strawberries")` القيمة `27`.

```js
assert.strictEqual(checkInventory('strawberries'), 27);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(checkInventory("apples"));
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  return foods[scannedItem];
}
```
