---
id: 587d78b2367417b2b2512b10
title: إزالة عناصر باستخدام splice()
challengeType: 1
forumTopicId: 301166
dashedName: remove-items-using-splice
---

# --description--

حسنًا، تعلمنا كيف نحذف عناصر من بداية ونهاية المصفوفات باستخدام `shift()` و `pop()`، لكن ماذا لو أردنا حذف عنصر من وسط المصفوفة؟ أو حذف أكثر من عنصر مرة واحدة؟ هنا يأتي دور `splice()`. تتيح لنا `splice()` فعل ذلك بالضبط: **حذف أي عدد من العناصر المتتالية** من أي مكان في المصفوفة.

يمكن لـ `splice()` أن تأخذ حتى 3 معلمات، لكننا سنركز الآن على أول معلمتين فقط. المعلمتان الأوليان لـ `splice()` هما أعداد صحيحة تمثل فهارس، أو مواقع، العناصر في المصفوفة التي يُستدعى عليها `splice()`. وتذكر أن المصفوفات تبدأ فهارسها من الصفر، لذا للإشارة إلى العنصر الأول في المصفوفة نستخدم `0`. المعلمة الأولى لـ `splice()` تمثل الفهرس في المصفوفة الذي نبدأ منه حذف العناصر، بينما المعلمة الثانية تشير إلى عدد العناصر التي يجب حذفها. على سبيل المثال:

```js
let array = ['today', 'was', 'not', 'so', 'great'];

array.splice(2, 2);
```

هنا نحذف عنصرين، بدءًا من العنصر الثالث (عند الفهرس 2). ستكون قيمة `array` هي `['today', 'was', 'great']`.

لا تقوم `splice()` فقط بتعديل المصفوفة التي تُستدعى عليها، بل تُرجع أيضًا مصفوفة جديدة تحتوي على العناصر المحذوفة:

```js
let array = ['I', 'am', 'feeling', 'really', 'happy'];

let newArray = array.splice(3, 2);
```

تحتوي `newArray` على القيمة `['really', 'happy']`.

# --instructions--

لقد أنشأنا مصفوفة `arr`. استخدم `splice()` لحذف عناصر من `arr` بحيث تحتوي فقط على عناصر مجموعها يساوي `10`.

# --hints--

لا يجب عليك تغيير السطر الأصلي `const arr = [2, 4, 5, 1, 7, 5, 2, 1];`.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/constarr=\[2,4,5,1,7,5,2,1\];?/)
);
```

يجب أن تحتوي `arr` فقط على عناصر مجموعها يساوي `10`.

```js
assert.strictEqual(
  arr.reduce((a, b) => a + b),
  10
);
```

يجب أن يستخدم شفرتك طريقة `splice()` على `arr`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(/));
```

يجب أن تُزيل الدالة splice() عناصر فقط من `arr` دون إضافة أي عناصر إضافية إلى `arr`.

```js
assert(
  !__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(\d+,\d+,\d+.*\)/g)
);
```

# --seed--

## --seed-contents--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
// Only change code below this line

// Only change code above this line
console.log(arr);
```

# --solutions--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
arr.splice(1, 4);
```
