---
id: 587d78b2367417b2b2512b0f
title: إزالة عناصر من مصفوفة باستخدام pop() و shift()
challengeType: 1
forumTopicId: 301165
dashedName: remove-items-from-an-array-with-pop-and-shift
---

# --description--

كل من `push()` و `unshift()` لهما طرق مقابلة تكاد تكون عكسية وظيفيًا: `pop()` و `shift()`. كما قد تكون خمّنت حتى الآن، بدلًا من الإضافة، تقوم `pop()` *بإزالة* عنصر من نهاية المصفوفة، بينما تزيل `shift()` عنصرًا من البداية. الفرق الرئيسي بين `pop()` و `shift()` وأقربائهما `push()` و `unshift()` هو أن أيًا من الطريقتين لا تأخذ معلمات، وكل واحدة تسمح بتعديل المصفوفة بعنصر واحد فقط في كل مرة.

لنلق نظرة:

```js
let greetings = ['whats up?', 'hello', 'see ya!'];

greetings.pop();
```

`greetings` ستكون قيمتها `['whats up?', 'hello']`.

```js
greetings.shift();
```

`greetings` ستكون قيمتها `['hello']`.

يمكننا أيضًا إرجاع قيمة العنصر المُزال باستخدام أي من الطريقتين هكذا:

```js
let popped = greetings.pop();
```

`greetings` ستكون قيمتها `[]`، و`popped` ستكون قيمتها `hello`.

# --instructions--

لقد عرّفنا دالة `popShift` التي تأخذ مصفوفة كمعلمة وتُرجع مصفوفة جديدة. عدّل الدالة باستخدام `pop()` و `shift()` لإزالة العنصر الأول والأخير من مصفوفة المعطى، وعيّن العناصر المُزالة إلى المتغيرات المقابلة لها، بحيث تحتوي المصفوفة المرجعة على قيمها.

# --hints--

`popShift(["challenge", "is", "not", "complete"])` يجب أن تُرجع `["challenge", "complete"]`

```js
assert.deepEqual(popShift(['challenge', 'is', 'not', 'complete']), [
  'challenge',
  'complete'
]);
```

يجب أن تستخدم دالة `popShift` الطريقة `pop()`

```js
assert.notStrictEqual(popShift.toString().search(/\.pop\(/), -1);
```

يجب أن تستخدم دالة `popShift` الطريقة `shift()`

```js
assert.notStrictEqual(popShift.toString().search(/\.shift\(/), -1);
```

# --seed--

## --seed-contents--

```js
function popShift(arr) {
  let popped; // Change this line
  let shifted; // Change this line
  return [shifted, popped];
}

console.log(popShift(['challenge', 'is', 'not', 'complete']));
```

# --solutions--

```js
function popShift(arr) {
  let popped = arr.pop(); // Change this line
  let shifted = arr.shift(); // Change this line
  return [shifted, popped];
}
```
