---
id: 587d7b7c367417b2b2512b19
title: تعديل كائن مسكن داخل كائن
challengeType: 1
forumTopicId: 301164
dashedName: modify-an-object-nested-within-an-object
---

# --description--

لنلقِ نظرة الآن على كائن أكثر تعقيدًا قليلاً. يمكن تسكين خصائص الكائن إلى عمق غير محدد، وقيمها يمكن أن تكون أي نوع من البيانات التي يدعمها JavaScript، بما في ذلك المصفوفات وحتى الكائنات الأخرى. انظر المثال التالي:

```js
let nestedObject = {
  id: 28802695164,
  date: 'December 31, 2016',
  data: {
    totalUsers: 99,
    online: 80,
    onlineStatus: {
      active: 67,
      away: 13,
      busy: 8
    }
  }
};
```

لدى `nestedObject` ثلاث خصائص: `id` (القيمة رقم)، و`date` (القيمة سلسلة نصية)، و`data` (القيمة كائن يحتوي على هيكل مسكن). رغم أن الهياكل قد تصبح معقدة بسرعة، يمكننا استخدام نفس الترميزات للوصول إلى المعلومات التي نحتاجها. لتعيين القيمة `10` إلى خاصية `busy` في الكائن المسكن `onlineStatus`، نستخدم ترميز النقطة للإشارة إلى الخاصية:

```js
nestedObject.data.onlineStatus.busy = 10;
```

# --instructions--

هنا عرّفنا كائنًا باسم `userActivity`، يحتوي على كائن آخر مسكن بداخله. عيّن قيمة المفتاح `online` إلى `45`.

# --hints--

يجب أن يحتوي `userActivity` على خصائص `id` و`date` و`data`.

```js
assert(
  'id' in userActivity && 'date' in userActivity && 'data' in userActivity
);
```

يجب أن يحتوي `userActivity` على مفتاح `data` معين إلى كائن يحتوي على المفاتيح `totalUsers` و`online`.

```js
assert('totalUsers' in userActivity.data && 'online' in userActivity.data);
```

يجب تعيين الخاصية `online` المسكنة في مفتاح `data` من `userActivity` إلى `45`.

```js
assert(userActivity.data.online === 45);
```

يجب تعيين الخاصية `online` باستخدام ترميز النقطة أو ترميز الأقواس.

```js
assert.strictEqual(__helpers.removeJSComments(code).search(/online: 45/), -1);
```

# --seed--

## --seed-contents--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

// Only change code below this line

// Only change code above this line

console.log(userActivity);
```

# --solutions--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

userActivity.data.online = 45;
```
