---
id: 587d7b7b367417b2b2512b15
title: التكرار خلال جميع عناصر المصفوفة باستخدام حلقات for
challengeType: 1
forumTopicId: 301161
dashedName: iterate-through-all-an-arrays-items-using-for-loops
---

# --description--

أحيانًا عند العمل مع المصفوفات، يكون من المفيد جدًا التكرار خلال كل عنصر للعثور على عنصر أو أكثر قد نحتاجها، أو للتلاعب بالمصفوفة بناءً على العناصر التي تحقق مجموعة معينة من المعايير. تقدم JavaScript عدة طرق مدمجة لكل منها طريقة مختلفة قليلًا في التكرار على المصفوفات لتحقيق نتائج مختلفة (مثل `every()`, `forEach()`, `map()`، وغيرها)، لكن التقنية الأكثر مرونة والتي تمنحنا أكبر قدر من التحكم هي حلقة `for` بسيطة.

انظر المثال التالي:

```js
function greaterThanTen(arr) {
  let newArr = [];
  for (let i = 0; i < arr.length; i++) {
    if (arr[i] > 10) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}

greaterThanTen([2, 12, 8, 14, 80, 0, 1]);
```

باستخدام حلقة `for`، تقوم هذه الدالة بالتكرار والوصول إلى كل عنصر في المصفوفة، وتخضعه لاختبار بسيط أنشأناه. بهذه الطريقة، حددنا بسهولة وبرمجيًا أي العناصر أكبر من `10`، وأرجعنا مصفوفة جديدة، `[12, 14, 80]`، تحتوي على تلك العناصر.

# --instructions--

لقد عرفنا دالة `filteredArray` التي تأخذ `arr`، وهي مصفوفة مسكنة، و`elem` كمعلمات، وتُرجع مصفوفة جديدة. تمثل `elem` عنصرًا قد يكون موجودًا أو غير موجود في واحدة أو أكثر من المصفوفات المسكنة داخل `arr`. عدل الدالة باستخدام حلقة `for` لإرجاع نسخة مفلترة من المصفوفة الممررة بحيث تُزال أي مصفوفة مسكنة داخل `arr` تحتوي على `elem`.

# --hints--

`filteredArray([[10, 8, 3], [14, 6, 23], [3, 18, 6]], 18)` يجب أن تُرجع `[[10, 8, 3], [14, 6, 23]]`

```js
assert.deepEqual(
  filteredArray(
    [
      [10, 8, 3],
      [14, 6, 23],
      [3, 18, 6]
    ],
    18
  ),
  [
    [10, 8, 3],
    [14, 6, 23]
  ]
);
```

`filteredArray([["trumpets", 2], ["flutes", 4], ["saxophones", 2]], 2)` يجب أن تُرجع `[["flutes", 4]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['trumpets', 2],
      ['flutes', 4],
      ['saxophones', 2]
    ],
    2
  ),
  [['flutes', 4]]
);
```

`filteredArray([["amy", "beth", "sam"], ["dave", "sean", "peter"]], "peter")` يجب أن تُرجع `[["amy", "beth", "sam"]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['amy', 'beth', 'sam'],
      ['dave', 'sean', 'peter']
    ],
    'peter'
  ),
  [['amy', 'beth', 'sam']]
);
```

`filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3)` يجب أن تُرجع `[]`

```js
assert.deepEqual(
  filteredArray(
    [
      [3, 2, 3],
      [1, 6, 3],
      [3, 13, 26],
      [19, 3, 9]
    ],
    3
  ),
  []
);
```

يجب أن تستخدم دالة `filteredArray` حلقة `for`

```js
assert.notStrictEqual(filteredArray.toString().search(/for/), -1);
```

# --seed--

## --seed-contents--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  // Only change code below this line

  // Only change code above this line
  return newArr;
}

console.log(filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3));
```

# --solutions--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  for (let i = 0; i<arr.length; i++) {
    if (arr[i].indexOf(elem) < 0) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}
```
