---
id: 587d78b3367417b2b2512b11
title: إضافة عناصر باستخدام splice()
challengeType: 1
forumTopicId: 301152
dashedName: add-items-using-splice
---

# --description--

تذكر في التحدٍّ السابق ذكرنا أن `splice()` يمكن أن تأخذ حتى ثلاث معلمات؟ حسنًا، يمكنك استخدام المعلمة الثالثة، التي تتكون من عنصر واحد أو أكثر، لإضافتها إلى المصفوفة. هذا مفيد جدًا لتبديل عنصر أو مجموعة عناصر بسرعة بعناصر أخرى.

```js
const numbers = [10, 11, 12, 12, 15];
const startIndex = 3;
const amountToDelete = 1;

numbers.splice(startIndex, amountToDelete, 13, 14);
console.log(numbers);
```

تم حذف الظهور الثاني لـ `12`، وأضفنا `13` و`14` في نفس الفهرس. ستكون مصفوفة `numbers` الآن `[ 10, 11, 12, 13, 14, 15 ]`.

هنا، نبدأ بمصفوفة من الأرقام. ثم نمرر إلى `splice()` ما يلي: الفهرس الذي نبدأ عنده حذف العناصر (3)، عدد العناصر التي سيتم حذفها (1)، والمعلمات المتبقية (13، 14) التي ستُدرج بدءًا من نفس الفهرس. لاحظ أنه يمكن أن يكون هناك أي عدد من العناصر (مفصولة بفواصل) بعد `amountToDelete`، وكل منها يُدرج.

# --instructions--

لقد عرّفنا دالة `htmlColorNames` التي تأخذ مصفوفة من ألوان HTML كمعلمة. عدّل الدالة باستخدام `splice()` لحذف أول عنصرين من المصفوفة وإضافة `'DarkSalmon'` و`'BlanchedAlmond'` في مواضعهما المناسبة.

# --hints--

يجب أن تُرجع `htmlColorNames` `["DarkSalmon", "BlanchedAlmond", "LavenderBlush", "PaleTurquoise", "FireBrick"]`

```js
assert.deepEqual(
  htmlColorNames([
    'DarkGoldenRod',
    'WhiteSmoke',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]),
  [
    'DarkSalmon',
    'BlanchedAlmond',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]
);
```

يجب أن تستخدم دالة `htmlColorNames` طريقة `splice()`

```js
assert(/.splice/.test(__helpers.removeJSComments(code)));
```

لا يجب أن تستخدم `shift()` أو `unshift()`.

```js
assert(!/shift|unshift/.test(__helpers.removeJSComments(code)));
```

لا يجب أن تستخدم ترميز الأقواس للمصفوفة.

```js
assert(!/\[\d\]\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function htmlColorNames(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(htmlColorNames(['DarkGoldenRod', 'WhiteSmoke', 'LavenderBlush', 'PaleTurquoise', 'FireBrick']));
```

# --solutions--

```js
function htmlColorNames(arr) {
  arr.splice(0,2,'DarkSalmon', 'BlanchedAlmond');
  return arr;
}
```
