---
id: 587d7b7a367417b2b2512b12
title: نسخ عناصر المصفوفة باستخدام slice()
challengeType: 1
forumTopicId: 301158
dashedName: copy-array-items-using-slice
---

# --description--

الطريقة التالية التي سنغطيها هي `slice()`. بدلًا من تعديل المصفوفة، تقوم `slice()` بنسخ أو *استخراج* عدد معين من العناصر إلى مصفوفة جديدة، مع ترك المصفوفة التي استُدعيت عليها دون تغيير. تأخذ `slice()` معلمتين فقط — الأولى هي الفهرس الذي يبدأ عنده الاستخراج، والثانية هي الفهرس الذي يتوقف عنده الاستخراج (سيتم الاستخراج حتى، ولكن دون تضمين العنصر عند هذا الفهرس). انظر المثال التالي:

```js
let weatherConditions = ['rain', 'snow', 'sleet', 'hail', 'clear'];

let todaysWeather = weatherConditions.slice(1, 3);
```

ستحمل `todaysWeather` القيمة `['snow', 'sleet']`، بينما ستظل `weatherConditions` تحتوي على `['rain', 'snow', 'sleet', 'hail', 'clear']`.

فعليًا، أنشأنا مصفوفة جديدة عن طريق استخراج عناصر من مصفوفة موجودة.

# --instructions--

عرّفنا دالة `forecast` التي تأخذ مصفوفة كمعلمة. عدّل الدالة باستخدام `slice()` لاستخراج معلومات من مصفوفة المعلمة وأرجع مصفوفة جديدة تحتوي على العناصر النصية `warm` و`sunny`.

# --hints--

يجب أن تُرجع `forecast` القيمة `["warm", "sunny"]`

```js
assert.deepEqual(
  forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']),
  ['warm', 'sunny']
);
```

يجب أن تستخدم دالة `forecast` طريقة `slice()`

```js
assert(/\.slice\(/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function forecast(arr) {
  // Only change code below this line

  return arr;
}

// Only change code above this line
console.log(forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']));
```

# --solutions--

```js
function forecast(arr) {
  return arr.slice(2,4);
}
```
