---
id: 587d7b7b367417b2b2512b17
title: دمج المصفوفات باستخدام معامل النشر
challengeType: 1
forumTopicId: 301156
dashedName: combine-arrays-with-the-spread-operator
---

# --description--

ميزة كبيرة أخرى لمعامل <dfn>النشر</dfn> هي القدرة على دمج المصفوفات، أو إدخال كل عناصر مصفوفة في أخرى، في أي فهرس. باستخدام الصيغ النحوية التقليدية، يمكننا ربط المصفوفات، لكن هذا يسمح فقط بدمج المصفوفات في نهاية واحدة وبداية أخرى. تجعل صيغة النشر العملية التالية سهلة جدًا:

```js
let thisArray = ['sage', 'rosemary', 'parsley', 'thyme'];

let thatArray = ['basil', 'cilantro', ...thisArray, 'coriander'];
```

سيكون لـ `thatArray` القيمة `['basil', 'cilantro', 'sage', 'rosemary', 'parsley', 'thyme', 'coriander']`.

باستخدام صيغة النشر، حققنا للتو عملية كانت لتكون أكثر تعقيدًا وأكثر إطالة لو استخدمنا الطرق التقليدية.

# --instructions--

لقد عرّفنا دالة `spreadOut` التي تُرجع المتغير `sentence`. عدّل الدالة باستخدام معامل <dfn>النشر</dfn> بحيث تُرجع المصفوفة `['learning', 'to', 'code', 'is', 'fun']`.

# --hints--

يجب أن تُرجع `spreadOut` المصفوفة `["learning", "to", "code", "is", "fun"]`

```js
assert.deepEqual(spreadOut(), ['learning', 'to', 'code', 'is', 'fun']);
```

يجب أن تستخدم دالة `spreadOut` صيغة النشر

```js
assert.notStrictEqual(spreadOut.toString().search(/[...]/), -1);
```

# --seed--

## --seed-contents--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence; // Change this line
  return sentence;
}

console.log(spreadOut());
```

# --solutions--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence = ['learning', ...fragment, 'is', 'fun'];
  return sentence;
}
```
