---
id: 587d7b7d367417b2b2512b1e
title: توليد مصفوفة بكل مفاتيح الكائن باستخدام Object.keys()
challengeType: 1
forumTopicId: 301160
dashedName: generate-an-array-of-all-object-keys-with-object-keys
---

# --description--

يمكننا أيضًا توليد مصفوفة تحتوي على كل المفاتيح المخزنة في كائن باستخدام طريقة `Object.keys()`. تأخذ هذه الطريقة كائنًا كمعلمة وتُرجع مصفوفة من السلاسل النصية تمثل كل خاصية في الكائن. مرة أخرى، لن يكون هناك ترتيب محدد للمدخلات في المصفوفة.

# --instructions--

أكمل كتابة الدالة `getArrayOfUsers` بحيث تُرجع مصفوفة تحتوي على كل الخصائص في الكائن الذي تستقبله كمعلمة.

# --hints--

يجب أن يحتوي كائن `users` فقط على المفاتيح `Alan` و `Jeff` و `Sarah` و `Ryan`

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

يجب أن تُرجع الدالة `getArrayOfUsers` مصفوفة تحتوي على كل المفاتيح في كائن `users`

```js
assert(
  (function () {
    users.Sam = {};
    users.Lewis = {};
    let R = getArrayOfUsers(users);
    return (
      R.indexOf('Alan') !== -1 &&
      R.indexOf('Jeff') !== -1 &&
      R.indexOf('Sarah') !== -1 &&
      R.indexOf('Ryan') !== -1 &&
      R.indexOf('Sam') !== -1 &&
      R.indexOf('Lewis') !== -1
    );
  })() === true
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  // Only change code below this line

  // Only change code above this line
}

console.log(getArrayOfUsers(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  return Object.keys(obj);
}

console.log(getArrayOfUsers(users));
```
