---
id: 5a661e0f1068aca922b3ef17
title: الوصول إلى محتويات المصفوفة باستخدام ترميز الأقواس
challengeType: 1
forumTopicId: 301149
dashedName: access-an-arrays-contents-using-bracket-notation
---

# --description--

الميزة الأساسية لأي هيكل بيانات هي، بالطبع، القدرة ليس فقط على تخزين البيانات، بل أيضًا استرجاع تلك البيانات عند الطلب. الآن بعد أن تعلمنا كيفية إنشاء مصفوفة، لنبدأ بالتفكير في كيفية الوصول إلى معلومات تلك المصفوفة.

عندما نُعرّف مصفوفة بسيطة كما هو موضح أدناه، تحتوي على 3 عناصر:

```js
let ourArray = ["a", "b", "c"];
```

في المصفوفة، لكل عنصر <dfn>فهرس</dfn>. هذا الفهرس يمثل أيضًا موقع ذلك العنصر في المصفوفة، وكيفية الرجوع إليه. من المهم ملاحظة أن مصفوفات JavaScript <dfn>تبدأ فهارسها من الصفر</dfn>، مما يعني أن العنصر الأول في المصفوفة يكون في الموضع ***الصفر*** وليس الأول. لاسترجاع عنصر من المصفوفة، يمكننا وضع الفهرس داخل أقواس وإلحاقه بنهاية المصفوفة، أو أكثر شيوعًا، بمتغير يشير إلى كائن المصفوفة. هذا يُعرف بـ <dfn>ترميز الأقواس</dfn>. على سبيل المثال، إذا أردنا استرجاع `a` من `ourArray` وتعيينها لمتغير، يمكننا فعل ذلك بالشيفرة التالية:

```js
let ourVariable = ourArray[0];
```

الآن يحتوي `ourVariable` على القيمة `a`.

بالإضافة إلى الوصول إلى القيمة المرتبطة بفهرس، يمكنك أيضًا *تعيين* قيمة لفهرس باستخدام نفس الترميز:

```js
ourArray[1] = "not b anymore";
```

باستخدام ترميز الأقواس، قمنا الآن بإعادة تعيين العنصر في الفهرس 1 من السلسلة النصية `b` إلى `not b anymore`. الآن `ourArray` هي `["a", "not b anymore", "c"]`.

# --instructions--

لإكمال هذا التحدٍّ، عيّن الموضع الثاني (الفهرس `1`) في `myArray` إلى أي قيمة تريدها، باستثناء الحرف `b`.

# --hints--

`myArray[0]` يجب أن تكون مساوية للحرف `a`

```js
assert.strictEqual(myArray[0], 'a');
```

`myArray[1]` يجب ألا تكون مساوية للحرف `b`

```js
assert.notStrictEqual(myArray[1], 'b');
```

`myArray[2]` يجب أن تكون مساوية للحرف `c`

```js
assert.strictEqual(myArray[2], 'c');
```

`myArray[3]` يجب أن تكون مساوية للحرف `d`

```js
assert.strictEqual(myArray[3], 'd');
```

# --seed--

## --seed-contents--

```js
let myArray = ["a", "b", "c", "d"];
// Only change code below this line

// Only change code above this line
console.log(myArray);
```

# --solutions--

```js
let myArray = ["a", "b", "c", "d"];
myArray[1] = "e";
```
