---
id: 66f17db06803d11a1bd19a20
title: 构建一个测验游戏
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `questions` 的数组。
1. `questions` 数组应至少包含五个对象，每个对象都应具有密钥 `category`、`question`、`choices` 和 `answer`。
1. `category` 密钥应具有表现问题类别的字串值。
1. `question` 密钥应具有表现问题的字串的值。
1. `choices` 密钥的值应为包含三个字串的数组，这些字串是该问题的备选答案。
1. `answer` 密钥应具有一个字串值，表现问题的正确答案。同时，`answer` 的值应包含在 `choices` 数组中。
1. 你应该有一个名为 `getRandomQuestion` 的函数，它接受一个问题数组作为参数，并从数组中返回一个随机的问题对象。
1. 你应该有一个名为 `getRandomComputerChoice` 的函数，它接受可用选项的数组作为参数，并返回对所选问题的随机答案。
1. 你应该有一个名为 `getResults` 的函数，第一个参数是问题对象，第二个参数是计算机的选择。如果答案正确，函数应返回 `The computer's choice is correct!`。否则，它返回 `The computer's choice is wrong. The correct answer is: <correct-answer>`，其中 `<correct-answer>` 是所选问题的正确答案的值。

# --hints--

你应该创建一个名为 `questions` 的数组。

```js
assert.isArray(questions);
```

`questions` 数组应至少包含五个对象，每个对象都应具有密钥 `category`、`question`、`choices` 和 `answer`。

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

`category` 密钥应具有表现问题类别的字串值。

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

`question` 密钥应具有表现问题的字串的值。

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

`choices` 密钥应具有一个包含三个彼此不同的字串的数组值。

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

`answer` 密钥应具有字串的值。

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

`answer` 的值应包含在 `choices` 数组中。

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

你应该有一个名为 `getRandomQuestion` 的函数，它接受一个问题数组作为参数，并从数组中返回一个随机的问题对象。

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

你应该有一个名为 `getRandomComputerChoice` 的函数，它接受可用选项的数组作为参数，并返回对所选问题的随机答案。

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

你应该有一个名为 `getResults` 的函数。

```js
assert.isFunction(getResults);
```

你的 `getResults` 函数应将问题对象作为第一个参数，将计算机的选择作为第二个参数。

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

如果计算机选择与答案匹配，`getResults` 应该返回 `The computer's choice is correct!`

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

如果计算机选择与答案不匹配，`getResults` 应该返回 `The computer's choice is wrong. The correct answer is: <correct-answer>`，其中 `<correct-answer>` 是所选问题的正确答案的值。

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

你的 `getResults` 函数应使用精确的相等性比较，而不是字串匹配。

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
