---
id: a6e40f1041b06c996f7b2406
title: crea un cercatore del primo elemento
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

In questo laboratorio, creerai una funzione che scorre un array e restituisce il primo elemento che supera una funzione di test (nota anche come "test di verità").

La funzione itererà su tutto l'array e testerà ogni elemento usando la funzione di test fornita. Alla fine, restituirà il primo elemento per cui la funzione di test restituisce `true`.

Esempio:

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**Obiettivo:** soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `findElement` che accetta un array e una funzione come argomenti.
2. La funzione dovrebbe restituire il primo elemento nell'array che supera un test di verità. Questo significa che, chiamando la funzione passata `func`, dato un elemento `x`, il test di verità è superato se `func(x)` è `true`.
3. Se nessun elemento supera il test, la funzione dovrebbe restituire `undefined`.

# --hints--

Dovresti avere una funzione `findElement`.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` dovrebbe restituire `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` dovrebbe restituire `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` dovrebbe restituire `3`.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` dovrebbe restituire `"javascript"`.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` dovrebbe restituire `undefined`.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` dovrebbe restituire `2`.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` dovrebbe restituire `undefined`.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
