---
id: 68e6698b29e562d21b1d260c
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将构建一个主要网页浏览器的列表。已为你提供了 `HTML` 公式化。

首先，在 `body` 元素内使用 `h1` 元素为你的页面添加一个标题，内容为 `List of Major Web Browsers`。

# --hints--

你应该在 `body` 元素内有一个 `h1` 元素。

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

你的 `h1` 元素应包含文本 `List of Major Web Browsers`。

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
