---
id: 68ecfa0f03d88990bb96a33c
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

正如你从之前的课程中回忆的那样，描述列表用于以有组织且易于阅读的形式展示术语和定义。 

这是一个示例：

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

在 `h1` 元素下方，创建一个 `dl` 元素。这个元素将包含浏览器的列表。

# --hints--

你应该有一个 `dl` 元素。

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

你的 `dl` 元素应位于 `h1` 元素下方。

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        --fcc-editable-region--
        
        --fcc-editable-region--
    </body> 
</html>
```

