---
id: 68ecfa1003d88990bb96a33d
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

描述列表由一组术语和详细信息组成。`dt` 元素表现描述术语，而 `dd` 元素表现描述详细信息。 

你将为你的描述列表添加的第一个浏览器是 Google Chrome。

在你的 `dl` 元素内，添加一个包含文本 `Google Chrome` 的 `dt` 元素。

在你的 `dt` 元素下方，添加一个带有以下文本的 `dd` 元素：

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

你的 `dl` 元素内应包含一个 `dt` 元素。

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

你的 `dt` 元素应包含 `Google Chrome` 文本。

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

你应该在你的 `dt` 元素下方有一个 `dd` 元素。

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

你的 `dd` 元素应包含文本 `This is a free web browser developed by Google and first released in 2008.`

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```

