---
id: 69038f5efad6f722def4a214
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

现在是将你的第四个浏览器添加到列表中的时候了。

为你的文档添加另一个包含文本 `Brave` 的 `dt` 元素。

在你的 `dt` 元素下方，添加一个带有以下文本的 `dd` 元素：

```md
This is a free web browser first released in 2016 that is based on the Chromium web browser.
```

# --hints--

你的代码中应该有四个 `dt` 元素。

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 4);
```

你的第四个 `dt` 元素的文本应为 `Brave`。

```js
const dt = document.querySelectorAll('dt')[3];
assert.strictEqual(dt?.textContent.trim(), "Brave");
```

你的代码中应该有四个 `dd` 元素。

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 4);
```

你的第四个 `dd` 元素应包含文本 `This is a free web browser first released in 2016 that is based on the Chromium web browser.`

```js
const ddElement = document.querySelectorAll('dd')[3];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser first released in 2016 that is based on the Chromium web browser.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
