---
id: 69038dd51a41ae215fcc030a
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

你将为 `Safari` 网络浏览器添加到列表中的第三个浏览器。

为你的 `dt` 元素添加另一个包含文本 `Safari` 的元素。

在你的 `dt` 元素下方，添加一个带有以下文本的 `dd` 元素：

```md
This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.
```

# --hints--

你的代码中应该有三个 `dt` 元素。

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 3);
```

你的第三个 `dt` 元素的 `text` 应该是 `Safari`。

```js
const dt = document.querySelectorAll('dt')[2];
assert.strictEqual(dt?.textContent.trim(), "Safari");
```

你的代码中应该有三个 `dd` 元素。

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 3);
```

你的第三个 `dd` 元素应包含文本 `This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.`

```js
const ddElement = document.querySelectorAll('dd')[2];
assert.strictEqual(ddElement?.textContent.trim(), "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
