---
id: 69038fe7da59ca23a5eba82d
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

你将添加到列表中的最后一个浏览器将是 Arc 浏览器。

为你的文档添加另一个包含文本 `Arc` 的 `dt` 元素。

在你的 `dt` 元素下方，添加一个带有以下文本的 `dd` 元素：

```md
This is a free Chromium based web browser first released in 2023 by The Browser Company.
```

通过最后的添加，你的浏览器列表已完成！

# --hints--

你的代码中应该有五个 `dt` 元素。

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 5);
```

你的第五个 `dt` 元素应该有 `Arc` 文本。

```js
const dt = document.querySelectorAll('dt')[4];
assert.strictEqual(dt?.textContent.trim(), "Arc");
```

你的代码中应该有五个 `dd` 元素。

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 5);
```

你的第五个 `dd` 元素应包含文本 `This is a free Chromium based web browser first released in 2023 by The Browser Company.`。

```js
const ddElement = document.querySelectorAll('dd')[4];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free Chromium based web browser first released in 2023 by The Browser Company.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            <dt>Brave</dt>
            <dd>This is a free web browser first released in 2016 that is based on the Chromium web browser.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>List of Browsers and Descriptions</title>
</head>
<body>
    <h1>List of Major Web Browsers</h1>
    <dl>
        <dt>Google Chrome</dt>
        <dd>This is a free web browser developed by Google and first released in 2008.</dd>

        <dt>Firefox</dt>
        <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

        <dt>Safari</dt>
        <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

        <dt>Brave</dt>
        <dd>This is a free web browser first released in 2016 that is based on the Chromium web browser.</dd>

        <dt>Arc</dt>
        <dd>This is a free Chromium based web browser first released in 2023 by The Browser Company.</dd>
    </dl>
</body>
</html>
```
