---
id: 5d5a813321b9e3db6c106a46
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript ist eine leistungsstarke Sprache, mit der Sie interaktive Websites erstellen können.

*Hinweis*: Für alle weiteren Projekte in diesem Curriculum benötigen Sie grundlegende Kenntnisse in HTML und CSS. Wenn Sie neu in HTML und CSS sind, gehen Sie bitte den [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/) durch.

Um zu beginnen, erstellen Sie Ihre Standard-HTML-Boilerplate mit einem `DOCTYPE`, `html`, `head` und `body`, und fügen Sie dann ein `meta`-Tag für das `charset` hinzu. Fügen Sie ein `title`-Element hinzu und verwenden Sie den Text `RPG - Dragon Repeller` dafür. Fügen Sie ein `link`-Tag für Ihr Stylesheet hinzu, um die `styles.css`-Datei zu verlinken.

Erstellen Sie schließlich ein `div`-Element mit `id`, das auf `game` innerhalb Ihres `body` gesetzt ist.

# --hints--

Ihr Code sollte die `DOCTYPE`-Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Sie sollten nach der `DOCTYPE`-Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Sie sollten den Dokumenttyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Sie sollten die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Ihr `html`-Element sollte ein öffnendes Tag haben. Vergessen Sie nicht das `lang`-Attribut.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Ihre `DOCTYPE`-Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Sie sollten ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Sie sollten ein schließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Sie sollten ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Sie sollten ein schließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Die `head`- und `body`-Elemente sollten Geschwister sein.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

Das `head`-Element sollte innerhalb des `html`-Elements sein.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

Das `body`-Element sollte innerhalb des `html`-Elements sein.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

Ihr Code sollte ein `meta`-Element haben.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Ihr `meta`-Element sollte ein `charset`-Attribut mit dem Wert `UTF-8` haben.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Ihr Code sollte ein `title`-Element haben.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Sie sollten ein schließendes `title`-Tag haben.

```js
assert(code.match(/<\/title\s*>/i));
```

Ihr Code sollte ein `link`-Element haben.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Sie sollten ein `div`-Element haben.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Ihr `div`-Element sollte ein `id`-Attribut mit dem Wert `game` haben.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Ihr `div`-Element sollte innerhalb des `body`-Elements sein.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
