---
id: 62a1166ed9a56d439c0770e7
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Erstellen Sie drei `span`-Elemente innerhalb Ihres `#stats`-Elements. Geben Sie jedem von ihnen die Klasse `stat`. Geben Sie dann dem ersten den Text `XP: 0`, dem zweiten den Text `Health: 100` und dem dritten den Text `Gold: 50`.

# --hints--

Sie sollten drei `span`-Elemente innerhalb des `#stats`-Elements haben.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Sie sollten den drei neuen `span`-Elementen die Klasse `stat` geben.

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

Ihr erstes `.stat`-Element sollte den vorgegebenen Text `XP: 0` haben.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

Ihr zweites `.stat`-Element sollte den vorgegebenen Text `Health: 100` haben.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

Ihr drittes `.stat`-Element sollte den vorgegebenen Text `Gold: 50` haben.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
