---
id: 62a115879a6d51422652cbfc
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

In Ihrem Rollenspiel können Nutzer ihre Werte verfolgen, Waffen kaufen und gegen Monster kämpfen. Bevor Sie mit dem interaktiven JavaScript-Teil des Spiels fortfahren, müssen Sie zuerst die HTML-Elemente erstellen, die die Spielinformationen anzeigen.

Erstellen Sie vier `div`-Elemente innerhalb Ihres `#game`-Elements. Geben Sie ihnen nacheinander die folgenden `id`-Werte: `stats`, `controls`, `monsterStats` und `text`.

# --hints--

Sie sollten vier neue `div`-Elemente erstellen.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

Einem der neuen `div`-Elemente sollten Sie eine `id` mit dem Wert `stats` geben.

```js
assert.exists(document.querySelector('div#stats'));
```

Einem der neuen `div`-Elemente sollten Sie eine `id` mit dem Wert `controls` geben.

```js
assert.exists(document.querySelector('div#controls'));
```

Einem der neuen `div`-Elemente sollten Sie eine `id` mit dem Wert `monsterStats` geben.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

Einem der neuen `div`-Elemente sollten Sie eine `id` mit dem Wert `text` geben.

```js
assert.exists(document.querySelector('div#text'));
```

Sie sollten die neuen `div`-Elemente in der richtigen Reihenfolge platzieren.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

Sie sollten die neuen `div`-Elemente innerhalb des `#game`-Elements platzieren.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
