---
id: 62a257659d0d1e2456f24ba2
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Bevor Sie mit dem Schreiben Ihres Projekt-Codes beginnen, sollten Sie ihn in eine eigene Datei verschieben, um alles übersichtlich zu halten.

Entfernen Sie Ihre `console.log("Hello World");`-Zeile. Geben Sie dann Ihrem nun leeren `script`-Element ein `src`-Attribut mit dem Wert `./script.js`.

# --hints--

Sie sollten keine `console.log("Hello World");`-Zeile in Ihrem Code haben.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Ihr `script`-Element sollte ein `src`-Attribut mit dem Wert `./script.js` haben.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
