---
id: 62a255dae245b52317da824a
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Eines der mächtigsten Werkzeuge ist Ihre Entwicklerkonsole. Je nach Browser können Sie diese durch Drücken von `F12` oder `Ctrl+Shift+I` öffnen. Auf dem Mac können Sie `Option + ⌘ + C` drücken und „Konsole“ auswählen. Sie können auch auf die Schaltfläche „Konsole“ oberhalb des Vorschaufensters klicken, um unsere eingebaute Konsole zu sehen.

Die Entwicklerkonsole zeigt Fehler an, die durch Ihren Code verursacht werden, aber Sie können sie auch verwenden, um Werte von Variablen in Ihrem Code zu sehen, was beim Debugging hilfreich ist.

Fügen Sie eine `console.log("Hello World");`-Zeile zwischen Ihre `script`-Tags ein. Klicken Sie dann auf die Schaltfläche „Konsole“, um die Konsole zu öffnen. Sie sollten den Text `"Hello World"` sehen.


# --hints--

Sie sollten eine `console.log("Hello World");`-Zeile zu Ihrem Code hinzufügen.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

Ihre `console.log("Hello World");`-Zeile sollte zwischen Ihren `script`-Tags stehen.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
