---
id: 62a23c1d505bfa13747c8a9b
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Fassen Sie die Zahlen `0`, `100` und `50` in `span`-Elemente ein, und fassen Sie diese neuen `span`-Elemente in `strong`-Elemente ein. Geben Sie dann Ihren neuen `span`-Elementen `id`-Werte von `xpText`, `healthText` und `goldText`.

Ihre Antwort sollte dieser Grundstruktur folgen:

```html
<span class="stat">TEXT <strong><span id="VALUE">TEXT</span></strong></span>
```

# --hints--

Sie sollten ein `strong`-Element in Ihrem ersten `.stat`-Element hinzufügen.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

Ihr erstes neues `strong`-Element sollte ein `span`-Element enthalten.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

Ihr erstes verschachteltes `span`-Element sollte das `id` von `xpText` haben.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'xpText');
```

Ihr erstes `span`-Element sollte den Text `0` umschließen.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '0');
```

Der Text Ihres ersten `.stat`-Elements sollte weiterhin `XP: 0` sein.

```js
const stat = document.querySelectorAll('.stat')[0];
assert.equal(stat.innerText.trimEnd(), 'XP: 0');
```

Sie sollten ein `strong`-Element in Ihrem zweiten `.stat`-Element hinzufügen.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

Ihr zweites neues `strong`-Element sollte ein `span`-Element enthalten.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

Ihr zweites verschachteltes `span`-Element sollte das `id` von `healthText` haben.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'healthText');
```

Ihr zweites `span`-Element sollte den Text `100` umschließen.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '100');
```

Der Text Ihres zweiten `.stat`-Elements sollte weiterhin `Health: 100` sein.

```js
const stat = document.querySelectorAll('.stat')[1];
assert.equal(stat.innerText.trimEnd(), 'Health: 100');
```

Sie sollten ein `strong`-Element in Ihrem dritten `.stat`-Element hinzufügen.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

Ihr drittes neues `strong`-Element sollte ein `span`-Element enthalten.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

Ihr drittes verschachteltes `span`-Element sollte das `id` von `goldText` haben.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'goldText');
```

Ihr drittes `span`-Element sollte den Text `50` umschließen.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '50');
```

Der Text Ihres dritten `.stat`-Elements sollte weiterhin `Gold: 50` sein.

```js
const stat = document.querySelectorAll('.stat')[2];
assert.equal(stat.innerText.trimEnd(), 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
--fcc-editable-region--
        <span class="stat">XP: 0</span>
        <span class="stat">Health: 100</span>
        <span class="stat">Gold: 50</span>
--fcc-editable-region--
      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
