---
id: 62a23d1c5f1c93161f3582ae
title: Schritt 19
challengeType: 0
dashedName: step-19
---

# --description--

Ähnlich wie Ihr `#stats`-Element benötigt Ihr `#monsterStats`-Element zwei `span`-Elemente. Geben Sie ihnen die Klasse `stat` und geben Sie dem ersten Element den Text `Monster Name: ` und dem zweiten den Text `Health: `. Fügen Sie nach dem Text in jedem ein `strong`-Element mit einem leeren verschachtelten `span`-Element hinzu. Geben Sie dem ersten inneren `span`-Element eine `id` von `monsterName` und dem zweiten inneren `span`-Element eine `id` von `monsterHealth`.

# --hints--

Ihr `#monsterStats`-Element sollte zwei `span`-Elemente haben.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

Ihre neuen `span`-Elemente sollten beide einen `class`-Wert von `stat` haben.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

Ihr erstes `span`-Element sollte ein `strong`-Element mit einem leeren verschachtelten `span`-Element haben.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Ihr zweites `span`-Element sollte ein `strong`-Element mit einem leeren verschachtelten `span`-Element haben.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Ihr erstes `span`-Element sollte den Text `Monster Name: ` haben. Achten Sie darauf, dass der Abstand korrekt ist.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+/);
```

Ihr zweites `span`-Element sollte den Text `Health: ` haben. Achten Sie darauf, dass der Abstand korrekt ist.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[1].innerHTML, /Health:\s+/);
```

Ihre `strong`- und `span`-Elemente sollten nach dem Text kommen.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+<strong>/);
assert.match(spans[1].innerHTML, /Health:\s+<strong>/);
```

Ihr erstes inneres `span`-Element sollte eine `id` von `monsterName` haben.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[0]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterName');
```

Ihr zweites inneres `span`-Element sollte eine `id` von `monsterHealth` haben.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[1]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterHealth');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
--fcc-editable-region--
      <div id="monsterStats">

      </div>
--fcc-editable-region--
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
