---
id: 62a39f5a5790eb27c1e5d4bf
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Deine Ansicht wurde auf deine neue `script.js` Datei umgestellt. Denke daran, dass du die oben genannten Registerkarten verwenden kannst, um zwischen den Dateien wechseln zu können.

Füge deine `console.log("Hello World");`-Zeile zu dieser Datei hinzu und sehe, wie sie in deiner Konsole erscheint.

# --hints--

Du solltest eine `console.log("Hello World");`-Zeile in deinem Code haben.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
