---
id: 62a3a7e4f1060e2fc5ffb34b
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Erstellen Sie eine weitere Variable namens `currentWeaponIndex` und setzen Sie sie auf `0`.

# --hints--

Sie sollten `let` verwenden, um eine Variable namens `currentWeaponIndex` zu deklarieren.

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

Sie sollten CamelCase verwenden, um Ihre Variable zu benennen.

```js
assert.match(code, /currentWeaponIndex/);
```

Ihre `currentWeaponIndex`-Variable sollte auf `0` gesetzt werden.

```js
assert.equal(currentWeaponIndex, 0);
```

Sie sollten Ihre Variable mit `0` initialisieren.

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
