---
id: 62a3b41c9494f937560640ab
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Deklarieren Sie zwei weitere Variablen mit den Namen `monsterHealth` und `inventory`.

Weisen Sie Ihrer `inventory`-Variablen den Wert eines Arrays zu, das den String `"stick"` enthält.

Denken Sie daran, dass Sie im vorherigen Projekt bereits mit Arrays so gearbeitet haben:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

Sie sollten `let` verwenden, um Ihre `monsterHealth`-Variable zu deklarieren.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Sie sollten CamelCase verwenden, um Ihre `monsterHealth`-Variable zu benennen.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` sollte keinen Wert haben.

```js
assert.isUndefined(monsterHealth);
```

Sie sollten Ihrer `monsterHealth`-Variablen keinen Wert zuweisen.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

Sie sollten `let` verwenden, um Ihre `inventory`-Variable zu deklarieren.

```js
assert.match(code, /let\s+inventory/i);
```

Ihre `inventory`-Variable sollte weiterhin ein Array sein.

```js
assert.isArray(inventory);
```

Ihre `inventory`-Variable sollte nur einen Wert haben.

```js
assert.lengthOf(inventory, 1);
```

Ihre `inventory`-Variable sollte den String `"stick"` enthalten.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
