---
id: 62a3b79d520a7f3d0e25afd6
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript interagiert mit dem HTML über das <dfn>Document Object Model</dfn>, oder DOM. Das DOM ist ein Baum von Objekten, der das HTML darstellt. Sie können auf das HTML über das `document`-Objekt zugreifen, das Ihr gesamtes HTML-Dokument darstellt.

Eine Methode, um bestimmte Elemente in Ihrem HTML zu finden, ist die Verwendung der `querySelector()`-Methode. Die `querySelector()`-Methode nimmt einen CSS-Selektor als Argument und gibt das erste Element zurück, das diesem Selektor entspricht. Zum Beispiel, um das `<h1>`-Element in Ihrem HTML zu finden, würden Sie schreiben:

```js
let h1 = document.querySelector("h1");
```

Beachten Sie, dass `h1` ein String ist und dem CSS-Selektor entspricht, den Sie verwenden würden.

Erstellen Sie eine `button1`-Variable und verwenden Sie `querySelector()`, um ihr Ihr Element mit dem `id` von `button1` zuzuweisen. Denken Sie daran, dass CSS-`id`-Selektoren mit einem `#` als Präfix versehen sind.

# --hints--

Sie sollten `let` verwenden, um eine `button1`-Variable zu deklarieren.

```js
assert.match(code, /let\s+button1/);
```

Sie sollten `document.querySelector()` verwenden.

```js
assert.match(code, /document\.querySelector/);
```

Sie sollten den `"#button1"`-Selektor verwenden.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
