---
id: 62a3a0a3c0a4b32915d26a6e
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Entfernen Sie Ihre `console.log("Hello World");`-Zeile, um mit dem Schreiben Ihres Projekt-Codes zu beginnen.

Im vorherigen Projekt haben Sie gelernt, wie man mit Variablen und dem Schlüsselwort `let` so arbeitet:

```js
let camperbot;
```

Sie haben auch gelernt, wie man eine Variable mit einem Wert initialisiert, so wie hier:

```js
let age = 32;
```

Verwenden Sie das Schlüsselwort `let`, um eine Variable namens `xp` zu deklarieren und ihr den Wert `0`, eine Zahl, zuzuweisen.

# --hints--

Sie sollten keine `console.log("Hello World");`-Zeile in Ihrem Code haben.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Sie sollten das Schlüsselwort `let` verwenden, um Ihre Variable zu deklarieren.

```js
assert.match(code, /let/);
```

Sie sollten eine Variable namens `xp` deklarieren.

```js
assert.match(code, /let\s+xp/);
```

`xp` sollte den Wert `0` haben.

```js
assert.equal(xp, 0);
```

Sie sollten die Variable `xp` mit `0` initialisieren.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
