---
id: 62a39f5a5790eb27c1e5d4bf
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Ihre Ansicht wurde auf Ihre neue `script.js`-Datei umgeschaltet. Denken Sie daran, dass Sie die Tabs oben verwenden können, um zwischen Dateien zu wechseln.

Fügen Sie Ihre `console.log("Hello World");`-Zeile in diese Datei ein und sehen Sie, wie sie in Ihrer Konsole erscheint.

# --hints--

Sie sollten eine `console.log("Hello World");`-Zeile in Ihrem Code haben.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
