---
id: 62b46e3a8d4be31be5af793d
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Wir sind auf ein kleines Problem gestoßen. Sie versuchen, auf Ihrer Seite ein Button-Element abzufragen, aber Ihr `script`-Tag befindet sich im `head` Ihres HTML. Das bedeutet, Ihr Code wird ausgeführt, bevor der Browser das HTML vollständig gelesen hat, und Ihr `document.querySelector()` wird den Button nicht sehen – weil der Browser ihn noch nicht verarbeitet hat.

Um das zu beheben, verschieben Sie Ihr `script`-Element aus dem `head`-Element heraus und platzieren Sie es ans Ende Ihres `body`-Elements (direkt vor dem schließenden `</body>`-Tag).

# --hints--

Ihr `script`-Element sollte sich nicht in Ihrem `head`-Element befinden.

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

Ihr `script`-Element sollte am Ende Ihres `body`-Elements stehen.

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
