---
id: 62b46e3a8d4be31be5af793d
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Wir sind auf ein kleines Problem gestoßen. Du versuchst, deine Seite nach einem Button-Element zu durchsuchen, aber dein `script`-Tag befindet sich im `head` deines HTMLs. Dies bedeutet, dass dein Code ausgeführt wird, bevor der Browser das HTML zu Ende gelesen hat, und dass dein `document.querySelector()` die Schaltfläche nicht sehen wird - da der Browser sie noch nicht verarbeitet hat.

To fix this, move your `script` element out of the `head` element, and place it at the end of your `body` element (just before the closing `</body>` tag).

# --hints--

Dein `script`-Element sollte sich nicht in deinem `head`-Element befinden.

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

Dein `script`-Element sollte am Ende deines `body`-Elements stehen.

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
