---
id: 62fc20387ef88d1d1998aac5
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Verwenden Sie `querySelector()`, um die anderen beiden `button`-Elemente anhand ihrer `id`s zu erhalten: `button2` und `button3`. Speichern Sie sie in Variablen mit den Namen `button2` und `button3`. Denken Sie daran, `const` zu verwenden.

# --hints--

Sie sollten eine `button2`-Variable mit `const` deklarieren.

```js
assert.match(code, /const\s+button2/);
```

Ihre `button2`-Variable sollte den Wert Ihres `"#button2"`-Elements haben.

```js
assert.match(code, /const\s+button2\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#button2\1\s*\)/);
```

Sie sollten eine `button3`-Variable mit `const` deklarieren.

```js
assert.match(code, /const\s+button3/);
```

Ihre `button3`-Variable sollte den Wert Ihres `"#button3"`-Elements haben.

```js
assert.match(code, /const\s+button3\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#button3\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");

--fcc-editable-region--
```
