---
id: 587d824a367417b2b2512c44
title: مدقق سعر السهم
challengeType: 4
forumTopicId: 301572
dashedName: stock-price-checker
---

# --description--

ابنِ تطبيق جافاسكريبت متكامل الوظائف يشبه هذا عمليًا: <a href="https://stock-price-checker.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker.freecodecamp.rocks/</a>.

بما أن جميع واجهات برمجة التطبيقات الموثوقة لأسعار الأسهم تتطلب مفتاح API، فقد أنشأنا حلًا بديلاً. استخدم <a href="https://stock-price-checker-proxy.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker-proxy.freecodecamp.rocks/</a> للحصول على معلومات محدثة عن أسعار الأسهم دون الحاجة للتسجيل للحصول على مفتاح خاص بك.

سيتطلب العمل على هذا المشروع كتابة شفرتك باستخدام إحدى الطرق التالية:

- استنساخ <a href="https://github.com/freeCodeCamp/boilerplate-project-stockchecker/" target="_blank" rel="noopener noreferrer nofollow">مستودع GitHub هذا</a> وإكمال مشروعك محليًا.
- استخدام منشئ مواقع من اختيارك لإكمال المشروع. تأكد من دمج جميع الملفات من مستودع GitHub الخاص بنا.

# --instructions--

1. عرّف متغير البيئة `NODE_ENV` إلى القيمة `test` بدون علامات اقتباس
2. أكمل المشروع في `routes/api.js` أو بإنشاء معالِج/متحكم
3. أضف أي ميزات أمان إلى `server.js`
4. أنشئ جميع الاختبارات الوظيفية في `tests/2_functional-tests.js`

**ملاحظة** اعتبارات الخصوصية: بسبب شرط قبول إعجاب واحد فقط لكل عنوان IP، يجب حفظ عناوين IP. من المهم الالتزام بقوانين خصوصية البيانات مثل اللائحة العامة لحماية البيانات. خيار واحد هو الحصول على إذن لحفظ بيانات المستخدم، لكن من الأسهل إخفاء هويتها. في هذا التحدٍّ، تذكر إخفاء هوية عناوين IP قبل حفظها في قاعدة البيانات. إذا كنت تحتاج أفكارًا، يمكنك تشفير البيانات، أو تقصيرها، أو تعيين جزء من عنوان IP إلى 0.

اكتب الاختبارات التالية في `tests/2_functional-tests.js`:

- عرض سهم واحد: طلب GET إلى `/api/stock-prices/`
- عرض سهم واحد والإعجاب به: طلب GET إلى `/api/stock-prices/`
- عرض نفس السهم والإعجاب به مرة أخرى: طلب GET إلى `/api/stock-prices/`
- عرض سهمين: طلب GET إلى `/api/stock-prices/`
- عرض سهمين والإعجاب بهما: طلب GET إلى `/api/stock-prices/`

# --hints--

يمكنك تقديم مشروعك الخاص، وليس عنوان URL المثال.

```js
  assert(
    !/.*\/stock-price-checker\.freecodecamp\.rocks/.test(code)
  );
```

يجب تعيين سياسات أمان المحتوى للسماح فقط بتحميل السكربتات وCSS من خادمك.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("script-src 'self'")
  );
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("style-src 'self'")
  );
```

يمكنك إرسال طلب `GET` إلى `/api/stock-prices`، مع تمرير رمز سهم NASDAQ كمعلمة استعلام `stock`. ستحتوي الكائن المعاد على خاصية باسم `stockData`.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  assert.property(parsed, 'stockData');
```

تتضمن خاصية `stockData` رمز `stock` كسلسلة نصية، و`price` كرقم، و`likes` كرقم.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker.price, 'number');
  assert.typeOf(ticker.likes, 'number');
  assert.typeOf(ticker.stock, 'string');
```

يمكنك أيضًا تمرير حقل `like` بقيمة `true` (قيمة منطقية) لإضافة إعجابك للسهم/الأسهم. يجب قبول إعجاب واحد فقط لكل عنوان IP.

```js

```

إذا مررت سهمين، ستكون القيمة المعادة مصفوفة بمعلومات عن كلا السهمين. بدلًا من `likes`، ستعرض `rel_likes` (الفرق بين الإعجابات على السهمين) لكلا كائنَي `stockData`.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG&stock=MSFT'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker, 'array');
  assert.property(ticker[0], 'rel_likes');
  assert.property(ticker[1], 'rel_likes');
```

اكتملت جميع الاختبارات الوظيفية الخمسة ونجحت.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 5);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
  });
```

