---
id: 587d824a367417b2b2512c45
title: لوحة الرسائل المجهولة
challengeType: 4
forumTopicId: 301568
dashedName: anonymous-message-board
---

# --description--

ابنِ تطبيق جافاسكريبت متكامل يشبه وظيفيًا هذا: <a href="https://anonymous-message-board.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://anonymous-message-board.freecodecamp.rocks/</a>.

سيتطلب منك العمل على هذا المشروع كتابة شفرتك باستخدام إحدى الطرق التالية:

- استنساخ <a href="https://github.com/freeCodeCamp/boilerplate-project-messageboard/" target="_blank" rel="noopener noreferrer nofollow">مستودع GitHub هذا</a> وإكمال مشروعك محليًا.
- استخدام منشئ مواقع من اختيارك لإكمال المشروع. تأكد من دمج جميع الملفات من مستودع GitHub الخاص بنا.

# --instructions--

1. عيّن `NODE_ENV` إلى test بدون علامات اقتباس عندما تكون مستعدًا لكتابة الاختبارات، وعيّن DB إلى سلسلة اتصال قاعدة البيانات الخاصة بك (في `.env`)
2. يُنصح بإنشاء المتحكمات/المعالجات والتعامل مع التوجيه في `routes/api.js`
3. ستضيف أي ميزات أمان إلى `server.js`

اكتب الاختبارات التالية في `tests/2_functional-tests.js`:

- إنشاء خيط جديد: طلب POST إلى `/api/threads/{board}`
- عرض أحدث 10 خيوط مع 3 ردود لكل منها: طلب GET إلى `/api/threads/{board}`
- حذف خيط بكلمة مرور خاطئة: طلب DELETE إلى `/api/threads/{board}` مع `delete_password` غير صحيح
- حذف خيط بكلمة مرور صحيحة: طلب DELETE إلى `/api/threads/{board}` مع `delete_password` صحيح
- الإبلاغ عن خيط: طلب PUT إلى `/api/threads/{board}`
- إنشاء رد جديد: طلب POST إلى `/api/replies/{board}`
- عرض خيط واحد مع كل الردود: طلب GET إلى `/api/replies/{board}`
- حذف رد بكلمة مرور خاطئة: طلب DELETE إلى `/api/replies/{board}` مع `delete_password` غير صحيح
- حذف رد بكلمة مرور صحيحة: طلب DELETE إلى `/api/replies/{board}` مع `delete_password` صحيح
- الإبلاغ عن رد: طلب PUT إلى `/api/replies/{board}`

# --hints--

يمكنك تقديم مشروعك الخاص، وليس عنوان URL المثال.

```js
  assert(
    !/.*\/anonymous-message-board\.freecodecamp\.rocks/.test(
      code
    )
  );
```

اسمح بتحميل موقعك فقط داخل iFrame على صفحاتك الخاصة.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-frame-options']?.includes('SAMEORIGIN'));
```

لا تسمح بعملية استباق DNS.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-dns-prefetch-control']?.includes('off'));
```

اسمح لموقعك فقط بإرسال المرجع لصفحاتك الخاصة.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['referrer-policy']?.includes('same-origin'));
```

يمكنك إرسال طلب POST إلى `/api/threads/{board}` مع بيانات نموذج تشمل `text` و`delete_password`. ستحتوي السجلات المحفوظة في قاعدة البيانات على الأقل على الحقول `_id`، `text`، `created_on` (تاريخ ووقت)، `bumped_on` (تاريخ ووقت، يبدأ مثل `created_on`)، `reported` (قيمة منطقية)، `delete_password`، و`replies` (مصفوفة).

```js
  const date = new Date();
  const text = `fcc_test_${date}`;
  const deletePassword = 'delete_me';
  const data = { text, delete_password: deletePassword };
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(url + '/api/threads/fcc_test');
    const parsed = await checkData.json();
    try {
      assert.equal(parsed[0].text, text);
      assert.isNotNull(parsed[0]._id);
      assert.equal(new Date(parsed[0].created_on).toDateString(), date.toDateString());
      assert.equal(parsed[0].bumped_on, parsed[0].created_on);
      assert.isArray(parsed[0].replies);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

يمكنك إرسال طلب POST إلى `/api/replies/{board}` مع بيانات نموذج تشمل `text`، `delete_password`، و`thread_id`. سيُحدّث هذا تاريخ `bumped_on` إلى تاريخ التعليق. في مصفوفة `replies` للخيط، سيُحفظ كائن يحتوي على الأقل على الخصائص `_id`، `text`، `created_on`، `delete_password`، و`reported`.

```js
  const url = code;
  const body = await fetch(url + '/api/threads/fcc_test');
  const thread = await body.json();

  const date = new Date();
  const text = `fcc_test_reply_${date}`;
  const delete_password = 'delete_me';
  const thread_id = thread[0]._id;
  const replyCount = thread[0].replies.length;

  const data = { text, delete_password, thread_id };
  const res = await fetch(url + '/api/replies/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
    const parsed = await checkData.json();
    try {
      assert.equal(parsed.replies.length, replyCount + 1);
      assert.equal(parsed.replies[0].text, text);
      assert.equal(parsed._id, thread_id);
      assert.equal(parsed.bumped_on, parsed.replies[0].created_on);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

يمكنك إرسال طلب GET إلى `/api/threads/{board}`. ستُرجع مصفوفة بأحدث 10 خيوط تم تحديثها على اللوحة مع أحدث 3 ردود فقط لكل منها. لن تُرسل حقول `reported` و`delete_password` إلى العميل.

```js
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test');

  if (res.ok) {
    const threads = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isAtMost(threads.length, 10);
      for (let i = 0; i < threads.length; i++) {
        assert.containsAllKeys(threads[i], ["_id", "text", "created_on", "bumped_on", "replies"]);
        assert.isAtMost(threads[i].replies.length, 3);
        assert.notExists(threads[i].delete_password);
        assert.notExists(threads[i].reported);
        for (let j = 0; j < threads[i].replies.length; j++) {
          assert.notExists(threads[i].replies[j].delete_password);
          assert.notExists(threads[i].replies[j].reported);
        }
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

يمكنك إرسال طلب GET إلى `/api/replies/{board}?thread_id={thread_id}`. ستُرجع الخيط بالكامل مع كل ردوده، مع استبعاد نفس الحقول من العميل كما في الاختبار السابق.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);

  if (res.ok) {
    const thread = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isObject(thread);
      assert.containsAllKeys(thread, ["_id", "text", "created_on", "bumped_on", "replies"]);
      assert.isArray(thread.replies);
      assert.notExists(thread.delete_password);
      assert.notExists(thread.reported);
      for (let i = 0; i < thread.replies.length; i++) {
        assert.notExists(thread.replies[i].delete_password);
        assert.notExists(thread.replies[i].reported);
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

يمكنك إرسال طلب DELETE إلى `/api/threads/{board}` وتمرير `thread_id` و`delete_password` لحذف الخيط. ستُرجع السلسلة `incorrect password` أو `success`.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  let data = { thread_id, delete_password: "wrong_password" };
  const res_invalid = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  data = { thread_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    const not_deleted = await res_invalid.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      assert.equal(not_deleted, "incorrect password");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

يمكنك إرسال طلب DELETE إلى `/api/replies/{board}` وتمرير `thread_id`، `reply_id`، و`delete_password`. ستُرجع السلسلة `incorrect password` أو `success`. عند النجاح، سيتغير نص `reply_id` إلى `[deleted]`.

```js
  const url = code;

  const thread_data = {
    text: "fcc_test_thread",
    delete_password: "delete_me",
  };
  await fetch(`${url}/api/threads/fcc_test`, {
    method: "POST",
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(thread_data)
  });
  let res = await fetch(`${url}/api/threads/fcc_test`);
  let threads = await res.json();
  const thread_id = threads[0]._id;
  
  const reply_data = { thread_id, text: "fcc_test_reply", delete_password: "delete_me" };
  await fetch(`${url}/api/replies/fcc_test`, {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(reply_data)
  });
  res = await fetch(`${url}/api/threads/fcc_test`);
  threads = await res.json();
  const reply_id = threads[0].replies[0]._id;

  const data = { thread_id, reply_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/replies/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
      const thread = await res.json();
      assert.equal(thread._id, thread_id);
      assert.equal(thread.replies[0]._id, reply_id);
      assert.equal(thread.replies[0].text, "[deleted]");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

يمكنك إرسال طلب PUT إلى `/api/threads/{board}` وتمرير `thread_id`. ستُرجع السلسلة النصية `reported`. ستتغير قيمة `reported` الخاصة بـ `thread_id` إلى `true`.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const data = { thread_id };
  
  res = await fetch(`${url}/api/threads/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

يمكنك إرسال طلب PUT إلى `/api/replies/{board}` وتمرير `thread_id` و `reply_id`. ستُعاد السلسلة النصية `reported`. ستتغير قيمة `reported` الخاصة بـ `reply_id` إلى `true`.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const reply_id = threads[0].replies[0]._id;
  const data = { thread_id, reply_id };
  
  res = await fetch(`${url}/api/replies/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

اكتملت جميع الاختبارات الوظيفية العشرة ونجحت.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 10);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
    assert.isAtLeast(test.assertions.length, 1);
  });
```

