---
id: 66ba762af611169359d9d369
title: 建置一篇報紙文章
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該將你的 HTML 文件的根 `font-size` 設定為 `24px`。
1. 你應該有一個 `newspaper` 類別的元素，該元素包含你所有其他元素。
1. 你的 `.newspaper` 元素應該有 `16px` 的 `font-size`，並使用 `Open Sans` 字型，備用字型為 `sans-serif`。
1. 在你的 `.newspaper` 元素中，你應該至少有七個以上的元素：一個用於報紙名稱且具有 `name` 類別，一個用於文章日期且具有 `date` 類別，一個用於標題且具有 `headline` 類別，一個用於副標題且具有 `sub-headline` 類別，一個用於作者且具有 `author` 類別，以及兩個段落，每個段落具有 `text` 類別。所有這些元素都應該填入你的文章資訊。
1. 你的 `.name` 元素的 `font-size` 應該是根元素 `font-size` 的兩倍，並且應該使用 `Times New Roman` 字型，備用字型為 `serif`。
1. 你的 `.name` 和 `.author` 元素應該使用 CSS 將所有字元轉成大寫。
1. 你的 `.headline` 元素的 `font-size` 應該是其父元素 `font-size` 的兩倍，並且應該是粗體。
1. 你的 `.sub-headline` 元素的 `font-weight` 應為 `100`，`font-size` 應為其父元素 `font-size` 的 1.5 倍，且應為斜體。
1. 你的 `.author` 應該使用 CSS 來設為粗體。
1. 你的 `.text` 元素應該有 `text-indent` 為 `20px`。
1. 你的 `.text` 元素的 `line-height` 應該是其父元素 `font-size` 的兩倍。
1. 你的 `.text` 元素的第一個字母應該是粗體，且大小為其父元素 `font-size` 的兩倍。為此請使用 `::first-letter` 選擇器。

# --hints--

你應該有一個 `html` 選擇器，將其 `font-size` 設定為 `24px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

你應該有一個類別為 `newspaper` 的元素。

```js
assert.exists(document.querySelector('.newspaper'));
```

你應該有一個 `.newspaper` 選擇器，將其元素的 `font-size` 設定為 `16px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

你應該有一個 `.newspaper` 選擇器，將其元素的 `font-family` 設定為 `'Open Sans', sans-serif`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

你應該在你的 `.newspaper` 元素中有一個具有 `name` 類別的元素。

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

你的 `.name` 元素應該包含你的報紙名稱。

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

你應該在你的 `.newspaper` 元素中有一個 `class` 為 `date` 的元素。

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

你的 `.date` 元素應該包含日期。

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

你應該在你的 `.newspaper` 元素中有一個 `headline` 類別的元素。

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

你的 `.headline` 元素應該包含你的文章標題。

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

你應該在你的 `.newspaper` 元素中有一個類別為 `sub-headline` 的元素。

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

你的 `.sub-headline` 元素應該包含你的文章副標題。

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

你應該在你的 `.newspaper` 元素中有一個具有 `author` 類別的元素。

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

你的 `.author` 元素應該包含你的文章作者。

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

你應該在你的 `.newspaper` 元素中至少有兩個段落元素，每個段落元素都有一個 `text` 類別。

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

你的 `.text` 元素應該包含你的文章文字。

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

你應該有一個 `.name` 選擇器，將其元素的 `font-size` 設定為 `2rem`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

你應該有一個 `.name` 選擇器，將其元素的 `font-family` 設定為 `'Times New Roman', serif`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

`.name` 元素應該將 `text-transform` 屬性設為 `uppercase`。

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

`.author` 元素的 `text-transform` 屬性應設定為 `uppercase`。

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

你應該有一個 `.headline` 選擇器，將其元素的 `font-size` 設定為 `2em`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

你應該有一個 `.headline` 選擇器，將其元素的 `font-weight` 設定為 `bold`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

你應該有一個 `.sub-headline` 選擇器，將其元素的 `font-weight` 設定為 `100`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

你應該有一個 `.sub-headline` 選擇器，將其元素的 `font-size` 設定為 `1.5em`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

你應該有一個 `.sub-headline` 選擇器，將其元素的 `font-style` 設為 `italic`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

你應該有一個 `.author` 選擇器，將其元素的 `font-weight` 設為 `bold`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

你應該有一個 `.text` 選擇器，將其元素的 `text-indent` 設為 `20px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

你應該有一個 `.text` 選擇器，將其元素的 `line-height` 設定為 `2em`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

你應該有一個 `.text::first-letter` 選擇器，將其元素的 `font-weight` 設為 `bold`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

你應該有一個 `.text::first-letter` 選擇器，將其元素的 `font-size` 設定為 `2em`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
