---
id: 68de40f53fcaf106f51cc20f
title: Einen HTML-Audio- und Video-Player erstellen
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `h1`-Element für den Haupttitel der Seite haben.
2. Sie sollten zwei `section`-Elemente haben.
3. Innerhalb des ersten `section`-Elements sollten Sie ein `h2`-Element für den Titel des abgespielten Videos haben.
4. Unterhalb des `h2`-Elements sollten Sie ein `video`-Element mit den `controls`- und `width`-Attributen haben. Das `width`-Attribut sollte auf `640` gesetzt sein.
5. Innerhalb des `video`-Elements sollten Sie ein `source`-Element mit einem `src`-Attribut, das auf eine Videodatei zeigt, und einem `type`-Attribut haben.
   - Sie können `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4` verwenden.
6. Innerhalb des zweiten `section`-Elements sollten Sie ein `h2`-Element für den Titel des abgespielten Songs haben.
7. Unterhalb des `h2`-Elements sollten Sie ein `audio`-Element mit den `controls`- und `loop`-Attributen sowie einem `src`-Attribut haben, das auf eine Audiodatei zeigt.
   - Sie können `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3` verwenden.
   - Oder `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`.

# --hints--

Sie sollten ein `h1`-Element für den Haupttitel der Seite haben.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

Ihr `h1` sollte nicht leer sein.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

Sie sollten nur ein `h1`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Sie sollten zwei `section`-Elemente haben.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

Innerhalb des ersten `section`-Elements sollten Sie ein `h2`-Element für den Titel des abgespielten Videos haben.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Unterhalb des `h2`-Elements sollten Sie ein `video`-Element haben.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

Das `video`-Element sollte ein `controls`-Attribut haben.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

Das `video`-Element sollte ein `width`-Attribut haben, das auf `640` gesetzt ist.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

Innerhalb des `video`-Elements sollten Sie ein `source`-Element haben.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

Das `source`-Element sollte ein `src`-Attribut haben, das auf eine Videodatei zeigt.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

Das `source`-Element sollte ein `type`-Attribut haben.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

Das `type`-Attribut sollte mit dem Medientyp der Dateiendung des `src`-Attributs übereinstimmen.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

Innerhalb des zweiten `section`-Elements sollten Sie ein `h2`-Element für den Titel des abgespielten Songs haben.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Alle Ihre `h2`-Elemente sollten Text enthalten.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

Unterhalb des `h2`-Elements sollten Sie ein `audio`-Element haben.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

Das `audio`-Element sollte ein `controls`-Attribut haben.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

Das `audio`-Element sollte ein `loop`-Attribut haben.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

Das `audio`-Element sollte ein `src`-Attribut haben, das auf eine Audiodatei zeigt.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
