---
id: 5e44412c903586ffb414c94c
title: 산술 포매터 프로젝트 만들기
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462359
dashedName: build-an-arithmetic-formatter-project
---

# --description--

초등학생들은 종종 산술 문제를 세로로 배열하여 더 쉽게 풉니다. 예를 들어, "235 + 52"는 다음과 같이 됩니다:

```py
  235
+  52
-----
```

산술 문제 문자열 목록을 받아서 문제들을 세로로 나란히 배열한 문자열을 반환하는 `arithmetic_arranger` 함수를 완성하세요. 함수는 선택적으로 두 번째 인수를 받을 수 있습니다. 두 번째 인수가 `True`로 설정되면 답도 표시해야 합니다.

## 예시

함수 호출:

```py
arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])
```

출력:

```py
   32      3801      45      123
+ 698    -    2    + 43    +  49
-----    ------    ----    -----
```

함수 호출:

```py
arithmetic_arranger(["32 + 8", "1 - 3801", "9999 + 9999", "523 - 49"], True)
```

출력:

```py
  32         1      9999      523
+  8    - 3801    + 9999    -  49
----    ------    ------    -----
  40     -3800     19998      474
```

## 규칙

제공된 문제가 올바르게 형식화되면 함수는 올바른 변환을 반환하고, 그렇지 않으면 사용자에게 의미 있는 오류를 설명하는 **문자열**을 **반환**해야 합니다.

- 오류를 반환하는 상황:
  - 함수에 너무 많은 문제가 제공된 경우. 제한은 **다섯 개**이며, 이를 초과하면 다음을 반환합니다:
    `'Error: Too many problems.'`
  - 함수가 허용하는 연산자는 **덧셈**과 **뺄셈**입니다. 곱셈과 나눗셈은 오류를 반환합니다. 이 항목에 언급되지 않은 다른 연산자는 테스트할 필요가 없습니다. 반환되는 오류는 다음과 같습니다:
    `"Error: Operator must be '+' or '-'."`
  - 각 숫자(피연산자)는 숫자만 포함해야 합니다. 그렇지 않으면 함수는 다음을 반환합니다:
    `'Error: Numbers must only contain digits.'`
  - 각 피연산자(연산자 양쪽의 숫자)는 최대 네 자리 숫자여야 합니다. 그렇지 않으면 반환되는 오류 문자열은 다음과 같습니다:
    `'Error: Numbers cannot be more than four digits.'`
- 사용자가 올바른 형식의 문제를 제공한 경우, 반환하는 변환은 다음 규칙을 따릅니다:
  - 연산자와 두 피연산자 중 더 긴 쪽 사이에는 한 칸의 공백이 있어야 하며, 연산자는 두 번째 피연산자와 같은 줄에 위치합니다. 두 피연산자는 제공된 순서대로 유지되어야 하며(첫 번째가 위, 두 번째가 아래).
  - 숫자는 오른쪽 정렬되어야 합니다.
  - 각 문제 사이에는 네 칸의 공백이 있어야 합니다.
  - 각 문제 아래에는 대시가 있어야 하며, 대시는 각 문제의 전체 길이에 걸쳐 있어야 합니다. (위 예시가 어떻게 보여야 하는지 나타냅니다.)

참고: F12를 눌러 브라우저 콘솔을 열면 테스트의 더 자세한 출력을 볼 수 있습니다.

# --hints--

`arithmetic_arranger(["3801 - 2", "123 + 49"])`는 `  3801      123\n-    2    +  49\n------    -----`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3801 - 2", "123 + 49"]), '  3801      123\\n-    2    +  49\\n------    -----')`);
  }
})
```

`arithmetic_arranger(["1 + 2", "1 - 9380"])`는 `  1         1\n+ 2    - 9380\n---    ------`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["1 + 2", "1 - 9380"]), '  1         1\\n+ 2    - 9380\\n---    ------')`);
  }
})
```

`arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"])`는 `    3      3801      45      123\n+ 855    -    2    + 43    +  49\n-----    ------    ----    -----`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"]), '    3      3801      45      123\\n+ 855    -    2    + 43    +  49\\n-----    ------    ----    -----')`);
  }
})
```

`arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"])`는 `  11      3801      1      123         1\n+  4    - 2999    + 2    +  49    - 9380\n----    ------    ---    -----    ------`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"]), '  11      3801      1      123         1\\n+  4    - 2999    + 2    +  49    - 9380\\n----    ------    ---    -----    ------')`);
  }
})
```

`arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"])`는 `'Error: Too many problems.'`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"]), 'Error: Too many problems.')`);
  }
})
```

`arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"])`는 `"Error: Operator must be '+' or '-'."`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Operator must be '+' or '-'.")`);
  }
})
```

`arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"])`는 `'Error: Numbers cannot be more than four digits.'`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers cannot be more than four digits.")`);
  }
})
```

`arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"])`는 `'Error: Numbers must only contain digits.'`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers must only contain digits.")`);
  }
})
```

`arithmetic_arranger(["3 + 855", "988 + 40"], True)`는 `    3      988\n+ 855    +  40\n-----    -----\n  858     1028`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "988 + 40"], True), "    3      988\\n+ 855    +  40\\n-----    -----\\n  858     1028")`);
  }
})
```

`arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True)`는 `   32         1      45      123      988\n- 698    - 3801    + 43    +  49    +  40\n-----    ------    ----    -----    -----\n -666     -3800      88      172     1028`를 반환해야 합니다.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True), "   32         1      45      123      988\\n- 698    - 3801    + 43    +  49    +  40\\n-----    ------    ----    -----    -----\\n -666     -3800      88      172     1028")`);
  }
})
```

# --seed--

## --seed-contents--

```py
def arithmetic_arranger(problems, show_answers=False):

    return problems

print(f'\n{arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])}')
```

# --solutions--

```py
def arithmetic_arranger(problems, result=False):

	lin1 = ""
	lin2 = ""
	lin3 = ""
	lin4 = ""

	if len(problems) > 5:
		return 'Error: Too many problems.'
	for problem in problems:
		[num1, sym, num2] = problem.split()
		sign = ['+', '-']
		if sym not in sign:
			return ("Error: Operator must be '+' or '-'.")
		if len(num1) > 4 or len(num2) > 4:
			return ("Error: Numbers cannot be more than four digits.")
		if not num1.isnumeric() or not num2.isnumeric():
			return ("Error: Numbers must only contain digits.")

		lin1 += "  " + num1 + "    " if len(num1) >= len(
		    num2) else " " * (len(num2) + 2 - len(num1)) + num1 + "    "
		lin2 += sym + " " + num2 + "    " if len(num2) >= len(
		    num1) else sym + " " * (len(num1) - len(num2) + 1) + num2 + "    "
		nmax = (len(num1) + 2) if len(num1) >= len(num2) else (len(num2) + 2)
		lin3 += "-" * nmax + "    "
		op = int(num1) + int(num2) if sym == "+" else int(num1) - int(num2)
		lin4 += (" " * (nmax - len(str(op)))) + str(op) + "    "

	arranged_problems = lin1.rstrip() + "\n" + lin2.rstrip(
	) + "\n" + lin3.rstrip()
	if result:
		arranged_problems += "\n" + lin4.rstrip()

	return arranged_problems
```
