---
id: 66ebd4ae2812430bb883c787
title: Construye un Hub de eventos
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

En este laboratorio, utilizarás elementos semánticos de HTML para crear la estructura de una página web. Añadirás contenido e imágenes para que parezca un verdadero centro de eventos.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debe tener un elemento `header`.
1. Dentro del elemento `header`, debe tener un elemento `h1` que contenga el texto `Event Hub`, y un elemento `nav`.
1. Dentro del elemento `nav`, debe tener una lista desordenada con dos elementos que contengan enlaces a diferentes secciones de la página. El primer elemento debe tener el texto `Upcoming Events`, y el segundo elemento debe tener el texto `Past Events`.
1. Cada enlace debe estar representado por un elemento `a` con un atributo `href` que enlace a la sección correspondiente de la página, `#upcoming-events` y `#past-events` respectivamente.
1. Debe tener un elemento `main` que contenga las diferentes secciones de la página.
1. Dentro del elemento `main`, debe tener dos elementos `section`.
1. El primer elemento `section` debe tener un atributo `id` con el valor `upcoming-events`.
1. Dentro de la sección `#upcoming-events`, debe tener:
   
   - Un elemento `h2` con el texto `Upcoming Events`.
   - Dos elementos `article`. Cada artículo debe representar un evento, y debe tener:
      - Un elemento `h3` para el título del evento.
      - Un elemento `p` para la descripción del evento. Puedes añadir una fecha al final si lo deseas.

1. El segundo elemento `section` debe tener un atributo `id` con el valor `past-events`.
1. Dentro de la sección `#past-events`, debe tener:

   - Un elemento `h2` con el texto `Past Events`.
   - Dos elementos `article`. Cada elemento de artículo debe representar un evento pasado, y debe tener:
     - Un elemento `h3` para el título del evento,
     - Un elemento `p` para la descripción del evento. Puedes añadir una fecha al final si lo deseas.
     - Un elemento de imagen con el atributo `src` apuntando a un archivo de imagen y el atributo `alt` con una descripción de la imagen.

**Nota:** Puede usar cualquier texto para las descripciones y fechas de los eventos. Puede usar las siguientes URL de imagen para las imágenes si lo desea:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

Debe tener un elemento `header`.

```js
assert.isNotNull(document.querySelector("header"));
```

Su elemento `header` debe venir después de la etiqueta de apertura `body`.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

Dentro del elemento `header`, debe tener un elemento `h1` que contenga el texto `Event Hub`.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

Dentro del elemento `header`, después del elemento `h1`, debe tener un elemento `nav`.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

Su elemento `nav` debe contener una lista desordenada con dos elementos.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

El primer elemento de la lista desordenada debe ser un enlace.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

El segundo elemento de la lista desordenada debe ser un enlace.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

El texto del primer elemento en la lista desordenada debe ser `Upcoming Events`.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

El primer elemento de la lista desordenada debe tener el `href` establecido a `#upcoming-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

El texto del segundo elemento en la lista desordenada debe ser `Past Events`.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

El segundo elemento de la lista desordenada debe tener el `href` establecido a `#past-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

Debe tener un elemento `main` después del cierre de la etiqueta `header`.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

Dentro del elemento `main`, debe tener dos elementos `section`.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

Su primer elemento `section` debe tener un atributo `id` con el valor `upcoming-events`.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

Su segundo elemento `section` debe tener un atributo `id` con el valor `past-events`.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

Dentro de la sección `#upcoming-events`, debe tener un elemento `h2` con el texto `Upcoming Events`.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

Dentro de la sección `#upcoming-events`, debe tener dos elementos `article` debajo del elemento `h2`.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Ambos elementos `article` dentro de la sección `#upcoming-events` deben tener un elemento `h3` para el título del evento.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Ambos elementos `article` dentro de la sección `#upcoming-events` deben tener un elemento de párrafo para la descripción del evento.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Dentro de la sección `#past-events`, debe tener un elemento `h2` con el texto `Past Events`.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

Dentro de la sección `#past-events`, debe tener dos elementos `article` debajo del elemento `h2`.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Ambos elementos `article` dentro de la sección `#past-events` deben tener un elemento `h3` para el título del evento.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Ambos elementos `article` dentro de la sección `#past-events` deben tener un elemento de párrafo para la descripción del evento.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Ambos elementos `article` dentro de la sección `#past-events` deben tener un elemento de imagen.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

Ambos elementos de imagen dentro de la sección `#past-events` deben tener el atributo `src` apuntando a un archivo de imagen.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

Ambos elementos de imagen dentro de la sección `#past-events` deben tener el atributo `alt` con una descripción de la imagen.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

Cada elemento `h3` debe tener el título del evento.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

Cada elemento `p` debe tener la descripción del evento.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```

