---
id: 6998b99a8e195cd0c1f211df
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه ورشة العمل، ستبني متتبعًا بسيطًا للحياة البرية باستخدام كائنات JavaScript.

إليك مثالًا على كائن:

```js
const animal = {}; // empty object
```

هذا ينشئ كائنًا فارغًا.

الآن عرّف متغيرًا باسم `tiger` وعيّن له كائنًا فارغًا.

# --hints--

يجب أن تعرّف متغيرًا باسم `tiger`.

```js
assert.exists(tiger);
```

يجب تعيين كائن فارغ للمتغير `tiger`.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
