---
id: 6999b07c34c2cde21ee6e1d8
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

في هذه الخطوة، ستُنشئ دالة تتحقق مما إذا كان الكائن يحتوي على خاصية معينة.

إليك مثالًا على استخدام `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

في هذا المثال، تستخدم الدالة `hasColor` طريقة `hasOwnProperty` للتحقق مما إذا كانت الخاصية `color` موجودة في الكائن.

الآن أنشئ دالة تُسمى `hasHabitat`. يجب أن تأخذ الدالة معلمة واحدة باسم `animal`.

استخدم طريقة `hasOwnProperty` لإرجاع `true` إذا كان كائن `animal` يحتوي على خاصية تسمى `"habitat"` و`false` خلاف ذلك.

بعد إنشاء الدالة، استخدم `console.log` لاستدعاء `hasHabitat(tiger)` و`hasHabitat(elephant)` حتى ترى النتيجتين في وحدة التحكم.

# --hints--

يجب أن تنشئ دالة باسم `hasHabitat`.

```js
assert.isFunction(hasHabitat);
```

يجب أن تحتوي الدالة `hasHabitat` على معلمة واحدة فقط، وهي `animal`.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تستخدم `hasHabitat` طريقة `hasOwnProperty` للتحقق من الخاصية `habitat`.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

يجب أن تسجل `hasHabitat(tiger)` في وحدة التحكم.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

يجب أن تسجل `hasHabitat(elephant)` في وحدة التحكم.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

يجب أن تُرجع استدعاء `hasHabitat(tiger)` القيمة `true` بعد إضافة موطن.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

يجب أن تُرجع استدعاء `hasHabitat(elephant)` القيمة `false` إذا لم يكن هناك موطن.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

يجب أن تستخدم `hasHabitat` معلمات الدالة وتعمل مع أي كائن.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
