---
id: 6999faf886634b72be7a819b
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

في هذا المتتبع للحياة البرية، حدّث كائنًا ليشمل خصائص جديدة.

إليك مثالًا على إنشاء كائن مع خاصية:

```js
const animal = {
  name: "Lion" // property: value
};
```

الآن حدّث كائن `tiger` ليشمل خاصية تسمى `species`.

عيّن قيمتها إلى `"Tiger"`.

# --hints--

يجب أن يكون `tiger` كائنًا.

```js
assert.isObject(tiger);
```

يجب أن يحتوي `tiger` على خاصية `species`.

```js
assert.property(tiger, "species");
```

يجب أن تكون قيمة `tiger.species` مساوية لـ `"Tiger"`.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
