---
id: 6999a0c1698b93d5d41156ed
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

أضف خاصية أخرى إلى كائن `tiger` باسم `isEndangered`.

عيّن الخاصية `isEndangered` إلى القيمة `true`.

# --hints--

يجب أن يحتوي `tiger` على خاصية `isEndangered`.

```js
assert.property(tiger, "isEndangered");
```

يجب أن تكون قيمة `tiger.isEndangered` مساوية لـ `true`.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
