---
id: 6999ac171928136168654774
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

في هذه الخطوة، ستُنشئ دالة للوصول إلى خاصية `species` لكائن.

إليك مثال صغير للوصول إلى خاصية من كائن:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

في هذا المثال، تأخذ الدالة كائنًا كمعلمة وتُرجع خاصية `species` الخاصة به.

الآن أنشئ دالة باسم `getSpecies`.

يجب أن تأخذ الدالة معلمة واحدة باسم `animal`.

أرجع خاصية `species` من `animal` باستخدام ترميز النقطة.

بعد إنشاء الدالة، استخدم `console.log` لاستدعاء `getSpecies(tiger)` لترى النتيجة في وحدة التحكم.

# --hints--

يجب أن تنشئ دالة باسم `getSpecies`.

```js
assert.isFunction(getSpecies);
```

يجب أن تحتوي دالة `getSpecies` على معلمة واحدة فقط، وهي `animal`.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تسجل `getSpecies(tiger)` في وحدة التحكم.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

يجب أن تصل دالة `getSpecies` إلى خاصية species باستخدام ترميز النقطة.

```js
assert.match(code, /animal\.species/);
```

يجب أن تُرجع استدعاء `getSpecies(tiger)` القيمة `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

يجب أن تستخدم دالة `getSpecies` معلمة الدالة وتعمل مع أي كائن.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
