---
id: 69999ff7f7862fdfb885ea5d
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

حدّث كائن `tiger` بإضافة خاصية جديدة تسمى `age`.

عيّن الخاصية `age` إلى القيمة `5`.

# --hints--

يجب أن يحتوي `tiger` على خاصية `age`.

```js
assert.property(tiger, "age");
```

يجب أن تكون قيمة `tiger.age` مساوية لـ `5`.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
