---
id: 6999b0d79ea938edcb032237
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

في هذه الخطوة النهائية من ورشة عمل متعقب الحياة البرية، ستُنشئ دالة تصل إلى خاصية في كائن باستخدام ترميز الأقواس.

إليك مثالًا على الوصول إلى خاصية باستخدام ترميز الأقواس:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

في هذا المثال، يسمح ترميز الأقواس بالوصول إلى خاصية باستخدام متغير.

الآن أنشئ دالة تسمى `getProperty`.

يجب أن تأخذ الدالة معلمتين: `animal` و `propertyName`.

أرجع قيمة الخاصية باستخدام ترميز الأقواس.

بعد إنشاء الدالة، استخدم `console.log` لاستدعاء `getProperty(tiger, "species")` و `getProperty(elephant, "age")` حتى ترى القيم المرجعة في وحدة التحكم.

وبذلك، تكتمل ورشة عمل متعقب الحياة البرية!

# --hints--

يجب أن تنشئ دالة باسم `getProperty`.

```js
assert.isFunction(getProperty);
```

يجب أن تحتوي دالة `getProperty` على معلمتين: `animal` و `propertyName`.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تستخدم `getProperty` ترميز الأقواس للوصول إلى الخاصية.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

يجب أن تسجل `getProperty(tiger, "species")` في وحدة التحكم.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

يجب أن تسجل `getProperty(elephant, "age")` في وحدة التحكم.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

يجب أن تُرجع استدعاء `getProperty(tiger, "species")` القيمة `"Tiger"`.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

يجب أن تُرجع استدعاء `getProperty(elephant, "age")` القيمة `12`.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

يجب أن تستخدم `getProperty` معلمات الدالة وتعمل مع أي كائن.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
