---
id: 6999adc72c02812a28c01e31
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

في هذه الخطوة، ستنشئ دالة تُحدّث خاصية في كائن.

إليك مثالًا على تحديث خاصية داخل دالة:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

في هذا المثال، تُحدّث الخاصية `age` إلى قيمة جديدة.

الآن أنشئ دالة باسم `updateAge`. يجب أن تأخذ الدالة معلمتين: `animal` و `newAge`.

داخل الدالة، حدّث الخاصية `age` في كائن `animal` إلى `newAge`. أرجع كائن `animal` المحدث.

بعد إنشاء الدالة، استخدم `console.log` لاستدعاء `updateAge(elephant, 12)` حتى ترى كائن `elephant` المحدث في وحدة التحكم.

# --hints--

يجب أن تنشئ دالة باسم `updateAge`.

```js
assert.isFunction(updateAge);
```

يجب أن تحتوي دالة `updateAge` على معلمتين: `animal` و `newAge`.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تستخدم `updateAge` ترميز النقطة لتحديث الخاصية `age`.

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

يجب أن تُرجع دالة `updateAge` كائن `animal` المحدث.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

يجب أن تسجل `updateAge(elephant, 12)` في وحدة التحكم.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

يجب أن يؤدي استدعاء `updateAge(elephant, 12)` إلى تحديث خاصية العمر.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

يجب أن تستخدم `updateAge` معلمات الدالة وتعمل مع أي كائن.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
