---
id: 6999b0198d18d402c36cc5fe
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

في هذه الخطوة، ستُنشئ دالة تزيل خاصية من كائن.

إليك مثال على إزالة خاصية باستخدام `delete`:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

في هذا المثال، تُزال الخاصية `canFly` من الكائن `bird`.

الآن أنشئ دالة تُسمى `removeEndangeredStatus`. يجب أن تأخذ الدالة معلمة واحدة باسم `animal`.

داخل الدالة، أزل الخاصية `isEndangered` من كائن `animal` باستخدام الكلمة المفتاحية `delete`. أرجع الكائن `animal` المحدث.

بعد إنشاء الدالة، استخدم `console.log` لاستدعاء `removeEndangeredStatus(tiger)` حتى ترى الكائن المحدث في وحدة التحكم.

# --hints--

يجب أن تنشئ دالة باسم `removeEndangeredStatus`.

```js
assert.isFunction(removeEndangeredStatus);
```

يجب أن تحتوي الدالة `removeEndangeredStatus` على معلمة واحدة فقط، وهي `animal`.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تُرجع الدالة `removeEndangeredStatus` الكائن `animal` المحدث.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

يجب أن تستخدم `removeEndangeredStatus` الكلمة المفتاحية `delete` لإزالة الخاصية `isEndangered`.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

يجب أن تسجل `removeEndangeredStatus(tiger)` في وحدة التحكم.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

يجب أن يؤدي استدعاء `removeEndangeredStatus(tiger)` إلى إزالة الخاصية `isEndangered`.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

يجب أن تستخدم `removeEndangeredStatus` معلمات الدالة وتعمل مع أي كائن.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
