---
id: 699a00de9f564bb0effaa14d
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

في هذه الخطوة، ستُنشئ دالة للوصول إلى خاصية `age` لكائن.

إليك مثال:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
};

console.log(getAnimalAge(cat)); // 3
```

في هذا المثال، تأخذ الدالة كائنًا كمعلمة وتُرجع خاصية `age` الخاصة به.

الآن أنشئ دالة باسم `getAge`.

يجب أن تأخذ الدالة معلمة واحدة باسم `animal`.

أرجع خاصية `age` من `animal` باستخدام ترميز النقطة.

بعد إنشاء الدالة، استخدم `console.log` لاستدعاء `getAge(tiger)` لترى النتيجة في وحدة التحكم.

# --hints--

يجب أن تنشئ دالة باسم `getAge`.

```js
assert.isFunction(getAge);
```

يجب أن تحتوي دالة `getAge` على معلمة واحدة فقط، وهي `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تصل دالة `getAge` إلى خاصية `age` باستخدام ترميز النقطة.

```js
assert.match(code, /animal\.age/);
```

يجب أن تسجل `getAge(tiger)` في وحدة التحكم.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

يجب أن تُرجع استدعاء `getAge(tiger)` القيمة `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

يجب أن تستخدم دالة `getAge` المعلمة وتعمل مع أي كائن.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
