---
id: 69ba8da6b296391d913096a0
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

أنشئ الآن كائنًا ثانيًا اسمه `elephant`.

أضف الخصائص التالية:

- خاصية `species` بقيمة `"Elephant"`
- خاصية `age` بقيمة `10`
- خاصية `isEndangered` بقيمة `true`

# --hints--

يجب أن تنشئ متغيرًا اسمه `elephant`.

```js
assert.exists(elephant);
```

يجب أن يحتوي `elephant` على خاصية `species` تساوي `"Elephant"`.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

يجب أن يحتوي `elephant` على خاصية `age` تساوي `10`.

```js
assert.propertyVal(elephant, "age", 10);
```

يجب أن يحتوي `elephant` على خاصية `isEndangered` تساوي `true`.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
