---
id: 686386e76197458b7dfd52ee
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `ul` 裡，創建兩個列表項目。在第一個列表項目裡，創建一個 `span` 元素，並將 HTML 物體 `&#10003;` 作為它的文字。緊接著該 `span` 但仍在列表項目內，添加文字 `Ad-supported streaming`。

在第二個列表條款內，創建一個 `span` 元素並將物體 `&#10003;` 添加為其文字。在該 `span` 之後，但仍在列表條款內，添加文字 `Curated playlists`。

# --hints--

你的 `ul` 元素內應該有兩個列表項目。

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

你的第一個 `li` 元素應該包含一個帶有 HTML 字元物體 `&#10003;` 的 `span` 元素。

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

你的第一個 `li` 元素應該在 `span` 元素之後有文字 `Ad-supported streaming`。

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

你的第二個 `li` 元素應該包含一個帶有 HTML 字元物體 `&#10003;` 的 `span` 元素。

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

你的第二個 `li` 元素應該在 `span` 元素之後有文字 `Curated playlists`。

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
