---
id: 686383abac6c625a84458ea9
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次工作坊中，你將透過使用 Tailwind CSS 的網格工具類別來建置一個定價組件，進一步了解 Tailwind CSS。

基本的 HTML 範本已為你提供。請先建立一個 `h1` 元素，並使用文字 `Choose your listening plan`。

# --hints--

你應該創建一個 `h1` 元素。

```js
assert.exists(document.querySelector("h1"));
```

你的 `h1` 元素應該有文字 `Choose your listening plan`。

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
