---
id: 686386e76197458b7dfd52ef
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

現在你將創建一個 CTA 按鈕。行動呼籲（CTA）是一個按鈕或連結，旨在促使用者採取特定行動，在此情況下，是開始收聽。

在最內層的 `div` 外，創建一個錨點元素，其 `href` 為 `#` 且文字為 `Start listening`。

這就是第一個定價卡的 HTML。

# --hints--

你應該創建一個錨點元素。

```js
assert.exists(document.querySelector("a"))
```

你的錨點元素應該設定 `href` 為 `#`。

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

你的錨點元素應該有文字 `Start listening`。

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
