---
id: 686386e76197458b7dfd52ed
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在你剛剛添加的 `p` 元素下，創建另一個帶有文字 `Start exploring millions of songs with basic features and ads.` 的 `p` 元素。在那之下，創建一個空的 `ul` 元素。

# --hints--

你應該創建另一個帶有文字 `Start exploring millions of songs with basic features and ads.` 的 `p` 元素。

```js
const paragraphs = document.querySelectorAll("p")

const targetP = Array.from(paragraphs).find(p =>
  p.textContent === "Start exploring millions of songs with basic features and ads."
)
assert.exists(targetP)
```

你應該在 `p` 元素下創建一個空的 `ul` 元素。

```js
const paragraphs = document.querySelectorAll("p")

const targetP = Array.from(paragraphs).find(p =>
  p.textContent === "Start exploring millions of songs with basic features and ads."
)

const ul = targetP.nextElementSibling
assert.equal(ul?.tagName, "UL")
assert.isEmpty(ul?.innerHTML.trim())
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
