---
id: 686386e76197458b7dfd52eb
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `h1` 下創建一個 `div` 元素。在該 `div` 內，創建另一個 `div` 元素，並在其中巢狀另一個 `div` 元素。

並且在最內層的 `div` 中，創建一個帶有文字 `Listener` 的 `h2` 元素。這將表述第一個定價方案。

# --hints--

你應該創建一個 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

你應該在外層的 `div` 裡面創建另一個 `div` 元素。

```js
assert.exists(document.querySelector("div > div"));
```

你應該在內部的 `div` 裡面創建另一個 `div` 元素。

```js
assert.exists(document.querySelector("div > div > div"));
```

你應該在最內層的 `div` 元素內創建一個 `h2`，其內含文字為 `Listener`。

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
