---
id: 686386e76197458b7dfd52ec
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在 `h2` 下，創建一個帶有文字 `$0` 的 `p` 元素。然後，在該 `p` 元素內且文字之後，創建一個帶有文字 `/month` 的 `span` 元素。

# --hints--

你應該創建一個帶有文字 `$0` 的 `p` 元素。

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

你應該在 `p` 元素內創建一個包含文字 `/month` 的 `span` 元素。

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
