---
id: 6864df7ec0fc508ac3828ab5
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

從一般的 CSS 你學到，指定元素的顯示為 `grid` 不會有任何作用，直到你指定網格的列和行。Tailwind 也是這樣運作的。

在 Tailwind 中，要指定網格列和網格行，你會使用它們各自的工具類別，如下所示：

```html
<div class="grid grid-cols-[num] grid-rows-[num]">
  <div>Grid Item</div>
  <div>Grid Item</div>
  <div>Grid Item</div>
</div>
```

為外覆 `div` 指定工具類別 `grid-cols-1` 和 `md:grid-cols-3`。

`grid-cols-1` 類別將三個項目排列在單一行中，垂直堆疊它們。在中等及更大螢幕上，`md:grid-cols-3` 將版面切換為三行，並排顯示這些項目。

# --hints--

你的 `div` 元素應該有 `grid-cols-1` 類別。

```js
const firstDiv = document.querySelector("div")

assert.isTrue(firstDiv.classList.contains("grid-cols-1"))
```

你的 `div` 元素應該有 `md:grid-cols-3` 類別。

```js
const firstDiv = document.querySelector("div")

assert.isTrue(firstDiv.classList.contains("md:grid-cols-3"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
--fcc-editable-region--
      <div class="grid ">
--fcc-editable-region--
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
