---
id: 68667227c6b33aaf41d4e086
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

要定位卡片中的項目，你需要再次使用 grid 類別。將工具類別 `grid`、`grid-rows-[1fr_auto]` 和 `gap-6` 指定到外包的 `div`。

此外，使用類別 `rounded-xl` 移除銳利邊緣。

# --hints--

你的高級方案外覆 `div` 元素應該具有類別 `grid`。

```js
const premiumPlanWrapper = document.querySelectorAll("div")[3]

assert.isTrue(premiumPlanWrapper.classList.contains("grid"))
```

你的高級方案的外覆 `div` 元素應該有類別 `grid-rows-[1fr_auto]`。

```js
const premiumPlanWrapper = document.querySelectorAll("div")[3]

assert.isTrue(premiumPlanWrapper.classList.contains("grid-rows-[1fr_auto]"))
```

你的高級方案的外包 `div` 元素應該有類別 `gap-6`。

```js
const premiumPlanWrapper = document.querySelectorAll("div")[3]

assert.isTrue(premiumPlanWrapper.classList.contains("gap-6"))
```

你的高級方案外覆 `div` 元素應該有 `rounded-xl` 類別。

```js
const premiumPlanWrapper = document.querySelectorAll("div")[3]

assert.isTrue(premiumPlanWrapper.classList.contains("rounded-xl"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
      <div class="grid grid-cols-1 md:grid-cols-3 max-w-6xl mx-auto gap-8 mt-16">
        <div class="bg-gray-100 ring-1 ring-gray-300 grid grid-rows-[1fr_auto] rounded-xl p-8 gap-6">
          <div class="grid grid-rows-[auto_auto_auto_1fr] gap-y-2">
            <h2 class="text-lg font-semibold text-indigo-600">Listener</h2>
            <p class="text-4xl font-bold text-gray-900">$0<span class="text-base font-medium text-gray-500">/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul class="mt-6 space-y-2 text-sm text-gray-700">
              <li>
                <span aria-hidden class="text-green-700 mr-2">&#10003;</span>Ad-supported streaming
              </li>
              <li>
                <span aria-hidden class="text-green-700 mr-2">&#10003;</span>Curated playlists
              </li>
            </ul>
          </div>
          <a
            href="#"
            class="block rounded-md bg-indigo-100 px-4 py-2 text-center font-semibold text-indigo-700 hover:bg-indigo-200"
          >
            Start listening
          </a>
        </div>
--fcc-editable-region--
        <div class="bg-gray-950 text-white ring-2 ring-fuchsia-500 p-8">
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
--fcc-editable-region--
        <div>
          <div>
            <div>
              <h2>Family</h2>
              <p>$14.99<span>/month</span></p>
              <p>
                Enjoy all of the features with a plan for up to 6 family members.
              </p>
              <ul>
                <li>
                  <span>&#10003;</span>All Premium features
                </li>
                <li>
                  <span>&#10003;</span>Up to 6 accounts
                </li>
                <li>
                  <span>&#10003;</span>Individual playlists &
                  libraries
                </li>
                <li>
                  <span>&#10003;</span>Family Mix playlists
                </li>
              </ul>
            </div>
            <a href="#">Start Family Plan</a>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
