---
id: 6823d9ac8bdc3853df65a1ff
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Kuna vipengele sita vya kichwa cha sehemu katika HTML: `h1` hadi `h6`. Vinatumika kuonyesha umuhimu wa sehemu kwenye ukurasa wako wa mtandao, ambapo `h1` ni muhimu zaidi na `h6` ni muhimu kidogo zaidi.

Chini ya kipengele chako cha `p`, ongeza kichwa cha sehemu cha `h3` chenye maandishi:

```md
Introduction to HTML
```

# --hints--

Kipengele chako cha `h3` kinapaswa kuwa na lebo ya ufunguzi ya `<h3>`.

```js
assert.exists(document.querySelector("h3"));
```

Kipengele chako cha `h3` kinapaswa kuwa na lebo ya kufunga ya `</h3>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

Maandishi ya kipengele cha `h3` yanapaswa kuwa `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
