---
id: 6823d6244511f252c8300eed
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Unapohitaji kuongeza aya kwenye ukurasa wa mtandao, unaweza kutumia kipengele cha `p` kama ifuatavyo:

```html
<p>This is a paragraph element.</p>
```

Badilisha `Learn the skills to become a full-stack developer` kuwa kipengele cha aya.

# --hints--

Kipengele chako cha `p` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

Kipengele chako cha `p` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina sintaksia hii: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

Maandishi ya kipengele chako cha `p` yanapaswa kuwa `Learn the skills to become a full-stack developer`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

Kipengele chako cha `p` kinapaswa kuwa chini ya kipengele chako cha `h2`.

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
