---
id: 6823ac607bfdbc46331b2559
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML inamaanisha HyperText Markup Language. Ni msimbo unaoelezea muundo na maudhui ya ukurasa wa mtandao. Huu ni mhariri wako wa msimbo, ambapo utaandika HTML.

Tafuta mstari wa 1 katika mhariri na andika maandishi haya:

```md
Welcome to freeCodeCamp
```

Ukipomaliza, bofya kitufe cha "angalia msimbo wako" kuona kama ni sahihi.

# --hints--

Unapaswa kuwa na maandishi `Welcome to freeCodeCamp` katika mhariri wako. Hakikisha umeangalia tahajia kwa makini.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
