---
id: 6823e637c1c0ed56f781b4fc
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Mwishowe, JavaScript hufanya ukurasa wako wa mtandao uwe wa mwingiliano — hukuruhusu kusema ukurasa ufanye nini mtu anapobofya kitufe, kuwasilisha fomu, au mambo mengine mengi.

Kwa hatua ya mwisho ya warsha, ongeza kipengele kingine cha `h3` na `p` kwenye ukurasa kinachoelezea JavaScript.

Kwanza, ongeza kipengele cha `h3` chenye maandishi:

```md
Introduction to JavaScript
```

Kisha, chini ya kipengele hicho cha `h3`, ongeza kipengele cha `p` chenye maandishi yafuatayo:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

Unapaswa kuongeza kipengele cha tatu cha `h3` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Maandishi ya kipengele chako cha `h3` yanapaswa kuwa `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

Unapaswa kuongeza kipengele cha nne cha `p` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi ya `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
