---
id: 6823e036ea4b71553558c01b
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Tambua kwamba kila kichwa cha sehemu kinaonekana tofauti kidogo katika muhtasari. Vivyo hivyo na aya.

Ifuatayo, ongeza kipengele kingine cha `p` kinachoonyesha maandishi yafuatayo:

```md
HTML represents the content and structure of a webpage
```

# --hints--

Unapaswa kuwa na kipengele cha pili cha `p`.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi ya `HTML represents the content and structure of a webpage`. Unaweza kunakili maandishi kutoka kwa maelekezo.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
