---
id: 682ba2318000b62f179bdf04
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML inaundwa na vipengele. Kwanza utatumia kipengele cha `h1`:

```html
<h1>Welcome to freeCodeCamp</h1>
```

Huanza na lebo ya kufungua (`<h1>`), huisha na lebo ya kufunga (`</h1>`), na ina maandishi ambayo yataonyeshwa kati ya lebo hizo.

Badilisha maandishi yako ya `Welcome to freeCodeCamp` kuwa kipengele cha `h1` kwa kuongeza lebo ya kufungua mbele yake, na lebo ya kufunga baada yake.

# --hints--

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufungua `<h1>`.

```js
assert.exists(document.querySelector("h1"));
```

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufunga `</h1>`.

```js
assert.match(code, /<\/h1\s*\>/);
```

Kipengele chako cha `h1` kinapaswa kuonekana kama hii: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
