---
id: 682cd206883fc7b25eb539c5
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Tambua kwamba HTML unayoandika katika mhariri inaonekana katika mtazamo wa awali. Katika warsha hii, utaandika HTML kwa ajili ya ukurasa wa mtandao wa mtaala wa sehemu.

Chini ya kipengele chako cha `h1`, andika yafuatayo kwenye mstari tupu:

```md
Full-Stack Curriculum
```

# --hints--

Unapaswa kuwa na maandishi ya `Full-Stack Curriculum` katika mhariri wako. Hakikisha kwa makini tahajia.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
