---
id: 683921c4769fd23dbadec2fe
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Wakati HTML inabainisha muundo na maudhui ya ukurasa wa mtandao, CSS hutumika kuongeza mtindo — vitu kama rangi, fonti, nafasi, na mpangilio.

Chini ya `h3` uliyoongeza tu, ongeza kipengele kingine cha aya chenye maandishi:

```md
CSS is used to style a webpage
```

# --hints--

Unapaswa kuwa na vipengele viwili tu vya `h3`, hakikisha usiondoe vilivyopo, na usiongeze vipya.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Haupaswi kubadilisha maandishi ya kipengele cha pili cha `h3`, inapaswa bado kusoma `Introduction to CSS`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

Unapaswa kuongeza kipengele cha tatu cha `p` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi `CSS is used to style a webpage`. Unaweza kunakili maandishi kutoka kwa maelekezo.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
--fcc-editable-region--

--fcc-editable-region--
```
