---
id: 5f33294a6af5e9188dbdb8f3
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop üben Sie die Grundlagen von CSS (Cascading Style Sheets), indem Sie ein Café-Menü erstellen.

Beginnen Sie damit, etwas Menüinhalt hinzuzufügen. Fügen Sie ein `main`-Element innerhalb des vorhandenen `body`-Elements ein. Es wird später Preisinformationen über Kaffee und Desserts enthalten, die das Café anbietet.

# --hints--

Ihr Code sollte ein öffnendes `<main>`-Tag enthalten.

```js
assert.match(code, /<main>/i);
```

Ihr Code sollte ein schließendes `</main>`-Tag enthalten.

```js
assert.match(code, /<\/main>/i);
```

Sie sollten Ihr `body`-Element nicht verändern. Achten Sie darauf, das schließende Tag nicht versehentlich zu löschen.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Ihr `main`-Tag sollte innerhalb Ihres `body`-Tags liegen.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
