---
id: 5f344f9c805cd193c33d829c
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man `CSS` Eigenschaften und Werte wie diese hinzufügt:

```css
element {
 property: value;
}
```

Zentrieren Sie den Inhalt des `h1` Elements, indem Sie seine `text-align` Eigenschaft auf den Wert `center` setzen.

# --hints--

Sie sollten einen `h1` Selektor in Ihrem `style` Element haben.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasSelector);
```

Ihre `text-align` Eigenschaft sollte den Wert `center` haben.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert.isTrue(hasTextAlign);
```

Ihr `h1` Selektor sollte die `text-align` Eigenschaft auf `center` setzen.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

