---
id: 5f3477ae34c1239cafe128be
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Sie haben jetzt drei Typselektoren mit derselben Formatierung. Sie können dieselbe Gruppe von Stilen auf viele Elemente anwenden, indem Sie eine Liste von Selektoren erstellen. Jeder Selektor wird durch Kommas getrennt, so:

```css
selector1, selector2 {
  property: value;
}
```

Löschen Sie die drei vorhandenen Typselektoren und ersetzen Sie sie durch eine Selektorenliste, die den Text für die `h1`, `h2` und `p` Elemente zentriert.

# --hints--

Sie sollten einen einzigen Typselektor für alle drei Elemente, `h1, h2, p`, verwenden. Achten Sie darauf, diese Reihenfolge einzuhalten.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert.exists(hasSelector);
```

Sie sollten nur einen Selektor in Ihrem `style` Element haben.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert.equal(selectors.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

