---
id: 5f344fad8bf01691e71a30eb
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Bisher hattest Du nur begrenzte Kontrolle über die Darstellung und das Aussehen der Inhalte. Um das zu ändern, füge ein `style`-Element innerhalb des `head`-Elements ein.

# --hints--

Dein Code sollte ein öffnendes `<style>`-Tag haben.

```js
assert.match(code, /<style\s*>/i);
```

Dein Code sollte ein schließendes `</style>`-Tag haben.

```js
assert.match(code, /<\/style\s*>/);
```

Dein `style`-Element sollte innerhalb deines `head`-Elements sein.

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
