---
id: 5f3477cbcb6ba47918c1da92
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Damit das Styling der Seite auf Mobilgeräten ähnlich aussieht wie auf einem Desktop oder Laptop, müssen Sie ein `meta`-Element mit einem speziellen `content`-Attribut hinzufügen.

Sie haben in vorherigen Lektionen das viewport `meta`-Element kennengelernt.

Hier ist ein Beispiel:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

Ihr Code sollte zwei `meta`-Elemente enthalten.

```js
assert.lengthOf(code.match(/<meta.*?\/?>/g),2);
```

Ihr `meta`-Element sollte ein `name`-Attribut mit dem Wert `viewport` haben.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/name=('|")viewport\1/) || meta[1]?.outerHTML?.match(/name=('|")viewport\1/));
```

Ihr `meta`-Element sollte ein `content`-Attribut mit dem Wert `width=device-width, initial-scale=1.0` haben.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
--fcc-editable-region--
    
--fcc-editable-region--
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
