---
id: 5f344fc1520b6719f2e35605
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Es wird zwei Abschnitte im Menü geben, einen für Kaffees und einen für Desserts. Fügen Sie ein `section`-Element innerhalb des `main`-Elements hinzu, damit Sie einen Platz haben, um alle verfügbaren Kaffees unterzubringen.

# --hints--

Sie sollten ein öffnendes `<section>`-Tag haben.

```js
assert.match(code, /<section\s*>/i);
```

Sie sollten ein schließendes `</section>`-Tag haben.

```js
assert.match(code, /<\/section\s*>/i);
```

Sie sollten Ihr bestehendes `main`-Element nicht ändern. Stellen Sie sicher, dass Sie das schließende Tag nicht gelöscht haben.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Ihr `section`-Element sollte sich innerhalb Ihres `main`-Elements befinden.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert.strictEqual(section.parentElement, main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
