---
id: 5f3477cb2e27333b1ab2b955
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Jetzt müssen Sie die `styles.css`-Datei verlinken, damit die Styles wieder angewendet werden. Fügen Sie innerhalb des `head`-Elements ein `link`-Element hinzu. Geben Sie ihm ein `rel`-Attribut mit dem Wert `"stylesheet"` und ein `href`-Attribut mit dem Wert `"styles.css"`.

**Hinweis:** Das `link`-Element ist ein Void-Element, das heißt, es hat kein schließendes Tag. Void-Elemente sollten als `<link>` und nicht als `<link></link>` geschrieben werden.

# --hints--

Ihr Code sollte ein `link`-Element enthalten.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

Das `link`-Element ist ein Void-Element, es sollte kein End-Tag `</link>` haben.

```js
assert.notMatch(code, /<\/link>/i);
```

Sie sollten Ihr bestehendes `head`-Element nicht ändern. Stellen Sie sicher, dass Sie das schließende Tag nicht gelöscht haben.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

Sie sollten ein `link`-Element haben.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
