---
id: 5f344fad8bf01691e71a30eb
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Bis jetzt hatten Sie nur begrenzte Kontrolle über die Darstellung und das Aussehen Ihres Inhalts. Um das zu ändern, fügen Sie ein `style`-Element innerhalb des `head`-Elements hinzu.

# --hints--

Ihr Code sollte ein `style`-Element enthalten.

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
```

Ihr Code sollte einen schließenden `</style>`-Tag enthalten.

```js
assert.match(code, /<\/style\s*>/i);
```

Ihr `style`-Element sollte in Ihrem `head`-Element verschachtelt sein.

```js
assert.exists(document.querySelector('head > style:not(.fcc-hide-header)'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
