---
id: 5f3477ae9675db8bb7655b30
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Im vorherigen Schritt haben Sie einen <dfn>Typselektor</dfn> verwendet, um das `h1`-Element zu gestalten. Zentrieren Sie den Inhalt der `h2`- und der `p`-Elemente, indem Sie für jedes einen neuen Typselektor zum bestehenden `style`-Element hinzufügen.

# --hints--

Sie sollten den bestehenden `h1`-Selektor nicht ändern.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasH1);
```

Sie sollten keinen neuen `style`-Tag hinzufügen. Fügen Sie die neuen CSS-Regeln zum bestehenden `style`-Tag hinzu.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

Sie sollten einen neuen `h2`-Selektor hinzufügen.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert.exists(hasH2);
```

Sie sollten einen neuen `p`-Selektor hinzufügen.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert.exists(hasP);
```

Ihr `h1`-Element sollte eine `text-align` von `center` haben.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

Ihr `h2`-Element sollte eine `text-align` von `center` haben.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert.equal(h2TextAlign, 'center');
```

Ihr `p`-Element sollte eine `text-align` von `center` haben.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert.equal(pTextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1 {
        text-align: center;
      }
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

