---
id: 5f356ed69db0a491745e2bb6
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

Um das Styling der Klasse auf das `div`-Element anzuwenden, entfernen Sie das `id`-Attribut und fügen Sie dem öffnenden Tag des `class`-Elements ein `div`-Attribut hinzu. Stellen Sie sicher, dass der Klassenwert auf `menu` gesetzt ist.

# --hints--

Ihr `div` sollte weiterhin gerendert werden. Achten Sie darauf, dass Sie das `<div>`-Tag nicht fehlerhaft verändert haben.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

Ihr `div`-Element sollte die `menu`-Klasse haben.

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

Ihr `div`-Element sollte die `menu`-ID nicht mehr haben.

```js
assert.lengthOf(document.querySelectorAll('div#menu'), 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

