---
id: 5f356ed6199b0cdef1d2be8f
title: Schritt 22
challengeType: 0
dashedName: step-22
---

# --description--

Bisher haben Sie Typ- und ID-Selektoren verwendet, um Elemente zu gestalten. Es ist jedoch üblicher, einen anderen Selektor zu verwenden, um Ihre Elemente zu stylen.

Sie haben in vorherigen Lektionen gelernt, wie man mit `class` Selektoren arbeitet, zum Beispiel so:

```css
.class-name {
  styles
}
```

Ändern Sie den vorhandenen `#menu` Selektor in einen Klassenselektor, indem Sie `#menu` durch eine Klasse namens `.menu` ersetzen.

# --hints--

Sie sollten einen `.menu` Klassenselektor haben.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert.exists(hasMenu);
```

Sie sollten keinen `#menu` Selektor haben.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert.notExists(hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
--fcc-editable-region--
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

