---
id: 5f356ed63c7807a4f1e6d054
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Das Ziel ist jetzt, dass das `div` nicht die gesamte Breite der Seite einnimmt. Die CSS-Eigenschaft `width` ist dafür perfekt geeignet.

Sie können den `id`-Selektor verwenden, um ein bestimmtes Element mit einem `id`-Attribut anzusprechen.

Sie haben bereits in vorherigen Lektionen gelernt, wie man mit dem `id`-Selektor arbeitet, zum Beispiel so:

```css 
#cat {
  width: 250px;
}
```

Verwenden Sie den `#menu`-Selektor, um Ihrem Element eine Breite von `300px` zu geben.

# --hints--

Sie sollten einen `#menu`-Selektor haben.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert.exists(hasDiv);
```

Sie sollten die `width`-Eigenschaft auf `300px` setzen.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert.isTrue(hasWidth);
```

Ihr `div` sollte eine Breite von 300px haben.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert.equal(divWidth, '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```

