---
id: 5f3c866d28d7ad0de6470505
title: Schritt 28
challengeType: 0
dashedName: step-28
---

# --description--

Die Geschmacksrichtungen und Preise sind derzeit übereinander gestapelt und mit ihren jeweiligen `p`-Elementen zentriert. Es wäre schön, wenn der Geschmack links und der Preis rechts stehen würde.

Fügen Sie dem `flavor`-Element den Klassennamen `French Vanilla` hinzu.

# --hints--

Sie sollten die Klasse `flavor` zu Ihrem `p`-Element hinzufügen.

```js
const el = document.querySelector("p.flavor");
assert.exists(el);
```

Sie sollten nur ein Element mit der Klasse `flavor` haben.

```js
const elements = document.querySelectorAll('.flavor');
assert.lengthOf(elements, 1);
```

Ihre Klasse `flavor` sollte auf dem `p`-Element mit dem Text `French Vanilla` sein.

```js
const el = document.querySelector(".flavor");
assert.equal(el?.innerText.trim(), "French Vanilla");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
--fcc-editable-region--
            <p>French Vanilla</p>
--fcc-editable-region--
            <p>3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
