---
id: 5cdafbb0291309899753167f
title: Crea una promesa de JavaScript
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

Una promesa en JavaScript es exactamente lo que parece: la usas para hacer una promesa de hacer algo, usualmente de forma asíncrona. Cuando la tarea termina, o cumples tu promesa o no lo haces. `Promise` es una función constructor, así que necesitas usar la palabra clave `new` para crear una. Toma una función como argumento, con dos parámetros: `resolve` y `reject`. Estos son métodos que se usan para determinar el resultado de la promesa. La sintaxis es así:

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

Crea una nueva promesa llamada `makeServerRequest`. Pásale una función con parámetros `resolve` y `reject` al constructor.

# --hints--

Debes asignar una promesa a una variable declarada con el nombre `makeServerRequest`.

```js
assert(makeServerRequest instanceof Promise);
```

Tu promesa debe recibir una función con `resolve` y `reject` como parámetros.

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
