---
id: 68814221fc4752e17b9696bd
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Quando o usuário clicar na nota, fizer edições e depois clicar fora da nota, você precisará de uma maneira de obter essa versão recém-editada.Você pode realizar isso usando o evento `blur`.Este evento é disparado quando um elemento perde o foco.Aqui está um exemplo:

```js
element.addEventListener("blur", () => {
  // do something here
});
```

Anexe um método `addEventListener` à variável `noteEl`.O primeiro argumento deve ser o evento `"blur"` e o segundo argumento deve ser uma função arrow.Dentro do corpo daquela arrow function, use `const` para criar uma nova variável chamada `newContent` e atribua a ela o valor de `noteEl.innerHTML`. 

`innerHTML` é usado aqui para ajudar a preservar o espaçamento e a formatação da nota. 

# --hints--

Você deve anexar um método `addEventListener` à variável `noteEl`.

```js
assert.match(code, /noteEl\.addEventListener/);
```

Seu método `addEventListener` deve ter um primeiro argumento do evento `"blur"`.

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],/);
```

Dentro do corpo da arrow function, você deve usar `const` para criar uma nova variável chamada `newContent` e atribuir a ela o valor de `noteEl.innerHTML`.

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],\s*\(\)\s*=>\s*{[^}]*const\s+newContent\s*=\s*noteEl\.innerHTML[^}]*}\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--
  
--fcc-editable-region--

window.addEventListener("DOMContentLoaded", () => {
  currentContent = noteEl.textContent;
});
```
