---
id: 688141e05e367de0e3ef7635
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

No momento, `currentContent` está inicializado com o valor de uma string vazia. Mas o resultado desejado seria que `currentContent` armazenasse o valor do texto da nota atual.

Em uma lição anterior, você aprendeu como trabalhar com o evento `DOMContentLoaded` assim:

```js
window.addEventListener("DOMContentLoaded", () => {
  // Do stuff here
});
```

Quando tudo no documento HTML tiver sido carregado e analisado, você poderá acessar o `noteEl` com segurança.

Comece anexando um método `addEventListener` ao objeto `window`. O primeiro argumento para o método `addEventListener` deve ser o evento `"DOMContentLoaded"`. O segundo argumento deve ser uma arrow function. Dentro do corpo da arrow function, atribua o valor de `noteEl.textContent` para `currentContent`.

# --hints--

Você deve anexar um método `addEventListener` ao objeto `window`.

```js
assert.match(code, /window\.addEventListener/);
```

Seu listener de evento deve escutar o evento `"DOMContentLoaded"`.

```js
assert.match(code, /window\.addEventListener\(['"]DOMContentLoaded['"]\,/);
```

Dentro do corpo da arrow function, atribua o valor de `noteEl.textContent` para `currentContent`.

```js
const event = new Event("DOMContentLoaded");
window.dispatchEvent(event);

assert.equal(currentContent, noteEl.textContent);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English.

      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.

      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--

--fcc-editable-region--
```
