---
id: 68813c1b28ec3fd41838daf6
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você praticará o uso dos atributos `aria-live` e `aria-label` construindo um aplicativo de anotações acessível. 

A maior parte do HTML e todo o CSS foram fornecidos para você. Os primeiros passos envolverão adicionar a marcação restante. 

Se você olhar a pré-visualização, há um texto que indica para você clicar no cartão para editá-lo. Mas se você tentar fazer isso, não vai funcionar. Para tornar esse elemento `div` editável, você pode usar o atributo `contenteditable` assim:

```html
<div contenteditable="true"></div>
```

No elemento de abertura `div`, adicione o atributo `contenteditable` e defina seu valor como `"true"`. Agora você deve ser capaz de clicar no elemento e editar o texto.

# --hints--

Sua tag de abertura `div` deve ter o atributo `contenteditable`.

```js
assert.isTrue(document.getElementById("note").hasAttribute("contenteditable"));
```

Seu elemento `div` deve ter o atributo `contenteditable` definido como `"true"`.

```js
assert.equal(document.getElementById("note").getAttribute("contenteditable"), "true");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    --fcc-editable-region--
    <div id="note" class="note">
    --fcc-editable-region--
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```
