---
id: 61695c4aad56f95497c19583
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Пам’ятайте, що елемент `title` надає пошуковим системам додаткову інформацію про сторінку. Він також відображає вміст цього елемента `title` у двох додаткових місцях:

* у рядку заголовка, коли сторінка відкрита  
* на вкладці браузера для сторінки, коли ви наводите на неї курсор. Навіть якщо ця вкладка неактивна, при наведенні на неї відображається текст `title`.

Всередині елемента `head` вкладіть елемент `title` з текстом `Colored Markers`.

# --hints--

Ви повинні мати відкриваючий тег `title`.

```js
assert(code.match(/<title\s*>/i));
```

Ви повинні мати закриваючий тег `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

У вашому проєкті має бути заголовок `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Пам’ятайте, що регістр і правопис мають значення для тексту заголовка.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
