---
id: 61695197ac34f0407e339882
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Як ви бачили в попередніх проєктах, вебсторінки повинні починатися з оголошення `DOCTYPE html`, за яким іде елемент `html`.

Додайте оголошення `DOCTYPE html` на початок документа, а після нього — елемент `html`. Додайте елементу `html` атрибут `lang` зі значенням `en`.

# --hints--

Ваш код має містити референс `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Після референса `DOCTYPE` має бути пробіл.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Ви маєте визначити тип документа як `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Оголошення `DOCTYPE` має бути закрито кінцевим тегом `>` після типу.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Ваш елемент `html` має мати початковий тег з атрибутом `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш елемент `html` має мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/gi));
```

Оголошення `DOCTYPE` має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
