---
id: 61695d1fbc003856628bf561
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Атрибут `charset` визначає кодування символів, яке використовується у документі. `utf-8` (Unicode Transformation Format – 8-bit) — це стандарт кодування символів, який застосовується для електронного обміну інформацією.

Всередині елемента `head` вкладіть елемент `meta` з атрибутом `charset`, встановленим у значення `"utf-8"`.

# --hints--

У вас має бути один елемент `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

Ваш елемент `meta` має бути порожнім елементом, він не повинен мати кінцевого тегу `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Ваш елемент `meta` має містити атрибут `charset`, встановлений у значення `"utf-8"`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Ваш елемент `meta` має бути всередині елемента `head`.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
