---
id: 61695ab9f6ffe951c16d03dd
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Вкладіть елемент `head` всередину елемента `html`. Одразу після елемента `head` додайте елемент `body`.

# --hints--

Ви повинні мати початковий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви повинні мати кінцевий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви повинні мати початковий тег `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви повинні мати кінцевий тег `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Повинен бути лише один початковий тег `head`.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

Повинен бути лише один кінцевий тег `head`.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

Повинен бути лише один початковий тег `body`.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

Повинен бути лише один кінцевий тег `body`.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

Ваш елемент `head` має бути порожнім.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

Ваш елемент `body` має бути порожнім.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

Ваш елемент `body` має бути розміщений після елемента `head`.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

Ваш елемент `head` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Ваш елемент `body` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
