---
id: 616965351e74d4689eb6de30
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Ви можете мати кілька `meta` елементів на вебсторінці. Кожен `meta` елемент додає інформацію про сторінку, яку не можна виразити іншими HTML елементами.

Додайте ще один `meta` елемент всередині `head`. Дайте йому атрибут `name`, встановлений у `"viewport"`, та атрибут `content`, встановлений у `"width=device-width, initial-scale=1.0"`, щоб ваша сторінка виглядала однаково на всіх пристроях.

# --hints--

У вас має бути два `meta` елементи.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

Ваш новий `meta` елемент має бути порожнім елементом, він не повинен мати кінцевого тегу `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```


Ваш новий `meta` елемент має мати атрибут `name`, встановлений у `"viewport"`, та атрибут `content`, встановлений у `"width=device-width, initial-scale=1.0"`.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

Ваші два `meta` елементи мають бути всередині `head` елемента.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
