---
id: 616d3a67ccf800ad94ec89ae
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Тепер ви додасте кілька елементів, які згодом стилізуєте у кольорові маркери.

Спочатку, всередині елемента `body`, додайте елемент `div` і встановіть його атрибут `class` у значення `container`. Переконайтеся, що елемент `div` знаходиться нижче за елемент `h1`.

# --hints--

Ви повинні мати відкриваючий тег `div`.

```js
assert(code.match(/<div.*>/i));
```

Ви повинні мати закриваючий тег `div`.

```js
assert(code.match(/<\/div\s*>/i));
```

Ваш елемент `div` має бути всередині елемента `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Ваш елемент `div` має бути після елемента `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Ваш елемент `div` має мати атрибут `class`, встановлений у `container`.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
