---
id: 616d55bd160a95e22453a081
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Тепер, коли у вас є один маркер, центрований із кольором, настав час додати інші маркери.

У `container` `div` додайте ще два елементи `div` і надайте кожному клас `marker`.

# --hints--

Ваш перший новий елемент `div` повинен мати початковий тег.

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

Ваш перший новий елемент `div` повинен мати кінцевий тег.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Ваш другий новий елемент `div` повинен мати початковий тег.

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

Ваш другий новий елемент `div` повинен мати кінцевий тег.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Ваші нові елементи `div` повинні бути всередині `div` з класом `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

Ваші нові елементи `div` повинні обидва мати клас `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
