---
id: 616d525007b8c5d8b3308b1c
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Ваш маркер виглядатиме краще, якщо він буде по центру сторінки. Легкий спосіб зробити це — використати скорочену властивість `margin`.

У попередньому проєкті ви окремо задавали поле елементів за допомогою властивостей, як-от `margin-top` і `margin-left`. Скорочена властивість `margin` дозволяє легко встановити кілька полів одночасно.

Щоб розмістити маркер по центру сторінки, встановіть його властивість `margin` у значення `auto`. Це встановить `margin-top`, `margin-right`, `margin-bottom` і `margin-left` усі в `auto`.

# --hints--

Ваше CSS-правило `.marker` має містити властивість `margin`, встановлену в `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
