---
id: 617bd6ec666b1da2587e4e37
title: Крок 34
challengeType: 0
dashedName: step-34
---

# --description--

Існує ще три третинні кольори: chartreuse green (зелений + жовтий), azure (синій + бірюзовий) і rose (червоний + маджента).

Щоб створити chartreuse green, оновіть функцію `rgb` у правилі CSS `.one`, щоб червоний був на `127`, а зелений встановіть на максимальне значення.

Для azure оновіть функцію `rgb` у правилі CSS `.two`, щоб зелений був на `127`, а синій — на максимальному значенні.

А для rose, який іноді називають яскраво-рожевим, оновіть функцію `rgb` у правилі CSS `.three`, щоб синій був на `127`, а червоний — на максимальному значенні.

# --hints--

Ваше правило CSS `.one` має містити властивість `background-color`, встановлену в `rgb(127, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

Ваше правило CSS `.two` має містити властивість `background-color`, встановлену в `rgb(0, 127, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

Ваше правило CSS `.three` має містити властивість `background-color`, встановлену в `rgb(255, 0, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
