---
id: 617b92b9de349513466f6156
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Хоча червоний і синій маркери виглядають однаково, зелений став значно світлішим, ніж раніше. Це тому, що ключове слово кольору `green` насправді є темнішим відтінком і знаходиться приблизно посередині між чорним і максимальним значенням для зеленого.

У правилі CSS `.two` встановіть значення зеленого в функції `rgb` на `127`, щоб знизити його інтенсивність.

# --hints--

Ваше правило CSS `.two` має містити властивість `background-color`, встановлену в `rgb(0, 127, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
