---
id: 61a489b8579e87364b2d2cdb
title: Крок 49
challengeType: 0
dashedName: step-49
---

# --description--

Ви вже вивчили кілька способів встановлення однотонних кольорів у CSS, але також можна використати колірний перехід, або <dfn>градієнт</dfn>, для елемента.

Градієнт — це коли один колір переходить в інший. CSS-функція `linear-gradient` дозволяє керувати напрямком переходу вздовж лінії та вибирати кольори.

Варто пам’ятати, що функція `linear-gradient` фактично створює елемент `image`, і зазвичай її використовують разом із властивістю `background`, яка може приймати зображення як значення.

У CSS-правилі `.red` змініть властивість `background-color` на `background`.

# --hints--

Ваше CSS-правило `.red` має містити властивість `background` зі значенням `rgb(255, 0, 0)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
