---
id: 61a49d15bdbb5e57cc6fd280
title: Крок 54
challengeType: 0
dashedName: step-54
---

# --description--

Зупинки кольору дозволяють точно налаштувати, де саме розташовуються кольори вздовж лінії градієнту. Вони є одиницею довжини, як `px` або відсотки, які йдуть після кольору у функції `linear-gradient`.

Наприклад, у цьому градієнті від червоного до чорного перехід від червоного до чорного відбувається в точці 90% вздовж лінії градієнту, тому червоний займає більшу частину доступного простору:

```css
linear-gradient(90deg, red 90%, black);
```

У функції `linear-gradient` додайте зупинку кольору `75%` після першого аргументу кольору червоного. Не додавайте зупинки кольору до інших аргументів кольору.

# --hints--

Ваш CSS-правило `.red` має мати властивість `background`, встановлену в `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
