---
id: 69a5f35669099ed52f8563b1
title: Costruisci un rifornitore intelligente per la dispensa
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

In questo laboratorio costruirai un piccolo programma per la gestione della dispensa usando concetti base di JavaScript come array, oggetti, loop e condizionali.

Simulerai la ricezione di una spedizione di prodotti per la dispensa, deciderai cosa fare con ogni elemento e organizzerai i risultati per lo stoccaggio.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

L’array `rawData` contiene stringhe separate da pipe con il formato `sku|name|qty|expires|zone`, dove `zone` è opzionale.

**Storie utente:**

1. Devi implementare una funzione `parseShipment(rawData)` che prende un array di stringhe e restituisce un array di oggetti con proprietà `{ sku, name, qty, expires, zone }`.

   - I valori `sku` duplicati nella spedizione devono essere ignorati.
   - Quando il segmento `zone` non è fornito, deve essere impostato su `"general"`.
   - Il valore `qty` deve essere convertito in un numero.

2. Devi implementare una funzione `planRestock(pantry, shipment)` che confronta la dispensa attuale con la spedizione in arrivo e restituisce un array di azioni nella forma `{ type, item }`, dove `type` è uno tra `"restock"`, `"discard"` o `"donate"`, e `item` è l’oggetto della spedizione analizzato.

   Il parametro `pantry` è un array di oggetti con la stessa forma di un elemento spedito analizzato (`{ sku, name, qty, expires, zone }`).

- Se un elemento della spedizione ha `qty` pari o inferiore a `0`, il `type` dell'azione deve essere `"discard"`, indipendentemente dal fatto che l’elemento esista nella dispensa.
   - Altrimenti, se lo `sku` dell’elemento spedito esiste già nella dispensa, il `type` dell'azione deve essere `"restock"`.
   - Altrimenti (lo `sku` dell’elemento spedito non esiste nella dispensa), il `type` dell'azione deve essere `"donate"`.

3. Devi implementare una funzione `groupByZone(actions)` che raggruppa le azioni in zone di stoccaggio basate sulla proprietà `zone` di ogni elemento. La funzione deve restituire un oggetto dove ogni chiave è il nome di una zona e il valore è un array di azioni appartenenti a quella zona. Per esempio, se le azioni contengono elementi con zone `"fridge"` e `"pantry"`, il risultato deve essere `{ fridge: [...], pantry: [...] }`.

4. Devi implementare una funzione `clonePantry(pantry)` che restituisce una copia profonda della dispensa in modo che le modifiche di pianificazione non influenzino la lista originale. Una copia profonda significa creare un nuovo array con nuovi oggetti, così modificare la copia non cambia la dispensa originale.

5. Devi usare tutte le funzioni insieme per processare una spedizione e stampare in console l’oggetto finale con le azioni raggruppate.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Devi definire una funzione chiamata `parseShipment` che accetta un array di stringhe chiamato `rawData` come parametro.

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

La tua funzione `parseShipment` dovrebbe convertire le stringhe della spedizione nell’array in oggetti con le proprietà: `sku`, `name`, `qty`, `expires` e `zone`.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

I codici SKU duplicati nella spedizione devono essere ignorati.

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

Quando `zone` non è presente in una stringa della spedizione, devi assegnarla a `general`.

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

Il valore `qty` deve essere convertito in un numero.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

Devi definire una funzione chiamata `planRestock` che accetta due parametri: `pantry` e `shipment`.

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

La tua funzione `planRestock` deve restituire un array di azioni con le proprietà `type` e `item`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

Gli elementi con quantità pari o inferiore a `0` nella spedizione devono creare un’azione di tipo `discard`, indipendentemente dal fatto che l’elemento sia presente nella dispensa.

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

Quando un elemento della spedizione esiste già nella dispensa (stesso `sku`), l’azione deve essere `restock`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

Se un elemento della spedizione non esiste nella dispensa, l’azione deve essere `"donate"`.

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

Devi definire una funzione chiamata `groupByZone` che accetta un parametro chiamato `actions`.

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

La tua funzione `groupByZone` deve restituire le azioni raggruppate per la proprietà `zone` di ogni oggetto.

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

La tua funzione `groupByZone` deve raggruppare correttamente le azioni con il contenuto e il conteggio giusti.

``` js
const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "donate",  item: { sku: "B1", zone: "fridge" } },
  { type: "discard", item: { sku: "C1", zone: "pantry" } },
];

const result = groupByZone(actions);

assert.lengthOf(result.fridge, 2);
assert.lengthOf(result.pantry, 1);
assert.strictEqual(result.fridge[0].item.sku, "A1");
assert.strictEqual(result.pantry[0].type, "discard");
```

Devi definire una funzione chiamata `clonePantry` che accetta un parametro chiamato `pantry`.

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` deve restituire un nuovo array invece dell’array originale della dispensa.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

Gli oggetti dentro la dispensa clonata devono essere anch’essi nuovi oggetti (copia profonda).

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

Le funzioni devono lavorare insieme per processare una spedizione e raggruppare le azioni risultanti.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

Devi stampare in console le azioni risultanti raggruppate per zone.

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js
const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01",
  "B21|Bananas|10|2027-01-01",
  "C32|Eggs|3|2027-01-01|fridge",
  "C32|Eggs|3|2027-01-01",
  "D43|Pineapples|0|2027-01-01",
  "E54|Peppers|-1|2027-01-01|fridge"
];
```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
