---
id: 63b60aaaa65f8922bfce6b7e
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Erstellen Sie in Ihrem `span`-Element ein `label`-Element für ein `entry-dropdown` und geben Sie ihm den Text `Add food or exercise:`. Erstellen Sie dann ein `select`-Element mit dem `id`, das auf `entry-dropdown` gesetzt ist, und einem `name`, das auf `options` gesetzt ist. Darunter fügen Sie ein `button`-Element mit dem `id`, das auf `add-entry` gesetzt ist, und dem Text `Add Entry` hinzu.

Geben Sie Ihrem `button`-Element ein `type`-Attribut mit dem Wert `button`, um eine automatische Formularübermittlung zu verhindern.

# --hints--

Sie sollten ein `label`-Element zu Ihrem `span`-Element hinzufügen.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

Ihr neues `label`-Element sollte ein `for`-Attribut mit dem Wert `entry-dropdown` haben.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

Ihr neues `label`-Element sollte den Text `Add food or exercise:` enthalten.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

Sie sollten ein `select`-Element zu Ihrem `span`-Element hinzufügen.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

Ihr `select`-Element sollte nach Ihrem `label`-Element kommen.

```js
assert.equal(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName, 'LABEL');
```

Ihr neues `select`-Element sollte ein `id`-Attribut mit dem Wert `entry-dropdown` haben.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

Ihr neues `select`-Element sollte ein `name`-Attribut mit dem Wert `options` haben.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

Ihr `select`-Element sollte leer sein und einen schließenden `</select>`-Tag haben.

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

Sie sollten ein `button`-Element zu Ihrem `span`-Element hinzufügen.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

Ihr `button`-Element sollte nach Ihrem `select`-Element kommen.

```js
assert.equal(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName, 'SELECT');
```

Ihr neues `button`-Element sollte ein `id`-Attribut mit dem Wert `add-entry` haben.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

Ihr neues `button`-Element sollte den Text `Add Entry` enthalten.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
