---
id: 5dc17dc8f86c76b9248c6eb5
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Colocar comentários permite que você deixe mensagens sem afetar a exibição do navegador. Também permite que você deixe o código inativo. Um comentário em HTML começa com `<!--`, contém qualquer número de linhas de texto e termina com `-->`. 

Aqui está um exemplo de comentário dizendo `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Adicione um comentário acima do elemento `p` com este texto:

`TODO: Add link to cat photos`

# --hints--

O comentário deve terminar com `<!--`. Um ou mais caracteres que definem o início de um comentário estão ausentes.

```js
assert(code.match(/<!--/));
```

Seu comentário deve terminar com `-->`. Está faltando um ou mais dos caracteres que definem o fim de um comentário.

```js
assert(code.match(/-->/));
```

O código não deve ter caracteres extras de abertura/fechamento de comentários. Você tem um `<!--` ou um `-->` a mais sendo exibido no navegador.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

O comentário deve estar acima do elemento `p`. Eles estão na ordem errada.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

O comentário deve conter o texto `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```

