---
id: 5dc174fcf86c76b9248c6eb2
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Elementos HTML têm tags de abertura como `<h1>` e tags de fechamento como `</h1>`.O texto que um elemento exibirá deverá estar entre as tags de abertura e de fechamento.

Altere o texto do elemento `h1` abaixo do `Hello World` para `CatPhotoApp` e veja
as mudanças na pré-visualização do navegador.

Quando estiver pronto, pressione o botão "Verifique seu código" para ver se está correto.

# --hints--

O texto `CatPhotoApp` deve estar presente no código.Não se esqueça de verificar a ortografia.

```js
assert(code.match(/catphotoapp/i));
```

O elemento `h1` deve ter uma tag de abertura.As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Suas tags `h1` devem estar em letras minúsculas. Por convenção, todas as tags HTML são escritas em letras minúsculas.

```js
assert.notMatch(code, /<\/?H1>/);
```

O elemento `h1` deve ter uma tag de fechamento.As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/h1\>/));
```

Você tem mais de um elemento `h1`.Remova o elemento `h1` a mais.

```js
assert(document.querySelectorAll('h1').length === 1);
```

O elemento `h1` deve conter o texto `CatPhotoApp`.Você omitiu o texto, digitou o texto incorretamente ou ele não está entre as tags de abertura e fechamento do elemento `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

