---
id: 5ef9b03c81a63668521804d7
title: Passo 38
challengeType: 0
dashedName: step-38
---

# --description--

O atributo `action` indica o destino dos dados do formulário. 

Aqui está um exemplo de um elemento `form` com o atributo `action`:

```html
<form action="/submit-url"></form>
```

No exemplo, `action="/submit-url"` informa ao navegador que os dados do formulário devem ser enviados para o caminho `/submit-url`.

Adicione um atributo `action` com o valor `https://freecatphotoapp.com/submit-cat-photo` ao elemento `form`.

# --hints--

O elemento `form` deve ter uma tag de abertura e uma de fechamento na ordem correta.Pode ser que você não tenha inserido uma delas, as duas, ou que elas estejam na ordem errada.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Seu elemento `form` aninhado no último elemento `section` deve estar abaixo do elemento `h2`. Você colocou o elemento `h2` e o elemento `form` na ordem errada.

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

Seu elemento `form` não deve conter conteúdo. Remova quaisquer elementos HTML ou texto entre as tags do elemento `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

O elemento `form` não tem um atributo `action`.Verifique se há um espaço depois do nome da tag de abertura e/ou se há espaços antes de todos os nomes dos atributos.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

O elemento `form` deve ter o atributo `action` com o valor `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Embora você tenha definido o valor do atributo `action` com o URL correto, é recomendável sempre cercar o valor de um atributo com aspas.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

