---
id: 5efb23e70dc218d6c85f89b1
title: Passo 40
challengeType: 0
dashedName: step-40
---

# --description--

Existem muitos tipos de entradas (inputs) que você pode criar usando o atributo `type`.Você pode criar facilmente um campo de senha, um botão de redefinição ou um controle para permitir que os usuários selecionem um arquivo armazenado no computador.

Crie um campo de texto para obter a entrada de texto de um usuário adicionando o atributo `type` com o valor `text` para o elemento `input`.

# --hints--

Você excluiu o elemento `input` ou ele tem uma sintaxe inválida.Se você adicionou atributos, certifique-se de que seus valores estejam cercados por aspas.

```js
assert(document.querySelectorAll('input').length);
```

Seu elemento `form` deve conter apenas o elemento `input`. Remova quaisquer elementos HTML extras ou texto entre as tags do elemento `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

O elemento `input` não tem o atributo `type` com o valor `text`.Verifique se há um espaço depois do nome da tag de abertura e/ou se há espaços antes de todos os nomes dos atributos.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

O elemento `input` deve ter o atributo `type` com o valor `text`.Você omitiu o valor ou tem um erro de digitação.Lembre-se de que os valores dos atributos devem estar cercados com aspas.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type').replace(/\s+/g, ' '),
  /^text$/i
);
```

Embora você tenha definido o atributo `type` do elemento `input` com o valor `text`, é recomendável sempre cercar o valor de um atributo com aspas.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

