---
id: 671141feba228a35cefba82d
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

Transforme o texto `cute cats` dentro do elemento âncora para:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

Deve haver um novo elemento anchor no primeiro elemento `p`.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

O texto do link deve ser `cute cats`.

```js
assert.strictEqual(document.querySelector('a')?.innerText.trim(), "cute cats");
```

O atributo `href` do novo elemento anchor deve ser `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

Deve haver apenas uma instância da frase `cute cats` em seu código.  

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

O texto do elemento `p` ainda deve ser `Everyone loves cute cats online!`.

```js
assert.strictEqual(document.querySelector('p')?.innerText.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
