---
id: 671141d8e32fe934c26fa1be
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Adicione as palavras `See more ` antes do elemento anchor e ` in our gallery` depois do elemento anchor.

# --hints--

Você deve usar o elemento de âncora existente, não crie um novo.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

O texto dentro do elemento de âncora, não pode ser mudado, ele precisa ser `cat photos`

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

Você deve ter `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` no seu código.

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
