---
id: 68a47dcdbce12193caf77012
title: Implementieren Sie einen Suchindex für sortierte Arrays
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die den kleinsten Index zurückgibt, an dem ein Wert in ein Array eingefügt werden sollte, nachdem dieses aufsteigend sortiert wurde.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `getIndexToIns`-Funktion haben, die zwei Argumente annimmt: ein Array und eine Zahl.  
2. Sie sollten die `sort`-Methode verwenden, um das Array aufsteigend zu sortieren.
3. Ihre `getIndexToIns`-Funktion sollte den kleinsten Index zurückgeben, an dem die Zahl eingefügt werden sollte, indem Sie die `findIndex`-Methode verwenden.
4. Ihre `getIndexToIns`-Funktion sollte immer eine Zahl zurückgeben.

**Hinweis:**  
Die `findIndex`-Methode ist eine eingebaute Array-Methode in JavaScript. Sie nimmt eine Callback-Funktion und gibt den Index des ersten Elements zurück, das die Bedingung erfüllt. Sowohl `findIndex` als auch `sort` sind Funktionen höherer Ordnung.

**Beispiele:**

- `getIndexToIns([1, 2, 3, 4], 1.5)` sollte `1` zurückgeben, weil `1.5` größer ist als `1` (Index `0`) und kleiner als `2` (Index `1`).
- `getIndexToIns([20, 3, 5], 19)` sollte `2` zurückgeben, weil nach dem Sortieren zu `[3, 5, 20]` `19` kleiner ist als `20` (Index `2`) und größer als `5` (Index `1`).


# --hints--

Sie sollten eine `getIndexToIns`-Funktion haben.

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` sollte immer eine Zahl zurückgeben.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

Ihre Funktion sollte die `sort`-Methode verwenden.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


Ihre Funktion sollte die `findIndex`-Methode verwenden.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` sollte `3` zurückgeben.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` sollte `2` zurückgeben.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` sollte `1` zurückgeben.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` sollte `0` zurückgeben.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` sollte `2` zurückgeben.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` sollte `2` zurückgeben.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` sollte `3` zurückgeben.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` sollte `0` zurückgeben.

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
```

