---
id: 587d7b7b367417b2b2512b16
title: Створення складних багатовимірних масивів
challengeType: 1
forumTopicId: 301159
dashedName: create-complex-multi-dimensional-arrays
---

# --description--

Чудово! Ви щойно дізналися багато про масиви! Це був досить загальний огляд, і попереду ще багато чого можна дізнатися про роботу з масивами, багато з чого ви побачите в наступних розділах. Але перш ніж перейти до вивчення <dfn>об’єктів</dfn>, давайте ще раз поглянемо, і побачимо, як масиви можуть стати трохи складнішими, ніж те, що ми бачили в попередніх завданнях.

Одна з найпотужніших функцій, коли думаєш про масиви як про структури даних, полягає в тому, що масиви можуть містити, або навіть повністю складатися з інших масивів. Ми вже бачили масиви, які містять масиви в попередніх завданнях, але досить прості. Однак масиви можуть містити нескінченну глибину масивів, які можуть містити інші масиви, кожен зі своїм довільним рівнем глибини, і так далі. Таким чином, масив дуже швидко може стати дуже складною структурою даних, відомою як <dfn>багатовимірний</dfn> або вкладений масив. Розгляньте наступний приклад:

```js
let nestedArray = [
  ['deep'],
  [
    ['deeper'], ['deeper'] 
  ],
  [
    [
      ['deepest'], ['deepest']
    ],
    [
      [
        ['deepest-est?']
      ]
    ]
  ]
];
```

Масив `deep` вкладений на 2 рівні глибини. Масиви `deeper` — на 3 рівні. Масиви `deepest` — на 4 рівні, а `deepest-est?` — на 5.

Хоч цей приклад може здатися заплутаним, такий рівень складності не є рідкістю або незвичайним, коли працюєш з великими обсягами даних. Однак ми все одно можемо дуже легко отримати доступ до найглибших рівнів такого складного масиву за допомогою дужкової нотації:

```js
console.log(nestedArray[2][1][0][0][0]);
```

Це виведе рядок `deepest-est?`. І тепер, коли ми знаємо, де знаходиться цей фрагмент даних, ми можемо його змінити, якщо потрібно:

```js
nestedArray[2][1][0][0][0] = 'deeper still';

console.log(nestedArray[2][1][0][0][0]);
```

Тепер він виводить `deeper still`.

# --instructions--

Ми визначили змінну `myNestedArray`, яка дорівнює масиву. Змініть `myNestedArray`, використовуючи будь-яку комбінацію <dfn>рядків</dfn>, <dfn>чисел</dfn> і <dfn>булевих</dfn> значень як елементи даних, так щоб він мав рівно п’ять рівнів глибини (пам’ятайте, що зовнішній масив — це рівень 1). Десь на третьому рівні включіть рядок `deep`, на четвертому рівні — рядок `deeper`, а на п’ятому рівні — рядок `deepest`.

# --hints--

`myNestedArray` повинен містити лише числа, булеві значення та рядки як елементи даних

```js
assert.strictEqual(
  (function (arr) {
    let flattened = (function flatten(arr) {
      const flat = [].concat(...arr);
      return flat.some(Array.isArray) ? flatten(flat) : flat;
    })(arr);
    for (let i = 0; i < flattened.length; i++) {
      if (
        typeof flattened[i] !== 'number' &&
        typeof flattened[i] !== 'string' &&
        typeof flattened[i] !== 'boolean'
      ) {
        return false;
      }
    }
    return true;
  })(myNestedArray),
  true
);
```

`myNestedArray` повинен мати рівно 5 рівнів глибини

```js
assert.strictEqual(
  (function (arr) {
    let depth = 0;
    function arrayDepth(array, i, d) {
      if (Array.isArray(array[i])) {
        arrayDepth(array[i], 0, d + 1);
      } else {
        depth = d > depth ? d : depth;
      }
      if (i < array.length) {
        arrayDepth(array, i + 1, d);
      }
    }
    arrayDepth(arr, 0, 0);
    return depth;
  })(myNestedArray),
  4
);
```

`myNestedArray` повинен містити рівно одне входження рядка `deep` у масиві, вкладеному на 3 рівні глибини

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deep').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deep')[0] === 2
);
```

`myNestedArray` повинен містити рівно одне входження рядка `deeper` у масиві, вкладеному на 4 рівні глибини

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deeper').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deeper')[0] === 3
);
```

`myNestedArray` повинен містити рівно одне входження рядка `deepest` у масиві, вкладеному на 5 рівні глибини

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deepest').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deepest')[0] === 4
);
```

# --seed--

## --seed-contents--

```js
let myNestedArray = [
  // Only change code below this line
  ['unshift', false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
  // Only change code above this line
];
```

# --solutions--

```js
let myNestedArray = [
  ['unshift', ['deep', ['deeper', ['deepest']]],false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
];
```
