---
id: 587d7b7e367417b2b2512b20
title: Використання масиву для збереження колекції даних
challengeType: 1
forumTopicId: 301167
dashedName: use-an-array-to-store-a-collection-of-data
---

# --description--

Нижче наведено приклад найпростішої реалізації структури даних масиву. Це відомо як <dfn>одновимірний масив</dfn>, що означає, що він має лише один рівень або що в ньому немає вкладених інших масивів. Зверніть увагу, що він містить <dfn>булеві</dfn>, <dfn>рядки</dfn> та <dfn>числа</dfn>, серед інших дійсних типів даних JavaScript:

```js
let simpleArray = ['one', 2, 'three', true, false, undefined, null];
console.log(simpleArray.length);
```

Виклик `console.log` відображає `7`.

У всіх масивів є властивість довжина, яку, як показано вище, дуже легко отримати за допомогою синтаксису `Array.length`. Нижче наведено більш складну реалізацію масиву. Це відомо як <dfn>багатовимірний масив</dfn>, або масив, що містить інші масиви. Зверніть увагу, що цей масив також містить JavaScript <dfn>об’єкти</dfn>, які ми детально розглянемо в наступному розділі, але наразі все, що вам потрібно знати, це те, що масиви також можуть зберігати складні об’єкти.

```js
let complexArray = [
  [
    {
      one: 1,
      two: 2
    },
    {
      three: 3,
      four: 4
    }
  ],
  [
    {
      a: "a",
      b: "b"
    },
    {
      c: "c",
      d: "d"
    }
  ]
];
```

# --instructions--

Ми визначили змінну з назвою `yourArray`. Завершіть інструкцію, призначивши змінній `yourArray` масив довжиною щонайменше 5 елементів. Ваш масив повинен містити принаймні один <dfn>рядок</dfn>, одне <dfn>число</dfn> та один <dfn>булевий</dfn>.

# --hints--

`yourArray` має бути масивом.

```js
assert.strictEqual(Array.isArray(yourArray), true);
```

`yourArray` має містити щонайменше 5 елементів.

```js
assert.isAtLeast(yourArray.length, 5);
```

`yourArray` має містити принаймні один `boolean`.

```js
assert(yourArray.filter((el) => typeof el === 'boolean').length >= 1);
```

`yourArray` має містити принаймні один `number`.

```js
assert(yourArray.filter((el) => typeof el === 'number').length >= 1);
```

`yourArray` має містити принаймні один `string`.

```js
assert(yourArray.filter((el) => typeof el === 'string').length >= 1);
```

# --seed--

## --seed-contents--

```js
let yourArray; // Change this line
```

# --solutions--

```js
let yourArray = ['a string', 100, true, ['one', 2], 'another string'];
```
