---
id: 587d7b7a367417b2b2512b12
title: Копіювання елементів масиву за допомогою slice()
challengeType: 1
forumTopicId: 301158
dashedName: copy-array-items-using-slice
---

# --description--

Наступний метод, який ми розглянемо, — це `slice()`. Замість того, щоб змінювати масив, `slice()` копіює або *витягує* задану кількість елементів у новий масив, залишаючи масив, до якого він застосовується, без змін. `slice()` приймає лише 2 параметри — перший це індекс, з якого починається витягування, а другий — індекс, на якому витягування припиняється (витягування відбувається до, але не включно з елементом за цим індексом). Розгляньте це:

```js
let weatherConditions = ['rain', 'snow', 'sleet', 'hail', 'clear'];

let todaysWeather = weatherConditions.slice(1, 3);
```

`todaysWeather` матиме значення `['snow', 'sleet']`, тоді як `weatherConditions` залишиться з `['rain', 'snow', 'sleet', 'hail', 'clear']`.

Фактично, ми створили новий масив, витягуючи елементи з існуючого масиву.

# --instructions--

Ми визначили функцію `forecast`, яка приймає масив як аргумент. Змініть функцію, використовуючи `slice()`, щоб витягти інформацію з аргументного масиву і повернути новий масив, що містить рядкові елементи `warm` і `sunny`.

# --hints--

`forecast` має повертати `["warm", "sunny"]`

```js
assert.deepEqual(
  forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']),
  ['warm', 'sunny']
);
```

Функція `forecast` має використовувати метод `slice()`

```js
assert(/\.slice\(/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function forecast(arr) {
  // Only change code below this line

  return arr;
}

// Only change code above this line
console.log(forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']));
```

# --solutions--

```js
function forecast(arr) {
  return arr.slice(2,4);
}
```
