---
id: 587d78b2367417b2b2512b0f
title: Видалення елементів з масиву за допомогою pop() і shift()
challengeType: 1
forumTopicId: 301165
dashedName: remove-items-from-an-array-with-pop-and-shift
---

# --description--

Обидва `push()` і `unshift()` мають відповідні методи, які майже є функціональними протилежностями: `pop()` і `shift()`. Як ви, мабуть, вже здогадалися, замість того, щоб додавати, `pop()` *видаляє* елемент з кінця масиву, тоді як `shift()` видаляє елемент з початку. Ключова різниця між `pop()` і `shift()` та їхніми "родичами" `push()` і `unshift()` полягає в тому, що жоден із цих методів не приймає параметрів, і кожен дозволяє змінювати масив лише по одному елементу за раз.

Давайте подивимося:

```js
let greetings = ['whats up?', 'hello', 'see ya!'];

greetings.pop();
```

`greetings` матиме значення `['whats up?', 'hello']`.

```js
greetings.shift();
```

`greetings` матиме значення `['hello']`.

Ми також можемо повернути значення видаленого елемента за допомогою будь-якого з цих методів так:

```js
let popped = greetings.pop();
```

`greetings` матиме значення `[]`, а `popped` матиме значення `hello`.

# --instructions--

Ми визначили функцію `popShift`, яка приймає масив як аргумент і повертає новий масив. Змініть функцію, використовуючи `pop()` і `shift()`, щоб видалити перший і останній елементи аргументного масиву, і призначте видалені елементи відповідним змінним, щоб повернутий масив містив їхні значення.

# --hints--

`popShift(["challenge", "is", "not", "complete"])` має повертати `["challenge", "complete"]`

```js
assert.deepEqual(popShift(['challenge', 'is', 'not', 'complete']), [
  'challenge',
  'complete'
]);
```

Функція `popShift` має використовувати метод `pop()`

```js
assert.notStrictEqual(popShift.toString().search(/\.pop\(/), -1);
```

Функція `popShift` має використовувати метод `shift()`

```js
assert.notStrictEqual(popShift.toString().search(/\.shift\(/), -1);
```

# --seed--

## --seed-contents--

```js
function popShift(arr) {
  let popped; // Change this line
  let shifted; // Change this line
  return [shifted, popped];
}

console.log(popShift(['challenge', 'is', 'not', 'complete']));
```

# --solutions--

```js
function popShift(arr) {
  let popped = arr.pop(); // Change this line
  let shifted = arr.shift(); // Change this line
  return [shifted, popped];
}
```
