---
id: 587d78b2367417b2b2512b0e
title: додати елементи до масиву за допомогою push() та unshift()
challengeType: 1
forumTopicId: 301151
dashedName: add-items-to-an-array-with-push-and-unshift
---

# --description--

Довжина масиву, як і типи даних, які він може містити, не є фіксованою. Масиви можна визначати з довжиною будь-якої кількості елементів, а елементи можна додавати або видаляти з часом; іншими словами, масиви є <dfn>змінними</dfn>. У цьому завданні ми розглянемо два методи, за допомогою яких можна програмно змінювати масив: `Array.push()` та `Array.unshift()`.

Обидва методи приймають один або кілька елементів як параметри і додають ці елементи до масиву, на якому виконується метод; метод `push()` додає елементи в кінець масиву, а `unshift()` додає елементи на початок. Розгляньте наступне:

```js
let twentyThree = 'XXIII';
let romanNumerals = ['XXI', 'XXII'];

romanNumerals.unshift('XIX', 'XX');
```

`romanNumerals` матиме значення `['XIX', 'XX', 'XXI', 'XXII']`.

```js
romanNumerals.push(twentyThree);
```

`romanNumerals` матиме значення `['XIX', 'XX', 'XXI', 'XXII', 'XXIII']`. Зверніть увагу, що ми також можемо передавати змінні, що дає нам ще більшу гнучкість у динамічній зміні даних нашого масиву.

# --instructions--

Ми визначили функцію `mixedNumbers`, якій передаємо масив як аргумент. Змініть функцію, використовуючи `push()` та `unshift()`, щоб додати `'I', 2, 'three'` на початок масиву і `7, 'VIII', 9` у кінець, щоб повернутий масив містив представлення чисел від 1 до 9 у порядку.

# --hints--

`mixedNumbers(["IV", 5, "six"])` тепер має повертати `["I", 2, "three", "IV", 5, "six", 7, "VIII", 9]`

```js
assert.deepEqual(mixedNumbers(['IV', 5, 'six']), [
  'I',
  2,
  'three',
  'IV',
  5,
  'six',
  7,
  'VIII',
  9
]);
```

Функція `mixedNumbers` має використовувати метод `push()`

```js
assert(mixedNumbers.toString().match(/\.push/));
```

Функція `mixedNumbers` має використовувати метод `unshift()`

```js
assert(mixedNumbers.toString().match(/\.unshift/));
```

# --seed--

## --seed-contents--

```js
function mixedNumbers(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(mixedNumbers(['IV', 5, 'six']));
```

# --solutions--

```js
function mixedNumbers(arr) {
  arr.push(7,'VIII',9);
  arr.unshift('I',2,'three');
  return arr;
}
```
