---
id: 587d7b7c367417b2b2512b1b
title: Використання ключового слова delete для видалення властивостей об’єкта
challengeType: 1
forumTopicId: 301168
dashedName: use-the-delete-keyword-to-remove-object-properties
---

# --description--

Тепер ви знаєте, що таке об’єкти, їхні базові функції та переваги. Коротко кажучи, це сховища пар ключ-значення, які забезпечують гнучкий, інтуїтивний спосіб структурування даних, ***і*** вони забезпечують дуже швидкий час пошуку. Протягом решти цих завдань ми опишемо кілька загальних операцій, які ви можете виконувати з об’єктами, щоб ви могли впевнено застосовувати ці корисні структури даних у своїх програмах.

У попередніх завданнях ми додавали та змінювали пари ключ-значення об’єкта. Тут ми побачимо, як можна *видалити* пару ключ-значення з об’єкта.

Давайте ще раз повернемося до нашого прикладу об’єкта `foods`. Якщо ми хочемо видалити ключ `apples`, ми можемо зробити це, використовуючи ключове слово `delete` ось так:

```js
delete foods.apples;
```

# --instructions--

Використайте ключове слово delete, щоб видалити ключі `oranges`, `plums` і `strawberries` з об’єкта `foods`.

# --hints--

Об’єкт `foods` повинен містити лише три ключі: `apples`, `grapes` і `bananas`.

```js
assert(
  !foods.hasOwnProperty('oranges') &&
    !foods.hasOwnProperty('plums') &&
    !foods.hasOwnProperty('strawberries') &&
    Object.keys(foods).length === 3
);
```

Ключі `oranges`, `plums` і `strawberries` потрібно видалити за допомогою `delete`.

```js
assert(
  __helpers.removeJSComments(code).search(/oranges:/) !== -1 &&
    __helpers.removeJSComments(code).search(/plums:/) !== -1 &&
    __helpers.removeJSComments(code).search(/strawberries:/) !== -1
);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

delete foods.oranges;
delete foods.plums;
delete foods.strawberries;

console.log(foods);
```
