---
id: 587d7b7b367417b2b2512b15
title: Ітерувати по всіх елементах масиву за допомогою циклів for
challengeType: 1
forumTopicId: 301161
dashedName: iterate-through-all-an-arrays-items-using-for-loops
---

# --description--

Іноді при роботі з масивами дуже зручно ітерувати по кожному елементу, щоб знайти один або кілька елементів, які нам можуть знадобитися, або щоб маніпулювати масивом на основі того, які елементи даних відповідають певному набору критеріїв. JavaScript пропонує кілька вбудованих методів, які ітерують по масивах трохи по-різному, щоб досягти різних результатів (наприклад, `every()`, `forEach()`, `map()` тощо), проте техніка, яка є найгнучкішою і дає нам найбільший контроль, — це простий цикл `for`.

Розгляньте наступне:

```js
function greaterThanTen(arr) {
  let newArr = [];
  for (let i = 0; i < arr.length; i++) {
    if (arr[i] > 10) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}

greaterThanTen([2, 12, 8, 14, 80, 0, 1]);
```

Використовуючи цикл `for`, ця функція ітерує по кожному елементу масиву, отримує до нього доступ і піддає простому тесту, який ми створили. Таким чином, ми легко і програмно визначили, які елементи даних більші за `10`, і повернули новий масив `[12, 14, 80]`, що містить ці елементи.

# --instructions--

Ми визначили функцію `filteredArray`, яка приймає `arr`, вкладений масив, і `elem` як аргументи, і повертає новий масив. `elem` представляє елемент, який може бути присутнім або відсутнім у одному або кількох масивах, вкладених у `arr`. Змініть функцію, використовуючи цикл `for`, щоб повернути відфільтровану версію переданого масиву так, щоб будь-який масив, вкладений у `arr`, що містить `elem`, був видалений.

# --hints--

`filteredArray([[10, 8, 3], [14, 6, 23], [3, 18, 6]], 18)` має повертати `[[10, 8, 3], [14, 6, 23]]`

```js
assert.deepEqual(
  filteredArray(
    [
      [10, 8, 3],
      [14, 6, 23],
      [3, 18, 6]
    ],
    18
  ),
  [
    [10, 8, 3],
    [14, 6, 23]
  ]
);
```

`filteredArray([["trumpets", 2], ["flutes", 4], ["saxophones", 2]], 2)` має повертати `[["flutes", 4]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['trumpets', 2],
      ['flutes', 4],
      ['saxophones', 2]
    ],
    2
  ),
  [['flutes', 4]]
);
```

`filteredArray([["amy", "beth", "sam"], ["dave", "sean", "peter"]], "peter")` має повертати `[["amy", "beth", "sam"]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['amy', 'beth', 'sam'],
      ['dave', 'sean', 'peter']
    ],
    'peter'
  ),
  [['amy', 'beth', 'sam']]
);
```

`filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3)` має повертати `[]`

```js
assert.deepEqual(
  filteredArray(
    [
      [3, 2, 3],
      [1, 6, 3],
      [3, 13, 26],
      [19, 3, 9]
    ],
    3
  ),
  []
);
```

Функція `filteredArray` має використовувати цикл `for`

```js
assert.notStrictEqual(filteredArray.toString().search(/for/), -1);
```

# --seed--

## --seed-contents--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  // Only change code below this line

  // Only change code above this line
  return newArr;
}

console.log(filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3));
```

# --solutions--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  for (let i = 0; i<arr.length; i++) {
    if (arr[i].indexOf(elem) < 0) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}
```
