---
id: 587d7b7d367417b2b2512b1d
title: Ітерувати по ключах об’єкта за допомогою інструкції for...in
challengeType: 1
forumTopicId: 301162
dashedName: iterate-through-the-keys-of-an-object-with-a-for---in-statement
---

# --description--

Іноді потрібно ітерувати по всіх ключах об’єкта. Для цього можна використати цикл <dfn>for...in</dfn>. Цикл for...in виглядає так:

```javascript
const refrigerator = {
  'milk': 1,
  'eggs': 12,
};

for (const food in refrigerator) {
  console.log(food, refrigerator[food]);
}
```

Цей код виводить у консоль `milk 1` та `eggs 12`, причому кожна пара ключ-значення виводиться на окремому рядку.

Ми визначили змінну `food` у заголовку циклу, і ця змінна на кожній ітерації приймала значення ключа об’єкта, в результаті чого в консоль виводилося ім’я кожної їжі.

**ПРИМІТКА:** Об’єкти не зберігають порядок ключів так, як це роблять масиви; тому позиція ключа в об’єкті або відносний порядок його появи не мають значення при зверненні або доступі до цього ключа.

# --instructions--

Ми визначили функцію `countOnline`, яка приймає один аргумент `allUsers`. Використайте інструкцію <dfn>for...in</dfn> всередині цієї функції, щоб ітерувати по об’єкту `allUsers` і повернути кількість користувачів, у яких властивість `online` має значення `true`. Нижче наведено приклад об’єкта, який може бути переданий у `countOnline`. Кожен користувач матиме властивість `online`, встановлену або в `true`, або в `false`.

```js
{
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
```

# --hints--

Функція `countOnline` має використовувати інструкцію `for in` для ітерації по ключах об’єкта, який їй передано.

```js
assert.match(
  __helpers.removeJSComments(code),
    /for\s*\(\s*(var|let|const)\s+[a-zA-Z_$]\w*\s+in\s+[a-zA-Z_$]\w*\s*\)/
);
```

Функція `countOnline` має повертати `1`, коли їй передано об’єкт `{ Alan: { online: false }, Jeff: { online: true }, Sarah: { online: false } }`

```js
const usersObj1 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj1), 1);
```

Функція `countOnline` має повертати `2`, коли їй передано об’єкт `{ Alan: { online: true }, Jeff: { online: false }, Sarah: { online: true } }`

```js
const usersObj2 = {
  Alan: {
    online: true
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: true
  }
}
assert.equal(countOnline(usersObj2), 2);
```

Функція `countOnline` має повертати `0`, коли їй передано об’єкт `{ Alan: { online: false }, Jeff: { online: false }, Sarah: { online: false } }`

```js
const usersObj3 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj3), 0);
```

# --seed--

## --seed-contents--

```js
const users = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}

function countOnline(allUsers) {
  // Only change code below this line

  // Only change code above this line
}

console.log(countOnline(users));
```

# --solutions--

```js
function countOnline(allUsers) {
  let numOnline = 0;
  for(const user in allUsers){
    if(allUsers[user].online) {
      numOnline++;
    }
  }
  return numOnline;
}
```
