---
id: 587d7b7c367417b2b2512b1a
title: Доступ до імен властивостей за допомогою дужкової нотації
challengeType: 1
forumTopicId: 301150
dashedName: access-property-names-with-bracket-notation
---

# --description--

У першому завданні з об’єктами ми згадували використання дужкової нотації як спосіб доступу до значень властивостей за допомогою оцінки змінної. Наприклад, уявіть, що наш об’єкт `foods` використовується в програмі для касового апарата супермаркету. У нас є функція, яка встановлює `selectedFood`, і ми хочемо перевірити об’єкт `foods` на наявність цього продукту. Це може виглядати так:

```js
let selectedFood = getCurrentFood(scannedItem);
let inventory = foods[selectedFood];
```

Цей код оцінить значення, збережене у змінній `selectedFood`, і поверне значення цього ключа в об’єкті `foods`, або `undefined`, якщо такого ключа немає. Дужкова нотація дуже корисна, бо іноді властивості об’єкта невідомі до часу виконання або нам потрібно отримати до них доступ більш динамічно.

# --instructions--

Ми визначили функцію `checkInventory`, яка приймає відсканований елемент як аргумент. Поверніть поточне значення ключа `scannedItem` в об’єкті `foods`. Ви можете припустити, що в `checkInventory` будуть передані лише дійсні ключі.

# --hints--

`checkInventory` має бути функцією.

```js
assert.strictEqual(typeof checkInventory, 'function');
```

Об’єкт `foods` повинен містити лише такі пари ключ-значення: `apples: 25`, `oranges: 32`, `plums: 28`, `bananas: 13`, `grapes: 35`, `strawberries: 27`.

```js
assert.deepEqual(foods, {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
});
```

`checkInventory("apples")` має повертати `25`.

```js
assert.strictEqual(checkInventory('apples'), 25);
```

`checkInventory("bananas")` має повертати `13`.

```js
assert.strictEqual(checkInventory('bananas'), 13);
```

`checkInventory("strawberries")` має повертати `27`.

```js
assert.strictEqual(checkInventory('strawberries'), 27);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(checkInventory("apples"));
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  return foods[scannedItem];
}
```
