---
id: 587d7b7c367417b2b2512b19
title: "Змінити об’єкт, вкладений у об’єкт"
challengeType: 1
forumTopicId: 301164
dashedName: modify-an-object-nested-within-an-object
---

# --description--

Тепер розглянемо трохи складніший об’єкт. Властивості об’єкта можуть бути вкладеними на довільну глибину, а їхні значення можуть бути будь-яким типом даних, підтримуваним JavaScript, включно з масивами та навіть іншими об’єктами. Розгляньмо наступне:

```js
let nestedObject = {
  id: 28802695164,
  date: 'December 31, 2016',
  data: {
    totalUsers: 99,
    online: 80,
    onlineStatus: {
      active: 67,
      away: 13,
      busy: 8
    }
  }
};
```

`nestedObject` має три властивості: `id` (значення — число), `date` (значення — рядок) і `data` (значення — об’єкт із вкладеною структурою). Хоча структури можуть швидко ставати складними, ми все одно можемо використовувати ті самі нотації для доступу до потрібної інформації. Щоб призначити значення `10` властивості `busy` вкладеного об’єкта `onlineStatus`, ми використовуємо точкову нотацію для посилання на властивість:

```js
nestedObject.data.onlineStatus.busy = 10;
```

# --instructions--

Тут ми визначили об’єкт `userActivity`, який містить інший об’єкт, вкладений у нього. Призначте значення ключу `online` рівним `45`.

# --hints--

`userActivity` має мати властивості `id`, `date` і `data`.

```js
assert(
  'id' in userActivity && 'date' in userActivity && 'data' in userActivity
);
```

`userActivity` має мати ключ `data`, призначений об’єктом із ключами `totalUsers` і `online`.

```js
assert('totalUsers' in userActivity.data && 'online' in userActivity.data);
```

Властивість `online`, вкладена в ключ `data` об’єкта `userActivity`, має бути призначена значенням `45`.

```js
assert(userActivity.data.online === 45);
```

Властивість `online` має бути призначена за допомогою точкової або дужкової нотації.

```js
assert.strictEqual(__helpers.removeJSComments(code).search(/online: 45/), -1);
```

# --seed--

## --seed-contents--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

// Only change code below this line

// Only change code above this line

console.log(userActivity);
```

# --solutions--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

userActivity.data.online = 45;
```
