---
id: 587d78b2367417b2b2512b10
title: Видалення елементів за допомогою splice()
challengeType: 1
forumTopicId: 301166
dashedName: remove-items-using-splice
---

# --description--

Отже, ми навчилися видаляти елементи з початку та кінця масивів за допомогою `shift()` і `pop()`, але що, якщо ми хочемо видалити елемент десь посередині? Або видалити одразу кілька елементів? Саме для цього існує `splice()`. `splice()` дозволяє зробити саме це: **видалити будь-яку кількість послідовних елементів** з будь-якого місця в масиві.

`splice()` може приймати до 3 параметрів, але наразі ми зосередимося лише на перших двох. Перші два параметри `splice()` — це цілі числа, які представляють індекси, або позиції, елементів у масиві, до якого виконується виклик `splice()`. І пам’ятайте, що масиви мають *нульовий індекс*, тому щоб вказати перший елемент масиву, ми використовуємо `0`. Перший параметр `splice()` означає індекс у масиві, з якого починається видалення елементів, а другий параметр вказує кількість елементів для видалення. Наприклад:

```js
let array = ['today', 'was', 'not', 'so', 'great'];

array.splice(2, 2);
```

Тут ми видаляємо 2 елементи, починаючи з третього елемента (індекс 2). `array` матиме значення `['today', 'was', 'great']`.

`splice()` не лише змінює масив, до якого він застосовується, а й повертає новий масив, що містить значення видалених елементів:

```js
let array = ['I', 'am', 'feeling', 'really', 'happy'];

let newArray = array.splice(3, 2);
```

`newArray` матиме значення `['really', 'happy']`.

# --instructions--

Ми ініціалізували масив `arr`. Використайте `splice()`, щоб видалити елементи з `arr`, так щоб у ньому залишилися лише елементи, сума яких дорівнює `10`.

# --hints--

Ви не повинні змінювати початковий рядок `const arr = [2, 4, 5, 1, 7, 5, 2, 1];`.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/constarr=\[2,4,5,1,7,5,2,1\];?/)
);
```

`arr` має містити лише елементи, сума яких дорівнює `10`.

```js
assert.strictEqual(
  arr.reduce((a, b) => a + b),
  10
);
```

Ваш код має використовувати метод `splice()` для `arr`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(/));
```

splice має лише видаляти елементи з `arr` і не додавати жодних нових елементів до `arr`.

```js
assert(
  !__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(\d+,\d+,\d+.*\)/g)
);
```

# --seed--

## --seed-contents--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
// Only change code below this line

// Only change code above this line
console.log(arr);
```

# --solutions--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
arr.splice(1, 4);
```
