---
id: 587d78b3367417b2b2512b11
title: додати елементи за допомогою splice()
challengeType: 1
forumTopicId: 301152
dashedName: add-items-using-splice
---

# --description--

Пам’ятаєте, у попередньому завданні ми згадували, що `splice()` може приймати до трьох параметрів? Насправді, ви можете використати третій параметр, який складається з одного або кількох елементів, щоб додати їх до масиву. Це дуже корисно, коли потрібно швидко замінити один елемент або множину елементів на інші.

```js
const numbers = [10, 11, 12, 12, 15];
const startIndex = 3;
const amountToDelete = 1;

numbers.splice(startIndex, amountToDelete, 13, 14);
console.log(numbers);
```

Другий випадок `12` видаляється, а `13` і `14` додаються на той самий індекс. Тепер масив `numbers` буде `[ 10, 11, 12, 13, 14, 15 ]`.

Тут ми починаємо з масиву чисел. Потім передаємо `splice()` такі аргументи: індекс, з якого починаємо видаляти елементи (3), кількість елементів для видалення (1), а решта аргументів (13, 14) буде вставлена, починаючи з того самого індексу. Зверніть увагу, що після `amountToDelete` може бути будь-яка кількість елементів (розділених комами), кожен з яких буде вставлений.

# --instructions--

Ми визначили функцію `htmlColorNames`, яка приймає масив HTML-кольорів як аргумент. Змініть функцію, використовуючи `splice()`, щоб видалити перші два елементи масиву і додати `'DarkSalmon'` та `'BlanchedAlmond'` на їхні відповідні місця.

# --hints--

`htmlColorNames` має повертати `["DarkSalmon", "BlanchedAlmond", "LavenderBlush", "PaleTurquoise", "FireBrick"]`

```js
assert.deepEqual(
  htmlColorNames([
    'DarkGoldenRod',
    'WhiteSmoke',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]),
  [
    'DarkSalmon',
    'BlanchedAlmond',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]
);
```

Функція `htmlColorNames` має використовувати метод `splice()`

```js
assert(/.splice/.test(__helpers.removeJSComments(code)));
```

Вам не слід використовувати `shift()` або `unshift()`.

```js
assert(!/shift|unshift/.test(__helpers.removeJSComments(code)));
```

Вам не слід використовувати дужкову нотацію масиву.

```js
assert(!/\[\d\]\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function htmlColorNames(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(htmlColorNames(['DarkGoldenRod', 'WhiteSmoke', 'LavenderBlush', 'PaleTurquoise', 'FireBrick']));
```

# --solutions--

```js
function htmlColorNames(arr) {
  arr.splice(0,2,'DarkSalmon', 'BlanchedAlmond');
  return arr;
}
```
