---
id: 587d7b7d367417b2b2512b1f
title: "змінити масив, збережений в об’єкті"
challengeType: 1
forumTopicId: 301163
dashedName: modify-an-array-stored-in-an-object
---

# --description--

Тепер ви ознайомилися з усіма базовими операціями для об’єктів JavaScript. Ви можете додавати, змінювати та видаляти пари ключ-значення, перевіряти наявність ключів і ітерувати над усіма ключами в об’єкті. Продовжуючи вивчати JavaScript, ви побачите ще більш універсальні застосування об’єктів. Крім того, уроки зі структур даних у розділі Підготовка до співбесіди з програмування навчальної програми також охоплюють об’єкти ES6 <dfn>Map</dfn> і <dfn>Set</dfn>, які схожі на звичайні об’єкти, але мають додаткові функції. Тепер, коли ви вивчили основи масивів і об’єктів, ви повністю готові розв’язувати складніші задачі за допомогою JavaScript!

# --instructions--

Погляньте на об’єкт, який ми надали в редакторі коду. Об’єкт `user` містить три ключі. Ключ `data` містить п’ять ключів, один із яких містить масив `friends`. З цього ви можете побачити, наскільки гнучкими є об’єкти як структури даних. Ми почали писати функцію `addFriend`. Завершіть її так, щоб вона приймала об’єкт `user`, додавала ім’я аргументу `friend` до масиву, збереженого в `user.data.friends`, і повертала цей масив.

# --hints--

Об’єкт `user` має містити ключі `name`, `age` і `data`.

```js
assert('name' in user && 'age' in user && 'data' in user);
```

Функція `addFriend` має приймати об’єкт `user` і рядок `friend` як аргументи та додавати друга до масиву `friends` в об’єкті `user`.

```js
assert(
  (function () {
    let L1 = user.data.friends.length;
    addFriend(user, 'Sean');
    let L2 = user.data.friends.length;
    return L2 === L1 + 1;
  })()
);
```

`addFriend(user, "Pete")` має повертати `["Sam", "Kira", "Tomo", "Pete"]`.

```js
assert.deepEqual(
  (function () {
    delete user.data.friends;
    user.data.friends = ['Sam', 'Kira', 'Tomo'];
    return addFriend(user, 'Pete');
  })(),
  ['Sam', 'Kira', 'Tomo', 'Pete']
);
```

# --seed--

## --seed-contents--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  // Only change code below this line

  // Only change code above this line
}

console.log(addFriend(user, 'Pete'));
```

# --solutions--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  userObj.data.friends.push(friend);
  return userObj.data.friends;
}
```
