---
id: 587d7b7c367417b2b2512b18
title: додати пари ключ-значення до об’єктів JavaScript
challengeType: 1
forumTopicId: 301153
dashedName: add-key-value-pairs-to-javascript-objects
---

# --description--

У найпростішому вигляді об’єкти — це просто колекції пар ключ-значення. Іншими словами, це фрагменти даних (значення), спроєктовані на унікальні ідентифікатори, які називаються властивостями (ключами). Подивіться на приклад:

```js
const tekkenCharacter = {
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true
};
```

Вищенаведений код визначає об’єкт персонажа з відеогри Tekken під назвою `tekkenCharacter`. Він має три властивості, кожна з яких відповідає певному значенню. Якщо ви хочете додати додаткову властивість, наприклад "origin", це можна зробити, призначивши `origin` об’єкту:

```js
tekkenCharacter.origin = 'South Korea';
```

Це використовує точкову нотацію. Якщо ви подивитесь на об’єкт `tekkenCharacter`, він тепер включатиме властивість `origin`. У Хоаранга також були характерні помаранчеве волосся. Ви можете додати цю властивість за допомогою дужкової нотації, зробивши так:

```js
tekkenCharacter['hair color'] = 'dyed orange';
```

Дужкова нотація потрібна, якщо у вашої властивості є пробіл або якщо ви хочете використати змінну для імені властивості. У наведеному випадку властивість укладена в лапки, щоб позначити її як рядок, і буде додана саме так, як показано. Без лапок вона буде оцінена як змінна, і ім’я властивості буде відповідати значенню цієї змінної. Ось приклад зі змінною:

```js
const eyes = 'eye color';

tekkenCharacter[eyes] = 'brown';
```

Після додавання всіх прикладів об’єкт виглядатиме так:

```js
{
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true,
  origin: 'South Korea',
  'hair color': 'dyed orange',
  'eye color': 'brown'
};
```

# --instructions--

Об’єкт `foods` створено з трьома записами. Використовуючи синтаксис на ваш вибір, додайте до нього ще три записи: `bananas` зі значенням `13`, `grapes` зі значенням `35` і `strawberries` зі значенням `27`.

# --hints--

`foods` має бути об’єктом.

```js
assert(typeof foods === 'object');
```

Об’єкт `foods` має містити ключ `bananas` зі значенням `13`.

```js
assert(foods.bananas === 13);
```

Об’єкт `foods` має містити ключ `grapes` зі значенням `35`.

```js
assert(foods.grapes === 35);
```

Об’єкт `foods` має містити ключ `strawberries` зі значенням `27`.

```js
assert(foods.strawberries === 27);
```

Визначення об’єкта `foods` не повинно змінюватися.

```js
assert(
  __helpers.removeJSComments(code).search(/let foods/) === -1 &&
  __helpers.removeJSComments(code).search(/const\s+foods\s*=\s*{\s*apples:\s*25,\s*oranges:\s*32,\s*plums:\s*28\s*};/
) !== -1
);
```

# --seed--

## --seed-contents--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

foods['bananas'] = 13;
foods['grapes']  = 35;
foods['strawberries'] = 27;
```
