---
id: 5a661e0f1068aca922b3ef17
title: Доступ до вмісту масиву за допомогою дужкової нотації
challengeType: 1
forumTopicId: 301149
dashedName: access-an-arrays-contents-using-bracket-notation
---

# --description--

Основна функція будь-якої структури даних — це, звичайно, не лише зберігання даних, а й можливість отримати ці дані за запитом. Отже, тепер, коли ми навчилися створювати масив, давайте почнемо думати про те, як ми можемо отримати інформацію з цього масиву.

Коли ми визначаємо простий масив, як показано нижче, у ньому є 3 елементи:

```js
let ourArray = ["a", "b", "c"];
```

У масиві кожен елемент має <dfn>індекс</dfn>. Цей індекс одночасно є позицією цього елемента в масиві та способом, як ви до нього звертаєтеся. Однак важливо зауважити, що масиви JavaScript є <dfn>нульовою індексацією</dfn>, тобто перший елемент масиву фактично знаходиться на ***нульовій*** позиції, а не на першій. Щоб отримати елемент із масиву, ми можемо взяти індекс у дужки і додати його в кінець масиву або, частіше, до змінної, яка посилається на об’єкт масиву. Це називається <dfn>дужковою нотацією</dfn>. Наприклад, якщо ми хочемо отримати `a` з `ourArray` і призначити це значення змінній, ми можемо зробити це за допомогою такого коду:

```js
let ourVariable = ourArray[0];
```

Тепер `ourVariable` має значення `a`.

Окрім доступу до значення за індексом, ви також можете *призначити* індексу значення, використовуючи ту саму нотацію:

```js
ourArray[1] = "not b anymore";
```

Використовуючи дужкову нотацію, ми тепер змінили елемент на індексі 1 із рядка `b` на `not b anymore`. Тепер `ourArray` — це `["a", "not b anymore", "c"]`.

# --instructions--

Щоб виконати це завдання, призначте 2-й позиції (індекс `1`) у `myArray` будь-яке значення, крім літери `b`.

# --hints--

`myArray[0]` має бути рівним літері `a`

```js
assert.strictEqual(myArray[0], 'a');
```

`myArray[1]` не має бути рівним літері `b`

```js
assert.notStrictEqual(myArray[1], 'b');
```

`myArray[2]` має бути рівним літері `c`

```js
assert.strictEqual(myArray[2], 'c');
```

`myArray[3]` має бути рівним літері `d`

```js
assert.strictEqual(myArray[3], 'd');
```

# --seed--

## --seed-contents--

```js
let myArray = ["a", "b", "c", "d"];
// Only change code below this line

// Only change code above this line
console.log(myArray);
```

# --solutions--

```js
let myArray = ["a", "b", "c", "d"];
myArray[1] = "e";
```
