---
id: 655cd899f8de09431eabb40c
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este proyecto, aprenderás estructuras de datos construyendo el algoritmo merge sort. 

Este es un algoritmo de ordenamiento que usa el principio de divide y vencerás para ordenar colecciones de datos. Es decir, 'divide' una colección en subpartes más pequeñas y 'vence' las subpartes ordenándolas de forma independiente, luego une las subpartes ordenadas.

Crea una función llamada `merge_sort`. Esta función se encargará de ordenar una lista de números. 

Usa la palabra clave `pass` en el cuerpo de la función.

# --hints--

Debes declarar una función llamada `merge_sort`. No olvides usar la palabra clave `pass` en el cuerpo de la función.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
