---
id: 6564aee9c077774ea49c3faf
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

The merge sort algorithm mainly performs three actions:

- Dividir una secuencia desordenada de elementos en subpartes
- Ordenar los elementos en las subpartes
- Unir las subpartes ordenadas

Lo anterior sucede recursivamente hasta que las subpartes se combinan en una secuencia completamente ordenada. Comencemos dividiendo la secuencia.

Primero, reemplaza la palabra clave `pass` con una variable `middle_point`. Utiliza el operador de división entera (`//`) para dividir la longitud de la lista `array` entre 2 y asigna el resultado a tu nueva variable `middle_point`. Recuerda sangrar tu código.

# --hints--

No debes tener `pass` en el cuerpo de la función.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

Debes declarar una variable `middle_point` y asignarle `len(array) // 2`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
