---
id: 6564a9fe51964c229d5b7f4c
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Necesitarás un parámetro que indique la colección de datos que se va a ordenar. Crea un parámetro llamado `array` en la función `merge_sort`.

# --hints--

Tu función `merge_sort` debe tomar un solo parámetro: `array`.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
