---
id: 656627b47bd2d2a4afbc945d
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

En el paso anterior, obtuviste el punto medio. Puedes usarlo para dividir `array` en dos y asignar cada parte a nuevas variables. 

Debes usar la sintaxis de segmento para extraer la mitad izquierda del `array` y asignarla a una variable llamada `left_part`.


# --hints--

Debes tener una variable llamada `left_part` en tu función `merge_sort`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

Debes asignar `array[:middle_point]` a la variable `left_part`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
