---
id: 656680b0fc79f2c38a34d90e
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Ahora es momento de ordenar y fusionar las listas (`left_part` y `right_part`) en el `array` original. 

Puedes hacer esto comparando elementos en ambas listas y fusionando el elemento más pequeño en la lista principal. Vas a hacer esta comparación para todos los índices en `left_part` y `right_part`.

Crea tres variables: `left_array_index`, `right_array_index` y `sorted_index` y establece sus valores en `0`. Estas variables te ayudarán a realizar un seguimiento de cada índice durante el proceso de ordenamiento.

# --hints--

Debes tener una variable llamada `left_array_index` dentro de tu función `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index/);
  }
})
```

El valor de `left_array_index` debe ser 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index\s*=\s*0/);
  }
})
```

Debes tener una variable llamada `right_array_index` dentro de tu función `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index/);
  }
})
```

El valor de `right_array_index` debe ser 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index\s*=\s*0/);
  }
})

```

Debes tener una variable llamada `sorted_index` dentro de tu función `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index/);
  }
})
```

El valor de `sorted_index` debe ser 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index\s*=\s*0/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)

--fcc-editable-region--

--fcc-editable-region--
```
