---
id: 656639771fed09b5c0e8fe71
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Ahora que has dividido la lista de `array` en dos listas separadas, continuarás dividiendo cada lista hasta que cada elemento esté solo en su propia lista. Una lista con un solo número siempre está ordenada.

Para hacer eso, llama recursivamente a `merge_sort` dentro de tu función. No pases ningún argumento por ahora.

# --hints--

Debes llamar a `merge_sort` al final del cuerpo de tu función.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
