---
id: 656702f8b4cbd8cbf0a433c6
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Dentro de tu funci칩n, crea un bucle `while` que compare un elemento en `left_part` con un elemento en `right_part`, y fusione el elemento m치s peque침o a la lista principal de `array`. 

Crea dos condiciones para el bucle: una que comprueba si la `left_array_index` es menor que la longitud de `left_part` y otra condici칩n que comprueba si `right_array_index` es menor que la longitud de `right_part`.

# --hints--

Tu ciclo `while` debe tener estas condiciones: `left_array_index < len(left_part) and right_array_index < len(right_part)`. No olvides usar la palabra clave `pass` en el cuerpo del ciclo.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

Usa la palabra clave `pass` en el cuerpo del bucle

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
