---
id: 656710d1e0ec62253426db24
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

El `while` que creaste compara un elemento de `left_part` con otro en `right_part`, y luego agrega el elemento menor a la lista principal `array`.

Continuará esta operación hasta que no queden elementos por comparar. Pero `left_part` aún puede tener elementos restantes mientras que `right_part` no tiene ninguno, y viceversa.

Crea otro ciclo `while` para copiar los elementos restantes en `left_part` dentro de la lista `array`. Usa `left_array_index < len(left_part)` para la condición del `while`.

# --hints--

Tu ciclo `while` debe tener esta condición: `left_array_index < len(left_part)`. No olvides usar la palabra clave `pass` en el cuerpo del ciclo.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):/m);
  }
})
```

Debe tener la palabra clave `pass` en el cuerpo de tu bucle `while`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
