---
id: 657f59751d5a5c9b069ae0f3
title: Paso 28
challengeType: 20
dashedName: step-28
---

# --description--

Puedes usar la variable `__name__` para determinar si un script de Python se está ejecutando como el programa principal o si se está importando como un módulo (código escrito en otro archivo de Python).

Si el valor de `__name__` está establecido en `'__main__'`, implica que el script actual es el programa principal, y no un módulo.

En este proyecto, usarás el script actual como el programa principal.

Crea una declaración `if` que verifique si el valor de `__name__` es `'__main__'`.

Utilice la palabra clave `pass` en el cuerpo de la declaración `if`.

# --hints--

Tu declaración `if` debería verificar si `__name__` tiene un valor de `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
