---
id: 68ed661e9f463c155bb677a1
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з гнучким контейнером CSS, створюючи множину кольорових блоків.

Для початку, всередині елемента `body` додайте елемент `header`. Вкладіть елемент `h1` всередину елемента `header` з текстом `Colored Boxes Layout`.

# --hints--

У вас має бути елемент `header` всередині елемента `body`.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

Ваш елемент `header` має містити всередині елемент `h1`.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Ваш елемент `h1` має містити текст `Colored Boxes Layout`.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
