---
id: 68ef0b3237fef991c92d9c19
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Тепер додайте елемент `div` з класом `flex-container` під вашим `header`. Потім створіть шість елементів `div` з класом `box` всередині `div`, який ви щойно створили.

Пам’ятайте, що атрибут класу часто використовується для вказівки імені класу у стилях. У цьому випадку всі елементи `div` будуть стилізовані однаково відповідно до визначення стилів `.flex-container` і `.box`.

# --hints--

У вас має бути елемент `div` під елементом `header`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

Ваш елемент `div` під елементом `header` має мати клас `flex-container`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

Усередині елемента `div` має бути шість вкладених елементів `.flex-container`.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

Всі шість елементів `div` всередині елемента `.flex-container` мають мати клас `box`.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
