---
id: 68f46b758eb8ea7043abd73c
title: Крок 33
challengeType: 0
dashedName: step-33
---

# --description--

Тепер ви організуєте блоки всередині свого гнучкого контейнера. Спочатку створіть селектор `.box1` і додайте до нього властивість `background` зі значенням `#f16e79`.

Потім додайте властивість `order` зі значенням `1`. Ви побачите, як перший блок переміститься в кінець. Це відбувається тому, що клас `box` вже має значення `order` зі значенням `0`, і елементи з більшими значеннями order з’являються пізніше.

Також додайте властивість `flex-grow` зі значенням `1`. Це дозволить блоку розширюватися, щоб заповнити додатковий простір на своїй лінії.

# --hints--

Ви повинні мати селектор `.box1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box1'));
```

Ваш селектор `.box1` має мати властивість `background` зі значенням `#f16e79`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('background'), 'rgb(241, 110, 121)');
```

Ваш селектор `.box1` має мати властивість `order` зі значенням `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('order'), '1');
```

Ваш селектор `.box1` має мати властивість `flex-grow` зі значенням `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--
```
