---
id: 68f91df5776e693049fb8942
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Тепер настав час додати кілька CSS правил у файл `styles.css`. Почніть зі створення селектора для елемента `h1`.

Вирівняйте вміст елемента `h1` по центру, встановивши його властивість `text-align` у значення `center`. Також створіть властивість `margin-bottom` зі значенням `10px`, щоб додати поле між низом елемента `h1` та будь-яким HTML елементом, який розташовується нижче.

# --hints--

Ви повинні мати селектор `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Ваш селектор `h1` має встановити властивість `text-align` у значення `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

Ваш селектор `h1` має встановити властивість `margin-bottom` у значення `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
