---
id: 692f22933ba1cb86cd75a70c
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Тепер настав час створити назву кольору для блоків. Для цього додайте елемент `p` під кожним елементом `h2` з такими назвами кольорів у порядку: `Red`, `Orange`, `Yellow`, `Green`, `Blue` та `Indigo`. По одному кольору для кожного блоку.

# --hints--

Ви повинні додати елемент `p` з текстом `Red` під елементом `h2` всередині першого елемента `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

Ви повинні додати елемент `p` з текстом `Orange` під елементом `h2` всередині другого елемента `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

Ви повинні додати елемент `p` з текстом `Yellow` під елементом `h2` всередині третього елемента `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

Ви повинні додати елемент `p` з текстом `Green` під елементом `h2` всередині четвертого елемента `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

Ви повинні додати елемент `p` з текстом `Blue` під елементом `h2` всередині п’ятого елемента `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

Ви повинні додати елемент `p` з текстом `Indigo` під елементом `h2` всередині шостого елемента `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
