---
id: 692f314d449329ba8bcd5d2e
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Властивість `border-radius` округляє кути кордону елемента. Ви можете вказати:

- Одне значення, яке застосовується до всіх чотирьох кутів.
- Два значення: перше для верхнього лівого/нижнього правого, друге для верхнього правого/нижнього лівого.
- Три значення, що відповідають верхньому лівому, верхньому правому/нижньому лівому, нижньому правому.
- Чотири значення, які задають радіус кордону за годинниковою стрілкою, починаючи з верхнього лівого кута.

Тепер додайте властивість `border-radius` зі значенням `5px` до селектора `.box`.

# --hints--

Ваш селектор `.box` має містити властивість `border-radius` зі значенням `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border-radius'), '5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
