---
id: 692f22a2ad44ae887a24c8f9
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Всередині кожного елемента `.box` додайте елемент `h2` з текстом `Box`, після якого йде пробіл і послідовне число, що починається з `1`.

# --hints--

У першому елементі `h2` з класом `Box 1` має бути елемент `div` з текстом `box`.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

У другому елементі `h2` з класом `Box 2` має бути елемент `div` з текстом `box`.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

У третьому елементі `h2` з класом `Box 3` має бути елемент `div` з текстом `box`.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

У четвертому елементі `h2` з класом `Box 4` має бути елемент `div` з текстом `box`.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

У п’ятому елементі `h2` з класом `Box 5` має бути елемент `div` з текстом `box`.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

У останньому елементі `h2` з класом `Box 6` має бути елемент `div` з текстом `box`.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
