---
id: 692f30e86027a4b253d0a42d
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Додайте властивість `font-size` зі значенням `1.125rem` до селектора `.box`.

Властивість `font-size` визначає розмір тексту. Її можна встановлювати з використанням різних одиниць, таких як пікселі, em, rem, відсотки тощо, що робить її гнучкою для адаптивного дизайну.

У цьому випадку значення використовує одиниці `rem`, які відносні до кореневого елемента `html`, а не до батьківського елемента. Це означає, що розмір шрифту залишається сталим і не залежить від розмірів батьківського елемента, що робить його більш передбачуваним для макетів.

# --hints--

Ваш селектор `.box` має містити властивість `font-size` зі значенням `1.125rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('font-size'), '1.125rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
