---
id: 692f2d37f34ed866ed42fcd6
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Властивість `flex-direction` контролює, як гнучкі елементи розташовуються всередині свого контейнера. Значення за замовчуванням — `row`, яке розміщує елементи горизонтально зліва направо. Ця властивість працює лише з елементами, що мають `display: flex` або `display: inline-flex`.

Зі значенням `column` елементи укладаються вертикально зверху вниз. Інші можливі значення — `row`, `row-reverse` та `column-reverse`.

Тепер додайте властивість `flex-direction` зі значенням `column` до селектора `.box`.

# --hints--

Ваш селектор `.box` має містити властивість `flex-direction` зі значенням `column`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex-direction'), 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
