---
id: 694b46a414bb2f5079868ed1
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Властивість `flex` контролює розмір і поведінку елементів всередині гнучкого контейнера. Вона складається з трьох властивостей: `flex-grow`, `flex-shrink` та `flex-basis`.

Тепер додайте властивість `flex` зі значенням `1 1 100px` до селектора `.box`. Це змусить кожен блок починатися з `100px`, дозволяючи їм розширюватися, щоб заповнити доступний простір, або звужуватися за потреби.

# --hints--

Ваш селектор `.box` має містити властивість `flex` зі значенням `1 1 100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex'), '1 1 100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
