---
id: 5a24c314108439a4d4036159
title: Den Spread-Operator auf Arrays anwenden
challengeType: 6
forumTopicId: 301452
dashedName: use-the-spread-operator-on-arrays
---

# --description--

Eine Lösung aus ES6, um die Unveränderlichkeit des Zustands in Redux durchzusetzen, ist der Spread-Operator: `...`. Der Spread-Operator hat verschiedene Anwendungsbereiche, einer davon passt gut zur vorherigen Aufgabe, ein neues Array aus einem bestehenden Array zu erzeugen. Diese Syntax ist relativ neu, aber weit verbreitet. Zum Beispiel, wenn Sie ein Array `myArray` haben und schreiben:

```js
let newArray = [...myArray];
```

`newArray` ist jetzt eine Kopie von `myArray`. Beide Arrays existieren weiterhin getrennt im Speicher. Wenn Sie eine Mutation wie `newArray.push(5)` durchführen, ändert sich `myArray` nicht. Der `...` *breitet* effektiv die Werte in `myArray` in ein neues Array aus. Um ein Array zu klonen und gleichzeitig zusätzliche Werte im neuen Array hinzuzufügen, könnten Sie `[...myArray, 'new value']` schreiben. Das würde ein neues Array zurückgeben, das aus den Werten in `myArray` und dem String `new value` als letztem Wert besteht. Die Spread-Syntax kann mehrfach bei der Array-Zusammensetzung verwendet werden, aber es ist wichtig zu beachten, dass sie nur eine flache Kopie des Arrays erstellt. Das heißt, sie bietet nur unveränderliche Array-Operationen für eindimensionale Arrays.

# --instructions--

Verwenden Sie den Spread-Operator, um eine neue Kopie des Zustands zurückzugeben, wenn eine To-do hinzugefügt wird.

# --hints--

Der Redux-Store sollte existieren und mit einem Zustand initialisiert werden, der gleich `["Do not mutate state!"]` ist.

```js
assert(
  (function () {
    const initialState = store.getState();
    return (
      Array.isArray(initialState) === true &&
      initialState[0] === 'Do not mutate state!'
    );
  })()
);
```

`addToDo` und `immutableReducer` sollten beide Funktionen sein.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Das Dispatchen einer Aktion vom Typ `ADD_TO_DO` im Redux-Store sollte ein `todo`-Element hinzufügen und den Zustand NICHT mutieren.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = ['Do not mutate state!', '__TEST__TO__DO__'];
assert(isFrozen)
assert.deepEqual(finalState, expectedState);
```

Der Spread-Operator sollte verwendet werden, um den neuen Zustand zurückzugeben.

```js
assert(code.includes('...state'));
```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      return [
        ...state,
        action.todo
      ];
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
