---
id: 5a24c314108439a4d4036154
title: Mehrere Reducer kombinieren
challengeType: 6
forumTopicId: 301436
dashedName: combine-multiple-reducers
---

# --description--

Wenn der Zustand Ihrer Anwendung komplexer wird, kann es verlockend sein, den Zustand in mehrere Teile aufzuteilen. Denken Sie jedoch an das erste Prinzip von Redux: Der gesamte Anwendungszustand wird in einem einzigen Zustandsobjekt im Store gehalten. Daher bietet Redux die Komposition von Reducern als Lösung für ein komplexes Zustandsmodell an. Sie definieren mehrere Reducer, die verschiedene Teile des Zustands Ihrer Anwendung verwalten, und setzen diese Reducer dann zu einem Root-Reducer zusammen. Der Root-Reducer wird anschließend an die Redux `createStore()`-Methode übergeben.

Um uns das Kombinieren mehrerer Reducer zu ermöglichen, stellt Redux die `combineReducers()`-Methode bereit. Diese Methode akzeptiert ein Objekt als Argument, in dem Sie Eigenschaften definieren, die Schlüssel mit bestimmten Reducer-Funktionen verknüpfen. Die Namen, die Sie den Schlüsseln geben, werden von Redux als Name für den zugehörigen Zustandsteil verwendet.

Typischerweise ist es eine gute Praxis, für jeden Zustandsteil der Anwendung einen eigenen Reducer zu erstellen, wenn diese sich in irgendeiner Weise unterscheiden oder einzigartig sind. Zum Beispiel könnte in einer Notiz-App mit Benutzer-Authentifizierung ein Reducer die Authentifizierung verwalten, während ein anderer den Text und die Notizen verwaltet, die der Benutzer absendet. Für eine solche Anwendung könnten wir die `combineReducers()`-Methode so schreiben:

```js
const rootReducer = Redux.combineReducers({
  auth: authenticationReducer,
  notes: notesReducer
});
```

Nun enthält der Schlüssel `notes` den gesamten Zustand, der mit unseren Notizen verbunden ist und von unserem `notesReducer` verwaltet wird. So können mehrere Reducer zusammengesetzt werden, um komplexeren Anwendungszustand zu verwalten. In diesem Beispiel wäre der Zustand im Redux-Store dann ein einzelnes Objekt, das die Eigenschaften `auth` und `notes` enthält.

# --instructions--

Im Code-Editor sind die Funktionen `counterReducer()` und `authReducer()` sowie ein Redux-Store vorhanden. Beenden Sie die Implementierung der `rootReducer()`-Funktion mithilfe der `Redux.combineReducers()`-Methode. Weisen Sie `counterReducer` einem Schlüssel namens `count` und `authReducer` einem Schlüssel namens `auth` zu.

# --hints--

Der `counterReducer` sollte den `state` inkrementieren und dekrementieren.

```js
assert(
  (function () {
    const initialState = store.getState().count;
    store.dispatch({ type: INCREMENT });
    store.dispatch({ type: INCREMENT });
    const firstState = store.getState().count;
    store.dispatch({ type: DECREMENT });
    const secondState = store.getState().count;
    return firstState === initialState + 2 && secondState === firstState - 1;
  })()
);
```

Der `authReducer` sollte die `state` von `authenticated` zwischen `true` und `false` umschalten.

```js
assert(
  (function () {
    store.dispatch({ type: LOGIN });
    const loggedIn = store.getState().auth.authenticated;
    store.dispatch({ type: LOGOUT });
    const loggedOut = store.getState().auth.authenticated;
    return loggedIn === true && loggedOut === false;
  })()
);
```

Das Store-`state` sollte zwei Schlüssel haben: `count`, der eine Zahl hält, und `auth`, der ein Objekt hält. Das `auth`-Objekt sollte eine Eigenschaft `authenticated` haben, die einen Boolean hält.

```js
assert(
  (function () {
    const state = store.getState();
    return (
      typeof state.auth === 'object' &&
      typeof state.auth.authenticated === 'boolean' &&
      typeof state.count === 'number'
    );
  })()
);
```

Der `rootReducer` sollte eine Funktion sein, die den `counterReducer` und den `authReducer` kombiniert.

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        typeof rootReducer === 'function' &&
        noWhiteSpace.includes('Redux.combineReducers')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = // Define the root reducer here

const store = Redux.createStore(rootReducer);
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = Redux.combineReducers({
  count: counterReducer,
  auth: authReducer
});

const store = Redux.createStore(rootReducer);
```
