---
id: 5a24c314108439a4d403615b
title: Ein Objekt mit Object.assign kopieren
challengeType: 6
forumTopicId: 301437
dashedName: copy-an-object-with-object-assign
---

# --description--

Die letzten Aufgaben arbeiteten mit Arrays, aber es gibt auch Möglichkeiten, die Unveränderlichkeit des Zustands zu gewährleisten, wenn der Zustand ein `object` ist. Ein nützliches Werkzeug zum Umgang mit Objekten ist das `Object.assign()`-Utility. `Object.assign()` nimmt ein Zielobjekt und Quellobjekte und weist Eigenschaften von den Quellobjekten dem Zielobjekt zu. Übereinstimmende Eigenschaften werden durch die Eigenschaften der Quellobjekte überschrieben. Dieses Verhalten wird häufig verwendet, um flache Kopien von Objekten zu erstellen, indem ein leeres Objekt als erstes Argument übergeben wird, gefolgt von dem oder den Objekten, die Sie kopieren möchten. Hier ist ein Beispiel:

```js
const newObject = Object.assign({}, obj1, obj2);
```

Dies erstellt `newObject` als neues `object`, das die Eigenschaften enthält, die derzeit in `obj1` und `obj2` existieren.

# --instructions--

Der Redux-Zustand und die Aktionen wurden angepasst, um ein `object` für das `state` zu verarbeiten. Bearbeiten Sie den Code so, dass für Aktionen mit dem Typ `state` ein neues `ONLINE`-Objekt zurückgegeben wird, das die Eigenschaft `status` auf den String `online` setzt. Versuchen Sie, `Object.assign()` zu verwenden, um die Aufgabe zu lösen.

# --hints--

Der Redux-Store sollte existieren und mit einem Zustand initialisiert werden, der dem auf Zeile 1 deklarierten `defaultState`-Objekt entspricht.

```js
const expectedState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};
const initialState = store.getState();
assert.deepEqual(expectedState, initialState);
```

`wakeUp` und `immutableReducer` sollten beide Funktionen sein.

```js
assert(typeof wakeUp === 'function' && typeof immutableReducer === 'function');
```

Das Dispatchen einer Aktion vom Typ `ONLINE` sollte die Eigenschaft `status` im Zustand auf `online` aktualisieren und den Zustand NICHT verändern.

```js

const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch({ type: 'ONLINE' });
const finalState = store.getState();
const expectedState = {
  user: 'CamperBot',
  status: 'online',
  friends: '732,982',
  community: 'freeCodeCamp'
};
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

`Object.assign` sollte verwendet werden, um den neuen Zustand zurückzugeben.

```js
assert(code.includes('Object.assign'));
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      return Object.assign({}, state, {
        status: 'online'
      });
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```
