---
id: 5a24c314108439a4d4036152
title: Verwenden Sie const für Action-Typen
challengeType: 6
forumTopicId: 301450
dashedName: use-const-for-action-types
---

# --description--

Eine gängige Praxis bei der Arbeit mit Redux ist es, Action-Typen als schreibgeschützte Konstanten zuzuweisen und diese Konstanten dann überall dort zu referenzieren, wo sie verwendet werden. Sie können den Code, an dem Sie arbeiten, so umgestalten, dass die Action-Typen als `const` Deklarationen geschrieben werden.

# --instructions--

Deklarieren Sie `LOGIN` und `LOGOUT` als `const` Werte und weisen Sie ihnen die Strings `'LOGIN'` bzw. `'LOGOUT'` zu. Bearbeiten Sie dann den `authReducer()` und die Action-Creators so, dass sie diese Konstanten anstelle von String-Werten referenzieren.

**Hinweis:** Es ist allgemein eine Konvention, Konstanten komplett in Großbuchstaben zu schreiben, und dies ist auch eine Standardpraxis in Redux.

# --hints--

Das Aufrufen der Funktion `loginUser` sollte ein Objekt zurückgeben, dessen `type` Eigenschaft auf den String `LOGIN` gesetzt ist.

```js
assert(loginUser().type === 'LOGIN');
```

Das Aufrufen der Funktion `logoutUser` sollte ein Objekt zurückgeben, dessen `type` Eigenschaft auf den String `LOGOUT` gesetzt ist.

```js
assert(logoutUser().type === 'LOGOUT');
```

Der Store sollte mit einem Objekt initialisiert werden, dessen Eigenschaft `login` auf `false` gesetzt ist.

```js
assert(store.getState().authenticated === false);
```

Das Dispatchen von `loginUser` sollte die `login` Eigenschaft im Store-State auf `true` aktualisieren.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

Das Dispatchen von `logoutUser` sollte die `login` Eigenschaft im Store-State auf `false` aktualisieren.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

Die `authReducer` Funktion sollte mehrere Action-Typen mit einer switch-Anweisung behandeln.

```js
assert(
    (function () {
      return (
        typeof authReducer === 'function' &&
        code.toString().includes('switch') &&
        code.toString().includes('case') &&
        code.toString().includes('default')
      );
    })()
  );
```

`LOGIN` und `LOGOUT` sollten als `const` Werte deklariert werden und die Strings `LOGIN` und `LOGOUT` zugewiesen bekommen.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(LOGIN === 'LOGIN' && LOGOUT === 'LOGOUT')
assert(noWhiteSpace.includes('const'))
```

Die Action-Creators und der Reducer sollten die Konstanten `LOGIN` und `LOGOUT` referenzieren.

```js
  assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(
        code.toString()
      );
      return (
        noWhiteSpace.includes('caseLOGIN:') &&
        noWhiteSpace.includes('caseLOGOUT:') &&
        noWhiteSpace.includes('type:LOGIN') &&
        noWhiteSpace.includes('type:LOGOUT')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js


const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {
    case 'LOGIN': 
      return {
        authenticated: true
      }
    case 'LOGOUT': 
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case LOGIN:
      return {
        authenticated: true
      }

    case LOGOUT:
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: LOGIN
  }
};

const logoutUser = () => {
  return {
    type: LOGOUT
  }
};
```
