---
id: 5a24c314108439a4d4036157
title: Einen Zähler mit Redux schreiben
challengeType: 6
forumTopicId: 301453
dashedName: write-a-counter-with-redux
---

# --description--

Jetzt haben Sie alle Grundsätze von Redux gelernt! Sie haben gesehen, wie man Aktionen und Action Creators erstellt, einen Redux Store anlegt, Aktionen an den Store dispatcht und Zustandsaktualisierungen mit reinen Reducern gestaltet. Sie haben sogar gesehen, wie man komplexen Zustand mit Reducer-Komposition verwaltet und asynchrone Aktionen behandelt. Diese Beispiele sind einfach gehalten, aber diese Konzepte sind die Grundsätze von Redux. Wenn Sie sie gut verstehen, sind Sie bereit, Ihre eigene Redux-App zu bauen. Die nächsten Aufgaben behandeln einige Details zur `state` Unveränderlichkeit, aber zuerst hier eine Übersicht über alles, was Sie bisher gelernt haben.

# --instructions--

In dieser Lektion implementieren Sie einen einfachen Zähler mit Redux von Grund auf. Die Grundlagen sind im Code-Editor vorgegeben, aber Sie müssen die Details ausfüllen! Verwenden Sie die vorgegebenen Namen und definieren Sie `incAction` und `decAction` Action Creators, die `counterReducer()`, `INCREMENT` und `DECREMENT` Aktionstypen und schließlich den Redux `store`. Sobald Sie fertig sind, sollten Sie in der Lage sein, `INCREMENT` oder `DECREMENT` Aktionen zu dispatchen, um den Zustand im `store` zu erhöhen oder zu verringern. Viel Erfolg beim Erstellen Ihrer ersten Redux-App!

# --hints--

Der Action Creator `incAction` sollte ein Aktionsobjekt zurückgeben, bei dem `type` den Wert von `INCREMENT` hat.

```js
assert(incAction().type === INCREMENT);
```

Der Action Creator `decAction` sollte ein Aktionsobjekt zurückgeben, bei dem `type` den Wert von `DECREMENT` hat.

```js
assert(decAction().type === DECREMENT);
```

Das Ausführen von `store.getState()` sollte eine Zahl zurückgeben.

```js
assert(typeof store.getState() === 'number');
```

Der Redux Store sollte mit einem `state` von 0 initialisiert werden.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    return _store.getState() === 0;
  })()
);
```

Das Dispatchen von `incAction` auf den Redux Store sollte den `state` um 1 erhöhen.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(incAction());
    const incState = _store.getState();
    return initialState + 1 === incState;
  })()
);
```

Das Dispatchen von `decAction` auf den Redux Store sollte den `state` um 1 verringern.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(decAction());
    const decState = _store.getState();
    return initialState - 1 === decState;
  })()
);
```

`counterReducer` sollte eine Funktion sein.

```js
assert(typeof counterReducer === 'function');
```

# --seed--

## --seed-contents--

```js
const INCREMENT = null; // Define a constant for increment action types
const DECREMENT = null; // Define a constant for decrement action types

const counterReducer = null; // Define the counter reducer which will increment or decrement the state based on the action it receives

const incAction = null; // Define an action creator for incrementing

const decAction = null; // Define an action creator for decrementing

const store = null; // Define the Redux store here, passing in your reducers
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const incAction = () => {
  return {
    type: INCREMENT
  }
};

const decAction = () => {
  return {
    type: DECREMENT
  }
};

const store = Redux.createStore(counterReducer);
```
