---
id: 5a24c314108439a4d403614c
title: Status aus dem Redux-Speichern abrufen
challengeType: 6
forumTopicId: 301443
dashedName: get-state-from-the-redux-store
---

# --description--

Das Redux-Speicherobjekt bietet mehrere Methoden, mit denen Sie interagieren können. Zum Beispiel können Sie den aktuellen `state` abrufen, der im Redux-Speicherobjekt mit der `getState()`-Methode gehalten wird.

# --instructions--

Der Code aus der vorherigen Aufgabe ist im Code-Editor kürzer neu geschrieben. Verwenden Sie `store.getState()`, um den `state` aus dem `store` abzurufen, und weisen Sie diesen einer neuen Variablen `currentState` zu.

# --hints--

Der Redux-Speicher sollte für den Anfangszustand den Wert 5 haben.

```js
assert(store.getState() === 5);
```

Eine Variable `currentState` sollte existieren und den aktuellen Status des Redux-Speichers zugewiesen bekommen.

```js
assert(currentState === 5 && code.includes('store.getState()'));
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
```

# --solutions--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
const currentState = store.getState();
```
