---
id: 5a24c314108439a4d4036158
title: Zustand niemals verändern
challengeType: 6
forumTopicId: 301445
dashedName: never-mutate-state
---

# --description--

Diese abschließenden Aufgaben beschreiben mehrere Methoden, um das wichtige Prinzip der Unveränderlichkeit des Zustands in Redux durchzusetzen. Unveränderlicher Zustand bedeutet, dass Sie den Zustand niemals direkt verändern, sondern stattdessen eine neue Kopie des Zustands zurückgeben.

Wenn Sie im Zeitverlauf einen Schnappschuss des Zustands einer Redux-App machen würden, sähe das etwa so aus `state 1`, `state 2`, `state 3`, `state 4`, `...` und so weiter, wobei jeder Zustand dem vorherigen ähnlich sein kann, aber jeder ein eigenständiges Datenstück ist. Diese Unveränderlichkeit ist tatsächlich die Grundlage für Funktionen wie das Time-Travel-Debugging, von dem Sie vielleicht schon gehört haben.

Redux erzwingt die Unveränderlichkeit des Zustands in seinem Store oder in den Reducern nicht aktiv, diese Verantwortung liegt beim Programmierer. Glücklicherweise bietet JavaScript (insbesondere ES6) mehrere nützliche Werkzeuge, mit denen Sie die Unveränderlichkeit Ihres Zustands durchsetzen können, egal ob es sich um ein `string`, `number`, `array` oder `object` handelt. Beachten Sie, dass Strings und Zahlen primitive Werte sind und von Natur aus unveränderlich sind. Anders gesagt: 3 ist immer 3. Sie können den Wert der Zahl 3 nicht ändern. Ein `array` oder `object` hingegen ist veränderlich. In der Praxis wird Ihr Zustand wahrscheinlich aus einem `array` oder `object` bestehen, da diese nützliche Datenstrukturen sind, um viele Arten von Informationen darzustellen.

# --instructions--

Im Code-Editor gibt es ein `store` und ein `reducer` zur Verwaltung von To-do-Elementen. Beenden Sie das Schreiben des `ADD_TO_DO`-Falls im Reducer, um ein neues To-do an den Zustand anzuhängen. Es gibt mehrere Möglichkeiten, dies mit Standard-JavaScript oder ES6 zu erreichen. Versuchen Sie, eine Methode zu finden, die ein neues Array zurückgibt, an dessen Ende das Element aus `action.todo` angehängt ist.

# --hints--

Der Redux-Store sollte existieren und mit einem Zustand initialisiert werden, der dem `todos`-Array im Code-Editor entspricht.

```js
assert(
  (function () {
    const todos = [
      'Go to the store',
      'Clean the house',
      'Cook dinner',
      'Learn to code'
    ];
    const initialState = store.getState();
    return (
      Array.isArray(initialState) && initialState.join(',') === todos.join(',')
    );
  })()
);
```

`addToDo` und `immutableReducer` sollten beide Funktionen sein.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Das Dispatchen einer Aktion vom Typ `ADD_TO_DO` im Redux-Store sollte ein `todo`-Element hinzufügen und den Zustand NICHT verändern.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
  '__TEST__TO__DO__'
];
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

# --seed--

## --seed-contents--

```js
const ADD_TO_DO = 'ADD_TO_DO';

// A list of strings representing tasks to do:
const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const ADD_TO_DO = 'ADD_TO_DO';

const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      return state.concat(action.todo);
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
