---
id: 5a24c314108439a4d4036151
title: "Eine switch-Anweisung verwenden, um mehrere Aktionen zu verarbeiten"
challengeType: 6
forumTopicId: 301449
dashedName: use-a-switch-statement-to-handle-multiple-actions
---

# --description--

Sie können dem Redux-Store mitteilen, wie er mit mehreren Aktionstypen umgehen soll. Angenommen, Sie verwalten die Benutzer-Authentifizierung in Ihrem Redux-Store. Sie möchten einen Zustandsdarstellung haben, wenn Benutzer eingeloggt sind und wenn sie ausgeloggt sind. Sie stellen dies mit einem einzelnen Zustandsobjekt mit der Eigenschaft `authenticated` dar. Sie benötigen außerdem Action Creators, die Aktionen für Benutzer-Login und Benutzer-Logout erzeugen, zusammen mit den Aktionsobjekten selbst.

# --instructions--

Der Code-Editor hat einen Store, Aktionen und Action Creators für Sie eingerichtet. Füllen Sie die `reducer`-Funktion aus, um mehrere Authentifizierungsaktionen zu verarbeiten. Verwenden Sie eine JavaScript-`switch`-Anweisung im `reducer`, um auf verschiedene Aktionsereignisse zu reagieren. Dies ist ein Standardmuster beim Schreiben von Redux-Reducer. Die switch-Anweisung sollte über `action.type` wechseln und den entsprechenden Authentifizierungszustand zurückgeben.

**Hinweis:** Machen Sie sich an dieser Stelle keine Sorgen um die Unveränderlichkeit des Zustands, da dieser in diesem Beispiel klein und einfach ist. Für jede Aktion können Sie ein neues Objekt zurückgeben – zum Beispiel `{authenticated: true}`. Vergessen Sie auch nicht, einen `default`-Fall in Ihrer switch-Anweisung zu schreiben, der den aktuellen `state` zurückgibt. Das ist wichtig, weil, sobald Ihre App mehrere Reducer hat, alle ausgeführt werden, wenn eine Aktion dispatcht wird, auch wenn die Aktion nicht mit diesem Reducer zusammenhängt. In einem solchen Fall möchten Sie sicherstellen, dass Sie den aktuellen `state` zurückgeben.

# --hints--

Das Aufrufen der Funktion `loginUser` sollte ein Objekt mit der Eigenschaft type zurückgeben, die auf den String `LOGIN` gesetzt ist.

```js
assert(loginUser().type === 'LOGIN');
```

Das Aufrufen der Funktion `logoutUser` sollte ein Objekt mit der Eigenschaft type zurückgeben, die auf den String `LOGOUT` gesetzt ist.

```js
assert(logoutUser().type === 'LOGOUT');
```

Der Store sollte mit einem Objekt initialisiert werden, das eine `authenticated`-Eigenschaft mit dem Wert `false` hat.

```js
assert(store.getState().authenticated === false);
```

Das Dispatchen von `loginUser` sollte die `authenticated`-Eigenschaft im Store-Zustand auf `true` aktualisieren.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

Das Dispatchen von `logoutUser` sollte die `authenticated`-Eigenschaft im Store-Zustand auf `false` aktualisieren.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

Die `authReducer`-Funktion sollte mehrere Aktionstypen mit einer `switch`-Anweisung verarbeiten.

```js
assert(
    code.toString().includes('switch') &&
      code.toString().includes('case') &&
      code.toString().includes('default')
  );
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case 'LOGIN':
      return {
        authenticated: true
      }

    case 'LOGOUT':
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```
