---
id: 5a24c314108439a4d403614e
title: Einen Action Creator definieren
challengeType: 6
forumTopicId: 301441
dashedName: define-an-action-creator
---

# --description--

Nachdem Sie eine Action erstellt haben, besteht der nächste Schritt darin, die Action an das Redux speichern zu senden, damit es seinen Zustand aktualisieren kann. In Redux definieren Sie Action Creators, um dies zu erreichen. Ein Action Creator ist einfach eine JavaScript-Funktion, die eine Action zurückgibt. Mit anderen Worten: Action Creators erstellen Objekte, die Action-Ereignisse darstellen.

# --instructions--

Definieren Sie eine Funktion namens `actionCreator()`, die das `action`-Objekt zurückgibt, wenn sie aufgerufen wird.

# --hints--

Die Funktion `actionCreator` sollte existieren.

```js
assert(typeof actionCreator === 'function');
```

Das Ausführen der Funktion `actionCreator` sollte das `action`-Objekt zurückgeben.

```js
assert(typeof action === 'object');
```

Das zurückgegebene `action` sollte eine Schlüssel-Eigenschaft `type` mit dem Wert `LOGIN` haben.

```js
assert(action.type === 'LOGIN');
```

# --seed--

## --seed-contents--

```js
const action = {
  type: 'LOGIN'
}
// Define an action creator here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
const actionCreator = () => {
  return action;
};
```
