---
id: 5a24c314108439a4d4036153
title: Einen Store-Listener registrieren
challengeType: 6
forumTopicId: 301446
dashedName: register-a-store-listener
---

# --description--

Eine weitere Methode, auf die Sie beim Redux `store` Objekt zugreifen können, ist `store.subscribe()`. Damit können Sie Listener-Funktionen beim Store abonnieren, die aufgerufen werden, sobald eine Aktion an den Store gesendet wird. Eine einfache Verwendung dieser Methode ist, eine Funktion beim Store zu abonnieren, die bei jedem Empfang einer Aktion und der Aktualisierung des Stores einfach eine Nachricht in der Konsole ausgibt.

# --instructions--

Schreiben Sie eine Callback-Funktion, die die globale Variable `count` jedes Mal erhöht, wenn der Store eine Aktion erhält, und übergeben Sie diese Funktion an die `store.subscribe()` Methode. Sie werden sehen, dass `store.dispatch()` dreimal hintereinander aufgerufen wird, wobei jedes Mal direkt ein Aktionsobjekt übergeben wird. Beobachten Sie die Konsolenausgabe zwischen den Aktions-Dispatches, um die Aktualisierungen zu sehen.

# --hints--

Das Dispatchen der `ADD` Aktion im Store sollte den Zustand um `1` erhöhen.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch({ type: 'ADD' });
    const newState = store.getState();
    return newState === initialState + 1;
  })()
);
```

Es sollte eine Listener-Funktion mit `store.subscribe` beim Store abonniert sein.

```js
assert.match(code, /store\s*\.\s*subscribe\(/gm);
```

`store.subscribe` sollte eine Funktion erhalten.

```js
assert.match(code, /(\s*function\s*)|(\s*\(\s*\)\s*=>)/gm);
```

Die an `store.subscribe` übergebene Funktion sollte nicht aufgerufen werden.

```js
assert.notMatch(code, /store\.subscribe\(.+\(\)\)/);
```

Der Callback an `store.subscribe` sollte ebenfalls die globale `count` Variable erhöhen, während der Store aktualisiert wird.

```js
assert.strictEqual(store.getState(), count);
```

# --seed--

## --seed-contents--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);

// Global count variable:
let count = 0;

// Change code below this line

// Change code above this line

store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
```

# --solutions--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);
 let count = 0;
// Change code below this line

store.subscribe( () =>
 {
 count++;
 }
);

// Change code above this line

store.dispatch({type: ADD});
store.dispatch({type: ADD});
store.dispatch({type: ADD});
```
