---
id: 5a24c314108439a4d4036156
title: "Middleware verwenden, um asynchrone Aktionen zu handhaben"
challengeType: 6
forumTopicId: 301451
dashedName: use-middleware-to-handle-asynchronous-actions
---

# --description--

Bisher haben diese Aufgaben vermieden, asynchrone Aktionen zu behandeln, aber sie sind ein unvermeidlicher Teil der Webentwicklung. Irgendwann müssen Sie asynchrone Endpunkte in Ihrer Redux-Anwendung aufrufen. Wie gehen Sie also mit diesen Arten von Anfragen um? Redux stellt Middleware bereit, die speziell für diesen Zweck entwickelt wurde, die Redux Thunk Middleware. Hier ist eine kurze Beschreibung, wie Sie diese mit Redux verwenden.

Um die Redux Thunk Middleware einzubinden, übergeben Sie sie als Argument an `Redux.applyMiddleware()`. Diese Anweisung wird dann als zweiter optionaler Parameter an die `createStore()`-Funktion übergeben. Schauen Sie sich den Code am unteren Rand des Editors an, um dies zu sehen. Um dann eine asynchrone Aktion zu erstellen, geben Sie in der Action-Creator-Funktion eine Funktion zurück, die `dispatch` als Argument erhält. Innerhalb dieser Funktion können Sie Aktionen dispatchen und asynchrone Anfragen ausführen.

In diesem Beispiel wird eine asynchrone Anfrage mit einem `setTimeout()`-Aufruf simuliert. Es ist üblich, vor dem Start einer asynchronen Aktion eine Aktion zu dispatchen, damit Ihr Anwendungs-state weiß, dass Daten angefordert werden (dieser state könnte beispielsweise ein Lade-Symbol anzeigen). Sobald Sie die Daten erhalten, dispatchen Sie eine weitere Aktion, die die Daten als Nutzlast zusammen mit der Information trägt, dass die Aktion abgeschlossen ist.

Denken Sie daran, dass Sie `dispatch` als Parameter an diesen speziellen Action-Creator übergeben. Diesen verwenden Sie, um Ihre Aktionen zu dispatchen: Sie übergeben die Aktion einfach direkt an dispatch, und die Middleware kümmert sich um den Rest.

# --instructions--

Schreiben Sie beide Dispatches im `handleAsync()`-Action-Creator. Dispatchen Sie `requestingData()` vor dem `setTimeout()` (dem simulierten API-Aufruf). Dann, nachdem Sie die (vorgeblichen) Daten erhalten haben, dispatchen Sie die `receivedData()`-Aktion und übergeben diese Daten. Nun wissen Sie, wie man asynchrone Aktionen in Redux handhabt. Alles andere verhält sich wie zuvor.

# --hints--

Der `requestingData`-Action-Creator sollte ein Objekt zurückgeben, dessen Typ dem Wert von `REQUESTING_DATA` entspricht.

```js
assert(requestingData().type === REQUESTING_DATA);
```

Der `receivedData`-Action-Creator sollte ein Objekt zurückgeben, dessen Typ dem Wert von `RECEIVED_DATA` entspricht.

```js
assert(receivedData('data').type === RECEIVED_DATA);
```

`asyncDataReducer` sollte eine Funktion sein.

```js
assert(typeof asyncDataReducer === 'function');
```

Das Dispatchen des `requestingData`-Action-Creators sollte die fetching-Eigenschaft im Store `state` auf `true` aktualisieren.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(requestingData());
    const reqState = store.getState();
    return initialState.fetching === false && reqState.fetching === true;
  })()
);
```

Das Dispatchen von `handleAsync` sollte die Datenanforderungsaktion dispatchen und dann nach einer Verzögerung die Aktion mit den empfangenen Daten dispatchen.

```js
assert(
  (function () {
    const noWhiteSpace = __helpers.removeWhiteSpace(handleAsync.toString());
    return (
      noWhiteSpace.includes('dispatch(requestingData())') === true &&
      noWhiteSpace.includes('dispatch(receivedData(data))') === true
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    // Dispatch request action here

    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      // Dispatch received data action here

    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```

# --solutions--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    dispatch(requestingData());
    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      dispatch(receivedData(data));
    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```
