---
id: 5a24c314108439a4d4036155
title: Aktionsdaten an das speichern senden
challengeType: 6
forumTopicId: 301448
dashedName: send-action-data-to-the-store
---

# --description--

Sie haben inzwischen gelernt, wie Sie Aktionen an das Redux speichern senden. Bisher enthielten diese Aktionen jedoch keine Informationen außer einem `type`. Sie können auch bestimmte Daten zusammen mit Ihren Aktionen senden. Das ist sogar sehr üblich, denn Aktionen entstehen meist durch eine Benutzerinteraktion und enthalten oft einige Daten. Das Redux speichern muss diese Daten häufig kennen.

# --instructions--

Im Code-Editor sind ein grundlegender `notesReducer()` und ein `addNoteText()` Action Creator definiert. Vervollständigen Sie den Körper der `addNoteText()`-Funktion, sodass sie ein `action`-Objekt zurückgibt. Das Objekt soll eine `type`-Eigenschaft mit dem Wert `ADD_NOTE` enthalten und außerdem eine `text`-Eigenschaft, die auf die `note`-Daten gesetzt ist, die an den Action Creator übergeben werden. Wenn Sie den Action Creator aufrufen, übergeben Sie spezifische Notizinformationen, auf die Sie im Objekt zugreifen können.

Vervollständigen Sie anschließend die `switch`-Anweisung in der `notesReducer()`. Sie müssen einen Fall hinzufügen, der `addNoteText()`-Aktionen behandelt. Dieser Fall soll ausgelöst werden, wenn eine Aktion vom Typ `ADD_NOTE` vorliegt, und er soll die `text`-Eigenschaft der eingehenden `action` als neuen `state` zurückgeben.

Die Aktion wird am Ende des Codes ausgelöst. Wenn Sie fertig sind, führen Sie den Code aus und beobachten Sie die Konsole. So einfach ist es, aktionsspezifische Daten an das speichern zu senden und sie beim Aktualisieren des `state`-Speichers zu verwenden.

# --hints--

Der Action Creator `addNoteText` sollte ein Objekt mit den Schlüsseln `type` und `text` zurückgeben.

```js
assert(
  (function () {
    const addNoteFn = addNoteText('__TEST__NOTE');
    return addNoteFn.type === ADD_NOTE && addNoteFn.text === '__TEST__NOTE';
  })()
);
```

Das Auslösen einer Aktion vom Typ `ADD_NOTE` mit dem `addNoteText`-Action Creator sollte den `state` auf den an den Action Creator übergebenen String aktualisieren.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(addNoteText('__TEST__NOTE'));
    const newState = store.getState();
    return initialState !== newState && newState === '__TEST__NOTE';
  })()
);
```

# --seed--

## --seed-contents--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line

    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello!'));
console.log(store.getState());
```

# --solutions--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line
    case ADD_NOTE:
      return action.text;
    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line
  return {
    type: ADD_NOTE,
    text: note
  }
  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello Redux!'));
console.log(store.getState());
```
