---
id: 564944c91be2204b269d51e3
title: jQuery를 사용해 요소 내부 텍스트 변경하기
challengeType: 6
forumTopicId: 16773
dashedName: change-text-inside-an-element-using-jquery
---

# --description--

jQuery를 사용하면 요소의 시작 태그와 끝 태그 사이에 있는 텍스트를 변경할 수 있습니다. HTML 마크업도 변경할 수 있습니다.

jQuery에는 `.html()`라는 함수가 있어 요소 내에 HTML 태그와 텍스트를 추가할 수 있습니다. 이 함수를 사용해 제공하는 내용으로 요소 내의 기존 모든 내용이 완전히 교체됩니다.

다음은 제목의 텍스트를 다시 작성하고 강조하는 방법입니다.

```js
$("h3").html("<em>jQuery Playground</em>");
```

jQuery에는 `.text()`라는 비슷한 함수도 있는데, 이 함수는 태그를 추가하지 않고 텍스트만 변경합니다. 즉, 이 함수는 전달된 HTML 태그를 평가하지 않고 기존 내용을 교체할 텍스트로 처리합니다.

id가 `target4`인 버튼의 텍스트를 강조해 변경하세요.

<a href="https://www.freecodecamp.org/news/html-elements-explained-what-are-html-tags/#em-element" target="_blank" rel="noopener noreferrer nofollow">우리 뉴스 기사에서 &lt;em&gt;</a> 태그와 `<i>`, `<em>`의 차이점과 용도를 확인하세요.

`<i>` 태그는 전통적으로 텍스트를 강조하는 데 사용되었지만, 현재는 아이콘 태그로도 채택되었습니다. `<em>` 태그가 이제 강조 태그로 널리 인정받고 있습니다. 이 도전 과제에서는 어느 쪽을 사용해도 됩니다.

# --hints--

`target4` 버튼의 텍스트를 HTML 태그를 추가해 강조하세요.

```js
assert.isTrue(
  /<em>|<i>\s*#target4\s*<\/em>|<\/i>/gi.test($('#target4').html())
);
```

텍스트는 그 외에는 변경하지 마세요.

```js
assert($('#target4') && $('#target4').text().trim() === '#target4');
```

다른 텍스트는 변경하지 않아야 합니다.

```js
assert.isFalse(/<em>|<i>/gi.test($('h3').html()));
```

`.html()`를 사용하고 `.text()`를 사용하지 마세요.

```js
assert(code.match(/\.html\(/g));
```

jQuery로 `button id="target4"`를 선택하세요.

```js
assert(code.match(/\$\(\s*?(\"|\')#target4(\"|\')\s*?\)\.html\(/));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target4").html('<em>#target4</em>');
  });
</script>

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
