---
id: bad87fee1348bd9aedc08826
title: jQuery로 클래스별 요소 선택하기
challengeType: 6
forumTopicId: 18316
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-class-using-jquery
---

# --description--

`button` 요소들이 모두 튕기도록 만든 것을 보셨나요? `$("button")`로 선택한 다음 `.addClass("animated bounce");`로 CSS 클래스를 추가했습니다.

지금 막 jQuery의 `.addClass()` 함수를 사용했는데, 이 함수는 요소에 클래스를 추가할 수 있게 해줍니다.

먼저 `div` 선택자를 사용해 클래스가 `well`인 `$(".well")` 요소들을 선택해 보세요.

CSS 선언과 마찬가지로 클래스 이름 앞에 `.`를 붙여야 한다는 점을 기억하세요.

그다음 jQuery의 `.addClass()` 함수를 사용해 `animated`와 `shake` 클래스를 추가하세요.

예를 들어, 클래스가 `text-primary`인 모든 요소를 흔들리게 하려면 `document ready function`에 다음을 추가할 수 있습니다:

```js
$(".text-primary").addClass("animated shake");
```

# --hints--

클래스가 `addClass()`인 모든 요소에 `animated`와 `shake` 클래스를 주려면 jQuery의 `well` 함수를 사용하세요.

```js
assert($('.well').hasClass('animated') && $('.well').hasClass('shake'));
```

이 클래스들은 요소에 jQuery로만 추가해야 합니다.

```js
assert(!code.match(/class\.\*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
