---
id: bad87fee1348bd9aeda08826
title: jQuery로 id 속성으로 요소 선택하기
challengeType: 6
forumTopicId: 18317
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-id-using-jquery
---

# --description--

id 속성으로도 요소를 선택할 수 있습니다.

먼저 `button` 요소를 `target3` id로 `$("#target3")` 선택자를 사용해 선택하세요.

CSS 선언과 마찬가지로 id 이름 앞에 `#`를 입력하는 점을 기억하세요.

그다음 jQuery의 `.addClass()` 함수를 사용해 `animated`와 `fadeOut` 클래스를 추가하세요.

다음은 `button` id를 가진 `target6` 요소를 서서히 사라지게(fade out) 만드는 방법입니다:

```js
$("#target6").addClass("animated fadeOut");
```

# --hints--

`button` `id`를 가진 `target3` 요소를 선택하고 jQuery `addClass()` 함수를 사용해 `animated` 클래스를 부여하세요.

```js
assert($('#target3').hasClass('animated'));
```

id가 `target3`인 요소를 선택하고 jQuery `addClass()` 함수를 사용해 `fadeOut` 클래스를 부여하세요.

```js
assert(
  ($('#target3').hasClass('fadeOut') || $('#target3').hasClass('fadeout')) &&
    code.match(/\$\(\s*.#target3.\s*\)/g)
);
```

이 클래스들은 jQuery로만 추가해야 합니다.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
    $("#target3").addClass("animated fadeOut");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
