---
id: bad87fee1348bd9acdd08826
title: 스크립트 태그와 문서 준비 작동 방식 배우기
challengeType: 6
forumTopicId: 18224
dashedName: learn-how-script-tags-and-document-ready-work
---

# --description--

이제 가장 인기 있는 JavaScript 도구인 jQuery를 배울 준비가 되었습니다.

jQuery를 사용하기 전에 HTML에 몇 가지를 추가해야 합니다.

먼저, 페이지 상단에 `script` 요소를 추가하세요. 다음 줄에서 닫는 태그를 꼭 작성하세요.

브라우저는 `script` 요소 안에 있는 모든 JavaScript를 실행하는데, 여기에는 jQuery도 포함됩니다.

`script` 요소 안에 이 코드를 추가하세요: `$(document).ready(function() {`를 `script`에 넣으세요. 그런 다음 다음 줄에서 (여전히 `script` 요소 안에) `});`로 닫으세요.

`functions`에 대해서는 나중에 더 배울 것입니다. 중요한 점은 이 `function` 안에 넣은 코드는 브라우저가 페이지를 불러오자마자 실행된다는 것입니다.

이것이 중요한 이유는 `document ready function`가 없으면 HTML이 렌더링되기 전에 코드가 실행되어 버그가 발생할 수 있기 때문입니다.

# --hints--

유효한 닫는 태그가 있는 `script` 요소를 만들어야 합니다.

```js
assert(
  code.match(/<\/script\s*>/g) &&
    code.match(
      /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
    ) &&
    code.match(/<\/script\s*>/g).length ===
      code.match(
        /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
      ).length
);
```

`$(document).ready(function() {` 요소의 시작 부분에 `script`를 추가해야 합니다.

```js
assert(
  code.match(
    /\$\s*?\(\s*?document\s*?\)\.ready\s*?\(\s*?function\s*?\(\s*?\)\s*?\{/g
  )
);
```

`$(document).ready(function() {` 함수는 `});`로 닫아야 합니다.

```js
assert(code.match(/\n*?\s*?\}\s*?\);/g));
```

# --seed--

## --seed-contents--

```html
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
  });
</script>
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
