---
id: bad87fee1348bd9aed908626
title: 여러 jQuery 선택자로 같은 요소를 선택하기
challengeType: 6
forumTopicId: 18322
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-the-same-element-with-multiple-jquery-selectors
---

# --description--

이제 요소를 선택하는 세 가지 방법을 알게 되었습니다: 타입별로 `$("button")`, 클래스별로 `$(".btn")`, 그리고 id별로 `$("#target1")`.

하나의 `.addClass()` 호출에서 여러 클래스를 추가할 수 있지만, 같은 요소에 *세 가지 별도의 방법*으로 클래스를 추가해 보세요.

`.addClass()`를 사용해 같은 요소에 한 번에 하나씩, 세 가지 다른 방법으로 클래스만 추가하세요:

타입이 `animated`인 모든 요소에 `button` 클래스를 추가하세요.

클래스가 `shake`인 모든 버튼에 `.btn` 클래스를 추가하세요.

id가 `btn-primary`인 버튼에 `#target1` 클래스를 추가하세요.

**참고:** 한 번에 하나의 요소만 선택하고 한 번에 하나의 클래스만 추가해야 합니다. 세 개의 개별 선택자는 결국 `shake`에 `animated`, `btn-primary`, `#target1` 세 클래스를 추가하게 됩니다.

# --hints--

코드는 `$("button")` 선택자를 사용해야 합니다.

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?button\s*?(?:'|")/gi));
```

코드는 `$(".btn")` 선택자를 사용해야 합니다.

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?\.btn\s*?(?:'|")/gi));
```

코드는 `$("#target1")` 선택자를 사용해야 합니다.

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?#target1\s*?(?:'|")/gi));
```

세 개의 선택자 각각으로 한 번에 하나의 클래스만 추가해야 합니다.

```js
assert(
  code.match(/addClass/g) &&
    code.match(/addClass\s*?\(\s*?('|")\s*?[\w-]+\s*?\1\s*?\)/g).length > 2
);
```

`#target1` 요소는 `animated`‚ `shake` 그리고 `btn-primary` 클래스를 가져야 합니다.

```js
assert(
  $('#target1').hasClass('animated') &&
    $('#target1').hasClass('shake') &&
    $('#target1').hasClass('btn-primary')
);
```

이 클래스들은 jQuery만 사용해 요소에 추가해야 합니다.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated");
    $(".btn").addClass("shake");
    $("#target1").addClass("btn-primary");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
