---
id: bad87fee1348bd9aed208826
title: jQuery를 사용해 요소의 자식 요소 선택하기
challengeType: 6
forumTopicId: 18320
dashedName: target-the-children-of-an-element-using-jquery
---

# --description--

HTML 요소가 다른 요소 바로 아래 한 단계에 위치하면 그 요소의 <dfn>자식</dfn>이라고 합니다. 예를 들어, 이 도전 과제에서 텍스트가 `#target1`, `#target2`, `#target3`인 버튼 요소들은 모두 `<div class="well" id="left-well">` 요소의 자식입니다.

jQuery에는 선택한 요소의 자식에 접근할 수 있는 `children()`라는 함수가 있습니다.

다음은 `children()` 함수를 사용해 `left-well` 요소의 자식들에게 `blue` 색상을 지정하는 예시입니다:

```js
$("#left-well").children().css("color", "blue")
```

# --instructions--

`right-well` 요소의 모든 자식에게 주황색을 지정하세요.

# --hints--

`#right-well`의 모든 자식은 주황색 텍스트여야 합니다.

```js
assert($('#right-well').children().css('color') === 'rgb(255, 165, 0)');
```

이 요소들을 수정하려면 `children()` 함수를 사용하세요.

```js
assert(code.match(/\.children\(\)\.css/g));
```

이 클래스들은 jQuery로만 요소에 추가해야 합니다.

```js
assert(code.match(/<div class="well" id="right-well">/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
    $("#right-well").children().css("color", "orange");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
