---
id: bad87fee1348bd9aed108826
title: jQuery를 사용해 특정 자식 요소 선택하기
challengeType: 6
forumTopicId: 18315
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-a-specific-child-of-an-element-using-jquery
---

# --description--

id 속성이 jQuery 선택자에서 얼마나 편리한지 보셨습니다. 하지만 항상 이렇게 깔끔한 id를 사용할 수 있는 것은 아닙니다.

다행히도 jQuery에는 올바른 요소를 선택할 수 있는 다른 방법들이 있습니다.

jQuery는 CSS 선택자를 사용해 요소를 선택합니다. `target:nth-child(n)` CSS 선택자를 사용하면 대상 클래스나 요소 유형을 가진 n번째 모든 요소를 선택할 수 있습니다.

각 well 요소에서 세 번째 요소에 bounce 클래스를 주는 방법은 다음과 같습니다.

```js
$(".target:nth-child(3)").addClass("animated bounce");
```

각 well 요소에서 두 번째 자식 요소가 bounce 효과를 갖도록 만드세요. 요소의 자식들은 `target` 클래스를 가지고 있어야 합니다.

# --hints--

`target` 요소들 중 두 번째 요소가 bounce 효과를 가져야 합니다.

```js
assert(
  $('.target:nth-child(2)').hasClass('animated') &&
    $('.target:nth-child(2)').hasClass('bounce')
);
```

bounce 효과가 적용되는 요소는 두 개뿐이어야 합니다.

```js
assert($('.animated.bounce').length === 2);
```

이 요소들을 수정하려면 `:nth-child()` 선택자를 사용하세요.

```js
assert(code.match(/\:nth-child\(/g));
```

이 클래스들은 jQuery를 사용해 요소에 추가해야 합니다.

```js
assert(
  code.match(/\$\(".target:nth-child\(2\)"\)/g) ||
    code.match(/\$\('.target:nth-child\(2\)'\)/g) ||
    code.match(/\$\(".target"\).filter\(":nth-child\(2\)"\)/g) ||
    code.match(/\$\('.target'\).filter\(':nth-child\(2\)'\)/g)
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
    $("#right-well").children().css("color", "orange");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
    $("#right-well").children().css("color", "orange");
    $(".target:nth-child(2)").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
