---
id: bad87fee1348bd9bedc08826
title: jQuery를 사용해 선택자로 HTML 요소 지정하기
challengeType: 6
forumTopicId: 18319
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-html-elements-with-selectors-using-jquery
---

# --description--

이제 `document ready` 함수를 가지고 있습니다.

이제 첫 번째 jQuery 문을 작성해 보겠습니다. 모든 jQuery 함수는 `$`로 시작하며, 보통 달러 기호 연산자 또는 블링이라고 부릅니다.

jQuery는 종종 <dfn>선택자</dfn>로 HTML 요소를 선택한 다음, 그 요소에 무언가를 수행합니다.

예를 들어, 모든 `button` 요소가 튕기도록 만들어 보겠습니다. 이 코드를 문서 준비 함수 안에 추가하세요:

```js
$("button").addClass("animated bounce");
```

jQuery 라이브러리와 Animate.css 라이브러리를 이미 배경에 포함했으니, 편집기에서 사용할 수 있습니다. 따라서 jQuery를 사용해 Animate.css `bounce` 클래스를 `button` 요소에 적용하는 것입니다.

# --hints--

jQuery `addClass()` 함수를 사용해 `animated`와 `bounce` 클래스를 `button` 요소에 추가해야 합니다.

```js
assert($('button').hasClass('animated') && $('button').hasClass('bounce'));
```

이 클래스들은 jQuery로만 요소에 추가해야 합니다.

```js
assert(!code.match(/class.*animated/g));
```

jQuery 코드는 `$(document).ready();` 함수 안에 있어야 합니다.

```js
assert(
  code.replace(/\s/g, '').match(/\$\(document\)\.ready\(function\(\)\{\$/g)
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
