---
id: bad87fee1348bd9aed508826
title: jQuery를 사용해 요소 복제하기
challengeType: 6
forumTopicId: 16780
dashedName: clone-an-element-using-jquery
---

# --description--

요소를 이동하는 것 외에도, 한 위치에서 다른 위치로 복사할 수도 있습니다.

jQuery에는 요소를 복사하는 `clone()`라는 함수가 있습니다.

예를 들어, `target2`에서 `left-well`을 `right-well`으로 복사하고 싶다면 다음과 같이 사용합니다:

```js
$("#target2").clone().appendTo("#right-well");
```

이것이 두 개의 jQuery 함수를 연결하는 것이라는 점을 눈치채셨나요? 이것을 <dfn>함수 체이닝</dfn>이라고 하며, jQuery로 작업할 때 편리한 방법입니다.

`target5` 요소를 복제하여 `left-well`에 추가하세요.

# --hints--

`target5` 요소가 `right-well` 안에 있어야 합니다.

```js
assert($('#right-well').children('#target5').length > 0);
```

`target5` 요소의 복사본도 `left-well` 안에 있어야 합니다.

```js
assert($('#left-well').children('#target5').length > 0);
```

이 요소들을 이동할 때는 jQuery만 사용하세요.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
