---
id: 669aff9f5488f1bea056416d
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, practicará trabajar con HTML semántico construyendo una página de blog dedicada al gato Sr. Whiskers. 

Para comenzar el proyecto, agrega el `<!DOCTYPE html>` y un elemento `html` con un atributo `lang` de `en`.

Recuerda que aprendiste cómo construir una plantilla básica de HTML como esta en el módulo anterior.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

# --hints--

Debes tener el `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Debes tener una etiqueta de apertura `html` con el lenguaje configurado en inglés.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Debes tener una etiqueta de cierre `html`.

```js
assert.match(code, /<\/html>/i);
```

Tu `DOCTYPE` debe ir antes del elemento `html`.

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
