---
id: 669fcb06c3034a39f5431a38
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Para preparar la creación de algún contenido real, añade un elemento `body` debajo del elemento `head`.

# --hints--

Debes tener una etiqueta de apertura `<body>`.

```js
assert.match(code, /<body>/i);
```

Debes tener una etiqueta de cierre `</body>`.

```js
assert.match(code, /<\/body>/i);
```

No debes modificar el elemento `head`. Asegúrese de que no eliminaste tu etiqueta de cierre.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Tu elemento `body` debe estar después del elemento `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  
  --fcc-editable-region--
</html>
```


# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  
  <body></body>
</html>
```
