---
id: 68778af56eb5d980a4e2daed
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

各個別の`Card`コンポーネントにプロップスを設定する代わりに、プロフィールデータの配列を作成し、そのデータを使って`Card`コンポーネントを動的にレンダリングしてください。

まず、`profiles`コンポーネント内に`App`配列を作成しましょう。

配列には`Mark`用のオブジェクトを1つ追加し、以下の値を設定してください：

- `id`: `1`
- `name`: `"Mark"`
- `title`: `"Front-End developer"`
- `bio`: `"I like to work with different front-end technologies and play video games."`

# --hints--

`profiles`コンポーネント内に`App`配列を定義してください。

```js
const explorer = await __helpers.Explorer(code);
const { profiles } = explorer.allFunctions.App?.variables ?? {};
assert.exists(profiles);
const profilesVal = eval(profiles?.value.toString());
assert.isArray(profilesVal);
```

`profiles` 配列には、必須のプロパティである `id`、`name`、`title`、`bio` をすべて含むオブジェクトを入れてください。

```js
const explorer = await __helpers.Explorer(code);
const { profiles } = explorer.allFunctions.App?.variables ?? {};
const profilesVal = eval(profiles?.value.toString());
assert.deepEqual(profilesVal, [{
  id: 1,
  name: "Mark",
  title: "Front-End developer",
  bio: "I like to work with different front-end technologies and play video games."
}]);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  );
}

export function App() {
  --fcc-editable-region--
     
  --fcc-editable-region--
  return (
    <div className="flex-container"></div>
  );
}
```

