---
id: a103376db3ba46b2d50db289
title: implementare un convertitore in spinal case
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

Lo spinal case è un formato di stringa in cui tutte le parole sono in minuscolo e separate da trattini. `"this-is-spinal-tap"` è un esempio di stringa in spinal case.

In questo laboratorio, creerai una funzione che converte una stringa data in spinal case.

**Obiettivo**: soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi creare una funzione chiamata `spinalCase`.
2. La funzione `spinalCase` deve prendere un singolo argomento, una stringa.
3. La funzione `spinalCase` deve restituire la stringa in formato spinal case. Per esempio, se l'argomento è `ProductLanding page`, la funzione deve restituire `product-landing-page`.

# --hints--

Devi creare una funzione chiamata `spinalCase`.

```js
assert.isFunction(spinalCase);
```

`spinalCase` deve prendere un singolo argomento.

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` deve restituire la stringa `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` deve restituire la stringa `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` deve restituire la stringa `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` deve restituire la stringa `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` deve restituire la stringa `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
