---
id: 6823f9df49cc206af5471a30
title: Налагодження сторінки профілю Camperbot
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot намагається створити сторінку профілю. Вони попросили друга переглянути свій код, і той сказав, що в ньому є помилки.

Ваше завдання — виправити всі помилки Camperbot, щоб вони могли продовжити створення своєї сторінки профілю. Виконайте пункти в історіях користувача нижче та натисніть "Run the Tests", щоб перевірити, чи ви виправили всі помилки.

**Історії користувача:**

1. Camperbot намагається використати елемент `heading2`, але такого елемента не існує. Виправте ці теги, щоб вони використовували правильний заголовок другого рівня.
2. Camperbot намагається додати два абзаци з `pp`, але їх теж не існує. Виправте їх, щоб вони використовували правильні теги абзацу.
3. Camperbot використовує елемент `h3` для підзаголовка `Background and Interests`, але в ньому є синтаксична помилка. Знайдіть проблему і виправте її.

# --hints--

У вас має бути рівно один елемент `h1` на сторінці.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Не видаляйте елемент `<h1>Hello from Camperbot!</h1>` зі сторінки.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

У вас має бути рівно один елемент `h2` на сторінці.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

У вас має бути елемент `h2` з текстом `About`. Ось приклад: `<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

У вас має бути елемент абзацу з текстом `My name is Camperbot and I love learning new things.`. Ось приклад: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

У вас має бути елемент абзацу з текстом `I enjoy solving puzzles.`. Ось приклад: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

У вас має бути рівно два елементи абзацу на сторінці.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

У вас має бути рівно один елемент `h3` на сторінці.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

У вас має бути елемент `h3` з текстом `Background and Interests`. Ось приклад: `<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

У вас має бути загалом п’ять елементів на сторінці.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
