---
id: 694ca14c8ef00056ad146e51
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Agrega el objeto `astronaut` al arreglo `crew` y registra `Added ${astronaut.name} as ${astronaut.role}` en la consola. Puedes usar literales de plantilla o concatenación de cadenas para el mensaje del registro.

# --hints--

Debes agregar el `astronaut` al arreglo `crew`.

```js
const crew = [];
const astronauts = [
  { id: 1, name: "Abraham", role: "Commander", isEVAEligible: true, priority: 1 },
  { id: 2, name: "Boris", role: "Pilot", isEVAEligible: false, priority: 2 },
  { id: 3, name: "Catarina", role: "Engineer", isEVAEligible: true, priority: 3 }
];

astronauts.forEach(astro => addCrewMember(crew, astro));

assert.strictEqual(crew.length, astronauts.length);
astronauts.forEach((astro, idx) => {
  assert.deepEqual(crew[idx], astro);
});
```

Debes registrar el `name` del astronauta en el siguiente formato, `Added {astronaut.name} as {astronaut.role}`.

```js
function captureConsoleLogs(fn) {
  const originalLog = console.log;
  const logs = [];

  console.log = function (...args) {
    logs.push(args.join(" "));
  };

  try {
    fn();
  } finally {
    console.log = originalLog;
  }

  return logs;
}

const crewForLogging = [];
const astronautsToLog = [
  { id: 1, name: "Abraham", role: "Commander" },
  { id: 2, name: "Boris", role: "Pilot" },
  { id: 3, name: "Catarina", role: "Engineer" }
];

const logs = captureConsoleLogs(() => {
  astronautsToLog.forEach(astro => addCrewMember(crewForLogging, astro));
});

assert.strictEqual(logs.length, astronautsToLog.length);

astronautsToLog.forEach((astro, idx) => {
  assert.match(
    logs[idx],
    new RegExp(`Added\\s*${astro.name}\\s*as\\s*${astro.role}`, 'i'),
  );
});
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
