---
id: 694c84330cbe604821b9deef
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

El control de misión requiere un método para agregar nuevos miembros a la lista. Crea una función vacía llamada `addCrewMember` que acepte dos parámetros, `crew` y `astronaut`.

# --hints--

Debes crear una función llamada `addCrewMember`.

```js
assert.isFunction(addCrewMember);
```

Tu función `addCrewMember` debe tener los parámetros `crew` y `astronaut`.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
