---
id: 694cabb9f1afa9446f36ffc0
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Ahora puedes empezar a construir tu lista. Usa `addCrewMember()` para agregar el `firstAstronaut` a tu `squad`.

# --hints--

Debes llamar a `addCrewMember()` con `squad` y `firstAstronaut`.

```js
assert.match(__helpers.removeJSComments(code), /addCrewMember\s*\(\s*squad\s*,\s*firstAstronaut\s*\)/);
assert.strictEqual(squad.length, 1);
assert.strictEqual(squad[0], firstAstronaut);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
  console.log(`Added ${astronaut.name} as ${astronaut.role}`);
}

--fcc-editable-region--

--fcc-editable-region--
```
