---
id: 694c9c186c089f7bdff0f3c6
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Generalmente, es buena práctica validar tus entradas dentro de una función. Usa un ciclo `for` para iterar a través del arreglo `crew` y verificar si algún miembro ya tiene el mismo `id` que el `astronaut` de entrada. Si se encuentra un duplicado, registra el mensaje de error `console.log("Duplicate ID: " + astronaut.id)` y luego llama a `return` para salir de la función anticipadamente. 

Aquí tienes un ejemplo de esta técnica comúnmente usada:

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

Tu función debe usar un ciclo `for` para escanear los miembros existentes de `crew` (cualquier forma válida de `for (...)` está bien).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

Dentro del ciclo `for`, debes verificar si hay un `id` duplicado. Si se encuentra uno, registra `"Duplicate ID: " + astronaut.id"` en la consola y luego llama a `return` para salir de la función.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Cuando se proporcionan entradas válidas, tu función `addCrewMember` no debe invocar un mensaje de error `console.log()` ni una salida temprana con `return`.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
