---
id: 694c84330cbe604821b9deef
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

El control de misión requiere un método para agregar nuevos miembros a la lista. Crea una función vacía llamada `addCrewMember` que acepte dos parámetros, `crew` y `astronaut`.

# --hints--

Debes crear una función llamada `addCrewMember`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

Tu función `addCrewMember` debe tener los parámetros `crew` y `astronaut`.

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
