---
id: 694c84106e9023f289c03b84
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Tu lista comienza con un solo astronauta comandante. Crea un objeto llamado `firstAstronaut` con las siguientes propiedades:

| Key | Value |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

Debes crear una variable `firstAstronaut`.

```js
assert.isDefined(firstAstronaut);
```

Tu `firstAstronaut` debe ser un objeto.

```js
assert.isObject(firstAstronaut);
```

Tu objeto `firstAstronaut` debe incluir los valores correctos de las propiedades.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
