---
id: 694c40d86fe6a78d20b6559c
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, usarás JavaScript para construir una lista de astronautas para una misión espacial. 

Para comenzar, crea un arreglo vacío para tu lista llamado `squad`.

# --hints--

Debes crear una variable `squad`.

```js
assert.isDefined(squad);
```

Tu variable `squad` debe ser un arreglo.

```js
assert.isArray(squad);
```

Tu arreglo `squad` debe estar vacío.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
