---
id: 694d8ddd106340237efb0c80
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

El método `splice()` puede modificar arreglos agregando o eliminando elementos en cualquier posición, incluso en el medio. Dado que `splice()` devuelve un arreglo que contiene los elementos eliminados, puedes usarlo para intercambiar dos elementos en un arreglo sin mutar el original. Esto se puede hacer en una línea usando la siguiente técnica:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

Aquí hay un ejemplo:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

La técnica, aplicada en la tercera línea, funciona de la siguiente manera:

- `splice(3, 1, copyArray[1])` elimina el elemento en el índice `3` (`55`)

- Inserta el elemento de `copyArray[1]` (`97`) en el índice `3`

- `splice()` devuelve `[55]`, un arreglo que contiene el elemento eliminado

- `[0]` extrae `55` de ese arreglo

- Ese valor (`55`) se asigna de nuevo a `copyArray[1]`, completando el intercambio

Usa esta técnica con el arreglo `updatedCrew` para intercambiar a los astronautas en `fromIndex` y `toIndex`.

# --hints--

Debes realizar el intercambio llamando al método `splice()`.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

Debes llamar a `splice()` con el arreglo `updatedCrew` para realizar el intercambio. Debes usar la técnica de una sola línea.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

Tu función `swapCrewMembers` no debe mutar el arreglo original `crew` de entrada.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
