---
id: 587d7b8e367417b2b2512b5c
title: Розуміння термінології функціонального програмування
challengeType: 1
forumTopicId: 301240
dashedName: understand-functional-programming-terminology
---

# --description--

Команда FCC змінила настрій і тепер хоче два види чаю: зелений і чорний. Загальний факт: зміни настрою клієнта — це досить поширене явище.

З цією інформацією нам потрібно повернутися до функції `getTea` з попереднього завдання, щоб обробляти різні запити на чай. Ми можемо змінити `getTea`, щоб вона приймала функцію як параметр і могла змінювати тип чаю, який готує. Це робить `getTea` більш гнучкою і дає програмісту більше контролю, коли запити клієнта змінюються.

Але спершу розглянемо деяку функціональну термінологію:

<dfn>Функції зворотного виклику</dfn> — це функції, які передаються або вставляються в іншу функцію, щоб визначити, коли викликати цю функцію. Ви могли бачити, як їх передають іншим методам, наприклад, у `filter` функція зворотного виклику повідомляє JavaScript критерії для фільтрації масиву.

Функції, які можна призначити змінній, передати в іншу функцію або повернути з іншої функції так само, як будь-яке інше звичайне значення, називаються <dfn>функціями першого класу</dfn>. У JavaScript усі функції є функціями першого класу.

Функції, які приймають функцію як аргумент або повертають функцію як повернене значення, називаються <dfn>функціями вищого порядку</dfn>.

Коли функції передаються в іншу функцію або повертаються з неї, тоді ці функції, які були передані або повернені, можна назвати <dfn>лямбдами</dfn>.

# --instructions--

Приготуйте 27 чашок зеленого чаю і 13 чашок чорного чаю та збережіть їх у змінних `tea4GreenTeamFCC` і `tea4BlackTeamFCC` відповідно. Зверніть увагу, що функція `getTea` була змінена і тепер приймає функцію як перший аргумент.

Примітка: Дані (кількість чашок чаю) передаються як останній аргумент. Про це ми поговоримо докладніше в наступних уроках.

# --hints--

Змінна `tea4GreenTeamFCC` має містити 27 чашок зеленого чаю для команди.

```js
assert(tea4GreenTeamFCC.length === 27);
```

Змінна `tea4GreenTeamFCC` має містити чашки зеленого чаю.

```js
assert(tea4GreenTeamFCC[0] === 'greenTea');
```

Змінна `tea4BlackTeamFCC` має містити 13 чашок чорного чаю.

```js
assert(tea4BlackTeamFCC.length === 13);
```

Змінна `tea4BlackTeamFCC` має містити чашки чорного чаю.

```js
assert(tea4BlackTeamFCC[0] === 'blackTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareGreenTea = () => 'greenTea';

// Function that returns a string representing a cup of black tea
const prepareBlackTea = () => 'blackTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (prepareTea, numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4GreenTeamFCC = null;
const tea4BlackTeamFCC = null;
// Only change code above this line

console.log(
  tea4GreenTeamFCC,
  tea4BlackTeamFCC
);
```

# --solutions--

```js
const prepareGreenTea = () => 'greenTea';
const prepareBlackTea = () => 'blackTea';

const getTea = (prepareTea, numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

const tea4BlackTeamFCC = getTea(prepareBlackTea, 13);
const tea4GreenTeamFCC = getTea(prepareGreenTea, 27);
```
