---
id: 587d7b8f367417b2b2512b62
title: Реалізувати map на прототипі
challengeType: 1
forumTopicId: 301230
dashedName: implement-map-on-a-prototype
---

# --description--

Як ви бачили, застосовуючи `Array.prototype.map()`, або просто `map()` раніше, метод `map` повертає масив такої ж довжини, як і той, до якого він був викликаний. Він також не змінює оригінальний масив, якщо тільки його функція зворотного виклику цього не робить.

Інакше кажучи, `map` — це чиста функція, і її вихідні дані залежать виключно від вхідних. Крім того, вона приймає іншу функцію як свій аргумент.

Ви можете багато чого дізнатися про метод `map`, якщо реалізуєте власну версію. Рекомендується використовувати цикл `for` або `Array.prototype.forEach()`.

# --instructions--

Напишіть власний `Array.prototype.myMap()`, який має поводитися точно так само, як `Array.prototype.map()`. Не слід використовувати вбудований метод `map`. Екземпляр `Array` можна отримати в методі `myMap` за допомогою `this`.

# --hints--

`[23, 65, 98, 5, 13].myMap(item => item * 2)` має дорівнювати `[46, 130, 196, 10, 26]`.

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item * 2;
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`["naomi", "quincy", "camperbot"].myMap(element => element.toUpperCase())` має повертати `["NAOMI", "QUINCY", "CAMPERBOT"]`.

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element.toUpperCase();
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`[1, 1, 2, 5, 2].myMap((element, index, array) => array[index + 1] || array[0])` має повертати `[1, 2, 5, 2, 1]`.

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array[index + 1] || array[0];
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

Ваш код не повинен використовувати метод `map`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?map/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    newArray.push(callback(this[i], i, this));
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const doubled_s = s.myMap(item => item * 2);
```
