---
id: 587d7b8e367417b2b2512b5d
title: Зрозуміти небезпеки використання імперативного коду
challengeType: 1
forumTopicId: 301241
dashedName: understand-the-hazards-of-using-imperative-code
---

# --description--

Функціональне програмування — це хороша звичка. Воно допомагає тримати ваш код легким для керування і захищає від підступних помилок. Але перед тим, як перейти до нього, давайте розглянемо імперативний підхід до програмування, щоб підкреслити, де можуть виникнути проблеми.

В англійській (та багатьох інших мовах) імперативний спосіб використовується для того, щоб давати накази. Аналогічно, імперативний стиль у програмуванні — це стиль, який дає комп’ютеру набір інструкцій для виконання завдання.

Часто ці інструкції змінюють стан програми, наприклад, оновлюють глобальні змінні. Класичним прикладом є написання `for` циклу, який дає точні вказівки, як ітерувати над індексами масиву.

Натомість функціональне програмування — це форма декларативного програмування. Ви кажете комп’ютеру, що хочете зробити, викликаючи метод або функцію.

JavaScript пропонує багато заздалегідь визначених методів, які виконують загальні завдання, тож вам не потрібно писати, як саме комп’ютер має їх виконувати. Наприклад, замість використання `for` циклу, згаданого вище, ви можете викликати метод `map`, який опікується деталями ітерації над масивом. Це допомагає уникнути семантичних помилок, таких як "Off By One Errors", які розглядалися в розділі налагодження.

Розглянемо ситуацію: ви переглядаєте веб у браузері і хочете відстежувати відкриті вкладки. Спробуємо змоделювати це за допомогою простого об’єктно-орієнтованого коду.

Об’єкт Window складається з вкладок, і зазвичай у вас відкрито більше ніж одне вікно. Заголовки кожного відкритого сайту в кожному об’єкті Window зберігаються в масиві. Після роботи в браузері (відкриття нових вкладок, об’єднання вікон і закриття вкладок) ви хочете вивести вкладки, які залишилися відкритими. Закриті вкладки видаляються з масиву, а нові вкладки (для простоти) додаються в кінець.

Редактор коду показує реалізацію цієї функціональності з функціями для `tabOpen()`, `tabClose()` і `join()`. Масив `tabs` є частиною об’єкта Window, який зберігає назви відкритих сторінок.

# --instructions--

Розгляньте код у редакторі. Він використовує метод, який має побічні ефекти в програмі, що спричиняє некоректну поведінку. Фінальний список відкритих вкладок, збережений у `finalTabs.tabs`, має бути `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']`, але список, який створює код, трохи відрізняється.

Змініть `Window.prototype.tabClose`, щоб він видаляв правильну вкладку.

# --hints--

`finalTabs.tabs` має бути `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']`

```js
assert.deepEqual(finalTabs.tabs, [
  'FB',
  'Gitter',
  'Reddit',
  'Twitter',
  'Medium',
  'new tab',
  'Netflix',
  'YouTube',
  'Vine',
  'GMail',
  'Work mail',
  'Docs',
  'freeCodeCamp',
  'new tab'
]);
```

# --seed--

## --seed-contents--

```js
// tabs is an array of titles of each site open within the window
const Window = function(tabs) {
  this.tabs = tabs; // We keep a record of the array inside the object
};

// When you join two windows into one window
Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

// When you open a new tab at the end
Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab'); // Let's open a new tab for now
  return this;
};

// When you close a tab
Window.prototype.tabClose = function(index) {

  // Only change code below this line

  const tabsBeforeIndex = this.tabs.splice(0, index); // Get the tabs before the tab
  const tabsAfterIndex = this.tabs.splice(index + 1); // Get the tabs after the tab

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex); // Join them together

  // Only change code above this line

  return this;
 };

// Let's create three browser windows
const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']); // Your mailbox, drive, and other work sites
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']); // Social sites
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']); // Entertainment sites

// Now perform the tab opening, closing, and other operations
const finalTabs = socialWindow
  .tabOpen() // Open a new tab for cat memes
  .join(videoWindow.tabClose(2)) // Close third tab in video window, and join
  .join(workWindow.tabClose(1).tabOpen());
console.log(finalTabs.tabs);
```

# --solutions--

```js
const Window = function(tabs) {
  this.tabs = tabs;
};

Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab');
  return this;
};

Window.prototype.tabClose = function(index) {
  const tabsBeforeIndex = this.tabs.slice(0, index);
  const tabsAfterIndex = this.tabs.slice(index + 1);

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex);
  return this;
 };

const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']);
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']);
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']);

const finalTabs = socialWindow
  .tabOpen()
  .join(videoWindow.tabClose(2))
  .join(workWindow.tabClose(1).tabOpen());
```
