---
id: 587d7da9367417b2b2512b69
title: Сортування масиву в алфавітному порядку за допомогою методу sort
challengeType: 1
forumTopicId: 18303
dashedName: sort-an-array-alphabetically-using-the-sort-method
---

# --description--

Метод `sort` сортує елементи масиву відповідно до функції зворотного виклику.

Наприклад:

```js
function ascendingOrder(arr) {
  return arr.sort(function(a, b) {
    return a - b;
  });
}

ascendingOrder([1, 5, 2, 3, 4]);
```

Це поверне значення `[1, 2, 3, 4, 5]`.

```js
function reverseAlpha(arr) {
  return arr.sort(function(a, b) {
    return a === b ? 0 : a < b ? 1 : -1;
  });
}

reverseAlpha(['l', 'h', 'z', 'b', 's']);
```

Це поверне значення `['z', 's', 'l', 'h', 'b']`.

За замовчуванням метод сортування в JavaScript відбувається за значенням Unicode символів рядка, що може дати несподівані результати. Тому рекомендується надавати функцію зворотного виклику, щоб вказати, як сортувати елементи масиву. Коли така функція зворотного виклику, зазвичай званa `compareFunction`, передається, елементи масиву сортуються відповідно до поверненого значення `compareFunction`: Якщо `compareFunction(a,b)` повертає значення менше за 0 для двох елементів `a` і `b`, тоді `a` буде розташований перед `b`. Якщо `compareFunction(a,b)` повертає значення більше за 0 для двох елементів `a` і `b`, тоді `b` буде розташований перед `a`. Якщо `compareFunction(a,b)` повертає значення, рівне 0, для двох елементів `a` і `b`, тоді `a` і `b` залишаться без змін.

# --instructions--

Використайте метод `sort` у функції `alphabeticalOrder`, щоб відсортувати елементи `arr` в алфавітному порядку. Функція має повертати відсортований масив.

# --hints--

Ваш код має використовувати метод `sort`.

```js
assert(__helpers.removeJSComments(code).match(/\.sort/g));
```

`alphabeticalOrder(["a", "d", "c", "a", "z", "g"])` має повертати `["a", "a", "c", "d", "g", "z"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'd', 'c', 'a', 'z', 'g'])) ===
    JSON.stringify(['a', 'a', 'c', 'd', 'g', 'z'])
);
```

`alphabeticalOrder(["x", "h", "a", "m", "n", "m"])` має повертати `["a", "h", "m", "m", "n", "x"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['x', 'h', 'a', 'm', 'n', 'm'])) ===
    JSON.stringify(['a', 'h', 'm', 'm', 'n', 'x'])
);
```

`alphabeticalOrder(["a", "a", "a", "a", "x", "t"])` має повертати `["a", "a", "a", "a", "t", "x"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'a', 'a', 'a', 'x', 't'])) ===
    JSON.stringify(['a', 'a', 'a', 'a', 't', 'x'])
);
```

# --seed--

## --seed-contents--

```js
function alphabeticalOrder(arr) {
  // Only change code below this line

  return arr
  // Only change code above this line
}

alphabeticalOrder(["a", "d", "c", "a", "z", "g"]);
```

# --solutions--

```js
function alphabeticalOrder(arr) {
  return arr.sort();
}
```
