---
id: 587d7da9367417b2b2512b67
title: додати елементи в кінець масиву за допомогою concat замість push
challengeType: 1
forumTopicId: 301226
dashedName: add-elements-to-the-end-of-an-array-using-concat-instead-of-push
---

# --description--

Функціональне програмування полягає у створенні та використанні немутуючих функцій.

Останнє завдання познайомило з методом `concat` як способом об’єднати масиви в новий масив без мутації оригінальних масивів. Порівняйте `concat` з методом `push`. `push` додає елементи в кінець того самого масиву, на якому викликається, що мутує цей масив. Ось приклад:

```js
const arr = [1, 2, 3];
arr.push(4, 5, 6);
```

`arr` матиме змінене значення `[1, 2, 3, 4, 5, 6]`, що не відповідає принципам функціонального програмування.

`concat` пропонує спосіб об’єднати нові елементи в кінець масиву без будь-яких мутаційних побічних ефектів.

# --instructions--

Змініть функцію `nonMutatingPush` так, щоб вона використовувала `concat` для об’єднання `newItem` в кінець `original` без мутації масивів `original` або `newItem`. Функція має повертати масив.

# --hints--

Ваш код має використовувати метод `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

Ваш код не має використовувати метод `push`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?push/g));
```

Масив `first` не повинен змінюватися.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

Масив `second` не повинен змінюватися.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingPush([1, 2, 3], [4, 5])` має повертати `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingPush([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingPush(original, newItem) {
  // Only change code below this line
  return original.push(newItem);

  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingPush(first, second);
```

# --solutions--

```js
function nonMutatingPush(original, newItem) {
  return original.concat(newItem);
}
const first = [1, 2, 3];
const second = [4, 5];
```
