---
id: 587d7dab367417b2b2512b70
title: Вступ до каррінгу та часткового застосування
challengeType: 1
forumTopicId: 301232
dashedName: introduction-to-currying-and-partial-application
---

# --description--

<dfn>Арність</dfn> функції — це кількість параметрів, які вона потребує. <dfn>Каррінг</dfn> функції означає перетворення функції з арністю N у N функцій з арністю 1.

Іншими словами, це перебудова функції так, щоб вона приймала один параметр, а потім повертала іншу функцію, яка приймає наступний параметр, і так далі.

Ось приклад:

```js
function unCurried(x, y) {
  return x + y;
}

function curried(x) {
  return function(y) {
    return x + y;
  }
}

const curried = x => y => x + y

curried(1)(2)
```

`curried(1)(2)` поверне `3`.

Це корисно у вашій програмі, якщо ви не можете передати всі параметри функції одночасно. Ви можете зберегти кожен виклик функції у змінну, яка триматиме референс на повернену функцію, що прийматиме наступний параметр, коли він стане доступним. Ось приклад із використанням каррінгової функції з прикладу вище:

```js
const funcForY = curried(1);
console.log(funcForY(2)); // 3
```

Аналогічно, <dfn>часткове застосування</dfn> можна описати як застосування кількох параметрів до функції одночасно з подальшим поверненням іншої функції, до якої застосовують більше параметрів. Ось приклад:

```js
function impartial(x, y, z) {
  return x + y + z;
}

const partialFn = impartial.bind(this, 1, 2);
partialFn(10); // 13
```

# --instructions--

Заповніть тіло функції `add`, щоб вона використовувала каррінг для додавання параметрів `x`, `y` та `z`.

# --hints--

`add(10)(20)(30)` має повертати `60`.

```js
assert(add(10)(20)(30) === 60);
```

`add(1)(2)(3)` має повертати `6`.

```js
assert(add(1)(2)(3) === 6);
```

`add(11)(22)(33)` має повертати `66`.

```js
assert(add(11)(22)(33) === 66);
```

Ваш код має містити фінальну інструкцію, яка повертає `x + y + z`.

```js
assert(__helpers.removeJSComments(code).match(/[xyz]\s*?\+\s*?[xyz]\s*?\+\s*?[xyz]/g));
```

# --seed--

## --seed-contents--

```js
function add(x) {
  // Only change code below this line


  // Only change code above this line
}

add(10)(20)(30);
```

# --solutions--

```js
const add = x => y => z => x + y + z
```
