---
id: 587d7b8d367417b2b2512b5b
title: Дізнайтеся про функціональне програмування
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

Функціональне програмування — це стиль програмування, де рішення представлені простими, ізольованими функціями, без будь-яких побічних ефектів поза межами області видимості функції: `INPUT -> PROCESS -> OUTPUT`

Функціональне програмування полягає в:

1) Ізольованих функціях — немає залежності від стану програми, включно з глобальними змінними, які можуть змінюватися

2) Чистих функціях — однакові вхідні дані завжди дають однакові вихідні дані

3) Функціях з обмеженими побічними ефектами — будь-які зміни або мутації стану програми поза межами функції ретельно контролюються

# --instructions--

Учасники freeCodeCamp дуже люблять чай.

У редакторі коду функції `prepareTea` і `getTea` вже визначені для вас. Викличте функцію `getTea`, щоб отримати 40 чашок чаю для команди, і збережіть їх у змінній `tea4TeamFCC`.

# --hints--

Змінна `tea4TeamFCC` має містити 40 чашок чаю для команди.

```js
assert(tea4TeamFCC.length === 40);
```

Змінна `tea4TeamFCC` має містити чашки зеленого чаю.

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40); 
```
