---
id: 587d7da9367417b2b2512b66
title: Об’єднання двох масивів за допомогою методу concat
challengeType: 1
forumTopicId: 301229
dashedName: combine-two-arrays-using-the-concat-method
---

# --description--

<dfn>Конкатенація</dfn> означає з’єднання елементів один за одним. JavaScript пропонує метод `concat` як для рядків, так і для масивів, який працює однаково. Для масивів метод викликається на одному масиві, а інший масив передається як аргумент `concat`, який додається в кінець першого масиву. Метод повертає новий масив і не змінює жоден із початкових масивів. Ось приклад:

```js
[1, 2, 3].concat([4, 5, 6]);
```

Повернутий масив буде `[1, 2, 3, 4, 5, 6]`.

# --instructions--

Використайте метод `concat` у функції `nonMutatingConcat`, щоб об’єднати `attach` у кінець `original`. Функція має повертати об’єднаний масив.

# --hints--

Ваш код має використовувати метод `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

Масив `first` не повинен змінюватися.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

Масив `second` не повинен змінюватися.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingConcat([1, 2, 3], [4, 5])` має повертати `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingConcat([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingConcat(original, attach) {
  // Only change code below this line


  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingConcat(first, second);
```

# --solutions--

```js
function nonMutatingConcat(original, attach) {
  return original.concat(attach);
}
const first = [1, 2, 3];
const second = [4, 5];
```
