---
id: 66b62d0ad68488dd76228d6c
title: Einen zufälligen Hintergrundfarbwechsler debuggen
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot lernt JavaScript und hat versucht, einen eigenen zufälligen Hintergrundfarbwechsler zu bauen. Dabei sind ihm jedoch einige Fehler unterlaufen.   

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, damit das Labor richtig funktioniert.

**User Stories:**

1. Sie sollten die Syntaxfehler im `darkColorsArr` Array beheben.
1. Sie sollten den `ReferenceError` beheben, der geworfen wird, wenn das `math` Objekt innerhalb der `getRandomIndex` Funktion referenziert wird.
1. Sie sollten `randomIndex` auf die nächstkleinere ganze Zahl abrunden.
1. Sie sollten den Wert der `body` Variablen aktualisieren, damit der korrekte Methodenname verwendet wird, und den dazugehörigen `TypeError` beheben.
1. Sie sollten `bgHexCodeSpanElement` so korrigieren, dass es das Element mit dem `id` von `bg-hex-code` richtig auswählt.
1. Sie sollten die `color` Variable innerhalb von `changeBackgroundColor` so ändern, dass sie auf eine zufällige Farbe aus dem `darkColorsArr` Array gesetzt wird.
1. Sie sollten die `btn` Variable aktualisieren, damit sie das Button-Element anspricht. Öffnen Sie dazu die `index.html`, um den korrekten `id` Namen für dieses Button-Element zu sehen.
1. Sie sollten die Zeile, die einen Ereignis-Listener für das `click` Ereignis an `btn` hinzufügt, so korrigieren, dass eine Referenz auf die `changeBackgroundColor` Funktion an die Methode übergeben wird.
1. Ihr Hintergrund sollte sich bei jedem Klick auf das `darkColorsArr` zufällig in eine der Farben aus `#btn` ändern.

# --hints--

Sie sollten die Syntaxfehler im `darkColorsArr` Array beheben. Denken Sie daran, dass Array-Elemente durch Kommas getrennt sein müssen.

```js
assert.isArray(darkColorsArr);
```

Sie sollten den Groß-/Kleinschreibungsfehler in der `math.random()` Zeile beheben.

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

Sie sollten `darkColorsArr.length * Math.random()` auf die nächstkleinere ganze Zahl abrunden.

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

Sie sollten den Wert Ihrer `body` Variablen aktualisieren, sodass `queryselector` camelCase verwendet.

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

Sie sollten `"bg-hex-code"` so korrigieren, dass es die `id` von `bg-hex-code` richtig referenziert. Denken Sie daran, dass `id` Selektoren in CSS mit einem Hash-Symbol (`#`) beginnen.

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

Sie sollten die `getRandomIndex` Funktion in der `darkColorsArr[getRandomIndex]` Zeile aufrufen.

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

Sie sollten den `id` Namen der `"#click-btn"` Zeile an den korrekten `id` Namen in der `index.html` Datei anpassen.

```js
assert.match(btn?.id, /btn/);
```

Sie sollten die `changeBackgroundColor` Funktion nicht aufrufen. Stattdessen sollten Sie die Funktionsreferenz an die `addEventListener` Methode übergeben.

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
