---
id: 69757cc0faae0152c1418aad
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、数値と数学的な演算を使って割り勘計算機を作る練習をします。このツールは、食事代とチップを加えた後に、各人がいくら支払うべきかを計算します。

まず、費用が加算されるたびに合計金額を記録する方法が必要です。Pythonでは、変数を使って時間とともに変わる整数（小数点のない数値）を保存できます。

例えば、次のように書くかもしれません。

```py
my_number = 2
```

`running_total` という名前の変数を作り、値 `0` を代入してください。

# --hints--

`running_total` という名前の変数を用意してください。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

整数 `0` を `running_total` 変数に代入してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
