---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

各コースには料金があります。これらの金額を後で使うために変数にストアする必要があります。これらの金額にはセントも含まれるため、小数点数を表すのに使われる float 型を使います。以下は float 値を持つ変数の例です。

```py
change = 2.35
```

4つの変数を作成してください：`appetizers` に `37.89` を代入し、`main_courses` に `57.34` を代入し、`desserts` に `39.39` を代入し、`drinks` に `64.21` を代入してください。

# --hints--

`appetizers` という名前の変数を定義し、値 `37.89` を代入してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

`main_courses` という名前の変数を定義し、値 `57.34` を代入してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

`desserts` という名前の変数を定義し、値 `39.39` を代入してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

`drinks` という名前の変数を定義し、値 `64.21` を代入してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
