---
id: 6976cc0b0c135686a4620b94
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

次に、割り勘する人数を考慮する必要があります。この値を、前のステップと同様に変数にストアしてください。

`num_of_friends` という名前の変数を作成し、`4` の値を代入してください。これはワークショップの後半で最終的な割り勘額を計算するために使います。

# --hints--

`num_of_friends` という名前の変数を用意してください。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

`4` 変数に整数の `num_of_friends` を代入してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
