---
id: 69779dcb44016eccc05c15a4
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

サービスが素晴らしかったので、グループは25％のチップを渡すことにしました。Pythonでパーセンテージを計算するには、合計にパーセンテージの小数表現を乗算すればいいです。

例えば、ある値の10％を求めるには、`0.10`に`*`演算子を使って乗算します。

```py
tax = total * 0.10
```

`tip`に`running_total`を乗算した結果を`0.25`という変数に代入してください。

最後に、`print()`を使って文字列`Tip amount:`とスペース、その後に`tip`変数の値を出力してください。

# --hints--

`tip`という名前の変数があるはずです。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

`tip`変数は`running_total * 0.25`の結果であるべきです。

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

文字列`Tip amount:`とスペース、その後に`tip`変数を出力してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
