---
id: 6977a8bbcbdd557886bd0718
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

チップが含まれたので、グループ全体の合計金額がわかりました。友達の人数で合計金額を割って、1人あたりの支払額を計算してください。

Pythonでは、割り算を行うのにスラッシュ `/` を使います。例えば：

```py
half = 10 / 2
```

`final_bill` を `running_total` で割った結果を `num_of_friends` という名前の変数に代入してください。

最後に、`print()` 関数を使って、文字列 `Bill per person:` と空白、その後に `final_bill` の値を出力してください。

# --hints--

`final_bill` という名前の変数を定義してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

`/` 演算子を使って `running_total` を `num_of_friends` で割り、その結果を `final_bill` 変数に代入してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

文字列 `Bill per person:` と空白、その後に `final_bill` 変数を出力してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
