---
id: 6977a4306bb3e856690a4890
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

チップを計算したので、最終的な請求額を求めるためにそれを`running_total`に加える必要があります。

Pythonでは、拡張代入演算子`+=`を使って、変数に値を加え、その変数を同時に更新できます。例えば、`total += 5`は`total = total + 5`と書く省略形です。

`+=`演算子を使って、`tip`の値を`running_total`に加えてください。最後に、`print()`を使って、文字列`Total with tip:`と空白、その後に`running_total`の値を出力してください。

# --hints--

`+=`演算子を使って、`tip`の値を`running_total`に加えるべきです。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

文字列`Total with tip:`と空白、その後に`running_total`変数を出力してください。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total with tip:', running_total)") or _Node(_code).has_call("print(f'Total with tip: {running_total}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
