---
id: 6982684f3a25f379e195a5fc
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

個別の費用をストアしたので、合計を計算できます。Pythonでは、加算演算子 `+` を使って値を合計します。

`+=` 演算子は、既存の変数に値を加えて同時に更新します。例えば：

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

`+=` 演算子を一度使って、`appetizers`、`main_courses`、`desserts`、`drinks` を `running_total` に加えてください。

最後に、`print()` を使って文字列 `Total bill so far:` と空白、その後に `running_total` の値を表示してください。

**注意：** 出力に予想より多くの10進数の桁が表示されることがあります。前のレッスンで学んだように、これは数値が2進数でストアされているためで、多くの10進数の値はこの形式で正確に表せず、丸め誤差が生じるためです。

# --hints--

`running_total` を `+=` 演算子で一度更新して、4つのコース費用をすべて加えてください。

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total += {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

文字列 `Total bill so far:` と空白、その後に `running_total` の値を出力してください。

```js
({
    test: () => runPython(`
    sol1 = "print('Total bill so far:', running_total)"
    sol2 = "print(f'Total bill so far: {running_total}')"
    assert _Node(_code).has_call(sol1) or _Node(_code).has_call(sol2)`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
