---
id: 68d275dd800f404d22a07564
title: تنفيذ خوارزمية البحث بالعمق أولاً
challengeType: 27
dashedName: lab-depth-first-search
---

# --description--

في هذا المختبر، ستنفذ خوارزمية استكشاف رسم بياني تُسمى <dfn>البحث بالعمق أولاً</dfn>.

بينما يبحث <dfn>البحث بالاتساع أولاً</dfn> على طول حواف متزايدة من العقدة المصدر، يذهب <dfn>البحث بالعمق أولاً</dfn> أولاً في مسار من الحواف لأبعد حد ممكن.

عندما يصل إلى نهاية مسار، يعود البحث إلى آخر عقدة بها مسار حافة غير مُزار ويواصل البحث.

على عكس البحث بالاتساع أولاً، في كل مرة تُزار عقدة، لا يزور كل جيرانها. بل يزور أولاً أحد جيرانها ويستمر في ذلك المسار حتى لا تبقى عقد لزيارتها في ذلك المسار.

لتنفيذ هذه الخوارزمية، ستحتاج إلى استخدام مكدس (مصفوفة يُزال منها العنصر الأخير المُضاف أولاً، وفق مبدأ <dfn>Last-In-First-Out</dfn>). المكدس مفيد في خوارزميات البحث بالعمق أولاً لأنك عند إضافة الجيران إلى المكدس، تريد زيارة الجيران المُضافين حديثًا أولاً وإزالتهم من المكدس.

الناتج البسيط لهذه الخوارزمية هو قائمة بالعقد التي يمكن الوصول إليها من عقدة معينة. لذلك، ستحتاج أيضًا إلى تتبع العقد التي تزورها.

**الهدف**: نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك دالة باسم `dfs`.
1. يجب أن تأخذ دالة `dfs` معلمتين:
   - مصفوفة مجاورة غير موجهة.
   - مسمى عقدة، وهو القيمة الرقمية للعقدة بين `0` و `n - 1`، حيث `n` هو العدد الكلي للعقد في الرسم البياني.
1. يجب أن تنفذ دالة `dfs` خوارزمية البحث بالعمق أولاً لإخراج قائمة بكل العقد التي يمكن الوصول إليها من العقدة المُمررة إليها.

# --hints--

يجب أن يكون لديك دالة باسم `dfs` تأخذ معلمتين.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs)
  sig = inspect.signature(dfs)
  assert len(sig.parameters) == 2
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` يجب أن تُرجع قائمة تحتوي على `1`، `2`، `3`، و `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {1, 2, 3, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` يجب أن تُرجع قائمة تحتوي على `1`، `2`، `3`، و `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {3, 2, 1, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` يجب أن تُرجع `[3]`.

```js
({ test: () => runPython(`
  assert dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3) == [3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` يجب أن تُرجع قائمة تحتوي على `3` و `2`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)
  assert result == [3, 2] or result == [2, 3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` يجب أن تُرجع قائمة تحتوي على `0` و `1`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)
  assert result == [0, 1] or result == [1, 0]
`) })
```

يجب أن تُرجع دالة `dfs` النتائج الصحيحة.

```js
({ test: () => runPython(`
  import random

  def solve(g, s):
      v, q = set(), [s]
      while q:
          u = q.pop()
          if u not in v:
              v.add(u)
              q.extend([i for i, x in enumerate(g[u]) if x == 1 and i not in v])
      return v

  random.seed(0)

  for _ in range(10):
      n = 6
      mat = [[0]*n for _ in range(n)]
      for i in range(n):
          for j in range(i+1, n):
              if random.random() > 0.5:
                  mat[i][j] = mat[j][i] = 1

      start = random.randint(0, n-1)
      assert set(dfs(mat, start)) == solve(mat, start), "Random graph check failed"
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs(graph, root):
    stack = []
    temp_v = None
    visited = []
    temp_v_neighbors = []
    stack.append(root)
    while stack:
        temp_v = stack.pop()
        if temp_v not in visited:
            visited.append(temp_v)
            temp_v_neighbors = graph[temp_v]
            for n, is_neighbor in enumerate(temp_v_neighbors):
                if is_neighbor == 1:
                    stack.append(n)
    return visited
```
