---
id: 68c343692285523aef0fc6bf
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

`scope` 屬性用於指定標題儲存格是否為列、行或多列或多行群組的標題。以下是一個範例：

```html
<th scope="col">Example Header</th>
```

這有助於螢幕閱讀器理解標題與資料儲存格之間的關係。

為所有 `th` 元素添加一個值為 `col` 的 `scope` 屬性。

# --hints--

每個 `th` 元素都應該有一個值為 `"col"` 的 `scope` 屬性。

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
