---
id: 68c346d03709343bff132707
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在接下來的幾個步驟中，你將建置表格的主體。

首先在 `thead` 元素下方添加一個 `tbody` 元素。然後在 `tbody` 內部添加一個 `tr` 元素。

# --hints--

你應該在 `thead` 元素下方有一個 `tbody` 元素。

```js
assert.exists(document.querySelector("tbody"));
```

你應該在 `tbody` 元素內有一個 `tr` 元素。

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
