---
id: 68c341332d98fe38fae733af
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

現在是開始建置表格的時候了。 

首先在你的 `h1` 元素下方添加一個 `table` 元素。在 `table` 元素內，添加一個帶有文字 `Schedule by Track and Time` 的 `caption` 元素。

# --hints--

你應該有一個 `table` 元素。

```js
assert.exists(document.querySelector("table"));
```

你的 `table` 元素應該有一個 `caption` 元素。

```js
assert.exists(document.querySelector("table caption"));
```

你的 `caption` 元素應該有文字 `Schedule by Track and Time`。

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
