---
id: 68c34d57ed874e4004f7b5d4
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

目前，`td` 元素中包含 `Break` 的文字內容只跨越一個行。但如果它能跨越全部三個行會更好。 

如你從先前的工作坊和課程中所回想，你可以使用 `colspan` 屬性讓表格儲存格跨越多個行。

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

為你的 `td` 元素添加 `colspan` 屬性並將其值設為 `3`。

# --hints--

你的 `td` 元素中，`Break` 的文字內容應該設定 `colspan` 屬性為 `3`。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
