---
id: 68c34f4aeb2b1a4170506751
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

接著，為表格添加第五列。先添加另一個 `tr` 元素。在該 `tr` 元素內，添加一個具有 `scope` 屬性且設為 `"row"` 的 `th` 元素，文字內容為 `12:30 PM`。然後在該 `th` 元素下方，添加一個文字內容為 `Lunch Break` 的 `td` 元素。你的 `td` 元素也應該有一個設為 `3` 的 `colspan` 屬性，以便跨越所有三個軌道。

# --hints--

你應該在 `tbody` 元素內有第五個 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

你應該在第五個 `tr` 元素內有一個 `th` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

你的 `th` 元素應該有 `12:30 PM` 的文字內容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

你的 `th` 元素應該有一個 `scope` 屬性設定為 `"row"`。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

你應該在第五個 `tr` 元素內有一個 `td` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

你的 `td` 元素應該有 `Lunch Break` 的文字內容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

你的 `td` 元素應該有一個設定為 `3` 的 `colspan` 屬性。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
