---
id: 68c34c608cf9e53f5d36b1d5
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

接著，為表格添加第三列。首先添加另一個 `tr` 元素。在該 `tr` 元素內，添加一個 `th` 元素，其 `scope` 屬性設為 `"row"`，並且文字內容為 `11:00 AM`。然後在該 `th` 元素下方，添加一個文字內容為 `Break` 的 `td` 元素。 

# --hints--

你應該在 `tbody` 元素內有第三個 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

你應該在第三個 `tr` 元素內有一個 `th` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

你的 `th` 元素應該有 `11:00 AM` 的文字內容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

你的 `th` 元素應該有一個 `scope` 屬性設定為 `"row"`。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

你應該在第三個 `tr` 元素內有一個 `td` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

你的 `td` 元素應該有 `Break` 的文字內容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
