---
id: 68c34329d430583a6cd5f7f1
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在你的 `tr` 元素內，為你的 `tr` 元素添加四個 `th` 元素。第一個 `th` 元素應該有 `Time` 的文字，第二個應該有 `Track A` 的文字，第三個應該有 `Track B` 的文字，第四個應該有 `Track C` 的文字。

# --hints--

你的 `tr` 元素應該有四個 `th` 元素。

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

你的第一個 `th` 元素應該有 `Time` 的文字。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

你的第二個 `th` 元素應該有 `Track A` 的文字。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

你的第三個 `th` 元素應該有 `Track B` 的文字。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

你的第四個 `th` 元素應該有 `Track C` 的文字。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
