---
id: 68c34e201b347c40acda8111
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

現在是為表格添加另一列的時候了。 

首先為你的 `tr` 元素添加另一個 `tr` 元素。在該 `tr` 元素內，添加一個帶有 `scope` 屬性且設為 `"row"` 的 `th` 元素，並且文字內容為 `11:30 AM`。 

然後，添加三個帶有以下文字內容的 `td` 元素：

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

你應該在 `tbody` 元素內有第四個 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

你應該在第四個 `tr` 元素內有一個 `th` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

你的 `th` 元素應該有 `11:30 AM` 的文字內容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

你的 `th` 元素應該有一個 `scope` 屬性設定為 `"row"`。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

你的第四個 `tr` 元素應該有三個 `td` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

你的第一個 `td` 元素應該有 `AR/VR in Education` 的文字內容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

你的第二個 `td` 元素應該有 `JavaScript Fundamentals` 的文字內容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

你的第三個 `td` 元素應該有 `Design Systems at Scale` 的文字內容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
