---
id: 68c3420cb0176439a95af1b8
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

下一步是為你的表格添加表頭。 

首先在 `caption` 元素下方添加一個 `thead` 元素。在 `thead` 內部添加一個 `tr` 元素。 

# --hints--

你應該有一個 `thead` 元素。

```js
assert.exists(document.querySelector("table thead"));
```

你應該在 `thead` 裡面有一個 `tr` 元素。

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
