---
id: 68c34a35c91b053dbae2a8c6
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

`scope` 屬性的另一個值是 `row`，它顯示表頭儲存格是其整個列的表頭。

在你的 `th` 元素內，為你的 `scope` 屬性添加值為 `row`。

# --hints--

你的 `th` 元素應該有一個值為 `row` 的 `scope` 屬性。

```js
assert.equal(document.querySelector("tbody tr th")?.getAttribute("scope"), "row");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        <th>9:00 AM</th>
        --fcc-editable-region--
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
