---
id: 68c3482278a14e3cf225e897
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

在你的 `tr` 元素內，為你的 `tr` 元素添加一個帶有 `9:00 AM` 文字的 `th` 元素。然後在該 `th` 元素下方，添加三個帶有 `Keynote: Tech Future`、`Intro to Web Dev` 和 `UX for All` 文字的 `td` 元素。 

# --hints--

你應該在你的 `tr` 元素內有一個 `th` 元素。

```js
assert.exists(document.querySelector("tbody tr th"));
```

你的 `th` 元素應該有 `9:00 AM` 的文字內容。

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

你應該在你的 `th` 元素下方有總共三個 `td` 元素。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

你的第一個 `td` 元素應該有 `Keynote: Tech Future` 的文字內容。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

你的第二個 `td` 元素應該有 `Intro to Web Dev` 的文字內容。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

你的第三個 `td` 元素應該有 `UX for All` 的文字內容。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
