---
id: 68c33e502cd97f2fd52cb63b
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次工作坊中，你將學習如何透過建置一個技術會議的 `schedule` 來製作可存取的表格元素。

首先為你的文件添加一個帶有文字 `Tech Conference 2025 Schedule` 的 `h1` 元素。

# --hints--

你應該有一個 `h1` 元素。

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

你的 `h1` 元素應該有文字 `Tech Conference 2025 Schedule`。

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
