---
id: a789b3483989747d63b0e427
title: Construye el Buscador del Número Más Grande
challengeType: 26
dashedName: build-the-largest-number-finder
---

# --description--

En este laboratorio, construirás una función que devuelve un arreglo que consiste en el número más grande de cada sub-arreglo proporcionado.

Recuerda, puedes iterar a través de un arreglo con un simple ciclo `for`, y acceder a cada elemento con la sintaxis de arreglo `arr[i]`.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función `largestOfAll` que tome un arreglo de arreglos como argumento.
2. La función debe devolver un arreglo que contenga el número más grande de cada sub-arreglo.

# --hints--

Debes crear una función llamada `largestOfAll`.

```js
assert.isFunction(largestOfAll);
```

`largestOfAll` debería tomar un solo parámetro.

```js
assert.lengthOf(largestOfAll, 1);
```

`largestOfAll([[4, 5, 1, 3], [13, 27, 18, 26], [32, 35, 37, 39], [1000, 1001, 857, 1]])` debería devolver un arreglo.

```js
assert.isArray(
  largestOfAll([
    [4, 5, 1, 3],
    [13, 27, 18, 26],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ])
);
```

`largestOfAll([[13, 27, 18, 26], [4, 5, 1, 3], [32, 35, 37, 39], [1000, 1001, 857, 1]])` debería devolver `[27, 5, 39, 1001]`.

```js
assert.deepEqual(
  largestOfAll([
    [13, 27, 18, 26],
    [4, 5, 1, 3],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ]),
  [27, 5, 39, 1001]
);
```

`largestOfAll([[4, 9, 1, 3], [13, 35, 18, 26], [32, 35, 97, 39], [1000000, 1001, 857, 1]])` debería devolver `[9, 35, 97, 1000000]`.

```js
assert.deepEqual(
  largestOfAll([
    [4, 9, 1, 3],
    [13, 35, 18, 26],
    [32, 35, 97, 39],
    [1000000, 1001, 857, 1]
  ]),
  [9, 35, 97, 1000000]
);
```

`largestOfAll([[17, 23, 25, 12], [25, 7, 34, 48], [4, -10, 18, 21], [-72, -3, -17, -10]])` debería devolver `[25, 48, 21, -3]`.

```js
assert.deepEqual(
  largestOfAll([
    [17, 23, 25, 12],
    [25, 7, 34, 48],
    [4, -10, 18, 21],
    [-72, -3, -17, -10]
  ]),
  [25, 48, 21, -3]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function largestOfAll(arr) {
  return arr.map(subArr => Math.max.apply(null, subArr));
}

largestOfAll([
  [4, 5, 1, 3],
  [13, 27, 18, 26],
  [32, 35, 37, 39],
  [1000, 1001, 857, 1]
]);
```
