---
id: 67d83df6f82eda3868dd2a84
title: RPGキャラクターを作成する
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

このラボでは、Pythonの基本を練習しながら、RPGアドベンチャー用のキャラクターを作成する小さなアプリを作ります。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `create_character`という名前の関数を作成してください。
1. 関数は、キャラクター名の後に、順番に力、知性、魅力の3つのステータスを引数として受け取る必要があります。
1. キャラクター名は次のように検証してください：
   - キャラクター名が文字列でない場合、関数は`The character name should be a string`を返す必要があります。
   - キャラクター名が空文字列の場合、関数は`The character should have a name`を返す必要があります。
   - キャラクター名が10文字を超える場合、関数は`The character name is too long`を返す必要があります。
   - キャラクター名にスペースが含まれている場合、関数は`The character name should not contain spaces`を返す必要があります。
1. ステータスも次のように検証してください：
   - 1つ以上のステータスが整数でない場合、関数は`All stats should be integers`を返す必要があります。
   - 1つ以上のステータスが1未満の場合、関数は`All stats should be no less than 1`を返す必要があります。
   - 1つ以上のステータスが4を超える場合、関数は`All stats should be no more than 4`を返す必要があります。
   - すべてのステータスの合計が7でない場合、関数は`The character should start with 7 points`を返す必要があります。
1. すべての値が検証を通過した場合、関数は4行の文字列を返す必要があります：
   - 1行目はキャラクター名を含みます
   - 2～4行目はステータスの略称、`STR`、`INT`、または`CHA`（この順番）で始まり、その後にスペース、そしてステータスの値に等しい数の満点のドット（`●`）、さらに合計10個になるように空点のドット（`○`）が続きます。例：力の値が3の場合、満点のドットが3つ、その後に空点のドットが7つ続きます。ドットはエディターに用意されています。

`create_character('ren', 4, 2, 1)`が返すべき文字列は以下の通りです：

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

注意：`str`と`int`はステータスの一般的な略称ですが、Pythonの予約キーワードなので変数名として使わないでください。

# --hints--

`create_character`という名前の関数を作成してください。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

`create_character`が最初の引数に**文字列以外**を受け取った場合、`The character name should be a string`を返す必要があります。

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

`create_character`が最初の引数に文字列を受け取った場合、`The character name should be a string`を返してはいけません。

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

`create_character`が最初の引数に空文字列を受け取った場合、`The character should have a name`を返す必要があります。

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

`create_character`が最初の引数に空文字列以外を受け取った場合、`The character should have a name`を返してはいけません。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

`create_character`が最初の引数に10文字を超える文字列を受け取った場合、`The character name is too long`を返す必要があります。

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

`create_character`関数は、10文字を超えない場合に「キャラクター名が長すぎる」と言ってはいけません。

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

`create_character`が最初の引数にスペースを含む文字列を受け取った場合、`The character name should not contain spaces`を返す必要があります。

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

`create_character`が最初の引数にスペースを含まない文字列を受け取った場合、`The character name should not contain spaces`を返してはいけません。

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

`create_character`が2番目、3番目、または4番目の引数に**整数以外**を受け取った場合、`All stats should be integers`を返す必要があります。

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

`create_character`が2番目、3番目、4番目の引数にすべて整数を受け取った場合、`All stats should be integers`を返してはいけません。

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

`create_character`が2番目、3番目、または4番目の引数に`1`未満の値を受け取った場合、`All stats should be no less than 1`を返す必要があります。

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

`create_character`が2番目、3番目、4番目の引数にすべて`1`以上の値を受け取った場合、`All stats should be no less than 1`を返してはいけません。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

`create_character`が2番目、3番目、または4番目の引数に`4`を超える値を受け取った場合、`All stats should be no more than 4`を返す必要があります。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

`create_character`が2番目、3番目、4番目の引数にすべて`4`以下の値を受け取った場合、`All stats should be no more than 4`を返してはいけません。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

`create_character`が2番目、3番目、4番目の引数の合計が`7`でない場合、`The character should start with 7 points`を返す必要があります。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

`create_character`が2番目、3番目、4番目の引数の合計が`7`の場合、`The character should start with 7 points`を返してはいけません。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)`は`ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`を返す必要があります。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

`create_character`が有効な値で呼び出された場合、要求通りにキャラクターのステータスを出力する必要があります。

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
