---
id: 698b42492a7f5f737895892e
title: 隣接リストから隣接行列への変換器を作成する
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

このラボでは、グラフの隣接リスト表現を隣接行列に変換する関数を作成します。隣接リストは、各キーがノードを表し、そのキーのノードに接続されているノードの配列が対応する値となっているオブジェクトです。隣接行列は2次元配列で、位置`[i][j]`の要素は、ノード`1`からノード`i`へのエッジがある場合は`j`、そうでなければ`0`となります。

例えば、次の隣接リストが与えられた場合：

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

対応する隣接行列は次のようになります：

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. 隣接リストを隣接行列に変換するための`adjacencyListToMatrix`という名前の関数を定義してください。
2. 関数は、重みなし（無向または有向）グラフの隣接リストを表すオブジェクトを引数として受け取る必要があります。
3. 関数は以下を行う必要があります：
   - 隣接リストを隣接行列に変換する。
   - 隣接行列の各行を出力する。
   - 隣接行列を返す。

例えば、`adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})`は次のように出力し：

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

`[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`を返すべきです。

# --hints--

`adjacencyListToMatrix`という名前の関数を定義してください。

```js
assert.isFunction(adjacencyListToMatrix);
```

`adjacencyListToMatrix`関数は1つのパラメータを持つべきです。

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

関数は隣接リストから正しくノード数を判別する必要があります。

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

関数は存在するエッジに対して行列の値を`1`に正しく設定する必要があります。

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

関数は行列の各行を出力する必要があります。

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

関数は隣接行列を返す必要があります。

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

隣接リスト`{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}`が与えられた場合、関数は`[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`を返すべきです。

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

隣接リスト`{0: [1], 1: [0]}`が与えられた場合、関数は`[[0, 1], [1, 0]]`を返すべきです。

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

隣接リスト`{0: [], 1: [], 2: []}`が与えられた場合、関数は`[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`を返すべきです。

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
