---
id: 696920c0c98a1ed58eb86293
title: Produktpräsentation erstellen
challengeType: 25
dashedName: lab-product-showcase
demoType: onClick
---

# --description--

Für dieses Labor wurden Ihnen bereits das gesamte HTML und CSS bereitgestellt. Sie werden eine Produktpräsentation mit TypeScript vervollständigen.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein Interface namens `Item` erstellen mit:
   - Einer `type`-Eigenschaft, die nur `"book"`, `"electronics"` oder `"clothing"` sein kann.
   - Einer `id`-Eigenschaft vom Typ `string`.
   - Einer `price`-Eigenschaft vom Typ `number`.

1. Sie sollten ein Interface namens `Book` erstellen, das `Item` erweitert. Seine `type`-Eigenschaft sollte auf das Literal `"book"` gesetzt sein. Es sollte außerdem `title`- und `author`-Eigenschaften vom Typ `string` haben.

1. Sie sollten ein Interface namens `Electronics` erstellen, das `Item` erweitert. Seine `type`-Eigenschaft sollte auf das Literal `"electronics"` gesetzt sein. Es sollte außerdem `item`- und `model`-Eigenschaften vom Typ `string` sowie eine optionale `warranty`-Eigenschaft vom Typ `number` haben.

1. Sie sollten ein Interface namens `Clothing` erstellen, das `Item` erweitert. Seine `type`-Eigenschaft sollte auf das Literal `"clothing"` gesetzt sein. Es sollte außerdem `item`- und `brand`-Eigenschaften vom Typ `string` sowie eine optionale `size`-Eigenschaft haben, die nur `"S"`, `"M"` oder `"L"` sein kann.

1. Sie sollten einen Union-Typ namens `Product` erstellen, der die Interfaces `Book`, `Electronics` und `Clothing` kombiniert.

1. Sie sollten eine generische Klasse namens `Collection` mit einem Typparameter `T` erstellen. Der Konstruktor sollte ein Array von Elementen vom Typ `T` akzeptieren und diese in einer Eigenschaft namens `items` speichern.

1. Die `Collection`-Klasse sollte eine `getAll()`-Methode haben, die alle Elemente in der Sammlung zurückgibt.

1. Die `Collection`-Klasse sollte eine `filter()`-Methode haben, die einen Callback akzeptiert und nur die Elemente aus dem `items`-Array zurückgibt, die die gegebene Bedingung erfüllen. Der Callback sollte ein Element vom Typ `T` erhalten und einen Boolean zurückgeben.

1. Sie sollten eine Funktion namens `renderProduct` erstellen mit einem Parameter vom Typ `Product`. Sie sollte einen HTML-String zurückgeben, der ein Produkt darstellt mit:
   - Einem Element mit der Klasse `item` und einer ID mit dem Wert der Produkt-`id`.
   - Einem Element mit der Klasse `price` und dem Text des Produkt-`price`.
   - Einem Element, das zusätzliche Produktinformationen enthält.

1. Die `renderProduct`-Funktion sollte Typverengung verwenden, um je nach Produkttyp unterschiedliches HTML zu rendern:

   - Für Bücher sollte das gerenderte HTML das Format `Book: [title] by [author]` haben, wobei `[title]` und `[author]` durch den Titel bzw. den Autor ersetzt werden.
   - Für Elektronik sollte das gerenderte HTML das Format `Electronics: [item] - [model]` haben. Wenn das Produkt eine Garantie hat, fügen Sie ` - Warranty: [warranty] year(s)` an das Format an. Ersetzen Sie `[item]`, `[model]` und `[warranty]` durch den Artikel, das Modell und die Garantie.
   - Für Kleidung sollte das gerenderte HTML das Format `Clothing: [item] by [brand]` haben. Wenn das Produkt eine Größe hat, fügen Sie ` - Size [size]` an das Format an. Ersetzen Sie `[item]`, `[brand]` und `[size]` durch den Artikel, die Marke und die Größe.
   - Wenn ein ungültiger Produkttyp gefunden wird, werfen Sie einen Fehler mit dem Format: `Unknown product type: [product]`, wobei `[product]` durch das ungültige Produktobjekt ersetzt wird. Verwenden Sie `JSON.stringify`, um das ungültige Produkt in einen String umzuwandeln.

1. Sie sollten eine `Collection<Product>`-Instanz erstellen, die mindestens einen Artikel für jeden Produkttyp enthält, und die Instanz einer Variablen namens `products` zuweisen.

1. Sie sollten eine `showProducts`-Funktion erstellen, die Produkte im `#output`-Element anzeigt. Sie sollte einen optionalen Filter akzeptieren, der einen Produkttyp darstellt; wenn dieser weggelassen wird, zeigen Sie alle Produkte an; wenn angegeben, zeigen Sie nur Produkte dieses Typs an.

1. Sie sollten Produkte rendern, indem Sie jedes Produkt mit der `renderProduct`-Funktion in einen HTML-String umwandeln und das kombinierte Ergebnis in das `#output`-Element einfügen.

1. Das Klicken auf die `#all`-, `#books`-, `#electronics`- und `#clothing`-Buttons sollte die entsprechende Produktmenge im `#output`-Element anzeigen.

1. Beim Laden der Seite sollten standardmäßig alle Produkte angezeigt werden.

# --hints--

Sie sollten ein Interface namens `Item` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Item);
```

Ihr `Item`-Interface sollte eine `type`-Eigenschaft haben, die nur `"book"`, `"electronics"` oder `"clothing"` sein kann.

```js
const explorer = await __helpers.Explorer(code);
const { Item } = explorer.interfaces;
const typeProp = Item.typeProps.type;
assert.isTrue(typeProp.isUnionOf(["'book'", "'electronics'", "'clothing'"]));
```

Ihr `Item`-Interface sollte eine `id`-Eigenschaft vom Typ `string` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "id", type: "string"}));
```

Ihr `Item`-Interface sollte eine `price`-Eigenschaft vom Typ `number` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "price", type: "number"}));
```

Sie sollten ein Interface namens `Book` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Book);
```

Ihr `Book`-Interface sollte `Item` erweitern.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.doesExtend("Item"));
```

Ihr `Book`-Interface sollte eine `type`-Eigenschaft haben, die auf `"book"` gesetzt ist.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "type", type: "'book'"}));
```

Ihr `Book`-Interface sollte eine `title`-Eigenschaft vom Typ `string` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "title", type: "string"}));
```

Ihr `Book`-Interface sollte eine `author`-Eigenschaft vom Typ `string` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "author", type: "string"}));
```

Sie sollten ein Interface namens `Electronics` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Electronics);
```

Ihr `Electronics`-Interface sollte `Item` erweitern.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.doesExtend("Item"));
```

Ihr `Electronics`-Interface sollte eine `type`-Eigenschaft haben, die auf `"electronics"` gesetzt ist.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "type", type: "'electronics'"}));
```

Ihr `Electronics`-Interface sollte eine `item`-Eigenschaft vom Typ `string` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "item", type: "string"}));
```

Ihr `Electronics`-Interface sollte eine `model`-Eigenschaft vom Typ `string` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "model", type: "string"}));
```

Ihr `Electronics`-Interface sollte eine optionale `warranty`-Eigenschaft vom Typ `number` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "warranty", type: "number", isOptional: true}));
```

Sie sollten ein Interface namens `Clothing` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Clothing);
```

Ihr `Clothing`-Interface sollte `Item` erweitern.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.doesExtend("Item"));
```

Ihr `Clothing`-Interface sollte eine `type`-Eigenschaft haben, die auf `"clothing"` gesetzt ist.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "type", type: "'clothing'"}));
```

Ihr `Clothing`-Interface sollte eine `item`-Eigenschaft vom Typ `string` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "item", type: "string"}));
```

Ihr `Clothing`-Interface sollte eine `brand`-Eigenschaft vom Typ `string` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "brand", type: "string"}));
```

Ihr `Clothing`-Interface sollte eine optionale `size`-Eigenschaft haben, die nur `"S"`, `"M"` oder `"L"` sein kann.

```js
const explorer = await __helpers.Explorer(code);
const { Clothing } = explorer.interfaces;
assert.isTrue(Clothing.hasTypeProps({name: "size", isOptional: true}));
assert.isTrue(Clothing.typeProps.size.isUnionOf(["'S'", "'M'", "'L'"]));
```

Sie sollten einen Union-Typ namens `Product` erstellen, der die Interfaces `Book`, `Electronics` und `Clothing` kombiniert.

```js
const explorer = await __helpers.Explorer(code);
const { Product } = explorer.types;
assert.isTrue(Product.isUnionOf(["Book", "Electronics", "Clothing"]));
```

Sie sollten eine `Collection`-Klasse mit einem Typparameter `T` haben.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.isTrue(Collection.typeParameters[0].matches("T"));
```

Der Konstruktor Ihrer `Collection`-Klasse sollte ein Array von Elementen vom Typ `T` akzeptieren und diese in einer Eigenschaft namens `items` speichern.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.items, arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.items, arr);
```

Ihre `Collection`-Klasse sollte eine `getAll`-Methode haben.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
console.log(Collection.methods)
console.log(Collection.methods.getAll)
assert.exists(Collection.methods.getAll);
```

Ihre `getAll`-Methode sollte alle Elemente in der Sammlung zurückgeben.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.getAll(), arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.getAll(), arr);
```

Ihre `Collection`-Klasse sollte eine `filter`-Methode haben.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.exists(Collection.methods.filter);
```

Ihre `filter`-Methode sollte eine Callback-Funktion akzeptieren und ein Array mit den Elementen der `items`-Eigenschaft zurückgeben, die die gegebene Bedingung erfüllen.

```js
let arr = [1, 2, 3, 4, 5, 6];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.filter(n => n % 2 === 0), [2, 4, 6]);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.filter(s => s === "c"), ["c"]);
```

Sie sollten eine Funktion namens `renderProduct` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.renderProduct);
```

`renderProduct(product)` sollte einen String zurückgeben, der ein Element mit der Klasse `item` und der ID gleich `product.id` enthält.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  assert.isNotNull(testEl.querySelector(`[class="item"][id="${p.id}"]`));
})
```

`renderProduct(product)` sollte einen String zurückgeben, der ein Element mit der Klasse `price` und dem Text gleich `product.price` enthält.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const price = testEl.querySelector(`[class="price"]`);
  assert.include(price.innerText, p.price.toString());
})
```

Wenn `renderProduct` mit einem `Book`-Produkt `p` aufgerufen wird, sollte es einen String zurückgeben, der ein Element mit dem Text im Format `Book: [title] by [author]` enthält. Ersetzen Sie `[title]` und `[author]` durch die entsprechenden Eigenschaften von `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const expectedText = `Book: ${p.title} by ${p.author}`;
  assert.include(testEl.innerText, expectedText);
})
```

Wenn `renderProduct` mit einem `Electronics`-Produkt `p` aufgerufen wird, sollte es einen String zurückgeben, der ein Element mit dem Text im Format `Electronics: [item] - [model]` oder `Electronics: [item] - [model] - Warranty: [warranty] year(s)` enthält, falls `p.warranty` nicht undefined ist. Ersetzen Sie `[item]`, `[model]` und `[warranty]` durch die entsprechenden Eigenschaften von `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Electronics: ${p.item} - ${p.model}`;
  const expectedText = !p.warranty ? baseFormat : baseFormat + ` - Warranty: ${p.warranty} year(s)`;
  assert.include(testEl.innerText, expectedText);
  if (!p.warranty) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

Wenn `renderProduct` mit einem `Clothing`-Produkt `p` aufgerufen wird, sollte es einen String zurückgeben, der ein Element mit dem Text im Format `Clothing: [item] by [brand]` oder `Clothing: [item] by [brand] - Size [size]` enthält, falls `p.size` nicht undefined ist. Ersetzen Sie `[item]`, `[brand]` und `[size]` durch die entsprechenden Eigenschaften von `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Clothing: ${p.item} by ${p.brand}`;
  const expectedText = !p.size ? baseFormat : baseFormat + ` - Size ${p.size}`;
  assert.include(testEl.innerText, expectedText);
  if (!p.size) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

Wenn `renderProduct` mit einem ungültigen Produkt aufgerufen wird, sollte es einen Fehler mit dem Format `Unknown product type: [p]` werfen. Ersetzen Sie `[p]` durch das ungültige Produkt, das mit `JSON.stringify` in einen String umgewandelt wurde.

```js
const invalidProduct = { item: "T-Shirt", price: 19.99 };
const expected = `Unknown product type: ${JSON.stringify(invalidProduct)}`;
try {
  renderProduct(invalidProduct);
  assert.fail();
} catch({name, message}) {
  if (name == "AssertionError") {
    assert.fail("renderProduct should throw with invalid product");
  } else {
    assert.equal(message, expected);
  }
}
```

Sie sollten eine `Collection<Product>`-Instanz erstellen und sie einer Variablen namens `products` zuweisen.

```js
assert.instanceOf(products, Collection);
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.products.value.typeArguments[0].matches("Product"));
```

Ihre `products`-Variable sollte mindestens einen Artikel für jeden Produkttyp enthalten.

```js
assert.isTrue(products.items.some(p => p.type === "book"));
assert.isTrue(products.items.some(p => p.type === "electronics"));
assert.isTrue(products.items.some(p => p.type === "clothing"));
```

Sie sollten eine Funktion namens `showProducts` haben.

```js
assert.isFunction(showProducts);
```

Das Klicken auf den `#all`-Button sollte alle Produkte im `#output`-Element anzeigen.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
books.dispatchEvent(new Event("click"));
all.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

Das Klicken auf den `#books`-Button sollte die entsprechende Produktmenge im `#output`-Element anzeigen.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "book") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
books.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Das Klicken auf den `#electronics`-Button sollte die entsprechende Produktmenge im `#output`-Element anzeigen.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "electronics") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
electronics.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Das Klicken auf den `#clothing`-Button sollte die entsprechende Produktmenge im `#output`-Element anzeigen.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "clothing") {
    count++;
  }
});
books.dispatchEvent(new Event("click"));
clothing.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Alle Produkte sollten beim Laden der Seite standardmäßig sichtbar sein. Verwenden Sie das `DOMContentLoaded`-Ereignis, um dies sicherzustellen.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
document.dispatchEvent(new Event("DOMContentLoaded"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts
interface Item {
  type: "book" | "electronics" | "clothing";
  id: string;
  price: number;
}

interface Book extends Item {
  type: "book";
  title: string;
  author: string;
}

interface Electronics extends Item {
  type: "electronics";
  item: string;
  model: string;
  warranty?: number;
}

interface Clothing extends Item {
  type: "clothing";
  item: string;
  brand: string;
  size?: "S" | "M" | "L";
}

type Product = Book | Electronics | Clothing;

class Collection<T> {
  items: T[];
  constructor(items: T[]) {
    this.items = items;
  }

  getAll(): T[] {
    return this.items;
  }

  filter(callback: (item: T) => boolean): T[] {
    return this.items.filter(callback);
  }
}

function renderProduct(p: Product): string {
  const createProductCard = (id: string, content: string, price: number): string => {
    return `
     <div class="item" id="${id}">
       ${content}
       <div class="price">$${price}</div>
     </div>
   `;
  };

  if (p.type === "book") {
    const content = `<strong>Book:</strong> ${p.title} by ${p.author}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "electronics") {
    const warranty = p.warranty ? ` - Warranty: ${p.warranty} year(s)` : "";
    const content = `<strong>Electronics:</strong> ${p.item} - ${p.model}${warranty}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "clothing") {
    const size = p.size ? ` - Size ${p.size}` : "";
    const content = `<strong>Clothing:</strong> ${p.item} by ${p.brand}${size}`;
    return createProductCard(p.id, content, p.price);
  }

  const _never: never = p;
  throw new Error(`Unknown product type: ${JSON.stringify(p)}`);
}


const products = new Collection<Product>([
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },

  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "b2", type: "book", title: "1984", author: "George Orwell", price: 9.99 },

  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 },

  { id: "b3", type: "book", title: "Brave New World", author: "Aldous Huxley", price: 11.99 },
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "e5", type: "electronics", item: "Smartwatch", model: "Chronex Pulse-Q", warranty: 1, price: 199.99 },

  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 },
]);

const output = document.getElementById("output");

// Track currently active filter category (or undefined for all)
let currentFilter: Product["type"] | undefined = undefined;

function showProducts(filter?: Product["type"]): void {
  if (!output) return;

  let itemsToShow: Product[];

  // Toggle logic
  if (currentFilter === filter || filter === undefined) {
    // Reset to all if same button clicked again or filter is undefined
    itemsToShow = products.getAll();
    currentFilter = undefined;
    setActiveButton(undefined);
  } else {
    itemsToShow = products.filter((p: Product): boolean => p.type === filter);
    currentFilter = filter;
    setActiveButton(filter);
  }

  output.innerHTML = itemsToShow.map(renderProduct).join("");
}

// Active Button Highlighting
function setActiveButton(filter: Product["type"] | undefined): void {
  const buttons: { [key: string]: HTMLElement | null } = {
    all: document.getElementById("all"),
    book: document.getElementById("books"),
    electronics: document.getElementById("electronics"),
    clothing: document.getElementById("clothing")
  };

  for (const key in buttons) {
    buttons[key]?.classList.remove("active");
  }

  if (filter === undefined) {
    buttons["all"]?.classList.add("active");
  } else {
    buttons[filter]?.classList.add("active");
  }
}

// Button Event Handlers
document.getElementById("all")?.addEventListener("click", (): void => showProducts());
document.getElementById("books")?.addEventListener("click", (): void => showProducts("book"));
document.getElementById("electronics")?.addEventListener("click", (): void => showProducts("electronics"));
document.getElementById("clothing")?.addEventListener("click", (): void => showProducts("clothing"));

// Initial Render
document.addEventListener("DOMContentLoaded", (): void => showProducts());
```
