---
id: 5688e62ea601b2482ff8422b
title: 建置一個評測查詢
challengeType:  26
dashedName: lab-profile-lookup
---

# --description--

在本實驗中，你將建置一個評測查詢，用於查詢聯絡人列表中有關人員的資訊。

在這個範例中，假設有一個名為 Akira Laine 的聯絡人，`lookUpProfile("Akira", "lastName")` 應該會傳回 `Laine`。 

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `lookUpProfile` 的函式（程式），該函式接受一個名稱和一個屬性作為引數。
2. 你應該從提供的 `contacts` 陣列中擷取聯絡資訊。
3. 如果函式接收到聯絡人名稱，且該屬性存在於相關聯絡人上，則應傳回該屬性的值。
4. 如果傳入函式的名稱與聯絡人陣列中的任何聯絡人都不符，則該函式應傳回 `"No such contact"`。
5. 如果在找到的聯絡人中不存在該屬性，則該函式應傳回 `"No such property"`。

# --before-each--

```js
const _contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

contacts = structuredClone(_contacts); 
```

# --hints--

`contacts` 陣列應該仍然存在於全域的範圍中。如果你刪除了它，請重設課程以恢復它。

```js
assert.isArray(contacts);
```

你應該有一個名為 `lookUpProfile` 的函式（程式）。

```js
assert.isFunction(lookUpProfile);
```

`lookUpProfile("Kristian", "lastName")` 應該傳回一個字串。

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

`lookUpProfile("Kristian", "lastName")` 應該傳回字串 `Vos`

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

`lookUpProfile("Sherlock", "likes")` 應該傳回 `["Intriguing Cases", "Violin"]`

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` 應該傳回一個陣列

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

`lookUpProfile("Bob", "number")` 應該傳回字串 `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

`lookUpProfile("Bob", "potato")` 應該傳回字串 `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

`lookUpProfile("Akira", "address")` 應該傳回字串 `No such property`

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

執行 `lookUpProfile` 後，`contacts` 陣列應保持不變。 

```js
lookUpProfile("Sherlock", "likes");
assert.deepEqual(contacts,_contacts);
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
