---
id: 579e2a2c335b9d72dd32e05c
title: Implementar el Algoritmo de Cortar y Ratificar
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

En este laboratorio, necesitas crear un algoritmo para fusionar dos arreglos.

Cumple con las historias de usuario a continuación y haz que todas las pruebas pasen para completar el laboratorio.

**Historias de Usuario:**

1. Crea una función `frankenSplice` que acepte dos arreglos y un índice.
1. Copia cada elemento del primer arreglo dentro del segundo arreglo, en orden, comenzando en el índice dado, y devuelve el arreglo resultante.
1. Los arreglos de entrada deben permanecer iguales después de que la función se ejecute.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` debería devolver `[4, 1, 2, 3, 5]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` debería devolver `["a", 1, 2, "b"]`.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` debería devolver `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

Todos los elementos del primer arreglo deben ser agregados al segundo arreglo en su orden original. `frankenSplice([1, 2, 3, 4], [], 0)` debería devolver `[1, 2, 3, 4]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

El primer arreglo debe permanecer igual después de que la función se ejecute.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

El segundo arreglo debe permanecer igual después de que la función se ejecute.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
