---
id: 66a207974c806a19d6607073
title: Tengeneza programu ya orodha ya vitabu
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kipengele cha `h1` chenye maandishi ya `Book Inventory`.
1. Unapaswa kuwa na kipengele cha `table` chenye safu zilizo na majina ya `Title`, `Author`, `Category`, `Status`, na `Rate`.
1. Jedwali lako linapaswa kuwa na angalau mistari minne, mstari wa kwanza kwa ajili ya vichwa vya safu na mingine ikiwa na taarifa za vitabu.
1. Kila mstari wa jedwali ndani ya mwili wa jedwali unapaswa kuwa na darasa la vitu la `read`, `to-read`, au `in-progress`.
1. Vipengele vya `td` vya safu ya `Status` vinapaswa kuwa na kipengele cha `span` chenye sifa ya kipengele cha `class` cha `status` kinachozunguka maandishi ya `Read`, `To Read`, au `In Progress`, kulingana na darasa la mstari huo.
1. Vipengele vya `td` vya safu ya `Rate` vinapaswa kuwa na kipengele cha `span` chenye sifa ya kipengele cha `class` cha `rate` kinachozunguka vipengele vitatu vya `span` vilivyo tupu.
1. Vipengele vya `.rate` vilivyowekwa ndani ya mistari ya `.read` vinapaswa kuwa na darasa la ziada lenye thamani ya `one`, `two`, au `three`, kulingana na kiwango binafsi. Thamani hii inapaswa kuja baada ya `rate`.
1. Unapaswa kuunda wachaguaji wa sifa za kipengele tatu ili kulenga vipengele vyenye darasa la `read`, `to-read`, na `in-progress`, na kuweka sifa yao ya `background-image` kutumia `linear-gradient` unayochagua.
1. Unapaswa kuweka sifa ya `display` ya kila kipengele cha `span` kuwa `inline-block`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `to-read` na kuweka sifa zao za `border` na `background-image`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `read` na kuweka sifa zao za `border` na `background-image`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `in-progress` na kuweka sifa zao za `border` na `background-image`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` na vipengele vya `span` vyenye thamani ya darasa kuanzia na `rate` na kuweka sifa zao za `height`, `width`, na `padding`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni watoto wa moja kwa moja wa vipengele vya `span` vyenye thamani ya `class` kuanzia na `rate` na kuweka sifa zao za `border`, `border-radius`, `margin`, `height`, `width`, na `background-color`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga mzazi wa kwanza wa vipengele vya `span` ambavyo vina neno `one` kama sehemu ya thamani yao ya `class` na kuweka sifa yake ya `background-image` kutumia `linear-gradient`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga wazazi wawili wa kwanza wa vipengele vya `span` ambavyo vina neno `two` kama sehemu ya thamani yao ya `class` na kuweka sifa zao za `background-image` kutumia `linear-gradient`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vitatu vya `span` ambavyo ni mzao wa vipengele vya `span` ambavyo vina neno `three` kama sehemu ya thamani yao ya `class` na kuweka sifa zao za `background-image` kutumia `linear-gradient`.

# --hints--

Unapaswa kuwa na kipengele cha `h1` chenye maandishi ya `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

Unapaswa kuwa na kipengele kimoja tu cha `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Unapaswa kuwa na kipengele cha `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Unapaswa kuwa na kipengele kimoja cha `thead` na kipengele kimoja cha `tbody` ndani ya `table`.

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

Ndani ya `thead` kunapaswa kuwa na kipengele kimoja cha `tr` chenye vipengele 5 vya `th`.

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

Safu yako ya kwanza inapaswa kuwa na maandishi ya `Title` kama kichwa cha sehemu.

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

Safu yako ya pili inapaswa kuwa na maandishi ya `Author` kama kichwa cha sehemu.

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

Safu yako ya tatu inapaswa kuwa na maandishi ya `Category` kama kichwa cha sehemu.

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

Safu yako ya nne inapaswa kuwa na maandishi ya `Status` kama kichwa cha sehemu.

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

Safu yako ya tano inapaswa kuwa na maandishi ya `Rate` kama kichwa cha sehemu.

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

Jedwali lako linapaswa kuwa na angalau mistari minne.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Kila mstari unapaswa kuwa na safu 5 kila wakati.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

Kila mstari wa jedwali isipokuwa mstari wa kichwa cha sehemu unapaswa kuwa na darasa la `read`, `to-read`, au `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

Vipengele vya `td` vya safu ya `Status` vinapaswa kuwa na kipengele cha `span`.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Kila kipengele cha `span` cha safu ya `Status` kinapaswa kuwa na darasa la `status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Kila kipengele cha `.status` kinapaswa kuwa na maandishi ya `Read`, `To Read`, au `In Progress`, kulingana na darasa la mstari wake.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

Vipengele vya `td` vya safu ya `Rate` vinapaswa kuwa na kipengele cha `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Kila kipengele cha `span` ambacho ni mtoto wa moja kwa moja wa kipengele cha `td` cha safu ya `Rate` kinapaswa kuwa na darasa la `rate` kama darasa la kwanza.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Kila kipengele cha `.rate` kinapaswa kuwa na vipengele vitatu vya `span` vilivyo tupu.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

Vipengele vya `.rate` vilivyowekwa ndani ya mistari ya `.read` vinapaswa kuwa na darasa la ziada baada ya darasa la `rate` lenye thamani ya `one`, `two`, au `three`, kulingana na kiwango binafsi.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga mistari yenye darasa la `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga mistari yenye darasa la `read` na kuweka sifa yao ya `background-image` kuwa mchanganyiko wa mstari wa rangi unayochagua.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga mistari yenye darasa la `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga mistari yenye darasa la `to-read` na kuweka sifa yao ya `background-image` kuwa mchanganyiko wa mstari wa rangi unayochagua.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga mistari yenye darasa la `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga mistari yenye darasa la `in-progress` na kuweka sifa yao ya `background-image` kuwa mchanganyiko wa mstari wa rangi unayochagua.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Unapaswa kuweka sifa ya `display` ya kila kipengele cha `span` kuwa `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `to-read` na kuweka sifa yao ya `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `to-read` na kuweka sifa yao ya `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `read` na kuweka sifa yao ya `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `read` na kuweka sifa yao ya `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `in-progress` na kuweka sifa yao ya `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `in-progress` na kuweka sifa yao ya `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` na vipengele vya `span` vyenye thamani ya darasa kuanzia na `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` na vipengele vya `span` vyenye thamani ya darasa kuanzia na `rate` na kuweka sifa yao ya `height`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` na vipengele vya `span` vyenye thamani ya darasa kuanzia na `rate` na kuweka sifa yao ya `width`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` na vipengele vya `span` vyenye thamani ya darasa kuanzia na `rate` na kuweka sifa yao ya `padding`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni watoto wa moja kwa moja wa vipengele vya `span` vyenye thamani ya `class` kuanzia na `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni watoto wa moja kwa moja wa vipengele vya `span` vyenye thamani ya `class` kuanzia na `rate` na kuweka sifa yao ya `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni watoto wa moja kwa moja wa vipengele vya `span` vyenye thamani ya `class` kuanzia na `rate` na kuweka sifa yao ya `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni watoto wa moja kwa moja wa vipengele vya `span` vyenye thamani ya `class` kuanzia na `rate` na kuweka sifa yao ya `margin`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni watoto wa moja kwa moja wa vipengele vya `span` vyenye thamani ya `class` kuanzia na `rate` na kuweka sifa yao ya `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni watoto wa moja kwa moja wa vipengele vya `span` vyenye thamani ya `class` kuanzia na `rate` na kuweka sifa yao ya `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni watoto wa moja kwa moja wa vipengele vya `span` vyenye thamani ya `class` kuanzia na `rate` na kuweka sifa yao ya `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga mzazi wa kwanza wa vipengele vya `span` ambavyo vina neno `one` kama sehemu ya thamani yao ya `class`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga mzazi wa kwanza wa vipengele vya `span` ambavyo vina neno `one` kama sehemu ya thamani yao ya `class` na kuweka sifa yake ya `background-image` kutumia `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga wazazi wawili wa kwanza wa vipengele vya `span` ambavyo vina neno `two` kama sehemu ya thamani yao ya `class`.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga wazazi wawili wa kwanza wa vipengele vya `span` ambavyo vina neno `two` kama sehemu ya thamani yao ya `class` na kuweka sifa zao za `background-image` kutumia `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

Unapaswa kuwa na kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni mzao wa vipengele vya `span` ambavyo vina neno `three` kama sehemu ya thamani yao ya `class`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni mzao wa vipengele vya `span` ambavyo vina neno `three` kama sehemu ya thamani yao ya `class` na kuweka sifa zao za `background-image` kutumia `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
