---
id: 67fccc2463253b213a000142
title: 构建一个多媒体播放器
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

在前面的课程中，你已经学习了如何使用 `audio` 和 `video` 元素。 在本实验中，你将构建一个多媒体播放器，它会展示一个 `audio` 曲目和一个带有字幕的 `video`。

对于 `audio` 元素，你需要包含一个 `source` 元素，用于指定所使用的媒体文件。

这是一个示例：

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

`source` 元素同样可以在 `video` 元素中使用，如下所示：

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个 `h1` 元素，作为页面的主标题。
2. 你应该有三个 `section` 元素。
3. 在第一个 `section` 元素中，应该包含一个 `h2` 元素，作为正在播放歌曲的标题。
4. 在该 `h2` 元素下方，应该有一个 `audio` 元素，带有 `controls` 属性和 `aria-label` 属性。
5. 在 `audio` 元素内部，应该有一个 `source` 元素，带有指向音频文件的 `src` 属性和 `type` 属性。 你可以使用此音频文件：`https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. 在第二个 `section` 元素中，应该包含一个 `h2` 元素，作为正在播放视频的标题。
7. 在该 `h2` 元素下方，应该有一个 `video` 元素，带有 `controls`、`width` 属性和 `aria-label` 属性。
8. 在 `video` 元素内部，应该有一个 `source` 元素，带有指向视频文件的 `src` 属性和 `type` 属性。 你可以使用此视频文件：`https://cdn.freecodecamp.org/curriculum/labs/What is the map method and how does it work.mp4`
9. 在 `source` 元素下方，应该有一个 `track` 元素，带有指向字幕文件的 `src` 属性、`kind` 属性、`srclang` 属性和 `label` 属性。
10. 在第三个 `section` 元素中，应该包含一个 `h2` 元素，作为该部分的标题，例如： "Transcript"。
11. 在该 `h2` 元素下方，应该有一个 `p` 元素，用来写视频的文稿。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.isNotNull(document.querySelector("h1"));
```

你的 `h1` 应包含一些代表页面主标题的文本。

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

你应该有三个 `section` 元素。

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

第一个 `section` 应该包含一个 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

你的 `h2` 应包含一些代表播放歌曲标题的文本。

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

第一个 `section` 应该包含一个 `audio` 元素。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

你的 `audio` 元素应该有 `controls` 属性。

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

你的 `audio` 元素应该有一个带描述性文本的 `aria-label` 属性。

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

你的 `audio` 元素应该包含一个 `source` 元素。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

你的 `source` 元素应该有一个指向音频文件的 `src` 属性。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

你的 `source` 元素应该有一个指定音频类型的 `type` 属性。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

第二个 `section` 应该包含一个 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

你的 `h2` 应包含一些代表播放视频标题的文本。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

第二个 `section` 应该包含一个 `video` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

你的 `video` 元素应该有 `controls` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

你的 `video` 元素应该有一个指定视频宽度的 `width` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

该 `video` 元素应该包含一个 `source` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

该 `source` 元素应该有一个指向视频文件的 `src` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

该 `source` 元素应该有一个指定视频类型的 `type` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

你的 `video` 元素应该包含一个 `track` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

你的 `track` 元素应该有一个指向字幕文件的 `src` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

你的 `track` 元素应该有一个描述文件类型的 `kind` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

你的 `track` 元素应该有一个描述字幕语言的 `srclang` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

你的 `track` 元素应该有一个描述字幕语言标签的 `label` 属性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

第三个 `section` 应该包含一个 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

你的 `h2` 应包含一些代表该部分标题的文本。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

第三个 `section` 应该包含一个 `p` 元素，用来写文稿。

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/What is the map method and how does it work.mp4" type="video/mp4">

      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
