---
id: 67ed0e7faf1f2d130a275a82
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

當使用者在輸入欄中輸入時，應該會出現符合他們查詢的水果列表。

為了控制代碼，創建一個名為 `results` 的狀態變數，初始值為空陣列（`[]`）。同時使用 `useState` 掛鉤定義其設值函式 `setResults`。

# --hints--

你應該為 `results` 和 `setResults` 變數使用陣列解構語法。

```js
assert.match(code, /(const|let|var)\s+\[\s*(results)\s*,\s*(setResults)\s*\]\s*/);
```

你應該使用 `useState` 掛鉤。

```js
  const abuseState = __helpers.spyOn(React, "useState");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.FruitsSearch);
  assert.isAtLeast(abuseState.calls.length, 2);
```

你的 `useState` 掛鉤應該有一個初始值為空陣列（`[]`）。

```js
	const abuseState = __helpers.spyOn(React, "useState");
	const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
	const exports = {};
	const _a = eval(script);
	const _b = await __helpers.prepTestComponent(exports.FruitsSearch);
	
	assert.deepEqual(abuseState.calls[1]?.[0], []);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

export function FruitsSearch() {
  --fcc-editable-region--
  const [query, setQuery] = useState('');
  
  --fcc-editable-region--
  return (
    <div id="search-container">
      <form>
        <label htmlFor="search-input">Search for fruits:</label>
        <input
          id="search-input"
          type="search"
        />
      </form>
    </div>
  );
}
```
