---
id: afd15382cdfb22c9efe8b7de
title: Реалізувати генератор пар ДНК
challengeType: 26
dashedName: implement-a-dna-pair-generator
---

# --description--

У подвійній спіралі ДНК основи завжди паруються: якщо на одному ланцюгу є основа <em>A</em>, то на іншому ланцюгу безпосередньо навпроти неї стоїть основа <em>T</em>, інша пара — це <em>C</em> і <em>G</em>.

У цій лабораторній роботі ви напишете функцію, яка підбере відсутні пари основ для заданого ланцюга ДНК. Для кожного символу у наданому рядку знайдіть пару основ.

Наприклад, для вхідних даних `ATCG` поверніть `[["A", "T"], ["T", "A"], ["C", "G"], ["G", "C"]]`

Основа <em>A</em> парується з основою <em>T</em>, основа <em>T</em> парується з основою <em>A</em>, основа <em>C</em> парується з основою <em>G</em>, а основа <em>G</em> парується з основою <em>C</em>.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати функцію `pairElement`, яка приймає рядок будь-якої довжини як аргумент.
1. Функція `pairElement` повинна повертати двовимірний масив, де кожен внутрішній масив містить два рядки: перший рядок — це одна основа з вхідних даних, а другий рядок — відповідна парна основа.
1. Якщо функції передати `A`, вона повинна повернути пару `T`.
1. Якщо функції передати `T`, вона повинна повернути пару `A`.
1. Якщо функції передати `C`, вона повинна повернути пару `G`.
1. Якщо функції передати `G`, вона повинна повернути пару `C`.

# --hints--

Вам слід створити функцію з ім’ям `pairElement`.

```js
assert.isFunction(pairElement);
```

`pairElement` має приймати один аргумент.

```js
assert.lengthOf(pairElement, 1);
```

`pairElement("ATCGA")` має повертати `[["A","T"],["T","A"],["C","G"],["G","C"],["A","T"]]`.

```js
assert.deepEqual(pairElement('ATCGA'), [
  ['A', 'T'],
  ['T', 'A'],
  ['C', 'G'],
  ['G', 'C'],
  ['A', 'T']
]);
```

`pairElement("TTGAG")` має повертати `[["T","A"],["T","A"],["G","C"],["A","T"],["G","C"]]`.

```js
assert.deepEqual(pairElement('TTGAG'), [
  ['T', 'A'],
  ['T', 'A'],
  ['G', 'C'],
  ['A', 'T'],
  ['G', 'C']
]);
```

`pairElement("CTCTA")` має повертати `[["C","G"],["T","A"],["C","G"],["T","A"],["A","T"]]`.

```js
assert.deepEqual(pairElement('CTCTA'), [
  ['C', 'G'],
  ['T', 'A'],
  ['C', 'G'],
  ['T', 'A'],
  ['A', 'T']
]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var lookup = Object.create(null);
lookup.A = 'T';
lookup.T = 'A';
lookup.C = 'G';
lookup.G = 'C';

function pairElement(str) {
 return str.split('').map(function(p) {return [p, lookup[p]];});
}
```
