---
id: 67172b43f84bcd2dec238a3d
title: Construye un Organizador de Libros
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio. 

**Historias de usuario:**

1. Deberías tener un arreglo de objetos llamado `books` donde cada objeto en el arreglo debe tener una cadena de texto `title`, otra cadena `authorName` y un número `releaseYear`.

2. Tu arreglo `books` debe tener un mínimo de tres objetos.

3. Deberías tener una función de devolución de llamada llamada `sortByYear` que acepte dos libros como parámetros para ordenar el arreglo.

4. La función `sortByYear` debe devolver `-1` si el `releaseYear` del primer libro es menor que el del segundo libro.

5. La función `sortByYear` debe devolver `1` si el `releaseYear` del primer libro es mayor que el del segundo libro.

6. La función `sortByYear` debe devolver `0` si ambos valores de `releaseYear` son iguales.

7. Deberías filtrar los libros escritos después de cierto año como 1950 del arreglo `books` y guardar el arreglo filtrado en un nuevo arreglo llamado `filteredBooks`.

8. Deberías ordenar los libros en el arreglo `filteredBooks` según su `releaseYear` en orden ascendente.  You learned in a prior lesson that the `sort()` method will sort the array in place. Esto significa que el arreglo `filteredBooks` será mutado.

# --hints--

Deberías tener una función `sortByYear` en tu código.

```js
assert.isFunction(sortByYear);
```

Tu función `sortByYear` debería aceptar dos parámetros.

```js
assert.lengthOf(sortByYear, 2);
```

Tu función `sortByYear` debería devolver `-1` si el `releaseYear` del objeto `book1` es menor que el del objeto `book2`, `1` si el año de lanzamiento del objeto `book1` es mayor que el del objeto `book2`, y `0` en todos los demás escenarios.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

Deberías tener un arreglo `books` en tu código.

```js
assert.isArray(books)
```

Tu arreglo `books` debe tener al menos tres objetos de libro.

```js
assert.isAtLeast(books.length, 3)
```

Tu arreglo `books` debe contener objetos, cada uno con una cadena de texto `title`, otra cadena `authorName` y un número `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

Deberías tener un arreglo `filteredBooks` en tu código.

```js
assert.isArray(filteredBooks)
```

El arreglo `filteredBooks` debería tener algunos de los libros filtrados del arreglo `books` y no estar vacío.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

Debes llamar a la función de orden superior `sort` pasando la función de devolución de llamada `sortByYear` en el arreglo `filteredBooks`.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
