---
id: 66bb6a9c2dd58b73cd759034
title: 构建一个成绩册应用
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个名为 `getAverage` 的函数，它接受一个测试分数数组作为参数并返回平均分数。
1. 你应该有一个名为 `getGrade` 的函数，它接受一个学生分数作为参数，并返回一个表现字母等级的字串，基于该分数。以下是分数及其对应的字母等级：分数范围 等级 `100``"A+"``90 - 99``"A"``80 - 89``"B"``70 - 79``"C"``60 - 69``"D"``0 - 59``"F"`

1. 你应该有一个名为 `hasPassingGrade` 的函数，它接受一个分数作为参数，并根据分数是否对应及格成绩返回 `true` 或 `false`。
1. `hasPassingGrade` 函数应使用 `getGrade` 函数获取字母成绩，并用它来判断成绩是否及格。及格成绩是任何不等于 `"F"` 的成绩。
1. You should have a function named `studentMsg` that takes an array of scores and a student score as the parameters. The function should return a string with the format:

    - `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`，如果学生通过了课程。
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`，如果学生未通过课程。

    Replace `average-goes-here` with the average of total scores and `grade-goes-here` with the student's grade.
    Use `getAverage` to get the average score and `getGrade` to get the student's grade.

# --hints--

你应该有一个名为 `getAverage` 的函数。

```js
assert.isFunction(getAverage);
```

你的 `getAverage` 函数应该返回一个数字。

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` 应该返回 `71.7`。

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` 应该返回 `85.4`。

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` 应该返回 `92.4`。

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` 应该返回 `47.25`。

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` 应该返回 `47.125`。

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

你的 `getAverage` 函数应该返回考试分数的平均分。

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

你应该有一个名为 `getGrade` 的函数。

```js
assert.isFunction(getGrade);
```

你的 `getGrade` 函数应在分数为 `100` 时返回 `"A+"`。

```js
assert.strictEqual(getGrade(100), "A+");
```

如果分数在 `90` 和 `99` 之间，你的 `getGrade` 函数应返回 `"A"`。

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

如果分数在 `80` 到 `89` 之间，你的 `getGrade` 函数应该返回 `"B"`。

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

如果分数在 `70` 和 `79` 之间，你的 `getGrade` 函数应返回 `"C"`。

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

如果分数在 `60` 到 `69` 之间，你的 `getGrade` 函数应该返回 `"D"`。

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

如果分数在 `0` 到 `59` 之间，你的 `getGrade` 函数应该返回 `"F"`。

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


你应该有一个名为 `hasPassingGrade` 的函数。

```js
assert.isFunction(hasPassingGrade);
```

你的 `hasPassingGrade` 函数应该返回一个布尔值。

```js
assert.isBoolean(hasPassingGrade(100));
```

如果评级是 `"A"`，你的 `hasPassingGrade` 函数应该返回 `true`。

```js
assert.isTrue(hasPassingGrade(100));
```

如果评级是 `"F"`，你的 `hasPassingGrade` 函数应该返回 `false`。

```js
assert.isFalse(hasPassingGrade(53));
```

如果评级合格，你的 `hasPassingGrade` 函数应该返回 `true`。

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

你应该有一个名为 `studentMsg` 的函数。

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` 应返回以下消息：`"Class average: 71.7. Your grade: F. You failed the course."`。

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

当调用 `studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` 时，你的函数应该返回以下信息：`"Class average: 50.8.   Your grade: A+. You passed the course."`。

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` 应返回以下消息：`"Class average: 48.25. Your grade: B. You passed the course."`。

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` 应返回以下消息：`"Class average: 45.625. Your grade: C. You passed the course."`。

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

你的 `studentMsg` 函数应该基于学生的成绩和班级平均分来返回正确的信息。

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
