---
id: 66a207974c806a19d6607073
title: Eine Buchinventar-App erstellen
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um die Aufgabe abzuschließen.

**User Stories:**

1. Sie sollten ein `h1`-Element mit dem Text `Book Inventory` haben.
1. Sie sollten ein `table`-Element mit Spalten namens `Title`, `Author`, `Category`, `Status` und `Rate` haben.
1. Ihre Tabelle sollte mindestens vier Zeilen haben, die erste für die Spaltenüberschriften und die übrigen mit Buchinformationen gefüllt.
1. Jede Tabellenzeile im Tabellenkörper sollte entweder die Klasse `read`, `to-read` oder `in-progress` haben.
1. `td`-Elemente der `Status`-Spalte sollten ein `span`-Element mit dem `class` von `status` enthalten, das den Text `Read`, `To Read` oder `In Progress` umgibt, abhängig von der Klasse dieser Zeile.
1. `td`-Elemente der `Rate`-Spalte sollten ein `span`-Element mit dem `class` von `rate` enthalten, das drei leere `span`-Elemente umschließt.
1. `.rate`-Elemente, die in `.read`-Zeilen platziert sind, sollten eine zusätzliche Klasse mit dem Wert `one`, `two` oder `three` haben, abhängig von der persönlichen Bewertung. Dieser Wert sollte nach `rate` kommen.
1. Sie sollten drei Attributselektoren erstellen, um die Elemente mit der Klasse `read`, `to-read` und `in-progress` anzusprechen und deren `background-image`-Eigenschaft so einstellen, dass sie einen `linear-gradient` Ihrer Wahl verwenden.
1. Sie sollten die `display`-Eigenschaft jedes `span`-Elements auf `inline-block` setzen.
1. Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `to-read` sind, und deren `border`- und `background-image`-Eigenschaften setzen.
1. Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `read` sind, und deren `border`- und `background-image`-Eigenschaften setzen.
1. Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `in-progress` sind, und deren `border`- und `background-image`-Eigenschaften setzen.
1. Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` und die `span`-Elemente mit dem Klassenwert, der mit `rate` beginnt, anzusprechen und deren `height`-, `width`- und `padding`-Eigenschaften setzen.
1. Sie sollten einen Attributselektor verwenden, um die `span`-Elemente anzusprechen, die direkte Kinder von `span`-Elementen mit dem `class`-Wert sind, der mit `rate` beginnt, und deren `border`-, `border-radius`-, `margin`-, `height`-, `width`- und `background-color`-Eigenschaften setzen.
1. Sie sollten einen Attributselektor verwenden, um den ersten Nachfahren von `span`-Elementen anzusprechen, die das Wort `one` als Teil ihres `class`-Werts haben, und dessen `background-image`-Eigenschaft so einstellen, dass sie einen `linear-gradient` verwendet.
1. Sie sollten einen Attributselektor verwenden, um die ersten zwei Nachfahren von `span`-Elementen anzusprechen, die das Wort `two` als Teil ihres `class`-Werts haben, und deren `background-image`-Eigenschaft so einstellen, dass sie einen `linear-gradient` verwendet.
1. Sie sollten einen Attributselektor verwenden, um die drei `span`-Elemente anzusprechen, die Nachfahren von `span`-Elementen sind, die das Wort `three` als Teil ihres `class`-Werts haben, und deren `background-image`-Eigenschaft so einstellen, dass sie einen `linear-gradient` verwendet.

# --hints--

Sie sollten ein `h1`-Element mit dem Text `Book Inventory` haben.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

Sie sollten nur ein `h1`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Sie sollten ein `table`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Sie sollten ein `thead`-Element und ein `tbody`-Element innerhalb von `table` haben.

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

Innerhalb von `thead` sollte ein `tr` mit 5 `th`-Elementen sein.

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

Ihre erste Spalte sollte den Text `Title` als Überschrift haben.

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

Ihre zweite Spalte sollte den Text `Author` als Überschrift haben.

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

Ihre dritte Spalte sollte den Text `Category` als Überschrift haben.

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

Ihre vierte Spalte sollte den Text `Status` als Überschrift haben.

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

Ihre fünfte Spalte sollte den Text `Rate` als Überschrift haben.

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

Ihre Tabelle sollte mindestens vier Zeilen haben.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Jede Zeile sollte immer 5 Spalten haben.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

Jede Tabellenzeile außer der Überschriftszeile sollte entweder die Klasse `read`, `to-read` oder `in-progress` haben.

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

`td`-Elemente der `Status`-Spalte sollten ein `span`-Element enthalten.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Jedes `span`-Element der `Status`-Spalte sollte die Klasse `status` haben.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Jedes `.status`-Element sollte den Text `Read`, `To Read` oder `In Progress` haben, abhängig von der Klasse seiner Zeile.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

`td`-Elemente der `Rate`-Spalte sollten ein `span`-Element enthalten.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Jedes `span`-Element, das ein direktes Kind eines `td`-Elements der `Rate`-Spalte ist, sollte als erste Klasse die Klasse `rate` haben.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Jedes `.rate`-Element sollte drei leere `span`-Elemente enthalten.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

`.rate`-Elemente, die in `.read`-Zeilen platziert sind, sollten eine zusätzliche Klasse nach der `rate`-Klasse mit dem Wert `one`, `two` oder `three` haben, abhängig von der persönlichen Bewertung.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

Sie sollten einen Attributselektor haben, um Zeilen anzusprechen, die die Klasse `read` haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Sie sollten einen Attributselektor verwenden, um Zeilen anzusprechen, die die Klasse `read` haben, und deren `background-image`-Eigenschaft auf einen linearen Farbverlauf Ihrer Wahl setzen.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

Sie sollten einen Attributselektor haben, um Zeilen anzusprechen, die die Klasse `to-read` haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Sie sollten einen Attributselektor verwenden, um Zeilen anzusprechen, die die Klasse `to-read` haben, und deren `background-image`-Eigenschaft auf einen linearen Farbverlauf Ihrer Wahl setzen.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

Sie sollten einen Attributselektor haben, um Zeilen anzusprechen, die die Klasse `in-progress` haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Sie sollten einen Attributselektor verwenden, um Zeilen anzusprechen, die die Klasse `in-progress` haben, und deren `background-image`-Eigenschaft auf einen linearen Farbverlauf Ihrer Wahl setzen.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Sie sollten die `display`-Eigenschaft jedes `span`-Elements auf `inline-block` setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

Sie sollten einen Attributselektor haben, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `to-read` sind.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `to-read` sind, und deren `border`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `to-read` sind, und deren `background-image`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

Sie sollten einen Attributselektor haben, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `read` sind.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `read` sind, und deren `border`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `read` sind, und deren `background-image`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

Sie sollten einen Attributselektor haben, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `in-progress` sind.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `in-progress` sind, und deren `border`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` anzusprechen, die Nachfahren von `tr`-Elementen mit der Klasse `in-progress` sind, und deren `background-image`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

Sie sollten einen Attributselektor haben, um die `span`-Elemente mit der Klasse `status` und die `span`-Elemente mit dem Klassenwert, der mit `rate` beginnt, anzusprechen.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` und die `span`-Elemente mit dem Klassenwert, der mit `rate` beginnt, anzusprechen und deren `height`-Eigenschaft setzen.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` und die `span`-Elemente mit dem Klassenwert, der mit `rate` beginnt, anzusprechen und deren `width`-Eigenschaft setzen.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente mit der Klasse `status` und die `span`-Elemente mit dem Klassenwert, der mit `rate` beginnt, anzusprechen und deren `padding`-Eigenschaft setzen.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

Sie sollten einen Attributselektor haben, um die `span`-Elemente anzusprechen, die direkte Kinder von `span`-Elementen mit dem `class`-Wert sind, der mit `rate` beginnt.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente anzusprechen, die direkte Kinder von `span`-Elementen mit dem `class`-Wert sind, der mit `rate` beginnt, und deren `border`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente anzusprechen, die direkte Kinder von `span`-Elementen mit dem `class`-Wert sind, der mit `rate` beginnt, und deren `border-radius`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente anzusprechen, die direkte Kinder von `span`-Elementen mit dem `class`-Wert sind, der mit `rate` beginnt, und deren `margin`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente anzusprechen, die direkte Kinder von `span`-Elementen mit dem `class`-Wert sind, der mit `rate` beginnt, und deren `height`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente anzusprechen, die direkte Kinder von `span`-Elementen mit dem `class`-Wert sind, der mit `rate` beginnt, und deren `width`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente anzusprechen, die direkte Kinder von `span`-Elementen mit dem `class`-Wert sind, der mit `rate` beginnt, und deren `background-color`-Eigenschaft setzen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

Sie sollten einen Attributselektor haben, um den ersten Nachfahren von `span`-Elementen anzusprechen, die das Wort `one` als Teil ihres `class`-Werts haben.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Sie sollten einen Attributselektor verwenden, um den ersten Nachfahren von `span`-Elementen anzusprechen, die das Wort `one` als Teil ihres `class`-Werts haben, und dessen `background-image`-Eigenschaft so einstellen, dass sie einen `linear-gradient` verwendet.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

Sie sollten einen Attributselektor haben, um die ersten zwei Nachfahren von `span`-Elementen anzusprechen, die das Wort `two` als Teil ihres `class`-Werts haben.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Sie sollten einen Attributselektor verwenden, um die ersten zwei Nachfahren von `span`-Elementen anzusprechen, die das Wort `two` als Teil ihres `class`-Werts haben, und deren `background-image`-Eigenschaft so einstellen, dass sie einen `linear-gradient` verwendet.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

Sie sollten einen Attributselektor haben, um die `span`-Elemente anzusprechen, die Nachfahren von `span`-Elementen sind, die das Wort `three` als Teil ihres `class`-Werts haben.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Sie sollten einen Attributselektor verwenden, um die `span`-Elemente anzusprechen, die Nachfahren von `span`-Elementen sind, die das Wort `three` als Teil ihres `class`-Werts haben, und deren `background-image`-Eigenschaft so einstellen, dass sie einen `linear-gradient` verwendet.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
