---
id: 6a89c69cbe1bb69845992d8b
title: Costruisci un generatore di intervalli numerici
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

In questo laboratorio costruirai una funzione `range_of_numbers` che usa la ricorsione per generare una lista di numeri all'interno di un intervallo specificato.

Per esempio:

- `range_of_numbers(3, 9)` dovrebbe restituire `[3, 4, 5, 6, 7, 8, 9]`.
- `range_of_numbers(5, 5)` dovrebbe restituire `[5]`.

Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione chiamata `range_of_numbers` con due parametri chiamati `start_num` e `end_num`.
1. La funzione dovrebbe restituire una lista di numeri interi consecutivi che inizia con `start_num` e termina con `end_num`, includendo entrambi i valori.
1. Il valore di `start_num` sarà sempre minore o uguale a `end_num`.
1. La tua funzione dovrebbe usare la ricorsione chiamando se stessa.
1. La tua funzione non dovrebbe usare loop `for` o `while`, comprensioni o la funzione `range()`.
1. Quando `start_num` è uguale a `end_num`, la funzione dovrebbe restituire una lista contenente `start_num`. Questo è il caso base.
1. Per il caso ricorsivo, la funzione dovrebbe chiamare se stessa con un argomento che si avvicina al caso base, poi aggiungere il numero corrente alla lista restituita.
1. La funzione non dovrebbe usare variabili globali per archiviare la lista restituita.

# --hints--

Dovresti avere una funzione chiamata `range_of_numbers` con parametri chiamati `start_num` e `end_num`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` dovrebbe restituire `[1, 2, 3, 4, 5]`.

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` dovrebbe restituire `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` dovrebbe restituire `[4]`.

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` dovrebbe restituire `[10, 11, 12, 13, 14, 15]`.

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

La tua funzione `range_of_numbers` dovrebbe chiamare se stessa.

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

La tua funzione `range_of_numbers` non dovrebbe usare loop, comprensioni o la funzione `range()`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

Chiamare `range_of_numbers(1, 3)` non dovrebbe influenzare una chiamata successiva. Dopo quella chiamata, `range_of_numbers(6, 9)` dovrebbe ancora restituire `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
