---
id: 69f0a37c3bf194a7c9c62492
title: بناء سجل إعارة أجهزة
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

أنت مطوِّر برمجيات مكلف بإنشاء تطبيق سجل لإدارة أجهزة تكنولوجيا المعلومات التي يوفرها مكتب خدمة شركتك. السجل هو كائن مفاتيحه هي أوسمة الأصول وقيمه هي كائنات:

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تمثل السجل ككائن حيث يُفهرس كل جهاز بواسطة وسم الأصل الخاص به ويحتوي على خصائص `type` و`status` و`borrower` (مع `name` و`email`) و`dueDate`.

1. يجب تنفيذ دالة باسم `checkoutDevice(ledger, assetTag, borrower)` تستقبل كائن ledger وتقوم باستنساخه باستخدام `serializeLedger` و`loadLedger` دون تعديل الأصل. في النسخة المستنسخة، يجب تعيين خاصيتي `name` و`email` في كائن `borrower`، وتحديث `status` إلى `"CheckedOut"`. ثم تُرجع كائنًا بالشكل `{ ledger: updatedLedger, message: confirmationString }`، حيث يحتوي `message` على وسم الأصل وكائن `borrower` وخاصية `name`.

   - إذا لم يكن `assetTag` موجودًا في السجل، يجب أن تُرجع الدالة السجل دون تغيير مع `message` تتضمن وسم الأصل وتوضح أنه لم يُعثر عليه.
   - إذا كان الجهاز بالفعل في حالة `"CheckedOut"`، يجب أن تُرجع الدالة السجل دون تغيير وتترك المستعير الحالي كما هو، مع `message` تتضمن وسم الأصل وتوضح أن الجهاز مُستعار بالفعل.

1. يجب تنفيذ دالة `checkinDevice(ledger, assetTag)` التي تستقبل كائن ledger وتنسخه باستخدام `serializeLedger` و`loadLedger`، دون تعديل الكائن الأصلي. على النسخة، يجب مسح خاصيتي `name` و`email` في كائن `borrower`، إعادة تعيين `dueDate` إلى `""`، تحديث `status` إلى `"CheckedIn"`، وإرجاع كائن بالشكل `{ ledger: updatedLedger, message: confirmationString }`، حيث تتضمن الخاصية `message` وسم الأصل.

   - إذا لم يكن `assetTag` موجودًا في السجل، يجب أن تُرجع الدالة السجل دون تغيير مع `message` تتضمن وسم الأصل وتشرح أنه لم يتم العثور عليه.

1. يجب أن تنفذ دالة `listOverdueDevices(ledger, today)` التي تأخذ كائن السجل وترجع مصفوفة من كائنات الأجهزة المتأخرة، مرتبة حسب `dueDate` تصاعديًا. الجهاز متأخر إذا كان `status` هو `"CheckedOut"` و`dueDate` قبل `today` بدقة. كل من `today` و`dueDate` هما سلاسل نصية بترتيب شهر/يوم/سنة، وقد يكون الشهر أو اليوم مضافًا له صفر أو لا (مثلاً، `"9/5/2025"` و`"09/05/2025"` كلاهما صحيح).

1. يجب ألا تستخدم كائن `Date` في جافاسكريبت في أي مكان في كودك.

1. يجب أن تنفذ دالة `serializeLedger(ledger)` التي تحول كائن السجل إلى سلسلة JSON، ودالة `loadLedger(json)` التي تحول سلسلة JSON إلى كائن جافاسكريبت.

# --hints--

يجب أن تنشئ دالة باسم `checkoutDevice`.

```js
assert.isFunction(checkoutDevice);
```

يجب أن تنشئ دالة باسم `checkinDevice`.

```js
assert.isFunction(checkinDevice);
```

يجب أن تنشئ دالة باسم `listOverdueDevices`.

```js
assert.isFunction(listOverdueDevices);
```

يجب أن تنشئ دالة باسم `serializeLedger`.

```js
assert.isFunction(serializeLedger);
```

يجب أن تنشئ دالة باسم `loadLedger`.

```js
assert.isFunction(loadLedger);
```

يجب ألا تعدل دالة `checkoutDevice` السجل الأصلي.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

يجب أن تستخدم دالة `checkoutDevice` دالتي `serializeLedger` و`loadLedger` لنسخ ledger.

```js
const source = __helpers.removeJSComments(checkoutDevice.toString());
assert.match(source, /\bserializeLedger\s*\(/);
assert.match(source, /\bloadLedger\s*\(/);
```

يجب أن تضبط دالة `checkoutDevice` `name` و`email` للمستعير، وتضبط `status` إلى `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

يجب أن ترجع دالة `checkoutDevice` رسالة تأكيد تتضمن وسم الأصل و`name` للمستعير.

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

يجب أن تحافظ دالة `checkoutDevice` على الأجهزة غير المتعلقة والحقول غير المعدلة في السجل المعاد.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

يجب أن ترجع دالة `checkoutDevice` السجل دون تغيير ورسالة تتضمن وسم الأصل عندما لا يُعثر على وسم الأصل.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

يجب أن ترجع دالة `checkoutDevice` السجل دون تغيير، وتترك المستعير الحالي كما هو، وترجع رسالة تتضمن وسم الأصل عندما يكون الجهاز معارًا بالفعل.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

يجب ألا تعدل دالة `checkinDevice` السجل الأصلي.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

يجب أن تستخدم دالة `checkinDevice` دالتي `serializeLedger` و`loadLedger` لنسخ ledger.

```js
const source = __helpers.removeJSComments(checkinDevice.toString());
assert.match(source, /\bserializeLedger\s*\(/);
assert.match(source, /\bloadLedger\s*\(/);
```

يجب أن تمسح دالة `checkinDevice` `name` و`email` للمستعير، تعيد تعيين `dueDate` إلى `""`، وتضبط `status` إلى `"CheckedIn"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

يجب أن ترجع دالة `checkinDevice` رسالة تأكيد تتضمن وسم الأصل.

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

يجب أن تحافظ دالة `checkinDevice` على الأجهزة غير المتعلقة والحقول غير المعدلة في السجل المعاد.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

يجب أن ترجع دالة `checkinDevice` السجل دون تغيير ورسالة تتضمن وسم الأصل عندما لا يُعثر على وسم الأصل.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

يجب أن ترجع دالة `listOverdueDevices` مصفوفة.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

يجب أن تُرجع دالة `listOverdueDevices` الأجهزة المتأخرة فقط التي تكون قيمة `status` فيها هي `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

يجب أن ترجع دالة `listOverdueDevices` الأجهزة المتأخرة مرتبة حسب `dueDate` تصاعديًا.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

يجب أن ترتب دالة `listOverdueDevices` بشكل صحيح حتى عندما لا يكون الشهر أو اليوم مضافًا له صفر.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

يجب ألا تستخدم كائن `Date` في جافاسكريبت في أي مكان في كودك.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

يجب أن ترجع دالة `serializeLedger` سلسلة JSON تمثل السجل الممرر إليها.

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

يجب أن ترجع دالة `loadLedger` كائن جافاسكريبت تمثله سلسلة JSON.

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = loadLedger(serializeLedger(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = loadLedger(serializeLedger(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
