---
id: 612e89d254fe5d3df7d6693d
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

`html` 박스 모델을 재설정했으니, 내부 요소에도 이 설정을 전달해야 합니다. 이를 위해 `box-sizing` 속성을 `inherit`로 설정하면, 대상 요소들이 부모 요소와 같은 값을 사용하도록 지시할 수 있습니다.

또한 선택자 뒤에 오는 특별한 키워드인 가상 요소도 선택해야 합니다. 사용할 두 가상 요소는 `::before`와 `::after` 가상 요소입니다.

`::before` 선택자는 선택된 요소의 첫 번째 자식인 가상 요소를 생성하고, `::after` 선택자는 선택된 요소의 마지막 자식인 가상 요소를 생성합니다. 이러한 가상 요소는 나중에 이 프로젝트에서 보게 될 장식용 콘텐츠를 만드는 데 자주 사용됩니다.

지금은 `*`를 가진 모든 요소를 선택하는 CSS 선택자를 만들고, `::before`와 `::after` 가상 요소도 포함하세요. `box-sizing` 속성을 `inherit`로 설정하세요.

# --hints--

`*, ::before, ::after` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

`*, ::before, ::after` 선택자에 `box-sizing` 속성이 `inherit`로 설정되어 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
