---
id: 612e804c54d5e7308d7ebe56
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`.keys` 요소 안에 `div` 요소 일곱 개를 추가하세요. 모두 `key` 클래스를 부여하세요.

# --hints--

새로운 `div` 요소 일곱 개를 만들어야 합니다.

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

새로운 `div` 요소 일곱 개는 `.keys` 요소 안에 있어야 합니다.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

새로운 `div` 요소 일곱 개 모두 `class`가 `key`로 설정되어야 합니다.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
