---
id: 612e77aba7ca691f598feb02
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

모바일 기기에 최적화하기 위한 `meta` 태그 하나와 페이지에서 허용되는 `charset`를 지정하는 태그 하나, 이렇게 두 개의 `meta` 태그를 추가하세요.

# --hints--

페이지에 두 개의 `meta` 요소를 추가해야 합니다.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

두 개의 `meta` 요소가 있어야 합니다.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

한 `meta` 요소는 `name` 속성이 `viewport`로 설정되어 있고, `content`가 `width=device-width, initial-scale=1.0`로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

다른 `meta` 요소는 `charset` 속성이 `UTF-8`로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
