---
id: 612e8279827a28352ce83a72
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이제 일곱 개의 `.key` 요소 집합을 복사해서 `.keys` div 안에 두 세트 더 붙여넣으세요.

# --hints--

`.key` 요소가 21개 있어야 합니다.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

`.black--key` 요소가 15개 있어야 합니다.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

모든 `.key` 요소는 `.keys` 요소 안에 있어야 합니다.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

모든 `.black--key` 요소는 `.keys` 요소 안에 있어야 합니다.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
