---
id: 612e7d1c29fb872d6384379c
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

기존 `div` 안에 두 번째 `div`를 중첩하고, `class`를 `keys`로 설정하세요.

# --hints--

두 번째 `div` 요소를 만들어야 합니다.

```js
const divDiv = document.querySelectorAll('div');
assert(divDiv?.length === 2);
```

새로운 `div` 요소는 기존 `div` 요소 안에 있어야 합니다.

```js
const div = document.querySelector('div');
assert(div?.children?.length === 1);
assert(div?.children?.[0]?.localName === 'div');
```

새로운 `div` 요소의 `class`는 `keys`로 설정해야 합니다.

```js
const div = document.querySelector('div');
assert(div?.children?.[0]?.className === 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
