---
id: 612e83ec2eca1e370f830511
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`link` 요소 안에 `head` 요소를 추가하세요. `rel` 속성을 `stylesheet` 값으로, `href` 속성을 `styles.css` 값으로 설정하세요.

# --hints--

코드에 `link` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('link'));
```

기존 `head` 태그를 변경하지 마세요. 닫는 태그를 삭제하지 않았는지 확인하세요.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

`link` 요소가 하나 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 요소는 `rel` 속성을 가지고 그 값은 `stylesheet` 이어야 합니다.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 요소는 `href` 속성을 가지고 그 값은 `styles.css` 이어야 합니다.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
