---
id: 612e78af05201622d4bab8aa
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

피아노 작업을 시작할 시간입니다. `div` 요소 안에 `body` 요소를 만들고 `id`를 `piano`으로 설정하세요.

# --hints--

새로운 `div` 요소를 만들어야 합니다.

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div`는 `body` 요소 안에 있어야 합니다.

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

`div`는 `id`를 `piano`으로 설정해야 합니다.

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
