---
id: 612e9d142affc44a453655db
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

키 자체에 스타일을 적용할 시간입니다. `class` 요소에 대해 `.key` 선택자를 만드세요. `background-color` 집합을 `#ffffff` 값으로, `position` 속성을 `relative`로, `width` 속성을 `41px`로, `height` 속성을 `175px`로 설정하세요.

# --hints--

`.key` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key'));
```

`.key` 선택자에 `background-color` 속성이 `#ffffff`으로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor === 'rgb(255, 255, 255)');
```

`.key` 선택자에 `position` 속성이 `relative`으로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.position === 'relative');
```

`.key` 선택자에 `width` 속성이 `41px`으로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.width === '41px');
```

`.key` 선택자에 `height` 속성이 `175px`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.height === '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
