---
id: 612e6afc009b450a437940a1
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

기본 HTML 구조부터 시작하세요. `DOCTYPE` 선언과 `html`, `head`, `body`, `title` 요소를 추가하세요.

이 페이지의 언어를 영어로 설정하세요. `title`를 `Piano`로 설정하세요.

# --hints--

코드에 `DOCTYPE` 참조가 포함되어야 합니다.

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 참조 뒤에 공백을 포함해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

문서 유형을 `html`로 정의해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

`DOCTYPE` 선언은 타입 뒤에 `>`로 닫아야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 요소는 `lang` 속성이 `en`인 여는 태그가 있어야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 선언은 HTML의 맨 앞에 있어야 합니다.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

여는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<head\s*>/i));
```

닫는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<\/head\s*>/i));
```

여는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<body\s*>/i));
```

닫는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<\/body\s*>/i));
```

`head`과 `body` 요소는 형제 관계여야 합니다.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

코드에 `title` 요소가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.exists(title);
```

프로젝트의 제목은 `Piano`이어야 합니다.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

제목의 대소문자와 철자가 정확해야 한다는 점을 기억하세요.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
