---
id: 612e813b3ba67633222cbe54
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

`class` 속성은 여러 값을 가질 수 있다는 점을 기억하세요. 흰 건반과 검은 건반을 구분하려면 두 번째 `class` 값으로 `black--key`를 추가하세요. 이 값을 두 번째, 세 번째, 다섯 번째, 여섯 번째, 일곱 번째 `.key` 요소에 추가하세요.

# --hints--

두 번째 `.key` 요소에도 `class`가 `black--key`이어야 합니다.

```js
const key = document.querySelectorAll('.key')?.[1];
assert(key?.className?.includes('black--key'));
```

세 번째 `.key`에는 `black--key`가 `class`에 있어야 합니다.

```js
const third = document.querySelectorAll('.key')?.[2];
assert(third?.classList?.contains('black--key'));
```

다섯 번째 `.key`에는 `black--key`가 `class`에 있어야 합니다.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert(fifth?.classList?.contains('black--key'));
```

여섯 번째 `.key`에는 `black--key`가 `class`에 있어야 합니다.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert(sixth?.classList?.contains('black--key'));
```

일곱 번째 `.key`에는 `black--key`가 `class`에 있어야 합니다.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert(seventh?.classList?.contains('black--key'));
```

`.black--key` 요소가 다섯 개 있어야 합니다.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 5);
```

`.key` 요소가 일곱 개 있어야 합니다.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
