---
id: 614090d5a22b6f0a5a6b464c
title: 13 단계
challengeType: 0
dashedName: step-13
---

<!-- TODO: I purposefully added the `nav` styles without Camper input -->

# --description--

The child combinator selector `>` is used between selectors to target only elements that match the second selector and are a direct child of the first selector.

이것은 요소가 깊이 중첩된 경우 스타일 범위를 제어하려 할 때 유용할 수 있습니다.

`>` 선택자를 사용하여 `nav` 요소 내의 순서 없는 목록 요소를 타겟으로 하고, _Flexbox_를 사용해 자식 요소들을 균등하게 배치하세요.

# --hints--

`nav > ul` 선택자를 사용해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('nav > ul'));
```

`nav > ul` 요소에 `display`: `flex`를 적용해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('nav > ul')?.display, 'flex');
```

`nav > ul` 요소에 `justify-content`: `space-evenly`를 적용해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('nav > ul')?.justifyContent, 'space-evenly');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

--fcc-editable-region--

--fcc-editable-region--
```
