---
id: 68da6138d70960c197ec5eed
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

接下来，你需要验证点对的数量至少为 1，因为你无法生成零个或负数点对的括号组合。

为你的代码添加另一个 `if` 语句来查看 `pairs` 是否小于 `1`。如果此条件为真，返回字串 `The number of pairs should be at least 1`。

# --hints--

你应该在你的 `gen_parentheses` 函数中有第二个 `if` 语句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

你的第二个 `if` 语句应查看 `pairs` 是否小于 `1`。

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

当 `pairs` 小于 `1` 时，你应该返回字串 `The number of pairs should be at least 1`。

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    
    return []
--fcc-editable-region--
```
