---
id: 68da6138d70960c197ec5eec
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

在实现核心算法之前，你需要验证输入。`pairs` 参数应为整数，因为它表现为要生成的括号点对数量。

为你的函数开头添加一个 `if` 语句来查看 `pairs` 是否不是整数。为此使用 `isinstance()` 函数。

如果条件为真，返回字串 `The number of pairs should be an integer`。

# --hints--

你应该在你的 `gen_parentheses` 函数中有一个 `if` 语句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

你的 `if` 语句应使用 `not isinstance(pairs, int)` 来查看 `pairs` 是否不是整数。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

当 `input` 不是整数时，你应该 `return` 字串 `The number of pairs should be an integer`。

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    
    return []
--fcc-editable-region--
```
