---
id: 68da6138d70960c197ec5ef1
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

现在你需要识别何时构建了一个完整的括号组合。完整的组合长度等于点对数量的两倍（因为每个点对包含一个开括号和一个闭括号）。

在你的 `while` 循环内，添加一个 `if` 语句来查看 `len(current) == 2 * pairs` 是否为真。当此条件为真时，将 `current` 追加到 `result` 列表。

# --hints--

你应该在你的 `while` 循环内有一个 `if` 语句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

你的 `if` 语句应查看 `len(current) == 2 * pairs` 是否成立。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

你应该在你的新建 `if` 语句内将 `current` 追加到 `result` 列表中。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
--fcc-editable-region--    
    return result

print(gen_parentheses(1))
```
