---
id: 68da6138d70960c197ec5ef2
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

如果当前字串尚未完成，你需要探索下一个可能的状态。为你的 `if` 语句添加一个 `else` 子句。

在 `else` 块中，你将处理添加左括号。添加一个 `if` 语句来查看 `opens_used < pairs`。这确保你不会使用超过允许数量的左括号。

如果此条件为真，则向 `queue` 追加一个新元组：`(current + '(', opens_used + 1, closes_used)`。这表现为添加一个左括号后的状态。

# --hints--

你应该在 `while` 循环内的 `if` 语句中有一个 `else` 子句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

你应该在 `else` 块内有一个 `if` 语句，检查 `opens_used < pairs`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

当条件为真时，你应该将 `(current + '(', opens_used + 1, closes_used)` 追加到队列。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
