---
id: 68da6138d70960c197ec5eef
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

对于广度优先查找方法，你将使用一个队列来跟踪在构建括号组合时的不同状态。每个状态将表现为包含三个元素的元组：

- 当前正在构建的字串
- 到目前为止使用的左括号数量
- 到目前为止使用的闭合括号数量

创建一个名为 `queue` 的变量，并用包含一个元组的列表初始化它：`('', 0, 0)`。这表现了起始状态，字串为空且使用了零个括号。

# --hints--

你应该在你的 `gen_parentheses` 函数中声明一个名为 `queue` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

你应该用包含元组 `('', 0, 0)` 的列表初始化 `queue`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    result = []
    return result
--fcc-editable-region--
```
