---
id: 68da6138d70960c197ec5ef0
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

现在你将实现主 BFS 循环。创建一个 `while` 循环，只要 `queue` 不为空（当 `queue` 计算为 `True` 时）就继续。

在循环内，打印 `queue`。

# --hints--

你的 `gen_parentheses` 函数中应该有一个 `while` 循环。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

你的 `while` 循环应有条件 `queue`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

你应该在你的 `while` 循环内打印 `queue`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
