---
id: 68da6138d70960c197ec5eee
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

现在你将设置数据结构来保存你的结果。创建一个名为 `result` 的变量，并将其初始化为空列表。该列表将保存你生成的所有有效括号组合。

更新你的返回语句以返回 `result`，而不是空列表。

# --hints--

你应该在你的 `gen_parentheses` 函数中声明一个名为 `result` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

你应该将 `result` 初始化为空列表。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("result").is_equivalent("result = []")
`) })
```

你应该在你的函数末尾返回 `result`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    return []
--fcc-editable-region--
```
