---
id: 68da604b8b6012a33e2344fb
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次研讨会中，你将实现一个函数，使用广度优先查找（BFS）方法生成所有有效的括号组合。例如，两个点对括号的有效组合是 `(())` 和 `()()`。

首先创建一个名为 `gen_parentheses` 的函数，带有一个参数 `pairs`。目前，从函数中返回一个空列表。

# --hints--

你应该定义一个名为 `gen_parentheses` 的函数。

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

你的 `gen_parentheses` 函数应有一个名为 `pairs` 的单个参数。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

你的函数应返回一个空列表。

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
