---
id: 68da6138d70960c197ec5ef3
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

现在你需要处理器添加闭合括号。关键规则是你只能添加闭合括号，如果它维护平衡，意味着使用的开括号必须多于闭括号。

仍在 `else` 块内，添加第二个 `if` 语句以查看 `closes_used < opens_used`。

如果此条件为真，则向 `queue` 追加另一个新元组：`(current + ')'，opens_used，closes_used + 1)`。这表现了添加一个闭括号后的状态。

# --hints--

你应该在你的 `else` 块内有第二个 `if` 语句。

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

你的第二个嵌套 `if` 语句应查看 `closes_used < opens_used`。

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

当条件为真时，你应该将 `(current + ')', opens_used, closes_used + 1)` 追加到队列。

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
