---
id: 68e00c5349cbddfef6dbc053
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

在你的 `while` 循环内，使用 `queue.pop(0)` 来移除并获取队列中的第一个元素。这实现了 BFS 特有的先进先出（FIFO）行为。

将这个元组解包到三个变量中：`current`、`opens_used` 和 `closes_used`。

# --hints--

你应该使用 `queue.pop(0)` 来获取队列中的第一个元素。

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

你应该将结果解包到三个变量：`current`、`opens_used` 和 `closes_used`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        
--fcc-editable-region--
    return result
```
