---
id: 6690e10ebe2181212abc9652
title: Crie um cartão de visita
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve vincular a folha de estilo `styles.css` utilizando a tag `link` na seção `head` do arquivo HTML.
1. O fundo da página deve ter a cor `rosybrown`. A fonte geral deve ser `Arial` com uma alternativa de `sans-serif`.
1. All of the `p` elements should have a top and bottom `margin` of `5px`.
1. Por padrão, todos os links na página não devem ser sublinhados. Quando você passar o mouse sobre os links, eles devem ser sublinhados.
1. You should have a `div` with the `class` `business-card`.
1. The `business-card` `class` should be `300px` wide and it should have any background color of your choice. There should be a `padding` of `20px` all around and a top `margin` of `100px`. O texto deve ser centralizado e o tamanho da fonte deve ser `16px`. As margens esquerda e direita devem ser definidas como `auto`.
1. Inside the `business-card` `div`, there should be an `img` element with the `class` `profile-image`. Você pode definir a fonte da imagem como `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`, se quiser. Deve haver um `alt` com uma descrição significativa.
1. The `profile-image` selector should have a `max-width` of `100%`.
1. Inside the `business-card` `div`, after the `img` element, you should have three `p` elements with the `class` `full-name`, `designation`, and `company`, respectively.
1. O primeiro elemento `p` deve conter seu nome.
1. O segundo elemento `p` deve conter sua designação.
1. O terceiro elemento `p` deve ter o nome da sua empresa.
1. There should be an `hr` element below the `p` element containing your company name.
1. Após o elemento `hr`, devem haver dois elementos `p`. No primeiro elemento `p`, você deve ter um endereço de e-mail como texto. No segundo elemento `p`, o texto deve ser um número de telefone. Após os dois elementos `p`, deve haver um elemento `a` com o texto `Portfolio` apontando para qualquer link válido. Adicione um elemento `hr` após o elemento `a` contendo o link do portfólio.
1. You should have another `div` with the `class` `social-media`. Dentro desse elemento, deve haver um elemento `h2` com o texto `Connect with me`.
1. Inside the `social-media` `div`, there should be three `a` elements. Os elementos `a` devem ter o texto `Twitter`, `LinkedIn` e `GitHub` respectivamente, com links para sites válidos.

# --hints--

O elemento `link` deve estar dentro do elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Você deve definir a cor de fundo da página no elemento `body` como `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body').getPropVal('background-color'), 'rosybrown');
```

Você deve definir a fonte da página como `Arial` no elemento `body` com uma alternativa de `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body').getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```


Your `p` elements should have a top and bottom `margin` of `5px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('p').getPropVal('margin');
const marginTop = new __helpers.CSSHelp(document).getStyle('p').getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p').getPropVal('margin-bottom');

const condition1 = margin === '5px 0';
const condition2 = marginTop === '5px' && marginBottom === '5px';
const condition3 = margin === '5px 0 5px 0';
assert.isTrue(condition1 || condition2 || condition3);
```

Os links da página não devem ser sublinhados por padrão.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a').getPropVal('text-decoration'), 'none');
```

Os links da página devem ser sublinhados quando passar o mouse sobre eles.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const a_hover = gs('a:hover')?.textDecoration;
assert.equal(a_hover, 'underline');
```

You should have one main `div` with the `class` `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

Your `.business-card` `class` should have a width of `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('width'), '300px');
```

Your `.business-card` `class` should have a background color of your choice.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Your `.business-card` `class` should have a `padding` of `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('padding'), '20px');
```

Your `.business-card` `class` should have a `margin` of `100px` from the top.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('margin-top'), '100px');
```

The text inside the `.business-card` `class` should be center-aligned.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('text-align'), 'center');
```

The font size of the text inside the `.business-card` `class` should be `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('font-size'), '16px');
```

The left and right margins of the `.business-card` `class` should be set to `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('margin-left'), 'auto');
```

Inside the `.business-card` element, there should be an `img` element with the `class` `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

A fonte da imagem deve ser definida para uma imagem válida.

```js
assert.isAbove(document.querySelector('.profile-image').getAttribute('src').length, 0);
```

O atributo `alt` da imagem deve ser definido para uma descrição significativa.

```js
assert.isAbove(document.querySelectorAll('img')[0].alt.length, 0);
```

Your `.profile-image` selector should have a `max-width` of `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image').getPropVal('max-width'), '100%');
```

Inside the `.business-card` element, after the `img` element, there should be a `p` element with the `class` `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

After the `full-name` `p` element, there should be a `p` element with the `class` `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

After the `designation` `p` element, there should be a `p` element with the `class` `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

The first `p` element with the `class` `full-name` should contain your name.

```js
assert.isAbove(document.querySelector(".full-name")?.innerText.length, 0);
```

The second `p` element with the `class` `designation` should contain your designation.

```js
assert.isAbove(document.querySelector(".designation")?.innerText.length, 0);
```

The third `p` element with the `class` `company` should contain your company name.

```js
assert.isAbove(document.querySelector(".company")?.innerText.length, 0);
```

After the third `p` element with the `class` `company`, there should be an `hr` element.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Após o primeiro elemento `hr`, deve haver um elemento `p` com um endereço de e-mail como seu texto.

```js
assert.isAbove(document.querySelectorAll('.business-card hr + p')[0].innerText.length, 0);
```

Após o elemento `p` do e-mail, deve haver outro elemento `p` com um número de telefone como seu texto.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Após o elemento `p` do número de telefone, deve haver um elemento `a` com o texto `Portfolio`

```js
assert.equal(document.querySelector('.business-card a').textContent, 'Portfolio');
```

O elemento `a` do portfólio deve ter um `href` válido.

```js
assert.isAbove(document.querySelectorAll('.business-card a')[0].href.length, 0);
```

Após o elemento `a` contendo seu portfólio, deve haver outro elemento `hr` como um divisor.

```js
assert.exists(document.querySelector('a + hr'));
```

After the second `hr` element, there should be a `div` element with the `class` `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

Dentro do elemento `.social-media`, deve haver um elemento `h2` com o texto `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2').textContent, 'Connect with me');
```

Dentro do elemento `.social-media`, deve haver um elemento `a` com o texto `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent, 'Twitter');
```

O Twitter `href` deve apontar para um link válido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

Dentro do elemento `.social-media`, deve haver um elemento `a` com o texto `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent, 'LinkedIn');
```

O LinkedIn `href` deve apontar para um link válido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

Dentro do elemento `.social-media`, deve haver um elemento `a` com o texto `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent, 'GitHub');
```

O `href` do GitHub deve apontar para um link válido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

a:hover {
    text-decoration: underline;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
