---
id: 6690e10ebe2181212abc9652
title: Crie um cartão de visita
challengeType: 14
dashedName: design-a-business-card
demoType: onClick
---

# --description--

Fulfill the user stories below and get all the tests to pass to complete the lab.

**Histórias de usuário:**

1. Você deve vincular a folha de estilo `styles.css` utilizando a tag `link` na seção `head` do arquivo HTML.
1. O fundo da página deve ter a cor `rosybrown`. A fonte geral deve ser `Arial` com uma alternativa de `sans-serif`.
1. Todos os elementos `p` devem ter uma margem superior e inferior de `5px`.
1. Por padrão, todos os links na página não devem ser sublinhados. Quando você passar o mouse sobre os links, eles devem ser sublinhados.
1. Você deve ter uma `div` com a classe de `business-card`.
1. A classe `business-card` deve ter `300px` de largura e deve ter qualquer cor de fundo de sua escolha. Deve haver um padding de `20px` por todos os lados e uma margem superior de `100px`. O texto deve ser centralizado e o tamanho da fonte deve ser `16px`. As margens esquerda e direita devem ser definidas como `auto`.
1. Dentro da div `business-card`, deve haver um elemento `img` com a classe `profile-image`. Você pode definir a fonte da imagem como `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`, se quiser. Deve haver um `alt` com uma descrição significativa.
1. Dentro da `div` `business-card`, após o elemento `img`, você deve ter três elementos `p` com as classes `full-name`, `designation` e `company`, respectivamente.
1. O primeiro elemento `p` deve conter seu nome.
1. O segundo elemento `p` deve conter sua designação.
1. O terceiro elemento `p` deve ter o nome da sua empresa.
1. Deve haver um elemento `hr` abaixo do elemento `p` contendo o nome da sua empresa. O elemento `hr`, em geral, deve ter uma borda superior de `1px` tracejada na cor `slategray`.
1. Após o elemento `hr`, devem haver dois elementos `p`. No primeiro elemento `p`, você deve ter um endereço de e-mail como texto. No segundo elemento `p`, o texto deve ser um número de telefone. Após os dois elementos `p`, deve haver um elemento `a` com o texto `Portfolio` apontando para qualquer link válido. Adicione um elemento `hr` após o elemento `a` contendo o link do portfólio.
1. Você deve ter outra `div` com a classe `social-media`. Dentro desse elemento, deve haver um elemento `h2` com o texto `Connect with me`.
1. Dentro da div `social-media`, deve haver três elementos `a`. Os elementos `a` devem ter o texto `Twitter`, `LinkedIn` e `GitHub` respectivamente, com links para sites válidos.

# --hints--

O elemento `link` deve estar dentro do elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Você deve definir a cor de fundo da página no elemento `body` como `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body').getPropVal('background-color'), 'rosybrown');
```

Você deve definir a fonte da página como `Arial` no elemento `body` com uma alternativa de `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body').getPropVal('font-family');
console.log(el);
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert(condition1 || condition2);
```


Os elementos `p` devem ter uma margem superior e inferior de `5px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('p').getPropVal('margin');
const marginTop = new __helpers.CSSHelp(document).getStyle('p').getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p').getPropVal('margin-bottom');

const condition1 = margin === '5px 0';
const condition2 = marginTop === '5px' && marginBottom === '5px';
const condition3 = margin === '5px 0 5px 0';
assert(condition1 || condition2 || condition3);
```

Os links da página não devem ser sublinhados por padrão.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a').getPropVal('text-decoration'), 'none');
```

Os links da página devem ser sublinhados quando passar o mouse sobre eles.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const a_hover = gs('a:hover')?.textDecoration;
assert.equal(a_hover, 'underline');
```

Você deve ter uma `div` principal com a classe `business-card`.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'business-card');
```

A classe `.business-card` deve ter uma largura de `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('width'), '300px');
```

A classe `.business-card` deve ter uma cor de fundo de sua escolha.

```js
assert(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

A classe `.business-card` deve ter um padding de `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('padding'), '20px');
```

A classe `.business-card` deve ter uma margem superior de `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('margin-top'), '100px');
```

O texto dentro da classe `.business-card` deve ser centralizado.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('text-align'), 'center');
```

O tamanho da fonte do texto dentro da classe `.business-card` deve ser `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('font-size'), '16px');
```

As margens esquerda e direita da classe `.business-card` devem ser definidas como `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('margin-left'), 'auto');
```

Dentro do elemento `.business-card`, deve haver um elemento `img` com a classe `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

A fonte da imagem deve ser definida para uma imagem válida.

```js
assert.isAbove(document.querySelector('.profile-image').getAttribute('src').length, 0);
```

O atributo `alt` da imagem deve ser definido para uma descrição significativa.

```js
assert.isAbove($('img')[0].alt.length, 0);
```

Dentro do elemento `.business-card`, após o elemento `img`, deve haver um elemento `p` com a classe `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Após o elemento `p` `full-name`, deve haver um elemento `p` com a classe `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Após o elemento `p` `designation`, deve haver um elemento `p` com a classe `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

O primeiro elemento `p` com a classe `full-name` deve conter seu nome.

```js
assert.isAbove(document.querySelector(".full-name")?.innerText.length, 0);
```

O segundo elemento `p` com a classe `designation` deve conter sua designação.

```js
assert.isAbove(document.querySelector(".designation")?.innerText.length, 0);
```

O terceiro elemento `p` com a classe `company` deve conter o nome da sua empresa.

```js
assert.isAbove(document.querySelector(".company")?.innerText.length, 0);
```

Após o terceiro elemento `p` com a classe `company`, deve haver um elemento `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Os elementos `hr` devem ter uma borda superior de `1px` tracejada na cor `slategray`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('hr').getPropVal('border-top'), '1px dashed slategray');
```

Após o primeiro elemento `hr`, deve haver um elemento `p` com um endereço de e-mail como seu texto.

```js
assert.isAbove($('.business-card hr + p')[0].innerText.length, 0);
```

Após o elemento `p` do e-mail, deve haver outro elemento `p` com um número de telefone como seu texto.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Após o elemento `p` do número de telefone, deve haver um elemento `a` com o texto `Portfolio`

```js
assert.equal(document.querySelector('.business-card a').textContent, 'Portfolio');
```

O elemento `a` do portfólio deve ter um `href` válido.

```js
assert.isAbove($('.business-card a')[0].href.length, 0);
```

Após o elemento `a` contendo seu portfólio, deve haver outro elemento `hr` como um divisor.

```js
assert.exists(document.querySelector('a + hr'));
```

Após o segundo elemento `hr`, deve haver um elemento `div` com a classe `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

Dentro do elemento `.social-media`, deve haver um elemento `h2` com o texto `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2').textContent, 'Connect with me');
```

Dentro do elemento `.social-media`, deve haver um elemento `a` com o texto `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0].textContent, 'Twitter');
```

O Twitter `href` deve apontar para um link válido.

```js
assert.isAbove($('.social-media a')[0].getAttribute('href').length, 0);
```

Dentro do elemento `.social-media`, deve haver um elemento `a` com o texto `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1].textContent, 'LinkedIn');
```

O LinkedIn `href` deve apontar para um link válido.

```js
assert.isAbove($('.social-media a')[1].getAttribute('href').length, 0);
```

Dentro do elemento `.social-media`, deve haver um elemento `a` com o texto `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2].textContent, 'GitHub');
```

O `href` do GitHub deve apontar para um link válido.

```js
assert.isAbove($('.social-media a')[2].getAttribute('href').length, 0);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

a:hover {
    text-decoration: underline;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    border-top: 1px dashed slategray;
    margin: 20px 0;
}
```
