---
id: 56105e7b514f539506016a5e
title: 使用 For 迴圈倒數計數
challengeType: 1
forumTopicId: 16808
dashedName: count-backwards-with-a-for-loop
---

# --description--

只要我們能定義正確的條件，for 迴圈也可以倒數計數。

為了每次迭代遞減 2，我們需要更改我們的初始化、條件和最終表達式。

我們將從 `i = 10` 開始，並在 `i > 0` 時執行迴圈。每次迴圈將以 `i -= 2` 將 `i` 減少 2。

```js
const ourArray = [];

for (let i = 10; i > 0; i -= 2) {
  ourArray.push(i);
}
```

`ourArray` 現在將包含 `[10, 8, 6, 4, 2]`。讓我們更改初始化和最終表達式，這樣我們就可以每次倒數兩個數字，來創建一個遞減的奇數陣列。

# --instructions--

使用 `for` 迴圈將 9 到 1 的奇數推入 `myArray`。

# --hints--

你應該為此使用 `for` 迴圈。

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

你應該使用陣列方法 `push`。

```js
assert(__helpers.removeJSComments(code).match(/myArray.push/));
```

`myArray` 應該與 `[9, 7, 5, 3, 1]` 相等。

```js
assert.deepEqual(myArray, [9, 7, 5, 3, 1]);
```

# --seed--

## --after-user-code--

```js
if(typeof myArray !== "undefined"){(function(){return myArray;})();}
```

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 9; i > 0; i -= 2) {
  myArray.push(i);
}
```
