---
id: 56104e9e514f539506016a5c
title: 使用 For 迴圈迭代奇數
challengeType: 1
forumTopicId: 18212
dashedName: iterate-odd-numbers-with-a-for-loop
---

# --description--

迴圈不必一次迭代一個。透過更改我們的 `final-expression`，我們可以以偶數計數。

我們將從 `i = 0` 開始，並在 `i < 10` 時執行迴圈。我們將在每個迴圈中以 `i += 2` 將 `i` 增加 2。

```js
const ourArray = [];

for (let i = 0; i < 10; i += 2) {
  ourArray.push(i);
}
```

`ourArray` 現在將包含 `[0, 2, 4, 6, 8]`。讓我們更改 `initialization`，以便我們可以按奇數計數。

# --instructions--

使用 `for` 迴圈將 1 到 9 之間的奇數推入 `myArray`。

# --hints--

你應該為此使用 `for` 迴圈。

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

`myArray` 應該與 `[1, 3, 5, 7, 9]` 相等。

```js
assert.deepEqual(myArray, [1, 3, 5, 7, 9]);
```

# --seed--

## --after-user-code--

```js
if(typeof myArray !== "undefined"){(function(){return myArray;})();}
```

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 10; i += 2) {
  myArray.push(i);
}
```
