---
id: 56533eb9ac21ba0edf2244b1
title: 使用增強乘法的複合指派
challengeType: 1
forumTopicId: 16662
dashedName: compound-assignment-with-augmented-multiplication
---

# --description--

`*=` 運算子會將變數乘以一個數字。

```js
myVar = myVar * 5;
```

將 `myVar` 乘以 `5`。這可以重寫為：

```js
myVar *= 5;
```

# --instructions--

將 `a`、`b` 和 `c` 的指派改為使用 `*=` 運算子。

# --hints--

`a` 應該相等於 `25`。

```js
assert(a === 25);
```

`b` 應該與 `36` 相等。

```js
assert(b === 36);
```

`c` 應該相等於 `46`。

```js
assert(c === 46);
```

你應該對每個變數使用 `*=` 運算子。

```js
assert(__helpers.removeJSComments(code).match(/\*=/g).length === 3);
```

你不應該修改指定註解上方的程式碼。

```js
assert(
  /let a = 5;/.test(__helpers.removeJSComments(code)) &&
    /let b = 12;/.test(__helpers.removeJSComments(code)) &&
    /let c = 4\.6;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --after-user-code--

```js
(function(a,b,c){ return "a = " + a + ", b = " + b + ", c = " + c; })(a,b,c);
```

## --seed-contents--

```js
let a = 5;
let b = 12;
let c = 4.6;

// Only change code below this line
a = a * 5;
b = 3 * b;
c = c * 10;
```

# --solutions--

```js
let a = 5;
let b = 12;
let c = 4.6;

a *= 5;
b *= 3;
c *= 10;
```
