---
id: 56533eb9ac21ba0edf2244cc
title: 存取巢狀物件
challengeType: 1
forumTopicId: 16161
dashedName: accessing-nested-objects
---

# --description--

物件的子屬性可以透過串鏈點號或中括號表示法來存取。

這是一個巢狀物件：

```js
const ourStorage = {
  "desk": {
    "drawer": "stapler"
  },
  "cabinet": {
    "top drawer": { 
      "folder1": "a file",
      "folder2": "secrets"
    },
    "bottom drawer": "soda"
  }
};

ourStorage.cabinet["top drawer"].folder2;
ourStorage.desk.drawer;
```

`ourStorage.cabinet["top drawer"].folder2` 會是字串 `secrets`，而 `ourStorage.desk.drawer` 會是字串 `stapler`。

# --instructions--

存取 `myStorage` 物件，並將 `glove box` 屬性的內容指定給 `gloveBoxContents` 變數。盡可能對所有屬性使用點記法，否則使用括號記法。

# --hints--

`gloveBoxContents` 應該與字串 `maps` 相等。

```js
assert(gloveBoxContents === 'maps');
```

你的程式碼應該盡可能使用點記法來存取 `myStorage`。

```js
assert.match(code, /myStorage\.car\.inside/);
```

`gloveBoxContents` 仍然應該用 `const` 宣告。

```js
assert.match(code, /const\s+gloveBoxContents\s*=/);
```

你不應該更改 `myStorage` 物件。

```js
const expectedMyStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
assert.deepStrictEqual(myStorage, expectedMyStorage);
```

# --seed--

## --after-user-code--

```js
(function(x) { 
  if(typeof x != 'undefined') { 
    return "gloveBoxContents = " + x;
  }
  return "gloveBoxContents is undefined";
})(gloveBoxContents);
```

## --seed-contents--

```js
const myStorage = {
  "car": {
    "inside": {
      "glove box": "maps",
      "passenger seat": "crumbs"
     },
    "outside": {
      "trunk": "jack"
    }
  }
};

const gloveBoxContents = undefined;
```

# --solutions--

```js
const myStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
const gloveBoxContents = myStorage.car.inside["glove box"];
```
