---
id: 56533eb9ac21ba0edf2244b6
title: 字串中的跳脫序列
challengeType: 1
forumTopicId: 17567
dashedName: escape-sequences-in-strings
---

# --description--

引號並不是字串中唯一可以被<dfn>跳脫</dfn>的字元。跳脫序列允許你使用在字串中可能無法使用的字元。

<table><thead><tr><th>Code</th><th>Output</th></tr></thead><tbody><tr><td><code>\'</code></td><td>single quote</td></tr><tr><td><code>\"</code></td><td>double quote</td></tr><tr><td><code>\\</code></td><td>backslash</td></tr><tr><td><code>\n</code></td><td>newline</td></tr><tr><td><code>\t</code></td><td>tab</td></tr><tr><td><code>\r</code></td><td>carriage return</td></tr><tr><td><code>\b</code></td><td>backspace</td></tr><tr><td><code>\f</code></td><td>form feed</td></tr></tbody></table>

*請注意，反斜線本身必須被轉義才能顯示為反斜線。*

# --instructions--

使用跳脫序列將以下三行文字指定到單一變數 `myStr`。

<pre>
FirstLine
    \SecondLine
ThirdLine
</pre>

你將需要使用跳脫序列來正確插入特殊字元。你也需要遵循上方的間距，不可在跳脫序列或字詞之間加入空格。

**注意：** `SecondLine` 的縮排是使用製表符跳脫字元達成的，而非空白字元。

# --hints--

`myStr` 不應包含任何空格

```js
assert(!/ /.test(myStr));
```

`myStr` 應該包含字串 `FirstLine`、`SecondLine` 和 `ThirdLine`（請記得大小寫敏感）。

```js
assert(
  /FirstLine/.test(myStr) && /SecondLine/.test(myStr) && /ThirdLine/.test(myStr)
);
```

`FirstLine` 應該接著換行字元 `\n`

```js
assert(/FirstLine\n/.test(myStr));
```

`myStr` 應該包含一個跟隨換行字元的製表符字元 `\t`

```js
assert(/\n\t/.test(myStr));
```

`SecondLine` 應該以反斜線字元 `\` 為前置。

```js
assert(/\\SecondLine/.test(myStr));
```

在 `SecondLine` 和 `ThirdLine` 之間應該有一個換行字元。

```js
assert(/SecondLine\nThirdLine/.test(myStr));
```

`myStr` 應該只包含說明中顯示的字元。

```js
assert(myStr === 'FirstLine\n\t\\SecondLine\nThirdLine');
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "FirstLine\n\t\\SecondLine\nThirdLine";
```
