---
id: 56533eb9ac21ba0edf2244ae
title: 在 JavaScript 中尋找餘數
challengeType: 1
forumTopicId: 18184
dashedName: finding-a-remainder-in-javascript
---

# --description--

<dfn>餘數</dfn>運算子 `%` 給出兩個數字相除的餘數。

**範例**

<pre>
5 % 2 = 1
5 / 2 = 2 remainder 1
2 * 2 = 4
5 - 4 = 1
</pre>

**用法**在數學中，可以透過檢查數字除以 `2` 的餘數來判斷該數字是偶數還是奇數。偶數的餘數為 `0`，而奇數的餘數為 `1`。

<pre>
17 % 2 = 1
48 % 2 = 0
</pre>

**注意：** <dfn>remainder</dfn> 運算子有時會被錯誤地取用為 modulus 運算子。它與 modulus 非常相似，但在負數上無法正確運作。

# --instructions--

將 `remainder` 設為 `11` 除以 `3` 的餘數，使用 <dfn>remainder</dfn> (`%`) 運算子。

# --hints--

變數 `remainder` 應該被初始化

```js
assert(/(const|let|var)\s+?remainder/.test(__helpers.removeJSComments(code)));
```

`remainder` 的值應該是 `2`

```js
assert(remainder === 2);
```

你應該使用 `%` 運算子

```js
assert(/\s+?remainder\s*?=\s*?.*%.*;?/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function (y) {
  return 'remainder = ' + y;
})(remainder);
```

## --seed-contents--

```js
const remainder = 0;
```

# --solutions--

```js
const remainder = 11 % 3;
```
