---
id: 56533eb9ac21ba0edf2244ed
title: 將變數附加到字串
challengeType: 1
forumTopicId: 16656
dashedName: appending-variables-to-strings
---

# --description--

正如我們可以透過多行字串 <dfn>字面值</dfn> 來建置字串，我們也可以使用加等於（`+=`）運算子將變數附加到字串。

範例：

```js
const anAdjective = "awesome!";
let ourStr = "freeCodeCamp is ";
ourStr += anAdjective;
```

`ourStr` 的值會是 `freeCodeCamp is awesome!`。

# --instructions--

將 `someAdjective` 設定為至少 3 個字元的字串，並使用 `+=` 運算子將它附加到 `myStr`。

# --hints--

`someAdjective` 應該設定為長度至少 3 個字元的字串。

```js
assert(typeof someAdjective !== 'undefined' && someAdjective.length > 2);
```

你應該使用 `+=` 運算子將 `someAdjective` 附加到 `myStr`。

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*someAdjective\s*/).length > 0);
```

# --seed--

## --after-user-code--

```js
(function(){
  var output = [];
  if(typeof someAdjective === 'string') {
    output.push('someAdjective = "' + someAdjective + '"');
  } else {
    output.push('someAdjective is not a string');
  }
  if(typeof myStr === 'string') {
    output.push('myStr = "' + myStr + '"');
  } else {
    output.push('myStr is not a string');
  }
  return output.join('\n');
})();
```

## --seed-contents--

```js
// Change code below this line
const someAdjective = "";
let myStr = "Learning to code is ";
```

# --solutions--

```js
const someAdjective = "neat";
let myStr = "Learning to code is ";
myStr += someAdjective;
```
