---
id: 56533eb9ac21ba0edf2244bb
title: 字謎
challengeType: 1
forumTopicId: 18377
dashedName: word-blanks
---

# --description--

你會得到一些缺少字詞（如名詞、動詞、形容詞和副詞）的句子。然後你用自己選擇的字詞填補缺漏，使完成的句子有意義。

重要：僅回傳上述翻譯文字。請勿添加任何註解、說明或解釋。

```md
It was really ____, and we ____ ourselves ____.
```

這句話缺少三個部分－一個形容詞、一個動詞和一個副詞，我們可以添加自己選擇的字詞來完成它。然後，我們可以將完成的句子指定給一個變數，如下所示：

```js
const sentence = "It was really " + "hot" + ", and we " + "laughed" + " ourselves " + "silly" + ".";
```

# --instructions--

在這個挑戰中，我們會提供你一個名詞、一個動詞、一個形容詞和一個副詞。你需要使用你選擇的字詞，連同我們提供的字詞，組成一個完整的句子。

你需要使用字串串接運算子 `+` 來建置一個新的字串，使用提供的變數：`myNoun`、`myAdjective`、`myVerb` 和 `myAdverb`。然後你將把形成的字串指定給 `wordBlanks` 變數。你不應該更改指定給變數的字詞。

你還需要考慮字串中的空格，以便最終句子中所有字之間都有空格。結果應該是一個完整的句子。

# --hints--

`wordBlanks` 應該是一個字串。

```js
assert(typeof wordBlanks === 'string');
```

你不應該更改指定給 `myNoun`、`myVerb`、`myAdjective` 或 `myAdverb` 的值。

```js
assert(
  myNoun === 'dog' &&
    myVerb === 'ran' &&
    myAdjective === 'big' &&
    myAdverb === 'quickly'
);
```

你不應該直接使用值 `dog`、`ran`、`big` 或 `quickly` 來創建 `wordBlanks`。

```js
const newCode = removeAssignments(__helpers.removeJSComments(code));
assert(
  !/dog/.test(newCode) &&
    !/ran/.test(newCode) &&
    !/big/.test(newCode) &&
    !/quickly/.test(newCode)
);
```

`wordBlanks` 應該包含所有指定給變數 `myNoun`、`myVerb`、`myAdjective` 和 `myAdverb` 的字，並以非字元字元（以及你選擇的任何額外字）分隔。

```js
assert(
  /\bdog\b/.test(wordBlanks) &&
    /\bbig\b/.test(wordBlanks) &&
    /\bran\b/.test(wordBlanks) &&
    /\bquickly\b/.test(wordBlanks)
);
```

# --seed--

## --after-user-code--

```js
const removeAssignments = str => str
  .replace(/myNoun\s*=\s*["']dog["']/g, '')
  .replace(/myAdjective\s*=\s*["']big["']/g, '')
  .replace(/myVerb\s*=\s*["']ran["']/g, '')
  .replace(/myAdverb\s*=\s*["']quickly["']/g, '');
```

## --seed-contents--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

// Only change code below this line
const wordBlanks = ""; // Change this line
// Only change code above this line
```

# --solutions--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

let wordBlanks = "Once there was a " + myNoun + " which was very " + myAdjective + ". ";
wordBlanks += "It " + myVerb + " " + myAdverb + " around the yard.";
```
