---
id: 56533eb9ac21ba0edf2244df
title: 在 Switch 述語中多個相同的選項
challengeType: 1
forumTopicId: 18242
dashedName: multiple-identical-options-in-switch-statements
---

# --description--

如果在 `switch` 述語的 `case` 中省略了 `break` 述語，則會執行後續的 `case` 述語，直到遇到 `break` 為止。如果你有多個輸入對應相同的輸出，你可以像這樣在 `switch` 述語中表述它們：

```js
let result = "";
switch (val) {
  case 1:
  case 2:
  case 3:
    result = "1, 2, or 3";
    break;
  case 4:
    result = "4 alone";
}
```

1、2 和 3 的情況都會產生相同的結果。

# --instructions--

撰寫一個 switch 述語，為以下範圍設定 `answer`：`1-3` - `Low``4-6` - `Mid``7-9` - `High`

**注意：**你需要為範圍內的每個數字都有一個 `case` 述語。

# --hints--

`sequentialSizes(1)` 應該傳回字串 `Low`

```js
assert(sequentialSizes(1) === 'Low');
```

`sequentialSizes(2)` 應該傳回字串 `Low`

```js
assert(sequentialSizes(2) === 'Low');
```

`sequentialSizes(3)` 應該傳回字串 `Low`

```js
assert(sequentialSizes(3) === 'Low');
```

`sequentialSizes(4)` 應該傳回字串 `Mid`

```js
assert(sequentialSizes(4) === 'Mid');
```

`sequentialSizes(5)` 應該傳回字串 `Mid`

```js
assert(sequentialSizes(5) === 'Mid');
```

`sequentialSizes(6)` 應該傳回字串 `Mid`

```js
assert(sequentialSizes(6) === 'Mid');
```

`sequentialSizes(7)` 應該傳回字串 `High`

```js
assert(sequentialSizes(7) === 'High');
```

`sequentialSizes(8)` 應該傳回字串 `High`

```js
assert(sequentialSizes(8) === 'High');
```

`sequentialSizes(9)` 應該傳回字串 `High`

```js
assert(sequentialSizes(9) === 'High');
```

你不應該使用任何 `if` 或 `else` 述語。

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

你應該有九個 `case` 述語。

```js
assert(__helpers.removeJSComments(code).match(/case/g).length === 9);
```

# --seed--

## --seed-contents--

```js
function sequentialSizes(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

sequentialSizes(1);
```

# --solutions--

```js
function sequentialSizes(val) {
  let answer = "";

  switch (val) {
    case 1:
    case 2:
    case 3:
      answer = "Low";
      break;
    case 4:
    case 5:
    case 6:
      answer = "Mid";
      break;
    case 7:
    case 8:
    case 9:
      answer = "High";
  }

  return answer;
}
```
