---
id: 56533eb9ac21ba0edf2244dd
title: 使用 Switch 述語從多個選項中選擇
challengeType: 1
forumTopicId: 18277
dashedName: selecting-from-many-options-with-switch-statements
---

# --description--

如果你需要將一個值與多個選項進行比對，你可以使用 <dfn>switch</dfn> 述語。`switch` 述語會將該值與定義各種可能值的 <dfn>case</dfn> 述語進行比較。任何有效的 JavaScript 述語都可以在 <dfn>case</dfn> 區塊內執行，並且會從第一個符合的 `case` 值開始執行，直到遇到 `break`。

這裡是一個 `switch` 述語的範例：

```js
switch (fruit) {
  case "apple":
    console.log("The fruit is an apple");
    break;
  case "orange":
    console.log("The fruit is an orange");
    break;
}
```

`case` 值會以嚴格相等性（`===`）進行測試。`break` 告訴 JavaScript 停止執行述語。如果省略 `break`，將會執行下一個述語。

# --instructions--

撰寫一個 switch 述語來測試 `val` 並根據以下條件設定 `answer`：`1` - `alpha``2` - `beta``3` - `gamma``4` - `delta`

# --hints--

`caseInSwitch(1)` 應該有字串 `alpha` 的值

```js
assert(caseInSwitch(1) === 'alpha');
```

`caseInSwitch(2)` 應該有字串 `beta` 的值

```js
assert(caseInSwitch(2) === 'beta');
```

`caseInSwitch(3)` 應該有字串 `gamma` 的值

```js
assert(caseInSwitch(3) === 'gamma');
```

`caseInSwitch(4)` 應該有字串 `delta` 的值

```js
assert(caseInSwitch(4) === 'delta');
```

你不應該使用任何 `if` 或 `else` 述語。

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

你應該至少有 3 個 `break` 述語

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function caseInSwitch(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

caseInSwitch(1);
```

# --solutions--

```js
function caseInSwitch(val) {
  let answer = "";

  switch (val) {
    case 1:
      answer = "alpha";
      break;
    case 2:
      answer = "beta";
      break;
    case 3:
      answer = "gamma";
      break;
    case 4:
      answer = "delta";
  }
  return answer;
}
```
