---
id: 56533eb9ac21ba0edf2244c0
title: 函式中的全域範圍與區域範圍
challengeType: 1
forumTopicId: 18194
dashedName: global-vs--local-scope-in-functions
---

# --description--

可以同時擁有名稱相同的 <dfn>區域的</dfn> 和 <dfn>全域的</dfn> 變數。當你這麼做時，區域的變數會優先於全域的變數。

在此範例中：

```js
const someVar = "Hat";

function myFun() {
  const someVar = "Head";
  return someVar;
}
```

函式（程式）`myFun` 將傳回字串 `Head`，因為區域的變數版本存在。

# --instructions--

為你的 `myOutfit` 函式添加一個區域的變數，以用字串 `sweater` 改寫 `outerWear` 的值。

# --hints--

你不應該更改全域的 `outerWear` 的值。

```js
assert(outerWear === 'T-Shirt');
```

`myOutfit` 應該傳回字串 `sweater`。

```js
assert(myOutfit() === 'sweater');
```

你不應該更改 `return` 述語。

```js
assert(/return outerWear/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const outerWear = "T-Shirt";

function myOutfit() {
  // Only change code below this line

  // Only change code above this line
  return outerWear;
}

myOutfit();
```

# --solutions--

```js
const outerWear = "T-Shirt";
function myOutfit() {
  const outerWear = "sweater";
  return outerWear;
}
```
