---
id: 56533eb9ac21ba0edf2244d3
title: 使用嚴格不相等運算子進行比較
challengeType: 1
forumTopicId: 16791
dashedName: comparison-with-the-strict-inequality-operator
---

# --description--

嚴格不相等運算子（`!==`）是嚴格相等運算子的邏輯相反。它表示「嚴格不相等」，並且在嚴格相等會傳回 `true` 的地方傳回 `false`，反之亦然。嚴格不相等運算子不會轉換資料型別。

**範例**

```js
3 !==  3  // false
3 !== '3' // true
4 !==  3  // true
```

# --instructions--

將嚴格不相等運算子添加到 `if` 述語中，這樣當 `val` 不嚴格相等於 `17` 時，函式將傳回字串 `Not Equal`。

# --hints--

`testStrictNotEqual(17)` 應該傳回字串 `Equal`

```js
assert(testStrictNotEqual(17) === 'Equal');
```

`testStrictNotEqual("17")` 應該傳回字串 `Not Equal`

```js
assert(testStrictNotEqual('17') === 'Not Equal');
```

`testStrictNotEqual(12)` 應該傳回字串 `Not Equal`

```js
assert(testStrictNotEqual(12) === 'Not Equal');
```

`testStrictNotEqual("bob")` 應該傳回字串 `Not Equal`

```js
assert(testStrictNotEqual('bob') === 'Not Equal');
```

你應該使用 `!==` 運算子

```js
assert(__helpers.removeJSComments(code).match(/(val\s*!==\s*\d+)|(\d+\s*!==\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrictNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testStrictNotEqual(10);
```

# --solutions--

```js
function testStrictNotEqual(val) {
  if (val !== 17) {
    return "Not Equal";
  }
  return "Equal";
}
```
