---
id: 56533eb9ac21ba0edf2244cd
title: 存取巢狀陣列
challengeType: 1
forumTopicId: 16160
dashedName: accessing-nested-arrays
---

# --description--

如同我們在先前範例中所見，物件可以包含巢狀物件和巢狀陣列。類似於存取巢狀物件，可以串鏈使用陣列括號表示法來存取巢狀陣列。

以下是如何存取巢狀陣列的範例：

```js
const ourPets = [
  {
    animalType: "cat",
    names: [
      "Meowzer",
      "Fluffy",
      "Kit-Cat"
    ]
  },
  {
    animalType: "dog",
    names: [
      "Spot",
      "Bowser",
      "Frankie"
    ]
  }
];

ourPets[0].names[1];
ourPets[1].names[0];
```

`ourPets[0].names[1]` 會是字串 `Fluffy`，而 `ourPets[1].names[0]` 會是字串 `Spot`。

# --instructions--

使用點記法和括號記法，將變數 `secondTree` 設定為 `myPlants` 陣列中第二個物件的 `list` 陣列中的第二個元素。

# --hints--

`secondTree` 應該相等於字串 `pine`。

```js
assert(secondTree === 'pine');
```

你的程式碼應該使用點記法和括號記法來存取 `myPlants`。

```js
assert(/=\s*myPlants\[1\].list\[1\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function(x) {
  if(typeof x != 'undefined') {
    return "secondTree = " + x;
  }
  return "secondTree is undefined";
})(secondTree);
```

## --seed-contents--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = "";
```

# --solutions--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = myPlants[1].list[1];
```
