---
id: 56533eb9ac21ba0edf2244c3
title: 帶有傳回值的指派
challengeType: 1
forumTopicId: 16658
dashedName: assignment-with-a-returned-value
---

# --description--

如果你還記得我們關於 <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/storing-values-with-the-assignment-operator" target="_blank" rel="noopener noreferrer nofollow">使用指定運算子儲存值</a> 的討論，等號右邊的所有內容都會在值被指定之前先被決議。這表示我們可以將函式的返回值指定給一個變數。

假設我們已經定義了一個函式（程式）`sum`，它會將兩個數字相加。

```js
ourSum = sum(5, 12);
```

呼叫 `sum` 函式（程式）並帶入引數 `5` 和 `12`，會產生返回值 `17`。此返回值會指定給 `ourSum` 變數。

# --instructions--

呼叫 `processArg` 函式（程式），並以引數 `7` 呼叫，將其返回值指定給變數 `processed`。

# --hints--

`processed` 應該有一個值為 `2`

```js
assert(processed === 2);
```

你應該將 `processArg` 指定到 `processed`。

```js
assert(/processed\s*=\s*processArg\(\s*7\s*\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function(){return "processed = " + processed})();
```

## --seed-contents--

```js
// Setup
let processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

// Only change code below this line

```

# --solutions--

```js
var processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

processed = processArg(7);
```
