---
id: 56533eb9ac21ba0edf2244d9
title: 使用邏輯或運算子的比較
challengeType: 1
forumTopicId: 16800
dashedName: comparisons-with-the-logical-or-operator
---

# --description--

<dfn>邏輯或</dfn> 運算子 (`||`) 若任一個 <dfn>運算元</dfn> 為 `true`，則傳回 `true`。否則，傳回 `false`。

<dfn>邏輯或</dfn>運算子由兩個管道符號組成：(`||`)。這通常位於你的 Backspace 和 Enter 鍵之間。

下面的樣式應該讓你想起之前的航點。

```js
if (num > 10) {
  return "No";
}
if (num < 5) {
  return "No";
}
return "Yes";
```

如果 `num` 在 `5` 和 `10`（包含 `5` 和 `10`）之間，這段程式碼將傳回 `Yes`。相同的邏輯可以用 <dfn>logical or</dfn> 運算子來撰寫。

```js
if (num > 10 || num < 5) {
  return "No";
}
return "Yes";
```

# --instructions--

將兩個 `if` 述語合併成一個述語，如果 `val` 不在包含 `10` 和 `20` 的範圍內，則傳回字串 `Outside`。否則，傳回字串 `Inside`。

# --hints--

你應該使用一次 `||` 運算子

```js
assert(__helpers.removeJSComments(code).match(/\|\|/g).length === 1);
```

你應該只使用一個 `if` 述語

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalOr(0)` 應該傳回字串 `Outside`

```js
assert(testLogicalOr(0) === 'Outside');
```

`testLogicalOr(9)` 應該傳回字串 `Outside`

```js
assert(testLogicalOr(9) === 'Outside');
```

`testLogicalOr(10)` 應該傳回字串 `Inside`

```js
assert(testLogicalOr(10) === 'Inside');
```

`testLogicalOr(15)` 應該傳回字串 `Inside`

```js
assert(testLogicalOr(15) === 'Inside');
```

`testLogicalOr(19)` 應該傳回字串 `Inside`

```js
assert(testLogicalOr(19) === 'Inside');
```

`testLogicalOr(20)` 應該傳回字串 `Inside`

```js
assert(testLogicalOr(20) === 'Inside');
```

`testLogicalOr(21)` 應該傳回字串 `Outside`

```js
assert(testLogicalOr(21) === 'Outside');
```

`testLogicalOr(25)` 應該傳回字串 `Outside`

```js
assert(testLogicalOr(25) === 'Outside');
```

# --seed--

## --seed-contents--

```js
function testLogicalOr(val) {
  // Only change code below this line

  if (val) {
    return "Outside";
  }

  if (val) {
    return "Outside";
  }

  // Only change code above this line
  return "Inside";
}

testLogicalOr(15);
```

# --solutions--

```js
function testLogicalOr(val) {
  if (val < 10 || val > 20) {
    return "Outside";
  }
  return "Inside";
}
```
