---
id: 56533eb9ac21ba0edf2244d8
title: 使用邏輯 And 運算子的比較
challengeType: 1
forumTopicId: 16799
dashedName: comparisons-with-the-logical-and-operator
---

# --description--

有時你需要同時測試多個條件。<dfn>邏輯且</dfn>運算子（`&&`）若且唯若其左右兩邊的<dfn>運算元</dfn>皆為 true 時，才會傳回 `true`。

相同的效果可以透過將一個 `if` 述語巢狀在另一個 `if` 述語內來達成。

```js
if (num > 5) {
  if (num < 10) {
    return "Yes";
  }
}
return "No";
```

如果 `num` 大於 5 且小於 10，這段程式碼將傳回 `Yes`。相同的邏輯可以用 <dfn>logical and</dfn> 運算子來撰寫。

```js
if (num > 5 && num < 10) {
  return "Yes";
}
return "No";
```

# --instructions--

將兩個 if 述語取代為一個述語，使用 `&&` 運算子，當 `val` 小於或相等於 50 且大於或相等於 25 時，將傳回字串 `Yes`。否則，將傳回字串 `No`。

# --hints--

你應該使用一次 `&&` 運算子

```js
assert(__helpers.removeJSComments(code).match(/&&/g).length === 1);
```

你應該只使用一個 `if` 述語

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalAnd(0)` 應該傳回字串 `No`

```js
assert(testLogicalAnd(0) === 'No');
```

`testLogicalAnd(24)` 應該傳回字串 `No`

```js
assert(testLogicalAnd(24) === 'No');
```

`testLogicalAnd(25)` 應該傳回字串 `Yes`

```js
assert(testLogicalAnd(25) === 'Yes');
```

`testLogicalAnd(30)` 應該傳回字串 `Yes`

```js
assert(testLogicalAnd(30) === 'Yes');
```

`testLogicalAnd(50)` 應該傳回字串 `Yes`

```js
assert(testLogicalAnd(50) === 'Yes');
```

`testLogicalAnd(51)` 應該傳回字串 `No`

```js
assert(testLogicalAnd(51) === 'No');
```

`testLogicalAnd(75)` 應該傳回字串 `No`

```js
assert(testLogicalAnd(75) === 'No');
```

`testLogicalAnd(80)` 應該傳回字串 `No`

```js
assert(testLogicalAnd(80) === 'No');
```

# --seed--

## --seed-contents--

```js
function testLogicalAnd(val) {
  // Only change code below this line

  if (val) {
    if (val) {
      return "Yes";
    }
  }

  // Only change code above this line
  return "No";
}

testLogicalAnd(10);
```

# --solutions--

```js
function testLogicalAnd(val) {
  if (val >= 25 && val <= 50) {
    return "Yes";
  }
  return "No";
}
```
