---
id: 56533eb9ac21ba0edf2244bd
title: 將值傳遞給帶有引數的函式（程式）
challengeType: 1
forumTopicId: 18254
dashedName: passing-values-to-functions-with-arguments
---

# --description--

<dfn>參數</dfn>是作為佔位符的變數，用來接收在呼叫函式時輸入的值。當定義函式時，通常會同時定義一個或多個參數。實際在呼叫函式時輸入（或<dfn>「傳遞」</dfn>）的值稱為<dfn>引數</dfn>。

這是一個有兩個參數 `param1` 和 `param2` 的函式（程式）：

```js
function testFun(param1, param2) {
  console.log(param1, param2);
}
```

然後我們可以這樣呼叫 `testFun`：`testFun("Hello", "World");`。我們已經傳遞了兩個字串引數，`Hello` 和 `World`。在函式內，`param1` 將會相等於字串 `Hello`，而 `param2` 將會相等於字串 `World`。請注意，你可以用不同的引數再次呼叫 `testFun`，參數將會取得新引數的值。

# --instructions--

<ol><li>Create a function called <code>functionWithArgs</code> that accepts two arguments and outputs their sum to the dev console.</li><li>Call the function with two numbers as arguments.</li></ol>

# --before-each--

```js
var logOutput;
```

# --hints--

`functionWithArgs` 應該是一個函式（程式）。

```js
assert(typeof functionWithArgs === 'function');
```

`functionWithArgs(1,2)` 應該輸出 `3`。

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(1, 2);
assert(logOutput == 3);
```

`functionWithArgs(7,9)` 應該輸出 `16`。

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(7, 9);
assert(logOutput == 16);
```

你應該在定義 `functionWithArgs` 之後，用兩個數字來呼叫它。

```js
assert(
  /functionWithArgs\([-+]?\d*\.?\d*,[-+]?\d*\.?\d*\)/.test(
    __helpers.removeJSComments(code).replace(/\s/g, '')
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function functionWithArgs(a, b) {
  console.log(a + b);
}
functionWithArgs(10, 5);
```
