---
id: 56533eb9ac21ba0edf2244ac
title: 使用 JavaScript 遞增數字
challengeType: 1
forumTopicId: 18201
dashedName: increment-a-number-with-javascript
---

# --description--

你可以輕鬆地使用 `++` 運算子對變數進行 <dfn>遞增</dfn> 或加一。

```js
i++;
```

是等價的

```js
i = i + 1;
```

**注意：**整行變成 `i++;`，不需要等號。

# --instructions--

將程式碼更改為對 `myVar` 使用 `++` 運算子。

# --hints--

`myVar` 應該相等於 `88`。

```js
assert(myVar === 88);
```

你不應該使用指定運算子。

```js
assert(
  /let\s+myVar\s*=\s*87;\s*\/*.*\s*([+]{2}\s*myVar|myVar\s*[+]{2})/.test(__helpers.removeJSComments(code))
);
```

你應該使用 `++` 運算子。

```js
assert(/[+]{2}\s*myVar|myVar\s*[+]{2}/.test(__helpers.removeJSComments(code)));
```

你不應該更改指定註解以上的程式碼。

```js
assert(/let myVar = 87;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'myVar = ' + z;})(myVar);
```

## --seed-contents--

```js
let myVar = 87;

// Only change code below this line
myVar = myVar + 1;
```

# --solutions--

```js
let myVar = 87;
myVar++;
```
