---
id: 56533eb9ac21ba0edf2244be
title: 全域的範圍與函式
challengeType: 1
forumTopicId: 18193
dashedName: global-scope-and-functions
---

# --description--

在 JavaScript 中，<dfn>scope</dfn> 指的是變數的可見性。在函式（程式）區塊外定義的變數具有 <dfn>Global</dfn> 範圍。這表示它們可以在你的整個 JavaScript 程式碼中被取用。

未使用 `let` 或 `const` 關鍵字宣告的變數會自動在 `global` 範圍內創建。這可能會在你的程式碼其他地方或再次執行函式時產生意想不到的後果。你應該始終使用 `let` 或 `const` 宣告你的變數。

# --instructions--

使用 `let` 或 `const`，在任何函式（程式）之外宣告一個名為 `myGlobal` 的全域變數。並將其初始化為 `10`。

在函式（程式）`fun1`內，***不使用*** `var`、`let` 或 `const` 關鍵字，將 `5` 指定給 `oopsGlobal`。

# --before-each--

```js
var oopsGlobal;
```

# --hints--

`myGlobal` 應該被定義

```js
assert(typeof myGlobal != 'undefined');
```

`myGlobal` 應該有一個值為 `10`

```js
assert(myGlobal === 10);
```

`myGlobal` 應該使用 `let` 或 `const` 關鍵字宣告

```js
assert(/(let|const)\s+myGlobal/.test(__helpers.removeJSComments(code)));
```

`oopsGlobal` 應該是全域變數，且其值為 `5`

```js
fun1();
assert(typeof oopsGlobal != 'undefined');
```

# --seed--

## --seed-contents--

```js
// Declare the myGlobal variable below this line


function fun1() {
  // Assign 5 to oopsGlobal here

}

// Only change code above this line

function fun2() {
  let output = "";
  if (typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if (typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```

# --solutions--

```js
const myGlobal = 10;

function fun1() {
  oopsGlobal = 5;
}

function fun2() {
  let output = "";
  if(typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if(typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```
