---
id: 56533eb9ac21ba0edf2244b0
title: 使用增強減法的複合指派
challengeType: 1
forumTopicId: 16660
dashedName: compound-assignment-with-augmented-subtraction
---

# --description--

像 `+=` 運算子一樣，`-=` 從變數中減去一個數字。

```js
myVar = myVar - 5;
```

將 `5` 從 `myVar` 中減去。這可以重寫為：

```js
myVar -= 5;
```

# --instructions--

將 `a`、`b` 和 `c` 的指派改為使用 `-=` 運算子。

# --hints--

`a` 應該相等於 `5`。

```js
assert(a === 5);
```

`b` 應該相等於 `-6`。

```js
assert(b === -6);
```

`c` 應該相等於 `2`。

```js
assert(c === 2);
```

你應該對每個變數使用 `-=` 運算子。

```js
assert(__helpers.removeJSComments(code).match(/-=/g).length === 3);
```

你不應該修改指定註解上方的程式碼。

```js
assert(
  /let a = 11;/.test(__helpers.removeJSComments(code)) && /let b = 9;/.test(__helpers.removeJSComments(code)) && /let c = 3;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --after-user-code--

```js
(function(a,b,c){ return "a = " + a + ", b = " + b + ", c = " + c; })(a,b,c);
```

## --seed-contents--

```js
let a = 11;
let b = 9;
let c = 3;

// Only change code below this line
a = a - 6;
b = b - 15;
c = c - 1;
```

# --solutions--

```js
let a = 11;
let b = 9;
let c = 3;

a -= 6;
b -= 15;
c -= 1;
```
