---
id: 56533eb9ac21ba0edf2244d5
title: 使用大於或相等於運算子進行比較
challengeType: 1
forumTopicId: 16785
dashedName: comparison-with-the-greater-than-or-equal-to-operator
---

# --description--

大於或相等運算子（`>=`）比較兩個數字的值。如果左邊的數字大於或相等於右邊的數字，它會傳回 `true`。否則，它會傳回 `false`。

像等號運算子一樣，大於或等於運算子在比較時會轉換資料型別。

**範例**

```js
6   >=  6  // true
7   >= '3' // true
2   >=  3  // false
'7' >=  9  // false
```

# --instructions--

將大於或相等於運算子添加到所指出的行，以使傳回述語有意義。

# --hints--

`testGreaterOrEqual(0)` 應該傳回字串 `Less than 10`

```js
assert(testGreaterOrEqual(0) === 'Less than 10');
```

`testGreaterOrEqual(9)` 應該傳回字串 `Less than 10`

```js
assert(testGreaterOrEqual(9) === 'Less than 10');
```

`testGreaterOrEqual(10)` 應該傳回字串 `10 or Over`

```js
assert(testGreaterOrEqual(10) === '10 or Over');
```

`testGreaterOrEqual(11)` 應該傳回字串 `10 or Over`

```js
assert(testGreaterOrEqual(11) === '10 or Over');
```

`testGreaterOrEqual(19)` 應該傳回字串 `10 or Over`

```js
assert(testGreaterOrEqual(19) === '10 or Over');
```

`testGreaterOrEqual(100)` 應該傳回字串 `20 or Over`

```js
assert(testGreaterOrEqual(100) === '20 or Over');
```

`testGreaterOrEqual(21)` 應該傳回字串 `20 or Over`

```js
assert(testGreaterOrEqual(21) === '20 or Over');
```

你應該至少使用兩次 `>=` 運算子

```js
assert(__helpers.removeJSComments(code).match(/val\s*>=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterOrEqual(val) {
  if (val) {  // Change this line
    return "20 or Over";
  }

  if (val) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}

testGreaterOrEqual(10);
```

# --solutions--

```js
function testGreaterOrEqual(val) {
  if (val >= 20) {  // Change this line
    return "20 or Over";
  }

  if (val >= 10) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}
```
