---
id: 56533eb9ac21ba0edf2244de
title: 在 Switch 述語中添加預設選項
challengeType: 1
forumTopicId: 16653
dashedName: adding-a-default-option-in-switch-statements
---

# --description--

在 `switch` 述語中，你可能無法指定所有可能的值作為 `case` 述語。相反地，你可以添加 `default` 述語，當找不到符合的 `case` 述語時，該述語將會被執行。可以將它視為 `if/else` 串鏈中的最後一個 `else` 述語。

`default` 述語應該是最後一個案例。

```js
switch (num) {
  case value1:
    statement1;
    break;
  case value2:
    statement2;
    break;
...
  default:
    defaultStatement;
    break;
}
```

# --instructions--

撰寫一個 switch 述語來設定 `answer`，條件如下：`a` - `apple`、`b` - `bird`、`c` - `cat`、`default` - `stuff`

# --hints--

`switchOfStuff("a")` 應該傳回字串 `apple`

```js
assert(switchOfStuff('a') === 'apple');
```

`switchOfStuff("b")` 應該傳回字串 `bird`

```js
assert(switchOfStuff('b') === 'bird');
```

`switchOfStuff("c")` 應該傳回字串 `cat`

```js
assert(switchOfStuff('c') === 'cat');
```

`switchOfStuff("d")` 應該傳回字串 `stuff`

```js
assert(switchOfStuff('d') === 'stuff');
```

`switchOfStuff(4)` 應該傳回字串 `stuff`

```js
assert(switchOfStuff(4) === 'stuff');
```

你不應該使用任何 `if` 或 `else` 述語。

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

你應該使用 `default` 述語

```js
assert(switchOfStuff('string-to-trigger-default-case') === 'stuff');
```

你應該至少有 3 個 `break` 述語

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function switchOfStuff(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

switchOfStuff(1);
```

# --solutions--

```js
function switchOfStuff(val) {
  let answer = "";

  switch(val) {
    case "a":
      answer = "apple";
      break;
    case "b":
      answer = "bird";
      break;
    case "c":
      answer = "cat";
      break;
    default:
      answer = "stuff";
  }
  return answer;
}
```
