---
id: 56533eb9ac21ba0edf2244d2
title: 使用不等號運算子進行比較
challengeType: 1
forumTopicId: 16787
dashedName: comparison-with-the-inequality-operator
---

# --description--

不等號運算子（`!=`）是等號運算子的相反。不等號表示不相等。當等號運算子會傳回 `true` 時，不等號運算子會傳回 `false`，反之亦然。與等號運算子相同，不等號運算子在比較時會轉換值的資料型別。

**範例**

```js
1 !=  2    // true
1 != "1"   // false
1 != '1'   // false
1 != true  // false
0 != false // false
```

# --instructions--

在 `if` 述語中為不等號運算子 `!=` 添加，這樣當 `val` 不等價於 `99` 時，函式將傳回字串 `Not Equal`。

# --hints--

`testNotEqual(99)` 應該傳回字串 `Equal`

```js
assert(testNotEqual(99) === 'Equal');
```

`testNotEqual("99")` 應該傳回字串 `Equal`

```js
assert(testNotEqual('99') === 'Equal');
```

`testNotEqual(12)` 應該傳回字串 `Not Equal`

```js
assert(testNotEqual(12) === 'Not Equal');
```

`testNotEqual("12")` 應該傳回字串 `Not Equal`

```js
assert(testNotEqual('12') === 'Not Equal');
```

`testNotEqual("bob")` 應該傳回字串 `Not Equal`

```js
assert(testNotEqual('bob') === 'Not Equal');
```

你應該使用 `!=` 運算子

```js
assert(__helpers.removeJSComments(code).match(/(?!!==)!=/));
```

# --seed--

## --seed-contents--

```js
// Setup
function testNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testNotEqual(10);
```

# --solutions--

```js
function testNotEqual(val) {
  if (val != 99) {
    return "Not Equal";
  }
  return "Equal";
}
```
