---
id: 56533eb9ac21ba0edf2244bc
title: 購物列表
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

在變數 `myList` 中創建購物列表。該列表應該是一個包含多個子陣列的多維陣列。

每個子陣列中的第一個元素應該包含一個字串，該字串為該條款的名稱。第二個元素應該是一個數字，用來表述數量，例如。

```js
["Chocolate Bar", 15]
```

列表中應該至少有 5 個子陣列。

# --hints--

`myList` 應該是一個陣列。

```js
assert(Array.isArray(myList));
```

你每個子陣列中的第一個元素都應該是字串。

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

你每個子陣列中的第二個元素都應該是數字。

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

你的列表中應該至少有 5 項目。

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
