---
id: 56533eb9ac21ba0edf2244af
title: 使用增強加法的複合指派
challengeType: 1
forumTopicId: 16661
dashedName: compound-assignment-with-augmented-addition
---

# --description--

在程式設計中，使用指派來修改變數的內容是很常見的。請記住，等號右邊的所有內容會先被評估，所以我們可以說：

```js
myVar = myVar + 5;
```

將 `5` 添加到 `myVar`。由於這是如此常見的樣式，有些運算子可以在一步驟中同時執行數學操作和指派。

其中一個這樣的運算子是 `+=` 運算子。

```js
let myVar = 1;
myVar += 5;
console.log(myVar);
```

`6` 會顯示在主控台。

# --instructions--

將 `a`、`b` 和 `c` 的指派改為使用 `+=` 運算子。

# --hints--

`a` 應該相等於 `15`。

```js
assert(a === 15);
```

`b` 應該相等於 `26`。

```js
assert(b === 26);
```

`c` 應該相等於 `19`。

```js
assert(c === 19);
```

你應該對每個變數使用 `+=` 運算子。

```js
assert(__helpers.removeJSComments(code).match(/\+=/g).length === 3);
```

你不應該修改指定註解上方的程式碼。

```js
assert(
  /let a = 3;/.test(__helpers.removeJSComments(code)) &&
    /let b = 17;/.test(__helpers.removeJSComments(code)) &&
    /let c = 12;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --after-user-code--

```js
(function(a,b,c){ return "a = " + a + ", b = " + b + ", c = " + c; })(a,b,c);
```

## --seed-contents--

```js
let a = 3;
let b = 17;
let c = 12;

// Only change code below this line
a = a + 12;
b = 9 + b;
c = c + 7;
```

# --solutions--

```js
let a = 3;
let b = 17;
let c = 12;

a += 12;
b += 9;
c += 7;
```
