---
id: 56533eb9ac21ba0edf2244b7
title: 使用加號運算子串接字串
challengeType: 1
forumTopicId: 16802
dashedName: concatenating-strings-with-plus-operator
---

# --description--

在 JavaScript 中，當 `+` 運算子用於 `String` 值時，它被稱為 <dfn>串接</dfn> 運算子。你可以透過將字串 <dfn>串接</dfn> 在一起，來建造一個新的字串。

**範例**

```js
'My name is Alan,' + ' I concatenate.'
```

**注意：** 注意空格。串接不會在串接的字串之間加入空格，所以你需要自己添加空格。

範例：

```js
const ourStr = "I come first. " + "I come second.";
```

字串 `I come first. I come second.` 會顯示在主控台。
# --instructions--

使用 `+` 運算子，從字串 `This is the start.` 和 `This is the end.` 建置 `myStr`。務必在兩個字串之間包含一個空格。

# --hints--

`myStr` 應該在兩個字串之間有一個空白字元。

```js
assert(/start\. This/.test(myStr));
```

`myStr` 應該有字串 `This is the start. This is the end.` 的值。

```js
assert(myStr === 'This is the start. This is the end.');
```

你應該使用 `+` 運算子來建置 `myStr`。

```js
assert(__helpers.removeJSComments(code).match(/(["']).*\1\s*\+\s*(["']).*\2/g));
```

`myStr` 應該使用 `const` 關鍵字創建。

```js
assert(/const\s+myStr/.test(__helpers.removeJSComments(code)));
```

你應該將結果指定到 `myStr` 變數。

```js
assert(/myStr\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function(){
  if(typeof myStr === 'string') {
    return 'myStr = "' + myStr + '"';
  } else {
    return 'myStr is not a string';
  }
})();
```

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "This is the start. " + "This is the end.";
```
