---
id: 56533eb9ac21ba0edf2244cb
title: 操作複雜物件
challengeType: 1
forumTopicId: 18208
dashedName: manipulating-complex-objects
---

# --description--

有時你可能想將資料儲存在彈性的 <dfn>資料結構</dfn> 中。JavaScript 物件是一種控制彈性資料的方式。它們允許任意組合 <dfn>字串</dfn>、<dfn>數字</dfn>、<dfn>布林值</dfn>、<dfn>陣列</dfn>、<dfn>函式</dfn> 和 <dfn>物件</dfn>。

以下是一個複雜資料結構的範例：

```js
const ourMusic = [
  {
    "artist": "Daft Punk",
    "title": "Homework",
    "release_year": 1997,
    "formats": [ 
      "CD", 
      "Cassette", 
      "LP"
    ],
    "gold": true
  }
];
```

這是一個包含一個物件的陣列。該物件具有關於專輯的各種<dfn>詮釋資料</dfn>。它也有一個巢狀的 `formats` 陣列。如果你想為你的專輯記錄添加更多記錄，你可以透過將記錄添加到頂層陣列來做到這一點。物件在屬性中保存資料，該屬性具有鍵值格式。在上述範例中，`"artist": "Daft Punk"` 是一個屬性，其鍵為 `artist`，值為 Daft Punk。

**注意：**你需要在陣列中每個物件後面加上逗號，除非它是陣列中的最後一個物件。

# --instructions--

為你的 `myMusic` 陣列添加一個新的專輯。添加 `artist` 和 `title` 字串、`release_year` 數字，以及一個包含字串的 `formats` 陣列。

# --hints--

`myMusic` 應該是一個陣列

```js
assert(Array.isArray(myMusic));
```

`myMusic` 應該至少有兩個元素

```js
assert(myMusic.length > 1);
```

`myMusic` 陣列中的元素應該是物件

```js
myMusic.forEach(object => {assert.typeOf(object, 'object')})
```

你的 `myMusic` 中的物件應該至少有 4 個屬性

```js
myMusic.forEach(object => {assert(Object.keys(object).length > 3); });
```

你在 `myMusic` 中的物件應該包含屬性 `artist`，該屬性為字串。

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['artist']);
  assert.typeOf(object.artist, 'string')
})
```

你在 `myMusic` 中的物件應該包含屬性 `title`，該屬性為字串。

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['title']);
  assert.typeOf(object.title, 'string')
})
```

你的 `myMusic` 中的物件應該包含數字型態的屬性 `release_year`。

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['release_year']);
  assert.typeOf(object.release_year, 'number')
})
```

你的 `myMusic` 中的物件應該包含一個 `formats` 屬性，該屬性為陣列

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['formats']);
  assert.typeOf(object.formats, 'array')
})
```

`formats` 應該是一個至少包含兩個元素的字串陣列。

```js
myMusic.forEach(object => {
  object.formats.forEach(format => {
    assert.typeOf(format, 'string')
  });
  assert.isAtLeast(object.formats.length, 2)
})
```

# --seed--

## --after-user-code--

```js
(function(x){ if (Array.isArray(x)) { return JSON.stringify(x); } return "myMusic is not an array"})(myMusic);
```

## --seed-contents--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CD",
      "8T",
      "LP"
    ],
    "gold": true
  }
];
```

# --solutions--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP" ],
    "gold": true
  },
  {
    "artist": "ABBA",
    "title": "Ring Ring",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP",
    "CD",
  ]
  }
];
```
