---
id: 56533eb9ac21ba0edf2244d6
title: 與小於運算子的比較
challengeType: 1
forumTopicId: 16789
dashedName: comparison-with-the-less-than-operator
---

# --description--

小於運算子（`<`）比較兩個數字的值。如果左邊的數字小於右邊的數字，它會傳回 `true`。否則，它會傳回 `false`。像等號運算子一樣，小於運算子在比較時會轉換資料型別。

**範例**

```js
2   < 5 // true
'3' < 7 // true
5   < 5 // false
3   < 2 // false
'8' < 4 // false
```

# --instructions--

將小於運算子添加到所顯示的行，以使傳回述語有意義。

# --hints--

`testLessThan(0)` 應該傳回字串 `Under 25`

```js
assert(testLessThan(0) === 'Under 25');
```

`testLessThan(24)` 應該傳回字串 `Under 25`

```js
assert(testLessThan(24) === 'Under 25');
```

`testLessThan(25)` 應該傳回字串 `Under 55`

```js
assert(testLessThan(25) === 'Under 55');
```

`testLessThan(54)` 應該傳回字串 `Under 55`

```js
assert(testLessThan(54) === 'Under 55');
```

`testLessThan(55)` 應該傳回字串 `55 or Over`

```js
assert(testLessThan(55) === '55 or Over');
```

`testLessThan(99)` 應該傳回字串 `55 or Over`

```js
assert(testLessThan(99) === '55 or Over');
```

你應該至少使用兩次 `<` 運算子。

```js
assert(__helpers.removeJSComments(code).match(/val\s*<\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessThan(val) {
  if (val) {  // Change this line
    return "Under 25";
  }

  if (val) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}

testLessThan(10);
```

# --solutions--

```js
function testLessThan(val) {
  if (val < 25) {  // Change this line
    return "Under 25";
  }

  if (val < 55) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}
```
