---
id: 56533eb9ac21ba0edf2244b8
title: 使用加等運算子串接字串
challengeType: 1
forumTopicId: 16803
dashedName: concatenating-strings-with-the-plus-equals-operator
---

# --description--

我們也可以使用 `+=` 運算子將字串 <dfn>串接</dfn> 到現有字串變數的末尾。這對於將長字串分成多行非常有幫助。

**注意：** 注意空格。串接不會在串接的字串之間加入空格，所以你需要自己添加空格。

範例：

```js
let ourStr = "I come first. ";
ourStr += "I come second.";
```

`ourStr` 現在的值是字串 `I come first. I come second.`。

# --instructions--

透過串接這兩個字串 `This is the first sentence.` 和 `This is the second sentence.`，使用 `+=` 運算子將 `myStr` 建置成多行。使用 `+=` 運算子的方法類似範例所示，並且務必在兩個字串之間包含一個空格。先將第一個字串指定給 `myStr`，然後再將第二個字串添加上去。

# --hints--

`myStr` 應該在兩個字串之間有一個空白字元。

```js
assert(/sentence\. This/.test(myStr));
```

`myStr` 應該具有字串 `This is the first sentence. This is the second sentence.` 的值。

```js
assert(myStr === 'This is the first sentence. This is the second sentence.');
```

你應該使用 `+=` 運算子來建置 `myStr`。

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*(["']).*\1/g));
```

# --seed--

## --after-user-code--

```js
(function(){
  if(typeof myStr === 'string') {
    return 'myStr = "' + myStr + '"';
  } else {
    return 'myStr is not a string';
  }
})();
```

## --seed-contents--

```js
let myStr;
```

# --solutions--

```js
let myStr = "This is the first sentence. ";
myStr += "This is the second sentence.";
```
