---
id: 56533eb9ac21ba0edf2244e0
title: 用 Switch 取代 If Else 串鏈
challengeType: 1
forumTopicId: 18266
dashedName: replacing-if-else-chains-with-switch
---

# --description--

如果你有許多選項，可選方案可供選擇，使用 `switch` 述語比多個串鏈的 `if`/`else if` 述語更容易撰寫。以下：

```js
if (val === 1) {
  answer = "a";
} else if (val === 2) {
  answer = "b";
} else {
  answer = "c";
}
```

可以取代為：

```js
switch (val) {
  case 1:
    answer = "a";
    break;
  case 2:
    answer = "b";
    break;
  default:
    answer = "c";
}
```

# --instructions--

將串鏈的 `if`/`else if` 述語改成 `switch` 述語。

# --hints--

你不應該在編輯器中的任何地方使用 `else` 述語。

```js
assert(!/else/g.test(__helpers.removeJSComments(code)));
```

你不應該在編輯器中的任何地方使用 `if` 述語。

```js
assert(!/if/g.test(__helpers.removeJSComments(code)));
```

你應該至少有四個 `break` 述語

```js
assert(__helpers.removeJSComments(code).match(/break/g).length >= 4);
```

`chainToSwitch("bob")` 應該傳回字串 `Marley`

```js
assert(chainToSwitch('bob') === 'Marley');
```

`chainToSwitch(42)` 應該傳回字串 `The Answer`

```js
assert(chainToSwitch(42) === 'The Answer');
```

`chainToSwitch(1)` 應該傳回字串 `There is no #1`

```js
assert(chainToSwitch(1) === 'There is no #1');
```

`chainToSwitch(99)` 應該傳回字串 `Missed me by this much!`

```js
assert(chainToSwitch(99) === 'Missed me by this much!');
```

`chainToSwitch(7)` 應該傳回字串 `Ate Nine`

```js
assert(chainToSwitch(7) === 'Ate Nine');
```

`chainToSwitch("John")` 應該傳回 `""`（空字串）

```js
assert(chainToSwitch('John') === '');
```

`chainToSwitch(156)` 應該傳回 `""`（空字串）

```js
assert(chainToSwitch(156) === '');
```

# --seed--

## --seed-contents--

```js
function chainToSwitch(val) {
  let answer = "";
  // Only change code below this line

  if (val === "bob") {
    answer = "Marley";
  } else if (val === 42) {
    answer = "The Answer";
  } else if (val === 1) {
    answer = "There is no #1";
  } else if (val === 99) {
    answer = "Missed me by this much!";
  } else if (val === 7) {
    answer = "Ate Nine";
  }

  // Only change code above this line
  return answer;
}

chainToSwitch(7);
```

# --solutions--

```js
function chainToSwitch(val) {
  let answer = "";

  switch (val) {
    case "bob":
      answer = "Marley";
      break;
    case 42:
      answer = "The Answer";
      break;
    case 1:
      answer = "There is no #1";
      break;
    case 99:
      answer = "Missed me by this much!";
      break;
    case 7:
      answer = "Ate Nine";
  }
  return answer;
}
```
