---
id: 56533eb9ac21ba0edf2244c6
title: 排隊
challengeType: 1
forumTopicId: 18307
dashedName: stand-in-line
---

# --description--

在電腦科學中，<dfn>queue</dfn> 是一種抽象的 <dfn>Data Structure</dfn>，其中項目以順序方式保存。新的項目可以添加到佇列的尾端，舊的項目則從佇列的前端取出。

# --instructions--

撰寫一個函式（程式）`nextInLine`，該函式接受一個陣列（`arr`）和一個數字（`item`）作為引數。

將數字添加到陣列的末端，然後移除陣列的第一個元素。

`nextInLine` 函式（程式）接著應該傳回被移除的元素。

# --hints--

`nextInLine([], 5)` 應該傳回一個數字。

```js
assert.isNumber(nextInLine([], 5));
```

`nextInLine([], 1)` 應該傳回 `1`

```js
assert(nextInLine([], 1) === 1);
```

`nextInLine([2], 1)` 應該傳回 `2`

```js
assert(nextInLine([2], 1) === 2);
```

`nextInLine([5,6,7,8,9], 1)` 應該傳回 `5`

```js
assert(nextInLine([5, 6, 7, 8, 9], 1) === 5);
```

# --seed--

## --seed-contents--

```js
function nextInLine(arr, item) {
  // Only change code below this line
  
  return item;
  // Only change code above this line
}

// Setup
let testArr = [1, 2, 3, 4, 5];

// Display code
console.log("Before: " + JSON.stringify(testArr));
console.log(nextInLine(testArr, 6));
console.log("After: " + JSON.stringify(testArr));
```

# --solutions--

```js
let testArr = [1, 2, 3, 4, 5];

function nextInLine(arr, item) {
    arr.push(item);
    return arr.shift();
}
```
