---
id: 56533eb9ac21ba0edf2244da
title: 介紹 Else 述語
challengeType: 1
forumTopicId: 18207
dashedName: introducing-else-statements
---

# --description--

當 `if` 述語的條件為真時，緊接著的程式碼區塊會被執行。當該條件為假時會怎樣？通常不會有任何動作。使用 `else` 述語時，可以執行另一個程式碼區塊。

```js
if (num > 10) {
  return "Bigger than 10";
} else {
  return "10 or Less";
}
```

# --instructions--

將多個 `if` 述語合併為單一的 `if/else` 述語。

# --hints--

你在編輯器中應該只使用一個 `if` 述語。

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

你應該使用 `else` 述語

```js
assert(/else/g.test(__helpers.removeJSComments(code)));
```

`testElse(4)` 應該傳回字串 `5 or Smaller`

```js
assert(testElse(4) === '5 or Smaller');
```

`testElse(5)` 應該傳回字串 `5 or Smaller`

```js
assert(testElse(5) === '5 or Smaller');
```

`testElse(6)` 應該傳回字串 `Bigger than 5`

```js
assert(testElse(6) === 'Bigger than 5');
```

`testElse(10)` 應該傳回字串 `Bigger than 5`

```js
assert(testElse(10) === 'Bigger than 5');
```

你不應該更改指定註解上方或下方的程式碼。

```js
assert(/let result = "";/.test(__helpers.removeJSComments(code)) && /return result;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function testElse(val) {
  let result = "";
  // Only change code below this line

  if (val > 5) {
    result = "Bigger than 5";
  }

  if (val <= 5) {
    result = "5 or Smaller";
  }

  // Only change code above this line
  return result;
}

testElse(4);
```

# --solutions--

```js
function testElse(val) {
  let result = "";
  if(val > 5) {
    result = "Bigger than 5";
  } else {
    result = "5 or Smaller";
  }
  return result;
}
```
