---
id: 56533eb9ac21ba0edf2244b4
title: 使用單引號引用字串
challengeType: 1
forumTopicId: 18260
dashedName: quoting-strings-with-single-quotes
---

# --description--

<dfn>字串</dfn>值在 JavaScript 中可以用單引號或雙引號來撰寫，只要你以相同型別的引號開始和結束即可。與其他一些程式設計語言不同，單引號和雙引號在 JavaScript 中的作用相同。

```js
const doubleQuoteStr = "This is a string"; 
const singleQuoteStr = 'This is also a string';
```

你可能想使用一種引號型別而非另一種的原因，是如果你想在字串中同時使用兩者。這可能發生在你想將對話存檔為字串，且對話中包含引號時。另一個用途是將帶有各種屬性的 `<a>` 頁籤以引號包住，全部放在字串中。

```js
const conversation = 'Finn exclaims to Jake, "Algebraic!"';
```

但是，如果你需要在其中使用最外層的引號，這就會成為問題。請記住，字串在開頭和結尾使用相同種類的引號。但如果你在中間某處有相同的引號，字串就會提早結束並發出錯誤。

```js
const goodStr = 'Jake asks Finn, "Hey, let\'s go on an adventure?"'; 
const badStr = 'Finn responds, "Let's go!"';
```

這裡 `badStr` 會發出錯誤。

在上述的 <dfn>goodStr</dfn> 中，你可以透過使用反斜線 `\` 作為跳脫字元來安全地使用雙引號和單引號。

**注意：** 反斜線 `\` 不應與正斜線 `/` 混淆。它們的功能不同。

# --instructions--

將提供的字串更改為以單引號開頭和結尾且不含跳脫字元的字串。

目前，字串中的 `<a>` 頁籤在所有地方都使用雙引號。你需要將外層引號改為單引號，這樣你才能移除跳脫字元。

# --hints--

你應該移除所有反斜線（`\`）。

```js
assert(
  !/\\/g.test(__helpers.removeJSComments(code)) &&
    myStr.match(
      '\\s*<a href\\s*=\\s*"http://www.example.com"\\s*target\\s*=\\s*"_blank">\\s*Link\\s*</a>\\s*'
    )
);
```

你應該有兩個單引號 `'` 和四個雙引號 `"`。

```js
assert(__helpers.removeJSComments(code).match(/"/g).length === 4 && __helpers.removeJSComments(code).match(/'/g).length === 2);
```

# --seed--

## --seed-contents--

```js
const myStr = "<a href=\"http://www.example.com\" target=\"_blank\">Link</a>";
```

# --solutions--

```js
const myStr = '<a href="http://www.example.com" target="_blank">Link</a>';
```
