---
id: 56533eb9ac21ba0edf2244b5
title: 在字串中轉義字面引號
challengeType: 1
forumTopicId: 17568
dashedName: escaping-literal-quotes-in-strings
---

# --description--

當你定義字串時，必須以單引號或雙引號開始和結束。當你需要在字串中使用字面上的引號：`"` 或 `'`，會發生什麼情況？

在 JavaScript 中，你可以透過在引號前放置一個 <dfn>backslash</dfn>（`\`）來 <dfn>escape</dfn> 該引號，使其不被視為字串結尾的引號。

```js
const sampleStr = "Alan said, \"Peter is learning JavaScript\".";
```

這表示給 JavaScript 的信號，告訴它接下來的引號不是字串的結尾，而應該出現在字串內部。所以如果你將這個列印到主控台，你會得到：

```js
Alan said, "Peter is learning JavaScript".
```

# --instructions--

使用 <dfn>反斜線</dfn> 為 `myStr` 變數指定字串，這樣如果你要將它列印到主控台，你會看到：

```js
I am a "double quoted" string inside "double quotes".
```

# --hints--

你應該使用兩個雙引號（`"`）和四個跳脫雙引號（`\"`）。

```js
assert(__helpers.removeJSComments(code).match(/\\"/g).length === 4 && __helpers.removeJSComments(code).match(/[^\\]"/g).length === 2);
```

變數 `myStr` 應該包含字串：`I am a "double quoted" string inside "double quotes".`

```js
assert(/I am a "double quoted" string inside "double quotes(\."|"\.)$/.test(myStr));
```

# --seed--

## --after-user-code--

```js
(function(){
  if(typeof myStr === 'string') {
    console.log("myStr = \"" + myStr + "\"");
  } else {
    console.log("myStr is undefined");
  }
})();
```

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "I am a \"double quoted\" string inside \"double quotes\".";
```
