---
id: 56533eb9ac21ba0edf2244ad
title: 使用 JavaScript 將數字遞減
challengeType: 1
forumTopicId: 17558
dashedName: decrement-a-number-with-javascript
---

# --description--

你可以輕鬆地使用 `--` 運算子將變數 <dfn>遞減</dfn> 或減少一。

```js
i--;
```

是等價的

```js
i = i - 1;
```

**注意：**整行變成 `i--;`，不需要等號。

# --instructions--

將程式碼更改為對 `myVar` 使用 `--` 運算子。

# --hints--

`myVar` 應該相等於 `10`。

```js
assert(myVar === 10);
```

應該更改 `myVar = myVar - 1;`。

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*myVar\s*[-]\s*1.*?;?/));
```

你不應該將 `myVar` 指定為 `10`。

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*10.*?;?/));
```

你應該在 `myVar` 上使用 `--` 運算子。

```js
assert(/[-]{2}\s*myVar|myVar\s*[-]{2}/.test(__helpers.removeJSComments(code)));
```

你不應該更改指定註解以上的程式碼。

```js
assert(/let myVar = 11;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'myVar = ' + z;})(myVar);
```

## --seed-contents--

```js
let myVar = 11;

// Only change code below this line
myVar = myVar - 1;
```

# --solutions--

```js
let myVar = 11;
myVar--;
```
