---
id: 56533eb9ac21ba0edf2244c9
title: 使用變數存取物件屬性
challengeType: 1
forumTopicId: 16165
dashedName: accessing-object-properties-with-variables
---

# --description--

方括號表示法在物件上的另一種用法是存取作為變數值所儲存的屬性。這在遍歷物件的屬性或存取查找表時非常有用。

以下是使用變數來存取屬性的範例：

```js
const dogs = {
  Fido: "Mutt",
  Hunter: "Doberman",
  Snoopie: "Beagle"
};

const myDog = "Hunter";
const myBreed = dogs[myDog];
console.log(myBreed);
```

字串 `Doberman` 將會顯示在主控台。

請注意，當使用變數來存取屬性時，我們*不*會在變數名稱周圍加上引號，因為我們使用的是變數的*值*，而不是*名稱*。

# --instructions--

將 `playerNumber` 變數指定為 `16`。然後，使用該變數查找玩家名稱並將其指定給 `player``。

# --hints--

`playerNumber` 應該是一個數字

```js
assert(typeof playerNumber === 'number');
```

變數 `player` 應該是一個字串

```js
assert(typeof player === 'string');
```

`player` 的值應該是字串 `Montana`

```js
assert(player === 'Montana');
```

你應該使用括號表示法來存取 `testObj`。

```js
assert(/testObj\s*?\[.*?\]/.test(__helpers.removeJSComments(code)));
```

你不應該直接將值 `Montana` 指定給變數 `player`。

```js
assert(!__helpers.removeJSComments(code).match(/player\s*=\s*"|\'\s*Montana\s*"|\'\s*;/gi));
```

你應該在你的括號表示法中使用變數 `playerNumber`。

```js
assert(/testObj\s*?\[\s*playerNumber\s*\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
if(typeof player !== "undefined"){(function(v){return v;})(player);}
```

## --seed-contents--

```js
// Setup
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};

// Only change code below this line
const playerNumber = 42;  // Change this line
const player = testObj;   // Change this line
```

# --solutions--

```js
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};
const playerNumber = 16;
const player = testObj[playerNumber];
```
