---
id: 56533eb9ac21ba0edf2244dc
title: 串鏈 If Else 述語
challengeType: 1
forumTopicId: 16772
dashedName: chaining-if-else-statements
---

# --description--

`if/else` 述語可以串鏈在一起以實現複雜邏輯。以下是多個串鏈的 `if` / `else if` 述語的 <dfn>pseudocode</dfn>：

```js
if (condition1) {
  statement1
} else if (condition2) {
  statement2
} else if (condition3) {
  statement3
. . .
} else {
  statementN
}
```

# --instructions--

撰寫串鏈的 `if`/`else if` 述語以符合以下條件：

`num < 5` - 傳回 `Tiny``num < 10` - 傳回 `Small``num < 15` - 傳回 `Medium``num < 20` - 傳回 `Large``num >= 20` - 傳回 `Huge`

# --hints--

你應該至少有四個 `else` 述語

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 3);
```

你應該至少有四個 `if` 述語。

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 3);
```

你應該至少有一個 `return` 述語

```js
assert(__helpers.removeJSComments(code).match(/return/g).length >= 1);
```

`testSize(0)` 應該傳回字串 `Tiny`

```js
assert(testSize(0) === 'Tiny');
```

`testSize(4)` 應該傳回字串 `Tiny`

```js
assert(testSize(4) === 'Tiny');
```

`testSize(5)` 應該傳回字串 `Small`

```js
assert(testSize(5) === 'Small');
```

`testSize(8)` 應該傳回字串 `Small`

```js
assert(testSize(8) === 'Small');
```

`testSize(10)` 應該傳回字串 `Medium`

```js
assert(testSize(10) === 'Medium');
```

`testSize(14)` 應該傳回字串 `Medium`

```js
assert(testSize(14) === 'Medium');
```

`testSize(15)` 應該傳回字串 `Large`

```js
assert(testSize(15) === 'Large');
```

`testSize(17)` 應該傳回字串 `Large`

```js
assert(testSize(17) === 'Large');
```

`testSize(20)` 應該傳回字串 `Huge`

```js
assert(testSize(20) === 'Huge');
```

`testSize(25)` 應該傳回字串 `Huge`

```js
assert(testSize(25) === 'Huge');
```

# --seed--

## --seed-contents--

```js
function testSize(num) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

testSize(7);
```

# --solutions--

```js
function testSize(num) {
  if (num < 5) {
    return "Tiny";
  } else if (num < 10) {
    return "Small";
  } else if (num < 15) {
    return "Medium";
  } else if (num < 20) {
    return "Large";
  } else {
    return "Huge";
  }
}
```
