---
id: 56533eb9ac21ba0edf2244c4
title: 函式的提前傳回樣式
challengeType: 1
forumTopicId: 18272
dashedName: return-early-pattern-for-functions
---

# --description--

當執行到 `return` 述語時，當前函式的執行會停止，並且控制元件會返回到呼叫位置。

**範例**

```js
function myFun() {
  console.log("Hello");
  return "World";
  console.log("byebye")
}
myFun();
```

上述會在主控台顯示字串 `Hello`，並傳回字串 `World`。字串 `byebye` 永遠不會在主控台顯示，因為函式會在 `return` 述語處離開。

# --instructions--

修改函式（程式）`abTest`，使得如果 `a` 或 `b` 小於 `0`，函式（程式）將立即以值 `undefined` 離開、退離。

**提示**請記住 <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-javascript/understanding-uninitialized-variables" target="_blank" rel="noopener noreferrer nofollow"><code>undefined</code> 是一個關鍵字</a>，而不是字串。

# --hints--

`abTest(2, 2)` 應該傳回一個數字

```js
assert(typeof abTest(2, 2) === 'number');
```

`abTest(2, 2)` 應該傳回 `8`

```js
assert(abTest(2, 2) === 8);
```

`abTest(-2, 2)` 應該傳回 `undefined`

```js
assert(abTest(-2, 2) === undefined);
```

`abTest(2, -2)` 應該傳回 `undefined`

```js
assert(abTest(2, -2) === undefined);
```

`abTest(2, 8)` 應該傳回 `18`

```js
assert(abTest(2, 8) === 18);
```

`abTest(3, 3)` 應該傳回 `12`

```js
assert(abTest(3, 3) === 12);
```

`abTest(0, 0)` 應該傳回 `0`

```js
assert(abTest(0, 0) === 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function abTest(a, b) {
  // Only change code below this line



  // Only change code above this line

  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}

abTest(2,2);
```

# --solutions--

```js
function abTest(a, b) {
  if(a < 0 || b < 0) {
    return undefined;
  }
  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}
```
