---
id: 56533eb9ac21ba0edf2244b2
title: 使用增強除法的複合指派
challengeType: 1
forumTopicId: 16659
dashedName: compound-assignment-with-augmented-division
---

# --description--

`/=` 運算子將一個變數除以另一個數字。

```js
myVar = myVar / 5;
```

將 `myVar` 除以 `5`。這可以重寫為：

```js
myVar /= 5;
```

# --instructions--

將 `a`、`b` 和 `c` 的指派改為使用 `/=` 運算子。

# --hints--

`a` 應該相等於 `4`。

```js
assert(a === 4);
```

`b` 應該相等於 `27`。

```js
assert(b === 27);
```

`c` 應該相等於 `3`。

```js
assert(c === 3);
```

你應該對每個變數使用 `/=` 運算子。

```js
assert(__helpers.removeJSComments(code).match(/\/=/g).length === 3);
```

你不應該修改指定註解上方的程式碼。

```js
assert(
  /let a = 48;/.test(__helpers.removeJSComments(code)) &&
    /let b = 108;/.test(__helpers.removeJSComments(code)) &&
    /let c = 33;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --after-user-code--

```js
(function(a,b,c){ return "a = " + a + ", b = " + b + ", c = " + c; })(a,b,c);
```

## --seed-contents--

```js
let a = 48;
let b = 108;
let c = 33;

// Only change code below this line
a = a / 12;
b = b / 4;
c = c / 11;
```

# --solutions--

```js
let a = 48;
let b = 108;
let c = 33;

a /= 12;
b /= 4;
c /= 11;
```
