---
id: 56533eb9ac21ba0edf2244d4
title: 與大於運算子的比較
challengeType: 1
forumTopicId: 16786
dashedName: comparison-with-the-greater-than-operator
---

# --description--

大於運算子（`>`）比較兩個數字的值。如果左邊的數字大於右邊的數字，它會傳回 `true`。否則，它會傳回 `false`。

像等號運算子一樣，大於運算子在比較時會轉換值的資料型別。

**範例**

```js
5   >  3  // true
7   > '3' // true
2   >  3  // false
'1' >  9  // false
```

# --instructions--

將大於運算子添加到被顯示的行，以便述語能夠正確傳回。

# --hints--

`testGreaterThan(0)` 應該傳回字串 `10 or Under`

```js
assert(testGreaterThan(0) === '10 or Under');
```

`testGreaterThan(10)` 應該傳回字串 `10 or Under`

```js
assert(testGreaterThan(10) === '10 or Under');
```

`testGreaterThan(11)` 應該傳回字串 `Over 10`

```js
assert(testGreaterThan(11) === 'Over 10');
```

`testGreaterThan(99)` 應該傳回字串 `Over 10`

```js
assert(testGreaterThan(99) === 'Over 10');
```

`testGreaterThan(100)` 應該傳回字串 `Over 10`

```js
assert(testGreaterThan(100) === 'Over 10');
```

`testGreaterThan(101)` 應該傳回字串 `Over 100`

```js
assert(testGreaterThan(101) === 'Over 100');
```

`testGreaterThan(150)` 應該傳回字串 `Over 100`

```js
assert(testGreaterThan(150) === 'Over 100');
```

你應該至少使用兩次 `>` 運算子

```js
assert(__helpers.removeJSComments(code).match(/val\s*>\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterThan(val) {
  if (val) {  // Change this line
    return "Over 100";
  }

  if (val) {  // Change this line
    return "Over 10";
  }

  return "10 or Under";
}

testGreaterThan(10);
```

# --solutions--

```js
function testGreaterThan(val) {
  if (val > 100) {  // Change this line
    return "Over 100";
  }
  if (val > 10) {  // Change this line
    return "Over 10";
  }
  return "10 or Under";
}
```
