---
id: 56533eb9ac21ba0edf2244a9
title: 使用指定運算子初始化變數
challengeType: 1
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

在宣告變數的同一行中，通常會將變數 <dfn>初始化</dfn> 為初始值。

```js
var myVar = 0;
```

創建一個名為 `myVar` 的新變數，並將其指定初始值為 `0`。

# --instructions--

使用 `var` 定義一個變數 `a` 並將其初始化為值 `9`。

# --hints--

你應該將 `a` 初始化為 `9`。

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
