---
id: 56533eb9ac21ba0edf2244db
title: 介紹 Else If 述語
challengeType: 1
forumTopicId: 18206
dashedName: introducing-else-if-statements
---

# --description--

如果你有多個需要處理的條件，你可以將多個 `if` 述語與 `else if` 述語串鏈起來。

```js
if (num > 15) {
  return "Bigger than 15";
} else if (num < 5) {
  return "Smaller than 5";
} else {
  return "Between 5 and 15";
}
```

# --instructions--

將邏輯轉換為使用 `else if` 述語。

# --hints--

你應該至少有兩個 `else` 述語

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 1);
```

你應該至少有兩個 `if` 述語

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 1);
```

你應該為每個 `if else` 程式碼區塊加上閉合和開啟的大括號。

```js
assert(
  __helpers.removeJSComments(code).match(
    /if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s+if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s*\{[\s\S]+\s*\}/
  )
);
```

`testElseIf(0)` 應該傳回字串 `Smaller than 5`

```js
assert(testElseIf(0) === 'Smaller than 5');
```

`testElseIf(5)` 應該傳回字串 `Between 5 and 10`

```js
assert(testElseIf(5) === 'Between 5 and 10');
```

`testElseIf(7)` 應該傳回字串 `Between 5 and 10`

```js
assert(testElseIf(7) === 'Between 5 and 10');
```

`testElseIf(10)` 應該傳回字串 `Between 5 and 10`

```js
assert(testElseIf(10) === 'Between 5 and 10');
```

`testElseIf(12)` 應該傳回字串 `Greater than 10`

```js
assert(testElseIf(12) === 'Greater than 10');
```

# --seed--

## --seed-contents--

```js
function testElseIf(val) {
  if (val > 10) {
    return "Greater than 10";
  }

  if (val < 5) {
    return "Smaller than 5";
  }

  return "Between 5 and 10";
}

testElseIf(7);
```

# --solutions--

```js
function testElseIf(val) {
  if(val > 10) {
    return "Greater than 10";
  } else if(val < 5) {
    return "Smaller than 5";
  } else {
    return "Between 5 and 10";
  }
}
```
