---
id: 56533eb9ac21ba0edf2244ca
title: 使用物件進行查找
challengeType: 1
forumTopicId: 18373
dashedName: using-objects-for-lookups
---

# --description--

物件可以被視為鍵/值儲存，就像字典一樣。如果你有表格式資料，你可以使用物件來查找值，而不是使用 `switch` 述語或 `if/else` 串鏈。當你知道你的輸入資料限制在特定範圍內時，這種方法最為有用。

這是一個文章物件的範例：

```js
const article = {
  "title": "How to create objects in JavaScript",
  "link": "https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/",
  "author": "Kaashan Hussain",
  "language": "JavaScript",
  "tags": "TECHNOLOGY",
  "createdAt": "NOVEMBER 28, 2018"
};

const articleAuthor = article["author"];
const articleLink = article["link"];

const value = "title";
const valueLookup = article[value];
```

`articleAuthor` 是字串 `Kaashan Hussain`，`articleLink` 是字串 `https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/`，而 `valueLookup` 是字串 `How to create objects in JavaScript`。

# --instructions--

將 switch 述語轉換成名為 `lookup` 的物件。使用它來查找 `val` 並將相關的字串指定到 `result` 變數。

# --hints--

`phoneticLookup("alpha")` 應該相等於字串 `Adams`

```js
assert(phoneticLookup('alpha') === 'Adams');
```

`phoneticLookup("bravo")` 應該相等於字串 `Boston`

```js
assert(phoneticLookup('bravo') === 'Boston');
```

`phoneticLookup("charlie")` 應該相等於字串 `Chicago`

```js
assert(phoneticLookup('charlie') === 'Chicago');
```

`phoneticLookup("delta")` 應該相等於字串 `Denver`

```js
assert(phoneticLookup('delta') === 'Denver');
```

`phoneticLookup("echo")` 應該相等於字串 `Easy`

```js
assert(phoneticLookup('echo') === 'Easy');
```

`phoneticLookup("foxtrot")` 應該相等於字串 `Frank`

```js
assert(phoneticLookup('foxtrot') === 'Frank');
```

`phoneticLookup("")` 應該相等於 `undefined`

```js
assert(typeof phoneticLookup('') === 'undefined');
```

你不應該修改 `return` 述語

```js
assert(__helpers.removeJSComments(code).match(/return\sresult;/));
```

你不應該使用 `case`、`switch` 或 `if` 述語。

```js
assert(
  !/case|switch|if/g.test(__helpers.removeJSComments(code).replace(/([/]{2}.*)|([/][*][^/*]*[*][/])/g, ''))
);
```

# --seed--

## --seed-contents--

```js
// Setup
function phoneticLookup(val) {
  let result = "";

  // Only change code below this line
  switch(val) {
    case "alpha":
      result = "Adams";
      break;
    case "bravo":
      result = "Boston";
      break;
    case "charlie":
      result = "Chicago";
      break;
    case "delta":
      result = "Denver";
      break;
    case "echo":
      result = "Easy";
      break;
    case "foxtrot":
      result = "Frank";
  }

  // Only change code above this line
  return result;
}

phoneticLookup("charlie");
```

# --solutions--

```js
function phoneticLookup(val) {
  let result = "";

  const lookup = {
    alpha: "Adams",
    bravo: "Boston",
    charlie: "Chicago",
    delta: "Denver",
    echo: "Easy",
    foxtrot: "Frank"
  };

  result = lookup[val];

  return result;
}
```
