---
id: 56533eb9ac21ba0edf2244c2
title: 從函式使用 Return 傳回一個值
challengeType: 1
forumTopicId: 18271
dashedName: return-a-value-from-a-function-with-return
---

# --description--

我們可以透過 <dfn>引數</dfn> 將值傳入函式。你可以使用 `return` 述語將值傳回函式外。

**範例**

```js
function plusThree(num) {
  return num + 3;
}

const answer = plusThree(5);
```

`answer` 的值是 `8`。

`plusThree` 接受一個 `num` 的 <dfn>引數</dfn> 並傳回一個與 `num + 3` 相等的值。

# --instructions--

創建一個函式（程式）`timesFive`，接受一個引數，將它乘以 `5`，並傳回新的值。

# --hints--

`timesFive` 應該是一個函式（程式）。

```js
assert(typeof timesFive === 'function');
```

`timesFive(5)` 應該傳回 `25`

```js
assert(timesFive(5) === 25);
```

`timesFive(2)` 應該傳回 `10`

```js
assert(timesFive(2) === 10);
```

`timesFive(0)` 應該傳回 `0`

```js
assert(timesFive(0) === 0);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function timesFive(num) {
  return num * 5;
}
timesFive(10);
```
