---
id: 56533eb9ac21ba0edf2244d1
title: 使用嚴格等號運算子進行比較
challengeType: 1
forumTopicId: 16790
dashedName: comparison-with-the-strict-equality-operator
---

# --description--

嚴格相等 (`===`) 是等號運算子 (`==`) 的對應運算子。然而，不同於嘗試將兩個被比較的值轉換為共同型別的等號運算子，嚴格相等運算子不會執行型別轉換。

如果被比較的值具有不同的型別，則它們被視為不相等，且嚴格等號運算子會傳回 false。

**範例**

```js
3 ===  3  // true
3 === '3' // false
```

在第二個範例中，`3` 是 `Number` 型別，而 `'3'` 是 `String` 型別。

# --instructions--

在 `if` 述語中使用嚴格等號運算子，這樣當 `val` 嚴格相等於 `7` 時，函式將傳回字串 `Equal`。

# --hints--

`testStrict(10)` 應該傳回字串 `Not Equal`

```js
assert(testStrict(10) === 'Not Equal');
```

`testStrict(7)` 應該傳回字串 `Equal`

```js
assert(testStrict(7) === 'Equal');
```

`testStrict("7")` 應該傳回字串 `Not Equal`

```js
assert(testStrict('7') === 'Not Equal');
```

你應該使用 `===` 運算子

```js
assert(__helpers.removeJSComments(code).match(/(val\s*===\s*\d+)|(\d+\s*===\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrict(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testStrict(10);
```

# --solutions--

```js
function testStrict(val) {
  if (val === 7) {
    return "Equal";
  }
  return "Not Equal";
}
```
