---
id: 56533eb9ac21ba0edf2244ab
title: 理解變數的大小寫敏感性
challengeType: 1
forumTopicId: 18334
dashedName: understanding-case-sensitivity-in-variables
---

# --description--

在 JavaScript 中，所有變數和函式名稱皆區分大小寫。這表示大小寫字母很重要。

`MYVAR` 與 `MyVar` 及 `myvar` 不同。可能會有多個名稱相同但大小寫不同的獨立變數。強烈建議為了清晰起見，你*不要*使用此語言特性。

**最佳實務**

在 JavaScript 中以 <dfn>camelCase</dfn> 撰寫變數名稱。在 <dfn>camelCase</dfn> 中，多字變數名稱的第一個字為小寫，後續每個字的第一個字母皆為大寫。

**範例：**

```js
var someVariable;
var anotherVariableName;
var thisVariableNameIsSoLong;
```

# --instructions--

修改現有的宣告式和指派，使它們的名稱使用 <dfn>camelCase</dfn>。

不要創建任何新的變數。

# --hints--

`studlyCapVar` 應該被定義並且其值為 `10`。

```js
assert(typeof studlyCapVar !== 'undefined' && studlyCapVar === 10);
```

`properCamelCase` 應該被定義並且具有字串 `A String` 的值。

```js
assert(
  typeof properCamelCase !== 'undefined' && properCamelCase === 'A String'
);
```

`titleCaseOver` 應該被定義並且其值為 `9000`。

```js
assert(typeof titleCaseOver !== 'undefined' && titleCaseOver === 9000);
```

`studlyCapVar` 在宣告和指派區段中都應使用 camelCase。

```js
assert(__helpers.removeJSComments(code).match(/studlyCapVar/g).length === 2);
```

`properCamelCase` 應該在宣告和指派區段中都使用 camelCase。

```js
assert(__helpers.removeJSComments(code).match(/properCamelCase/g).length === 2);
```

`titleCaseOver` 在宣告和指派區段中都應使用 camelCase。

```js
assert(__helpers.removeJSComments(code).match(/titleCaseOver/g).length === 2);
```

# --seed--

## --seed-contents--

```js
// Variable declarations
var StUdLyCapVaR;
var properCamelCase;
var TitleCaseOver;

// Variable assignments
STUDLYCAPVAR = 10;
PRoperCAmelCAse = "A String";
tITLEcASEoVER = 9000;
```

# --solutions--

```js
var studlyCapVar;
var properCamelCase;
var titleCaseOver;

studlyCapVar = 10;
properCamelCase = "A String";
titleCaseOver = 9000;
```
