---
id: 56533eb9ac21ba0edf2244e1
title: 巢狀 For 迴圈
challengeType: 1
forumTopicId: 18248
dashedName: nesting-for-loops
---

# --description--

如果你有一個多維陣列，你可以使用與先前路徑點相同的邏輯來迴圈遍歷陣列及任何子陣列。以下是一個範例：

```js
const arr = [
  [1, 2], [3, 4], [5, 6]
];

for (let i = 0; i < arr.length; i++) {
  for (let j = 0; j < arr[i].length; j++) {
    console.log(arr[i][j]);
  }
}
```

這會一次輸出 `arr` 中的每個子元素。請注意，對於內層迴圈，我們正在檢查 `arr[i]` 的 `.length`，因為 `arr[i]` 本身就是一個陣列。

# --instructions--

修改函式（程式）`multiplyAll`，使其傳回 `arr` 中子陣列所有數字的乘積。

# --hints--

`multiplyAll([[1], [2], [3]])` 應該傳回 `6`

```js
assert(multiplyAll([[1], [2], [3]]) === 6);
```

`multiplyAll([[1, 2], [3, 4], [5, 6, 7]])` 應該傳回 `5040`

```js
assert(
  multiplyAll([
    [1, 2],
    [3, 4],
    [5, 6, 7]
  ]) === 5040
);
```

`multiplyAll([[5, 1], [0.2, 4, 0.5], [3, 9]])` 應該傳回 `54`

```js
assert(
  multiplyAll([
    [5, 1],
    [0.2, 4, 0.5],
    [3, 9]
  ]) === 54
);
```

# --seed--

## --seed-contents--

```js
function multiplyAll(arr) {
  let product = 1;
  // Only change code below this line

  // Only change code above this line
  return product;
}

multiplyAll([[1, 2], [3, 4], [5, 6, 7]]);
```

# --solutions--

```js
function multiplyAll(arr) {
  let product = 1;
  for (let i = 0; i < arr.length; i++) {
    for (let j = 0; j < arr[i].length; j++) {
      product *= arr[i][j];
    }
  }
  return product;
}
```
