---
id: 56533eb9ac21ba0edf2244bf
title: 區域的範圍與函式
challengeType: 1
forumTopicId: 18227
dashedName: local-scope-and-functions
---

# --description--

在函式中宣告的變數，以及函式參數，具有<dfn>區域的</dfn>範圍。這表示它們僅在該函式內可見。

這裡有一個函式（程式）`myTest`，其中有一個名為 `loc` 的區域的變數。

```js
function myTest() {
  const loc = "foo";
  console.log(loc);
}

myTest();
console.log(loc);
```

`myTest()` 函式的呼叫會在主控台顯示字串 `foo`。`console.log(loc)` 這一行（在 `myTest` 函式外）會發出錯誤，因為 `loc` 在函式外未定義。

# --instructions--

編輯器有兩個 `console.log` 幫助你了解發生了什麼。當你撰寫程式碼時，檢查主控台以查看它如何變化。在 `myLocalScope` 內宣告一個區域的變數 `myVar` 並執行測試。

**注意：** 主控台仍會顯示 `ReferenceError: myVar is not defined`，但這不會導致測試失敗。

# --hints--

程式碼中不應包含全域的 `myVar` 變數。

```js
function declared() {
  myVar;
}

assert.throws(declared, ReferenceError);
```

你應該為你的區域的 `myVar` 變數添加。

```js
assert(
  /functionmyLocalScope\(\)\{.*(var|let|const)myVar[\s\S]*}/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

# --seed--

## --seed-contents--

```js
function myLocalScope() {
  // Only change code below this line

  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```

# --solutions--

```js
function myLocalScope() {
  // Only change code below this line
  let myVar;
  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```
