---
id: 56533eb9ac21ba0edf2244ba
title: 了解字串不變性
challengeType: 1
forumTopicId: 18331
dashedName: understand-string-immutability
---

# --description--

在 JavaScript 中，`String` 值是 <dfn>不可變的</dfn>，這表示它們一旦創建後就無法被更改。

例如，以下程式碼會產生錯誤，因為字串 `Bob` 中的字母 `B` 無法更改為字母 `J`：

```js
let myStr = "Bob";
myStr[0] = "J";
```

請注意，這並*不*代表 `myStr` 無法被重新指定。唯一改變 `myStr` 的方法是用新值指定它，像這樣：

```js
let myStr = "Bob";
myStr = "Job";
```

# --instructions--

修正對 `myStr` 的指派，使其包含 `Hello World` 的字串值，方法如上例所示。

# --hints--

`myStr` 應該有字串 `Hello World` 的值。

```js
assert(myStr === 'Hello World');
```

你不應該更改指定註解上方的程式碼。

```js
assert(/myStr = "Jello World"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function(v){return "myStr = " + v;})(myStr);
```

## --seed-contents--

```js
// Setup
let myStr = "Jello World";

// Only change code below this line
myStr[0] = "H"; // Change this line
// Only change code above this line
```

# --solutions--

```js
let myStr = "Jello World";
myStr = "Hello World";
```
