---
id: 56533eb9ac21ba0edf2244aa
title: 理解未初始化的變數
challengeType: 1
forumTopicId: 18335
dashedName: understanding-uninitialized-variables
---

# --description--

當 JavaScript 變數被宣告時，它們的初始值為 `undefined`。如果你對 `undefined` 變數執行數學操作，結果將會是 `NaN`，這代表 <dfn>「不是數字」</dfn>。如果你將字串與 `undefined` 變數串接，你將得到一個 `undefined` 的 <dfn>字串</dfn>。

# --instructions--

將三個變數 `a`、`b` 和 `c` 分別初始化為 `5`、`10` 和 `"I am a"`，以避免它們為 `undefined`。

# --hints--

`a` 應該被定義並且最終值為 `6`。

```js
assert(typeof a === 'number' && a === 6);
```

`b` 應該被定義並且最終值為 `15`。

```js
assert(typeof b === 'number' && b === 15);
```

`c` 不應包含 `undefined`，並且最終值應為字串 `I am a String!`

```js
assert(!/undefined/.test(c) && c === 'I am a String!');
```

你不應該更改指定註解以下的程式碼。

```js
assert(
  /a = a \+ 1;/.test(__helpers.removeJSComments(code)) &&
    /b = b \+ 5;/.test(__helpers.removeJSComments(code)) &&
    /c = c \+ " String!";/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --after-user-code--

```js
(function(a,b,c){ return "a = " + a + ", b = " + b + ", c = '" + c + "'"; })(a,b,c);
```

## --seed-contents--

```js
// Only change code below this line
var a;
var b;
var c;
// Only change code above this line

a = a + 1;
b = b + 5;
c = c + " String!";
```

# --solutions--

```js
var a = 5;
var b = 10;
var c = "I am a";
a = a + 1;
b = b + 5;
c = c + " String!";
```
