---
id: 56533eb9ac21ba0edf2244c7
title: 使用點記法存取物件屬性
challengeType: 1
forumTopicId: 16164
dashedName: accessing-object-properties-with-dot-notation
---

# --description--

有兩種方式可以存取物件的屬性：點記法（`.`）和括號記法（`[]`），類似於陣列。

當你事先知道你想要存取的屬性名稱時，使用點記法。

以下是使用點記法（`.`）讀取物件屬性的範例：

```js
const myObj = {
  prop1: "val1",
  prop2: "val2"
};

const prop1val = myObj.prop1;
const prop2val = myObj.prop2;
```

`prop1val` 的值會是字串 `val1`，而 `prop2val` 的值會是字串 `val2`。

# --instructions--

使用點記法讀取 `testObj` 的屬性值。將變數 `hatValue` 設為物件的屬性 `hat`，並將變數 `shirtValue` 設為物件的屬性 `shirt`。

# --hints--

`hatValue` 應該是一個字串

```js
assert(typeof hatValue === 'string');
```

`hatValue` 的值應該是字串 `ballcap`

```js
assert(hatValue === 'ballcap');
```

`shirtValue` 應該是一個字串

```js
assert(typeof shirtValue === 'string');
```

`shirtValue` 的值應該是字串 `jersey`

```js
assert(shirtValue === 'jersey');
```

你應該使用兩次點記法

```js
assert(__helpers.removeJSComments(code).match(/testObj\.\w+/g).length > 1);
```

# --seed--

## --after-user-code--

```js
(function(a,b) { return "hatValue = '" + a + "', shirtValue = '" + b + "'"; })(hatValue,shirtValue);
```

## --seed-contents--

```js
// Setup
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

// Only change code below this line
const hatValue = testObj;      // Change this line
const shirtValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

const hatValue = testObj.hat;
const shirtValue = testObj.shirt;
```
