---
id: 56533eb9ac21ba0edf2244a8
title: 使用指定運算子儲存值
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

在 JavaScript 中，你可以使用 <dfn>指派</dfn> 運算子（`=`）將值儲存在變數中。

```js
myVariable = 5;
```

這會將 `Number` 值 `5` 指定給 `myVariable`。

如果在 `=` 運算子右側有任何計算，這些計算會在將值指定給運算子左側的變數之前執行。

```js
var myVar;
myVar = 5;
```

首先，這段程式碼創建了一個名為 `myVar` 的變數。然後，程式碼將 `5` 指定給 `myVar`。現在，如果 `myVar` 再次出現在程式碼中，程式會將它視為 `5`。

# --instructions--

將值 `7` 指定給變數 `a`。

# --hints--

你不應該更改指定註解以上的程式碼。

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a` 應該有一個值為 7。

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
