---
id: 56533eb9ac21ba0edf2244c8
title: 使用括號符號存取物件屬性
challengeType: 1
forumTopicId: 16163
dashedName: accessing-object-properties-with-bracket-notation
---

# --description--

存取物件屬性的第二種方式是括號表示法（`[]`）。如果你嘗試存取的物件屬性名稱中有空格，你將需要使用括號表示法。

不過，你仍然可以在物件屬性上使用不帶空格的括號表示法。

以下是使用括號表示法讀取物件屬性的範例：

```js
const myObj = {
  "Space Name": "Kirk",
  "More Space": "Spock",
  "NoSpace": "USS Enterprise"
};

myObj["Space Name"];
myObj['More Space'];
myObj["NoSpace"];
```

`myObj["Space Name"]` 會是字串 `Kirk`，`myObj['More Space']` 會是字串 `Spock`，而 `myObj["NoSpace"]` 會是字串 `USS Enterprise`。

請注意，屬性名稱中有空格時，必須使用引號（單引號或雙引號）。

# --instructions--

使用括號表示法讀取 `testObj` 的屬性 `an entree` 和 `the drink` 的值，並分別指定給 `entreeValue` 和 `drinkValue`。

# --hints--

`entreeValue` 應該是一個字串

```js
assert(typeof entreeValue === 'string');
```

`entreeValue` 的值應該是字串 `hamburger`

```js
assert(entreeValue === 'hamburger');
```

`drinkValue` 應該是一個字串

```js
assert(typeof drinkValue === 'string');
```

`drinkValue` 的值應該是字串 `water`

```js
assert(drinkValue === 'water');
```

你應該使用兩次括號表示法

```js
assert(__helpers.removeJSComments(code).match(/testObj\s*?\[('|")[^'"]+\1\]/g).length > 1);
```

# --seed--

## --after-user-code--

```js
(function(a,b) { return "entreeValue = '" + a + "', drinkValue = '" + b + "'"; })(entreeValue,drinkValue);
```

## --seed-contents--

```js
// Setup
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};

// Only change code below this line
const entreeValue = testObj;   // Change this line
const drinkValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};
const entreeValue = testObj["an entree"];
const drinkValue = testObj['the drink'];
```
