---
id: 56533eb9ac21ba0edf2244d0
title: 使用等號運算子進行比較
challengeType: 1
forumTopicId: 16784
dashedName: comparison-with-the-equality-operator
---

# --description--

JavaScript 中有許多 <dfn>比較運算子</dfn>。所有這些運算子都會傳回布林值 `true` 或 `false`。

最基本的運算子是等號運算子 `==`。等號運算子比較兩個值，若它們等價則傳回 `true`，若不等價則傳回 `false`。請注意，相等性不同於指派（`=`），指派會將運算子右邊的值指定給左邊的變數。

```js
function equalityTest(myVal) {
  if (myVal == 10) {
    return "Equal";
  }
  return "Not Equal";
}
```

如果 `myVal` 相等於 `10`，等號運算子會傳回 `true`，因此大括號中的程式碼將會執行，且函式會傳回 `Equal`。否則，函式會傳回 `Not Equal`。為了讓 JavaScript 比較兩個不同的 <dfn>資料型別</dfn>（例如，`numbers` 和 `strings`），它必須將一種型別轉換成另一種型別。這稱為型別強制轉換。然而，一旦完成，它就可以如下比較項目：

```js
1   ==  1  // true
1   ==  2  // false
1   == '1' // true
"3" ==  3  // true
```

# --instructions--

為指出的行添加等號運算子，使函式在 `val` 等價於 `12` 時傳回字串 `Equal`。

# --hints--

`testEqual(10)` 應該傳回字串 `Not Equal`

```js
assert(testEqual(10) === 'Not Equal');
```

`testEqual(12)` 應該傳回字串 `Equal`

```js
assert(testEqual(12) === 'Equal');
```

`testEqual("12")` 應該傳回字串 `Equal`

```js
assert(testEqual('12') === 'Equal');
```

你應該使用 `==` 運算子

```js
assert(__helpers.removeJSComments(code).match(/==/g) && !__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function testEqual(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testEqual(10);
```

# --solutions--

```js
function testEqual(val) {
  if (val == 12) {
    return "Equal";
  }
  return "Not Equal";
}
```
