---
id: 56533eb9ac21ba0edf2244b9
title: 使用變數構造字串
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

有時你需要建置一個字串。透過使用串接運算子（`+`），你可以將一個或多個變數插入到你正在建置的字串中。

範例：

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

`ourStr` 的值會是字串 `Hello, our name is freeCodeCamp, how are you?`。

# --instructions--

將 `myName` 設為一個與你的名字相等的字串，並建置 `myStr`，使 `myName` 位於字串 `My name is ` 和 ` and I am well!` 之間。

# --hints--

`myName` 應該設定為長度至少 3 個字元的字串。

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

你應該使用兩個 `+` 運算子來建置包含 `myName` 的 `myStr`。

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --after-user-code--

```js
(function(){
  var output = [];
  if(typeof myName === 'string') {
    output.push('myName = "' + myName + '"');
  } else {
    output.push('myName is not a string');
  }
  if(typeof myStr === 'string') {
    output.push('myStr = "' + myStr + '"');
  } else {
    output.push('myStr is not a string');
  }
  return output.join('\n');
})();
```

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
