---
id: 56592a60ddddeae28f7aa8e1
title: 使用索引存取多維陣列
challengeType: 1
forumTopicId: 16159
dashedName: access-multi-dimensional-arrays-with-indexes
---

# --description--

將 <dfn>多維</dfn> 陣列視為 *陣列的陣列* 是一種思考方式。當你使用括號來存取你的陣列時，第一組括號是指最外層（第一層級）陣列中的項目，而每一組額外的括號則是指內部下一層級的項目。

**範例**

```js
const arr = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14]
];

const subarray = arr[3];
const nestedSubarray = arr[3][0];
const element = arr[3][0][1];
```

在此範例中，`subarray` 的值為 `[[10, 11, 12], 13, 14]`，`nestedSubarray` 的值為 `[10, 11, 12]`，而 `element` 的值為 `11`。

**注意：** 陣列名稱和方括號之間不應該有任何空格，例如 `array [0][0]`，甚至連這種 `array [0] [0]` 也不允許。雖然 JavaScript 能正確處理這種情況，但這可能會讓其他程式設計師在閱讀你的程式碼時感到困惑。

# --instructions--

使用括號表示法從 `myArray` 中選取一個元素，使 `myData` 等於 8。

# --hints--

`myData` 應該要相等於 `8`。

```js
assert(myData === 8);
```

你應該使用括號表示法從 `myArray` 讀取正確的值。

```js
assert(/myData=myArray\[2\]\[1\]/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --after-user-code--

```js
if(typeof myArray !== "undefined"){(function(){return "myData: " + myData + " myArray: " + JSON.stringify(myArray);})();}
```

## --seed-contents--

```js
const myArray = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14],
];

const myData = myArray[0][0];
```

# --solutions--

```js
const myArray = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [[10, 11, 12], 13, 14]];
const myData = myArray[2][1];
```
