---
id: 565bbe00e9cc8ac0725390f4
title: 計數牌
challengeType: 1
forumTopicId: 16809
dashedName: counting-cards
---

# --description--

在賭場遊戲 Blackjack 中，玩家可以透過記錄牌堆中剩餘的高牌和低牌的相對數量，來判斷他們在下一手牌中是否對莊家有優勢。這稱為 Card Counting。

牌堆中剩餘較多高點數牌對玩家有利。每張牌依照下表被指定一個數值。當計數為正時，玩家應該下注較高；當計數為零或負時，玩家應該下注較低。

<table><thead><tr><th>Count Change</th><th>Cards</th></tr></thead><tbody><tr><td>+1</td><td>2, 3, 4, 5, 6</td></tr><tr><td>0</td><td>7, 8, 9</td></tr><tr><td>-1</td><td>10, 'J', 'Q', 'K', 'A'</td></tr></tbody></table>

你將撰寫一個計算牌數的函式（程式）。它會接收一個 `card` 參數，該參數可以是數字或字串，並根據牌的數值（參見表格）遞增或遞減全域的 `count` 變數。函式接著會傳回一個字串，內容包含目前的計數值，若計數值為正數則附加字串 `Bet`，若計數值為零或負數則附加字串 `Hold`。目前的計數值與玩家的決策（`Bet` 或 `Hold`）之間應以一個半形空格分隔。

**範例輸出：** `-3 Hold` 或 `5 Bet`

**提示**當值為 7、8 或 9 時，切勿將 `count` 重設為 0。切勿傳回陣列。切勿在輸出中包含引號（單引號或雙引號）。

# --hints--

你的函式應該傳回 count 和文字（`Bet` 或 `Hold`），兩者之間有一個空白字元。

```js
assert(//
  (function () {
    count = 0;
    let out = cc(10);
    const hasSpace = /-?\d+ (Bet|Hold)/.test('' + out);
    return hasSpace;
  })()
);
```

牌組序列 2、3、4、5、6 應該傳回字串 `5 Bet`

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(3);
    cc(4);
    cc(5);
    var out = cc(6);
    if (out === '5 Bet') {
      return true;
    }
    return false;
  })()
);
```

牌組序列 7、8、9 應該傳回字串 `0 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(7);
    cc(8);
    var out = cc(9);
    if (out === '0 Hold') {
      return true;
    }
    return false;
  })()
);
```

牌序 10、J、Q、K、A 應該傳回字串 `-5 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(10);
    cc('J');
    cc('Q');
    cc('K');
    var out = cc('A');
    if (out === '-5 Hold') {
      return true;
    }
    return false;
  })()
);
```

牌組順序 3、7、Q、8、A 應該傳回字串 `-1 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(7);
    cc('Q');
    cc(8);
    var out = cc('A');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

牌組順序 2、J、9、2、7 應該傳回字串 `1 Bet`

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc('J');
    cc(9);
    cc(2);
    var out = cc(7);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

牌組序列 2、2、10 應該傳回字串 `1 Bet`

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(2);
    var out = cc(10);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

牌組順序 3、2、A、10、K 應該傳回字串 `-1 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(2);
    cc('A');
    cc(10);
    var out = cc('K');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

# --seed--

## --seed-contents--

```js
let count = 0;

function cc(card) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

cc(2); cc(3); cc(7); cc('K'); cc('A');
```

# --solutions--

```js
let count = 0;
function cc(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
