---
id: 5664820f61c48e80c9fa476c
title: 高爾夫程式碼
challengeType: 1
forumTopicId: 18195
dashedName: golf-code
---

# --description--

在高爾夫遊戲中，每個洞都有一個 `par`，代表高爾夫球手預期需要的平均 `strokes` 數，以將球擊入洞中完成比賽。根據你的 `strokes` 超過或低於 `par` 的程度，會有不同的暱稱。

你的函式（程式）將會接收 `par` 和 `strokes` 引數。根據此表格傳回正確的字串，該表格依優先權順序列出擊球次數；由上（最高）到下（最低）：

<table><thead><tr><th>Strokes</th><th>Return</th></tr></thead><tbody><tr><td>1</td><td>"Hole-in-one!"</td></tr><tr><td>&#x3C;= par - 2</td><td>"Eagle"</td></tr><tr><td>par - 1</td><td>"Birdie"</td></tr><tr><td>par</td><td>"Par"</td></tr><tr><td>par + 1</td><td>"Bogey"</td></tr><tr><td>par + 2</td><td>"Double Bogey"</td></tr><tr><td>>= par + 3</td><td>"Go Home!"</td></tr></tbody></table>

`par` 和 `strokes` 將永遠是數字且為正數。我們已為你的方便添加了一個包含所有名稱的陣列。

# --hints--

`golfScore(4, 1)` 應該傳回字串 `Hole-in-one!`

```js
assert(golfScore(4, 1) === 'Hole-in-one!');
```

`golfScore(4, 2)` 應該傳回字串 `Eagle`

```js
assert(golfScore(4, 2) === 'Eagle');
```

`golfScore(5, 2)` 應該傳回字串 `Eagle`

```js
assert(golfScore(5, 2) === 'Eagle');
```

`golfScore(4, 3)` 應該傳回字串 `Birdie`

```js
assert(golfScore(4, 3) === 'Birdie');
```

`golfScore(4, 4)` 應該傳回字串 `Par`

```js
assert(golfScore(4, 4) === 'Par');
```

`golfScore(1, 1)` 應該傳回字串 `Hole-in-one!`

```js
assert(golfScore(1, 1) === 'Hole-in-one!');
```

`golfScore(5, 5)` 應該傳回字串 `Par`

```js
assert(golfScore(5, 5) === 'Par');
```

`golfScore(4, 5)` 應該傳回字串 `Bogey`

```js
assert(golfScore(4, 5) === 'Bogey');
```

`golfScore(4, 6)` 應該傳回字串 `Double Bogey`

```js
assert(golfScore(4, 6) === 'Double Bogey');
```

`golfScore(4, 7)` 應該傳回字串 `Go Home!`

```js
assert(golfScore(4, 7) === 'Go Home!');
```

`golfScore(5, 9)` 應該傳回字串 `Go Home!`

```js
assert(golfScore(5, 9) === 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

function golfScore(par, strokes) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

golfScore(5, 4);
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
