---
id: 5675e877dbd60be8ad28edc6
title: 使用 For 迴圈迭代陣列
challengeType: 1
forumTopicId: 18216
dashedName: iterate-through-an-array-with-a-for-loop
---

# --description--

在 JavaScript 中，一個常見的任務是迭代陣列的內容。做到這一點的一種方法是使用 `for` 迴圈。這段程式碼會將陣列 `arr` 的每個元素輸出到主控台：

```js
const arr = [10, 9, 8, 7, 6];

for (let i = 0; i < arr.length; i++) {
   console.log(arr[i]);
}
```

請記住，陣列的索引是從零開始，這表示陣列的最後一個索引是 `length - 1`。我們這個迴圈的條件是 `i < arr.length`，當 `i` 相等於 `length` 時迴圈會停止。在這個例子中，最後一次迭代是 `i === 4`，也就是當 `i` 等於 `arr.length - 1` 時，並且會將 `6` 輸出到主控台。接著 `i` 增加到 `5`，因為 `i < arr.length` 為 `false`，迴圈因此終止。

# --instructions--

宣告並初始化一個變數 `total` 為 `0`。使用 `for` 迴圈將 `myArr` 陣列中每個元素的值加到 `total`。

# --hints--

`total` 應該被宣告並初始化為 0。

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s*?total\s*=\s*0.*?;?/));
```

`total` 應該相等 20。

```js
assert(total === 20);
```

你應該使用 `for` 迴圈來迭代 `myArr`。

```js
assert(/for\s*\(/g.test(__helpers.removeJSComments(code)) && /myArr\s*\[/g.test(__helpers.removeJSComments(code)));
```

你不應該嘗試直接將值 20 指定給 `total`。

```js
assert(!__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/total[=+-]0*[1-9]+/gm));
```

# --seed--

## --after-user-code--

```js
(function(){if(typeof total !== 'undefined') { return "total = " + total; } else { return "total is undefined";}})()
```

## --seed-contents--

```js
// Setup
const myArr = [2, 3, 4, 5, 6];

// Only change code below this line

```

# --solutions--

```js
const myArr = [2, 3, 4, 5, 6];
let total = 0;

for (let i = 0; i < myArr.length; i++) {
  total += myArr[i];
}
```
