---
id: 5679ceb97cbaa8c51670a16b
title: 從函式傳回布林值
challengeType: 1
forumTopicId: 18273
dashedName: returning-boolean-values-from-functions
---

# --description--

你可能會回想起<a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/comparison-with-the-equality-operator" target="_blank" rel="noopener noreferrer nofollow">使用等號運算子進行比較</a>，所有比較運算子都會傳回布林值 `true` 或 `false`。

有時候人們會使用 `if/else` 述語來進行比較，就像這樣：

```js
function isEqual(a, b) {
  if (a === b) {
    return true;
  } else {
    return false;
  }
}
```

但是有更好的方法可以做到這一點。由於 `===` 傳回 `true` 或 `false`，我們可以傳回比較的結果：

```js
function isEqual(a, b) {
  return a === b;
}
```

# --instructions--

修正函式（程式）`isLess`，以移除 `if/else` 述語。

# --hints--

`isLess(10, 15)` 應該傳回 `true`

```js
assert(isLess(10, 15) === true);
```

`isLess(15, 10)` 應該傳回 `false`

```js
assert(isLess(15, 10) === false);
```

你不應該使用任何 `if` 或 `else` 述語。

```js
assert(!/if|else/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function isLess(a, b) {
  // Only change code below this line
  if (a < b) {
    return true;
  } else {
    return false;
  }
  // Only change code above this line
}

isLess(10, 15);
```

# --solutions--

```js
function isLess(a, b) {
  return a < b;
}
```
