---
id: 567af2437cbaa8c51670a16c
title: 測試物件的屬性
challengeType: 1
forumTopicId: 18324
dashedName: testing-objects-for-properties
---

# --description--

要檢查給定物件上的屬性是否存在，你可以使用 `.hasOwnProperty()` 方法。`someObject.hasOwnProperty(someProperty)` 會根據該屬性是否存在於物件上傳回 `true` 或 `false`。

**範例**

```js
function checkForProperty(object, property) {
  return object.hasOwnProperty(property);
}

checkForProperty({ top: 'hat', bottom: 'pants' }, 'top'); // true
checkForProperty({ top: 'hat', bottom: 'pants' }, 'middle'); // false
```

第一次 `checkForProperty` 函式呼叫傳回 `true`，而第二次傳回 `false`。

# --instructions--

修改函式 `checkObj`，以測試傳遞給函式參數 `obj` 的物件是否包含傳遞給函式參數 `checkProp` 的特定屬性。如果在 `obj` 上找到傳遞給 `checkProp` 的屬性，則傳回該屬性的值。如果沒有，則傳回 `Not Found`。

# --hints--

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "gift")` 應該傳回字串 `pony`。

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'gift') === 'pony'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "pet")` 應該傳回字串 `kitten`。

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'pet') === 'kitten'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "house")` 應該傳回字串 `Not Found`。

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'house') ===
    'Not Found'
);
```

`checkObj({city: "Seattle"}, "city")` 應該傳回字串 `Seattle`。

```js
assert(checkObj({ city: 'Seattle' }, 'city') === 'Seattle');
```

`checkObj({city: "Seattle"}, "district")` 應該傳回字串 `Not Found`。

```js
assert(checkObj({ city: 'Seattle' }, 'district') === 'Not Found');
```

`checkObj({pet: "kitten", bed: "sleigh"}, "gift")` 應該傳回字串 `Not Found`。

```js
assert(checkObj({ pet: 'kitten', bed: 'sleigh' }, 'gift') === 'Not Found');
```

# --seed--

## --seed-contents--

```js
function checkObj(obj, checkProp) {
  // Only change code below this line
  return "Change Me!";
  // Only change code above this line
}
```

# --solutions--

```js
function checkObj(obj, checkProp) {
  if(obj.hasOwnProperty(checkProp)) {
    return obj[checkProp];
  } else {
    return "Not Found";
  }
}
```
