---
id: 5690307fddb111c6084545d7
title: If Else 陳述式中的邏輯順序
challengeType: 1
forumTopicId: 18228
dashedName: logical-order-in-if-else-statements
---

# --description--

在 `if`、`else if` 述語中，順序很重要。

函式（程式）是從上到下執行的，所以你要小心哪個述語先出現。

以這兩個函式（程式）為例。

這是第一個：

```js
function foo(x) {
  if (x < 1) {
    return "Less than one";
  } else if (x < 2) {
    return "Less than two";
  } else {
    return "Greater than or equal to two";
  }
}
```

第二個只是交換述語的順序：

```js
function bar(x) {
  if (x < 2) {
    return "Less than two";
  } else if (x < 1) {
    return "Less than one";
  } else {
    return "Greater than or equal to two";
  }
}
```

雖然這兩個函式看起來幾乎相同，但如果我們傳入一個數字，兩者會得到不同的輸出。

```js
foo(0)
bar(0)
```

`foo(0)` 會傳回字串 `Less than one`，而 `bar(0)` 會傳回字串 `Less than two`。

# --instructions--

改變函式中邏輯的順序，使其在所有情況下都能傳回正確的述語。

# --hints--

`orderMyLogic(4)` 應該傳回字串 `Less than 5`

```js
assert(orderMyLogic(4) === 'Less than 5');
```

`orderMyLogic(6)` 應該傳回字串 `Less than 10`

```js
assert(orderMyLogic(6) === 'Less than 10');
```

`orderMyLogic(11)` 應該傳回字串 `Greater than or equal to 10`

```js
assert(orderMyLogic(11) === 'Greater than or equal to 10');
```

# --seed--

## --seed-contents--

```js
function orderMyLogic(val) {
  if (val < 10) {
    return "Less than 10";
  } else if (val < 5) {
    return "Less than 5";
  } else {
    return "Greater than or equal to 10";
  }
}

orderMyLogic(7);
```

# --solutions--

```js
function orderMyLogic(val) {
  if(val < 5) {
    return "Less than 5";
  } else if (val < 10) {
    return "Less than 10";
  } else {
    return "Greater than or equal to 10";
  }
}
```
