---
id: 56bbb991ad1ed5201cd392ca
title: 使用索引存取陣列資料
challengeType: 1
forumTopicId: 16158
dashedName: access-array-data-with-indexes
---

# --description--

我們可以使用 <dfn>索引值</dfn> 存取陣列內的資料。

陣列索引使用與字串相同的括號表示法，不同的是它們指定的是陣列中的一個項目，而不是字元。像字串一樣，陣列使用<dfn>從零開始</dfn>的索引，因此陣列中的第一個元素索引為 `0`。

<br>

**範例**

```js
const array = [50, 60, 70];
console.log(array[0]);
const data = array[1];
```

`console.log(array[0])` 列印 `50`，而 `data` 的值是 `60`。

# --instructions--

創建一個名為 `myData` 的變數，並使用括號表示法將其設定為 `myArray` 的第一個值。

# --hints--

變數 `myData` 應該與 `myArray` 的第一個值相等。

```js
assert(
  (function () {
    if (
      typeof myArray !== 'undefined' &&
      typeof myData !== 'undefined' &&
      myArray[0] === myData
    ) {
      return true;
    } else {
      return false;
    }
  })()
);
```

應該使用括號表示法來存取變數 `myArray` 中的資料。

```js
assert(
  (function () {
    if (__helpers.removeJSComments(code).match(/\s*=\s*myArray\[0\]/g)) {
      return true;
    } else {
      return false;
    }
  })()
);
```

# --seed--

## --after-user-code--

```js
if(typeof myArray !== "undefined" && typeof myData !== "undefined"){(function(y,z){return 'myArray = ' + JSON.stringify(y) + ', myData = ' + JSON.stringify(z);})(myArray, myData);}
```

## --seed-contents--

```js
const myArray = [50, 60, 70];


```

# --solutions--

```js
const myArray = [50, 60, 70];
const myData = myArray[0];
```
