---
id: 56bbb991ad1ed5201cd392cc
title: 使用 pop 方法操作陣列
challengeType: 1
forumTopicId: 18236
dashedName: manipulate-arrays-with-pop
---

# --description--

另一種更改陣列中資料的方法是使用 `.pop()` 函式（程式）。

`.pop()` 用來從陣列的尾端彈出一個值。我們可以透過將這個彈出的值指定給一個變數來儲存它。換句話說，`.pop()` 會移除陣列的最後一個元素並傳回該元素。

任何型別的項目都可以從陣列中彈出－數字、字串，甚至是巢狀陣列。

```js
const threeArr = [1, 4, 6];
const oneDown = threeArr.pop();
console.log(oneDown);
console.log(threeArr);
```

第一個 `console.log` 會顯示值 `6`，第二個會顯示值 `[1, 4]`。

# --instructions--

使用 `.pop()` 函式（程式）從 `myArray` 移除最後一個項目，並將彈出值指定到新的變數 `removedFromMyArray`。

# --hints--

`myArray` 應該只包含 `[["John", 23]]`。

```js
assert(
  (function (d) {
    if (d[0][0] == 'John' && d[0][1] === 23 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

你應該在 `myArray` 上使用 `pop()`。

```js
assert(/removedFromMyArray\s*=\s*myArray\s*.\s*pop\s*(\s*)/.test(__helpers.removeJSComments(code)));
```

`removedFromMyArray` 應該只包含 `["cat", 2]`。

```js
assert(
  (function (d) {
    if (d[0] == 'cat' && d[1] === 2 && d[2] == undefined) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --after-user-code--

```js
if (typeof removedFromMyArray !== 'undefined') (function(y, z){return 'myArray = ' + JSON.stringify(y) + ' & removedFromMyArray = ' + JSON.stringify(z);})(myArray, removedFromMyArray);
```

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
const removedFromMyArray = myArray.pop();
```
