---
id: 56bbb991ad1ed5201cd392d0
title: 建置 JavaScript 物件
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

你可能聽過 `object` 這個詞。

物件類似於 `arrays`，但你不是使用索引來存取和修改它們的資料，而是透過所謂的 `properties` 來存取物件中的資料。

物件對於以結構化方式儲存資料非常有用，並且可以表述現實世界的物件，例如一隻貓。

這是一個範例貓咪物件：

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

在這個範例中，所有屬性都以字串形式儲存，例如 `name`、`legs` 和 `tails`。不過，你也可以使用數字作為屬性。你甚至可以省略單字字串屬性的引號，如下所示：

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

但是，如果你的物件有任何非字串屬性，JavaScript 將自動將它們強制轉型為字串。

# --instructions--

建立一個名為 `myDog` 的物件，該物件包含屬性 `name`（一個字串）、`legs`、`tails` 和 `friends`。

你可以將這些物件屬性設定為你想要的任何值，只要 `name` 是字串，`legs` 和 `tails` 是數字，且 `friends` 是陣列。

# --hints--

`myDog` 應該包含屬性 `name`，且該屬性應該是 `string`。

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` 應該包含屬性 `legs`，且該屬性應該是 `number`。

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` 應該包含屬性 `tails`，且該屬性應該是 `number`。

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` 應該包含屬性 `friends`，且它應該是一個 `array`。

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` 應該只包含所有給定的屬性。

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --after-user-code--

```js
(function(z){return z;})(myDog);
```

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
