---
id: 56bbb991ad1ed5201cd392cb
title: 使用 push 方法操作陣列
challengeType: 1
forumTopicId: 18237
dashedName: manipulate-arrays-with-push
---

# --description--

將資料附加到陣列末端的一個簡單方法是使用 `push()` 方法。

`push()` 方法接受一個或多個 <dfn>引數</dfn>，並依照它們出現的順序將它們附加到陣列的末端。它會傳回陣列的新長度。

範例：

```js
const arr1 = [1, 2, 3];
arr1.push(4, 5);

const arr2 = ["Stimpson", "J", "cat"];
arr2.push(["happy", "joy"]);
```

`arr1` 現在的值是 `[1, 2, 3, 4, 5]`，而 `arr2` 的值是 `["Stimpson", "J", "cat", ["happy", "joy"]]`。

# --instructions--

將 `["dog", 3]` 推入 `myArray` 變數的末端。

# --hints--

`myArray` 現在應該相等於 `[["John", 23], ["cat", 2], ["dog", 3]]`。

```js
assert(
  (function (d) {
    if (
      d[2] != undefined &&
      d[0][0] == 'John' &&
      d[0][1] === 23 &&
      d[2][0] == 'dog' &&
      d[2][1] === 3 &&
      d[2].length == 2
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --after-user-code--

```js
(function(z){return 'myArray = ' + JSON.stringify(z);})(myArray);
```

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
myArray.push(["dog",3]);
```
