---
id: 56bbb991ad1ed5201cd392d3
title: 從 JavaScript 物件中刪除屬性
challengeType: 1
forumTopicId: 17560
dashedName: delete-properties-from-a-javascript-object
---

# --description--

我們也可以這樣從物件中刪除屬性：

```js
delete ourDog.bark;
```

範例：

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"],
  "bark": "bow-wow"
};

delete ourDog.bark;
```

在上面顯示的最後一行之後，`ourDog` 看起來像：

```js
{
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
}
```

# --instructions--

從 `myDog` 刪除 `tails` 屬性。你可以使用點記法或括號記法。

# --hints--

你應該從 `myDog` 中刪除屬性 `tails`。

```js
assert(typeof myDog === 'object' && myDog.tails === undefined);
```

你不應該修改 `myDog` 設定。

```js
assert(__helpers.removeJSComments(code).match(/"tails": 1/g).length > 0);
```

# --seed--

## --after-user-code--

```js
(function(z){return z;})(myDog);
```

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};
delete myDog.tails;
```
