---
id: 56bbb991ad1ed5201cd392d2
title: 為你的 JavaScript 物件添加新屬性
challengeType: 1
forumTopicId: 301169
dashedName: add-new-properties-to-a-javascript-object
---

# --description--

你可以用修改它們的相同方式，為現有的 JavaScript 物件添加新屬性。

以下是我們如何為 `ourDog` 添加 `bark` 屬性：

```js
ourDog.bark = "bow-wow";
```

或

```js
ourDog["bark"] = "bow-wow";
```

現在當我們評估 `ourDog.bark` 時，我們會得到他的叫聲，`bow-wow`。

範例：

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};

ourDog.bark = "bow-wow";
```

# --instructions--

為你的 `myDog` 添加一個 `bark` 屬性，並將其設為狗叫聲，例如 "woof"。你可以使用點記法或括號記法。

# --hints--

你應該為你的 `myDog` 添加屬性 `bark`。

```js
assert(myDog.bark !== undefined);
```

你不應該將 `bark` 添加到 `myDog` 的初始化中。

```js
assert(!/bark[^\n]:/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function(z){return z;})(myDog);
```

## --seed-contents--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};


```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.bark = "Woof Woof";
```
