---
id: 56bbb991ad1ed5201cd392d1
title: 更新物件屬性
challengeType: 1
forumTopicId: 18336
dashedName: updating-object-properties
---

# --description--

在你創建 `JavaScript` 物件之後，你可以隨時像更新其他變數一樣更新它的屬性。你可以使用點記法或括號記法來更新。

例如，讓我們來看看 `ourDog`：

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```

因為他是一隻特別快樂的狗，所以我們將他的名字改為字串 `Happy Camper`。以下是如何更新他的物件名稱屬性：`ourDog.name = "Happy Camper";` 或 `ourDog["name"] = "Happy Camper";` 現在當我們評估 `ourDog.name` 時，不會得到 `Camper`，而是會得到他的新名字，`Happy Camper`。

# --instructions--

更新 `myDog` 物件的 name 屬性。讓我們將她的名字從 `Coder` 改成 `Happy Coder`。你可以使用點記法或括號記法。

# --hints--

你應該更新 `myDog` 的 `name` 屬性，使其相等於字串 `Happy Coder`。

```js
assert(/happy coder/gi.test(myDog.name));
```

你不應該編輯 `myDog` 定義。

```js
assert(/"name": "Coder"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function(z){return z;})(myDog);
```

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.name = "Happy Coder";
```
