---
id: 56bbb991ad1ed5201cd392ce
title: 使用 unshift 方法操作陣列
challengeType: 1
forumTopicId: 18239
dashedName: manipulate-arrays-with-unshift
---

# --description--

你不僅可以從陣列的開頭 `shift` 元素，還可以 `unshift` 元素到陣列的開頭，也就是在陣列前面添加元素。

`.unshift()` 的作用與 `.push()` 完全相同，但 `unshift()` 是將元素添加到陣列的開頭，而不是陣列的末端。

範例：

```js
const ourArray = ["Stimpson", "J", "cat"];
ourArray.shift();
ourArray.unshift("Happy");
```

在 `shift` 之後，`ourArray` 的值會是 `["J", "cat"]`。在 `unshift` 之後，`ourArray` 的值會是 `["Happy", "J", "cat"]`。

# --instructions--

使用 `unshift()` 將 `["Paul", 35]` 添加到 `myArray` 變數的開頭。

# --hints--

`myArray` 現在應該有 `[["Paul", 35], ["dog", 3]]`。

```js
assert(
  (function (d) {
    if (
      typeof d[0] === 'object' &&
      d[0][0] == 'Paul' &&
      d[0][1] === 35 &&
      d[1][0] != undefined &&
      d[1][0] == 'dog' &&
      d[1][1] != undefined &&
      d[1][1] == 3
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --after-user-code--

```js
(function(y, z){return 'myArray = ' + JSON.stringify(y);})(myArray);
```

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();
myArray.unshift(["Paul", 35]);
```
