---
id: 56bbb991ad1ed5201cd392cf
title: 使用函式（程式）撰寫可重複使用的 JavaScript
challengeType: 1
forumTopicId: 18378
dashedName: write-reusable-javascript-with-functions
---

# --description--

在 JavaScript 中，我們可以將程式碼分割成可重複使用的零件，稱為 <dfn>functions</dfn>。

這是一個函式（程式）的範例：

```js
function functionName() {
  console.log("Hello World");
}
```

你可以使用函式名稱後接小括弧來呼叫或<dfn>喚起</dfn>這個函式，如此：`functionName();` 每次呼叫該函式時，會在開發主控台列印出訊息 `Hello World`。每次呼叫函式時，大括弧之間的所有程式碼都會被執行。

# --instructions--

<ol>
  <li>
    Create a function called <code>reusableFunction</code> which prints the string <code>Hi World</code> to the dev console.
  </li>
  <li>
    Call the function.
  </li>
</ol>

# --hints--

`reusableFunction` 應該是一個函式（程式）。

```js
assert(typeof reusableFunction === 'function');
```

如果呼叫 `reusableFunction`，它應該將字串 `Hi World` 輸出到主控台。

```js
assert(testConsole());
```

你應該在定義 `reusableFunction` 後呼叫它一次。

```js
const functionStr = reusableFunction && __helpers.removeWhiteSpace(reusableFunction.toString());
const codeWithoutFunction = __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).replace(/reusableFunction\(\)\{/g, '');
assert(/reusableFunction\(\)/.test(codeWithoutFunction));
```

# --seed--

## --after-user-code--

```js

function testConsole() {
  var logOutput = "";
  var originalConsole = console;
  var nativeLog = console.log;
  var hiWorldWasLogged = false;
  console.log = function (message) {
    if(message === 'Hi World')  {
      console.warn(message)
      hiWorldWasLogged = true;
    }
    if(message && message.trim) logOutput = message.trim();
    if(nativeLog.apply) {
      nativeLog.apply(originalConsole, arguments);
    } else {
      var nativeMsg = Array.prototype.slice.apply(arguments).join(' ');
      nativeLog(nativeMsg);
    }
  };
  reusableFunction();
  console.log = nativeLog;
  return hiWorldWasLogged;
}

```

## --seed-contents--

```js

```

# --solutions--

```js
function reusableFunction() {
  console.log("Hi World");
}
reusableFunction();
```
