---
id: 56bbb991ad1ed5201cd392cd
title: 使用 shift 方法操作陣列
challengeType: 1
forumTopicId: 18238
dashedName: manipulate-arrays-with-shift
---

# --description--

`pop()` 總是移除陣列的最後一個元素。如果你想移除第一個元素怎麼辦？

這就是 `.shift()` 的用武之地。它的作用就像 `.pop()`，只是移除的是第一個元素，而不是最後一個。

範例：

```js
const ourArray = ["Stimpson", "J", ["cat"]];
const removedFromOurArray = ourArray.shift();
```

`removedFromOurArray` 的值會是字串 `Stimpson`，而 `ourArray` 會是 `["J", ["cat"]]`。

# --instructions--

使用 `.shift()` 函式（程式）從 `myArray` 移除第一個項目，並將「被移除」的值指定到新的變數 `removedFromMyArray`。

# --hints--

`myArray` 現在應該相等於 `[["dog", 3]]`。

```js
assert(
  (function (d) {
    if (d[0][0] == 'dog' && d[0][1] === 3 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`removedFromMyArray` 應該包含 `["John", 23]`。

```js
assert(
  (function (d) {
    if (
      d[0] == 'John' &&
      d[1] === 23 &&
      typeof removedFromMyArray === 'object'
    ) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --after-user-code--

```js
if (typeof removedFromMyArray !== 'undefined') (function(y, z){return 'myArray = ' + JSON.stringify(y) + ' & removedFromMyArray = ' + JSON.stringify(z);})(myArray, removedFromMyArray);
```

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line
const removedFromMyArray = myArray.shift();
```
