---
id: 587d7b87367417b2b2512b3f
title: 探索 `var` 和 `let` 關鍵字之間的差異
challengeType: 1
forumTopicId: 301202
dashedName: explore-differences-between-the-var-and-let-keywords
---

# --description--

使用 `var` 關鍵字宣告變數的最大問題之一是你很容易覆寫變數宣告：

```js
var camper = "James";
var camper = "David";
console.log(camper);
```

在上述程式碼中，`camper` 變數最初宣告為 `James`，然後被覆寫為 `David`。主控台接著顯示字串 `David`。

在小型應用程式中，你可能不會遇到這種類型的問題。但隨著你的程式碼庫變得更大，你可能會不小心覆寫你不打算覆寫的變數。因為這種行為不會發出錯誤，搜尋和修正錯誤會變得更困難。

在 ES6（JavaScript 的重大更新）中，介紹了一個名為 `let` 的關鍵字，以解決 `var` 關鍵字的這個潛在問題。你將在後續的挑戰中學習其他 ES6 特性。

如果你在上述程式碼中將 `var` 取代為 `let`，將會導致錯誤：

```js
let camper = "James";
let camper = "David";
```

你可以在瀏覽器主控台中看到該錯誤。

所以與 `var` 不同，當你使用 `let` 時，名稱相同的變數只能宣告一次。

# --instructions--

更新程式碼，使其僅使用 `let` 關鍵字。

# --hints--

`var` 不應該存在於程式碼中。

```js
assert.notMatch(code, /var/g);
```

`catName` 應該是字串 `Oliver`。

```js
assert.equal(catName, 'Oliver');
```

`catSound` 應該是字串 `Meow!`

```js
assert.equal(catSound, 'Meow!');
```

# --seed--

## --seed-contents--

```js
var catName = "Oliver";
var catSound = "Meow!";
```

# --solutions--

```js
let catName = "Oliver";
let catSound = "Meow!";
```
