---
id: 587d7b7e367417b2b2512b24
title: 使用條件（三元）運算子
challengeType: 1
forumTopicId: 301181
dashedName: use-the-conditional-ternary-operator
---

# --description--

<dfn>條件運算子</dfn>，也稱為<dfn>三元運算子</dfn>，可以用作單行的 if-else 表達式。

語法是 `a ? b : c`，其中 `a` 是條件，當條件傳回 `true` 時執行的程式碼是 `b`，當條件傳回 `false` 時執行的程式碼是 `c`。

以下函式使用 `if/else` 述語來檢查條件：

```js
function findGreater(a, b) {
  if(a > b) {
    return "a is greater";
  }
  else {
    return "b is greater or equal";
  }
}
```

這可以使用條件運算子重新撰寫：

```js
function findGreater(a, b) {
  return a > b ? "a is greater" : "b is greater or equal";
}
```

# --instructions--

在 `checkEqual` 函式中使用條件運算子來檢查兩個數字是否相等。該函式應該傳回字串 `Equal` 或字串 `Not Equal`。

# --hints--

`checkEqual` 應該使用條件運算子

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?/.test(__helpers.removeJSComments(code)));
```

`checkEqual(1, 2)` 應該傳回字串 `Not Equal`

```js
assert(checkEqual(1, 2) === 'Not Equal');
```

`checkEqual(1, 1)` 應該傳回字串 `Equal`

```js
assert(checkEqual(1, 1) === 'Equal');
```

`checkEqual(1, -1)` 應該傳回字串 `Not Equal`

```js
assert(checkEqual(1, -1) === 'Not Equal');
```

# --seed--

## --seed-contents--

```js
function checkEqual(a, b) {

}

checkEqual(1, 2);
```

# --solutions--

```js
function checkEqual(a, b) {
  return a === b ? "Equal" : "Not Equal";
}
```
