---
id: 587d7b7e367417b2b2512b23
title: 使用 `parseInt` 函式（程式）
challengeType: 1
forumTopicId: 301183
dashedName: use-the-parseint-function
---

# --description--

`parseInt()` 函式解析字串並傳回整數。以下是一個範例：

```js
const a = parseInt("007");
```

上述函式將字串 `007` 轉換為整數 `7`。如果字串中的第一個字元無法轉換成數字，則會傳回 `NaN`。

# --instructions--

在 `convertToInteger` 函式中使用 `parseInt()`，使其將輸入字串 `str` 解析成整數，並傳回該整數。

# --hints--

`convertToInteger` 應該使用 `parseInt()` 函式（程式）。

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("56")` 應該傳回一個數字

```js
assert(typeof convertToInteger('56') === 'number');
```

`convertToInteger("56")` 應該傳回 56

```js
assert(convertToInteger('56') === 56);
```

`convertToInteger("77")` 應該傳回 77

```js
assert(convertToInteger('77') === 77);
```

`convertToInteger("JamesBond")` 應該傳回 `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("56");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str);
}
```
