---
id: 587d7b7e367417b2b2512b21
title: 使用多重條件（三元）運算子
challengeType: 1
forumTopicId: 301179
dashedName: use-multiple-conditional-ternary-operators
---

# --description--

在前一個挑戰中，你使用了單一條件運算子。你也可以將它們串鏈起來以檢查多個條件。

以下函式使用 `if`、`else if` 和 `else` 述語來檢查多個條件：

```js
function findGreaterOrEqual(a, b) {
  if (a === b) {
    return "a and b are equal";
  }
  else if (a > b) {
    return "a is greater";
  }
  else {
    return "b is greater";
  }
}
```

上述函式（程式）可以使用多個條件運算子重新撰寫：

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" 
    : (a > b) ? "a is greater" 
    : "b is greater";
}
```

最佳做法是將多個條件運算子格式化，使每個條件各自位於一行，如上所示。若未適當縮排多個條件運算子，可能會使你的程式碼難以閱讀。例如：

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" : (a > b) ? "a is greater" : "b is greater";
}
```

# --instructions--

在 `checkSign` 函式中，使用多個條件運算子－遵循 `findGreaterOrEqual` 中使用的建議格式－來檢查一個數字是正數、負數還是零。該函式應該傳回 `positive`、`negative` 或 `zero`。

# --hints--

`checkSign` 應該使用多個條件運算子

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?\s*?\?\s*?.+?\s*?:\s*?.+?/gi.test(__helpers.removeJSComments(code)));
```

`checkSign(10)` 應該傳回字串 `positive`。請注意大小寫有區別。

```js
assert(checkSign(10) === 'positive');
```

`checkSign(-12)` 應該傳回字串 `negative`。請注意大小寫很重要。

```js
assert(checkSign(-12) === 'negative');
```

`checkSign(0)` 應該傳回字串 `zero`。請注意大小寫很重要。

```js
assert(checkSign(0) === 'zero');
```

# --seed--

## --seed-contents--

```js
function checkSign(num) {

}

checkSign(10);
```

# --solutions--

```js
function checkSign(num) {
  return (num > 0) ? 'positive' : (num < 0) ? 'negative' : 'zero';
}
```
