---
id: 587d7b7e367417b2b2512b22
title: 使用帶有基數的 `parseInt` 函式（程式）
challengeType: 1
forumTopicId: 301182
dashedName: use-the-parseint-function-with-a-radix
---

# --description--

`parseInt()` 函式解析字串並傳回整數。它接受第二個引數 radix，該引數指定字串中數字的進位制。radix 可以是介於 2 和 36 之間的整數。

函式（程式）呼叫看起來像：

```js
parseInt(string, radix);
```

以下是一個範例：

```js
const a = parseInt("11", 2);
```

radix 變數表示 `11` 是在二進位系統，或基數 2。此範例將字串 `11` 轉換為整數 `3`。

# --instructions--

在 `convertToInteger` 函式（程式）中使用 `parseInt()`，使其將二進位數字轉換為整數並傳回。

# --hints--

`convertToInteger` 應該使用 `parseInt()` 函式（程式）。

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("10011")` 應該傳回一個數字

```js
assert(typeof convertToInteger('10011') === 'number');
```

`convertToInteger("10011")` 應該傳回 19

```js
assert(convertToInteger('10011') === 19);
```

`convertToInteger("111001")` 應該傳回 57

```js
assert(convertToInteger('111001') === 57);
```

`convertToInteger("JamesBond")` 應該傳回 `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("10011");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str, 2);
}
```
