---
id: 598e8944f009e646fc236146
title: 理解從函式傳回的未定義值
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

函式（程式）可以包含 `return` 述語，但不一定要包含。如果函式（程式）沒有 `return` 述語，當你呼叫它時，函式（程式）會處理內部程式碼，但傳回值是 `undefined`。

**範例**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum` 是一個沒有 `return` 述語的函式。該函式會改變全域的 `sum` 變數，但函式傳回的值是 `undefined`。

# --instructions--

創建一個沒有任何引數的函式（程式）`addFive`。這個函式（程式）會將 5 加到 `sum` 變數，但它傳回的值是 `undefined`。

# --hints--

`addFive` 應該是一個函式（程式）。

```js
assert(typeof addFive === 'function');
```

當兩個函式（程式）都執行完畢後，`sum` 應該要相等於 `8`。

```js
assert(sum === 8);
```

從 `addFive` 傳回的值應該是 `undefined`。

```js
assert(addFive() === undefined);
```

在 `addFive` 函式（程式）中，你應該將 `5` 添加到 `sum` 變數。

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
