---
id: 598e8944f009e646fc236146
title: 函數也可以返回 undefined
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

A function can include the `return` statement but it does not have to. In the case that the function doesn't have a `return` statement, when you call it, the function processes the inner code but the returned value is `undefined`.

**例如：**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum` 是一個沒有 `return` 語句的函數。 該函數將更改全局變量 `sum`，函數的返回值爲 `undefined`。

# --instructions--

創建一個沒有任何參數的函數 `addFive`。 此函數使 `sum` 變量加 5，但其返回值是 `undefined`。

# --hints--

`addFive` 應該是一個函數。

```js
assert(typeof addFive === 'function');
```

兩個函數運行後，`sum` 應該等於 `8`。

```js
assert(sum === 8);
```

`addFive` 的返回值應該是 `undefined`。

```js
assert(addFive() === undefined);
```

在 `addFive` 函數中，應該給 `sum` 變量增加 `5`。

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
