---
id: 599a789b454f2bbd91a3ff4d
title: 練習比較不同的值
challengeType: 1
forumTopicId: 301174
dashedName: practice-comparing-different-values
---

# --description--

在前兩個挑戰中，我們學習了等號運算子（`==`）和嚴格等號運算子（`===`）。讓我們快速複習並多練習使用這些運算子。

如果被比較的值不是相同的型別，等號運算子會執行型別轉換，然後計算這些值。然而，嚴格等號運算子會直接比較資料型別和數值，不會將其中一種型別轉換成另一種。

**範例**

`3 == '3'` 傳回 `true`，因為 JavaScript 會從字串轉換為數字的型別。`3 === '3'` 傳回 `false`，因為型別不同且不進行型別轉換。

**注意：** 在 JavaScript 中，你可以使用 `typeof` 運算子來判斷變數或值的型別，如下所示：

```js
typeof 3
typeof '3'
```

`typeof 3` 傳回字串 `number`，而 `typeof '3'` 傳回字串 `string`。

# --instructions--

編輯器中的 `compareEquality` 函式（程式）使用等號運算子比較兩個值。修改該函式，使其僅在兩個值嚴格相等時傳回字串 `Equal`。

# --hints--

`compareEquality(10, "10")` 應該傳回字串 `Not Equal`

```js
assert(compareEquality(10, '10') === 'Not Equal');
```

`compareEquality("20", 20)` 應該傳回字串 `Not Equal`

```js
assert(compareEquality('20', 20) === 'Not Equal');
```

你應該使用 `===` 運算子

```js
assert(__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function compareEquality(a, b) {
  if (a == b) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

compareEquality(10, "10");
```

# --solutions--

```js
function compareEquality(a,b) {
  if (a === b) {
    return "Equal";
  }
  return "Not Equal";
}
```
