---
id: 5cc0bd7a49b71cb96132e54c
title: 使用遞迴來創建數字範圍
challengeType: 1
forumTopicId: 301180
dashedName: use-recursion-to-create-a-range-of-numbers
---

# --description--

接續前一個挑戰，我們為你提供另一個機會來創建遞迴函式以解決問題。

# --instructions--

我們定義了一個名為 `rangeOfNumbers` 的函式（程式），具有兩個參數。該函式應該傳回一個整數陣列，該陣列以由 `startNum` 參數表述的數字開始，以由 `endNum` 參數表述的數字結束。起始數字將始終小於或相等於結束數字。你的函式必須使用遞迴，透過呼叫自身來實現，且不得使用任何形式的迴圈。它也應該適用於 `startNum` 和 `endNum` 相同的情況。

# --hints--

你的函式應該傳回一個陣列。

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

你的程式碼不應該使用任何迴圈語法（`for` 或 `while` 或較高階的函式（程式）如 `forEach`、`map`、`filter` 或 `reduce`）。

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

`rangeOfNumbers` 應該使用遞迴（呼叫自身）來解決此挑戰。

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

`rangeOfNumbers(1, 5)` 應該傳回 `[1, 2, 3, 4, 5]`。

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)` 應該傳回 `[6, 7, 8, 9]`。

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)` 應該傳回 `[4]`。

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

不應該使用全域變數來快取陣列。

```js
rangeOfNumbers(1, 3)
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js
function rangeOfNumbers(startNum, endNum) {
  return [];
};
```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
