---
id: 5cd9a70215d3c4e65518328f
title: 使用遞迴來創建倒數計時
challengeType: 1
forumTopicId: 305925
dashedName: use-recursion-to-create-a-countdown
---

# --description--

在<a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/replace-loops-using-recursion" target="_blank" rel="noopener noreferrer nofollow">先前的挑戰</a>中，你學會了如何使用遞迴來取代 `for` 迴圈。現在，讓我們來看看一個更複雜的函式，它會傳回一個從 `1` 開始到傳入函式的數字的連續整數陣列。

如同前一個挑戰中提到的，會有一個 <dfn>基底情況</dfn>。基底情況告訴遞迴函式何時不再需要呼叫自己。這是一個簡單的情況，返回值已經知道。也會有一個 <dfn>遞迴呼叫</dfn>，它會以不同的引數執行原始函式。如果函式寫得正確，最終會達到基底情況。

例如，假設你想要撰寫一個遞迴函式（程式），該函式會傳回一個包含數字 `1` 到 `n` 的陣列。這個函式需要接受一個引數 `n`，用來表述最後的數字。然後它需要以逐漸減小的 `n` 值呼叫自己，直到達到 `1`。你可以將函式寫成如下：

```javascript
function countup(n) {
  if (n < 1) {
    return [];
  } else {
    const countArray = countup(n - 1);
    countArray.push(n);
    return countArray;
  }
}
console.log(countup(5));
```

值 `[1, 2, 3, 4, 5]` 將會顯示在主控台。

一開始，這看起來似乎違反直覺，因為 `n` 的值*減少*，但最終陣列中的值卻是*遞增*的。這是因為 `push` 發生在最後，在遞迴呼叫傳回之後。在將 `n` 推入陣列的時候，`countup(n - 1)` 已經被評估並傳回了 `[1, 2, ..., n - 1]`。

# --instructions--

我們定義了一個名為 `countdown` 的函式，帶有一個參數（`n`）。該函式應該使用遞迴來傳回一個包含從 `n` 到 1 的整數的陣列，根據 `n` 參數。如果函式以小於 1 的數字呼叫，函式應該傳回一個空陣列。例如，當以 `n = 5` 呼叫此函式時，應該傳回陣列 `[5, 4, 3, 2, 1]`。你的函式必須透過呼叫自身來使用遞迴，且不得使用任何形式的迴圈。

# --hints--

`countdown(-1)` 應該傳回一個空的陣列。

```js
assert.isEmpty(countdown(-1));
```

`countdown(10)` 應該傳回 `[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(10), [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

`countdown(5)` 應該傳回 `[5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

你的程式碼不應該依賴任何類型的迴圈（`for`、`while` 或較高階的函式（程式）例如 `forEach`、`map`、`filter` 和 `reduce`）。

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

你應該使用遞迴來解決這個問題。

```js
assert(
  countdown.toString().match(/countdown\s*\(.+\)/)
);
```

不應該使用全域變數來快取陣列。

```js
countdown(1)
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
function countdown(n){
  return;
}
// Only change code above this line
```

# --solutions--

```js
function countdown(n){
   return n < 1 ? [] : [n].concat(countdown(n - 1));
}
```
