---
id: 5ee127a03c3b35dd45426493
title: 將一個變數的值指定給另一個變數
challengeType: 1
forumTopicId: 418265
dashedName: assigning-the-value-of-one-variable-to-another
---

# --description--

在使用 <dfn>assignment</dfn> 運算子將值指定給變數之後，你可以使用 <dfn>assignment</dfn> 運算子將該變數的值指定給另一個變數。

```js
var myVar;
myVar = 5;
var myNum;
myNum = myVar;
```

上述宣告了一個沒有值的 `myVar` 變數，然後將它指定為值 `5`。接著，宣告了一個名為 `myNum` 的沒有值的變數。然後，將 `myVar` 的內容（即 `5`）指定給變數 `myNum`。現在，`myNum` 也有值為 `5`。

# --instructions--

將 `a` 的內容指定到變數 `b`。

# --hints--

你不應該更改指定註解以上的程式碼。

```js
assert(/var a;/.test(__helpers.removeJSComments(code)) && /a = 7;/.test(__helpers.removeJSComments(code)) && /var b;/.test(__helpers.removeJSComments(code)));
```

`b` 應該有一個值為 `7`。

```js
assert(typeof b === 'number' && b === 7);
```

應該使用 `=` 將 `a` 指定、設值給 `b`。

```js
assert(/b\s*=\s*a\s*/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
var a;
a = 7;
var b;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
var b;
b = a;
```
