---
id: 6995d0325fd7da40854fac40
title: 第一步驟
challengeType: 1
dashedName: step-1
---

# --description--

在之前的課程中，你學習了二分搜尋法以及它與線性搜尋的不同之處。

在這個工作坊中，你將透過建置一個二分搜尋演算法來實作這些知識。

首先，創建一個 `binarySearch` 函式，並帶有 `searchList` 和 `value` 兩個參數。

`searchList` 是函式將要搜尋的陣列，而 `value` 是函式將嘗試在 `searchList` 中尋找的標的。

# --hints--

你應該創建一個 `binarySearch` 函式。

```js
assert.isFunction(binarySearch);
```

你的 `binarySearch` 函式應該有 `searchList` 和 `value` 兩個參數。

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
