---
id: 6995d0325fd7da40854fac40
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在之前的課程中，你學習了二分搜尋法以及它與線性搜尋的不同之處。

在這個工作坊中，你將透過建置二分搜尋法演算法來實作這個概念。

首先，創建一個 `binarySearch` 函式，並帶有 `searchList` 和 `value` 兩個參數。

`searchList` 是函式將要搜尋的陣列，而 `value` 是函式將嘗試在 `searchList` 中尋找的標的。

# --hints--

你應該創建一個 `binarySearch` 函式。

```js
assert.isFunction(binarySearch);
```

你的 `binarySearch` 函式應該有 `searchList` 和 `value` 兩個參數。

```js
const regex = __helpers.functionRegex('binarySearch', ['searchList', 'value']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
