---
id: 6996cb5869acf58de751ae81
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

回想一下，這個函式會接收一個名為 `searchList` 的陣列和你要尋找的 `value`。

現在，使用 `binarySearch([1, 2, 3, 4, 5], 3)` 和 `binarySearch([1, 2, 3, 4, 5, 9], 4)` 呼叫這個函式，並將兩個述語都輸出到主控台。

# --hints--

你應該使用 `binarySearch([1, 2, 3, 4, 5], 3)` 呼叫函式，並將結果輸出到主控台。

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*\]\s*,\s*3\s*\)\s*\)/
);
```

你應該使用 `binarySearch([1, 2, 3, 4, 5, 9], 4)` 呼叫函式，並將結果輸出到主控台。

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*,\s*9\s*\]\s*,\s*4\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--  

--fcc-editable-region--  
```
