---
id: 6996d18b31aab5cbbffbc731
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

你還可以添加的最後一個小改進是同時顯示找到標的 `value` 的索引。

所以，不要只在 `if` 述語中傳回 `pathToTarget`，而是傳回一個陣列，裡面包含 `pathToTarget` 和一個帶有訊息 `Value found at index ${mid}` 的模板字串。

這樣，你的二分搜尋法演算法工作坊就完成了！

# --hints--

你應該更新 `if` 中的 `return` 述語，使函式傳回一個陣列，第一個元素是 `pathToTarget`，第二個元素是 `Value found at index ${mid}` 的模板字串。記得模板字串要用反引號包住。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch: binarySearchFn } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearchFn?.toString());
assert.match(
  codeWithoutComments,
  /return\s*\[\s*pathToTarget\s*,\s*(?:`Value found at index\s*\$\{\s*mid\s*\}`|['"]Value found at index\s*['"]\s*\.\s*concat\(\s*mid\s*\))\s*\]/
);
```

你的函式應該會傳回預期的結果。

```js
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5], 3), [[3], "Value found at index 2"]
);
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5, 9], 4), [[3, 5, 4], "Value found at index 3"]
);
assert.deepEqual(
  binarySearch([1, 3, 5, 9, 14, 22], 10), [[], "Value not found"]
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
--fcc-editable-region--
      return pathToTarget;
--fcc-editable-region--
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```

# --solutions--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return [pathToTarget, `Value found at index ${mid}`];
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
