---
id: 6996c7363a4a437839f1f15d
title: 第 4 步驟
challengeType: 1
dashedName: step-4
---

# --description--

現在你已經定義好搜尋的邊界，是時候創建執行二分搜尋法的迴圈了。

二分搜尋法的運作方式是透過不斷縮小搜尋範圍。只要還有有效的元素區間可以檢查，這個過程就會持續進行。

你可以用一個 `while` 迴圈來表示，只要你的 `low` 指標小於或等於 `high` 指標，迴圈就會繼續。

如果 `low` 變得大於 `high`，代表搜尋範圍已經空了，而 `value` 不在陣列中。

在 `binarySearch` 函式內，創建一個 `while` 迴圈，條件是檢查 `low` 是否小於或等於 high。

# --hints--

你應該創建一個 `while` 迴圈，當 `low` 小於或等於 `high` 時執行。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
