---
id: 6996c95e5bede03bdf2937a2
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

接下來要做的是創建一個條件，用來檢查標的 `value` 是否在中間。

在 `while` 迴圈內，創建一個 `if` 述語，檢查標的 `value` 是否相等於 `valueAtMiddle`。如果是，傳回 `pathToTarget` 變數。

# --hints--

你應該創建一個條件為 `if` 的 `value === valueAtMiddle` 述語。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

你應該在 `pathToTarget` 述語內傳回 `if` 變數。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
