---
id: 6996c796385b6ea4e222d5e8
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

要開始使用迴圈，你需要找到目前搜尋空間的中間索引。

你可以透過計算 `low` 和 `high` 變數的平均值來取得中間索引。你必須在將兩個值相加後使用向下取整除法來取得平均值，因此答案將永遠是向下取整的整數，因為索引必須是整數。

在你的 `while` 迴圈中，宣告一個 `mid` 變數，並將其設為 `low` 和 `high` 的平均值。

# --hints--

在 while 迴圈中，你應該宣告一個 `mid` 變數，並將其設為 `Math.floor((low + high) / 2)`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
