---
id: 6996c8f94fe6b250f11a2a1e
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

現在，將 `valueAtMiddle` 加入 `pathToTarget` 陣列中，以追蹤搜尋過程中的每一步，不論是否為標的。你可以使用 `push()` 方法來完成這件事。

# --hints--

你應該使用 `push()` 方法，將 `valueAtMiddle` 添加到 `pathToTarget` 陣列中。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
