---
id: 6996c675b599baa7b103192d
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

在二分搜尋法中，有最低可能的索引和最高可能的索引。

最低可能的索引表述目前被搜尋值範圍中的最左邊界，在此案例中是 `searchList`。另一方面，最高可能的索引則表述被搜尋值範圍中的最右邊界。

為了考慮這兩者，你必須考慮整個正在搜尋的陣列。所以，在 `binarySearch` 函式（程式）內，宣告一個值為 `0` 的 `low` 變數，以及一個值為該陣列最後索引的 `high` 變數。

# --hints--

你應該在你的 `binarySearch` 函式（程式）中宣告一個名為 `low` 且設為 `0` 的變數。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

你應該在你的 `binarySearch` 函式（程式）中宣告一個名為 `high` 且設為 `searchList.length - 1` 的變數。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
