---
id: 6996d1032feed88e2957dc8e
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

你可以看到最後一次呼叫列印出一個空陣列。為了更明確表示該值不在被搜尋的陣列中，請更新 `return` 述語，使其傳回一個包含兩個元素的陣列：一個空陣列，以及字串 `Value not found`。

# --hints--

你應該更新 `return` 述語，使其傳回 `[], 'Value not found'`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /return\s*\[\s*\[\s*\]\s*,\s*('|"|`)Value not found\1\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
    return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
