---
id: 6996c60281feca6718cc15c1
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在你的 `binarySearch` 函式（程式）內，創建一個 `pathToTarget` 變數並初始化為空陣列，作為開始。

# --hints--

你應該在 `pathToTarget` 函式（程式）內有一個 `[]` 變數設定為空陣列 `binarySearch`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
