---
id: 6996ccadf9048c79b857bb15
title: 第 11 步驟
challengeType: 1
dashedName: step-11
---

# --description--

你可以看到第一次呼叫傳回 `[3]`。那是因為 `3` 是 `[1, 2, 3, 4, 5]` 中間的值。

另一方面，第二次呼叫傳回一個空陣列。這是因為在檢查中間值後，迴圈目前缺少透過更新 `low` 或 `high` 變數來縮小搜尋範圍的邏輯。它只執行一次中點檢查，然後結束，沒有進一步搜尋。

為了讓二分搜尋法能繼續縮小搜尋範圍，請新增一個 `else if` 述語，檢查 `value` 是否大於 `valueAtMiddle`。

# --hints--

你應該新增一個 `else if` 述語，條件為 `value > valueAtMiddle`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
