---
id: 6996ca21a74104721f2c7e5d
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

你需要測試一些東西，這樣你才能在這個初始階段理解演算法的流程。

為此，你必須先跳出迴圈。原因是目前的實作只允許執行一次迭代，如果條件不符合，將會產生無窮迴圈。

在 `if` 述語之後，使用 `break` 關鍵字跳出 `while` 迴圈。接著，在 `while` 迴圈之後，傳回一個空陣列以表示沒有找到該值。

# --hints--

你應該在 `break` 述語之後，使用 `while` 關鍵字跳出 `if` 迴圈。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

你應該在 `while` 迴圈之後傳回一個空陣列。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
