---
id: 6996c675b599baa7b103192d
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

在二分搜尋法中，有最低可能的索引和最高可能的索引。

最低可能的索引表述目前被搜尋值範圍中的最左邊界，在此案例中是 `searchList`。另一方面，最高可能的索引則表述被搜尋值範圍中的最右邊界。

為了考慮這兩者，你必須考慮整個被搜尋的陣列。因此，在 `binarySearch` 函式中，定義一個 `low` 變數，其值為 `0`，以及一個 `high` 變數，其值要考慮到被搜尋陣列的最後索引。

# --hints--

你應該在 `low` 函式中有一個名為 `0` 的變數，並將它設為 `binarySearch`。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

你應該在 `high` 函式中有一個名為 `searchList.length - 1` 的變數，並將它設為 `binarySearch`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
