---
id: 6996ca21a74104721f2c7e5d
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

你需要測試各種事物，這樣你才能在這個初始階段了解演算法的流程。

為了做到這一點，你需要先跳出迴圈。那是因為目前的實作只允許一次迭代，所以如果條件不符合，將會有無窮迴圈。

在 `if` 述語之後，使用 `break` 關鍵字跳出 `while` 迴圈。接著，在 `while` 迴圈之後，傳回一個空陣列以表示沒有找到該值。

# --hints--

你應該在 `if` 述語之後使用 `break` 關鍵字來跳出 `while` 迴圈。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

你應該在 `while` 迴圈之後傳回一個空陣列。

```js
assert.deepEqual(binarySearch([], NaN), []);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
