---
id: 6996cfb6c01e9a802add38be
title: 第 13 步驟
challengeType: 1
dashedName: step-13
---

# --description--

最後，如果 `value` 沒有在中間被找到，且也不大於 `valueAtMiddle`，那麼 `value` 必須小於 `valueAtMiddle`，並且必須在左邊。最終的 `else` 區塊將處理這個情況。

所以，透過移除 `break` 關鍵字並添加一個 `else` 區塊來完成迴圈。在其中，透過從 `mid` 減去 `1` 來更新 `high` 變數。

# --hints--

你應該將 `break` 關鍵字替換成一個 `else` 區塊，該區塊會將 `high` 變數更新為 `mid - 1`。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
assert.notMatch(codeWithoutComments, /\bbreak\b/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
--fcc-editable-region--
    } 
    break;
--fcc-editable-region--
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
