---
id: 6996d0a2c2ac55041145a13f
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

你有注意到第二次函式呼叫現在在列表中有 `3, 5, 4` 這些值嗎？那顯示二分搜尋法正在如預期運作。

事情是這樣發生的：演算法首先將 `3` 視為初始列表的中間。由於 `4` 大於 `3`，搜尋轉移到右半部。接著它將 `5` 視為新的中間。因為 `4` 小於 `5`，搜尋移向左邊，最終將 `4` 識別為最終範圍的中間。

要再次測試這個函式，請用 `[1, 3, 5, 9, 14, 22], 10` 呼叫它並將結果輸出到主控台。這是個找不到該值的情況。

# --hints--

你應該用 `binarySearch([1, 3, 5, 9, 14, 22], 10)` 呼叫函式並將結果輸出到主控台。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,3,5,9,14,22],10))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
