---
id: 6996c60281feca6718cc15c1
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在你的 `binarySearch` 函式（程式）內，創建一個 `pathToTarget` 變數並初始化為空陣列，作為開始。

# --hints--

你應該在你的 `binarySearch` 函式中宣告一個 `pathToTarget` 變數，並將它初始化為空陣列。

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
