---
id: 6996c7363a4a437839f1f15d
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

既然你已經定義了搜尋的邊界，現在是時候創建將執行二分搜尋法的迴圈。

二分搜尋法透過反覆縮小搜尋範圍來運作。只要有有效的元素區間可供檢查，這個處理程序就會持續進行。

你可以用一個 `while` 迴圈來表示，只要你的 `low` 指標小於或相等於你的 `high` 指標，迴圈就會繼續。

如果 `low` 變得大於 `high`，代表搜尋範圍已經空了，而 `value` 不在陣列中。

在 `binarySearch` 函式（程式）內，創建一個 `while` 迴圈，條件是檢查 `low` 是否小於或相等於 `high`。

# --hints--

你應該創建一個 `while` 迴圈，當 `low` 小於或相等於 `high` 時執行。

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
