---
id: 6996c8375f4301004459ee85
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

現在你已經計算出 `mid` 索引，接下來你需要從 `searchList` 中實際取得該索引位置的值。這個值將用來與你的 `value` 參數，也就是你正在搜尋的標的，比較。

創建一個 `valueAtMiddle` 變數，並將 `searchList` 中位於 `mid` 索引的元素指定給它。

# --hints--

你應該創建一個 `valueAtMiddle` 變數並設值為 `searchList[mid]`。

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
