---
id: 690ddb5e6ee94eb73ed172b0
title: Налагодження форми пожертвування
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

Місцева благодійна організація створила вебсайт з формою пожертвування, але є кілька проблем, які потрібно виправити. Форма недоступна і містить деякі синтаксичні помилки HTML.

Ваше завдання — виправити всі помилки, щоб форма працювала правильно і була доступна для всіх користувачів. Виконайте пункти в історіях користувача нижче і натисніть "Run the Tests", щоб перевірити, чи ви виправили всі помилки.

**Історії користувача:**

1. Елементи `input` є порожніми елементами і не повинні мати кінцевих тегів. Видаліть усі кінцеві теги `</input>` з форми.
2. Додайте елементи `label` для кожного поля введення форми, щоб користувачі знали, для чого кожне поле. Текст мітки повинен відповідати тому, що зараз знаходиться поруч із кожним полем введення.
3. Тип введення `Email Address:` повинен бути `email` замість `text`.
4. Ви повинні пов’язати кожен елемент `label` з відповідним елементом `input` за допомогою атрибута `for` на `label` і відповідного атрибута `id` на `input`.
5. Додайте атрибут `required` до полів введення тексту, е-пошти та числа (але не до прапорця або кнопки відправки), щоб гарантувати, що користувачі заповнюють необхідну інформацію.

# --hints--

У вашому коді не повинно бути кінцевих тегів `</input>`.

```js
assert.notMatch(code, /<\/input>/);
```

У вашій формі має бути рівно п’ять елементів `input`.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

У вашій формі має бути рівно чотири елементи `label`.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Ваш перший `label` повинен містити текст `Full Name:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

Ваш перший `input` повинен мати атрибут `required`.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

Ваш другий `label` повинен містити текст `Email Address:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

Ваш другий `input` повинен мати атрибут `required`.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

Ваш третій `label` повинен містити текст `Donation Amount ($):`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

Ваш третій `input` повинен мати атрибут `required`.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

Ваш четвертий `label` повинен містити текст `Subscribe`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

Ви повинні пов’язати кожен елемент `input` з елементом `label` за допомогою атрибутів `for` і `id`.

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

Ваш `input` для адреси е-пошти повинен бути типу `email`.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

Ваш прапорець `input` не повинен мати атрибута `required`.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

Ваша кнопка відправки `button` не повинна мати атрибута `required`.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
