---
id: 5a24c314108439a4d4036143
title: استخراج منطق الحالة إلى Redux
challengeType: 6
forumTopicId: 301429
dashedName: extract-state-logic-to-redux
---

# --description--

الآن بعد أن أنهيت مكوِّن React، تحتاج إلى نقل المنطق الذي ينفذه محليًا في `state` إلى Redux. هذه هي الخطوة الأولى لربط تطبيق React البسيط بـ Redux. الوظيفة الوحيدة لتطبيقك هي إضافة رسائل جديدة من المستخدم إلى قائمة غير مرتبة. المثال بسيط ليُظهر كيف يعمل React وRedux معًا.

# --instructions--

أولًا، عرّف نوع إجراء `ADD` واجعله ثابتًا `ADD`. بعد ذلك، عرّف منشئ إجراء `addMessage()` الذي ينشئ الإجراء لإضافة رسالة. ستحتاج إلى تمرير `message` إلى منشئ الإجراء هذا وتضمين الرسالة في `action` المُعاد.

ثم أنشئ مخفّضًا يسمى `messageReducer()` يتعامل مع الحالة الخاصة بالرسائل. يجب أن تكون الحالة الابتدائية مصفوفة فارغة. يجب أن يضيف هذا المخفّض رسالة إلى مصفوفة الرسائل المحتفظ بها في الحالة، أو يعيد الحالة الحالية. أخيرًا، أنشئ مخزن Redux ومرره المخفّض.

# --hints--

يجب أن يوجد الثابت `ADD` ويحمل قيمة تساوي السلسلة النصية `ADD`

```js
assert(ADD === 'ADD');
```

يجب أن يعيد منشئ الإجراء `addMessage` كائنًا يحتوي على `type` يساوي `ADD` و`message` تساوي الرسالة التي تم تمريرها.

```js
assert(
  (function () {
    const addAction = addMessage('__TEST__MESSAGE__');
    return addAction.type === ADD && addAction.message === '__TEST__MESSAGE__';
  })()
);
```

يجب أن يكون `messageReducer` دالة.

```js
assert(typeof messageReducer === 'function');
```

يجب أن يوجد المخزن وأن تكون الحالة الابتدائية مصفوفة فارغة.

```js
assert(
  (function () {
    const initialState = store.getState();
    return typeof store === 'object' && initialState.length === 0;
  })()
);
```

يجب أن يضيف استدعاء `addMessage` إلى المخزن رسالة جديدة إلى مصفوفة الرسائل المحتفظ بها في الحالة بطريقة غير قابلة للتغيير.

```js
assert(
  (function () {
    const initialState = store.getState();
    const isFrozen = DeepFreeze(initialState);
    store.dispatch(addMessage('__A__TEST__MESSAGE'));
    const addState = store.getState();
    return isFrozen && addState[0] === '__A__TEST__MESSAGE';
  })()
);
```

يجب أن يعيد `messageReducer` الحالة الحالية إذا تم استدعاؤه مع أي إجراءات أخرى.

```js
assert(
  (function () {
    const addState = store.getState();
    store.dispatch({ type: 'FAKE_ACTION' });
    const testState = store.getState();
    return addState === testState;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Define ADD, addMessage(), messageReducer(), and store here:
```

# --solutions--

```jsx
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);
```
