---
id: 5a24c314108439a4d4036141
title: البدء مع React Redux
challengeType: 6
forumTopicId: 301430
dashedName: getting-started-with-react-redux
---

# --description--

تُقدّم هذه السلسلة من التحديات كيفية استخدام Redux مع React. أولًا، إليك مراجعة لبعض المبادئ الأساسية لكل تقنية. React هي مكتبة عرض تزودها بالبيانات، ثم تعرض العرض بطريقة فعالة ومتوقعة. Redux هو إطار عمل لإدارة الحالة يمكنك استخدامه لتبسيط إدارة حالة تطبيقك. عادةً، في تطبيق React Redux، تنشئ مخزن Redux واحد يدير حالة التطبيق بأكمله. تشترك مكونات React فقط في أجزاء البيانات في المخزن التي تتعلق بدورها. ثم تُرسل الإجراءات مباشرة من مكونات React، مما يؤدي إلى تحديثات في المخزن.

على الرغم من أن مكونات React يمكنها إدارة حالتها محليًا، عندما يكون لديك تطبيق معقد، من الأفضل عمومًا الاحتفاظ بحالة التطبيق في مكان واحد باستخدام Redux. هناك استثناءات عندما يكون للمكونات الفردية حالة محلية خاصة بها فقط. أخيرًا، لأن Redux غير مصمم للعمل مع React مباشرة، تحتاج إلى استخدام حزمة `react-redux`. توفر هذه الحزمة طريقة لتمرير `state` و`dispatch` من Redux إلى مكونات React كـ `props`.

في التحديات القادمة، ستنشئ أولًا مكون React بسيط يسمح لك بإدخال رسائل نصية جديدة. تُضاف هذه الرسائل إلى مصفوفة تُعرض في العرض. يجب أن تكون هذه مراجعة جيدة لما تعلمته في دروس React. بعد ذلك، ستنشئ مخزن Redux وإجراءات تدير حالة مصفوفة الرسائل. أخيرًا، ستستخدم `react-redux` لربط مخزن Redux بمكونك، مما ينقل الحالة المحلية إلى مخزن Redux.

# --instructions--

ابدأ بمكون `DisplayMessages`. أضف دالة بناء لهذا المكون وابدأ حالته بخصائص اثنتين: `input` التي تُعيّن إلى سلسلة نصية فارغة، و`messages` التي تُعيّن إلى مصفوفة فارغة.

# --hints--

يجب أن يعرض مكون `DisplayMessages` عنصر `div` فارغ.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(DisplayMessages));
    return mockedComponent.find('div').text() === '';
  })()
);
```

يجب استدعاء دالة البناء `DisplayMessages` بشكل صحيح باستخدام `super`، مع تمرير `props`.

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        noWhiteSpace.includes('constructor(props)') &&
        noWhiteSpace.includes('super(props')
      );
    })()
  );
```

يجب أن يكون للمكون `DisplayMessages` حالة أولية تساوي `{input: "", messages: []}`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(DisplayMessages));
    const initialState = mockedComponent.state();
    return (
      typeof initialState === 'object' &&
      initialState.input === '' &&
      Array.isArray(initialState.messages) &&
      initialState.messages.length === 0
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class DisplayMessages extends React.Component {
  // Change code below this line

  // Change code above this line
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
  }
  render() {
    return <div/>
  }
};
```
