---
id: 5a24c314108439a4d4036144
title: استخدام Provider لربط Redux بـ React
challengeType: 6
forumTopicId: 301435
dashedName: use-provider-to-connect-redux-to-react
---

# --description--

في التحدٍّ السابق، أنشأت مخزن Redux للتعامل مع مصفوفة الرسائل وأنشأت إجراءً لإضافة رسائل جديدة. الخطوة التالية هي توفير وصول React إلى مخزن Redux والإجراءات التي يحتاجها لإرسال التحديثات. يوفر React Redux حزمه `react-redux` لمساعدتك في إنجاز هذه المهام.

يوفر React Redux واجهة برمجة تطبيقات صغيرة بميزتين رئيسيتين: `Provider` و `connect`. يغطي تحدٍّ آخر `connect`. `Provider` هو مكوِّن غلاف من React Redux يلف تطبيق React الخاص بك. يسمح هذا الغلاف بعد ذلك بالوصول إلى `store` و `dispatch` الخاصين بـ Redux عبر شجرة المكونات. يأخذ `Provider` خاصيتين، مخزن Redux والمكونات الفرعية لتطبيقك. قد يبدو تعريف `Provider` لمكوِّن App كما يلي:

```jsx
<Provider store={store}>
  <App/>
</Provider>
```

# --instructions--

يعرض محرر الشفرة الآن كل شفرتك الخاصة بـ Redux و React من التحديات السابقة. يشمل مخزن Redux، الإجراءات، ومكوِّن `DisplayMessages`. القطعة الجديدة الوحيدة هي مكوِّن `AppWrapper` في الأسفل. استخدم هذا المكوِّن الأعلى لعرض `Provider` من `ReactRedux`، ومرر مخزن Redux كخاصية. ثم اعرض مكوِّن `DisplayMessages` كمكوِّن فرعي. بعد الانتهاء، يجب أن ترى مكوِّن React معروضًا على الصفحة.

**ملاحظة:** React Redux متوفر كمتغير عام هنا، لذا يمكنك الوصول إلى Provider باستخدام ترميز النقطة. تستفيد الشفرة في المحرر من هذا وتعيّن `Provider` إلى ثابت لاستخدامه في طريقة العرض لمكوِّن `AppWrapper`.

# --hints--

يجب أن يعرض مكوِّن `AppWrapper`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('AppWrapper').length === 1;
  })()
);
```

يجب أن يحتوي مكوِّن الغلاف `Provider` على خاصية `store` تُمرر إليه، تساوي مخزن Redux.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
      return __helpers
        .removeWhiteSpace(code)
        .includes('<Providerstore={store}>');
    })()
  );
```

يجب أن يعرض `DisplayMessages` كمكوِّن فرعي لـ `AppWrapper`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('AppWrapper').find('DisplayMessages').length === 1
    );
  })()
);
```

يجب أن يعرض مكوِّن `DisplayMessages` عناصر `h2`، و`input`، و`button`، و`ul`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h2').length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('ul').length === 1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};



const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
    this.handleChange = this.handleChange.bind(this);
    this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {  
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Render the Provider below this line

  // Change code above this line
};
```

# --solutions--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
 this.handleChange = this.handleChange.bind(this);
 this.submitMessage = this.submitMessage.bind(this);
 }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };  
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Change code below this line
  render() {
    return (
      <Provider store = {store}>
        <DisplayMessages/>
      </Provider>
    );
  }
  // Change code above this line
};
```
