---
id: 5a24c314108439a4d4036145
title: تعيين الحالة إلى الخصائص
challengeType: 6
forumTopicId: 301433
dashedName: map-state-to-props
---

# --description--

يتيح لك مكوِّن `Provider` توفير `state` و`dispatch` لمكوِّنات React الخاصة بك، لكن يجب أن تحدد بالضبط الحالة والإجراءات التي تريدها. بهذه الطريقة، تضمن أن كل مكوِّن يصل فقط إلى الحالة التي يحتاجها. تحقق ذلك بإنشاء دالتين: `mapStateToProps()` و`mapDispatchToProps()`.

في هاتين الدالتين، تعلن عن أجزاء الحالة التي تريد الوصول إليها وأي منشئي إجراءات تحتاج إلى القدرة على تنفيذها. بمجرد وجود هذه الدوال، سترى كيف تستخدم طريقة `connect` من React Redux لربطها بمكوناتك في تحدٍّ آخر.

**ملاحظة:** خلف الكواليس، يستخدم React Redux طريقة `store.subscribe()` لتنفيذ `mapStateToProps()`.

# --instructions--

أنشئ دالة `mapStateToProps()`. يجب أن تأخذ هذه الدالة `state` كمعلمة، ثم تُرجع كائنًا يربط تلك الحالة بأسماء خصائص محددة. ستصبح هذه الخصائص متاحة لمكوِّنك عبر `props`. بما أن هذا المثال يحتفظ بالحالة الكاملة للتطبيق في مصفوفة واحدة، يمكنك تمرير تلك الحالة كاملة إلى مكوِّنك. أنشئ خاصية `messages` في الكائن المُرجع، وعيّنها إلى `state`.

# --hints--

يجب أن تكون القيمة الثابتة `state` مصفوفة فارغة.

```js
assert(Array.isArray(state) && state.length === 0);
```

يجب أن تكون `mapStateToProps` دالة.

```js
assert(typeof mapStateToProps === 'function');
```

يجب أن تُرجع `mapStateToProps` كائنًا.

```js
assert(typeof mapStateToProps() === 'object');
```

تمرير مصفوفة كحالة إلى `mapStateToProps` يجب أن يُرجع هذه المصفوفة مُعيَّنة إلى مفتاح `messages`.

```js
assert(mapStateToProps(['messages']).messages.pop() === 'messages');
```

# --seed--

## --seed-contents--

```jsx
const state = [];

// Change code below this line
```

# --solutions--

```jsx
const state = [];

// Change code below this line

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};
```
