---
id: 5a24c314108439a4d4036147
title: ربط Redux بـ React
challengeType: 6
forumTopicId: 301426
dashedName: connect-redux-to-react
---

# --description--

الآن بعد أن كتبت كل من الدالتين `mapStateToProps()` و `mapDispatchToProps()`، يمكنك استخدامهما لتعيين `state` و `dispatch` إلى الخصائص `props` لأحد مكونات React الخاصة بك. يمكن لطريقة `connect` من React Redux التعامل مع هذه المهمة. تأخذ هذه الطريقة معلمتين اختياريتين، `mapStateToProps()` و `mapDispatchToProps()`. هما اختياريتان لأنك قد تملك مكوّنًا يحتاج فقط إلى وصول إلى `state` دون الحاجة إلى تنفيذ أي إجراءات، أو العكس.

لاستخدام هذه الطريقة، مرر الدالتين كمعلمات، ثم استدعِ النتيجة فورًا على مكونك. هذه الصيغة غير معتادة قليلًا وتبدو كالتالي:

```js
connect(mapStateToProps, mapDispatchToProps)(MyComponent)
```

**ملاحظة:** إذا أردت حذف إحدى المعلمات لطريقة `connect`، مرر `null` بدلًا منها.

# --instructions--

يحتوي محرر الشفرة على الدالتين `mapStateToProps()` و`mapDispatchToProps()` ومكوّن React جديد يُسمى `Presentational`. اربط هذا المكوّن بـ Redux باستخدام الطريقة `connect` من الكائن العام `ReactRedux`، واستدعِها مباشرةً على المكوّن `Presentational`. أسنِد النتيجة إلى ثابت `const` جديد يُسمى `ConnectedComponent` يمثّل المكوّن المتصل. هذا كل شيء، أنت الآن متصل بـ Redux! جرّب تغيير أيٍّ من وسيطتي `connect` إلى `null` ولاحظ نتائج الاختبارات.

# --hints--

يجب أن يعرض مكوّن `Presentational`.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('Presentational').length === 1;
  })()
);
```

يجب أن يستقبل مكوّن `Presentational` خاصية `messages` عبر `connect`.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return props.messages === '__INITIAL__STATE__';
  })()
);
```

يجب أن يستقبل مكوّن `Presentational` خاصية `submitNewMessage` عبر `connect`.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return typeof props.submitNewMessage === 'function';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line

const ConnectedComponent = connect(mapStateToProps, mapDispatchToProps)(Presentational);
```
