---
id: 5a24c314108439a4d403614a
title: التقدم من هنا
challengeType: 6
forumTopicId: 301434
dashedName: moving-forward-from-here
---

# --description--

تهانينا! أنهيت دروس React و Redux. هناك نقطة أخيرة تستحق الذكر قبل المتابعة. عادةً، لن تكتب تطبيقات React في محرر شفرة مثل هذا. هذا التحدٍّ يمنحك لمحة عن شكل الصيغة النحوية إذا كنت تعمل مع نظام ملفات على جهازك الخاص. يجب أن تبدو الشفرة مشابهة، باستثناء استخدام عبارات `import` (هذه تستورد كل التبعيات التي تم توفيرها لك في التحديات).

أخيرًا، كتابة شفرة React و Redux تتطلب عادةً بعض الإعدادات. قد تصبح معقدة بسرعة. إذا كنت مهتمًا بالتجربة على جهازك الخاص، فإن <a href="https://www.freecodecamp.org/news/install-react-with-create-react-app/" target="_blank" rel="noopener noreferrer nofollow">Create React App</a> يأتي مُعدًا وجاهزًا للاستخدام.

بدلًا من ذلك، يمكنك تفعيل Babel كمُعالج مسبق للغة JavaScript في CodePen، وإضافة React و ReactDOM كمصادر JavaScript خارجية، والعمل هناك أيضًا.

# --instructions--

سجل الرسالة `'Now I know React and Redux!'` في وحدة التحكم.

# --hints--

يجب تسجيل الرسالة `Now I know React and Redux!` في وحدة التحكم.

```js
assert(
    /console\s*\.\s*log\s*\(\s*('|"|`)Now I know React and Redux!\1\s*\)/.test(
      code)
  );
```

# --seed--

## --seed-contents--

```jsx
/*
import React from 'react'
import ReactDOM from 'react-dom'
import { Provider, connect } from 'react-redux'
import { createStore, combineReducers, applyMiddleware } from 'redux'
import thunk from 'redux-thunk'

import rootReducer from './redux/reducers'
import App from './components/App'

const store = createStore(
  rootReducer,
  applyMiddleware(thunk)
);

ReactDOM.render(
  <Provider store={store}>
    <App/>
  </Provider>,
  document.getElementById('root')
);
*/

// Only change code below this line
```

# --solutions--

```jsx
console.log('Now I know React and Redux!');
```
