---
id: 5a24c314108439a4d4036146
title: ربط الاستدعاء بالخصائص
challengeType: 6
forumTopicId: 301432
dashedName: map-dispatch-to-props
---

# --description--

تُستخدم الدالة `mapDispatchToProps()` لتوفير منشئي الإجراءات المحددين لمكوِّنات React الخاصة بك حتى تتمكن من استدعاء الإجراءات على مخزن Redux. هي مشابهة في البنية للدالة `mapStateToProps()` التي كتبتها في التحدي السابق. تُرجع كائنًا يربط إجراءات الاستدعاء بأسماء الخصائص، التي تصبح خصائص `props` للمكوِّن. ولكن بدلًا من إرجاع جزء من الـ `state`، تُرجع كل خاصية دالة تستدعي `dispatch` مع منشئ الإجراء وأي بيانات إجراء ذات صلة. لديك وصول إلى هذا `dispatch` لأنه يُمرر إلى `mapDispatchToProps()` كمعلمة عند تعريف الدالة، تمامًا كما مررت `state` إلى `mapStateToProps()`. خلف الكواليس، يستخدم React Redux دالة `store.dispatch()` الخاصة بـ Redux لتنفيذ هذه الاستدعاءات عبر `mapDispatchToProps()`. وهذا مشابه لكيفية استخدامه `store.subscribe()` للمكوِّنات المرتبطة بـ `state`.

على سبيل المثال، لديك منشئ إجراء `loginUser()` يأخذ `username` كحمولة إجراء. الكائن الذي تُرجعه `mapDispatchToProps()` لهذا المنشئ سيكون مشابهًا لما يلي:

```jsx
{
  submitLoginUser: function(username) {
    dispatch(loginUser(username));
  }
}
```

# --instructions--

يوفر محرر الشفرة منشئ إجراء يسمى `addMessage()`. اكتب الدالة `mapDispatchToProps()` التي تأخذ `dispatch` كمعلمة، ثم تُرجع كائنًا. يجب أن يحتوي الكائن على خاصية `submitNewMessage` مرتبطة بدالة الاستدعاء، التي تأخذ معلمة للرسالة الجديدة التي تُضاف عند استدعاء `addMessage()`.

# --hints--

يجب أن تُرجع `addMessage` كائنًا يحتوي على المفاتيح `type` و `message`.

```js
assert(
  (function () {
    const addMessageTest = addMessage();
    return (
      addMessageTest.hasOwnProperty('type') &&
      addMessageTest.hasOwnProperty('message')
    );
  })()
);
```

يجب أن تكون `mapDispatchToProps` دالة.

```js
assert(typeof mapDispatchToProps === 'function');
```

يجب أن تُرجع `mapDispatchToProps` كائنًا.

```js
assert(typeof mapDispatchToProps() === 'object');
```

يجب أن يُرجع استدعاء `addMessage` مع `submitNewMessage` من `mapDispatchToProps` رسالة إلى دالة الاستدعاء.

```js
assert(
  (function () {
    let testAction;
    const dispatch = (fn) => {
      testAction = fn;
    };
    let dispatchFn = mapDispatchToProps(dispatch);
    dispatchFn.submitNewMessage('__TEST__MESSAGE__');
    return (
      testAction.type === 'ADD' && testAction.message === '__TEST__MESSAGE__'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

// Change code below this line

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: function(message) {
      dispatch(addMessage(message));
    }
  }
};
```
