---
id: 615f3b091162165948e1cb33
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Якщо ви перевірите свій `.label` елемент за допомогою інструментів розробника у браузері, ви можете помітити, що його ширина насправді становить 288 пікселів замість 270. Це тому, що за замовчуванням браузер враховує кордон і відступ при визначенні розміру елемента.

Щоб це виправити, скиньте модель блоку, створивши селектор `*` і задавши йому властивість `box-sizing` зі значенням `border-box`.

# --hints--

Ви повинні створити селектор `*`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

Ваш селектор `*` має мати властивість `box-sizing`, встановлену в `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
