---
id: 685239f0cad43c7a909cbd98
title: Налагодьте сторінку притулку для тварин
challengeType: 25
dashedName: lab-debug-pet-adoption-page
---

# --description--

Саллі — власниця притулку для тварин — вперше створила вебсторінку, але зіткнулась з помилками.

Ваше завдання — розв’язати помилки, щоб Саллі продовжила працювати.

Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Саллі хоче використати фото декількох котів, але вони не відтворюються правильно. В елементі `img` потрібно виправити наступне:
   - замініть атрибут `href` на атрибут, який вказує джерело зображення;
   - замініть атрибут `att` на атрибут, який коротко описує зображення;
   - видаліть кінцевий тег `</img>`, оскільки елементи `img` є порожніми і не потребують кінцевих тегів.
2. Саллі хоче використати посилання, щоб направити користувачів на сторінки з собаками чи котами. Але посилання не працюють. В елементах `a` потрібно виправити наступне:
   - замініть обидва атрибути `src` на атрибути, які вказують посилання.

# --hints--

Елемент `img` повинен мати атрибут `src`, а не `href`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

Елемент `img` повинен мати атрибут `alt`, а не `att`, якого не існує.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

Елемент `img` не повинен мати кінцевий тег `</img>`.

```js
assert.notMatch(code, /<\/img>/);
```

Елемент `a` з текстом `Visit cats page` повинен мати атрибут `href`, а не `src`.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

Елемент `a` з текстом `Visit dogs page` повинен мати атрибут `href`, а не `src`.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Five cats looking around a field."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
