---
id: 66db57ad34c7089b9b41bfd6
title: Construir un visor de Lightbox
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

Un lightbox se utiliza en sitios web para mostrar contenido multimedia de una manera más atractiva a través del uso de una ventana emergente o modal sobre el contenido principal de la página.

En este laboratorio, crearás una galería de lightbox que muestra imágenes de tamaño completo cuando se hace clic en una miniatura. Para cada imagen, se proporcionan dos versiones: una miniatura y una imagen de tamaño completo. La imagen de tamaño completo es la misma que la miniatura, pero sin el sufijo  ` -thumbnail `.

Cumple con las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes tener un ` div ` con una clase de  ` gallery ` dentro de tu ` body `.
1. Dentro del elemento  ` .gallery `, debes tener tres miniaturas de imagen, cada una con una clase de  ` gallery-item `. Debes usar los siguientes enlaces para imágenes en miniatura:

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. Debes tener un ` div ` con una clase de  ` lightbox ` dentro de tu ` body `.
1. Debes tener un ` botón ` con un ` id ` de  ` close-btn ` dentro de tu elemento  ` .lightbox `. Puedes usar ` &times; ` como su texto si lo deseas.
1. Debes tener un ` img ` con un ` id ` de  ` lightbox-image ` dentro de tu elemento  ` .lightbox `.
1. Tu elemento  ` .lightbox ` debe tener una posición fija para que la vista previa se abra encima de las imágenes actuales.
1. Tu elemento  ` .lightbox ` debe cubrir toda la ventana visual ajustando la altura y el ancho al 100% del contenedor. Deberías asegurarte de que el elemento  ` .lightbox ` comience en la esquina superior izquierda del contenedor.
1. ` .lightbox ` debería tener un color de fondo. Inicialmente, su propiedad  ` display ` debería estar configurada en  ` none ` para ocultarlo.
1. Cuando hagas clic en uno de tus elementos  ` .gallery-item `, la propiedad  ` display ` del elemento  ` .lightbox ` se debe configurar en  ` flex ` para hacer visible el elemento  ` .lightbox ` y los dos elementos dentro de él.
1. Cuando haces clic en uno de tus elementos  ` .gallery-item `, el atributo  ` src ` del elemento  ` #lightbox-image ` debe configurarse en una versión de tamaño completo de la imagen clicada eliminando  ` -thumbnail ` del atributo  ` src ` de la imagen. Las imágenes de tamaño completo están ubicadas en los siguientes enlaces:

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`

1. Cuando tu elemento  ` .lightbox ` esté visible y hagas clic en el  ` #close-btn ` o el elemento  ` .lightbox `, el  ` display ` del elemento  ` .lightbox ` se debe ajustar nuevamente a  ` none `.

Nota: Asegúrate de enlazar tu hoja de estilos y el archivo JavaScript en tu HTML.

# --hints--

Debes tener un ` div ` con la clase de  ` gallery ` dentro de tu elemento  ` body `.

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

Dentro del elemento  ` .gallery `, debes tener tres elementos  ` img ` con la clase de  ` gallery-item `.

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

Dentro del elemento  ` .gallery `, debes tener un elemento  ` img ` con el  ` src ` configurado en  ` https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg `.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Dentro del elemento  ` .gallery `, debes tener un elemento  ` img ` con el  ` src ` configurado en  ` https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg `.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Dentro del elemento  ` .gallery `, debes tener un elemento  ` img ` con el  ` src ` configurado en  ` https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg `.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Debes tener un elemento  ` div ` con la clase de  ` lightbox ` dentro de tu elemento  ` body `.

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

Dentro de tu elemento  ` .lightbox `, debes tener un elemento  ` botón ` con el  ` id ` configurado en  ` close-btn `.

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

Dentro de tu elemento  ` .lightbox `, debes tener un elemento  ` img ` con el  ` id ` configurado en  ` lightbox-image `.

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

Tu elemento  ` .lightbox ` debería tener un posicionamiento fijo.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

Tu elemento  ` .lightbox ` debería cubrir toda la ventana gráfica.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.width, "100%");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.height, "100%");
```

Tu elemento  ` .lightbox ` debería estar alineado con la esquina superior izquierda del contenedor.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

Tu elemento  ` .lightbox ` debería tener un color de fondo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

Tu elemento  ` .lightbox ` debería estar oculto inicialmente.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

Cuando hagas clic en uno de tus elementos ` .gallery-item `, la propiedad ` display ` del elemento ` .lightbox ` se debe configurar en ` flex ` para hacer que ` .lightbox ` y los dos elementos dentro de él sean visibles.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

Cuando hagas clic en uno de tus elementos ` .gallery-item `, el atributo ` src ` del elemento ` #lightbox-image ` debe configurarse en una versión de tamaño completo de la imagen clicada eliminando ` -thumbnail ` del atributo ` src ` de la imagen.

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

Cuando tu elemento ` .lightbox ` esté visible y hagas clic en el botón ` #close-btn `, el ` display ` del elemento ` .lightbox ` se debe ajustar nuevamente a ` none `.

```js
const lightbox = document.querySelector(".lightbox");
const background = document.getElementById("close-btn"); 

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

lightbox.style.display = "flex";

background.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

Cuando tu elemento ` .lightbox ` esté visible y hagas clic en el elemento ` .lightbox `, el ` display ` del elemento ` .lightbox ` se debe ajustar nuevamente a ` none `.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

lightbox.style.display = "flex";

lightbox.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```
