---
id: 61438ec09438696391076d6a
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

`.heading` 요소 아래에 `section`가 `class`로 설정된 새로운 `text` 요소를 만드세요. 그 안에 `p`가 `class`로 설정된 `first-paragraph` 요소와 다음 텍스트를 만드세요:

```markup
Soon the freeCodeCamp curriculum will be 100% project-driven learning. Instead of a series of coding challenges, you'll learn through building projects - step by step. Before we get into the details, let me emphasize: we are not changing the certifications. All 6 certifications will still have the same 5 required projects. We are only changing the optional coding challenges.
```

# --hints--

새로운 `section` 요소를 만들어야 합니다.

```js
assert(document.querySelectorAll('section')?.length === 2);
```

새로운 `section` 요소는 `.heading` 요소 뒤에 와야 합니다.

```js
assert(document.querySelectorAll('section')?.[1]?.previousElementSibling?.className === 'heading');
```

새로운 `section` 요소는 `class`가 `text`로 설정되어야 합니다.

```js
assert(document.querySelectorAll('section')?.[1]?.className === 'text');
```

`p` 요소 안에 새로운 `.text` 요소를 만들어야 합니다.

```js
assert(document.querySelector('.text')?.querySelectorAll('p')?.length === 1);
```

새로운 `p` 요소는 `class`가 `first-paragraph`로 설정되어야 합니다.

```js
assert(document.querySelector('.text p')?.className === 'first-paragraph');
```

새로운 `p` 요소는 제공된 텍스트를 포함해야 합니다.

```js
assert(document.querySelector('.first-paragraph')?.innerText === 'Soon the freeCodeCamp curriculum will be 100% project-driven learning. Instead of a series of coding challenges, you\'ll learn through building projects - step by step. Before we get into the details, let me emphasize: we are not changing the certifications. All 6 certifications will still have the same 5 required projects. We are only changing the optional coding challenges.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/freecodecamp/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/freecodecamp">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/freecodecamp">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
