---
id: 61437d575fb98f57fa8f7f36
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

표준 HTML 보일러플레이트로 시작하세요. `DOCTYPE` 선언, 이 페이지가 영어임을 지정하는 `html` 요소, `head` 요소, 그리고 `body` 요소를 추가하세요.

`<meta>` 요소 안에 적절한 `charset`를 가진 `<meta>` 태그와 모바일 반응형을 위한 `head` 태그를 추가하세요.

# --hints--

코드에 `DOCTYPE` 참조가 포함되어야 합니다.

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 참조 뒤에는 공백을 포함해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

문서 유형을 `html`으로 정의해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

`DOCTYPE` 선언은 타입 뒤에 `>`로 닫아야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`<html>` 여는 태그가 있어야 하며, `lang`의 `en`를 가져야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 선언은 HTML의 맨 앞에 있어야 합니다.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`<head>` 여는 태그가 있어야 합니다.

```js
assert(code.match(/<head\s*>/i));
```

`</head>` 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/head\s*>/i));
```

`<body>` 여는 태그가 있어야 합니다.

```js
assert(code.match(/<body\s*>/i));
```

`</body>` 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/body\s*>/i));
```

`head`와 `body` 요소는 형제 관계여야 합니다.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

`meta` 요소가 두 개 있어야 합니다.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

한 `meta` 요소는 `name`가 `viewport`로, `content`가 `width=device-width, initial-scale=1.0`로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

다른 `meta` 요소는 `charset` 속성이 `UTF-8`로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
