---
id: 614385513d91ae5c251c2052
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

텍스트가 `title`인 `Magazine` 요소, `link` 폰트 스타일시트를 위한 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` 요소, `link` FontAwesome 스타일시트를 위한 `https://use.fontawesome.com/releases/v5.8.2/css/all.css` 요소, 그리고 `link` 스타일시트를 위한 `./styles.css` 요소를 추가하세요.

폰트 스타일시트는 `Anton`, `Baskervville`, `Raleway` 세 가지 폰트를 불러옵니다.

# --hints--

코드에는 `link` 요소가 세 개 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

`link` 요소들은 `head` 요소 안에 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

세 개의 `link` 요소는 `rel` 속성을 가지고 그 값은 `stylesheet` 이어야 합니다.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

링크 요소 중 하나는 `href` 속성이 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` 으로 설정되어야 합니다.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

링크 요소 중 하나는 `href` 속성이 `https://use.fontawesome.com/releases/v5.8.2/css/all.css` 으로 설정되어야 합니다.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

`link` 요소 중 하나는 `href` 속성이 `styles.css` 으로 설정되어야 합니다.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

코드에는 `title` 요소가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.isNotNull(title);
```

코드에는 `title` 요소가 하나 있어야 합니다.

```js
const title = document.querySelectorAll('title');
assert.strictEqual(title?.length, 1);
```

`title` 요소는 `head` 요소 안에 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

프로젝트 제목은 `Magazine` 이어야 합니다. 제목의 대소문자와 철자가 정확해야 합니다.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
