---
id: 6165d3b702a5d92ad970b30c
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`img` 요소 다음에 `h1` 요소를 추가하고 `class`를 `hero-title`로 설정한 뒤 텍스트를 `OUR NEW CURRICULUM`로 설정하세요. 그다음 `p` 요소를 추가하고 `class`를 `hero-subtitle`로 설정한 뒤 텍스트를 `Our efforts to restructure our curriculum with a more project-based focus`로 설정하세요.

# --hints--

`h1` 요소를 생성해야 합니다.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 요소는 `img` 요소 다음에 와야 합니다.

```js
assert(document.querySelector('h1')?.previousElementSibling?.localName === 'img');
```

`h1` 요소는 `class`를 `hero-title`로 설정해야 합니다.

```js
assert(document.querySelector('h1')?.className === 'hero-title');
```

`h1` 요소는 텍스트를 `OUR NEW CURRICULUM`로 설정해야 합니다.

```js
assert(document.querySelector('h1')?.textContent === 'OUR NEW CURRICULUM');
```

새로운 `p` 요소를 생성해야 합니다.

```js
assert.exists(document.querySelector('p'));
```

`p` 요소는 `h1` 요소 다음에 와야 합니다.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

`p` 요소는 `class`를 `hero-subtitle`로 설정해야 합니다.

```js
assert(document.querySelector('p')?.className === 'hero-subtitle');
```

`p` 요소는 텍스트를 `Our efforts to restructure our curriculum with a more project-based focus`로 설정해야 합니다.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
