---
id: 6169cd8a558aa8434e0ad7f6
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`Referer` HTTP 헤더는 사용자가 방문할 수 있는 페이지의 주소나 URL에 관한 정보를 포함합니다. 이 정보는 예를 들어, 사용자가 귀하의 페이지에서 freecodecamp.org로 얼마나 많이 방문하는지 추적하는 분석에 사용될 수 있습니다. `rel` 속성을 `noreferrer`로 설정하면 이 정보가 HTTP 요청에서 생략됩니다. `a` 요소에 `rel` 속성을 `noreferrer`로 설정하세요.

# --hints--

`a` 요소에 `rel`를 `noreferrer`로 설정해야 합니다.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('rel') === 'noreferrer');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
