---
id: 698c50535b54b2d75fa985fb
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Атрибут `loop` перезапустить аудіо після завершення відтворення. Уявіть інтернет-мем, який повторює відтворення. Якщо опустити атрибут `loop`, аудіо відтвориться лише один раз.

Атрибут `loop` є булевим і не потребує значення.

Додайте атрибут `loop` до елемента `audio`.

# --hints--

Ваш елемент `audio` має містити атрибут `loop`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
