---
id: 698c4f4e965bd853ad5084e1
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Під `h1` додайте елемент `h2` з текстом `Can't Stay Down`, це назва першої пісні.

Під елементом `h2` додайте елемент `p` з текстом `Artist: Quincy Larson`.

# --hints--

У вас має бути елемент `h2`.

```js
assert.exists(document.querySelector('h2'));
```

Ваш елемент `h2` має бути під елементом `h1`.

```js
assert.exists(document.querySelector('h1 + h2'));
```

Ваш елемент `h2` має містити текст `Can't Stay Down`.

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

У вас має бути елемент `p`.

```js
assert.exists(document.querySelector('p'));
```

Ваш елемент `p` має містити текст `Artist: Quincy Larson`.

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
