---
id: 698c4e69abe14a9746c72312
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви створите HTML-плеєр для музики. Для вас уже підготовлено шаблонний код HTML.

Створіть елемент `h1` і додайте до нього текст `freeCodeCamp Tunes`.

# --hints--

У вас має бути елемент `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ваш елемент `h1` повинен містити текст `freeCodeCamp Tunes`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'freeCodeCamp Tunes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
