---
id: 698c4fe24f1f1a17be74a23e
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Далі створіть елемент `audio` під елементом `p`. Протягом наступних кількох кроків ви додасте необхідні атрибути, щоб аудіоелемент відтворював музику.

# --hints--

У вас має бути елемент `audio`.

```js
assert.exists(document.querySelector('audio'));
```

Новий елемент `audio` має бути розташований під елементом `p`.

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
