---
id: 698c58745b73289bb6c89ece
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Остання пісня має назву `Scratching the Surface`, виконавець — Quincy Larson, а файл для використання — `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

Додайте останню пісню, щоб завершити музичний плеєр, дотримуючись того самого шаблону, який ви використали для перших двох пісень.

# --hints--

У вас має бути третій елемент `h2`.

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

Ваш третій елемент `h2` має містити текст `Scratching the Surface`.

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

У вас має бути елемент `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

Ваш елемент `p` має містити текст `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

У вас має бути третій елемент `audio`.

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

Новий елемент `audio` має бути розташований нижче елемента `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

Ваш елемент `audio` має мати атрибут `src`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

Ваш елемент `audio` має мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

Ваш елемент `audio` має мати атрибут `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Ваш елемент `audio` має мати атрибут `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
