---
id: 612e89562043183c86df287c
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

瀏覽器經常會對特定元素套用預設的外邊距和內邊距值。為了確保你的鋼琴版面正確顯示，你需要重設盒模型。

為你的 CSS 檔案添加一個 `html` 規則選擇器，並將 `box-sizing` 屬性設為 `border-box`。

# --hints--

你應該有一個 `html` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

你的 `html` 選擇器應該將 `box-sizing` 屬性設為 `border-box`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
