---
id: 612ea4c4993aba52ab4aa32e
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

現在是時候使用你先前設定的偽選擇器。為了創建黑鍵，為 `.key.black--key::after` 新增一個選擇器。這會標的同時具有 `key` 和 `black--key` 類別的元素，並選擇在 HTML 中這些元素之後生成的偽元素。

在新的選擇器中，將 `background-color` 設定為 `#1d1e22`。同時將 `content` 屬性設定為 `""`。這將使偽元素為空。

`content` 屬性用於設定或改寫元素的內容。預設情況下，由 `::before` 和 `::after` 選擇器產生的偽元素是空的，這表示它們不會在頁面上呈現。將 `content` 屬性設定為空字串 `""`，你可以確保偽元素被呈現，同時仍然看起來是空的。

如果你想要嘗試，請移除 `background-color` 屬性並為 `content` 屬性設定不同的值，例如 `"♥"`。完成後請記得還原這些更改，以便測試通過。

# --hints--

你應該有一個 `.key.black--key::after` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

你的 `.key.black--key::after` 選擇器應該有一個 `background-color` 屬性設為 `#1d1e22`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor, 'rgb(29, 30, 34)');
```

你的 `.key.black--key::after` 選擇器應該有一個 `content` 屬性設定為 `""`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content, '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
