---
id: 612e8279827a28352ce83a72
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

現在，複製七個 `.key` 元素的集合，並將另外兩組貼到 `.keys` div 中。

# --hints--

你應該有 21 個 `.key` 元素。

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

你應該有 15 個 `.black--key` 元素。

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

所有 `.key` 元素應該位於你的 `.keys` 元素內。

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

所有 `.black--key` 元素應該位於你的 `.keys` 元素內。

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
