---
id: 612e813b3ba67633222cbe54
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

請記住，`class` 屬性可以包含多個值。為了區分你的白鍵和黑鍵，將第二個 `class` 值 `black--key` 添加到你的第二、第三、第五、第六和第七個 `.key` 元素。

# --hints--

你的第二個 `.key` 元素也應該有 `class` 為 `black--key`。

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

你的第三個 `.key` 應該在 `class` 中有 `black--key`。

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

你的第五個 `.key` 應該在 `class` 中有 `black--key`。

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

你的第六個 `.key` 應該在 `class` 中有 `black--key`。

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

你的第七個 `.key` 應該在 `class` 中有 `black--key`。

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

你應該有五個 `.black--key` 元素。

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

你應該有七個 `.key` 元素。

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
