---
id: 612ebcba99bfa46a15370b11
title: 第 25 步
challengeType: 0
dashedName: step-25
---

# --description--

`@media` at-rule，也稱為媒體查詢，用於有條件地套用 CSS。媒體查詢通常用於根據視窗寬度使用 `max-width` 和 `min-width` 屬性來套用 CSS。

在以下範例中，當視窗寬度為 `960px` 或以下時，會將內距套用到 `.card` 類別。

```css
@media (max-width: 960px) {
  .card {
    padding: 2rem;
  }
}
```

新增一個媒體查詢，當視窗寬度為 `768px` 或以下時套用。

# --hints--

你應該為你的程式碼添加一個新的 `@media` 查詢。

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('media'), 1);
```

你的 `@media` 查詢應該有 `max-width` 為 `768px`。

```js
assert.equal(new __helpers.CSSHelp(document).getCSSRules('media')[0]?.media?.mediaText, '(max-width: 768px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <img class="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg" alt="freeCodeCamp Logo" />
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
  position: relative;
  border-radius: 10px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
  border-radius: 0 0 3px 3px;
}

.key.black--key::after {
  background-color: #1d1e22;
  content: "";
  position: absolute;
  left: -18px;
  width: 32px;
  height: 100px;
  border-radius: 0 0 3px 3px;
}

.logo {
  width: 200px;
  position: absolute;
  top: 23px;
}

--fcc-editable-region--

--fcc-editable-region--
```
