---
id: 587d825b367417b2b2512c8d
title: Crea una mappa JavaScript ES6
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

La nuova versione di JavaScript ci offre un oggetto Map integrato che fornisce gran parte delle funzionalità che abbiamo scritto a mano nella sfida precedente. Questo oggetto Map, sebbene simile ai normali oggetti JavaScript, offre alcune funzionalità utili che gli oggetti normali non hanno. Per esempio, una Map ES6 tiene traccia dell’ordine di inserimento degli elementi aggiunti. Ecco una panoramica più completa dei suoi metodi:

- `.has(key)` restituisce true o false in base alla presenza di una chiave
- `.get(key)` restituisce il valore associato a una chiave
- `.set(key, value)` imposta una nuova coppia chiave-valore
- `.delete(key)` rimuove una coppia chiave-valore
- `.clear()` rimuove tutte le coppie chiave-valore
- `.keys()` restituisce un nuovo iteratore mappa che contiene tutte le chiavi nell’ordine di inserimento
- `.values()` restituisce un nuovo iteratore mappa che contiene tutti i valori nell’ordine di inserimento
- `.entries()` restituisce un nuovo iteratore mappa che contiene tutte le coppie chiave-valore nell’ordine di inserimento

# --instructions--

Definisci un oggetto Map JavaScript e assegnalo a una variabile chiamata myMap. Aggiungi la coppia chiave-valore `freeCodeCamp`, `Awesome!`.

# --hints--

L’oggetto `myMap` dovrebbe esistere.

```js
assert(typeof myMap === 'object');
```

`myMap` dovrebbe contenere la coppia chiave-valore `freeCodeCamp`, `Awesome!`.

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
