---
id: 587d8250367417b2b2512c60
title: Creare una classe queue
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--

Come gli stack, le queue sono una raccolta di elementi. Ma a differenza degli stack, le queue seguono il principio FIFO (First-In First-Out). Gli elementi aggiunti a una queue vengono inseriti in coda, o alla fine, della queue, e solo l'elemento in testa alla queue può essere rimosso.

Potremmo usare un array per rappresentare una queue, ma proprio come per gli stack, vogliamo limitare il controllo che abbiamo sulle nostre queue.

I due metodi principali di una classe queue sono il metodo enqueue e il metodo dequeue. Il metodo enqueue inserisce un elemento in coda alla queue, e il metodo dequeue rimuove e restituisce l'elemento in testa alla queue. Altri metodi utili sono i metodi front, size e isEmpty.

# --instructions--

Scrivi un metodo `enqueue` che inserisca un elemento in coda alla queue, un metodo `dequeue` che rimuova e restituisca l'elemento in testa, un metodo `front` che ci permetta di vedere l'elemento in testa, un metodo `size` che mostri la lunghezza, e un metodo `isEmpty` per verificare se la queue è vuota.

# --hints--

La tua classe `Queue` dovrebbe avere un metodo `enqueue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

La tua classe `Queue` dovrebbe avere un metodo `dequeue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

La tua classe `Queue` dovrebbe avere un metodo `front`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

La tua classe `Queue` dovrebbe avere un metodo `size`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

La tua classe `Queue` dovrebbe avere un metodo `isEmpty`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Il metodo `dequeue` dovrebbe rimuovere e restituire l'elemento in testa alla queue

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

Il metodo `front` dovrebbe restituire il valore dell'elemento in testa alla queue

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

Il metodo `size` dovrebbe restituire la lunghezza della queue

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

Il metodo `isEmpty` dovrebbe restituire `false` se ci sono elementi nella queue

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```
