---
id: 587d8255367417b2b2512c75
title: Creare una coda circolare
challengeType: 1
forumTopicId: 301625
dashedName: create-a-circular-queue
---

# --description--

In questa sfida creerai una coda circolare. Una coda circolare è una coda che scrive alla fine di una collezione e poi inizia a sovrascrivere se stessa all'inizio della collezione. Questo tipo di struttura dati è utile in certe situazioni. Per esempio, una coda circolare può essere usata per lo streaming di media. Una volta che la coda è piena, i nuovi dati media sovrascriveranno i dati vecchi.

Un buon modo per illustrare questo concetto è con un array di lunghezza `5`:

```js
[null, null, null, null, null]
 ^Read @ 0
 ^Write @ 0
```

Qui il puntatore di lettura e quello di scrittura sono entrambi in posizione `0`. Ora la coda riceve 3 nuovi record `a`, `b` e `c`. La nostra coda ora appare così:

```js
[a, b, c, null, null]
 ^Read @ 0
          ^Write @ 3
```

Man mano che il puntatore di lettura legge, può rimuovere i valori o mantenerli:

```js
[null, null, null, null, null]
                   ^Read @ 3
                   ^Write @ 3
```

Ora scriviamo i valori `d`, `e` e `f` nella coda. Una volta che il puntatore di scrittura raggiunge la fine dell'array, torna all'inizio:

```js
[f, null, null, d, e]
                ^Read @ 3
    ^Write @ 1
```

Questo approccio richiede una quantità costante di memoria ma permette di processare file di dimensioni molto più grandi.

# --instructions--

In questa sfida implementeremo una coda circolare. La coda circolare dovrebbe fornire i metodi `enqueue` e `dequeue` che ti permettono di leggere dalla coda e scrivere nella coda. La classe stessa dovrebbe anche accettare un argomento intero che puoi usare per specificare la dimensione della coda al momento della creazione. Abbiamo scritto la versione iniziale di questa classe per te nell'editor di codice.

Quando inserisci elementi nella coda, il puntatore di scrittura dovrebbe avanzare in avanti e tornare all'inizio una volta raggiunta la fine della coda. Il metodo `enqueue` dovrebbe restituire l'elemento inserito se ha successo; altrimenti restituirà `null`.

Allo stesso modo, il puntatore di lettura dovrebbe avanzare in avanti mentre estrai elementi dalla coda. Quando estrai un elemento, quell'elemento dovrebbe essere restituito. Se non puoi estrarre un elemento, dovresti restituire `null`.

Il puntatore di scrittura non dovrebbe poter superare il puntatore di lettura (la nostra classe non ti permetterà di sovrascrivere dati che non hai ancora letto) e il puntatore di lettura non dovrebbe poter avanzare oltre i dati che hai scritto.

# --hints--

Il metodo `enqueue` dovrebbe aggiungere elementi alla coda circolare.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Non dovresti inserire elementi oltre il puntatore di lettura.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    test.enqueue(13);
    test.enqueue(25);
    test.enqueue(59);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Il metodo `dequeue` dovrebbe estrarre elementi dalla coda.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 && test.dequeue() === 32 && test.dequeue() === 591
    );
  })()
);
```

Dopo che un elemento è stato estratto, la sua posizione nella coda dovrebbe essere resettata a `null`.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(672);
    test.dequeue();
    test.dequeue();
    var print = test.print();
    return print[0] === null && print[1] === null && print[2] === 672;
  })()
);
```

Tentare di estrarre oltre il puntatore di scrittura dovrebbe restituire `null` e non avanzare il puntatore di scrittura.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 &&
      test.dequeue() === 32 &&
      test.dequeue() === 591 &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.enqueue(100) === 100 &&
      test.dequeue() === 100
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
class CircularQueue {
  constructor(size) {

    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line

    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line

    // Only change code above this line
  }
}
```

# --solutions--

```js
class CircularQueue {
  constructor(size) {
    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line
    console.log(this.write, this.max);
    if (this.queue[this.write] === null) {
      this.queue[this.write++] = item;

      if (this.write > this.max) {
        this.write = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line
    if (this.queue[this.read] !== null) {
      let item = this.queue[this.read];
      this.queue[this.read++] = null;
      if (this.read > this.max) {
        this.read = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }
}
```
