---
id: 587d8258367417b2b2512c80
title: Eliminare un nodo foglia in un albero binario di ricerca
challengeType: 1
forumTopicId: 301637
dashedName: delete-a-leaf-node-in-a-binary-search-tree
---

# --description--

Questa è la prima di tre sfide in cui implementeremo un'operazione più difficile negli alberi binari di ricerca: l'eliminazione. L'eliminazione è difficile perché rimuovere i nodi rompe i collegamenti nell'albero. Questi collegamenti devono essere ristabiliti con cura per garantire che la struttura dell'albero binario venga mantenuta. Per alcune eliminazioni, questo significa che l'albero deve essere riorganizzato. In generale, incontrerai uno di tre casi quando cerchi di eliminare un nodo: Nodo foglia: l'obiettivo da eliminare non ha figli. Un figlio: l'obiettivo da eliminare ha solo un figlio. Due figli: l'obiettivo da eliminare ha due nodi figli. Rimuovere un nodo foglia è facile, lo eliminiamo semplicemente. Eliminare un nodo con un figlio è anche relativamente facile, lo eliminiamo e colleghiamo il suo genitore al figlio del nodo eliminato. Rimuovere un nodo con due figli è invece più difficile, perché questo crea due nodi figli che devono essere ricollegati all'albero genitore. Vedremo come gestire questo caso nella terza sfida. Inoltre, devi fare attenzione ad alcuni casi limite quando gestisci l'eliminazione. Cosa succede se l'albero è vuoto? Cosa succede se il nodo da eliminare è la radice? Cosa succede se ci sono solo due elementi nell'albero? Per ora, occupiamoci del primo caso in cui eliminiamo un nodo foglia.

# --instructions--

Crea un metodo nel nostro albero binario chiamato `remove`. Costruiremo qui la logica per la nostra operazione di eliminazione. Per prima cosa, vorrai creare una funzione all'interno di remove che trovi il nodo che stiamo cercando di eliminare nell'albero corrente. Se il nodo non è presente nell'albero, `remove` dovrebbe restituire `null`. Ora, se il nodo obiettivo è un nodo foglia senza figli, allora il riferimento del genitore a esso dovrebbe essere impostato su `null`. Questo elimina effettivamente il nodo dall'albero. Per fare questo, dovrai anche tenere traccia del genitore del nodo che stiamo cercando di eliminare. Sarà inoltre utile creare un modo per tenere traccia del numero di figli che il nodo obiettivo ha, poiché questo determinerà a quale caso appartiene la nostra eliminazione. Gestiremo il secondo e il terzo caso nelle sfide successive. Buona fortuna!

# --hints--

La struttura dati `BinarySearchTree` dovrebbe esistere.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

L'albero binario di ricerca dovrebbe avere un metodo chiamato `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Provare a rimuovere un elemento da un albero vuoto dovrebbe restituire `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

Provare a rimuovere un elemento che non esiste dovrebbe restituire `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(30);
    return test.remove(100) == null;
  })()
);
```

Se il nodo radice non ha figli, eliminarlo dovrebbe impostare la radice su `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

Il metodo `remove` dovrebbe rimuovere i nodi foglia dall'albero.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
}
```

# --solutions--

```js
// solution required
```
