---
id: 587d8257367417b2b2512c7e
title: Usa la ricerca in profondità in un albero binario di ricerca
challengeType: 1
forumTopicId: 301719
dashedName: use-depth-first-search-in-a-binary-search-tree
---

# --description--

Sappiamo come cercare un valore specifico in un albero binario di ricerca. Ma cosa succede se vogliamo semplicemente esplorare l'intero albero? O se non abbiamo un albero ordinato e dobbiamo solo cercare un valore? Qui introdurremo alcuni metodi di attraversamento dell'albero che possono essere usati per esplorare strutture dati ad albero. Il primo è la ricerca in profondità. Nella ricerca in profondità, un sottoalbero dato viene esplorato il più a fondo possibile prima che la ricerca continui su un altro sottoalbero. Ci sono tre modi per farlo: In ordine: inizia la ricerca dal nodo più a sinistra e termina al nodo più a destra. Pre-ordine: esplora tutte le radici prima delle foglie. Post-ordine: esplora tutte le foglie prima delle radici. Come puoi immaginare, puoi scegliere metodi di ricerca diversi a seconda del tipo di dati che il tuo albero memorizza e di cosa stai cercando. Per un albero binario di ricerca, un attraversamento in ordine restituisce i nodi in ordine ordinato.

# --instructions--

Qui creeremo questi tre metodi di ricerca sul nostro albero binario di ricerca. La ricerca in profondità è un'operazione intrinsecamente ricorsiva che continua a esplorare ulteriori sottoalberi finché sono presenti nodi figli. Una volta compreso questo concetto base, puoi semplicemente riorganizzare l'ordine in cui esplori i nodi e i sottoalberi per produrre uno qualsiasi dei tre tipi di ricerca sopra descritti. Per esempio, nella ricerca post-ordine vorresti ricorrere fino a un nodo foglia prima di iniziare a restituire i nodi stessi, mentre nella ricerca pre-ordine vorresti restituire prima i nodi e poi continuare a ricorrere lungo l'albero. Definisci i metodi `inorder`, `preorder` e `postorder` sul nostro albero. Ognuno di questi metodi dovrebbe restituire un array di elementi che rappresentano l'attraversamento dell'albero. Assicurati di restituire i valori interi di ogni nodo nell'array, non i nodi stessi. Infine, restituisci `null` se l'albero è vuoto.

# --hints--

La struttura dati `BinarySearchTree` dovrebbe esistere.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

L'albero binario di ricerca dovrebbe avere un metodo chiamato `inorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.inorder == 'function';
  })()
);
```

L'albero binario di ricerca dovrebbe avere un metodo chiamato `preorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.preorder == 'function';
  })()
);
```

L'albero binario di ricerca dovrebbe avere un metodo chiamato `postorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.postorder == 'function';
  })()
);
```

Il metodo `inorder` dovrebbe restituire un array dei valori dei nodi risultanti da un attraversamento in ordine.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.inorder().join('') == '012345678910';
  })()
);
```

Il metodo `preorder` dovrebbe restituire un array dei valori dei nodi risultanti da un attraversamento pre-ordine.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.preorder().join('') == '710325469810';
  })()
);
```

Il metodo `postorder` dovrebbe restituire un array dei valori dei nodi risultanti da un attraversamento post-ordine.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.postorder().join('') == '024653181097';
  })()
);
```

Il metodo `inorder` dovrebbe restituire `null` per un albero vuoto.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    return test.inorder() == null;
  })()
);
```

Il metodo `preorder` dovrebbe restituire `null` per un albero vuoto.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    return test.preorder() == null;
  })()
);
```

Il metodo `postorder` dovrebbe restituire `null` per un albero vuoto.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    return test.postorder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.result = [];

  this.inorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.inorder(node.left);
    this.result.push(node.value);
    if (node.right) this.inorder(node.right);
    return this.result;
  };
  this.preorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    this.result.push(node.value);
    if (node.left) this.preorder(node.left);
    if (node.right) this.preorder(node.right);
    return this.result;
  };
  this.postorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.postorder(node.left);
    if (node.right) this.postorder(node.right);
    this.result.push(node.value);

    return this.result;
  };
}
```
