---
id: 587d8251367417b2b2512c63
title: Rimuovere elementi da una lista concatenata
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--

Il prossimo metodo importante che qualsiasi implementazione di una lista concatenata dovrà avere è un metodo `remove`. Questo metodo dovrebbe prendere come argomento l'elemento che vogliamo rimuovere, quindi cercare nella lista per trovare e rimuovere il nodo che contiene quell'elemento.

Ogni volta che rimuoviamo un nodo da una lista concatenata, è importante non lasciare accidentalmente orfano il resto della lista. Ricorda che la proprietà `next` di ogni nodo punta al nodo che lo segue nella lista. Se stiamo rimuovendo l'elemento centrale, per esempio, dovremo assicurarci di avere una connessione dalla proprietà `next` del nodo precedente a quella `next` dell'elemento centrale (che è il nodo successivo nella lista!)

Questo potrebbe sembrare molto confuso, quindi torniamo all'esempio della fila del conga per avere un buon modello concettuale. Immagina di essere in una fila del conga, e la persona direttamente davanti a te lascia la fila. La persona che ha appena lasciato la fila non ha più le mani su nessuno nella fila — e tu non hai più le mani sulla persona che è uscita. Fai un passo avanti e metti le mani sulla persona che vedi dopo.

Se l'elemento che vogliamo rimuovere è il primo elemento `head`, riassegniamo `head` al secondo nodo della lista concatenata.

# --instructions--

Scrivi un metodo `remove` che prende un elemento e lo rimuove dalla lista concatenata.

**Nota:** La `length` della lista dovrebbe diminuire di uno ogni volta che un elemento viene rimosso dalla lista concatenata.

# --hints--

La tua classe `LinkedList` dovrebbe avere un metodo `remove`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

Il tuo metodo `remove` dovrebbe riassegnare `head` al secondo nodo quando il primo nodo viene rimosso.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

Il tuo metodo `remove` dovrebbe diminuire la `length` della lista concatenata di uno per ogni nodo rimosso.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

Il tuo metodo `remove` dovrebbe riassegnare il riferimento del nodo precedente al nodo rimosso al riferimento `next` del nodo rimosso.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

Il tuo metodo `remove` non dovrebbe modificare la lista concatenata se l'elemento non esiste nella lista concatenata.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```
