---
id: 587d8251367417b2b2512c61
title: lavorare con i nodi in una lista concatenata
challengeType: 1
forumTopicId: 301721
dashedName: work-with-nodes-in-a-linked-list
---

# --description--

Un'altra struttura dati comune che incontrerai nell'informatica è la <dfn>lista concatenata</dfn>. Una lista concatenata è una collezione lineare di elementi dati, chiamati "nodi", ognuno dei quali punta al successivo. Ogni <dfn>nodo</dfn> in una lista concatenata contiene due informazioni chiave: il `element` stesso e un riferimento al `node` successivo.

Immagina di essere in una fila a conga. Hai le mani sulla persona davanti a te nella fila, e la persona dietro di te ha le mani su di te. Puoi vedere la persona subito davanti a te, ma loro bloccano la vista delle altre persone più avanti nella fila. Un nodo è proprio come una persona in una fila a conga: sa chi è e può vedere solo la persona successiva nella fila, ma non è consapevole delle altre persone avanti o dietro di sé.

# --instructions--

Nel nostro editor di codice, abbiamo creato due nodi, `Kitten` e `Puppy`, e abbiamo collegato manualmente il nodo `Kitten` al nodo `Puppy`.

Crea un nodo `Cat` e un nodo `Dog` e aggiungili manualmente alla fila.

# --hints--

Il tuo nodo `Puppy` dovrebbe avere un riferimento a un nodo `Cat`.

```js
assert(Puppy.next.element === 'Cat');
```

Il tuo nodo `Cat` dovrebbe avere un riferimento a un nodo `Dog`.

```js
assert(Cat.next.element === 'Dog');
```

# --seed--

## --seed-contents--

```js
var Node = function(element) {
  this.element = element;
  this.next = null;
};
var Kitten = new Node('Kitten');
var Puppy = new Node('Puppy');

Kitten.next = Puppy;
// Only change code below this line
```

# --solutions--

```js
// solution required
```
