---
id: 587d825c367417b2b2512c90
title: Ricerca in ampiezza
challengeType: 1
forumTopicId: 301622
dashedName: breadth-first-search
---

# --description--

Finora abbiamo imparato diversi modi per creare rappresentazioni di grafi. E adesso? Una domanda naturale è: quali sono le distanze tra due nodi qualsiasi nel grafo? Ecco che entrano in gioco gli <dfn>algoritmi di attraversamento del grafo</dfn>.

Gli <dfn>algoritmi di attraversamento</dfn> sono algoritmi per attraversare o visitare i nodi di un grafo. Un tipo di algoritmo di attraversamento è l'algoritmo di ricerca in ampiezza.

Questo algoritmo parte da un nodo e visita tutti i suoi vicini che sono a una distanza di un arco. Poi procede a visitare ciascuno dei loro vicini e così via fino a raggiungere tutti i nodi.

Una struttura dati importante che aiuta a implementare l'algoritmo di ricerca in ampiezza è la coda. Si tratta di un array in cui puoi aggiungere elementi da un'estremità e rimuoverli dall'altra. Questa struttura è anche nota come <dfn>FIFO</dfn> o <dfn>First-In-First-Out</dfn>.

Visivamente, questo è ciò che fa l'algoritmo:
<img alt="animazione che mostra l'algoritmo di ricerca in ampiezza" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/breadth-first-search.gif'>

L'ombreggiatura grigia rappresenta un nodo che viene aggiunto alla coda, mentre l'ombreggiatura nera rappresenta un nodo che viene rimosso dalla coda. Nota come ogni volta che un nodo viene rimosso dalla coda (il nodo diventa nero), tutti i suoi vicini vengono aggiunti alla coda (il nodo diventa grigio).

Per implementare questo algoritmo, dovrai fornire in input una struttura dati grafo e un nodo da cui partire.

Per prima cosa, devi tenere conto delle distanze dal nodo di partenza, cioè del numero di archi che lo separano da esso. Dovresti inizializzare tutte le distanze con un numero grande, come `Infinity`. Questo evita problemi di conteggio nel caso in cui un nodo non sia raggiungibile dal nodo di partenza. Poi, devi passare dal nodo di partenza ai suoi vicini. Questi vicini sono a una distanza di un arco e a questo punto devi aggiungere un’unità di distanza alle distanze che stai tracciando.

# --instructions--

Scrivi una funzione `bfs()` che prende come parametri un grafo rappresentato da una matrice di adiacenza (un array bidimensionale) e un'etichetta nodo root. L'etichetta del nodo sarà semplicemente il valore intero del nodo compreso tra `0` e `n - 1`, dove `n` è il numero totale di nodi nel grafo.

La tua funzione deve restituire un oggetto JavaScript con coppie chiave-valore in cui la chiave è il nodo e il valore è la sua distanza dal nodo root. Se un nodo non è raggiungibile, la sua distanza deve essere `Infinity`.

# --hints--

Il grafo in input `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` con nodo di partenza `1` dovrebbe restituire `{0: 1, 1: 0, 2: 1, 3: 2}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 1);
assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: 2 });
```

Il grafo in input `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` con nodo di partenza `1` dovrebbe restituire `{0: 1, 1: 0, 2: 1, 3: Infinity}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 0],
  [0, 0, 0, 0]
];
var results = bfs(graph, 1);
 assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: Infinity });
```

Il grafo in input `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` con nodo di partenza `0` dovrebbe restituire `{0: 0, 1: 1, 2: 2, 3: 3}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 0);
 assert.deepEqual(results, { 0: 0, 1: 1, 2: 2, 3: 3 });
```

Il grafo in input `[[0, 1], [1, 0]]` con nodo di partenza `0` dovrebbe restituire `{0: 0, 1: 1}`

```js
var graph = [
  [0, 1],
  [1, 0]
];
var results = bfs(graph, 0);
assert.deepEqual(results, { 0: 0, 1: 1 });
```

# --seed--

## --seed-contents--

```js
function bfs(graph, root) {
  var nodesLen = {};

  return nodesLen;
};

var exBFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(bfs(exBFSGraph, 3));
```

# --solutions--

```js
function bfs(graph, root) {
  var nodesLen = {};
  // Set all distances to infinity
  for (var i = 0; i < graph.length; i++) {
    nodesLen[i] = Infinity;
  }
  nodesLen[root] = 0; // ...except root node
  var queue = [root]; // Keep track of nodes to visit
  var current; // Current node traversing
  // Keep on going until no more nodes to traverse
  while (queue.length !== 0) {
    current = queue.shift();
    // Get adjacent nodes from current node
    var curConnected = graph[current]; // Get layer of edges from current
    var neighborIdx = []; // List of nodes with edges
    var idx = curConnected.indexOf(1); // Get first edge connection
    while (idx !== -1) {
      neighborIdx.push(idx); // Add to list of neighbors
      idx = curConnected.indexOf(1, idx + 1); // Keep on searching
    }
    // Loop through neighbors and get lengths
    for (var j = 0; j < neighborIdx.length; j++) {
      // Increment distance for nodes traversed
      if (nodesLen[neighborIdx[j]] === Infinity) {
        nodesLen[neighborIdx[j]] = nodesLen[current] + 1;
        queue.push(neighborIdx[j]); // Add new neighbors to queue
      }
    }
  }
  return nodesLen;
}
```
