---
id: 587d8251367417b2b2512c65
title: Rimuovere elementi da una lista concatenata per indice
challengeType: 1
forumTopicId: 301711
dashedName: remove-elements-from-a-linked-list-by-index
---

# --description--

Prima di passare a un'altra struttura dati, facciamo un paio di ultimi esercizi pratici con le liste concatenate.

Scriviamo un metodo `removeAt` che rimuove il `element` a un dato `index`. Il metodo dovrebbe chiamarsi `removeAt(index)`. Per rimuovere un `element` a un certo `index`, dobbiamo tenere un conteggio progressivo di ogni nodo mentre attraversiamo la lista concatenata.

Una tecnica comune per iterare sugli elementi di una lista concatenata prevede un <dfn>'runner'</dfn>, o sentinella, che 'punta' ai nodi che il tuo codice sta confrontando. Nel nostro caso, partendo dal `head` della lista, iniziamo con una variabile `currentIndex` che parte da `0`. Il `currentIndex` dovrebbe incrementare di uno per ogni nodo che attraversiamo.

Proprio come nel nostro metodo `remove(element)`, che <a href="/learn/coding-interview-prep/data-structures/remove-elements-from-a-linked-list" target="_blank" rel="noopener noreferrer nofollow">abbiamo visto in una lezione precedente</a>, dobbiamo fare attenzione a non lasciare orfano il resto della lista quando rimuoviamo il nodo nel nostro metodo `removeAt(index)`. Manteniamo i nodi contigui assicurandoci che il nodo che ha il riferimento al nodo rimosso punti al nodo successivo.

# --instructions--

Scrivi un metodo `removeAt(index)` che rimuove e restituisce un nodo a un dato `index`. Il metodo dovrebbe restituire `null` se il `index` dato è negativo oppure maggiore o uguale alla `length` della lista concatenata.

**Nota:** Ricorda di tenere il conto della `currentIndex`.

# --hints--

La tua classe `LinkedList` dovrebbe avere un metodo `removeAt`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.removeAt === 'function';
  })()
);
```

Il tuo metodo `removeAt` dovrebbe ridurre la `length` della lista concatenata di uno.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.removeAt(1);
    return test.size() === 2;
  })()
);
```

Il tuo metodo `removeAt` dovrebbe rimuovere l'elemento all'indice specificato dalla lista concatenata.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.add('bird');
    test.removeAt(1);
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"kitten","next":{"element":"bird","next":null}}}'
    );
  })()
);
```

Quando è presente un solo elemento nella lista concatenata, il tuo metodo `removeAt` dovrebbe rimuovere e restituire l'elemento all'indice specificato e ridurre la lunghezza della lista concatenata.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    var removedItem = test.removeAt(0);
    return test.head() === null && test.size() === 0 && removedItem === 'cat';
  })()
);
```

Il tuo metodo `removeAt` dovrebbe restituire l'elemento del nodo rimosso.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.removeAt(1) === 'dog';
  })()
);
```

Il tuo metodo `removeAt` dovrebbe restituire `null` e la lista concatenata non dovrebbe cambiare se l'indice dato è minore di `0`.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(-1);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

Il tuo metodo `removeAt` dovrebbe restituire `null` e la lista concatenata non dovrebbe cambiare se l'indice dato è maggiore o uguale alla `length` della lista.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(3);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
      head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function (element) {
    this.element = element;
    this.next = null;
  };

  this.size = function () {
    return length;
  };

  this.head = function () {
    return head;
  };

  this.add = function (element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.removeAt = function (index) {
    var currentNode = head;
    var previous = head;
    var count = 0;
    if (index >= length || index < 0) {
      return null;
    }
    if (index === 0) {
      var removed = head.element;
      head = currentNode.next;
    } else {
      while (count < index) {
        previous = currentNode;
        currentNode = currentNode.next;
        count++;
      }
      var removed = previous.next.element;
      previous.next = currentNode.next;
    }
    length--;
    return removed;
  };
}
```
