---
id: 587d825b367417b2b2512c8c
title: implementare heap sort con un heap minimo
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--

Ora che possiamo aggiungere e rimuovere elementi, vediamo alcune delle applicazioni per cui si possono usare gli heap. Gli heap sono comunemente usati per implementare code di priorità perché archiviano sempre un elemento di valore massimo o minimo in prima posizione. Inoltre, sono usati per implementare un algoritmo di ordinamento chiamato heap sort. Vedremo come farlo qui. Heap sort usa un heap minimo, l'opposto di un heap massimo. Un heap minimo archivia sempre l'elemento di valore più basso nella posizione radice.

Heap sort funziona prendendo un array non ordinato, aggiungendo ogni elemento dell'array in un heap minimo, e poi estraendo ogni elemento dall'heap minimo in un nuovo array. La struttura dell'heap minimo garantisce che il nuovo array conterrà gli elementi originali in ordine dal più piccolo al più grande. Questo è uno degli algoritmi di ordinamento più efficienti con prestazioni medie e nel peggior caso di O(nlog(n)).

# --instructions--

Implementiamo heap sort con un heap minimo. Sentiti libero di adattare qui il codice del tuo heap massimo. Crea un oggetto `MinHeap` con i metodi `insert`, `remove` e `sort`. Il metodo `sort` dovrebbe restituire un array con tutti gli elementi nell'heap minimo ordinati dal più piccolo al più grande.

# --hints--

La struttura dati `MinHeap` dovrebbe esistere.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MinHeap` dovrebbe avere un metodo chiamato `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MinHeap` dovrebbe avere un metodo chiamato `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`MinHeap` dovrebbe avere un metodo chiamato `sort`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

Il metodo `sort` dovrebbe restituire un array contenente tutti gli elementi aggiunti all'heap minimo in ordine ordinato.

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
