---
id: 587d825d367417b2b2512c96
title: Ricerca in profondità
challengeType: 1
forumTopicId: 301640
dashedName: depth-first-search
---

# --description--

Simile alla <dfn>ricerca in ampiezza</dfn>, qui impareremo un altro algoritmo di attraversamento di grafi chiamato <dfn>ricerca in profondità</dfn>.

Mentre la ricerca in ampiezza esplora progressivamente i nodi a distanza crescente dal nodo sorgente, la <dfn>ricerca in profondità</dfn> scende prima lungo un percorso di archi finché può.

Una volta raggiunta una estremità del percorso, la ricerca torna indietro all’ultimo nodo con un percorso di archi non visitato e continua la ricerca.

L’animazione qui sotto mostra come funziona l’algoritmo. L’algoritmo parte dal nodo in alto e visita i nodi nell’ordine numerato.

<img alt="animation showing the depth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/depth-first-search.gif'>

Nota come, a differenza della ricerca in ampiezza, ogni volta che un nodo viene visitato, non visita tutti i suoi vicini. Invece, visita prima uno dei suoi vicini e continua lungo quel percorso finché non ci sono più nodi da visitare su quel percorso.

Per implementare questo algoritmo, vorrai usare uno stack. Uno stack è un array in cui l’ultimo elemento aggiunto è il primo a essere rimosso. Questo è anche noto come struttura dati <dfn>Last-In-First-Out</dfn>. Uno stack è utile negli algoritmi di ricerca in profondità perché, mentre aggiungiamo i vicini allo stack, vogliamo visitare prima i vicini aggiunti più di recente e rimuoverli dallo stack.

Un output semplice di questo algoritmo è una lista di nodi raggiungibili da un nodo dato. Perciò, vorrai anche tenere traccia dei nodi che visiti.

# --instructions--

Scrivi una funzione `dfs()` che prende come parametri una matrice di adiacenza non orientata `graph` e un’etichetta di nodo `root`. L’etichetta del nodo sarà semplicemente il valore numerico del nodo compreso tra `0` e `n - 1`, dove `n` è il numero totale di nodi nel grafo.

La tua funzione dovrebbe restituire un array con tutti i nodi raggiungibili da `root`.

# --hints--

Il grafo di input `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` con nodo iniziale `1` dovrebbe restituire un array con `0`, `1`, `2` e `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })(),
  [0, 1, 2, 3]
);
```

Il grafo di input `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` con nodo iniziale `3` dovrebbe restituire un array con `3`, `2`, `1` e `0`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3, 2, 1, 0]
);
```

Il grafo di input `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` con nodo iniziale `1` dovrebbe restituire un array con quattro elementi.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })().length === 4
);
```

Il grafo di input `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` con nodo iniziale `3` dovrebbe restituire un array con `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3]
);
```

Il grafo di input `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` con nodo iniziale `3` dovrebbe restituire un array con un elemento.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })().length === 1
);
```

Il grafo di input `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` con nodo iniziale `3` dovrebbe restituire un array con `2` e `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [2, 3]
);
```

Il grafo di input `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` con nodo iniziale `3` dovrebbe restituire un array con due elementi.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })().length === 2
);
```

Il grafo di input `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` con nodo iniziale `0` dovrebbe restituire un array con `0` e `1`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })(),
  [0, 1]
);
```

Il grafo di input `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` con nodo iniziale `0` dovrebbe restituire un array con due elementi.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })().length === 2
);
```

# --seed--

## --seed-contents--

```js
function dfs(graph, root) {

}

var exDFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(dfs(exDFSGraph, 3));
```

# --solutions--

```js
function dfs(graph, root) {
    var stack = [];
    var tempV;
    var visited = [];
    var tempVNeighbors = [];
    stack.push(root);
    while (stack.length > 0) {
        tempV = stack.pop();
        if (visited.indexOf(tempV) == -1) {
            visited.push(tempV);
            tempVNeighbors = graph[tempV];
            for (var i = 0; i < tempVNeighbors.length; i++) {
                if (tempVNeighbors[i] == 1) {
                    stack.push(i);
                }
            }
        }
    }
    return visited;
}
```
