---
id: 587d8252367417b2b2512c67
title: aggiungere elementi a un indice specifico in una lista concatenata
challengeType: 1
forumTopicId: 301619
dashedName: add-elements-at-a-specific-index-in-a-linked-list
---

# --description--

Creiamo un metodo addAt(index,element) che aggiunge un elemento a un indice specifico. Proprio come quando rimuoviamo elementi a un indice specifico, dobbiamo tenere traccia dell'indice corrente mentre attraversiamo la lista concatenata. Quando l'indice corrente corrisponde all'indice dato, dobbiamo riassegnare la proprietà next del nodo precedente per riferirsi al nuovo nodo aggiunto. E il nuovo nodo dovrebbe riferirsi al nodo successivo all'indice corrente. Tornando all'esempio della fila conga, una nuova persona vuole unirsi alla fila, ma vuole inserirsi in mezzo. Sei nel mezzo della fila, quindi togli le mani dalla persona davanti a te. La nuova persona si avvicina e mette le mani sulla persona che avevi prima, e ora hai le mani sulla nuova persona.

# --instructions--

Crea un metodo `addAt(index,element)` che aggiunge un elemento a un indice specifico. Ritorna falso se non è stato possibile aggiungere un elemento. **Nota:** Ricorda di controllare se l'indice dato è negativo o è più lungo della lunghezza della lista concatenata.

# --hints--

Il tuo metodo `addAt` dovrebbe riassegnare `head` al nuovo nodo quando l'indice dato è 0.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'fish');
    return test.head().element === 'fish' && test.head().next.element === 'cat';
  })()
);
```

Il tuo metodo `addAt` dovrebbe aumentare la lunghezza della lista concatenata di uno per ogni nuovo nodo aggiunto alla lista concatenata.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'cat');
    return test.size() === 3;
  })()
);
```

Il tuo metodo `addAt` dovrebbe ritornare `false` se un nodo non è stato possibile aggiungerlo.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.addAt(4, 'cat') === false;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element) {
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function() {
    return head;
  };

  this.add = function(element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if (head === null){
        head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };
  this.addAt = function (index, element) {
    if (index > length || index < 0) {
      return false;
    }
    var newNode = new Node(element);
    var currentNode = head;
    if (index === 0) {
      head = newNode;
    } else {
      var previousNode = null;
      var i = 0;
      while (currentNode && i < index) {
        previousNode = currentNode;
        currentNode = currentNode.next;
        i++;
      }
      previousNode.next = newNode;
    }
    newNode.next = currentNode;
    length++;
  }
}
```
