---
id: 587d8251367417b2b2512c62
title: crea una classe lista concatenata
challengeType: 1
forumTopicId: 301628
dashedName: create-a-linked-list-class
---

# --description--

Creiamo una classe `linked list`. Ogni lista concatenata dovrebbe iniziare con alcune proprietà di base: un `head` (il primo elemento della tua lista) e un `length` (numero di elementi nella tua lista). A volte vedrai implementazioni di liste concatenate che includono un `tail` per l'ultimo elemento della lista, ma per ora ci limiteremo a queste due. Ogni volta che aggiungiamo un elemento alla lista concatenata, la nostra proprietà `length` dovrebbe aumentare di uno.

Vogliamo avere un modo per aggiungere elementi alla nostra lista concatenata, quindi il primo metodo che vogliamo creare è il metodo `add`.

Se la nostra lista è vuota, aggiungere un elemento alla lista concatenata è abbastanza semplice: basta incapsulare quell'elemento in una classe `Node` e assegnare quel nodo al `head` della nostra lista concatenata.

Ma cosa succede se la nostra lista ha già uno o più membri? Come aggiungiamo un elemento alla lista? Ricorda che ogni nodo in una lista concatenata ha una proprietà `next`. Per aggiungere un nodo alla lista, trova l'ultimo nodo nella lista e punta la proprietà `next` di quell'ultimo nodo al nostro nuovo nodo. (Suggerimento: sai di aver raggiunto la fine di una lista concatenata quando la proprietà `next` di un nodo è `null`.)

# --instructions--

Scrivi un metodo add che assegni il primo nodo che aggiungi alla lista concatenata al `head`; dopo di che, ogni volta che aggiungi un nodo, ogni nodo dovrebbe essere referenziato dalla proprietà `next` del nodo precedente.

Nota

La proprietà `length` della tua lista dovrebbe aumentare di uno ogni volta che un elemento viene aggiunto alla lista concatenata.

# --hints--

La tua classe `LinkedList` dovrebbe avere un metodo `add`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.add === 'function';
  })()
);
```

La tua classe `LinkedList` dovrebbe assegnare `head` al primo nodo aggiunto.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    return test.head().element === 'cat';
  })()
);
```

Il precedente `node` nella tua classe `LinkedList` dovrebbe avere un riferimento al nodo più recente creato.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('fish');
    return test.head().next.element === 'dog' && test.head().next.next.element === 'fish';
  })()
);
```

La `size` della tua classe `LinkedList` dovrebbe essere uguale al numero di nodi nella lista concatenata.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.size() === 2;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() { 
  var length = 0; 
  var head = null; 

  var Node = function(element){
    this.element = element; 
    this.next = null; 
  }; 

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line
    if (head == null) {
      head = new Node(element);
    } 
    else {
      let currentNode = head;
      while (currentNode.next != null) {
        // currentNode.next will be last node of linked list after loop
        currentNode = currentNode.next;
      }
      currentNode.next = new Node(element);
    }
    length++;
    // Only change code above this line
  };
}
```
