---
id: 587d8257367417b2b2512c7c
title: Verificare se un elemento è presente in un albero binario di ricerca
challengeType: 1
forumTopicId: 301623
dashedName: check-if-an-element-is-present-in-a-binary-search-tree
---

# --description--

Ora che abbiamo un'idea generale di cosa sia un albero binario di ricerca, parliamone un po' più nel dettaglio. Gli alberi binari di ricerca offrono un tempo logaritmico per le operazioni comuni di ricerca, inserimento e cancellazione nel caso medio, e un tempo lineare nel caso peggiore. Perché succede questo? Ognuna di queste operazioni di base richiede di trovare un elemento nell'albero (o nel caso dell'inserimento di trovare dove dovrebbe andare) e, grazie alla struttura ad albero, a ogni nodo genitore si decide se andare a sinistra o a destra escludendo di fatto metà della dimensione dell'albero rimanente. Questo rende la ricerca proporzionale al logaritmo del numero di nodi nell'albero, creando così un tempo logaritmico per queste operazioni nel caso medio. Ok, ma nel caso peggiore? Considera di costruire un albero con i seguenti valori, aggiungendoli da sinistra a destra: `10`, `12`, `17`, `25`. Seguendo le regole per un albero binario di ricerca, aggiungeremo `12` a destra di `10`, `17` a destra di questo, e `25` a destra di questo. Ora il nostro albero assomiglia a una lista concatenata e attraversarlo per trovare `25` richiederebbe di scorrere tutti gli elementi in modo lineare. Da qui il tempo lineare nel caso peggiore. Il problema è che l'albero è sbilanciato. Approfondiremo cosa significa nelle sfide successive.

# --instructions--

In questa sfida creeremo un'utilità per il nostro albero. Scrivi un metodo `isPresent` che prende come input un valore intero e restituisce un valore booleano che indica la presenza o l'assenza di quel valore nell'albero binario di ricerca.

# --hints--

La struttura dati `BinarySearchTree` dovrebbe esistere.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

L'albero binario di ricerca dovrebbe avere un metodo chiamato `isPresent`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isPresent == 'function';
  })()
);
```

Il metodo `isPresent` dovrebbe controllare correttamente la presenza o l'assenza degli elementi aggiunti all'albero.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    test.add(4);
    test.add(7);
    test.add(411);
    test.add(452);
    return (
      test.isPresent(452) &&
      test.isPresent(411) &&
      test.isPresent(7) &&
      !test.isPresent(100)
    );
  })()
);
```

`isPresent` dovrebbe gestire i casi in cui l'albero è vuoto.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    return test.isPresent(5) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.isPresent = function (value) {
    var current = this.root
    while (current) {
      if (value === current.value) {
        return true;
      }
      current = value < current.value ? current.left : current.right;
    }
    return false;
  }
}
```
