---
id: 587d8257367417b2b2512c7b
title: Aggiungere un nuovo elemento a un albero binario di ricerca
challengeType: 1
forumTopicId: 301618
dashedName: add-a-new-element-to-a-binary-search-tree
---

# --description--

Questa serie di sfide introdurrà la struttura dati ad albero. Gli alberi sono una struttura dati importante e versatile nell'informatica. Naturalmente, il loro nome deriva dal fatto che, quando vengono visualizzati, assomigliano molto agli alberi che conosciamo nel mondo naturale. Una struttura dati ad albero inizia con un nodo, solitamente chiamato radice, e da qui si dirama verso altri nodi, ognuno dei quali può avere ulteriori nodi figli, e così via. La struttura dati viene solitamente visualizzata con il nodo radice in alto; puoi pensare a un albero naturale capovolto.

Per prima cosa, descriviamo alcuni termini comuni che incontreremo con gli alberi. Il nodo radice è la cima dell'albero. I punti dati nell'albero si chiamano nodi. I nodi con rami che portano ad altri nodi sono detti genitori del nodo a cui il ramo conduce (il figlio). Si applicano anche altri termini familiari più complessi, come ti aspetteresti. Un sottoalbero si riferisce a tutti i discendenti di un nodo particolare, i rami possono essere chiamati archi, e i nodi foglia sono nodi alla fine dell'albero che non hanno figli. Infine, nota che gli alberi sono strutture dati intrinsecamente ricorsive. Cioè, ogni figlio di un nodo è genitore del proprio sottoalbero, e così via. La natura ricorsiva degli alberi è importante da capire quando si progettano algoritmi per operazioni comuni sugli alberi.

Per cominciare, parleremo di un tipo particolare di albero, l'albero binario. In effetti, discuteremo di un albero binario particolare, un albero binario di ricerca. Vediamo cosa significa. Mentre una struttura dati ad albero può avere un numero qualsiasi di rami in un singolo nodo, un albero binario può avere solo due rami per ogni nodo. Inoltre, un albero binario di ricerca è ordinato rispetto ai sottoalberi figli, in modo che il valore di ogni nodo nel sottoalbero sinistro sia minore o uguale al valore del nodo genitore, e il valore di ogni nodo nel sottoalbero destro sia maggiore o uguale al valore del nodo genitore. È molto utile visualizzare questa relazione per capirla meglio:

<div style='width: 100%; display: flex; justify-content: center; align-items: center;'><img alt="an example of a binary search tree" style='width: 100%; max-width: 350px; background-color: var(--gray-05);' src='https://user-images.githubusercontent.com/18563015/32136009-1e665d98-bbd6-11e7-9133-63184f9f9182.png'></div>

Ora questa relazione ordinata è molto facile da vedere. Nota che ogni valore a sinistra di 8, il nodo radice, è minore di 8, e ogni valore a destra è maggiore di 8. Nota anche che questa relazione si applica a ciascuno dei sottoalberi. Per esempio, il primo figlio sinistro è un sottoalbero. 3 è il nodo genitore, e ha esattamente due nodi figli — secondo le regole che governano gli alberi binari di ricerca, sappiamo senza nemmeno guardare che il figlio sinistro di questo nodo (e di tutti i suoi figli) sarà minore di 3, e il figlio destro (e tutti i suoi figli) sarà maggiore di 3 (ma anche minore del valore radice della struttura), e così via.

Gli alberi binari di ricerca sono strutture dati molto comuni e utili perché forniscono un tempo logaritmico nel caso medio per diverse operazioni comuni come ricerca, inserimento e cancellazione.

# --instructions--

Iniziamo con qualcosa di semplice. Abbiamo definito lo scheletro di una struttura ad albero binario di ricerca qui, oltre a una funzione per creare nodi per il nostro albero. Osserva che ogni nodo può avere un valore sinistro e destro. Questi saranno assegnati ai sottoalberi figli se esistono. Nel nostro albero binario di ricerca, creerai un metodo per aggiungere nuovi valori all'albero. Il metodo dovrebbe chiamarsi `add` e dovrebbe accettare un valore intero da aggiungere all'albero. Fai attenzione a mantenere l'invariante di un albero binario di ricerca: il valore in ogni figlio sinistro dovrebbe essere minore o uguale al valore del genitore, e il valore in ogni figlio destro dovrebbe essere maggiore o uguale al valore del genitore. Qui, facciamo in modo che il nostro albero non possa contenere valori duplicati. Se proviamo ad aggiungere un valore che esiste già, il metodo dovrebbe restituire `null`. Altrimenti, se l'aggiunta ha successo, dovrebbe essere restituito `undefined`.

**Suggerimento:** gli alberi sono strutture dati naturalmente ricorsive!

# --hints--

La struttura dati `BinarySearchTree` dovrebbe esistere.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

L'albero binario di ricerca dovrebbe avere un metodo chiamato `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

Il metodo add dovrebbe aggiungere elementi secondo le regole dell'albero binario di ricerca.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    const expectedResult = [1, 4, 7, 8, 34, 45, 73, 87];
    const result = test.inOrder();
    return expectedResult.toString() === result.toString();
  })()
);
```

Aggiungere un elemento che esiste già dovrebbe restituire `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    return test.add(4) == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    inOrder() {
      if (!this.root) {
        return null;
      }
      var result = new Array();
      function traverseInOrder(node) {
        node.left && traverseInOrder(node.left);
        result.push(node.value);
        node.right && traverseInOrder(node.right);
      }
      traverseInOrder(this.root);
      return result;
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.add = function(element) {
    let current = this.root;
    if (!current) {
      this.root = new Node(element);
      return;
    } else {
      const searchTree = function(current) {
        if (current.value > element) {
          if (current.left) {
            return searchTree(current.left);
          } else {
            current.left = new Node(element);
            return;
          }
        } else if (current.value < element) {
          if (current.right) {
            return searchTree(current.right);
          } else {
            current.right = new Node(element);
            return;
          }
        } else {
          return null;
        }
      };
      return searchTree(current);
    }
  };
}
```
