---
id: 587d825b367417b2b2512c8e
title: Creare una tabella hash
challengeType: 1
forumTopicId: 301627
dashedName: create-a-hash-table
---

# --description--

In questa sfida impareremo cosa sono le tabelle hash. Una tabella hash viene usata per implementare array associativi, o mappature di coppie chiave-valore, come gli oggetti e le Mappe che abbiamo appena studiato. Un oggetto JavaScript potrebbe essere implementato come una tabella hash, per esempio (la sua implementazione effettiva dipenderà dall’ambiente in cui viene eseguito). Il funzionamento di una tabella hash consiste nel prendere una chiave come input e calcolare un hash di questa chiave in modo deterministico, ottenendo un valore numerico. Questo valore numerico viene poi usato come chiave effettiva sotto cui viene archiviato il valore associato. Quindi, se provi a consultare la stessa chiave di nuovo, la funzione di hashing elaborerà la chiave, restituirà lo stesso risultato numerico, che verrà usato per cercare il valore associato. Questo garantisce un tempo di consultazione molto efficiente, in media O(1).

Le tabelle hash possono essere implementate come array con funzioni di hashing che producono indici di array entro un intervallo specificato. In questo metodo, la scelta della dimensione dell’array è importante, così come la funzione di hashing. Per esempio, cosa succede se la funzione di hashing produce lo stesso valore per due chiavi diverse? Questo si chiama collisione. Un modo per gestire le collisioni è semplicemente archiviare entrambe le coppie chiave-valore in quell’indice. Poi, durante la consultazione di una delle due, dovresti iterare sul gruppo di elementi per trovare la chiave che stai cercando. Una buona funzione di hashing minimizzerà le collisioni per mantenere un tempo di ricerca efficiente.

Qui non ci preoccuperemo dei dettagli dell’hashing o dell’implementazione della tabella hash, cercheremo solo di capire in generale come funzionano.

# --instructions--

Creiamo la funzionalità base di una tabella hash. Abbiamo creato per te una funzione di hashing semplice da usare. Puoi passare una stringa alla funzione `hash` e questa restituirà un valore hash che puoi usare come chiave per l’archivio. Archivia gli elementi basandoti su questo valore hash nell’oggetto `this.collection`. Crea questi tre metodi: `add`, `remove` e `lookup`. Il primo deve accettare una coppia chiave-valore da aggiungere alla tabella hash. Il secondo deve rimuovere una coppia chiave-valore quando gli viene passata una chiave. Il terzo deve accettare una chiave e restituire il valore associato oppure `null` se la chiave non è presente.

Assicurati di scrivere il codice tenendo conto delle collisioni!

**Nota:** I test del metodo `remove` non passeranno finché i metodi `add` e `lookup` non saranno implementati correttamente.

# --hints--

La funzione `hash` dovrebbe essere valida.

```js
let calls = 0;
const ourHash = string => {
  calls++;
  let hashCode = 0;
  for (let i = 0; i < string.length; i++) {
    hashCode += string.charCodeAt(i);
  }
  return hashCode;
};

assert.strictEqual(hash('yek'),ourHash('yek'));

assert.strictEqual(hash('key'),ourHash('key'));
assert.strictEqual(hash('key1'),ourHash('key1'));
assert.strictEqual(hash('key2'),ourHash('key2'));
assert.strictEqual(hash('key3'),ourHash('key3'));

assert.strictEqual(hash('1key'),ourHash('1key'));
assert.strictEqual(hash('ke1y'),ourHash('ke1y'));
assert.strictEqual(hash('altKey'),ourHash('altKey'));

assert.strictEqual(called,calls); 
```

La struttura dati `HashTable` dovrebbe esistere.

```js
assert.isDefined(HashTable); 
let newHashTable = new HashTable(); 
assert.isObject(newHashTable); 
```

`HashTable` dovrebbe avere un metodo `add`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.add); 
```

`HashTable` dovrebbe avere un metodo `lookup`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.lookup); 
```

`HashTable` dovrebbe avere un metodo `remove`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.remove); 
```

Il metodo `add` dovrebbe aggiungere coppie chiave-valore e il metodo `lookup` dovrebbe restituire i valori associati a una data chiave.

```js
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
assert.strictEqual(newHashTable.lookup('key'),'value'); 
```

Il metodo `remove` dovrebbe accettare una chiave come input e rimuovere la coppia chiave-valore associata.

```js
let hashValue = hash('key');
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
newHashTable.remove('key'); 
assert.notProperty(newHashTable.collection,hashValue); 
```

Il metodo `remove` dovrebbe rimuovere solo la coppia chiave-valore corretta.

```js
let newHashTable = new HashTable(); 
let hashValue = hash('key');
newHashTable.add('key','value');
newHashTable.add('yek','value');
newHashTable.add('altKey','value');
newHashTable.remove('yek');

assert.notExists(newHashTable.lookup('yek'));
assert.exists(newHashTable.lookup('altKey'));
assert.exists(newHashTable.lookup('key'));

newHashTable.remove('key');

assert.notProperty(newHashTable.collection,hashValue); 
assert.exists(newHashTable.lookup('altKey'));

```

Gli elementi dovrebbero essere aggiunti usando la funzione di hashing.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1', 'value1')
newHashTable.add('key2', 'value2');
newHashTable.add('key3', 'value3');
assert.strictEqual(called,3); 
```

La tabella hash dovrebbe gestire le collisioni.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1','value1');
newHashTable.add('1key', 'value2'); 
newHashTable.add('ke1y', 'value3'); 
assert.strictEqual(newHashTable.lookup('key1'),'value1');
assert.strictEqual(newHashTable.lookup('1key'),'value2');
assert.strictEqual(newHashTable.lookup('ke1y'),'value3');
```

# --seed--

## --seed-contents--

```js
var called = 0;
var hash = string => {
  called++;
  var hashed = 0;
  for (var i = 0; i < string.length; i++) {
    hashed += string.charCodeAt(i);
  }
  return hashed;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
var called = 0;
var hash = (string) => {
  called++;
  var hash = 0;
  for (var i = 0; i < string.length; i++) { hash += string.charCodeAt(i); }
  return hash;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line

  this.add = function(key, val) {
    var theHash = hash(key);
    if (!this.collection.hasOwnProperty(theHash)) {
      this.collection[theHash] = {};
    }
    this.collection[theHash][key] = val;
  }

  this.remove = function(key) {
    var theHash = hash(key);
    var hashedObj = this.collection[theHash];
    if (hashedObj.hasOwnProperty(key)) {
      delete hashedObj[key];
    }
    if (!Object.keys(hashedObj).length) {
      delete this.collection[theHash];
    }
  }

  this.lookup = function(key) {
    var theHash = hash(key);
    if (this.collection.hasOwnProperty(theHash)) {
      return this.collection[theHash][key];
    }
    return null
  }
  // Only change code above this line
};
```
