---
id: 587d8255367417b2b2512c74
title: creare una classe coda di priorità
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--

In questa sfida creerai una coda di priorità. Una coda di priorità è un tipo speciale di coda in cui gli elementi possono avere informazioni aggiuntive che specificano la loro priorità. Questo può essere rappresentato semplicemente con un numero intero. La priorità dell'elemento sovrascriverà l'ordine di inserimento nel determinare la sequenza con cui gli elementi vengono estratti. Se un elemento con priorità più alta viene inserito dopo elementi con priorità più bassa, l'elemento con priorità più alta verrà estratto prima di tutti gli altri.

Per esempio, immaginiamo di avere una coda di priorità con tre elementi:

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Qui il secondo valore (un numero intero) rappresenta la priorità dell'elemento. Se inseriamo `['human', 1]` con una priorità di `1` (supponendo che le priorità più basse abbiano precedenza) allora sarà il primo elemento ad essere estratto. La collezione apparirebbe così:

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Abbiamo iniziato a scrivere una `PriorityQueue` nell'editor del codice. Dovrai aggiungere un metodo `enqueue` per aggiungere elementi con una priorità, un metodo `dequeue` per rimuovere e restituire elementi, un metodo `size` per restituire il numero di elementi nella coda, un metodo `front` per restituire l'elemento in testa alla coda, e infine un metodo `isEmpty` che restituirà `true` se la coda è vuota o `false` se non lo è.

Il metodo `enqueue` dovrebbe accettare elementi con il formato mostrato sopra (`['human', 1]`) dove `1` rappresenta la priorità. `dequeue` e `front` dovrebbero restituire solo il nome dell'elemento, non la sua priorità.

# --hints--

La tua classe `PriorityQueue` dovrebbe avere un metodo `enqueue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

La tua classe `PriorityQueue` dovrebbe avere un metodo `dequeue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

La tua classe `PriorityQueue` dovrebbe avere un metodo `size`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

La tua classe `PriorityQueue` dovrebbe avere un metodo `front`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

La tua classe `PriorityQueue` dovrebbe avere un metodo `isEmpty`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

La tua classe `PriorityQueue` dovrebbe tenere correttamente traccia del numero attuale di elementi usando il metodo `size` mentre gli elementi vengono inseriti ed estratti.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

Il metodo `front` dovrebbe restituire l'elemento corretto in testa alla coda mentre gli elementi vengono inseriti ed estratti.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

Il metodo `isEmpty` dovrebbe restituire `true` quando la coda è vuota.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

La coda di priorità dovrebbe restituire gli elementi con priorità più alta prima degli elementi con priorità più bassa e restituire gli elementi in ordine FIFO altrimenti.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```
