---
id: 587d8255367417b2b2512c72
title: Usa .has e .size su un set ES6
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

Vediamo i metodi .has e .size disponibili sull'oggetto Set di ES6.

Per prima cosa, crea un Set ES6

```js
var set = new Set([1,2,3]);
```

Il metodo .has verifica se un valore è contenuto all'interno del set.

```js
var hasTwo = set.has(2);
```

Il metodo .size restituisce un intero che rappresenta la dimensione del Set

```js
var howBig = set.size;
```

# --instructions--

In questo esercizio passeremo un array e un valore alla funzione checkSet(). La tua funzione dovrebbe creare un set ES6 dall'argomento array. Trova se il set contiene l'argomento valore. Trova la dimensione del set. E restituisci questi due valori in un array.

# --hints--

`checkSet([4, 5, 6], 3)` dovrebbe restituire [ falso, 3 ]

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
