---
id: 67051431a73f1ca25d9a6f25
title: Crea una bacheca per idee di progetto
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi definire una costante oggetto chiamata `projectStatus` con le tre chiavi: `PENDING`, `SUCCESS` e `FAILURE`. A ogni stato deve essere assegnato un oggetto con una proprietà `description` con il valore `Pending Execution`, `Executed Successfully` e `Execution Failed`, rispettivamente.
2. Devi definire una classe chiamata `ProjectIdea` con un `constructor` che prende `title` e `description` come parametri. Inizializza le proprietà `title` e `description` con i parametri forniti. La classe deve avere anche una proprietà chiamata `status` impostata al valore `projectStatus.PENDING` per impostazione predefinita.
3. Devi definire un metodo chiamato `updateProjectStatus` all’interno della classe `ProjectIdea`. Questo metodo deve accettare un parametro `newStatus` e aggiornare la proprietà `status` con il valore dato.
4. Devi definire una classe `ProjectIdeaBoard` con un `constructor` che accetta un `title` e inizializza un array vuoto chiamato `ideas` per contenere istanze della classe `ProjectIdea`.
5. Devi definire un metodo chiamato `pin` all’interno della classe `ProjectIdeaBoard` che accetta un’istanza della classe `ProjectIdea` e aggiunge l’istanza data all’array `ideas`.
6. Devi definire un metodo chiamato `unpin` all’interno della classe `ProjectIdeaBoard`. Questo metodo deve accettare un’istanza della classe `ProjectIdea` e rimuoverla dall’array `ideas`.
7. Devi definire un metodo chiamato `count` che restituisce il numero di idee di progetto nell’array `ProjectIdeaBoard` dato.
8. Devi definire un metodo chiamato `formatToString` che restituisce il nome delle idee di progetto, la loro descrizione e stato nel formato:

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

Devi definire una costante oggetto chiamata `projectStatus` con gli stati di progetto e le descrizioni corrette.

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

Devi avere una classe `ProjectIdea`.

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

La tua classe `ProjectIdea` deve inizializzare la proprietà `title` e la proprietà `description` in base ai parametri passati. Deve anche impostare lo stato su `projectStatus.PENDING`.

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

La tua classe `ProjectIdea` deve avere un metodo `updateProjectStatus`.

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

Devi avere una classe `ProjectIdeaBoard`.

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

La tua `ProjectIdeaBoard` deve inizializzare la proprietà `title` in base al parametro passato e inizializzare un array vuoto chiamato `ideas` per contenere istanze della classe `ProjectIdea`.

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

La tua classe `ProjectIdeaBoard` deve avere un metodo `pin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

La tua classe `ProjectIdeaBoard` deve avere un metodo `unpin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

La tua classe `ProjectIdeaBoard` deve avere un metodo `count`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

La tua classe `ProjectIdeaBoard` deve avere un metodo `formatToString`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")` deve restituire `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }`.

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

Chiamare `updateProjectStatus(projectStatus.SUCCESS)` su `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")` deve aggiornare lo stato a `{ description: 'Executed Successfully' }`.

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

Chiamare `updateProjectStatus(projectStatus.FAILURE)` su `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")` deve aggiornare lo stato a `{ description: 'Execution Failed' }`.

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

Chiamare `updateProjectStatus(projectStatus.SUCCESS)` su `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")` deve aggiornare lo stato a `{ description: 'Executed Successfully' }`.

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

Devi poter fissare un oggetto `ProjectIdea` alla tua `ProjectIdeaBoard` usando il metodo `pin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

Devi poter rimuovere un oggetto `ProjectIdea` dalla tua `ProjectIdeaBoard` usando il metodo `unpin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

Quando `new ProjectIdeaBoard("Empty Board")` è vuoto, `emptyBoard.formatToString()` deve restituire `Empty Board has 0 idea(s)\n`.

```js

let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

// Create an empty board
let emptyBoard = new ProjectIdeaBoard("Empty Board");

// Attempt to unpin a project from an empty board
emptyBoard.unpin(smartHome);

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

Quando fissi `new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")` a `new ProjectIdeaBoard("Tech Projects Board")`, `techProjects.formatToString()` deve restituire `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n`.

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
