---
id: 616965351e74d4689eb6de30
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Você pode ter vários elementos `meta` em uma página da web. Cada elemento `meta` adiciona informações sobre a página que não podem ser expressas por outros elementos do HTML.

Adicione outro elemento `meta` dentro do elemento `head`. Dê a ele o atributo `name` definido como `"viewport"` e o atributo `content` definido como `"width=device-width, initial-scale=1.0"` para que a página tenha a mesma aparência em todos os dispositivos.

# --hints--

Você deve ter dois elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

Seu novo elemento `meta` deve ser um elemento vazio, ele não deve ter uma tag de fechamento `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```


O novo elemento `meta` deve ter o atributo `name` definido como `"viewport"` e o atributo `content` definido como `"width=device-width, initial-scale=1.0"`.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

Os dois elementos `meta` devem estar dentro do elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
