---
id: 61695c4aad56f95497c19583
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Lembre-se de que o elemento `title` fornece aos mecanismos de busca informações extras sobre a página.Ele também exibe o conteúdo desse elemento `title` de mais duas maneiras:

* na barra de título quando a página está aberta  
* na aba do navegador para a página quando você passa o mouse sobre ela. Mesmo que essa aba não esteja ativa, assim que você passa o mouse sobre a aba, o texto `title` é exibido.

Dentro do elemento `head`, aninhe um elemento `title` com o texto `Colored Markers`.

# --hints--

Você deve acrescentar uma tag de abertura`title`.

```js
assert(code.match(/<title\s*>/i));
```

Você deve acrescentar uma tag de fechamento para `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

O projeto deve ter o título `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Lembre-se, maiúsculas e minúsculas e ortografia são importantes para o texto do título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
