---
id: 60f02e7361b68405e27b62a5
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви навчитеся стилізувати форми, створюючи форму реєстрації.

Весь шаблонний код HTML (`DOCTYPE`, `html`, `head`, і `body`) уже надано.

У межах `body` додайте заголовок, щоб надати контекст формі, використовуючи елемент `h1` з текстом `Registration Form`.

# --hints--

Ви повинні додати `h1` всередину `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Ви повинні надати `h1` текст `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
