---
id: 60f5d2776c854e069560fbe6
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Перший `fieldset` міститиме поля для імені, е-пошти та пароля. Почніть із додавання чотирьох `label` елементів до першого `fieldset`.

# --hints--

Всього має бути три `fieldset` елементи.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

`fieldset` елементи мають бути всіми прямими дочірніми елементами `form` елемента.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

Чотири `label` елементи мають бути всіми прямими дочірніми елементами першого `fieldset` елемента. Переконайтеся, що ви закриваєте `label` елементи.

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

Останні два `fieldset` елементи мають бути порожніми.

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

Переконайтеся, що ви закриваєте `label` елементи.

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
