---
id: 60f80e0081e0f2052ae5b505
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Вказання атрибута `type` елемента `input` важливе, щоб браузер знав, який тип даних очікувати. Якщо `type` не вказано, браузер за замовчуванням використовуватиме `text`.

Дайте першим двом елементам `input` атрибут `type` зі значенням `text`, третьому — атрибут `type` зі значенням `email`, а четвертому — атрибут `type` зі значенням `password`.

Тип `email` дозволяє лише е-пошту з `@` та `.` у домені.
Тип `password` приховує вхідні дані і попереджає, якщо сайт не використовує HTTPS.

# --hints--

Ви повинні дати першому елементу `input` атрибут `type` зі значенням `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Ви повинні дати другому елементу `input` атрибут `type` зі значенням `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Ви повинні дати третьому елементу `input` атрибут `type` зі значенням `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

Ви повинні дати четвертому елементу `input` атрибут `type` зі значенням `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
      <label for="email">Enter Your Email: <input id="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
