---
id: 60fab4a123ce4b04526b082b
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Наразі користувачі можуть відправляти форму, не вибираючи жоден із радіокнопок. Хоча ви раніше використовували атрибут `required` для позначення обов’язкового вхідного даного, це не спрацює в цьому випадку, бо додавання `required` до обох вхідних даних передасть користувачам неправильну інформацію.

Щоб це виправити, ви можете надати контекст того, що потрібно, додавши елемент `legend` з текстом `Account type (required)` перед елементами `label` у другому `fieldset`. Потім додайте атрибут `checked` до вхідних даних `Personal`, щоб гарантувати, що форма буде відправлена з необхідними даними.

# --hints--

Ваші елементи `input` мають залишатися `type` відносно `radio`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

Ви маєте додати елемент `legend` всередину другого елемента `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

Ви маєте додати текст `Account type (required)` до елемента `legend`.

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

Ви маєте надати атрибут `checked` вхідним даним `Personal`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
      <label><input type="radio" name="account-type" value="personal" /> Personal</label>
      <label><input type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
