---
id: 63541ef4f96cd82e8e6c788a
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Атрибут `method` визначає, як надсилати дані форми на URL, вказаний у атрибуті `action`. Дані форми можна надсилати через `GET` запит як параметри URL (за допомогою `method="get"`) або через `POST` запит як дані в тілі запиту (за допомогою `method="post"`).

Встановіть атрибут `method`, щоб надсилати дані форми через `POST` запит.

**ПРИМІТКА**: Ви дізнаєтеся більше про методи `GET` і `POST` у наступних уроках і практичних заняттях.

# --hints--

Вам не потрібно додавати новий елемент `form`.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

Ваш елемент `form` має містити атрибут `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

Атрибут `method` має бути встановлений у `post`.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
