---
id: 5c6c06847491271903d37cfd
title: Tumia sifa ya kipengele ya value na vitufe vya radio na visanduku vya kuchagua
challengeType: 0
forumTopicId: 301099
dashedName: use-the-value-attribute-with-radio-buttons-and-checkboxes
---

# --description--

Wakati fomu inapotumwa, data hutumwa kwa seva na inajumuisha maingizo ya chaguzi zilizochaguliwa. Ingizo za aina `radio` na `checkbox` hutorosha thamani zao kutoka kwa sifa ya kipengele ya `value`.

Kwa mfano:

```html
<label for="indoor">
  <input id="indoor" value="indoor" type="radio" name="indoor-outdoor">Indoor
</label>
<label for="outdoor">
  <input id="outdoor" value="outdoor" type="radio" name="indoor-outdoor">Outdoor
</label>
```

Hapa, una ingizo mbili za `radio`. Mtumiaji anapowasilisha fomu akiwa amechagua chaguo la `indoor`, data ya fomu itajumuisha mstari: `indoor-outdoor=indoor`. Hii inatokana na sifa za kipengele za `name` na `value` za ingizo la "indoor".

Ikiwa utapuuza sifa ya kipengele ya `value`, data ya fomu itatumia thamani ya msingi, ambayo ni `on`. Katika hali hii, ikiwa mtumiaji atabofya chaguo la "indoor" na kuwasilisha fomu, data itakayotokea itakuwa `indoor-outdoor=on`, ambayo si ya manufaa. Hivyo sifa ya kipengele ya `value` inahitaji kuwekwa ili kutambua chaguo hilo.

# --instructions--

Toa sifa ya kipengele ya `radio` kwa kila ingizo la `checkbox` na `value` lililopo. Usitengeneze vitufe vipya vya radio au visanduku vya kuchagua. Tumia maandishi ya lebo ya ingizo, kwa herufi ndogo, kama thamani ya sifa ya kipengele.

# --hints--

Moja ya vitufe vyako vya radio inapaswa kuwa na sifa ya kipengele ya `value` yenye thamani ya `indoor`.

```js
const indoorRadioButton = [...document.querySelectorAll('label > input[type="radio"]')].filter(x => x.value === "indoor");
assert.notEmpty(indoorRadioButton)
```

Moja ya vitufe vyako vya radio inapaswa kuwa na sifa ya kipengele ya `value` yenye thamani ya `outdoor`.

```js
const outdoorRadioButton = [...document.querySelectorAll('label > input[type="radio"]')].filter(x => x.value === "outdoor");
assert.notEmpty(outdoorRadioButton);
```

Moja ya visanduku vyako vya kuchagua inapaswa kuwa na sifa ya kipengele ya `value` yenye thamani ya `loving`.

```js
const lovingCheckbox = [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "loving");
assert.notEmpty(lovingCheckbox); 
```

Moja ya visanduku vyako vya kuchagua inapaswa kuwa na sifa ya kipengele ya `value` yenye thamani ya `lazy`.

```js
const lazyCheckbox = [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "lazy");
assert.notEmpty(lazyCheckbox); 
```

Moja ya visanduku vyako vya kuchagua inapaswa kuwa na sifa ya kipengele ya `value` yenye thamani ya `energetic`.

```js
const energeticCheckbox =  [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "energetic");
assert.notEmpty(energeticCheckbox);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
