---
id: bad87fee1348bd9aedf08804
title: Toa maoni kwenye HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cGyGbca'
forumTopicId: 16782
dashedName: comment-out-html
---

# --description--

Kumbuka kwamba ili kuanza maoni, unahitaji kutumia `<!--` na kumaliza maoni, unahitaji kutumia `-->`

Hapa utahitaji kumaliza maoni kabla ya kipengele chako cha `h2` kuanza.

# --instructions--

Toa maoni kwenye kipengele chako cha `h1` na kipengele chako cha `p`, lakini usitoe maoni kwenye kipengele chako cha `h2`.

# --hints--

Kipengele chako cha `h1` kinapaswa kutoa maoni ili kisionekane kwenye ukurasa.

```js
assert.isEmpty(document.querySelectorAll('h1'));
```

Kipengele chako cha `h2` hakipaswi kutoa maoni ili kiweze kuonekana kwenye ukurasa.

```js
assert.isNotEmpty(document.querySelectorAll('h2'));
```

Kipengele chako cha `p` kinapaswa kutoa maoni ili kisionekane kwenye ukurasa.

```js
assert.isEmpty(document.querySelectorAll('p'));
```

Kila maoni yako yanapaswa kufungwa na `-->`.

```js
assert.isAbove(code.match(/[^fc]-->/g).length, 1);
```

Haupasi kubadilisha mpangilio wa vipengele vya `h1`, `h2`, au `p` katika msimbo.

```js
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[0],'<h1>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[1],'<h2>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[2],'<p>');
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<!--<h1>Hello World</h1>-->
<h2>CatPhotoApp</h2> 
<!--<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p> -->
```
