---
id: bad87fee1348bd9aedf08820
title: Geuza picha kuwa kiungo
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cRdBnUr'
forumTopicId: 18327
dashedName: turn-an-image-into-a-link
---

# --description--

Unaweza kufanya vipengele kuwa viungo kwa kuzipanga ndani ya kipengele cha `a`.

Panga picha yako ndani ya kipengele cha `a`. Hapa kuna mfano:

```html
<a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
```

Kumbuka kutumia `#` kama sifa ya `a` ya kipengele chako cha `href` ili kuibadilisha kuwa kiungo kisicho na maana.

# --instructions--

Weka kipengele cha picha kilichopo ndani ya kipengele cha `a` (*tagi ya nanga*).

Mara tu utakapofanya hivyo, pita kidole cha panya juu ya picha yako. Kidole cha panya kitabadilika kutoka kidole cha kawaida kuwa kidole cha kubofya kiungo. Picha sasa ni kiungo.

# --hints--

Kipengele cha `img` kilichopo kinapaswa kupangwa ndani ya kipengele cha `a`.

```js
const anchor = document.querySelectorAll('a')[1];
const children = anchor.querySelectorAll("img");
assert.notEmpty(children);
```

Kipengele chako cha `a` kinapaswa kuwa kiungo kisicho na maana chenye sifa ya `href` iliyowekwa kuwa `#`.

```js
const anchor = document.querySelectorAll('a')[1];
const parentHREF = anchor.querySelector("img").parentNode.getAttribute('href');
assert.match(parentHREF,new RegExp('#'));
```

Kila kipengele chako cha `a` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length, code.match(/<a/g).length)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
