---
id: bad87fee1348bd9aede08835
title: Panga ndani vipengele vingi ndani ya kipengele kimoja cha div
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cNW4kC3'
forumTopicId: 18246
dashedName: nest-many-elements-within-a-single-div-element
---

# --description--

Kipengele cha `div`, kinachojulikana pia kama kipengele cha mgawanyo, ni kontena la matumizi ya jumla kwa vipengele vingine.

Kipengele cha `div` ndicho kipengele cha HTML kinachotumika zaidi kati ya vyote.

Kama kipengele kingine chochote kisichojifunga yenyewe, unaweza kufungua kipengele cha `div` kwa `<div>` na kukifunga kwenye mstari mwingine kwa `</div>`.

# --instructions--

Panga ndani orodha zako za "Mambo ambayo paka wanapenda" na "Mambo 3 ya juu ambayo paka wanakataa" zote ndani ya kipengele kimoja cha `div`.

Kidokezo: Jaribu kuweka lebo yako ya kufungua `div` juu ya kipengele chako cha `p` cha "Mambo ambayo paka wanapenda" na lebo yako ya kufunga `div` baada ya lebo yako ya kufunga `ol` ili orodha zako zote ziwe ndani ya `div` moja.

# --hints--

Vipengele vyako vya `p` vinapaswa kupangwa ndani ya kipengele chako cha `div`.

```js
const div = document.querySelector('div');
const children = div.querySelectorAll('p');
assert.isAbove(children.length,1);
```

Kipengele chako cha `ul` kinapaswa kupangwa ndani ya kipengele chako cha `div`.

```js
const div = document.querySelector('div');
const children = div.querySelectorAll('ul');
assert.notEmpty(children);
```

Kipengele chako cha `ol` kinapaswa kupangwa ndani ya kipengele chako cha `div`.

```js
const div = document.querySelector('div');
const children = div.querySelectorAll('ol');
assert.notEmpty(children);
```

Kipengele chako cha `div` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/div>/g);
assert.strictEqual(code.match(/<\/div>/g).length,code.match(/<div>/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>

  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving" checked> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving" checked> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
