---
id: bad87fee1348bd9aedf08816
title: Unganisha na kurasa za nje kwa kutumia vipengele vya nanga
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c8EkncB'
forumTopicId: 18226
dashedName: link-to-external-pages-with-anchor-elements
---

# --description--

Unaweza kutumia vipengele vya `a` (*tagi ya nanga*) kuunganisha na maudhui yaliyoko nje ya ukurasa wako wa mtandao.

Vipengele vya `a` vinahitaji anwani ya mtandao inayoitwa sifa ya kipengele ya `href`. Pia vinahitaji maandishi ya nanga. Hapa kuna mfano:

```html
<a href="https://www.freecodecamp.org">this links to freecodecamp.org</a>
```

Kisha kivinjari chako kitaonyesha maandishi ya `this links to freecodecamp.org` kama kiungo unachoweza kubofya. Na kiungo hicho kitakupeleka kwenye anwani ya mtandao ya `https://www.freecodecamp.org`.

# --instructions--

Tengeneza kipengele cha `a` kinachounganisha na `https://www.freecatphotoapp.com` na kina maandishi ya nanga "picha za paka".

# --hints--

Kipengele chako cha `a` kinapaswa kuwa na maandishi ya nanga ya `cat photos`.

```js
assert.match(document.querySelector('a').textContent,/cat photos/gi);
```

Unahitaji kipengele cha `a` kinachounganisha na `https://www.freecatphotoapp.com`

```js
assert.match(document.querySelector('a').getAttribute('href'),/^https?:\/\/(www\.)?freecatphotoapp\.com\/?$/i);
```

Kipengele chako cha `a` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/a>/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>



  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <a href="https://www.freecatphotoapp.com">cat photos</a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
