---
id: bad87fee1348bd9aedf08827
title: Unda orodha isiyopangwa yenye alama za risasi
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cDKVPuv'
forumTopicId: 16814
dashedName: create-a-bulleted-unordered-list
---

# --description--

HTML ina kipengele maalum cha kuunda <dfn>orodha isiyopangwa</dfn>, au orodha za mtindo wa alama za risasi.

Orodha isiyopangwa huanza na kipengele cha kufungua `<ul>`, kifuatiwe na idadi yoyote ya vipengele vya `<li>`. Mwishowe, orodha isiyopangwa hufungwa na `</ul>`.

Kwa mfano:

```html
<ul>
  <li>milk</li>
  <li>cheese</li>
</ul>
```

ingatengeneza orodha ya mtindo wa alama za risasi ya `milk` na `cheese`.

# --instructions--

Ondoa vipengele viwili vya mwisho vya `p` na unda orodha isiyopangwa ya vitu vitatu vinavyopendwa na paka chini ya ukurasa.

# --hints--

Unda kipengele cha `ul`.

```js
assert.isNotEmpty(document.querySelectorAll('ul'));
```

Unapaswa kuwa na vipengele vitatu vya `li` ndani ya kipengele chako cha `ul`.

```js
assert.lengthOf(document.querySelectorAll('ul li'),3);
```

Kipengele chako cha `ul` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/ul>/gi);
assert.match(code,/<ul/gi);
assert.strictEqual(code.match(/<\/ul>/gi).length,code.match(/<ul/gi).length);
```

Vipengele vyako vya `li` vinapaswa kuwa na lebo za kufunga.

```js
assert.match(code,/<\/li>/gi);
assert.match(code,/<li[\s>]/gi);
assert.strictEqual(code.match(/<\/li>/gi).length,code.match(/<li[\s>]/gi).length);
```

Vipengele vyako vya `li` havipaswi kuwa na mfuatano wa herufi tupu au nafasi tu nyeupe.

```js
assert.isEmpty([...document.querySelectorAll('ul li')].filter((item) => item.textContent.trim() === ""));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <ul>
    <li>milk</li>
    <li>mice</li>
    <li>catnip</li>
  </ul>
</main>
```
