---
id: bad87fee1348bd9aedf08817
title: Tengeneza viungo visivyo na maana kwa kutumia alama ya hash
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMdkytL'
forumTopicId: 18230
dashedName: make-dead-links-using-the-hash-symbol
---

# --description--

Wakati mwingine unataka kuongeza vipengele vya `a` kwenye tovuti yako kabla hujajua vinavyotakuwa na viungo vipi.

Hii pia ni rahisi wakati unabadilisha tabia ya kiungo kwa kutumia `JavaScript`, ambayo tutajifunza baadaye.

# --instructions--

Thamani ya sasa ya sifa ya kipengele ya `href` ni kiungo kinachoelekeza kwa "`https://www.freecatphotoapp.com`". Badilisha thamani ya sifa ya kipengele ya `href` na `#`, inayojulikana pia kama alama ya hash, ili kuunda kiungo kisicho na maana.

Kwa mfano: `href="#"`

# --hints--

Kipengele chako cha `a` kinapaswa kuwa kiungo kisicho na maana na thamani ya sifa ya kipengele ya `href` iwe "#".

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'),'#');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#" target="_blank">cat photos</a>.</p>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
