---
id: bad87fee1348bd9aedf08812
title: Ongeza picha kwenye tovuti yako
challengeType: 0
forumTopicId: 16640
dashedName: add-images-to-your-website
---

# --description--

Unaweza kuongeza picha kwenye tovuti yako kwa kutumia kipengele cha `img`, na kuelekeza kwenye URL ya picha maalum kwa kutumia sifa ya kipengele ya `src`.

Mfano wa hili utakuwa:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg">
```

Kumbuka kuwa `img` ni kipengele kisicho na maudhui.

Vipengele vyote vya `img` **vinapaswa** kuwa na sifa ya kipengele ya `alt`. Maandishi yaliyomo ndani ya sifa ya kipengele ya `alt` hutumika kwa kisomaji cha skrini kuboresha Ufikikaji na huonyeshwa ikiwa picha haitapakia.

**Kumbuka:** Ikiwa picha ni ya mapambo tu, kutumia sifa ya kipengele ya `alt` isiyo na maudhui ni mbinu bora.

Kwa kawaida sifa ya kipengele ya `alt` haipaswi kuwa na herufi maalum isipokuwa inahitajika.

Tuweke sifa ya kipengele ya `alt` kwenye mfano wetu wa `img` hapo juu:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg" alt="freeCodeCamp logo">
```

# --instructions--

Tujaribu kuongeza picha kwenye tovuti yetu:

Ndani ya kipengele cha `main` kilichopo, weka kipengele cha `img` kabla ya vipengele vya `p` vilivyopo.

Sasa weka sifa ya kipengele ya `src` ili ielekeze kwenye URL `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

Mwishowe, usisahau kumpa kipengele chako cha `img` sifa ya kipengele ya `alt` yenye maandishi yanayofaa.

# --hints--

Ukurasa wako unapaswa kuwa na kipengele cha picha.

```js
assert.exists(document.querySelector('img'));
```

Picha yako inapaswa kuwa na sifa ya kipengele ya `src` inayolekeza kwenye picha ya paka mdogo.

```js
const url = document.querySelector('img').getAttribute('src');
assert.match(url,/^https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg$/i);
```

Sifa ya kipengele ya `alt` ya kipengele chako cha picha haipaswi kuwa tupu.

```js
assert.exists(document.querySelector('img').getAttribute('alt'));
assert.isNotEmpty(document.querySelector('img').getAttribute('alt'));
assert.match(__helpers.removeWhiteSpace(code),/<(?:img|IMG)\S*alt=(['"])(?!\1|>)\S+\1\S*\/?>/)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>


  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
