---
id: bad87fee1348bd9aedf08802
title: Futa maoni ya HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBmG9T7'
forumTopicId: 18329
dashedName: uncomment-html
---

# --description--

Kuweka maoni ni njia ambayo unaweza kuacha maoni kwa wasanidi programu wengine ndani ya msimbo wako bila kuathiri matokeo yanayoonyeshwa kwa mtumiaji wa mwisho.

Kuweka maoni pia ni njia rahisi ya kufanya msimbo usitumike bila kulazimika kufuta kabisa.

Maoni katika HTML huanza na `<!--` na kumalizika na `-->`

# --instructions--

Futa maoni ya vipengele vyako vya `h1`, `h2` na `p`.

# --hints--

Kipengele chako cha `h1` kinapaswa kuonekana kwenye ukurasa kwa kufuta maoni yake.

```js
assert.notEmpty(document.querySelectorAll('h1'));
```

Kipengele chako cha `h2` kinapaswa kuonekana kwenye ukurasa kwa kufuta maoni yake.

```js
assert.notEmpty(document.querySelectorAll('h2'));
```

Kipengele chako cha `p` kinapaswa kuonekana kwenye ukurasa kwa kufuta maoni yake.

```js
assert.notEmpty(document.querySelectorAll('p'));
```

Hakuna lebo za maoni zinazofuata zinapaswa kuonekana kwenye ukurasa (yaani `-->`).

```js
const elements = document.querySelectorAll('*');
const potentialComments = Array.from(elements).filter(el => el.textContent.includes('-->'));
assert.notExists(potentialComments[1])
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
