---
id: bad87fee1348bd9aedf08828
title: Unda orodha iliyopangwa
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cQ3B8TM'
forumTopicId: 16824
dashedName: create-an-ordered-list
---

# --description--

HTML ina kipengele kingine maalum cha kuunda <dfn>orodha zilizopangwa</dfn>, au orodha zilizo na nambari.

Orodha zilizopangwa huanza na kipengele cha kufungua `<ol>`, kifuatiwe na idadi yoyote ya vipengele vya `<li>`. Mwisho, orodha zilizopangwa hufungwa kwa lebo ya kufunga `</ol>`.

Kwa mfano:

```html
<ol>
  <li>Garfield</li>
  <li>Sylvester</li>
</ol>
```

ingatengeneza orodha yenye nambari ya `Garfield` na `Sylvester`.

# --instructions--

Unda orodha iliyopangwa ya vitu 3 vya juu ambavyo paka wanavichukia zaidi.

# --hints--

Unapaswa kuwa na orodha iliyopangwa kwa ajili ya `Top 3 things cats hate:`

```js
const previousElement = document.querySelector('ol').previousElementSibling; 
assert.match(previousElement.textContent,/Top 3 things cats hate:/i);
```

Unapaswa kuwa na orodha isiyopangwa kwa ajili ya `Things cats love:`

```js
const previousElement = document.querySelector('ul').previousElementSibling; 
assert.match(previousElement.textContent,/Things cats love:/i);
```

Unapaswa kuwa na kipengele kimoja tu cha `ul`.

```js
assert.lengthOf(document.querySelectorAll('ul'), 1);
```

Unapaswa kuwa na kipengele kimoja tu cha `ol`.

```js
assert.lengthOf(document.querySelectorAll('ol'), 1);
```

Unapaswa kuwa na vipengele vitatu vya `li` ndani ya kipengele chako cha `ul`.

```js
assert.lengthOf(document.querySelectorAll('ul li'), 3);
```

Unapaswa kuwa na vipengele vitatu vya `li` ndani ya kipengele chako cha `ol`.

```js
assert.lengthOf(document.querySelectorAll('ol li'), 3);
```

Kipengele chako cha `ul` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/ul>/g);
assert.strictEqual(code.match(/<\/ul>/g).length ,code.match(/<ul>/g).length);
```

Kipengele chako cha `ol` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/ol>/g);
assert.strictEqual(code.match(/<\/ol>/g).length ,code.match(/<ol>/g).length);
```

Kipengele chako cha `li` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/li>/g);
assert.match(code,/<li>/g);
assert.strictEqual(code.match(/<\/li>/g).length ,code.match(/<li>/g).length);
```

Vipengele vya `li` katika orodha isiyopangwa havipaswi kuwa tupu.

```js
[...document.querySelectorAll('ul li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

Vipengele vya `li` katika orodha iliyopangwa havipaswi kuwa tupu.

```js
[...document.querySelectorAll('ol li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>

</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>hate 1</li>
    <li>hate 2</li>
    <li>hate 3</li>
  </ol>
</main>
```
