---
id: bad87fee1348bd9aedc08830
title: Tumia HTML5 kuhitaji sehemu ya fomu
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMd4EcQ'
forumTopicId: 18360
dashedName: use-html5-to-require-a-field
---

# --description--

Unaweza kuhitaji sehemu maalum za fomu ili mtumizi asiweze kuwasilisha fomu yako hadi atakapojaza sehemu hizo.

Kwa mfano, ikiwa ungependa kufanya sehemu ya ingizo la maandishi iwe ya lazima, unaweza tu kuongeza sifa ya kipengele `required` ndani ya kipengele chako cha `input`, kama hii: `<input type="text" required>`

# --instructions--

Fanya `input` yako ya maandishi kuwa sehemu ya `required` ya lazima, ili mtumizi asiweze kuwasilisha fomu bila kujaza sehemu hii.

Kisha jaribu kuwasilisha fomu bila kuingiza maandishi yoyote. Tazama jinsi fomu yako ya HTML5 inavyokujulisha kuwa sehemu hii ni ya lazima?

# --hints--

Kipengele chako cha `input` cha maandishi kinapaswa kuwa na sifa ya kipengele `required`.

```js
assert.isTrue(document.querySelector('input').required);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" required placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```
