---
id: bad87fee1348bd9aedf08830
title: Ongeza maandishi ya kiwakilishi kwenye uwanja wa maandishi
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cKdJDhg'
forumTopicId: 16647
dashedName: add-placeholder-text-to-a-text-field
---

# --description--

Maandishi ya kiwakilishi ni yale yanayoonyeshwa katika kipengele chako cha `input` kabla mtumizi hajatoa ingizo lolote.

Unaweza kuunda maandishi ya kiwakilishi kama ifuatavyo:

```html
<input type="text" placeholder="this is placeholder text">
```

**Kumbuka:** Kumbuka kwamba `input` ni kipengele kisicho na maudhui.

# --instructions--

Weka thamani ya `placeholder` ya kipengele chako cha maandishi cha `input` kuwa "cat photo URL".

# --hints--

Unapaswa kuongeza sifa ya kipengele ya `placeholder` kwenye kipengele cha maandishi cha `input` kilichopo.

```js
assert.notEmpty(document.querySelectorAll('input[placeholder]'));
```

Unapaswa kuweka thamani ya sifa ya kipengele ya `placeholder` kuwa `cat photo URL`.

```js
assert.exists(document.querySelector('input'));
assert.exists(document.querySelector('input').getAttribute('placeholder'));
const placeholder = document.querySelector('input').getAttribute('placeholder');
assert.match(placeholder,/cat\s+photo\s+URL/gi);
```

Kipengele cha `input` kilichokamilika hakipaswi kuwa na lebo ya kufunga.

```js
assert.notMatch(code,/<input.*\/?>.*<\/input>/gi);
```

Kipengele cha `input` kilichokamilika kinapaswa kuwa na sintaksia halali.

```js
assert.notEmpty(document.querySelectorAll('input[type=text]'));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```
