---
id: bad87fee1348bd9aedf08801
title: Taarifu kwa kutumia kipengele cha aya
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ceZ7DtN'
forumTopicId: 18202
dashedName: inform-with-the-paragraph-element
---

# --description--

Kipengele cha `p` ndicho kipengele kinachopendekezwa kwa maandishi ya aya kwenye tovuti. `p` ni kifupi cha "aya".

Unaweza kuunda kipengele cha aya kama hii:

```html
<p>I'm a p tag!</p>
```

# --instructions--

Unda kipengele cha `p` chini ya kipengele chako cha `h2`, na kipa maandishi ya `Hello Paragraph`.

**Kumbuka:** Kulingana na mkataba, lebo zote za HTML huandikwa kwa herufi ndogo, kwa mfano `<p></p>` na si `<P></P>`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele halali cha `p`.

```js
assert.lengthOf(document.querySelectorAll('p'),1);
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi ya `Hello Paragraph`.

```js
assert.match(document.querySelector('p').textContent,/hello(\s)+paragraph/gi);
```

Kipengele chako cha `p` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
```

# --solutions--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
<p>Hello Paragraph</p>
```
