---
id: bad87fee1348bd9aede08830
title: Unda kipengele cha fomu
challengeType: 0
forumTopicId: 16817
dashedName: create-a-form-element
---

# --description--

Unaweza kujenga fomu za mtandao ambazo zinawasilisha data kwa seva kwa kutumia HTML safi tu. Unaweza kufanya hivi kwa kuweka sifa ya kipengele `action` kwenye kipengele chako cha `form`.

Kwa mfano:

```html
<form action="url-where-you-want-to-submit-form-data">
  <input>
</form>
```

# --instructions--

Panga ndani kipengele kilichopo cha `input` ndani ya kipengele cha `form` na weka thamani `"https://www.freecatphotoapp.com/submit-cat-photo"` kwenye sifa ya kipengele `action` ya kipengele cha `form`.

# --hints--

Kipengele kilichopo cha `input` kinapaswa kupangwa ndani ya kipengele cha `form`.

```js
const inputElem = document.querySelector('form input');
assert.strictEqual(inputElem.getAttribute('type'),'text'); 
assert.strictEqual(inputElem.getAttribute('placeholder'),'cat photo URL');
```

`form` yako inapaswa kuwa na sifa ya kipengele `action` ambayo imewekwa kuwa `https://www.freecatphotoapp.com/submit-cat-photo`.

```js
const action = document.querySelector('form').getAttribute('action');
assert.match(action,/^https:\/\/(www\.)?freecatphotoapp\.com\/submit-cat-photo$/i);
```

Kipengele chako cha `form` kinapaswa kuwa na lebo za kufungua na kufunga zilizo sawa.

```js
assert.match(code,/<\/form>/g);
assert.match(code,/<form [^<]*>/g);
assert.strictEqual(code.match(/<\/form>/g).length,code.match(/<form [^<]*>/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
  </form>
</main>
```
