---
id: bad88fee1348bd9aedf08816
title: Unganisha na Sehemu za Ndani za Ukurasa kwa Kutumia Vipengele vya Nanga
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cyrDRUL'
forumTopicId: 301098
dashedName: link-to-internal-sections-of-a-page-with-anchor-elements
---

# --description--

Vipengele vya `a` (*tagi ya nanga*) vinaweza pia kutumika kuunda viungo vya ndani vya kuruka kwenda sehemu tofauti ndani ya ukurasa wa mtandao.

Ili kuunda kiungo cha ndani, weka thamani ya sifa ya `href` ya kiungo kuwa alama ya hash `#` pamoja na thamani ya sifa ya `id` ya kipengele unachotaka kuunganisha ndani, kawaida kilicho chini zaidi ya ukurasa. Kisha unahitaji kuongeza sifa hiyo hiyo ya `id` kwa kipengele unachounganishwa nacho. Sifa ya `id` ni sifa inayofafanua kipengele kwa kipekee.

Hapo chini ni mfano wa kiungo cha nanga cha ndani na kipengele kinacholengwa:

```html
<a href="#contacts-header">Contacts</a>
...
<h2 id="contacts-header">Contacts</h2>
```

Wakati watumizi wanapobofya kiungo cha `Contacts`, wataelekezwa kwenye sehemu ya ukurasa wa mtandao yenye kipengele cha kichwa cha sehemu cha **Contacts**.

# --instructions--

Badilisha kiungo chako cha nje kuwa kiungo cha ndani kwa kubadilisha sifa ya `href` kuwa `#footer` na maandishi kutoka `cat photos` hadi `Jump to Bottom`.

Ondoa sifa ya `target="_blank"` kutoka kwenye tagi ya nanga kwa sababu hii husababisha hati iliyounganishwa kufunguliwa kwenye kichupo kipya cha dirisha.

Kisha ongeza sifa ya `id` yenye thamani ya `footer` kwa kipengele cha `<footer>` kilicho chini ya ukurasa.

# --hints--

Kwenye ukurasa wako inapaswa kuwepo na tagi moja tu ya nanga.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Kwenye ukurasa wako inapaswa kuwepo na tagi moja tu ya `footer`.

```js
assert.lengthOf(document.querySelectorAll('footer'),1);
```

Tagi ya `a` inapaswa kuwa na sifa ya `href` iliyowekwa kuwa "#footer".

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'), '#footer');
```

Tagi ya `a` haipaswi kuwa na sifa ya `target`.

```js
assert.notExists(document.querySelector('a').getAttribute('target'));
```

Maandishi ya `a` yanapaswa kuwa "Jump to Bottom".

```js
assert.match(document.querySelector('a').textContent,/Jump to Bottom/gi);
```

Tagi ya `footer` inapaswa kuwa na sifa ya `id` iliyowekwa kuwa "footer".

```js
assert.strictEqual(document.querySelector('footer').getAttribute('id'),'footer');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer>Copyright Cat Photo App</footer>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="#footer">Jump to Bottom</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer id="footer">Copyright Cat Photo App</footer>
```
