---
id: bad87fee1348bd9aecf08801
title: Utangulizi wa vipengele vya HTML5
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBkZGpt7'
forumTopicId: 301097
dashedName: introduction-to-html5-elements
---

# --description--

HTML5 inaleta lebo za HTML zenye maelezo zaidi. Hizi ni pamoja na `main`, `header`, `footer`, `nav`, `video`, `article`, `section` na zingine.

Leboz hizi hutoa muundo wa kueleweka zaidi kwa HTML yako, hufanya HTML yako iwe rahisi kusoma, na husaidia katika Uboreshaji wa Injini za Utafutaji (SEO) na Ufikikaji. Lebo ya `main` ya HTML5 husaidia injini za utafutaji na wasanidi programu wengine kupata maudhui kuu ya ukurasa wako.

Mfano wa matumizi, kipengele cha `main` chenye vipengele viwili vya watoto vilivyopangwa ndani yake:

```html
<main> 
  <h1>Hello World</h1>
  <p>Hello Paragraph</p>
</main>
```

**Kumbuka:** Lebo nyingi mpya za HTML5 na faida zao zinajadiliwa katika sehemu ya Ufikikaji Umewekwa.

# --instructions--

Tengeneza kipengele cha pili cha `p` chenye maandishi ya kitty ipsum yafuatayo: `Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.`

Kisha, tengeneza kipengele cha `main` na upange ndani yake tu vipengele viwili vya `p` ndani ya kipengele cha `main`.

# --hints--

Unapaswa kuwa na vipengele viwili vya `p` vyenye maandishi ya Kitty Ipsum.

```js
assert.lengthOf(document.querySelectorAll('p'),2);
```

Kipengele kila kimoja cha `p` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

Kipengele chako cha `p` kinapaswa kuwa na maneno machache ya kwanza ya maandishi ya ziada ya `kitty ipsum` yaliyotolewa.

```js
assert.match(document.querySelectorAll('p')[1].textContent,/Purr\s+jump\s+eat/gi);
```

Msimbo wako unapaswa kuwa na kipengele kimoja cha `main`.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Kipengele cha `main` kinapaswa kuwa na vipengele viwili vya aya kama watoto.

```js
const main = document.querySelector('main');
const children = main.querySelectorAll("p"); 
assert.lengthOf(children,2);
```

Lebo ya kufungua ya `main` inapaswa kuja kabla ya lebo ya aya ya kwanza.

```js
assert.match(code,(/<main>\s*?<p>/g));
```

Lebo ya kufunga ya `main` inapaswa kuja baada ya lebo ya kufunga ya aya ya pili.

```js
assert.match(code,(/<\/p>\s*?<\/main>/g));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
