---
id: bad87fee1348bd9aedf08833
title: Jaza nafasi tupu na maandishi ya kiwakilishi
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cgR7Dc7'
forumTopicId: 18178
dashedName: fill-in-the-blank-with-placeholder-text
---

# --description--

Wasanidi programu wa mtandao kwa kawaida hutumia <dfn>maandishi ya lorem ipsum</dfn> kama maandishi ya kiwakilishi. Maandishi ya lorem ipsum hupatikana kwa bahati nasibu kutoka sehemu maarufu ya Cicero wa Roma ya Kale.

Maandishi ya lorem ipsum yamekuwa yakitumika kama maandishi ya kiwakilishi na wapangaji maandishi tangu karne ya 16, na desturi hii inaendelea kwenye mtandao.

Vizuri, karne 5 ni muda mrefu wa kutosha. Kwa kuwa tunajenga CatPhotoApp, tumia kitu kinachoitwa maandishi ya "kitty ipsum".

# --instructions--

Badilisha maandishi ndani ya kipengele chako cha `p` na maneno machache ya kwanza ya maandishi haya ya kitty ipsum: `Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.`

# --hints--

Kipengele chako cha `p` kinapaswa kuwa na maneno machache ya kwanza ya maandishi ya "kitty ipsum" yaliyotolewa.

```js
assert.match(document.querySelector('p').textContent,/Kitty(\s)+ipsum/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Hello Paragraph</p>
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff</p>
```
