---
id: bad87fee1348bd9aedf08835
title: Unda seti ya visanduku vya kuchagua
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cqrkJsp'
forumTopicId: 16821
dashedName: create-a-set-of-checkboxes
---

# --description--

Fomu hutumia mara nyingi <dfn>visanduku vya kuchagua</dfn> kwa maswali ambayo yanaweza kuwa na majibu zaidi ya moja.

Visanduku vya kuchagua ni aina ya `input`.

Kila kisanduku chako cha kuchagua kinaweza kupangwa ndani ya kipengele chake cha `label`. Kwa kuzungusha kipengele cha `input` ndani ya kipengele cha `label` itahusisha moja kwa moja ingizo la kisanduku cha kuchagua na kipengele cha lebo kinachokizunguka.

Visanduku vyote vya kuchagua vinavyohusiana vinapaswa kuwa na sifa ya kipengele `name` sawa.

Inachukuliwa kuwa mbinu bora kufafanua wazi uhusiano kati ya `input` ya kisanduku cha kuchagua na `label` inayolingana nayo kwa kuweka sifa ya kipengele `for` kwenye kipengele cha `label` ili iendane na sifa ya kipengele `id` ya kipengele cha `input` kinachohusiana.

Hapa kuna mfano wa kisanduku cha kuchagua:

```html
<label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
```

# --instructions--

Ongeza kwenye fomu yako seti ya visanduku vitatu vya kuchagua. Kila kisanduku cha kuchagua kinapaswa kupangwa ndani ya kipengele chake cha `label`. Vyote vitatu vinapaswa kushiriki sifa ya kipengele `name` ya `personality`.

# --hints--

Ukurasa wako unapaswa kuwa na vipengele vitatu vya kisanduku cha kuchagua.

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'),3);
```

Kila kipengele chako cha kisanduku cha kuchagua kinapaswa kupangwa ndani ya kipengele chake cha `label`.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="checkbox"]:only-child'),3);
```

Hakikisha kila kipengele chako cha `label` kina lebo ya kufunga.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length)
```

Visanduku vyako vya kuchagua vinapaswa kupewa sifa ya kipengele `name` ya `personality`.

```js
assert.lengthOf([...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.name === "personality"),3);
```

Kila kisanduku chako cha kuchagua kinapaswa kuongezwa ndani ya lebo ya `form`.

```js
assert.strictEqual(document.querySelector('label').parentNode.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="playful"><input id="playful" type="checkbox" name="personality">Playful</label>
    <label for="lazy"><input id="lazy" type="checkbox" 
name="personality">Lazy</label>
    <label for="evil"><input id="evil" type="checkbox" 
name="personality">Evil</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
