---
id: 6133d11ef548f51f876149e3
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Navigation is a core part of accessibility, and screen readers rely on you to provide the structure of your page. This is accomplished with semantic HTML elements.

Aggiungi un elemento `header` e un elemento `main` alla pagina.

L'elemento `header` verrà utilizzato per introdurre la pagina, oltre a fornire un menù di navigazione.

L'elemento `main` conterrà il contenuto principale della tua pagina.

# --hints--

Dovresti aggiungere un elemento `header` all'interno del `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Dovresti aggiungere un elemento `main` all'interno del `body`.

```js
assert.exists(document.querySelector('body > main'));
```

L'elemento `header` dovrebbe precedere l'elemento `main`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
