---
id: 6690e10ebe2181212abc9652
title: створіть візитну картку
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати `div` з атрибутом `class` зі значенням `business-card`, який міститиме всі інші елементи.  
1. Всередині елемента `.business-card` має бути елемент `img` з атрибутом `class` зі значенням `profile-image`. За бажанням можна встановити джерело зображення `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`. Має бути `alt` із змістовним описом.  
1. Всередині елемента `.business-card`, після елемента `img`, має бути три елементи `p` з атрибутом `class`, які відповідно мають значення `full-name`, `designation` і `company`.  
1. Перший елемент `p` має містити ваше ім’я.  
1. Другий елемент `p` має містити вашу посаду.  
1. Третій елемент `p` має містити назву вашої компанії.  
1. Під третім елементом `hr` має бути елемент `p`.  
1. Після елемента `hr` мають бути два елементи `p`. У першому елементі `p` має бути текст з адресою е-пошти. У другому елементі `p` має бути текст з номером телефону.  
1. Після двох елементів `p` має бути елемент `a` з класом `portfolio-link`, текстом `Portfolio` і посиланням на дійсний URL.  
1. Після елемента `hr` має бути елемент `a` із посиланням на портфоліо.  
1. Має бути ще один елемент `div` з атрибутом `class` зі значенням `social-media`. Усередині цього елемента має бути елемент `h2` з текстом `Connect with me`.  
1. Усередині елементів `.social-media` має бути три елементи `a`. Елементи `a` мають відповідно містити тексти `Twitter`, `LinkedIn` і `GitHub` з посиланнями на дійсні вебсайти.  
1. Ви маєте підключити таблицю стилів `styles.css` за допомогою тега `link` у розділі `head` HTML-файлу.  
1. Фон вашої сторінки має бути `rosybrown`. Загальний шрифт має бути `Arial` із запасним варіантом `sans-serif`.  
1. Селектор `.business-card` має мати властивості, щоб встановити елемент шириною `300px` і фон на ваш вибір. Також потрібно встановити поле `padding` розміром `20px` з усіх боків і верхнє поле `margin` розміром `100px`. Текст має бути вирівняний по центру, а розмір шрифту — `16px`. Ліве та праве поле має бути встановлене на `auto`.  
1. Селектор `.profile-image` має мати властивість `max-width` зі значенням `100%`.  
1. Усі елементи `p` мають мати верхнє та нижнє поле `margin` розміром `5px`.  
1. Усі посилання на сторінці не повинні бути підкреслені.

# --hints--

Ви повинні мати один головний елемент `div` з атрибутом `class` зі значенням `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

Всередині елемента `.business-card` має бути елемент `img` з атрибутом `class` зі значенням `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

Джерело зображення має бути встановлене на дійсне зображення.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

Атрибут `alt` зображення має бути встановлений на змістовний опис.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

Всередині елемента `.business-card`, після елемента `img`, має бути елемент `p` з атрибутом `class` зі значенням `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Після елемента `.full-name` має бути елемент `p` з атрибутом `class` зі значенням `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Після елемента `.designation` має бути елемент `p` з атрибутом `class` зі значенням `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

Перший елемент `p`, той, що має клас `full-name`, має містити ваше ім’я.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

Другий елемент `p`, той, що має атрибут `class` зі значенням `designation`, має містити вашу посаду.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

Третій елемент `p`, той, що має атрибут `class` зі значенням `company`, має містити назву вашої компанії.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

Після третього елемента `p`, того, що має атрибут `class` зі значенням `company`, має бути елемент `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Після першого елемента `hr` має бути елемент `p` з текстом адреси е-пошти.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

Після елемента `p` з адресою е-пошти має бути ще один елемент `p` з текстом номера телефону.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Після елемента `p` з номером телефону має бути елемент `a` з класом `portfolio-link`.

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

Елемент `.portfolio-link` має містити текст `Portfolio`.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

Ваш елемент `.portfolio-link` має мати дійсний `href`.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

Після елемента `a`, що містить ваше портфоліо, має бути ще один елемент `hr` як роздільник.

```js
assert.exists(document.querySelector('a + hr'));
```

Після другого елемента `hr` має бути елемент `div` з атрибутом `class` зі значенням `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

Всередині елемента `.social-media` має бути елемент `h2` з текстом `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

Всередині елемента `.social-media` має бути три елементи `a`.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

Перший елемент `a` має містити текст `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

Атрибут `href` першого елемента `a` має направляти на дійсне посилання.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

Другий елемент `a` має містити текст `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

Атрибут `href` другого елемента `a` має направляти на дійсне посилання.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

Третій елемент `a` має містити текст `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

Атрибут `href` третього елемента `a` має направляти на дійсне посилання.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

Ваш елемент `link` має бути всередині вашого елемента `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Ваш елемент `link` має мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Ваш елемент `link` має мати атрибут `href` зі значенням `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Ви маєте встановити колір фону сторінки в селекторі `body` на `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

Ви маєте встановити шрифт сторінки `Arial` у елементі `body` із запасним варіантом `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

Ваш селектор `.business-card` має мати властивість `width` зі значенням `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

Ваш селектор `.business-card` має мати властивість `background-color` з дійсним значенням кольору.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Ваш селектор `.business-card` має мати властивість `padding` зі значенням `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

Ваш селектор `.business-card` має мати верхнє поле `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

Текст усередині елемента `.business-card` має бути вирівняний по центру.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

Розмір шрифту тексту всередині елемента `.business-card` має бути `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

Ліве та праве поле елемента `.business-card` має бути встановлене на `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

Ваш селектор `.profile-image` має мати властивість `max-width` зі значенням `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

Ваші елементи `p` мають мати верхнє та нижнє поле `margin` розміром `5px`.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

Посилання на сторінці не повинні бути підкреслені.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
