---
id: 66fbcf750a62784cf11f5630
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Você deve agora inserir os três objetos no array `recipes`. Para fazer isso, você pode usar o método `push()`.

Use o método `push()` para adicionar todos os objetos de receita ao array `recipes`. Certifique-se de enviar `recipe1`, `recipe2` e `recipe3` nessa ordem.

Também exclua as variáveis `recipe1Name`, `recipe2CookingTime` e `recipe3Ingredients` e as declarações `console.log` que registram essas variáveis.

# --hints--

Você deve adicionar todos os objetos de receita no array `recipes`. Certifique-se de adicioná-los na ordem em que são declarados.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

Você deve remover a variável `recipe1Name`.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

Você deve remover a instrução `console.log(recipe1Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

Você deve remover a variável `recipe2CookingTime`.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

Você deve remover a instrução `console.log(recipe2CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

Você deve remover a variável `recipe3Ingredients`.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

Você deve remover a instrução `console.log(recipe3Ingredients);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
