---
id: 66fbcf750a62784cf11f562e
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Crie um objeto `recipe2` com as seguintes propriedades e valores:

| Key | Value   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

Você deve criar um objeto chamado `recipe2`.

```js
assert.isObject(recipe2);
```

Seu objeto `recipe2` deve ter uma propriedade `name`.

```js
assert.property(recipe2, "name");
```

A propriedade `name` deve ser definida como `Chicken Curry`.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

Seu objeto `recipe2` deve ter uma propriedade `ingredients`.

```js
assert.property(recipe2, "ingredients");
```

A sua propriedade `ingredients` deve ter um valor do tipo array.

```js
assert.isArray(recipe2.ingredients);
```

O valor do array da sua propriedade `ingredients` deve conter `chicken breast`, `coconut milk`, `curry powder`, `onion` e `garlic`.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

Seu objeto `recipe2` deve ter uma propriedade `cookingTime`.

```js
assert.property(recipe2, "cookingTime");
```

O valor da sua propriedade `cookingTime` deve ser um número.

```js
assert.isNumber(recipe2.cookingTime);
```

Sua propriedade `cookingTime` deve ser definida como `42`.

```js
assert.equal(recipe2.cookingTime, 42);
```

Seu objeto `recipe2` deve ter uma propriedade `totalIngredients`.

```js
assert.property(recipe2, "totalIngredients");
```

A sua propriedade `totalIngredients` deve ser definida como `null`.

```js
assert.isNull(recipe2.totalIngredients);
```

Seu objeto `recipe2` deve ter uma propriedade `difficultyLevel`.

```js
assert.property(recipe2, "difficultyLevel");
```

A sua propriedade `difficultyLevel` deve ser definida como uma string vazia.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
