---
id: 66fbcf750a62784cf11f5633
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Crie uma função `getDifficultyLevel` que recebe um número indicando o tempo de cozimento como parâmetro.

Se o tempo de cozimento for menor ou igual a `30`, a função deve retornar `"easy"`.Se for menor ou igual a `60`, a função deve retornar `"medium"`.Caso contrário, a função deve retornar `"hard"`.

# --hints--

Você deve criar uma função `getDifficultyLevel`.

```js
assert.isFunction(getDifficultyLevel)
```

Sua função `getDifficultyLevel` deve ter um único parâmetro.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

Sua função `getDifficultyLevel` deve retornar `"easy"` quando o tempo de cozimento for menor ou igual a `30`.

```js
assert.strictEqual(getDifficultyLevel(10), "easy")
assert.strictEqual(getDifficultyLevel(20), "easy")
assert.strictEqual(getDifficultyLevel(29), "easy")
assert.strictEqual(getDifficultyLevel(30), "easy")
```

Sua função `getDifficultyLevel` deve retornar `"medium"` quando o tempo de cozimento for maior que `31` e menor ou igual a `60`.

```js
assert.strictEqual(getDifficultyLevel(31), "medium")
assert.strictEqual(getDifficultyLevel(40), "medium")
assert.strictEqual(getDifficultyLevel(50), "medium")
assert.strictEqual(getDifficultyLevel(60), "medium")
```

Sua função `getDifficultyLevel` deve retornar `"hard"` quando o tempo de cozimento for maior que `60`.

```js
assert.strictEqual(getDifficultyLevel(61), "hard")
assert.strictEqual(getDifficultyLevel(65), "hard")
assert.strictEqual(getDifficultyLevel(70), "hard")
assert.strictEqual(getDifficultyLevel(75), "hard")
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return total / ratings.length;
}

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
