---
id: 66fbcf750a62784cf11f5634
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

É hora de testar cada uma das funções.Você pode usar qualquer uma das receitas para isso, mas este tutorial começará com `recipe1`.

Crie três novas variáveis: `recipe1AverageRating`, `recipe1TotalIngredients` e `recipe1DifficultyLevel`.Atribua a eles os valores chamando a função correspondente para cada variável e passando a propriedade `recipe1` apropriada.

Finalmente, registre cada variável no console para ver os resultados.

# --hints--

Você deve criar uma variável `recipe1AverageRating`.

```js
assert.isNotNull(recipe1AverageRating)
```

Sua variável `recipe1AverageRating` deve ser definida como `getAverageRating(recipe1.ratings)`.

```js
assert.equal(recipe1AverageRating, getAverageRating(recipe1.ratings));
```

Você deve registrar `recipe1AverageRating` no console.

```js
assert.match(code, /console\.log\(recipe1AverageRating\);?/)
```

Você deve criar uma variável `recipe1TotalIngredients`.

```js
assert.isNotNull(recipe1TotalIngredients)
```

Sua variável `recipe1TotalIngredients` deve ser definida como `getTotalIngredients(recipe1.ingredients)`.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

Você deve registrar `recipe1TotalIngredients` no console.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/)
```

Você deve criar uma variável `recipe1DifficultyLevel`.

```js
assert.isNotNull(recipe1DifficultyLevel)
```

Sua variável `recipe1DifficultyLevel` deve ser definida como `getDifficultyLevel(recipe1.cookingTime)`.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

Você deve registrar `recipe1DifficultyLevel` no console.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return total / ratings.length;
}

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return 'easy';
  } else if (cookingTime <= 60) {
    return 'medium';
  } else {
    return 'hard';
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
