---
id: 66fbcdfed6848e48947dbbc9
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você criará um rastreador de receitas usando objetos JavaScript.

Comece criando um array vazio chamado `recipes`.Este é o array no qual você irá adicionar os objetos de receita mais tarde.

# --hints--

Você deve criar um array chamado `recipes`.

```js
assert.isArray(recipes)
```

Seu array `recipes` deve estar vazio.

```js
assert.isEmpty(recipes)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
