---
id: 66fbcf750a62784cf11f562d
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Crie um objeto chamado `recipe1`. Dentro do objeto `recipe1`, crie uma propriedade `name` com o valor `"Spaghetti Carbonara"`.

Também dentro do objeto `recipe1`, crie uma propriedade `ingredients` com um array como valor. O array deve conter `spaghetti`, `Parmesan cheese`, `pancetta` e `black pepper`.

# --hints--

Você deve criar um objeto chamado `recipe1`.

```js
assert.isObject(recipe1);
```

Seu objeto `recipe1` deve ter uma propriedade `name`.

```js
assert.property(recipe1, 'name');
```

A sua propriedade `name` deve ser definida como `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

Seu objeto `recipe1` deve ter uma propriedade `ingredients`.

```js
assert.property(recipe1, 'ingredients');
```

A sua propriedade `ingredients` deve ter um valor do tipo array.

```js
assert.isArray(recipe1.ingredients);
```

O valor do array da sua propriedade `ingredients` deve conter `spaghetti`, `Parmesan cheese`, `pancetta` e `black pepper`.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
