---
id: 66fbcf750a62784cf11f5635
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Agora você pode preencher cada item do array `recipes` com valores para as propriedades `totalIngredients` e `difficultyLevel`.

Por enquanto, acesse o `totalIngredients` e o `difficultyLevel` de `recipe1` e defina-os com os resultados apropriados das chamadas de função e argumentos.

# --hints--

Você deve acessar a propriedade `totalIngredients` de `recipe1`.

```js
assert.isNotNull(recipe1.totalIngredients);
```

Você deve atribuir o resultado da chamada de `getTotalIngredients` com `recipe1.ingredients` à propriedade `totalIngredients` de `recipe1`.

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

Você deve acessar a propriedade `difficultyLevel` de `recipe1`.

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

Você deve atribuir o resultado da chamada de `getDifficultyLevel` com `recipe1.cookingTime` à propriedade `cookingTime` de `recipe1`.

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
