---
id: 66fbcf750a62784cf11f5631
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Nos próximos passos você trabalhará no cálculo da avaliação média, do total de ingredientes e do nível de dificuldade para cada receita no array `recipes`.

Comece criando uma função `getAverageRating` que recebe um único argumento, que é um array com avaliações.Dentro da função, calcule a avaliação média a partir do array passado para a função.

Sua função `getAverageRating` deve retornar um número.

# --hints--

Você deve criar uma função `getAverageRating`.

```js
assert.isFunction(getAverageRating)
```

Sua função `getAverageRating` deve ter um único parâmetro.

```js
assert.lengthOf(getAverageRating, 1);
```

Sua função `getAverageRating` deve retornar um número.

```js
assert.isNumber(getAverageRating(recipe1.ratings))
```

`getAverageRating(recipe1.ratings)` deve retornar `4.5`.

```js
assert.strictEqual(getAverageRating(recipe1.ratings), 4.5)
```

`getAverageRating(recipe2.ratings)` deve retornar `4.75`.

```js
assert.strictEqual(getAverageRating(recipe2.ratings), 4.75)
```

`getAverageRating(recipe3.ratings)` deve retornar `4`.

```js
assert.strictEqual(getAverageRating(recipe3.ratings), 4)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
