---
id: 66fbcf750a62784cf11f5632
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Crie uma função `getTotalIngredients` que recebe um único argumento, representando um array com ingredientes, e retorna o número de ingredientes do array passado para a função.

# --hints--

Você deve criar uma função `getTotalIngredients`.

```js
assert.isFunction(getTotalIngredients);
```

Sua função `getTotalIngredients` deve ter um único parâmetro.

```js
assert.lengthOf(getTotalIngredients, 1);
```

Sua função `getTotalIngredients` deve retornar um número.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)` deve retornar `4`.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` deve retornar `5`.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` deve retornar `3`.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
