---
id: 66ff02792b3042a2b811f475
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Adicione as seguintes propriedades ao objeto `recipe1`:

| Chave | Valor   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

As propriedades com valores `null` e string vazia serão atualizadas depois que você calculá-las.

# --hints--

Seu objeto `recipe1` deve ter uma propriedade `cookingTime`.

```js
assert.property(recipe1, "cookingTime");
```

A sua propriedade `cookingTime` deve ser definida como `22`.

```js
assert.equal(recipe1.cookingTime, 22);
```

Seu objeto `recipe1` deve ter uma propriedade `totalIngredients`.

```js
assert.property(recipe1, "totalIngredients");
```

A sua propriedade `totalIngredients` deve ser definida como `null`.

```js
assert.isNull(recipe1.totalIngredients);
```

Seu objeto `recipe1` deve ter uma propriedade `difficultyLevel`.

```js
assert.property(recipe1, "difficultyLevel");
```

A sua propriedade `difficultyLevel` deve ser definida como uma string vazia.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
