---
id: 670e4f45f7116c0f216a5177
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Antes de continuar, você deve praticar como acessar propriedades de um objeto.

Você pode usar a notação de ponto (`.`) ou colchete (`[]`) para fazer isso. Aqui está um exemplo:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Acesse a propriedade `name` de `recipe1` e atribua-a à variável `recipe1Name`.

Em seguida, acesse a propriedade `cookingTime` de `recipe2` e atribua-a à variável `recipe2CookingTime`.

Finalmente, acesse a propriedade `ingredients` de `recipe3` e atribua-a à variável `recipe3Ingredients`.

Certifique-se de que todas as variáveis que você criou estejam registradas no console.

# --hints--

Você deve criar uma variável `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

Você deve atribuir o valor da propriedade `name` de `recipe1` à sua variável `recipe1name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Você deve registrar `recipe1Name` no console.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Você deve criar uma variável `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

Você deve atribuir o valor da propriedade `cookingTime` de `recipe2` à sua variável `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Você deve registrar `recipe2CookingTime` no console.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

Você deve criar uma variável `recipe3Ingredients`.

```js
assert.isDefined(recipe3Ingredients);
```

Você deve atribuir o valor da propriedade `ingredients` de `recipe3` à sua variável `recipe3Ingredients`.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

Você deve registrar `recipe3Ingredients` no console.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
