---
id: 670e4f45f7116c0f216a5177
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Antes de continuar, você deve praticar como acessar propriedades de um objeto.

Você pode usar a notação de ponto (`.`) ou colchete (`[]`) para fazer isso.Aqui está um exemplo:

```js
const person = {
  name: 'John',
  age: 30,
  job: 'Software Engineer'
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Acesse as propriedades `name` de `recipe1` e `recipe2` e atribua-as às variáveis `recipe1Name` e `recipe2Name`, respectivamente.

Em seguida, acesse as propriedades `cookingTime` de ambas as receitas e atribua-as às variáveis `recipe1CookingTime` e `recipe2CookingTime`, respectivamente.

Certifique-se de que todas as variáveis que você criou estejam registradas no console.

# --hints--

Você deve criar uma variável `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

Você deve atribuir o valor da propriedade `name` de `recipe1` à sua variável `recipe1name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Você deve registrar `recipe1Name` no console.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Você deve criar uma variável `recipe2Name`.

```js
assert.isDefined(recipe2Name);
```

Você deve atribuir o valor da propriedade `name` de `recipe2` à sua variável `recipe2name`.

```js
assert.strictEqual(recipe2Name, recipe2.name);
```

Você deve registrar `recipe2Name` no console.

```js
assert.match(code, /console\.log\(recipe2Name\)/);
```

Você deve criar uma variável `recipe1CookingTime`.

```js
assert.isDefined(recipe1CookingTime);
```

Você deve atribuir o valor da propriedade `cookingTime` de `recipe1` à sua variável `recipe1CookingTime`.

```js
assert.strictEqual(recipe1CookingTime, recipe1.cookingTime);
```

Você deve registrar `recipe1CookingTime` no console.

```js
assert.match(code, /console\.log\(recipe1CookingTime\)/);
```

Você deve criar uma variável `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

Você deve atribuir o valor da propriedade `cookingTime` de `recipe2` à sua variável `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Você deve registrar `recipe2CookingTime` no console.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

--fcc-editable-region--

--fcc-editable-region--
```
