---
id: 69161f4c0c9c8c48b3c843c5
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Crie um objeto `recipe3` com as seguintes propriedades e valores:

|Chave|Valor|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

Você deve criar um objeto chamado `recipe3`.

```js
assert.isObject(recipe3);
```

Seu objeto `recipe3` deve ter uma propriedade `name`.

```js
assert.property(recipe3, "name");
```

A sua propriedade `name` deve ser definida como `Vegetable Stir Fry`.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

Seu objeto `recipe3` deve ter uma propriedade `ingredients`.

```js
assert.property(recipe3, "ingredients");
```

A sua propriedade `ingredients` deve ter um valor do tipo array.

```js
assert.isArray(recipe3.ingredients);
```

O valor do array da sua propriedade `ingredients` deve conter `broccoli`, `carrot` e `bell pepper`.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

Seu objeto `recipe3` deve ter uma propriedade `cookingTime`.

```js
assert.property(recipe3, "cookingTime");
```

O valor da sua propriedade `cookingTime` deve ser um número.

```js
assert.isNumber(recipe3.cookingTime);
```

A sua propriedade `cookingTime` deve ser definida como `15`.

```js
assert.equal(recipe3.cookingTime, 15);
```

Seu objeto `recipe3` deve ter uma propriedade `totalIngredients`.

```js
assert.property(recipe3, "totalIngredients");
```

A sua propriedade `totalIngredients` deve ser definida como `null`.

```js
assert.isNull(recipe3.totalIngredients);
```

Seu objeto `recipe3` deve ter uma propriedade `difficultyLevel`.

```js
assert.property(recipe3, "difficultyLevel");
```

A sua propriedade `difficultyLevel` deve ser definida como uma string vazia.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
