---
id: 68e6698b29e562d21b1d260c
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop costruirai una lista dei principali browser web. Il codice boilerplate HTML è già fornito.

Inizia aggiungendo un’intestazione alla tua pagina con il testo `List of Major Web Browsers` usando un elemento `h1` all’interno dell’elemento `body`.

# --hints--

Dovresti avere un elemento `h1` all’interno dell’elemento `body`.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

Il tuo elemento `h1` dovrebbe contenere il testo `List of Major Web Browsers`.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
