---
id: 68ecfa1003d88990bb96a33d
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Le liste di descrizione sono composte da un elenco di termini e dettagli. L'elemento `dt` rappresenta il termine della descrizione, mentre l'elemento `dd` rappresenta il dettaglio della descrizione.

Il primo browser che aggiungerai alla tua lista di descrizione sarà Google Chrome.

All'interno del tuo elemento `dl`, aggiungi un elemento `dt` contenente il testo `Google Chrome`.

Sotto il tuo elemento `dt`, aggiungi un elemento `dd` con il seguente testo:

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

Dovresti avere un elemento `dt` all'interno del tuo elemento `dl`.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

Il tuo elemento `dt` dovrebbe avere il testo `Google Chrome`.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

Dovresti avere un elemento `dd` sotto il tuo elemento `dt`.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

Il tuo elemento `dd` dovrebbe contenere il testo `This is a free web browser developed by Google and first released in 2008.`

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```
