---
id: 68ecfa1003d88990bb96a33e
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Ora è il momento di aggiungere il tuo secondo browser alla lista.

Sotto la tua voce `Google Chrome`, aggiungi un altro elemento `dt` contenente il testo `Firefox`.

Sotto il tuo elemento `dt`, aggiungi un elemento `dd` con il seguente testo:

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

Dovresti avere due elementi `dt` nel tuo codice.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

Il tuo secondo elemento `dt` dovrebbe avere il testo `Firefox`.

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

Dovresti avere due elementi `dd` nel tuo codice.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

Il tuo secondo elemento `dd` dovrebbe contenere il testo `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
