---
id: 68ecfa0f03d88990bb96a33c
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Come ricorderai da una lezione precedente, le liste di descrizione si usano per presentare termini e definizioni in un formato organizzato e facile da leggere.

Ecco un esempio:

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

Sotto l'elemento `h1`, crea un elemento `dl`. Questo conterrà la lista dei browser.

# --hints--

Dovresti avere un elemento `dl`.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

Il tuo elemento `dl` dovrebbe essere sotto l'elemento `h1`.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        --fcc-editable-region--
        
        --fcc-editable-region--
    </body> 
</html>
```
