---
id: 69038f5efad6f722def4a214
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Ora è il momento di aggiungere il tuo quarto browser alla lista.

Aggiungi un altro elemento `dt` contenente il testo `Brave`.

Sotto il tuo elemento `dt`, aggiungi un elemento `dd` con il seguente testo:

```md
This is a free web browser first released in 2016 that is based on the Chromium web browser.
```

# --hints--

Dovresti avere quattro elementi `dt` nel tuo codice.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 4);
```

Il tuo quarto elemento `dt` dovrebbe avere il testo `Brave`.

```js
const dt = document.querySelectorAll('dt')[3];
assert.strictEqual(dt?.textContent.trim(), "Brave");
```

Dovresti avere quattro elementi `dd` nel tuo codice.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 4);
```

Il tuo quarto elemento `dd` dovrebbe contenere il testo `This is a free web browser first released in 2016 that is based on the Chromium web browser.`

```js
const ddElement = document.querySelectorAll('dd')[3];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser first released in 2016 that is based on the Chromium web browser.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
