---
id: 66f17db06803d11a1bd19a20
title: Створіть тестову гру
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити масив з назвою `questions`.
1. Масив `questions` повинен містити щонайменше п’ять об’єктів, кожен з яких має ключі `category`, `question`, `choices` і `answer`.
1. Ключ `category` повинен мати значення рядка, що представляє категорію питання.
1. Ключ `question` повинен мати значення рядка, що представляє питання.
1. Ключ `choices` повинен мати значення масиву, який містить три рядки — альтернативні відповіді на питання.
1. Ключ `answer` повинен мати значення рядка, що представляє правильну відповідь на питання. Також значення `answer` має бути включене в масив `choices`.
1. Ви повинні мати функцію з назвою `getRandomQuestion`, яка приймає масив питань як параметр і повертає випадковий об’єкт питання з масиву.
1. Ви повинні мати функцію з назвою `getRandomComputerChoice`, яка приймає масив доступних варіантів як параметр і повертає випадкову відповідь на вибране питання.
1. Ви повинні мати функцію з назвою `getResults`, яка приймає об’єкт питання як перший параметр і вибір комп’ютера як другий параметр. Функція має повертати `The computer's choice is correct!`, якщо відповідь правильна. Інакше вона повертає `The computer's choice is wrong. The correct answer is: <correct-answer>`, де `<correct-answer>` — це значення правильної відповіді на вибране питання.

# --hints--

Ви повинні створити масив з назвою `questions`.

```js
assert.isArray(questions);
```

Масив `questions` повинен містити щонайменше п’ять об’єктів, кожен з яких має ключі `category`, `question`, `choices` і `answer`.

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

Ключ `category` повинен мати значення рядка, що представляє категорію питання.

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

Ключ `question` повинен мати значення рядка, що представляє питання.

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

Ключ `choices` повинен мати значення масиву, який містить три різні рядки.

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

Ключ `answer` повинен мати значення рядка.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

Значення `answer` має бути включене в масив `choices`.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

Ви повинні мати функцію з назвою `getRandomQuestion`, яка приймає масив питань як параметр і повертає випадковий об’єкт питання з масиву.

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

Ви повинні мати функцію з назвою `getRandomComputerChoice`, яка приймає масив доступних варіантів як параметр і повертає випадкову відповідь на вибране питання.

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

Ви повинні мати функцію з назвою `getResults`.

```js
assert.isFunction(getResults);
```

Ваша функція `getResults` повинна приймати об’єкт питання як перший параметр і вибір комп’ютера як другий параметр.

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

Якщо вибір комп’ютера співпадає з відповіддю, `getResults` має повернути `The computer's choice is correct!`.

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

Якщо вибір комп’ютера не співпадає з відповіддю, `getResults` має повернути `The computer's choice is wrong. The correct answer is: <correct-answer>`, де `<correct-answer>` — це значення правильної відповіді на вибране питання.

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

Ваша функція `getResults` має використовувати точне порівняння рівності, а не пошук підрядка.

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
