---
id: 68e448ec4ac4205968c0a805
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

문장 `"love"`에서 `"I love learning!"`라는 단어가 세 번 반복되면 좋겠습니다.

이전 수업에서 기억하듯이, `repeat()` 메서드를 사용해 문자열을 특정 횟수만큼 반복할 수 있습니다. 예를 들어 다음과 같습니다:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

문자열은 불변이므로 이 메서드는 원래 문자열을 수정하지 않습니다. 반복된 내용을 포함한 새 문자열을 반환합니다.

`repeatedLove`라는 변수를 만들고 `"love ".repeat(3)`를 할당하세요. 그런 다음 결과를 확인할 수 있도록 `repeatedLove` 변수를 콘솔에 출력하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`repeatedLove`라는 변수가 있어야 합니다.

```js
assert.exists(repeatedLove);
```

`repeatedLove` 변수는 문자열이어야 합니다.

```js
assert.isString(repeatedLove);
```

`"love ".repeat(3)`를 `repeatedLove` 변수에 할당해야 합니다.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

`repeatedLove` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
