---
id: 68e42de6396f5c49d90d1951
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이전 수업에서 문자열에서 부분 문자열을 찾아 다른 값으로 바꾸는 데 사용하는 `replace()` 메서드에 대해 배웠습니다.

예시는 다음과 같습니다:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

문자열은 불변이라는 점을 기억하세요. 즉, 이 예시에서 원래 `text`은 수정되지 않습니다. 대신 새 문자열이 만들어집니다.

`replacedString`라는 새 변수를 만들고 `originalString.replace("cats", "dogs")`을 할당하세요.

# --hints--

`replacedString`라는 변수가 있어야 합니다.

```js
assert.exists(replacedString);
```

`replacedString` 변수는 문자열이어야 합니다.

```js
assert.isString(replacedString);
```

`originalString.replace("cats", "dogs")`을 `replacedString` 변수에 할당해야 합니다.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
