---
id: 68e4417b7317835208145323
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

다음으로, `dogsOnlySentence`라는 새 변수를 만들고 `exampleSentence.replace("cats", "dogs")`을 할당하세요.

그 변수 아래에 `"Replacing all occurrences of cats with dogs:"`를 콘솔에 출력하세요. 그 아래에 `console.log()`를 두고 `console.log()` 변수에 대해 또 다른 `dogsOnlySentence`를 추가하세요.

이제 콘솔에 `"I love dogs and cats are so much fun!"`라는 텍스트가 출력되는 것을 볼 수 있습니다. 하지만 원하는 결과는 아닙니다. 다음 단계에서 텍스트가 `"I love dogs and dogs are so much fun!"`로 나오도록 고치는 방법을 배우게 됩니다.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`dogsOnlySentence` 변수가 있어야 합니다.

```js
assert.exists(dogsOnlySentence);
```

`dogsOnlySentence` 변수는 문자열이어야 합니다.

```js
assert.isString(dogsOnlySentence);
```

`exampleSentence.replace("cats", "dogs")`을 `dogsOnlySentence` 변수에 할당해야 합니다.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

`"Replacing all occurrences of cats with dogs:"`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

`dogsOnlySentence` 변수를 콘솔에 출력해야 합니다.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
