---
id: 68e42b62d7093345c6975f22
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

`originalString` 변수 아래에 `console.log()`을 `"Original string:"` 문자열로 추가하세요.

그다음 `console.log()` 아래에 `console.log()` 변수를 가진 또 다른 `originalString`를 추가하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

콘솔에 `"Original string:"` 문자열을 출력해야 합니다.

```js
assert.equal(getLogs()[0], "Original string:");
```

콘솔에 `originalString` 변수를 출력해야 합니다.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
