---
id: 68e43e884ef09d4d2ab05631
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이제 `replace()` 메서드를 사용한 결과를 확인할 차례입니다.

먼저 값이 `console.log()`인 `"After using the replace() method:"`을 추가하세요.

그다음 그 아래 `console.log()`에 `console.log()` 변수를 가진 또 다른 `replacedString`를 추가하세요.

콘솔을 확인하면 새 문장이 원래 `"I love dogs."` 대신 `"I love cats."`이라고 표시되는 것을 볼 수 있습니다.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`"After using the replace() method:"`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

`replacedString` 변수를 콘솔에 출력해야 합니다.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
