---
id: 68e44519c769bf545bb3f22a
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

원하는 출력인 `"I love dogs and dogs are so much fun!"`을 얻으려면 `replaceAll()` 메서드 대신 `replace()` 메서드를 사용해야 합니다. 이 메서드는 부분 문자열과 일치하는 모든 항목에 대해 새 문자열을 반환합니다.

`exampleSentence.replace("cats", "dogs")`을 `replaceAll()` 메서드를 사용하도록 업데이트하세요. 이제 콘솔에서 올바른 텍스트가 보여야 합니다.

# --hints--

`exampleSentence.replace("cats", "dogs")` 코드 줄을 `replaceAll()` 메서드를 사용하도록 업데이트하세요.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
