---
id: 68e4465fba0bd256868b7deb
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

이제 새로운 문자열 메서드를 사용해 볼 시간입니다.

먼저 `learningSentence`라는 변수를 만들고 문자열 `"I love learning!"`을 할당하세요.

그다음 그 변수 아래에 `"Original learning sentence:"`를 콘솔에 출력하세요. 그 아래 `console.log()` 다음에 `console.log()` 변수를 위한 또 다른 `learningSentence`를 추가하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`learningSentence` 변수가 있어야 합니다.

```js
assert.exists(learningSentence);
```

`learningSentence` 변수는 문자열이어야 합니다.

```js
assert.isString(learningSentence);
```

`"I love learning!"` 변수에 문자열 `learningSentence`을 할당해야 합니다.

```js
assert.equal(learningSentence, "I love learning!");
```

`"Original learning sentence:"`을 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

`learningSentence` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
