---
id: 68e44abea678a65b877864dc
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

문장에 반복된 단어가 보이면 좋겠습니다. `newSentence`라는 새 변수를 만들고, `"I "` 단어와 `repeatedLove` 변수를 이어 붙인 결과에 문자열 `" learning."`를 이어 붙인 값을 할당하세요. 그런 다음 `newSentence` 변수를 콘솔에 출력하세요. 문자열 연결에는 `+` 연산자나 템플릿 리터럴 중 하나를 선택해 사용해도 됩니다.

이제 `"I love love love  learning."` 텍스트가 보여야 합니다.

*참고*: 마지막 `"love"` 뒤에 공백이 하나 더 있는데, 이것이 올바른 상태입니다. 마지막 단계에서 그 공백을 잘라낼 것입니다.

# --hints--

`newSentence`라는 변수가 있어야 합니다.

```js
assert.exists(newSentence);
```

`newSentence` 변수는 문자열이어야 합니다.

```js
assert.isString(newSentence);
```

`"I "` 단어와 `repeatedLove` 변수를 이어 붙인 다음 문자열 `" learning."`를 이어 붙여야 합니다. 이 결과를 `newSentence` 변수에 할당하세요. 공백에 주의하세요.

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

`"I love love love  learning."`라는 문자열 리터럴을 `newSentence` 변수에 할당하면 안 됩니다. 올바른 방법을 위해 지침을 다시 읽어보세요.

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
