---
id: 68e43f94365a964f26c88f78
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

이제 새로운 문자열 메서드를 사용해 볼 시간입니다.

먼저 `exampleSentence`라는 새 변수를 만들고 문자열 `"I love cats and cats are so much fun!"`을 할당하세요.

그다음 그 변수 아래에 `"Original sentence:"`를 콘솔에 출력하세요. 그리고 그 아래 `console.log()` 다음에 `console.log()` 변수를 사용해 또 다른 `exampleSentence`를 추가하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`exampleSentence` 변수가 있어야 합니다.

```js
assert.exists(exampleSentence);
```

`exampleSentence` 변수는 문자열이어야 합니다.

```js
assert.isString(exampleSentence);
```

`"I love cats and cats are so much fun!"` 변수에 문자열 `exampleSentence`을 할당해야 합니다.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

`"Original sentence:"`을 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

`exampleSentence` 변수를 콘솔에 출력해야 합니다.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
