---
id: 68e44c305f1f935d949d76bb
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이전 단계에서 언급했듯이, 문장 `"love"`의 마지막 `"I love love love  learning."` 뒤에 여분의 공백이 있습니다. 그 여분의 공백을 제거하려면 `trimEnd()` 메서드를 사용할 수 있습니다.

`"love ".repeat(3)` 메서드 끝에 `trimEnd()` 메서드를 연결하세요. 메서드는 이렇게 연결할 수 있습니다:

```js
.firstMethod().secondMethod()
```

이제 콘솔을 확인하면 여분의 공백이 제거된 것을 볼 수 있습니다.

이 마지막 변경으로 워크숍이 완료되었습니다!

# --hints--

`trimEnd()` 메서드에 `"love ".repeat(3)` 메서드를 연결해야 합니다.

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
