---
id: 68e41084bb275539f72de9fe
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 `replace()`, `replaceAll()` 그리고 `repeat()` 메서드를 사용하여 문자열 변환기 앱을 만드는 연습을 합니다.

먼저, `originalString`라는 변수를 만들고 `"I love cats."` 값을 할당하세요.

# --hints--

`originalString`라는 변수가 있어야 합니다.

```js
assert.exists(originalString);
```

`originalString` 변수는 문자열이어야 합니다.

```js
assert.isString(originalString);
```

`"I love cats."` 변수에 문자열 `originalString`을 할당해야 합니다.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
