---
id: 691f7bbf76229bb97c60827d
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utazoea kutumia mizunguko ya `for...of` kwa kujenga kitendakazi kinachohesabu mara ngapi mfuatano wa herufi huonekana katika safu ya data za mfuatano wa herufi.

Kabla hujaunda kitendakazi hicho, utaanzisha kitendakazi rahisi zaidi kinachochapisha kila herufi ya mfuatano wa herufi kwenye konsoli.

Anza kwa kufafanua kitendakazi tupu kinachoitwa `printCharacters` chenye kigezo `str`.

# --hints--

Unapaswa kuunda kitendakazi cha `printCharacters`.

```js
assert.isFunction(printCharacters);
```

Kitendakazi chako cha `printCharacters` kinapaswa kuwa na kigezo kimoja kinachoitwa `str`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
