---
id: 6921e10ae2681205b516fbe6
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Sasa ni wakati wa kuongeza mantiki kwenye mzunguko wako ambayo huongeza `count` ipasavyo.

Ndani ya mzunguko, tumia kauli ya masharti kuongeza `count` kwa `1` ikiwa kigezo `word` kinafanana na kigezo `match`. Vinginevyo, acha thamani ya `count` isibadilike.

# --hints--

Kuitisha `getMatchedWordCount([], "bar")`, `getMatchedWordCount(["foo"], "bar")`, `getMatchedWordCount(["foo", "bar", "foo"], "bar")`, na `getMatchedWordCount(["foo", "bar", "foo"], "foo")` kwa mpangilio huo kunapaswa kurudisha thamani ya aina `"number"`; badala yake, ilirudisha thamani za aina `--fcc-actual--`.

```js
const obj = { method: (arg1, arg2) => typeof getMatchedWordCount(arg1, arg2) };
const spy = __helpers.spyOn(obj, 'method');
obj.method([], "bar");
obj.method(["foo"], "bar");
obj.method(["foo", "bar", "foo"], "bar");
obj.method(["foo", "bar", "foo"], "foo");
assert.deepEqual(spy.returns, ["number", "number", "number", "number"]);
spy.restore();
```

`getMatchedWordCount([], "bar")` inapaswa kurudisha `0`, lakini inarudisha `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount([], "bar"));
assert.strictEqual(resultString, "0");
```

`getMatchedWordCount(["foo"], "bar")` inapaswa kurudisha `0`, lakini inarudisha `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo"], "bar"));
assert.equal(resultString, "0");
```

`getMatchedWordCount(["foo", "bar", "foo"], "bar")` inapaswa kurudisha `1`, lakini inarudisha `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "bar"));
assert.equal(resultString, "1");
```

`getMatchedWordCount(["foo", "bar", "foo"], "foo")` inapaswa kurudisha `2`, lakini inarudisha `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "foo"));
assert.equal(resultString, "2");
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
