---
id: 6921d90b5ce50185505d2004
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Kwa kuwa kitendakazi cha `getMatchedWordCount` kinahitaji kurudisha hesabu ya nambari, utaanza kwa kuanzisha kigezo cha kuhesabu.

Tengeneza kigezo kinachoitwa `count` kwa kutumia neno kuu la `let`, kianzishe na thamani ya `0`, na ongeza kauli ya `return` inayorudisha kigezo cha `count`.

# --hints--

Unapaswa kutangaza kigezo cha `count` ndani ya kitendakazi cha `getMatchedWordCount`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getMatchedWordCount?.variables.count);
```

Ili kuruhusu kigezo cha `count` kuongezwa baadaye, kitangazwe kwa neno kuu la `let`.

```js
const explorer = await __helpers.Explorer(code);
const countVar = explorer.allFunctions.getMatchedWordCount?.variables.count;
assert.isTrue(countVar?.toString().startsWith('let'));
```

Kitendakazi chako kinapaswa kuwa na mgawo wa thamani `0` kwa kigezo cha `count`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.variables.count?.value.matches('0'));
```

Kitendakazi chako kinapaswa kuwa na kauli ya `return` inayorudisha kigezo cha `count`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.hasReturn('count'));
```

Kitendakazi chako cha `getMatchedWordCount` kinapaswa kurudisha thamani ya `0`.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
