---
id: 6921dae10700e300f26994e3
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Ili kuona matokeo ya kitendakazi chako cha `getMatchedWordCount` katika hali yake ya sasa, ongeza tamko la `console.log` chini yake. Ndani ya `console.log()`, fanya wito wa kitendakazi hicho kwa hoja `["I", "really", "really", "really", "like", "to", "code"]` kwa kigezo cha `sentence` na `"really"` kwa kigezo cha `match`.

# --hints--

Unapaswa kufanya wito wa kitendakazi `getMatchedWordCount` ndani ya tamko la `console.log`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

Unapaswa kufanya wito wa kitendakazi `getMatchedWordCount` kwa hoja `["I", "really", "really", "really", "like", "to", "code"]` kwa kigezo cha `sentence` na `"really"` kwa kigezo cha `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
