---
id: 6921de8d2a111d0341cff702
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Nzuri, kitendakazi chako kinarudisha `0`, lakini bado hakihesabu chochote.

Ili kurekebisha hili, ndani ya `getMatchedWordCount`, tengeneza mzunguko wa `for...of` wenye sehemu tupu ya msimbo unaorudia juu ya kila neno katika `sentence`.

# --hints--

Kitendakazi chako cha `getMatchedWordCount` kinapaswa kuwa na mzunguko wa `for...of` ndani ya sehemu yake ya msimbo.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

Mzunguko wako wa `for...of` unapaswa kurudia juu ya hoja ya `sentence` ya kitendakazi chako.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;

--fcc-editable-region--
  
--fcc-editable-region--

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
