---
id: 6921d688b8d58581d34459b9
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Katika hatua chache zijazo, utaunda kitendakazi kinachohesabu mara ngapi mfuatano wa herufi huonekana katika safu ya data za mfuatano wa herufi.

Kuanza, fafanua kitendakazi tupu chenye jina `getMatchedWordCount` na vigezo `sentence` na `match` kwa mpangilio huo.

# --hints--

Unapaswa kuunda kitendakazi cha `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

Kitendakazi chako cha `getMatchedWordCount` kinapaswa kuwa na vigezo viwili: `sentence` na `match`, kwa mpangilio huo.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.getMatchedWordCount;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), 'sentence');
assert.equal(parameters[1].toString(), 'match');
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
