---
id: 6978bd465213832659f3338b
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Sasa, ndani ya mzunguko, andika usemi wa template literal ufuatao kwenye konsoli: `Checking "${word}" against "${match}" | Running count: ${count}`

# --hints--

Unapaswa kuwa na sentensi moja tu ya `console.log` ndani ya mzunguko katika kitendakazi chako cha `getMatchedWordCount`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
console.log(normalizedFunctionCode);

assert.match(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?console\.log\(\S*?\)\S*?\}/);
assert.notMatch(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?(console\.log\(\S*?\))\S*?(console\.log\(\S*?\))\S*?\}/);
```

Sentensi yako ya `console.log` ndani ya mzunguko katika kitendakazi chako cha `getMatchedWordCount` inapaswa kutumia usemi wa template literal wenye viwakilishi vifuatavyo: `${word}`, `${match}`, na `${count}`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
const consoleLogArguments = normalizedFunctionCode.match(/for\(constwordofsentence\)\{\S*?console\.log\((\S*?)\)\S*?\}/)[1]
const templateLiteral = consoleLogArguments.match(/(^\`\S*\`$)/)[1];

assert.include(templateLiteral, "${word}");
assert.include(templateLiteral, "${match}");
assert.include(templateLiteral, "${count}");
```

Wakati kitendakazi cha `getMatchedWordCount` kinapoitwa na `(['foo'], 'bar')`, sentensi ya `console.log()` ndani ya kitendakazi inapaswa kuitwa mara moja na mfuatano wa herufi `Checking "foo" against "bar" | Running count: 0`. Badala yake, ilitumiwa na kila mfuatano wa herufi katika orodha ifuatayo iliyo na koma: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
getMatchedWordCount(['foo'], 'bar');
assert.deepEqual(consoleLogSpy.calls, [['Checking "foo" against "bar" | Running count: 0']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
