---
id: 56533eb9ac21ba0edf2244b3
title: Побудуйте конвертер із Цельсія в Фаренгейт
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

У цій лабораторній роботі ви напишете функцію, яка конвертує температуру з Цельсія у Фаренгейт. Формула для конвертації з Цельсія у Фаренгейт:

```js
fahrenheit = celsius * (9/5) + 32
```

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити функцію з назвою `convertCtoF`.
2. `convertCtoF` має приймати один числовий аргумент — температуру в Цельсіях.
3. `convertCtoF` має повертати число.

# --hints--

Ви повинні створити функцію з назвою `convertCtoF`.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` має приймати один параметр.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` має повертати число.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` має повертати значення `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` має повертати значення `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` має повертати значення `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` має повертати значення `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` має повертати значення `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
