---
id: 66c051d13a6a20255a963695
title: スタイライズされたTo-Doリストを作成する
challengeType: 25
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

このラボでは、リンクに対してホバー、フォーカス、クリック、訪問時に適用できるさまざまなスタイルを練習します。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. クラスが`todo-list`の番号なしリストを1つ持っている必要があります。
2. その番号なしリストの中に4つのリストアイテムがある必要があります。
3. 各リストアイテムの中には以下が含まれている必要があります：

   - タイプが`input`で、`checkbox`が一意の値に設定された`id`要素。
   - 対応する`label`要素の`for`を参照する`input`属性が設定された`id`要素。
   - クラスが`sub-item`の番号なしリスト。
   - その中にアンカー要素を含むリストアイテム。アンカーはクラス`sub-item-link`を持ち、有効な`href`値と、新しいタブでリンクを開く`target`値を持つ必要があります。

4. `a`要素にはテキスト装飾がないこと。
5. 未訪問リンクのテキストカラーは任意の色に設定してください。
6. リンクが訪問済みの場合、色が別の任意の色に変わること。
7. リンクにホバーしたとき、色が別の任意の色に変わること。
8. リンクにフォーカスが当たったとき、リンクの周りに色付きのアウトラインが表示されること。
9. リンクをクリックしたとき、色が別の任意の色に変わること。

**注意:** スタイルシートをHTMLにリンクし、CSSを適用してください。

# --hints--

クラスが`todo-list`の番号なしリストを1つ持っている必要があります。

```js
assert.exists(document.querySelector('ul.todo-list'));
assert.lengthOf(document.querySelectorAll('ul.todo-list'), 1);
```

番号なしリストの中に4つのリストアイテムがある必要があります。

```js
const li = document.querySelectorAll('ul.todo-list > li');
assert.lengthOf(li, 4);
```

クラスが`li`の`ul`内の`todo-list`は、`input`の`type`を持つ`checkbox`要素を含む必要があります。

```js
const checkboxes = document.querySelectorAll(
  'ul.todo-list li input[type="checkbox"]'
);
assert.lengthOf(checkboxes, 4);
```

クラスが`li`の`ul`内の`todo-list`は、`label`要素を含む必要があります。

```js
const labels = document.querySelectorAll(
  'ul.todo-list li label'
);
assert.lengthOf(labels, 4);
```

すべての`input`要素は`id`を持つ必要があります。

```js
const inputs = document.querySelectorAll('input');
assert.isAbove(inputs.length, 0);
for (let input of inputs) {
  assert.isAbove(input?.id.length, 0);
}
```

すべての`label`要素は`for`属性を持つ必要があります。

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.htmlFor.length, 0);
}
```

すべての`label`要素はテキストを持つ必要があります。

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.innerText.length, 0);
}
```

`id`と`for`要素のペアの`input`および`label`属性は対応する値を持つ必要があります。

```js
const labels = document.querySelectorAll("label");
assert.isNotEmpty(labels);

document.querySelectorAll("input").forEach((input, index) => {
  assert.equal(input.id, labels[index].htmlFor);
});
```

`label`要素の後に、クラスが`sub-item`の番号なしリストがある必要があります。

```js
const labelSbItems = document.querySelectorAll('ul.todo-list li label + ul.sub-item');
assert.lengthOf(labelSbItems, 4);
```

クラスが`li`の`ul`内の`sub-item`は、クラス`sub-item-link`のアンカー要素を持つ必要があります。

```js
const allItems = document.querySelectorAll(
  'ul.todo-list li label + ul.sub-item li a.sub-item-link'
);
assert.lengthOf(allItems, 4);
```

すべてのアンカー要素は有効な`href`属性を持つ必要があります。

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.href.length, 0);
}
```

すべてのアンカー要素はテキストを持つ必要があります。

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.innerText.length, 0);
}
```

各`a`要素は`target`の値を持つ`_blank`属性を持つ必要があります。

```js
const anchors = document.querySelectorAll('a');
assert.isAbove(anchors.length, 0);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

`a`要素にはテキスト装飾がないこと。

```js
assert.isTrue(
  new __helpers.CSSHelp(document).getStyle('a.sub-item-link')
    ?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('a')?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('.sub-item-link')
      ?.textDecoration === 'none'
);
```

未訪問リンクは任意の色にテキストカラーが設定されている必要があります。

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:link')?.color
);
```

リンクはホバー時に色が変わる必要があります。

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:hover')?.color
);
```

リンクはクリック時に色が変わる必要があります。

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:active')?.color
);
```

リンクはフォーカス時にアウトラインが表示される必要があります。

```js
const css = new __helpers.CSSHelp(document);
assert.exists(
  css.getStyle('.sub-item-link:focus')?.outline ||
    css.getStyle('a:focus')?.outline ||
    css.getStyle('a.sub-item-link:focus')?.outline ||
    css.getStyle('.sub-item-link:focus')?.border ||
    css.getStyle('a:focus')?.border ||
    css.getStyle('a.sub-item-link:focus')?.border
);
```

リンクは訪問済み時に色が変わる必要があります。

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:visited')?.color
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Styled To-Do List</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Styled To-Do List</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div class="container">
      <h1>My To-Do List</h1>
      <ul class="todo-list">
        <li>
          <input type="checkbox" id="task1" />
          <label for="task1">Explore gaming keyboards</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link">Store link</a>
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task2" />
          <label for="task2">Finish the report</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >Request Access</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task3" />
          <label for="task3">View Phone's Warranty</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Receipts</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task4" />
          <label for="task4">Check Processor Specs</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Model Number</a
              >
            </li>
          </ul>
        </li>
      </ul>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: #f0f0f0;
  margin: 0;
  padding: 20px;
}

.container {
  max-width: 400px;
  margin: 0 auto;
  background-color: #fff;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
  color: #333;
}

.todo-list {
  list-style-type: none;
  padding: 0;
}

.todo-list li {
  margin-bottom: 15px;
  line-height: 1.6;
  padding: 10px;
  background-color: #e5f1e5;
  border-radius: 10px;
  position: relative;
}

a.sub-item-link {
  margin-left: 10px;
  text-decoration: none;
}

.sub-item-link:link {
  color: blue;
}

.sub-item-link:visited {
  color: purple;
}

.sub-item-link:hover {
  color: red;
}

.sub-item-link:focus {
  outline: 2px solid blue;
  outline-offset: 2px;
}

.sub-item-link:active {
  color: green;
}

.sub-item {
  list-style-type: square;
}
```
