---
id: acda2fb1324d9b0fa741e6b5
title: Tengeneza chombo cha kuthibitisha mwisho wa mfuatano wa herufi
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

Katika maabara hii, utatekeleza kitendakazi kinachokagua kama mfuatano wa herufi unamalizika na mfuatano wa herufi lengwa uliotolewa.

**Lengo:** Timiza hadithi za mtumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za mtumizi:**

1. Unapaswa kuunda kitendakazi chenye jina `confirmEnding` kinachopokea vigezo viwili: mfuatano wa herufi wa kukagua na mfuatano wa herufi wa kulinganisha.
2. Kitendakazi kinapaswa kurudisha `true` ikiwa mfuatano wa herufi wa kwanza unamalizika na mfuatano wa herufi wa pili, na `false` vinginevyo.
3. Hupaswi kutumia njia ya `.endsWith()`; badala yake, tumia mojawapo ya njia za JavaScript za substring kufanikisha hili.

# --hints--

Unapaswa kuunda kitendakazi chenye jina `confirmEnding`.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` inapaswa kupokea vigezo 2.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` inapaswa kurudisha `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` inapaswa kurudisha `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` inapaswa kurudisha `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` inapaswa kurudisha `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` inapaswa kurudisha `false`.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Msimbo wako haupaswi kutumia njia iliyojengwa ndani ya `.endsWith()` kutatua maabara hii.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
