---
id: 587d825c367417b2b2512c8f
title: تنفيذ خوارزمية الدمج
challengeType: 1
forumTopicId: 301614
dashedName: implement-merge-sort
---

# --description--

خوارزمية فرز شائعة ومتوسطة المستوى هي خوارزمية الدمج. مثل خوارزمية الفرز السريع، تستخدم خوارزمية الدمج أيضًا منهجية التقسيم والتغلب، مع استدعاء ذاتي لفرز مصفوفة. تستفيد من حقيقة أنه من السهل نسبيًا فرز مصفوفتين طالما أن كل واحدة منهما مرتبة في الأصل. لكننا سنبدأ بمصفوفة واحدة فقط كمُدخَل، فكيف نحصل على مصفوفتين مرتبتين من ذلك؟ يمكننا تقسيم المُدخَل الأصلي إلى نصفين بشكل متكرر حتى نصل إلى الحالة الأساسية وهي مصفوفة تحتوي على عنصر واحد. المصفوفة ذات العنصر الواحد مرتبة بطبيعتها، ثم نبدأ بالدمج. هذا الدمج سيُفكك الاستدعاءات الذاتية التي قسمت المصفوفة الأصلية، مما ينتج في النهاية مصفوفة مرتبة نهائية لجميع العناصر. إذن خطوات خوارزمية الدمج هي:

**1)** قسم المصفوفة المُدخَلة إلى نصفين بشكل متكرر حتى تنتج مصفوفة فرعية تحتوي على عنصر واحد فقط.

**2)** دمج كل مصفوفة فرعية مرتبة معًا لإنتاج المصفوفة النهائية المرتبة.

خوارزمية الدمج طريقة فرز فعالة، وتعقيدها الزمني هو *O(nlog(n))*. هذه الخوارزمية شائعة لأنها ذات أداء جيد وسهلة التنفيذ نسبيًا.

كملاحظة جانبية، ستكون هذه آخر خوارزمية فرز نغطيها هنا. لكن لاحقًا في قسم هياكل بيانات الأشجار سنشرح خوارزمية فرز الحزمة، وهي طريقة فرز فعالة أخرى تتطلب حزمة ثنائية في تنفيذها.

**التعليمات:** اكتب دالة `mergeSort` تأخذ مصفوفة من الأعداد الصحيحة كمُدخَل وتُرجع مصفوفة مرتبة من هذه الأعداد من الأصغر إلى الأكبر. طريقة جيدة للتنفيذ هي كتابة دالة واحدة، مثل `merge`، مسؤولة عن دمج مصفوفتين مرتبتين، ودالة أخرى، مثل `mergeSort`، مسؤولة عن الاستدعاء الذاتي الذي ينتج مصفوفات ذات عنصر واحد لتغذّي دالة الدمج. حظًا موفقًا!

# --hints--

`mergeSort` يجب أن تكون دالة.

```js
assert.isFunction(mergeSort);
```

`mergeSort` يجب أن تُرجع مصفوفة مرتبة (من الأصغر إلى الأكبر).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    mergeSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`mergeSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` يجب أن تُرجع مصفوفة لم تتغير إلا في الترتيب.

```js
assert.sameMembers(
  mergeSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`mergeSort` يجب ألا تستخدم الطريقة المدمجة `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    mergeSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function mergeSort(array) {
  if (array.length === 1) {
    return array;
  } else {
    const splitIndex = Math.floor(array.length / 2);
    return merge(
      mergeSort(array.slice(0, splitIndex)),
      mergeSort(array.slice(splitIndex))
    );
  }

  // Merge two sorted arrays
  function merge(array1, array2) {
    let merged = [];
    while (array1.length && array2.length) {
      if (array1[0] < array2[0]) {
        merged.push(array1.shift());
      } else if (array1[0] > array2[0]) {
        merged.push(array2.shift());
      } else {
        merged.push(array1.shift(), array2.shift());
      }
    }

    // After looping ends, one array is empty, and other array contains only
    // values greater than all values in `merged`
    return [...merged, ...array1, ...array2];
  }
}

mergeSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]);
```
