---
id: 587d8259367417b2b2512c85
title: تنفيذ فرز الاختيار
challengeType: 1
forumTopicId: 301616
dashedName: implement-selection-sort
---

# --description--

هنا سننفذ فرز الاختيار. يعمل فرز الاختيار عن طريق اختيار القيمة الأدنى في قائمة وتبديلها مع القيمة الأولى في القائمة. ثم يبدأ من الموضع الثاني، يختار أصغر قيمة في القائمة المتبقية، ويبدلها مع العنصر الثاني. يستمر في التكرار عبر القائمة وتبديل العناصر حتى يصل إلى نهاية القائمة. الآن القائمة مرتبة. فرز الاختيار له تعقيد زمني تربيعي في جميع الحالات.

**التعليمات**: اكتب دالة `selectionSort` تأخذ مصفوفة من الأعداد الصحيحة كمُدخَل وتُرجع مصفوفة من هذه الأعداد مرتبة من الأصغر إلى الأكبر.

# --hints--

يجب أن تكون `selectionSort` دالة.

```js
assert.isFunction(selectionSort);
```

يجب أن تُرجع `selectionSort` مصفوفة مرتبة (من الأصغر إلى الأكبر).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}

assert.isTrue(
  isSorted(
    selectionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

يجب أن تُرجع `selectionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` مصفوفة لم تتغير إلا في الترتيب.

```js
assert.sameMembers(
  selectionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

يجب ألا تستخدم `selectionSort` الطريقة المدمجة `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function selectionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
