---
id: 587d825a367417b2b2512c89
title: تنفيذ خوارزمية الفرز السريع
challengeType: 1
forumTopicId: 301615
dashedName: implement-quick-sort
---

# --description--

سننتقل هنا إلى خوارزمية فرز متوسطة المستوى: الفرز السريع. الفرز السريع هو طريقة فعالة تعتمد على التقسيم والتغلب بشكل تكراري لفرز مصفوفة. في هذه الطريقة، يُختار قيمة محورية في المصفوفة الأصلية. ثم تُقسّم المصفوفة إلى مصفوفتين فرعيتين تحتويان على القيم الأقل والأكبر من القيمة المحورية. بعد ذلك، ندمج نتائج استدعاء خوارزمية الفرز السريع بشكل تكراري على كلتا المصفوفتين الفرعيتين. تستمر هذه العملية حتى نصل إلى الحالة الأساسية وهي مصفوفة فارغة أو تحتوي على عنصر واحد فقط، والتي نُعيدها. عند فك استدعاءات التكرار، نحصل على المصفوفة المرتبة.

الفرز السريع هو طريقة فرز فعالة جدًا، حيث يوفر أداء *O(nlog(n))* في المتوسط. كما أنه سهل التنفيذ نسبيًا. هذه الصفات تجعله طريقة فرز شائعة ومفيدة.

**التعليمات:** اكتب دالة `quickSort` تأخذ مصفوفة من الأعداد الصحيحة كمُدخَل وتُعيد مصفوفة من هذه الأعداد مرتبة من الأصغر إلى الأكبر. رغم أهمية اختيار القيمة المحورية، أي قيمة محورية ستفي بالغرض هنا. للتبسيط، يمكن استخدام العنصر الأول أو الأخير.

# --hints--

يجب أن تكون `quickSort` دالة.

```js
assert(typeof quickSort == 'function');
```

يجب أن تُعيد `quickSort` مصفوفة مرتبة (من الأصغر إلى الأكبر).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    quickSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

يجب أن تُعيد `quickSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` مصفوفة لم تتغير إلا في ترتيب عناصرها.

```js
assert.sameMembers(
  quickSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

يجب ألا تستخدم `quickSort` الطريقة المدمجة `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quickSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function quickSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function quickSort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quickSort(lesser), ...equal, ...quickSort(greater)];
  }
}
```
