---
id: 587d8259367417b2b2512c86
title: تنفيذ فرز الإدراج
challengeType: 1
forumTopicId: 301613
dashedName: implement-insertion-sort
---

# --description--

طريقة الفرز التالية التي سننظر إليها هي فرز الإدراج. تعمل هذه الطريقة عن طريق بناء مصفوفة مرتبة في بداية القائمة. تبدأ المصفوفة المرتبة بالعنصر الأول. ثم تفحص العنصر التالي وتبدله إلى الخلف داخل المصفوفة المرتبة حتى يكون في الموضع المرتب الصحيح. تستمر في التكرار عبر القائمة وتبديل العناصر الجديدة إلى الخلف داخل الجزء المرتب حتى تصل إلى النهاية. هذه الخوارزمية لها تعقيد زمني تربيعي في الحالات المتوسطة والأسوأ.

**التعليمات:** اكتب دالة `insertionSort` تأخذ مصفوفة من الأعداد الصحيحة كمُدخَل وتُرجع مصفوفة من هذه الأعداد مرتبة من الأصغر إلى الأكبر.

# --hints--

يجب أن تكون `insertionSort` دالة.

```js
assert.isFunction(insertionSort);
```

يجب أن تُرجع `insertionSort` مصفوفة مرتبة (من الأصغر إلى الأكبر).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    insertionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

يجب أن تُرجع `insertionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` مصفوفة لم تتغير إلا في الترتيب.

```js
assert.sameMembers(
  insertionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

يجب أن تُرجع `insertionSort([5, 4, 33, 2, 8])` `[2, 4, 5, 8, 33]`.

```js
assert.deepEqual(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

يجب ألا تستخدم `insertionSort` الطريقة المدمجة `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function insertionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
