---
id: 61abc7ebf3029b56226de5b6
title: تنفيذ البحث الثنائي
challengeType: 1
forumTopicId: 487618
dashedName: implement-binary-search
---

# --description--

البحث الثنائي هو خوارزمية ذات كفاءة **O(log(n))** للبحث في مصفوفة مرتبة لإيجاد عنصر. يعمل باستخدام الخطوات التالية:

1. ابحث عن `value` الأوسط في مصفوفة مرتبة. إذا كان `value == target` فأرجع `true` (تم العثور على القيمة واكتمل البحث).
1. إذا كان `value < target` الأوسط، ابحث في النصف الأيمن من المصفوفة في المقارنة التالية.
1. إذا كان `value > target` الأوسط، ابحث في النصف الأيسر من المصفوفة في المقارنة التالية.
1. إذا لم توجد القيمة بعد البحث في كامل المصفوفة، فأرجع `false` (تم البحث في المصفوفة والقيمة غير موجودة فيها).

كما ترى، أنت تقسم المصفوفة إلى نصفين بشكل متتالي، وهذا يمنحك كفاءة log(n). في هذا التحدٍّ، نريدك أن تُظهر عملك - كيف وصلت إلى القيمة المستهدفة... المسار الذي اتبعته!

# --instructions--

اكتب دالة `binarySearch` تنفذ خوارزمية البحث الثنائي على مصفوفة، وتُرجع المسار الذي اتبعته (كل مقارنة بالقيمة الوسطى) لإيجاد الهدف في المصفوفة.

تأخذ الدالة مصفوفة مرتبة من أعداد صحيحة وقيمة هدف كمُدخَل. تُرجع مصفوفة تحتوي (بالترتيب) القيمة الوسطى التي وجدتها عند كل تقسيم للمصفوفة الأصلية حتى تجد القيمة الهدف. يجب أن تكون القيمة الهدف العنصر الأخير في المصفوفة المُعادة. إذا لم تُعثر على القيمة، أعد السلسلة النصية `Value Not Found`.

على سبيل المثال، `binarySearch([1,2,3,4,5,6,7], 5)` ستُرجع `[4,6,5]`.

في هذا التحدٍّ، عند التقسيم إلى نصفين، يجب عليك استخدام `Math.floor()` عند القسمة: `Math.floor(x/2)`. هذا سيعطي مسارًا ثابتًا وقابلًا للاختبار.

**ملاحظة:** ستُستخدم المصفوفة التالية في الاختبارات:

```js
const testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
```

# --hints--

يجب أن تكون `binarySearch` دالة.

```js
assert.isFunction(binarySearch);
```

`binarySearch(testArray, 0)` يجب أن تُرجع `[13, 5, 2, 0]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 0), [13, 5, 2, 0]);
```

`binarySearch(testArray, 1)` يجب أن تُرجع `[13, 5, 2, 0, 1]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 1), [13, 5, 2, 0, 1]);
```


`binarySearch(testArray, 2)` يجب أن تُرجع `[13, 5, 2]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 2), [13, 5, 2]);
```

`binarySearch(testArray, 6)` يجب أن تُرجع السلسلة النصية `Value Not Found`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.strictEqual(binarySearch(_testArray, 6), 'Value Not Found');
```

`binarySearch(testArray, 11)` يجب أن تُرجع `[13, 5, 10, 11]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 11), [13, 5, 10, 11])
```

`binarySearch(testArray, 13)` يجب أن تُرجع `[13]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 13), [13]);
```

`binarySearch(testArray, 70)` يجب أن تُرجع `[13, 19, 22, 49, 70]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 70), [13, 19, 22, 49, 70]);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let arrayPath = [];
  return arrayPath;
}
```



# --solutions--

```js
let binarySearch = (searchList, value) => {
  let arrayPath = [];

  // set initial L - M - R
  let left = 0;
  let right = searchList.length - 1;
  let middle = Math.floor(right / 2);

  // if first comparison finds value
  if (searchList[middle] == value) {
    arrayPath.push(searchList[middle]);
    return arrayPath;
  }

  while (searchList[middle] !== value) {
    // add to output array
    arrayPath.push(searchList[middle]);

    // not found
    if (right < left) {
      return 'Value Not Found';
    }
    // value is in left or right portion of array
    // update L - M - R
    if (searchList[middle] > value) {
      right = middle - 1;
      middle = left + Math.floor((right - left) / 2);
    } else {
      left = middle + 1;
      middle = left + Math.floor((right - left) / 2);
    }

    // if found update output array and exit
    if (searchList[middle] == value) {
      arrayPath.push(searchList[middle]);

      break;
    }
  }
  return arrayPath;
};
```
