---
id: a3f503de51cfab748ff001aa
title: أزواج العناصر
challengeType: 1
forumTopicId: 301617
dashedName: pairwise
---

# --description--

معطى مصفوفة `arr`، ابحث عن أزواج العناصر التي يكون مجموعها مساويًا للمعلمة الثانية `arg` وأرجع مجموع فهارسها.

يمكنك استخدام أزواج متعددة تحتوي على نفس العناصر الرقمية لكن بفهارس مختلفة. يجب أن يستخدم كل زوج أقل الفهارس المتاحة. بمجرد استخدام عنصر لا يمكن إعادة استخدامه للزوج مع عنصر آخر. على سبيل المثال، `pairwise([1, 1, 2], 3)` ينشئ الزوج `[2, 1]` باستخدام الرقم 1 عند الفهرس 0 بدلاً من الرقم 1 عند الفهرس 1، لأن 0+2 < 1+2.

على سبيل المثال، `pairwise([7, 9, 11, 13, 15], 20)` تُرجع `6`. الأزواج التي مجموعها 20 هي `[7, 13]` و`[9, 11]`. بعد ذلك يمكننا كتابة المصفوفة مع فهارسها وقيمها.

<div style='margin-left: 2em;'>

| الفهرس | 0 | 1 | 2  | 3  | 4  |
| ------ | - | - | -- | -- | -- |
| القيمة | 7 | 9 | 11 | 13 | 15 |

</div>

في الأسفل سنأخذ الفهارس المقابلة لها ونجمعها.

<div style='margin-left: 2em;'>

7 + 13 = 20 → الفهارس 0 + 3 = 3  
9 + 11 = 20 → الفهارس 1 + 2 = 3  
3 + 3 = 6 → أرجع `6`

</div>

# --hints--

`pairwise([1, 4, 2, 3, 0, 5], 7)` يجب أن تُرجع 11.

```js
assert.deepEqual(pairwise([1, 4, 2, 3, 0, 5], 7), 11);
```

`pairwise([1, 3, 2, 4], 4)` يجب أن تُرجع 1.

```js
assert.deepEqual(pairwise([1, 3, 2, 4], 4), 1);
```

`pairwise([1, 1, 1], 2)` يجب أن تُرجع 1.

```js
assert.deepEqual(pairwise([1, 1, 1], 2), 1);
```

`pairwise([0, 0, 0, 0, 1, 1], 1)` يجب أن تُرجع 10.

```js
assert.deepEqual(pairwise([0, 0, 0, 0, 1, 1], 1), 10);
```

`pairwise([], 100)` يجب أن تُرجع 0.

```js
assert.deepEqual(pairwise([], 100), 0);
```

# --seed--

## --seed-contents--

```js
function pairwise(arr, arg) {
  return arg;
}

pairwise([1,4,2,3,0,5], 7);
```

# --solutions--

```js
function pairwise(arr, arg) {
  var sum = 0;
  arr.forEach(function(e, i, a) {
    if (e != null) {
      var diff = arg-e;
      a[i] = null;
      var dix = a.indexOf(diff);
      if (dix !== -1) {
        sum += dix;
        sum += i;
        a[dix] = null;
      }
    }
  });
  return sum;
}

pairwise([1,4,2,3,0,5], 7);
```
