---
id: a7bf700cd123b9a54eef01d5
title: لا تكرارات من فضلك
challengeType: 1
forumTopicId: 16037
dashedName: no-repeats-please
---

# --description--

أرجع عدد التباديل الكلي لسلسلة نصية معطاة لا تحتوي على حروف متكررة متتالية. افترض أن كل الحروف في السلسلة النصية المعطاة فريدة.

على سبيل المثال، يجب أن تُرجع `aab` العدد 2 لأنها تحتوي على 6 تباديل إجمالية (`aab`، `aab`، `aba`، `aba`، `baa`، `baa`)، لكن فقط 2 منها (`aba` و`aba`) لا تحتوي على تكرار نفس الحرف (في هذه الحالة `a`).

# --hints--

يجب أن تُرجع الدالة `permAlone("aab")` رقمًا.

```js
assert.isNumber(permAlone('aab'));
```

يجب أن تُرجع الدالة `permAlone("aab")` العدد 2.

```js
assert.strictEqual(permAlone('aab'), 2);
```

يجب أن تُرجع الدالة `permAlone("aaa")` العدد 0.

```js
assert.strictEqual(permAlone('aaa'), 0);
```

يجب أن تُرجع الدالة `permAlone("aabb")` العدد 8.

```js
assert.strictEqual(permAlone('aabb'), 8);
```

يجب أن تُرجع الدالة `permAlone("abcdefa")` العدد 3600.

```js
assert.strictEqual(permAlone('abcdefa'), 3600);
```

يجب أن تُرجع الدالة `permAlone("abfdefa")` العدد 2640.

```js
assert.strictEqual(permAlone('abfdefa'), 2640);
```

يجب أن تُرجع الدالة `permAlone("zzzzzzzz")` العدد 0.

```js
assert.strictEqual(permAlone('zzzzzzzz'), 0);
```

يجب أن تُرجع الدالة `permAlone("a")` العدد 1.

```js
assert.strictEqual(permAlone('a'), 1);
```

يجب أن تُرجع الدالة `permAlone("aaab")` العدد 0.

```js
assert.strictEqual(permAlone('aaab'), 0);
```

يجب أن تُرجع الدالة `permAlone("aaabb")` العدد 12.

```js
assert.strictEqual(permAlone('aaabb'), 12);
```

# --seed--

## --seed-contents--

```js
function permAlone(str) {
  return str;
}

permAlone('aab');
```

# --solutions--

```js
function permAlone(str) {
  return permuter(str).filter(function(perm) {
    return !perm.match(/(.)\1/g);
  }).length;
}

function permuter(str) {
  // http://staff.roguecc.edu/JMiller/JavaScript/permute.html
  //permArr: Global array which holds the list of permutations
  //usedChars: Global utility array which holds a list of "currently-in-use" characters
  var permArr = [], usedChars = [];
  function permute(input) {
    //convert input into a char array (one element for each character)
    var i, ch, chars = input.split("");
    for (i = 0; i < chars.length; i++) {
      //get and remove character at index "i" from char array
      ch = chars.splice(i, 1);
      //add removed character to the end of used characters
      usedChars.push(ch);
      //when there are no more characters left in char array to add, add used chars to list of permutations
      if (chars.length === 0) permArr[permArr.length] = usedChars.join("");
      //send characters (minus the removed one from above) from char array to be permuted
      permute(chars.join(""));
      //add removed character back into char array in original position
      chars.splice(i, 0, ch);
      //remove the last character used off the end of used characters array
      usedChars.pop();
    }
  }
  permute(str);
  return permArr;
}

permAlone('aab');
```
