---
id: 5688e62ea601b2482ff8422b
title: 프로필 조회 기능 만들기
challengeType:  26
dashedName: lab-profile-lookup
---

# --description--

이 수업에서는 연락처 목록에 있는 사람들의 정보를 조회하는 프로필 조회 기능을 만듭니다.

예를 들어 연락처에 Akira Laine이라는 사람이 있다고 가정하면, `lookUpProfile("Akira", "lastName")`는 `Laine`을 반환해야 합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 수업을 완료하세요.

**사용자 스토리:**

1. 이름과 속성을 인수로 받는 `lookUpProfile`라는 함수를 만들어야 합니다.
2. 제공된 `contacts` 배열에서 연락처 정보를 가져와야 합니다.
3. 함수가 연락처 이름과 해당 연락처에 속성이 존재하면, 그 속성의 값을 반환해야 합니다.
4. 함수에 전달된 이름이 연락처 배열에 없으면 `"No such contact"`를 반환해야 합니다.
5. 찾은 연락처에 속성이 없으면 `"No such property"`를 반환해야 합니다.

# --before-each--

```js
const _contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

contacts = structuredClone(_contacts); 
```

# --hints--

`contacts` 배열은 전역 범위에 그대로 있어야 합니다. 삭제했다면 수업을 초기화하여 복구하세요.

```js
assert.isArray(contacts);
```

`lookUpProfile`라는 함수가 있어야 합니다.

```js
assert.isFunction(lookUpProfile);
```

`lookUpProfile("Kristian", "lastName")`는 문자열을 반환해야 합니다.

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

`lookUpProfile("Kristian", "lastName")`는 문자열 `Vos`을 반환해야 합니다.

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

`lookUpProfile("Sherlock", "likes")`는 `["Intriguing Cases", "Violin"]`를 반환해야 합니다.

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")`는 배열을 반환해야 합니다.

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

`lookUpProfile("Bob", "number")`는 문자열 `No such contact`를 반환해야 합니다.

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

`lookUpProfile("Bob", "potato")`는 문자열 `No such contact`를 반환해야 합니다.

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

`lookUpProfile("Akira", "address")`는 문자열 `No such property`를 반환해야 합니다.

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

`contacts` 배열은 `lookUpProfile` 실행 후에도 변경되지 않아야 합니다.

```js
lookUpProfile("Sherlock", "likes");
assert.deepEqual(contacts,_contacts);
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
