---
id: 64dc963d866fff3dd0329993
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Chapisha variable uliyotangaza katika hatua iliyopita na jisikie huru kubadilisha idadi ya diski kuona jinsi kiwango cha chini kinavyoongezeka kwa haraka kwa idadi ya hatua zinazohitajika.

# --hints--

Unapaswa kuchapisha variable ya `number_of_moves`.

```js
({ test: () => assert.match(code, /print\s*\(\s*number_of_moves\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move():
    pass
```
