---
id: 64dc92a9718fb53ca3d1bad7
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Changamoto ya Tower of Hanoi inaweza kutatuliwa kwa hatua 2<sup>n</sup> - 1, ambapo n ni idadi ya diski. Tangaza kigezo kinachoitwa `number_of_moves` na weka thamani ya jumla ya hatua hizi katika kigezo hiki.

Kipengele cha nguvu katika Python ni `**`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `number_of_moves`.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

Thamani ya `number_of_moves` inapaswa kuwa usemi wa kihesabu unaohesabu idadi ya hatua.

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
