---
id: 64dc9004f658183afa85f4f1
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Sasa angalia aina ya data ya funguo yako ya `'A'` kwa kuipitisha kwenye kitendakazi cha `type()` na chapisha matokeo kwenye terminali.

# --hints--

Unapaswa kupitisha funguo yako ya `'A'` kwenye kitendakazi cha `type()`.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

Unapaswa kuchapisha aina ya `rods['A']` kwa kutumia kitendakazi cha `print()`.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
