---
id: 64dc97ab25730b3e9de86ea5
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Mwisho wa mradi huu, utaunda suluhisho la kurudia (recursive) kwa fumbo la Tower of Hanoi, lakini sasa utaangalia njia ya kurudia (iterative) kwa changamoto hii.

Anza kwa kuongeza mzunguko wa `for` kwenye kitendakazi chako unaorudia kupitia `number_of_moves` na kuchapisha nambari ya mzunguko wa sasa.

# --hints--

Unapaswa kuandika mzunguko wa `for` kurudia kupitia idadi ya hatua. Tumia kitendakazi cha `range()` kwa ajili hiyo.

```js
({ test: () => {
    const rgs = [
        /for\s+\w+\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

Unapaswa kuchapisha nambari ya hatua ya sasa kila mzunguko.

```js
({ test: () => {
    const rgs = [
        /for\s+(\w+)\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:\s+print\s*\(\s*\1\s*\)/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
