---
id: 64dc922df2919e3c38cead61
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Juu kabisa ya msimbo wako, tangaza kigezo kinachoitwa `NUMBER_OF_DISKS` kuhifadhi idadi ya diski na kipa thamani ya `3`. Kisha, badilisha hoja ya kwanza inayopitishwa kwenye kitendakazi cha `range()` na kigezo chako kipya.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `NUMBER_OF_DISKS`.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

Kigezo cha `NUMBER_OF_DISKS` kinapaswa kuwa na thamani `3`.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

Unapaswa kubadilisha thamani ya hoja ya kwanza ya kitendakazi cha `range()` kuwa kigezo cha `NUMBER_OF_DISKS`.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
