---
id: 64dc90949ff85e3b37be40d0
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Kitendakazi cha `range()` kinarejesha mfululizo usiobadilika wa nambari. Kama unavyoona, aina ya data ya `rods['A']` ni `range`, lakini unataka iwe orodha.

Pitia wito wako wa `range()` kwa kitendakazi cha `list()` ili kufanya hivyo.

# --hints--

Unapaswa kupitisha wito lako la `range()` kwa kitendakazi cha `list()`.

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
