---
id: 64dc8b1faf9c60304ca8a9b6
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika mradi huu, utatatua fumbo la kihesabu linalojulikana kama Mnara wa Hanoi. Fumbo hili lina nguzo tatu na diski kadhaa za ukubwa tofauti.

Lengo la fumbo hili ni kusogeza diski kutoka kwenye nguzo ya kwanza hadi nguzo ya tatu, ukifuata sheria maalum zinazozuia kuweka diski kubwa juu ya diski ndogo.

Anza kwa kuunda kamusi tupu iitwayo `rods` ili wakilisha nguzo.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `rods`.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

Kigezo chako cha `rods` kinapaswa kuwa kamusi tupu.

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
