---
id: 64dc916138dbe03bdfe6cec6
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Lengo la Tower of Hanoi ni kuhamisha diski zote hadi fimbo ya mwisho. Ili kufanya hivyo, lazima ufuate sheria tatu rahisi:

1. Unaweza kuhamisha diski zilizo juu kabisa tu
2. Unaweza kuhamisha diski moja kwa wakati mmoja tu
3. Huwezi kuweka diski kubwa juu ya diski ndogo

Chini ya msimbo wako uliopo, tangaza kitendakazi tupu kinachoitwa `move`. Baadaye, utatumia kitendakazi hicho kuhamisha diski kati ya fimbo. Kwa sasa, ili kuepuka makosa, tumia neno kuu la `pass` ndani ya mwili wa kitendakazi.

# --hints--

Unapaswa kutangaza kitendakazi tupu kinachoitwa `move`. Kumbuka kutumia neno kuu la `pass` ndani ya mwili wa kitendakazi kwa uelekeo sahihi.

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
