---
id: 64dc8ea01436383a88256d7f
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Changamoto inaanza na diski zilizokusanywa kwenye fimbo ya kwanza, kwa ukubwa unaopungua, ambapo diski ndogo zaidi iko juu na diski kubwa zaidi iko chini. Unahitaji kujaza orodha yako ya kwanza na nambari zinazowakilisha ukubwa mbalimbali wa diski.

Badala ya kuongeza vitu kwa mkono kwenye orodha ya kwanza, tengeneza mfululizo wa nambari unaopungua kutoka `3` hadi `1` kwa kutumia kitendakazi cha `range()` na weka thamani yake kwenye `rods['A']`. Hapa, `3` inawakilisha diski kubwa zaidi iliyo chini ya mkusanyiko na `1` inawakilisha diski ndogo zaidi iliyo juu ya mkusanyiko.

Sintaksia ni `range(x, y, h)`, ambapo `x` ni nambari ya kuanzia (ikiwa ni pamoja na), `y` ni nambari ya mwisho (isiyo jumuishi), na `h` ni tofauti kati ya nambari moja na nyingine inayofuata katika mfululizo.

# --hints--

Unapaswa kutumia kitendakazi cha `range()` kuweka thamani ya mfululizo wa nambari kwenye `rods['A']`. Sintaksia ya kuita kitendakazi cha `range()` ni `range(x, y, h)`.

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

Unapaswa kutumia kitendakazi cha `range()` kuweka thamani ya mfululizo wa nambari kutoka `3` hadi `1` kwenye `rods['A']`.

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
