---
id: 64dc985e6720a23edac63b51
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Miondoko ya diski inayoruhusiwa kati ya nguzo inaonyesha muundo unaojirudia kila miondoko mitatu. Kwa mfano, miondoko kati ya nguzo `A` na nguzo `C` inaruhusiwa katika miondoko ya kwanza, ya nne, na ya saba, ambapo mabaki ya mgawanyo kati ya nambari ya miondoko na 3 ni sawa na 1.

Ndani ya mzunguko wa `for` uliotengenezwa awali, badilisha wito wa `print()` uliopo na tamko la `if` linalochochewa wakati `(i + 1) % 3 == 1`. Ndani ya tamko la `if`, chapisha `f'Move {i + 1} allowed between {source} and {target}'` kwa kutumia f-string. Tafadhali, kumbuka kuwa `i + 1` ni nambari ya miondoko tangu `i` iwe ni sifuri wakati wa mzunguko wa kwanza.

# --hints--

Unapaswa kuwa na tamko la `if` ambalo linapaswa kutekelezwa wakati `(i + 1) % 3 == 1`.

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

Unapaswa kuchapisha mfuatano wa herufi `f'Move {i + 1} allowed between {source} and {target}'` kwa kutumia f-string.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
