---
id: 64dc97005bc7943e2978df0a
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Sasa unaweza kuondoa wito wako wa `print()`. Kisha, ndani ya kitendakazi cha `move()`, ondoa neno kuu la `pass` na chapisha yaliyomo ya kamusi yako ya `rods`.

# --hints--

Haupaswi kuwa na `print(number_of_moves)` katika msimbo wako.

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

Unapaswa kuondoa neno kuu la `pass` na kuchapisha kamusi ya `rods`.

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
