---
id: 64de1469fccce61940f498c4
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Kamusi ya `rods` itawakilisha fimbo tatu zilizo na diski zao. Ipe mfuatano wa herufi `'A'`, `'B'`, na `'C'` kama funguo na weka kila moja yao kuwa orodha tupu.

# --hints--

Kamusi yako ya `rods` inapaswa kuwa na ufunguo wa `'A'`.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` inapaswa kuwa orodha tupu.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

Kamusi yako ya `rods` inapaswa kuwa na ufunguo wa `'B'`.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` inapaswa kuwa orodha tupu.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

Kamusi yako ya `rods` inapaswa kuwa na ufunguo wa `'C'`.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` inapaswa kuwa orodha tupu.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
