---
id: 64df45a3ad4f8719e5355244
title: Hatua 46
challengeType: 20
dashedName: step-46
---

# --description--

Hatua za kuhama diski za `n - 1` zinaweza kugawanywa zaidi hadi diski moja tu itakapotazamwa. Hii itakuwa ni hatua ya kwanza inayotokea. Baada ya hatua ya kwanza kutokea, hatua zinazofuata hutengenezwa na kuondolewa kwa wito wa kurudia. Kumbuka kwamba katika kila hatua ya kurudia, jukumu linalochezwa na fimbo hubadilika kati ya chanzo, lengo, na msaada.

Kwa sasa, kila wito wa kurudia huchapisha kamusi ya `rods` bila kufanya mabadiliko yoyote kwenye orodha. Kabla ya wito wa `print()`, ondoa kipengele cha mwisho kutoka kwenye orodha ya `rods[source]` na uiongeze kwenye orodha ya `rods[target]`.

# --hints--

Unapaswa kuondoa kipengele cha mwisho kutoka kwenye orodha ya `rods[source]` na kuiongeza kwenye orodha ya `rods[target]` kabla ya wito wa `print`.

```js
({ test: () => assert.isTrue(runPython(`
_log = []
def capture_print(*args):
    _log.append(repr(args))

old_print = print
print = capture_print

old_rods = rods
rods2 = {
    'A': list(range(5, 0, -1)),
    'B': [],
    'C': []
}
rods = rods2

move(5, 'A', 'B', 'C')
rods = old_rods
print = old_print

_expected_prints = [
r"({'A': [5, 4, 3, 2], 'B': [], 'C': [1]}, '\\n')",
r"({'A': [5, 4, 3], 'B': [], 'C': [1, 2]}, '\\n')",
r"({'A': [5, 4], 'B': [], 'C': [1, 2, 3]}, '\\n')",
r"({'A': [5], 'B': [], 'C': [1, 2, 3, 4]}, '\\n')",
r"({'A': [], 'B': [], 'C': [1, 2, 3, 4, 5]}, '\\n')",
]

for args in _log:
    if not _expected_prints:
        break
    if args == _expected_prints[0]:
        _expected_prints.pop(0)

not rods2['A'] and rods2['C'] == list(range(1, 6)) and not _expected_prints
`))})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # display starting configuration
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
