---
id: 64df496c6a8ddf1b38db1ed6
title: Hatua 50
challengeType: 20
dashedName: step-50
---

# --description--

Katika hatua iliyopita, uliandika msimbo wa kusogeza diski kubwa zaidi ya changamoto ndogo kwenda kwenye fimbo lengwa.

Sasa, unachohitaji kufanya ni kuongeza wito mwingine wa kurudia (recursive) kusogeza diski `n - 1` ulizozihamisha tayari. Nakili wito wa kwanza wa kurudia na ubandike mwishoni mwa kipengele cha `if`.

Kumbuka kuwa hoja za kitendakazi haziko katika mpangilio sahihi. Jaribu kubaini mpangilio sahihi.

# --hints--

Unapaswa kuita `move(n - 1,  auxiliary, source, target)` mwishoni mwa kitendakazi chako cha `move`.

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
