---
id: 64df47b32b92301a815d5ef8
title: Hatua 49
challengeType: 20
dashedName: step-49
---

# --description--

Mwanzo, wito wa kurudia (recursive call) uliouongeza unashughulikia changamoto ndogo ya kuhamisha `n - 1` diski kwenda fimbo ya pili.

Kwa sababu hiyo, hoja ya `target` inahusiana na fimbo yako ya pili, wakati hoja ya `auxiliary` ni fimbo ya tatu. Kumbuka kwamba hizo zitabadilishana nafasi kadri kurudia kunavyoendelea.

Rekebisha mpangilio wa hoja kwa kubadilisha `target` na `auxiliary` katika wito wako wa kurudia.

# --hints--

Unapaswa kubadilisha mpangilio wa hoja katika wito lako la `move(n - 1, source, auxiliary, target)`.

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
