---
id: 616968c2c94c8071b349c146
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dentro del `body`, anida un elemento `h1` con el texto `CSS Color Markers`.

# --hints--

Tu código debe tener un elemento `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Debe haber una etiqueta de apertura `h1`.

```js
assert.match(code, /<h1\s*>/i);
```

Debe haber una etiqueta de cierre `h1`.

```js
assert.match(code, /<\/h1\s*>/i);
```

Tu elemento `h1` debe estar dentro del elemento `body`.

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

Tu elemento `h1` debe tener el texto `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Recuerda que las mayúsculas y la ortografía son importantes para el texto del `h1`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
