---
id: 616d595270d902f0e7086e18
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Mientras tienes tres elementos `div` marcadores separados, parecen un gran rectángulo. Debes añadir algo de espacio entre ellos para facilitar ver cada elemento.

Cuando la propiedad abreviada `margin` tiene dos valores, establece `margin-top` y `margin-bottom` al primer valor, y `margin-left` y `margin-right` al segundo valor.

En tu regla CSS `.marker`, establece la propiedad `margin` en `10px auto`.

# --hints--

Tu regla CSS `.marker` debe tener una propiedad `margin` establecida en `10px auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
