---
id: 616d4a84b756d9c4b8255093
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Es hora de añadir algo de color al marcador. Recuerda que una manera de añadir color a un elemento es usar una <dfn>palabra clave de color</dfn> como `black`, `cyan` o `yellow`.

Como recordatorio, así es como puedes dirigir a la clase `freecodecamp`:

```css
.freecodecamp {

}
```

Crea una nueva regla CSS que dirija a la clase `marker` y establece su propiedad `background-color` a `red`.

**Nota:** No verás cambios después de añadir el CSS.

# --hints--

Deberías crear un selector de clase para dirigir a la clase `marker`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Tu regla CSS `.marker` debe tener la propiedad `background-color` establecida a `red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
