---
id: 61a49d15bdbb5e57cc6fd280
title: Paso 49
challengeType: 0
dashedName: step-49
---

# --description--

Los puntos de color te permiten afinar dónde se colocan los colores a lo largo de la línea de degradado. Se utilizan en la función `linear-gradient` para especificar dónde termina un color y comienza la transición al siguiente color. Se puede usar un porcentaje o un valor de longitud para definir la posición del punto de color.

Por ejemplo, en este degradado rojo-negro, la transición de rojo a negro ocurre en el punto 90% a lo largo de la línea de degradado, por lo que el rojo ocupa la mayor parte del espacio disponible:

```css
linear-gradient(90deg, red 90%, black);
```

En la función `linear-gradient`, añade un punto de color `75%` después del primer argumento de color rojo. No añadas puntos de color a los otros argumentos de color.

# --hints--

La regla CSS `.red` debe tener una propiedad `background` configurada en `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
