---
id: 66a9bcf00f13a418368a272e
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Елемент заголовка таблиці складається з елемента <dfn>рядок таблиці</dfn>, `tr`, який містить елементи <dfn>клітинка заголовка таблиці</dfn>, `th`.

Ось приклад із використанням елементів `tr` та `th` для спортивної таблиці:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Всередині вашого елемента `thead` додайте елемент `tr`.

Всередині вашого елемента `tr` додайте три елементи `th`.

Перший елемент `th` має містити текст `Last Name`. Другий елемент `th` має містити текст `First Name`. Третій елемент `th` має містити текст `Grade`.

# --hints--

У вас має бути відкриваючий тег `tr`.

```js
assert.match(code, /<tr>/i);
```

У вас має бути закриваючий тег `tr`.

```js
assert.match(code, /<\/tr>/i);
```

Ваш елемент `tr` має бути всередині вашого елемента `thead`.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

Перший елемент `th` має містити текст `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

Другий елемент `th` має містити текст `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

Третій елемент `th` має містити текст `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

Всі ваші елементи `th` мають бути всередині вашого елемента `tr`.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
