---
id: 66a9b8f14b963916a3baa732
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Щоб додати підпис до `table`, ви можете використати елемент <dfn>table caption</dfn>.

Ось приклад із використанням елемента `caption`:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

Всередині вашого елемента `table` вкладіть елемент `caption` з текстом `Calculus Final Exam Grades`.

# --hints--

У вас має бути відкриваючий тег `caption`.

```js
assert.match(code, /<caption>/i);
```

У вас має бути закриваючий тег `caption`.

```js
assert.match(code, /<\/caption>/i);
```

Ваш елемент `caption` має бути всередині вашого елемента `table`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

Ваш `caption` має містити текст `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
