---
id: 66a9c006e4f58819396605b3
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Тепер, коли ви завершили розділ head, настав час додати <dfn>тіло таблиці</dfn>, `tbody`. Тіло таблиці представить усі імена студентів та їхні оцінки.

Додайте елемент тіла таблиці до своєї таблиці.

# --hints--

Ви повинні мати відкриваючий тег `tbody`.

```js
assert.match(code, /<tbody>/i);
```

Ви повинні мати закриваючий тег `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

Ваш відкриваючий тег `tbody` має йти перед закриваючим тегом `table`.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
