---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Тепер настав час додати до таблиці ще двох студентів.

Дотримуючись того самого шаблону, що й у попередньому кроці, додайте другий рядок таблиці для студента. Використайте такі дані для елементів таблиці:

- Прізвище: `Doe`
- Ім’я: `Samantha`
- Оцінка: `92`

Для третього рядка таблиці студента використайте такі дані для елементів таблиці:

- Прізвище: `Rodriguez`
- Ім’я: `Marcus`
- Оцінка: `88`

# --hints--

У вашому елементі `tr` має бути другий `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

У другому елементі `td` має бути елемент `Doe` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

У другому елементі `td` має бути елемент `Samantha` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

У другому елементі `td` має бути елемент `92` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

У вашому елементі `tr` має бути третій `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

У третьому елементі `td` має бути елемент `Rodriguez` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

У третьому елементі `td` має бути елемент `Marcus` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

У третьому елементі `td` має бути елемент `88` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
