---
id: 66a9c14e3b34c719e34bae20
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Щоб додати дані студентів до таблиці, потрібно використати елементи рядка таблиці та даних таблиці.

Елемент <dfn>table data</dfn>, `td`, використовується для створення клітинки в таблиці.

Ось приклад використання елемента `td` для таблиці спортивних гравців:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Всередині елемента тіла таблиці додайте елемент рядка таблиці, `tr`, з трьома елементами даних таблиці, `td`.

Перший елемент даних таблиці має містити прізвище `Davis`.

Другий елемент даних таблиці має містити ім’я `Alex`.

Третій елемент даних таблиці має містити оцінку `54`.

# --hints--

Усередині вашого елемента `tr` має бути елемент `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Усередині елемента `td` має бути елемент `Davis` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

Усередині елемента `td` має бути елемент `Alex` з текстом `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

Усередині елемента `td` має бути елемент `54` з текстом `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

Усередині елемента `td` мають бути всі три елементи `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
