---
id: 66a9b8f14b963916a3baa732
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Щоб додати підпис до таблиці, можна використати елемент <dfn>caption</dfn>.

Ось приклад використання елемента `caption`:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

Додайте елемент `caption` з текстом `Calculus Final Exam Grades` до елемента `table`.

# --hints--

Ви повинні мати початковий тег `caption`.

```js
assert.match(code, /<caption>/i);
```

Ви повинні мати кінцевий тег `caption`.

```js
assert.match(code, /<\/caption>/i);
```

Елемент `caption` має бути в межах елемента `table`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

Елемент `caption` повинен мати текст `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--

    --fcc-editable-region--      
    </table>
  </body>
</html>
```
