---
id: 66a9cb26ec6bd41cf6c82bc5
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Останнім розділом, який потрібно додати до таблиці, буде елемент <dfn>table foot</dfn>, `tfoot`. Елемент table foot використовуватиметься для відображення середнього балу для всіх студентів у таблиці.

Додайте елемент `tfoot` до таблиці.

Всередині елемента `tfoot` додайте елемент `tr`.

Всередині елемента `tr` додайте два елементи `td`.

Перший елемент `td` має містити текст `Average Grade`. Другий елемент `td` має містити текст `78.8`.

# --hints--

Ви повинні мати відкриваючий тег `tfoot`.

```js
assert.match(code, /<tfoot>/);
```

Ви повинні мати закриваючий тег `tfoot`.

```js
assert.match(code, /<\/tfoot>/);
```

У вашому елементі `tr` має бути елемент `tfoot`.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

Ваш перший елемент `td` має містити текст `Average Grade`. Переконайтеся, що він знаходиться всередині елемента `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

Ваш другий елемент `td` має містити текст `78.8`. Переконайтеся, що він знаходиться всередині елемента `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

Ваш елемент `tr` має містити два елементи `td` всередині.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
