---
id: 66a9bb8578459a175432e7d0
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Для першого розділу таблиці потрібно згрупувати вміст заголовка, який представляє мітки стовпців для імені студента, прізвища та оцінки за фінальний екзамен.

Елемент <dfn>table head</dfn>, `thead`, використовується для групування вмісту заголовка в таблиці.

Ось приклад із використанням елемента `thead`:

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

Під вашим елементом `caption` додайте елемент table head.

# --hints--

Ви повинні мати відкриваючий тег `thead`.

```js
assert.match(code, /<thead>/i);
```

Ви повинні мати закриваючий тег `thead`.

```js
assert.match(code, /<\/thead>/i);
```

Ваш елемент `thead` має бути всередині вашого елемента `table`.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
