---
id: 66a9ca099e3ecb1be10a2696
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Тепер настав час додати до таблиці ще двох студентів.

Дотримуючись того самого шаблону, що й у попередніх кроках, додайте четвертий рядок таблиці для студента. Використайте такі дані для елементів таблиці:

- Прізвище: `Thompson`
- Ім’я: `Jane`
- Оцінка: `77`

Для п’ятого рядка таблиці студента використайте такі дані для елементів таблиці:

- Прізвище: `Williams`
- Ім’я: `Natalie`
- Оцінка: `83`

# --hints--

У вашому елементі `tr` має бути четвертий `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

У четвертому елементі `td` має бути елемент `Thompson` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

У четвертому елементі `td` має бути елемент `Jane` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

У четвертому елементі `td` має бути елемент `77` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

У вашому елементі `tr` має бути п’ятий `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

У п’ятому елементі `td` має бути елемент `Williams` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

У п’ятому елементі `td` має бути елемент `Natalie` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

У п’ятому елементі `td` має бути елемент `83` з текстом `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
