---
id: 66a9b48a211a73155621d0df
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з HTML-таблицями, створюючи таблицю фінального екзамену для групи студентів.

У попередніх уроках ви навчились працювати з елементом `table`, щоб представити табличні дані.

Всередині вашого елемента `body` вкладіть елемент `table`.

# --hints--

У вас має бути відкриваючий тег `table`.

```js
assert.match(code, /<table>/i);
```

У вас має бути закриваючий тег `table`.

```js
assert.match(code, /<\/table>/i);
```

Ваш елемент `table` має бути всередині вашого елемента `body`.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
