---
id: ab6137d4e35944e21037b769
title: 제목 대문자 변환기 만들기
challengeType: 26
dashedName: build-a-title-case-converter
---

# --description--

이 실습에서는 문자열을 제목 대문자 형식으로 변환하는 함수를 만듭니다. 제목 대문자 형식은 각 단어의 첫 글자는 대문자로, 나머지 글자는 소문자로 표기하는 것을 의미합니다.

`"Web Development Is Awesome"`는 제목 대문자 형식의 문자열 예시입니다.

**목표**: 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `titleCase` 함수가 문자열을 인수로 받아야 합니다.
1. `titleCase` 함수는 각 단어의 첫 글자는 대문자, 나머지는 소문자인 문자열을 반환해야 합니다.
1. `titleCase("I like to code")`는 `"I Like To Code"`를 반환해야 합니다.
1. `titleCase("javaScript is fun")`는 `"Javascript Is Fun"`를 반환해야 합니다.

# --hints--

`titleCase`라는 이름의 함수를 만드세요.

```js
assert.isFunction(titleCase);
```

`titleCase`는 하나의 인수를 받아야 합니다.

```js
assert.lengthOf(titleCase, 1);
```


`titleCase("I'm a little tea pot")`는 문자열을 반환해야 합니다.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")`는 문자열 `I'm A Little Tea Pot`를 반환해야 합니다.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")`는 문자열 `Short And Stout`를 반환해야 합니다.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")`는 문자열 `Here Is My Handle Here Is My Spout`를 반환해야 합니다.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}
```
