---
id: 6998b99a8e195cd0c1f211df
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將使用 JavaScript 物件建造一個簡單的野生動物追蹤器。

這是一個物件的範例：

```js
const animal = {}; // empty object
```

這會創建一個空物件。

現在創建一個名為 `tiger` 的變數，並將它指定為一個空物件。

# --hints--

你應該創建一個名為 `tiger` 的變數。

```js
assert.exists(tiger);
```

`tiger` 變數應該被指定為一個空物件。

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
