---
id: 69999ff7f7862fdfb885ea5d
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

現在更新你的 `tiger` 物件，新增一個名為 `age` 的屬性。

將 `age` 屬性設定為值 `5`。

# --hints--

`tiger` 應該有一個 `age` 屬性。

```js
assert.property(tiger, "age");
```

`tiger.age` 應該相等於 `5`。

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
