---
id: 6999b07c34c2cde21ee6e1d8
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

在這一步，你將創建一個函式，用來檢查物件是否擁有特定的屬性。

以下是使用 `hasOwnProperty` 的範例：

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

在這個範例中，`hasColor` 函式使用 `hasOwnProperty` 來檢查物件上是否存在 `color` 屬性。

現在創建一個名為 `hasHabitat` 的函式。該函式應該接受一個名為 `animal` 的參數。

使用 `hasOwnProperty` 方法，如果 `true` 物件擁有名為 `animal` 的屬性，則傳回 `"habitat"`，否則傳回 `false`。

創建函式後，使用 `console.log` 呼叫 `hasHabitat(tiger)` 和 `hasHabitat(elephant)`，以便你能在主控台看到兩個結果。

# --hints--

你應該創建一個名為 `hasHabitat` 的函式。

```js
assert.isFunction(hasHabitat);
```

`hasHabitat` 函式應該有一個參數，名為 `animal`。

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` 應該使用 `hasOwnProperty` 方法來檢查 `habitat` 屬性。

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

你應該將 `hasHabitat(tiger)` 記錄到主控台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

你應該將 `hasHabitat(elephant)` 記錄到主控台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

呼叫 `hasHabitat(tiger)` 在添加 habitat 後應該傳回 `true`。

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

呼叫 `hasHabitat(elephant)` 如果沒有 habitat，應該傳回 `false`。

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` 應該使用函式參數並能處理任何物件。

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
