---
id: 6999adc72c02812a28c01e31
title: 第 9 步驟
challengeType: 1
dashedName: step-9
---

# --description--

在這個步驟中，你將創建一個函式，用來更新物件的屬性。

以下是一個在函式內更新屬性的範例：

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

在這個範例中，`age` 屬性被更新為一個新值。

現在創建一個名為 `updateAge` 的函式。該函式應該接受兩個參數：`animal` 和 `newAge`。

在函式內，將 `age` 物件的 `animal` 屬性更新為 `newAge`。傳回更新後的 `animal` 物件。

創建函式後，使用 `console.log` 呼叫 `updateAge(elephant, 12)`，以便你可以在主控台看到更新後的 `elephant` 物件。

# --hints--

你應該創建一個名為 `updateAge` 的函式。

```js
assert.isFunction(updateAge);
```

`updateAge` 函式應該有兩個參數：`animal` 和 `newAge`。

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

`updateAge` 應該使用點記法來更新 `age` 屬性。

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

`updateAge` 函式應該傳回更新後的 `animal` 物件。

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

你應該將 `updateAge(elephant, 12)` 記錄到主控台。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

呼叫 `updateAge(elephant, 12)` 應該會更新 age 屬性。

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` 應該使用函式參數並能處理任何物件。

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
