---
id: 6999b0198d18d402c36cc5fe
title: 第 10 步驟
challengeType: 1
dashedName: step-10
---

# --description--

在這個步驟中，你將創建一個函式，用來從物件中移除一個屬性。

以下是使用 `delete` 移除屬性的範例：

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

在這個範例中，`canFly` 屬性從 `bird` 物件中被移除。

現在創建一個名為 `removeEndangeredStatus` 的函式。該函式應該接受一個名為 `animal` 的參數。

在函式內，使用 `isEndangered` 關鍵字從 `animal` 物件中移除 `delete` 屬性。傳回更新後的 `animal` 物件。

創建函式後，使用 `console.log` 呼叫 `removeEndangeredStatus(tiger)`，以便在主控台中查看更新後的物件。

# --hints--

你應該創建一個名為 `removeEndangeredStatus` 的函式。

```js
assert.isFunction(removeEndangeredStatus);
```

`removeEndangeredStatus` 函式應該有一個參數，名為 `animal`。

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`removeEndangeredStatus` 函式應該傳回更新後的 `animal` 物件。

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` 應該使用 `delete` 關鍵字來移除 `isEndangered` 屬性。

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

你應該將 `removeEndangeredStatus(tiger)` 記錄到主控台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

呼叫 `removeEndangeredStatus(tiger)` 應該會移除 `isEndangered` 屬性。

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` 應該使用函式參數並能處理任何物件。

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
