---
id: 6999b0d79ea938edcb032237
title: 第 12 步驟
challengeType: 1
dashedName: step-12
---

# --description--

在 Wildlife Tracker 工作坊的最後一步，你將創建一個函式，使用中括號表示法來存取物件的屬性。

以下是使用中括號表示法存取屬性的範例：

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

在這個範例中，中括號表示法讓你可以使用變數來存取屬性。

現在創建一個名為 `getProperty` 的函式。

該函式應該接受兩個參數：`animal` 和 `propertyName`。

使用中括號表示法傳回該屬性的值。

創建函式後，使用 `console.log` 呼叫 `getProperty(tiger, "species")` 和 `getProperty(elephant, "age")`，以便你能在主控台看到傳回的值。

這樣，Wildlife Tracker 工作坊就完成了！

# --hints--

你應該創建一個名為 `getProperty` 的函式。

```js
assert.isFunction(getProperty);
```

`getProperty` 函式應該有兩個參數：`animal` 和 `propertyName`。

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` 應該使用中括號表示法來存取屬性。

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

你應該將 `getProperty(tiger, "species")` 輸出到主控台。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

你應該將 `getProperty(elephant, "age")` 輸出到主控台。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

呼叫 `getProperty(tiger, "species")` 應該傳回 `"Tiger"`。

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

呼叫 `getProperty(elephant, "age")` 應該傳回 `12`。

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` 應該使用函式參數並能處理任何物件。

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
