---
id: 6999a0c1698b93d5d41156ed
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

現在為 `tiger` 物件新增另一個名為 `isEndangered` 的屬性。

將 `isEndangered` 屬性設定為 `true`。

# --hints--

`tiger` 應該有一個 `isEndangered` 屬性。

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` 應該相等於 `true`。

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
