---
id: 6999faf886634b72be7a819b
title: 第 2 步驟
challengeType: 1
dashedName: step-2
---

# --description--

在這個野生動物追蹤器中，你將更新一個物件，加入新的屬性。

以下是一個建立帶有屬性的物件範例：

```js
const animal = {
  name: "Lion" // property: value
};
```

現在更新 `tiger` 物件，使其包含一個名為 `species` 的屬性。

將它的值設為 `"Tiger"`。

# --hints--

`tiger` 應該是一個物件。

```js
assert.isObject(tiger);
```

`tiger` 應該有一個 `species` 屬性。

```js
assert.property(tiger, "species");
```

`tiger.species` 應該相等於 `"Tiger"`。

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
