---
id: 6999ac171928136168654774
title: 第 6 步驟
challengeType: 1
dashedName: step-6
---

# --description--

在這個步驟中，你將創建一個函式來存取物件的 `species` 屬性。

以下是一個從物件存取屬性的簡單範例：

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

在這個範例中，函式接受一個物件作為參數，並傳回它的 `species` 屬性。

現在創建一個名為 `getSpecies` 的函式。

該函式應該接受一個名為 `animal` 的參數。

使用點記法傳回 `species` 的 `animal` 屬性。

創建函式後，使用 `console.log` 呼叫 `getSpecies(tiger)`，以便你可以在主控台看到結果。

# --hints--

你應該創建一個名為 `getSpecies` 的函式。

```js
assert.isFunction(getSpecies);
```

`getSpecies` 函式應該有一個參數，名為 `animal`。

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

你應該將 `getSpecies(tiger)` 輸出到主控台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` 應該使用點記法存取 species 屬性。

```js
assert.match(code, /animal\.species/);
```

呼叫 `getSpecies(tiger)` 應該傳回 `"Tiger"`。

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` 應該使用函式參數，並能處理任何物件。

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
