---
id: 699a00de9f564bb0effaa14d
title: 第 7 步驟
challengeType: 1
dashedName: step-7
---

# --description--

在這個步驟中，你將創建一個函式來存取物件的 `age` 屬性。

以下是一個範例：

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

在這個範例中，函式接受一個物件作為參數，並傳回它的 `age` 屬性。

現在創建一個名為 `getAge` 的函式。

該函式應該接受一個名為 `animal` 的參數。

使用點記法傳回 `age` 的 `animal` 屬性。

創建函式後，使用 `console.log` 呼叫 `getAge(tiger)`，以便你可以在主控台看到結果。

# --hints--

你應該創建一個名為 `getAge` 的函式。

```js
assert.isFunction(getAge);
```

`getAge` 函式應該有一個參數，名為 `animal`。

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` 應該使用點記法存取 `age` 屬性。

```js
assert.match(code, /animal\.age/);
```

你應該將 `getAge(tiger)` 記錄到主控台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

呼叫 `getAge(tiger)` 應該傳回 `5`。

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` 應該使用函式參數，並能處理任何物件。

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
