---
id: 69ba8da6b296391d913096a0
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

現在創建第二個名為 `elephant` 的物件。

添加以下屬性：

- `species`，值為 `"Elephant"`
- `age`，值為 `10`
- `isEndangered`，值為 `true`

# --hints--

你應該創建一個名為 `elephant` 的變數。

```js
assert.exists(elephant);
```

`elephant` 應該有一個 `species` 屬性，其值相等於 `"Elephant"`。

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` 應該有一個 `age` 屬性，其值相等於 `10`。

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` 應該有一個 `isEndangered` 屬性，其值相等於 `true`。

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
