---
id: 669fcb06c3034a39f5431a38
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare creating some actual content, add a `body` element below the `head` element.

# --hints--

Du solltest ein öffnendes `<body>`-Tag haben.

```js
assert.match(code, /<body>/i);
```

Du solltest ein schließendes `</body>`-Tag haben.

```js
assert.match(code, /<\/body>/i);
```

Du solltest dein `head`-Element nicht ändern. Stelle sicher, dass du dein schließendes Tag nicht gelöscht hast.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Dein `body`-Element sollte nach deinem `head`-Element stehen.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>

  --fcc-editable-region--
</html>
```
