---
id: 669fdd3965f36f40af9615e5
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Da Sie nun die Kopfzeile der Seite fertiggestellt haben, müssen Sie beginnen, Ihren Hauptinhalt hinzuzufügen.

Fügen Sie unter Ihrem `header` ein `main`-Element hinzu.

# --hints--

Ihr `main`-Element sollte ein öffnendes Tag haben.

```js
assert.match(code, /<main\>/);
```

Ihr `main`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/main\>/);
```

Ihr `main`-Element sollte unter Ihrem `header`-Element stehen.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
