---
id: 669fdfc9a5e3da42d2376609
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Jetzt, da Sie den Abschnitt „about“ hinzugefügt haben, versuchen Sie, auf den `About`-Link zu klicken, um zu sehen, wie die Seite zu diesem Abschnitt springt.

Der nächste Abschnitt auf der Blogseite wird eine Liste von Beiträgen sein, die über Mr. Whiskers sprechen.

Fügen Sie ein weiteres `section`-Element mit einem `id` hinzu, das auf `"posts"` gesetzt ist.

Fügen Sie innerhalb des `section`-Elements ein `h2`-Element mit dem Text `Posts` hinzu.

# --hints--

Ihr Code sollte ein öffnendes `section`-Tag enthalten.

```js
assert.match(code, /<section.*>/i);
```

Ihr Code sollte ein schließendes `section`-Tag enthalten.

```js
assert.match(code, /<\/section>/i);
```

Ihr `section`-Element sollte ein `id` haben, das auf `"posts"` gesetzt ist.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

Ihr `section`-Element sollte sich innerhalb Ihres `main`-Elements befinden.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Ihr `h2`-Element sollte sich innerhalb des `section`-Elements befinden.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

Ihr `h2`-Element sollte den Text `Posts` enthalten.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
