---
id: 669fdc11c9b0263fe0814a7a
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Innerhalb jedes `li`-Elements müssen Sie ein Ankerelement haben.

Für das erste Ankerelement sollte der Text `About` sein und der Wert des `href`-Attributs sollte `"#about"` sein. Das Rautezeichen vor `about` steht für einen `id`-Namen, der später im Projekt hinzugefügt wird.

Für das zweite Ankerelement sollte der Text `Posts` sein und der Wert des `href`-Attributs sollte `"#posts"` sein.

Für das dritte Ankerelement sollte der Text `Contact` sein und der Wert des `href`-Attributs sollte `"#contact"` sein.

# --hints--

Ihr erstes Ankerelement sollte den Text `About` innerhalb Ihres ersten `li`-Elements enthalten.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

Ihr erstes Ankerelement sollte ein `href`-Attribut mit dem Wert `"#about"` haben.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Ihr zweites Ankerelement sollte den Text `Posts` innerhalb Ihres zweiten `li`-Elements enthalten.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

Ihr zweites Ankerelement sollte ein `href`-Attribut mit dem Wert `"#posts"` haben.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Ihr drittes Ankerelement sollte den Text `Contact` innerhalb Ihres dritten `li`-Elements enthalten.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

Ihr drittes Ankerelement sollte ein `href`-Attribut mit dem Wert `"#contact"` haben.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
