---
id: 669fd003cb89ee3c2402e041
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

In diesem einführenden Inhalt möchten Sie ein Bild von Mr. Whiskers mit einer Bildunterschrift anzeigen.

Fügen Sie unter dem `h1`-Element zunächst ein `figure`-Element hinzu.

Fügen Sie innerhalb des `figure`-Elements ein `img`-Element hinzu.

Das `src`-Attribut des `img` sollte den Wert `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` haben, und der `alt`-Text sollte den Wert `"a cat in the garden"` haben.

Fügen Sie unter Ihrem `img`-Element ein `figcaption` mit dem Text `Mr. Whiskers in the Garden` hinzu.

# --hints--

Ihr `figure`-Element sollte ein öffnendes Tag haben.

```js
assert.match(code, /<figure\>/);
```

Ihr `figure`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/figure\>/);
```

Ihr `figure`-Element sollte sich innerhalb Ihres `header`-Elements befinden.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

Sie sollten ein `img`-Element haben.

```js
assert.exists(document.querySelector("img"));
```

Ihr `img`-Element sollte sich innerhalb Ihres `figure`-Elements befinden.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

Ihr `img`-Element sollte ein `src`-Attribut mit dem Wert `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` haben.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Ihr `img`-Element sollte ein `alt`-Attribut mit dem Wert `"a cat in the garden"` haben.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

Ihr `figcaption`-Element sollte ein öffnendes Tag haben.

```js
assert.match(code, /<figcaption\>/);
```

Ihr `figcaption`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/figcaption\>/);
```

Ihr `figcaption`-Element sollte sich innerhalb Ihres `figure`-Elements befinden.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Ihr `figcaption`-Element sollte den Text `Mr. Whiskers in the Garden` haben.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

Ihr `figcaption` sollte unter dem `img`-Element sein. Sie haben sie in der falschen Reihenfolge.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
