---
id: 669fcb5106eeb13ab7afcd09
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie den Umgang mit semantischem HTML, indem Sie eine Blogseite erstellen, die Herrn Whiskers, der Katze, gewidmet ist.

Der erste Abschnitt, den Sie erstellen, ist die Kopfzeile der Seite.

Das `header`-Element wird verwendet, um einleitende Inhalte wie die Seitennavigation und andere einleitende Informationen darzustellen.

Hier ist ein Beispiel mit dem `header`-Element:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Fügen Sie innerhalb des `body`-Elements ein `header`-Element hinzu.

# --hints--

Sie sollten ein öffnendes `header`-Tag haben.

```js
assert.match(code, /<header>/i);
```

Sie sollten ein schließendes `header`-Tag haben.

```js
assert.match(code, /<\/header>/i);
```

Ihr `header`-Element sollte sich innerhalb des `body`-Elements befinden.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
