---
id: 669fde2081f65141ad703fe4
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Der erste Abschnitt auf der Seite wird der Über-mich-Abschnitt sein. Der Abschnitt stellt Mr. Whiskers vor und gibt den Nutzenden eine Vorstellung davon, worum es in diesem Blog geht.

Fügen Sie innerhalb Ihres `main`-Elements ein `section`-Element mit dem `id`-Attribut hinzu, das auf `"about"` gesetzt ist.

Fügen Sie innerhalb des `section`-Elements ein `h2`-Element mit dem Text `About` hinzu.

# --hints--

Ihr `section`-Element sollte ein öffnendes Tag haben.

```js
assert.match(code, /<section.*>/);
```

Ihr `section`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/section\>/);
```

Ihr `section`-Element sollte ein `id`-Attribut haben, das auf `"about"` gesetzt ist.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

Ihr `section`-Element sollte sich innerhalb Ihres `main`-Elements befinden.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Ihr `h2`-Element sollte ein öffnendes Tag haben.

```js
assert.match(code, /<h2\>/);
```

Ihr `h2`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/h2\>/);
```

Ihr `h2`-Element sollte den Text `About` enthalten. Überprüfen Sie Ihre Rechtschreibung.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

Ihr `h2`-Element sollte sich innerhalb Ihres `section`-Elements befinden.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
