---
id: 669fd6fd12918e3de87854d4
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Für Ihren Blog sollte es eine Möglichkeit geben, dass Nutzer zu verschiedenen Abschnitten auf der Seite navigieren können.

Das `nav`-Element wird verwendet, um Navigationslinks zu anderen Abschnitten im Dokument oder anderen Abschnitten auf der Website bereitzustellen. Sehr oft sehen Sie das `nav`-Element für Menüs oder Inhaltsverzeichnisse verwendet.

Hier ist ein Beispiel für die Verwendung des `nav`-Elements:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

Fügen Sie unter Ihrem `figure`-Element ein `nav`-Element mit einem verschachtelten `ul`-Element hinzu.

Fügen Sie innerhalb des `ul`-Elements drei `li`-Elemente hinzu.

# --hints--

Ihr `nav`-Element sollte ein öffnendes Tag haben.

```js
assert.match(code, /<nav\>/);
```

Ihr `nav`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/nav\>/);
```

Ihr `nav`-Element sollte innerhalb Ihres `header`-Elements sein.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

Ihr `ul`-Element sollte ein öffnendes Tag haben.

```js
assert.match(code, /<ul\>/);
```

Ihr `ul`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/ul\>/);
```

Ihr `ul`-Element sollte innerhalb Ihres `nav`-Elements sein.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

Ihr `ul`-Element sollte drei `li`-Elemente enthalten.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
