---
id: 66a49f685961e997e337cab1
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Das `header` ist dafür verantwortlich, den Haupttitel, das Bild und die Seiten-Navigation für den Blog anzuzeigen.

Fügen Sie innerhalb des `header`-Elements ein `h1` mit dem Text von `Welcome to Mr. Whiskers' Blog Page!` hinzu.

# --hints--

Sie sollten ein öffnendes `h1`-Tag haben.

```js
assert.match(code, /<h1>/i);
```

Sie sollten ein schließendes `h1`-Tag haben.

```js
assert.match(code, /<\/h1>/i);
```

Ihr `h1`-Element sollte den Text von `Welcome to Mr. Whiskers' Blog Page!` enthalten. Überprüfen Sie sorgfältig auf Rechtschreibfehler.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

Ihr `h1`-Element sollte innerhalb Ihres `header`-Elements verschachtelt sein.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
