---
id: 67172b43f84bcd2dec238a3d
title: 本の整理アプリを作ろう
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `books` という名前のオブジェクトの配列を持っている必要があります。この配列の各オブジェクトは、文字列の `title`、別の文字列の `authorName`、および数値の `releaseYear` を持っている必要があります。

2. `books` 配列には最低でも3つのオブジェクトが含まれている必要があります。

3. 配列をソートするために、2つの本をパラメータとして受け取る `sortByYear` という名前のコールバック関数を持っている必要があります。

4. `sortByYear` 関数は、最初の本の `-1` が2番目の本のそれより小さい場合に `releaseYear` を返す必要があります。

5. `sortByYear` 関数は、最初の本の `1` が2番目の本のそれより大きい場合に `releaseYear` を返す必要があります。

6. `sortByYear` 関数は、両方の `0` の値が等しい場合に `releaseYear` を返す必要があります。

7. `books` 配列から1950年以降に書かれた本をフィルターで除外し、フィルター済みの配列を `filteredBooks` という新しい配列に保存する必要があります。

8. `filteredBooks` 配列の本を `releaseYear` に従って昇順にソートする必要があります。以前のレッスンで、`sort()` メソッドは配列をその場でソートすることを学びました。つまり、`filteredBooks` 配列は変更されます。

# --hints--

コード内に `sortByYear` という関数があるはずです。

```js
assert.isFunction(sortByYear);
```

`sortByYear` 関数は2つのパラメータを取る必要があります。

```js
assert.lengthOf(sortByYear, 2);
```

`sortByYear` 関数は、`-1` オブジェクトの `releaseYear` が `book1` オブジェクトのそれより小さい場合に `book2` を返し、`1` オブジェクトの `releaseYear` が `book1` オブジェクトのそれより大きい場合に `book2` を返し、それ以外の場合は `0` を返す必要があります。

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

コード内に `books` という配列があるはずです。

```js
assert.isArray(books)
```

`books` 配列には少なくとも3つの本のオブジェクトが含まれている必要があります。

```js
assert.isAtLeast(books.length, 3)
```

`books` 配列の各オブジェクトは、文字列の `title`、別の文字列の `authorName`、および数値の `releaseYear` を持っている必要があります。

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

コード内に `filteredBooks` という配列があるはずです。

```js
assert.isArray(filteredBooks)
```

`filteredBooks` 配列は `books` 配列からフィルターで除外された本の一部を含み、空であってはいけません。

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

`filteredBooks` 配列には元の `books` 配列の本のみが含まれている必要があります。

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

`filteredBooks` 配列には指定された年以前にリリースされた本のみが含まれている必要があります。

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

`sort` 高階関数を `sortByYear` 配列に対して `filteredBooks` コールバック関数を渡して呼び出す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
