---
id: 60eebd07ea685b0e777b5583
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

مرحبًا بك في مشروع نموذج التسجيل! ابدأ بإضافة الإعلان `!DOCTYPE html` في أعلى الوثيقة ليعرف المتصفح نوع الوثيقة التي يقرأها.

# --hints--

يجب أن تحتوي شفرتك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب تضمين مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب تحديد نوع الوثيقة لتكون `html`.

```js
assert(code.match(/html/gi));
```

يجب إغلاق الإعلان `DOCTYPE` بعلامة `>` بعد النوع.

```js
assert(code.match(/html\s*>/gi));
```

يجب أن يكون إعلان `DOCTYPE` في بداية شفرة HTML الخاصة بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
