---
id: 60f02e7361b68405e27b62a5
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

ضمن الـ `body`، أضف عنوانًا للمحتوى بإضافة `h1` يحتوي على النص `Registration Form`.

# --hints--

يجب إضافة الـ `h1` داخل الـ `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

يجب أن يحتوي الـ `h1` على النص `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
