---
id: 60f0286404aefb0562a4fdf9
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

أضف عنصر `title` وعنصر `meta` داخل عنصر `head`. امنح مشروعك عنوانًا هو `Registration Form`، وأضف الخاصية `charset` بقيمة `utf-8` إلى عنصر `meta` الخاص بك.

# --hints--

يجب أن يحتوي شفرتك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون عنصر `title` داخل عنصر `head`.

```js
assert.exists(document.querySelector('head > title'));
```

يجب أن يكون لمشروعك عنوان `Registration Form`. تذكر أن حالة الأحرف والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

يجب أن تنشئ عنصر `meta` داخل عنصر `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

يجب أن تعطي عنصر `meta` الخاصية `charset` بقيمة `utf-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

يجب أن يكون عنصر `meta` عنصرًا فارغًا، فلا يحتوي على وسم إغلاق `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
