---
id: 60f027099a15b00485563dd2
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أسفل `DOCTYPE`، أضف عنصر `html` مع خاصية `lang` بقيمة `en`، لتجد مكانًا تبدأ فيه بوضع بعض الشفرة.

# --hints--

يجب أن يكون إعلان `DOCTYPE` في بداية ملف HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يحتوي عنصر `html` على وسم بداية، وخاصية `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم الإغلاق.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن تكون وسوم `html` بالترتيب الصحيح.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

يجب أن يكون لديك عنصر `html` واحد فقط.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
