---
id: 60f1922fcbd2410527b3bd89
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--
وحدة `vh` تعني ارتفاع إطار العرض، وتساوي 1% من `height` لإطار العرض. هذا يجعلها نسبة إلى ارتفاع إطار العرض.

حان الوقت لتحسين المشروع باستخدام CSS. ابدأ بإعطاء العنصر `body` عرضًا `width` بقيمة `100%` وارتفاعًا `height` بقيمة `100vh`.

# --hints--

يجب استخدام محدِّد عنصر `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

يجب إعطاء العنصر `body` عرضًا `width` بقيمة `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

يجب إعطاء العنصر `body` ارتفاعًا `height` بقيمة `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
