---
id: 60f8618d191b940d62038513
title: الخطوة 27
challengeType: 0
dashedName: step-27
---

# --description--

تريد أن يكون من الممكن اختيار مُدخَل راديو واحد فقط في كل مرة. مع ذلك، النموذج لا يعرف أن مُدخَلات الراديو مرتبطة ببعضها.

لربط مُدخَلات الراديو، أعطها نفس خاصية `name` بقيمة `account-type`. الآن، لا يمكن اختيار كلا مُدخَلي الراديو في نفس الوقت.

# --hints--

يجب أن تعطي مُدخَل الراديو الأول خاصية `name` بقيمة `account-type`.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(1) input[type="radio"]')?.name, 'account-type');
```

يجب أن تعطي مُدخَل الراديو الثاني خاصية `name` بقيمة `account-type`.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(2) input[type="radio"]')?.name, 'account-type');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label><input type="radio" /> Personal</label>
        <label><input type="radio" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
