---
id: 60fabf0dd4959805dbae09e6
title: الخطوة 33
challengeType: 0
dashedName: step-33
---

# --description--

أضف `label` آخر بعد الأول، مع النص `Input your age (years): `. ثم، قم بتسكين `input` مع `type` من نوع `number`.

بعد ذلك، أضف خاصية `min` إلى الـ `input` بقيمة `13` لأن المستخدمين الذين تقل أعمارهم عن 13 سنة لا يجب أن يسجلوا. كما أن المستخدمين على الأرجح لن يكونوا فوق عمر 120؛ أضف خاصية `max` بقيمة `120`.

الآن، إذا حاول أحدهم إرسال النموذج بقيم خارج النطاق، ستظهر تحذيرات ولن يُرسل النموذج. جرّب ذلك.

# --hints--

يجب أن تضيف عنصر `label` إلى عنصر `fieldset` الثالث، بعد عنصر `label` الموجود.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

يجب أن تعطي الـ `label` النص `Input your age (years): `.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

يجب أن يحتوي الـ `label` على عنصر `input` يكون له `type` بقيمة `number`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

يجب أن تُعطِي الـ `input` خاصية `min` بقيمة `13`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

يجب أن تُعطِي الـ `input` خاصية `max` بقيمة `120`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        
      </fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
      </label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
