---
id: 60fab8367d35de04e5cb7929
title: الخطوة 31
challengeType: 0
dashedName: step-31
---

# --description--

أضف النص `I accept the terms and conditions` مباشرة بعد عنصر الـ `input` في الـ `label` المضاف حديثًا. ثم اربط النص `terms and conditions` بالموقع التالي:

```md
https://www.freecodecamp.org/news/terms-of-service/
```

# --hints--

يجب إضافة نص `I accept the terms and conditions` إلى الـ `label` التي تلي مجموعة الحقول الثالثة `fieldset`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label')?.innerText.trim(), 'I accept the terms and conditions');
```

يجب استخدام عنصر `a` لربط نص الشروط والأحكام.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label a'));
```

ضع النص الجديد مباشرة بعد عنصر الـ `input` في الـ `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input + a'));
```

يجب إعطاء عنصر الـ `a` الخاصية `href` بالقيمة `https://www.freecodecamp.org/news/terms-of-service/`.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + label > input + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

يجب تغليف عنصر الـ `a` فقط حول النص `terms and conditions`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent.trim(), 'terms and conditions');
```

النص داخل عنصر وسم الرابط يحتوي على فراغات زائدة في البداية أو النهاية. يجب أن تكون الفراغات الوحيدة في نص وسم الرابط بين الكلمات `terms` و`and` و`conditions`.

```js
const nestedAnchor = document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent;
const innerContent = nestedAnchor?.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
