---
id: 62cc5b1779e4d313466f73c5
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

قم بتسكين عنصر `link` داخل عنصر `head`. أضف له خاصية `rel` بقيمة `stylesheet` وخاصية `href` بقيمة `styles.css`.

# --hints--

يجب أن يحتوي شفرتك على عنصر `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

لا تغير وسم `head` الموجود لديك. تأكد من عدم حذف وسم الإغلاق.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

يجب أن يكون لديك عنصر `link` واحد.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

يجب أن يكون عنصر `link` عنصرًا فارغًا، فلا يجب أن يحتوي على وسم إغلاق `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
