---
id: 63541ef4f96cd82e8e6c788a
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

تحدد الخاصية `method` كيفية إرسال بيانات النموذج إلى عنوان URL المحدد في الخاصية `action`. يمكن إرسال بيانات النموذج عبر طلب `GET` كمعلمات في عنوان URL (باستخدام `method="get"`) أو عبر طلب `POST` كبيانات في جسم الطلب (باستخدام `method="post"`).

عيّن الخاصية `method` لإرسال بيانات النموذج عبر طلب `POST`.

# --hints--

لا تضف عنصر `form` جديدًا.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

يجب أن يحتوي عنصر `form` على خاصية `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

يجب تعيين الخاصية `method` إلى `post`.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
