---
id: a5deed1811a43193f9f1c841
title: implementare un saltatore di elementi
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

In questo laboratorio creerai una funzione che salta elementi in un array finché non trova uno accettabile basato su una specifica funzione di test.

Per esempio, per un array come `[1, 1, 1, 2, 1, 1, 1]` e una funzione di test `function(n) {return n === 2}`, il primo elemento accettabile è quello all'indice 3, quindi tutti gli elementi precedenti devono essere scartati e l'output dovrebbe essere gli elementi rimanenti `[2, 1, 1, 1]`.

**Obiettivo**: soddisfa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente**

1. Dovresti avere una funzione `dropElements` che accetta un array (`arr`) e una funzione (`func`) come argomenti.
1. La funzione `dropElements` dovrebbe iterare sull'array e rimuovere gli elementi a partire dal primo finché `func` non restituisce `true` per un elemento.
1. La funzione `dropElements` dovrebbe restituire gli elementi rimanenti nell'array se la condizione è soddisfatta.
1. Se la condizione non viene mai soddisfatta, dovrebbe restituire un array vuoto.

# --hints--

Dovresti avere una funzione `dropElements`.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` dovrebbe restituire `[3, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` dovrebbe restituire `[1, 0, 1]`.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` dovrebbe restituire `[1, 2, 3]`.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` dovrebbe restituire `[]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` dovrebbe restituire `[7, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` dovrebbe restituire `[3, 9, 2]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
