---
id: 66e2eab8a5638f57b637b7cc
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер ви маєте порахувати кількість розділових знаків.

Створіть функцію `getPunctuationCount` з параметром `sentence`.

Всередині функції створіть цикл, щоб порахувати кількість розділових знаків у `sentence`, який буде переданий у функцію під час її виклику. Для наших цілей розділовим знаком вважається будь-який символ, який не є пробілом (`" "`) або літерою.

Ваша функція `getPunctuationCount` має повертати число.

# --hints--

Ви маєте створити функцію `getPunctuationCount`.

```js
assert.isFunction(getPunctuationCount);
```

Ваша функція `getPunctuationCount` має мати параметр `sentence`.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

Ваша функція `getPunctuationCount` має повертати число.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

Коли речення є `"What's going on here?"`, функція `getPunctuationCount` має повертати `2`.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

Коли речення є `"What????!"`, функція `getPunctuationCount` має повертати `5`.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

Ваша функція `getPunctuationCount` має повертати правильну кількість розділових знаків для будь-якого речення.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
