---
id: 66e2f06b191e305c00574e4d
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Нарешті, порахуйте кількість слів, створивши функцію `getWordCount` з параметром `sentence`. Функція має повертати загальну кількість слів у рядку, який передається під час виклику.

# --hints--

Ви маєте створити функцію `getWordCount`.

```js
assert.isFunction(getWordCount);
```

Ваша функція `getWordCount` має мати параметр `sentence`.

```js
assert.match(getWordCount.toString(), /sentence/);
```

Ваша функція `getWordCount` має повертати число.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

Коли рядок `"When are you gonna start learning to code?"`, функція `getWordCount` має повертати `8`.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

Коли рядок `"What's going on?"`, функція `getWordCount` має повертати `3`.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

Підрахунок слів має бути нечутливим до регістру.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

Ваша функція `getWordCount` має повертати правильну кількість слів для будь-якого рядка.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both front-end and back-end development"), 7);
```

Ваша функція `getWordCount` має повертати правильну кількість слів для порожнього рядка або рядка, що містить лише пробіли.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
