---
id: 66e2edb23013ca596f8d9568
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Тепер настав час протестувати вашу функцію `getPunctuationCount`.

Створіть змінну `punctuationCount` і призначте їй результат виклику функції `getPunctuationCount` з аргументом `"WHAT?!?!?!?!?"`.

Після цього виведіть у консоль наступне: `"Punctuation Count: [Punctuation count goes here]"`. Замініть `[Punctuation count goes here]` на фактичну назву змінної. Тут ви можете обрати використання шаблонних рядків або конкатенації рядків з оператором `+`.

# --hints--

Ви повинні створити змінну `punctuationCount`.

```js
assert.isNotNull(punctuationCount)
```

Ваша змінна `punctuationCount` має бути призначена результатом `getPunctuationCount("WHAT?!?!?!?!?")`.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

Ви повинні вивести у консоль наступне: `"Punctuation Count: [Punctuation count goes here]"`. Замініть `[Punctuation count goes here]` на фактичну назву змінної. Переконайтеся, що використовуєте правильний синтаксис конкатенації рядків.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
