---
id: 66e2d680e129e1423116a541
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите аналізатор речень, який прийматиме речення і визначатиме кількість слів, голосних, приголосних та розділових знаків.

Для початку створіть функцію `getVowelCount` з параметром `sentence`. Ваша функція має повертати загальну кількість голосних у реченні.

# --hints--

Ви повинні створити функцію `getVowelCount`.

```js
assert.isFunction(getVowelCount);
```

Ваша функція `getVowelCount` має мати параметр `sentence`.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

Ваша функція `getVowelCount` має повертати число.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

Коли речення `"Apples are tasty fruits"`, функція `getVowelCount` має повертати `7`.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

Коли речення `"Hello, World!"`, функція має повертати `3`.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

Підрахунок голосних має бути нечутливим до регістру.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

Ваша функція `getVowelCount` має повертати правильну кількість голосних для будь-якого речення.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
