---
id: 66e2e5d890ac8a54c9d703bd
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Настав час порахувати приголосні. Створіть `getConsonantCount` функцію з параметром `sentence`.

Всередині функції використайте цикл, щоб порахувати кількість приголосних у `sentence`, який буде переданий у функцію під час її виклику. Приголосна — це будь-яка літера, яка не є одним із наступних символів: `"aeiou"`.

Ваша `getConsonantCount` функція має повертати число.

# --hints--

Ви повинні створити `getConsonantCount` функцію.

```js
assert.isFunction(getConsonantCount);
```

Ваша `getConsonantCount` функція має мати параметр `sentence`.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

Ваша `getConsonantCount` функція має повертати число.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

Коли речення є `"Coding is fun"`, функція `getConsonantCount` має повертати `7`.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

Коли речення є `"Hello, World!"`, функція `getConsonantCount` має повертати `7`.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

Підрахунок приголосних має бути нечутливим до регістру.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

Ваша `getConsonantCount` функція має повертати правильну кількість приголосних для будь-якого речення.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
