---
id: 66e2df8c5676fe51cead8ec8
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Тепер настав час протестувати вашу функцію `getVowelCount`.

Створіть змінну `vowelCount` і призначте їй результат виклику функції `getVowelCount` з аргументом `"Apples are tasty fruits"`.

Після цього виведіть у консоль наступне: `"Vowel Count: [vowel count goes here]"`. Замініть `[vowel count goes here]` на фактичну назву змінної. Ви можете використати шаблонні рядки або конкатенацію рядків з оператором `+`.

# --hints--

Ви повинні створити змінну `vowelCount`.

```js
assert.isNotNull(vowelCount)
```

Ваша змінна `vowelCount` має бути призначена результатом `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

Ви повинні вивести змінну `vowelCount` у консоль з параметром `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
