---
id: 66e2e99caa0cb3570f4822b4
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер настав час протестувати вашу функцію `getConsonantCount`.

Створіть змінну `consonantCount` і призначте їй результат виклику функції `getConsonantCount` з аргументом `"Coding is fun"`.

Після цього виведіть у консоль наступне: `"Consonant Count: [Consonant count goes here]"`. Замініть `[Consonant count goes here]` на фактичну назву змінної. Ви можете використати шаблонні рядки або конкатенацію рядків з оператором `+`.

# --hints--

Ви повинні створити змінну `consonantCount`.

```js
assert.isNotNull(consonantCount)
```

Ваша змінна `consonantCount` має бути призначена результатом `getConsonantCount("Coding is fun")`.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

Ви повинні вивести змінну `consonantCount` у консоль з параметром `"Consonant Count: ${consonantCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
