---
id: acda2fb1324d9b0fa741e6b5
title: Ein Tool zum Überprüfen des Endes eines Strings erstellen
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

In diesem Labor implementieren Sie eine Funktion, die überprüft, ob ein String mit dem angegebenen Zielstring endet.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und lassen Sie alle Tests erfolgreich durchlaufen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `confirmEnding` erstellen, die zwei Parameter entgegennimmt: den zu überprüfenden String und den String, mit dem verglichen werden soll.
2. Die Funktion sollte `true` zurückgeben, wenn der erste String mit dem zweiten String endet, und `false` sonst.
3. Sie sollten die `.endsWith()`-Methode nicht verwenden; stattdessen verwenden Sie eine der JavaScript-Substring-Methoden, um dies zu erreichen.

# --hints--

Sie sollten eine Funktion namens `confirmEnding` erstellen.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` sollte 2 Parameter entgegennehmen.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` sollte `false` zurückgeben.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` sollte `false` zurückgeben.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` sollte `false` zurückgeben.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` sollte `false` zurückgeben.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` sollte `false` zurückgeben.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Ihr Code sollte die eingebaute Methode `.endsWith()` nicht verwenden, um das Labor zu lösen.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
