---
id: 683ec95f2f8781355556ff81
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Você precisa testar as coisas para entender o fluxo do algoritmo nesta fase inicial.

Para fazer isso, você precisa primeiro sair do loop.Isso ocorre porque a implementação atual permitirá apenas uma iteração, então se a condição não for atendida, haverá um loop infinito.

Logo após a declaração `if`, use a palavra-chave `break` para sair do loop `while`.Então, após o loop `while`, retorne uma lista vazia para indicar que o valor não foi encontrado.

# --hints--

Você deve usar a palavra-chave `break` para sair do loop `while` após a instrução `if`.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

Você deve retornar uma lista vazia após o seu loop `while`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```
