---
id: 683ec95f2f8781355556ff7a
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

O algoritmo de busca binária que você está construindo não vai apenas retornar `True` ou `False` para indicar se o valor foi encontrado ou não. Ele também retornará o caminho que ele percorre para encontrar esse valor. E para isso, você terá uma variável `path_to_target` para rastrear o caminho.

Dentro da sua função `binary_search`, substitua a palavra-chave `pass` por uma variável `path_to_target` e inicialize-a como uma lista vazia, para começar.

# --hints--

Você deve criar uma variável `path_to_target` definida como uma lista vazia dentro da sua função `binary_search`. Não esqueça de remover a palavra-chave `pass`.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```
