---
id: 683ec39565dbc7e3375263b1
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Em uma lição anterior, você aprendeu sobre `binary search` e como ele difere do `linear search`.

Neste workshop, você colocará isso em prática construindo um algoritmo de busca binária.

Para começar, crie uma função `binary_search` com os parâmetros `search_list` e `value`.

`search_list` é a lista de elementos dentro da qual a função irá buscar, e `value` é o alvo que a função tentará encontrar dentro do `search_list`.

Adicione a palavra-chave `pass` dentro da função por enquanto para que o teste possa passar.

# --hints--

Você deve criar uma função `binary_search`.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

Sua função `binary_search` deve ter os parâmetros `search_list` e `value`.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
