---
id: 683ec95f2f8781355556ff82
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Lembre-se de que a função recebe uma lista chamada `search_list` e o `value` que você está procurando.

Agora, chame a função com `binary_search([1, 2, 3, 4, 5], 3)` e `binary_search([1, 2, 3, 4, 5, 9], 4)` e imprima as chamadas imediatamente.

# --hints--

Você deve chamar a função com `binary_search([1, 2, 3, 4, 5], 3)` e imprimir imediatamente.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

Você deve chamar a função com `binary_search([1, 2, 3, 4, 5, 9], 4)` e imprimir imediatamente.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        break
    return []
          
--fcc-editable-region--

--fcc-editable-region--
```
