---
id: 683ec95f2f8781355556ff7d
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Para começar com o loop, você precisa encontrar o índice do meio do espaço de busca atual.

Você pode obter o índice do meio calculando a média das variáveis `low` e `high`.Você terá que usar divisão inteira para obter a média após somar os dois valores, então a resposta será sempre um número inteiro arredondado para baixo, pois os índices devem ser inteiros.

Remova a palavra-chave `pass` do seu loop `while` e crie uma variável `mid` definida como a média de `low` e `high`.

# --hints--

Você deve remover a palavra-chave `pass` do loop while e criar uma variável `mid` definida como `(low + high) // 2`.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```
