---
id: 683ec95f2f8781355556ff7e
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Agora que você calculou o índice `mid`, você precisa realmente recuperar o valor da `search_list` nesse índice.Este valor é o que você irá comparar com seu parâmetro `value`, que é o alvo que você está procurando.

Crie uma variável `value_at_middle` e atribua a ela o elemento de `search_list` que está localizado no índice `mid`. 

# --hints--

Você deve criar uma variável `value_at_middle` definida como `search_list[mid]`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2

        
--fcc-editable-region--
```
