---
id: 683ec95f2f8781355556ff84
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Se a condição no `elif` for verdadeira, então remova a palavra-chave `pass` e atualize o valor da variável `low` adicionando `1` à variável `mid`.

Isso vai estender a busca para a metade direita das áreas de busca atuais na lista, porque se o `value` for maior que o `value_at_middle`, isso significa que o `value` deve estar na metade direita da área de busca atual.

# --hints--

Você deve remover a palavra-chave `pass` e atualizar a variável `low` para `mid + 1`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid + 1")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        elif value > value_at_middle:
            pass
--fcc-editable-region--
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
```
