---
id: 683ec95f2f8781355556ff7b
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Na busca binária, existe o menor índice possível e o maior índice possível.

O índice mais baixo possível representa o limite mais à esquerda na porção atual dos valores que estão sendo pesquisados, neste caso, `search_list`. O índice mais alto possível, por outro lado, representa o limite mais à direita nos valores que estão sendo pesquisados.

Para levar em conta esses dois, você precisa considerar toda a lista que está sendo pesquisada. Então, dentro da função `binary_search`, defina uma variável `low` com o valor `0` e uma variável `high` com um valor que corresponda ao último índice na lista que está sendo pesquisada.

# --hints--

Você deve ter uma variável chamada `low` na sua função `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

Sua variável `low` deve ser configurada para `0`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

Você deve ter uma variável chamada `high` na sua função `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

Sua variável `high` deve ser definida como `len(search_list) - 1`.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []
    
--fcc-editable-region--
```
