---
id: 683ec95f2f8781355556ff83
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Você pode ver que a primeira chamada retorna `[3]`.Isso porque `3` é o valor do meio em `[1, 2, 3, 4, 5]`.

Por outro lado, a segunda chamada retorna uma lista vazia.Isso acontece porque, após verificar o valor do meio, o loop atualmente não possui a lógica para restringir o intervalo de busca atualizando as variáveis `low` ou `high`.Ele atua apenas como uma verificação única de ponto médio e então termina sem continuar a busca.

Para permitir que a busca binária continue restringindo sua busca, adicione um bloco `elif` que verifica se `value` é maior que `value_at_middle`.Adicione a palavra-chave `pass` dentro do `elif` por enquanto.

# --hints--

Você deve adicionar um bloco `elif` para quando `value` for maior que `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--          
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
