---
id: 683ec95f2f8781355556ff7c
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora que você definiu os limites da sua busca, é hora de criar o loop que executará o algoritmo de busca binária.

A busca binária funciona reduzindo repetidamente o espaço de busca. Esse processo continua enquanto houver um intervalo válido de elementos para verificar.

Você pode expressar isso com um loop `while` que continuará enquanto seu ponteiro `low` for menor ou igual ao seu ponteiro `high`.

Se `low` alguma vez se tornar maior que `high`, isso significa que o espaço de busca ficou vazio e o `value` não está na lista.

Dentro da função `binary_search`, crie um loop `while` com uma condição que verifica se `low` é menor ou igual a high. Dentro do loop, adicione a palavra-chave `pass` como um espaço reservado por enquanto.

# --hints--

Você deve criar um loop `while` que execute enquanto `low` for menor ou igual a `high`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--
    
--fcc-editable-region--
```
