---
id: 683ec95f2f8781355556ff80
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

A próxima coisa a fazer é criar uma condição que verificará se o `value` alvo está no meio. 

Dentro do loop `while`, crie uma declaração `if` que verifica se o `value` alvo é igual a `value_at_middle`. Se for, retorne a variável `path_to_target`.

# --hints--

Você deve criar uma declaração `if` que verifique se o `value` é igual a `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

Você deve retornar a variável `path_to_target` dentro da declaração `if`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
--fcc-editable-region--
        
--fcc-editable-region--
```
