---
id: 684169b00a073c9b9f0a703c
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Você percebeu que a segunda chamada da função agora tem os valores `3, 5, 4` na lista?Isso indica que a busca binária está funcionando conforme o esperado.

Foi assim que aconteceu: o algoritmo primeiro verificou `3` como o meio da lista inicial.Como `4` é maior que `3`, a busca mudou para a metade direita.Em seguida, examinou `5` como o novo meio.Porque `4` é menor que `5`, a busca se moveu para a esquerda, identificando finalmente `4` como o meio do intervalo final.

Para testar a função novamente, chame-a com `[1, 3, 5, 9, 14, 22], 10` e imprima a chamada imediatamente.Esta é uma situação em que o valor não será encontrado.

# --hints--

Você deve chamar a função com `binary_search([1, 3, 5, 9, 14, 22], 10)` e imprimir imediatamente.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []
    
print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
