---
id: 5688e62ea601b2482ff8422b
title: Побудова пошуку профілю
challengeType:  26
dashedName: lab-profile-lookup
---

# --description--

У цій лабораторній роботі ви створите пошук профілю, який шукає інформацію про людей у списку контактів.

Для цього прикладу уявіть, що є контакт на ім’я Akira Laine, `lookUpProfile("Akira", "lastName")` має повернути `Laine`.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви маєте створити функцію з назвою `lookUpProfile`, яка приймає ім’я та властивість як аргументи.
2. Ви маєте отримувати інформацію про контакт із наданого масиву `contacts`.
3. Якщо функція отримує ім’я контакту і властивість існує у відповідного контакту, тоді має повертатися значення цієї властивості.
4. Якщо ім’я, передане функції, не співпадає з жодним контактом у масиві контактів, функція має повертати `"No such contact"`.
5. Якщо властивість не існує у знайденого контакту, функція має повертати `"No such property"`.

# --before-each--

```js
const _contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

contacts = structuredClone(_contacts); 
```

# --hints--

Масив `contacts` має залишатися у глобальній області видимості. Скиньте урок, щоб відновити його, якщо ви його видалили.

```js
assert.isArray(contacts);
```

Ви маєте створити функцію з назвою `lookUpProfile`.

```js
assert.isFunction(lookUpProfile);
```

`lookUpProfile("Kristian", "lastName")` має повертати рядок.

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

`lookUpProfile("Kristian", "lastName")` має повертати рядок `Vos`.

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

`lookUpProfile("Sherlock", "likes")` має повертати `["Intriguing Cases", "Violin"]`.

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` має повертати масив.

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

`lookUpProfile("Bob", "number")` має повертати рядок `No such contact`.

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

`lookUpProfile("Bob", "potato")` має повертати рядок `No such contact`.

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

`lookUpProfile("Akira", "address")` має повертати рядок `No such property`.

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

Масив `contacts` має залишатися незмінним після запуску `lookUpProfile`.

```js
lookUpProfile("Sherlock", "likes");
assert.deepEqual(contacts,_contacts);
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
