---
id: 6976db7ecfa770bf21307b20
title: costruire un motore di remix per playlist
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

In questo laboratorio costruirai un programma che crea una singola playlist remix da più playlist inviate dagli ascoltatori.

Ogni ascoltatore fornisce una lista di canzoni che vuole sentire. Alcune canzoni possono apparire più di una volta e alcuni artisti possono comparire troppe volte. Il tuo compito è lavorare su queste playlist passo dopo passo: combinarle in un'unica lista, assegnare un punteggio a ogni canzone, rimuovere le canzoni duplicate, limitare quante volte lo stesso artista appare e infine creare un ordine di riproduzione finale.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione chiamata `flattenPlaylists` che accetta un array di playlist dove ogni playlist è un array di oggetti con le seguenti proprietà: `trackId`, `artist`, `title`, `votes`, `bpm`. Se l'input non è un array, `flattenPlaylists` dovrebbe restituire un array vuoto. Ti è stato fornito un esempio di playlist che puoi usare per testare la tua funzione.

2. `flattenPlaylists` dovrebbe restituire un array piatto di oggetti traccia, dove ogni oggetto include tutte le proprietà originali della traccia più una proprietà `source` impostata su un array con l'indice della playlist e l'indice della traccia che indicano da dove proviene la traccia.

3. Dovresti creare una funzione chiamata `scoreTracks` che accetta un array di oggetti traccia come restituiti da `flattenPlaylists` (ognuno con le proprietà `trackId`, `artist`, `title`, `votes`, `bpm` e `source`) e restituisce un nuovo array di oggetti traccia, ognuno con una proprietà `score` aggiunta usando la formula: `votes * 10 - Math.abs(bpm - 120)`.

4. Dovresti creare una funzione chiamata `dedupeTracks` che accetta un array di oggetti traccia come restituiti da `scoreTracks` e restituisce un nuovo array con le voci duplicate di `trackId` rimosse, mantenendo solo la prima occorrenza di ciascuna.

5. Dovresti creare una funzione chiamata `enforceArtistQuota` che accetta un array di oggetti traccia come restituiti da `dedupeTracks` e un numero che rappresenta il massimo numero consentito di occorrenze per artista. La funzione dovrebbe restituire un nuovo array in cui nessun artista appare più volte del numero dato, mantenendo le occorrenze più precoci.

6. Dovresti creare una funzione chiamata `buildSchedule` che accetta un array di oggetti traccia come restituiti da `enforceArtistQuota` e restituisce un nuovo array di oggetti `{ slot, trackId }`, dove `slot` è un indice basato su 1 che rappresenta la posizione di ogni traccia nell'ordine di trasmissione.

7. Dovresti creare una funzione chiamata `remixPlaylist` che accetta un array di playlist e il numero massimo di occorrenze consentite per artista. La funzione dovrebbe restituire il programma finale di trasmissione come un array di oggetti `{ slot, trackId }`, chiamando `flattenPlaylists`, `scoreTracks`, `dedupeTracks`, `enforceArtistQuota` e `buildSchedule` in ordine.

# --hints--

Dovresti avere una funzione chiamata `flattenPlaylists`.

```js
assert.isFunction(flattenPlaylists);
```

Dovresti restituire un array vuoto da `flattenPlaylists` quando l'input non è un array.

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

Ogni traccia restituita da `flattenPlaylists` dovrebbe includere un campo `source` che è un array contenente l'indice della playlist e l'indice della traccia.

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

Dovresti avere una funzione chiamata `scoreTracks`.

```js
assert.isFunction(scoreTracks)
```

Ogni traccia restituita da `scoreTracks` dovrebbe includere un campo numerico `score`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

Dovresti calcolare `score` usando un BPM obiettivo di `120` e questa formula: `votes * 10 - Math.abs(bpm - 120)`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

Dovresti avere una funzione chiamata `dedupeTracks`.

```js
assert.isFunction(dedupeTracks)
```

Quando esistono valori duplicati di `trackId`, `dedupeTracks` dovrebbe mantenere solo la prima occorrenza della traccia.

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

Dovresti avere una funzione chiamata `enforceArtistQuota`.

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` dovrebbe assicurarsi che nessun artista appaia più di `maxPerArtist` volte rimuovendo le tracce extra mantenendo le più precoci.

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

Dovresti avere una funzione chiamata `buildSchedule`.

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` dovrebbe restituire un array di oggetti con la forma `{ slot, trackId }`, dove `slot` inizia da `1`.

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

Dovresti avere una funzione chiamata `remixPlaylist`.

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` dovrebbe chiamare le funzioni di supporto in ordine per produrre il programma finale.

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
