---
id: 69899123889dcf5b19be1847
title: 深さ優先探索アルゴリズムの実装
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

このラボでは、<dfn>深さ優先探索</dfn>（DFS）と呼ばれるグラフトラバーサルアルゴリズムを実装します。

<dfn>幅優先探索</dfn>（BFS）が隣接ノードをレベルごとに探索するのに対し、<dfn>深さ優先探索</dfn>は一つの経路をできるだけ深く進み、行き止まりに達したらバックトラックします。

アルゴリズムが経路の終端（未訪問の隣接ノードがないノード）に到達すると、未探索の辺が残っている直近のノードまで戻り、そこから探索を続けます。

このアルゴリズムを反復的に実装するには、**スタック**を使うのが適しています。JavaScriptでは、標準の`Array`が`.push()`と`.pop()`メソッドを使ってスタックとして完璧に機能し、**後入れ先出し（LIFO）**の原則に従います。これにより、スタックに最後に追加された隣接ノードが次に探索されます。

このアルゴリズムの単純な出力は、あるノードから到達可能なノードのリストです。したがって、訪問したノードを追跡する必要もあります。

**目的**：以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `dfs`という名前の関数を定義してください。
1. `dfs`関数は2つの引数を取る必要があります：

    - `graph`：隣接行列（配列の配列）。
    - `root`：数値のノードラベル（開始インデックス）で、`0`と`n - 1`の間のノードの数値です。ここで`n`はグラフ内のノードの総数です。

1. 関数は開始ノードから到達可能なすべてのノードラベルを含む配列を返す必要があります。

# --hints--

`dfs`という名前の関数が2つの引数を取る必要があります。

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)`は`1`、`2`、`3`、`0`を含むリストを返す必要があります。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)`は`1`、`2`、`3`、`0`を含むリストを返す必要があります。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)`は`[3]`を返す必要があります。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)`は`3`と`2`を含むリストを返す必要があります。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)`は`0`と`1`を含むリストを返す必要があります。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
