---
id: 587d78af367417b2b2512b03
title: بناء نموذج استبيان
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301145
dashedName: build-a-survey-form
---

# --description--

**الهدف:** بناء تطبيق يشبه وظيفيًا <a href="https://survey-form.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://survey-form.freecodecamp.rocks</a>. **لا تنسخ هذا المشروع التجريبي**.

**قصص المستخدم:**

1. يجب أن يكون لديك عنوان صفحة في عنصر `h1` يحمل `id` بقيمة `title`
1. يجب أن يكون لديك شرح قصير في عنصر `p` يحمل `id` بقيمة `description`
1. يجب أن يكون لديك عنصر `form` يحمل `id` بقيمة `survey-form`
1. داخل عنصر النموذج، يجب عليك **إدخال اسمك** في حقل `input` يحمل `id` بقيمة `name` و`type` بقيمة `text`
1. داخل عنصر النموذج، يجب عليك **إدخال بريدك الإلكتروني** في حقل `input` يحمل `id` بقيمة `email`
1. إذا أدخلت بريدًا إلكترونيًا غير منسق بشكل صحيح، سترى خطأ تحقق HTML5
1. داخل النموذج، يمكنك إدخال رقم في حقل `input` يحمل `id` بقيمة `number`
1. يجب ألا يقبل حقل الرقم غير الأرقام، إما بمنعك من كتابتها أو بإظهار خطأ تحقق HTML5 (حسب المتصفح).
1. إذا أدخلت أرقامًا خارج نطاق حقل الرقم، المحدد بواسطة خاصيتي `min` و`max`، سترى خطأ تحقق HTML5
1. لحقول الاسم والبريد الإلكتروني والرقم، يجب أن ترى عناصر `label` مقابلة في النموذج، تصف غرض كل حقل بالمعرفات التالية: `id="name-label"`، `id="email-label"`، و`id="number-label"`
1. لحقول الاسم والبريد الإلكتروني والرقم، يجب أن ترى نص عنصر نائب يصف أو يوجه كل حقل
1. داخل عنصر النموذج، يجب أن يكون لديك عنصر قائمة منسدلة `select` يحمل `id` بقيمة `dropdown` ويحتوي على خيارين على الأقل للاختيار
1. داخل عنصر النموذج، يمكنك اختيار خيار من مجموعة تحتوي على زري اختيار على الأقل، مجمعة باستخدام خاصية `name`
1. داخل عنصر النموذج، يمكنك اختيار عدة حقول من سلسلة مربعات اختيار، يجب أن يحتوي كل منها على خاصية `value`
1. داخل عنصر النموذج، يُعرض عليك مربع نص `textarea` لتعليقات إضافية
1. داخل عنصر النموذج، يُعرض عليك زر يحمل `id` بقيمة `submit` لإرسال جميع المُدخلات

نفذ قصص المستخدم واجتز جميع الاختبارات أدناه لإكمال هذا المشروع. أضف لمستك الشخصية. برمجة سعيدة!

**ملاحظة:** تأكد من إضافة `<link rel="stylesheet" href="styles.css">` في ملف HTML لربط ورقة الأنماط وتطبيق CSS الخاص بك

# --hints--

يجب أن يكون لديك عنصر `h1` يحمل `id` بقيمة `title`.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

يجب ألا يكون `#title` فارغًا.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يكون لديك عنصر `p` يحمل `id` بقيمة `description`.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

يجب ألا يكون `#description` فارغًا.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يكون لديك عنصر `form` يحمل `id` بقيمة `survey-form`.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

يجب أن يكون لديك عنصر `input` يحمل `id` بقيمة `name`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

يجب أن يكون لـ `#name` خاصية `type` بقيمة `text`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

يجب أن يتطلب `#name` إدخالًا.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

يجب أن يكون `#name` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

يجب أن يكون لديك عنصر `input` يحمل `id` بقيمة `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

يجب أن يكون لـ `#email` خاصية `type` بقيمة `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

يجب أن يتطلب `#email` إدخالًا.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

يجب أن يكون `#email` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

يجب أن يكون لديك عنصر `input` يحمل `id` بقيمة `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

يجب أن يكون `#number` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

يجب أن يكون لـ `#number` خاصية `type` بقيمة `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

يجب أن يحتوي `#number` على خاصية `min` بقيمة رقمية.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

يجب أن يحتوي `#number` على خاصية `max` بقيمة رقمية.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

يجب أن يكون لديك عنصر `label` يحمل `id` بقيمة `name-label`.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

يجب أن يكون لديك عنصر `label` يحمل `id` بقيمة `email-label`.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

يجب أن يكون لديك عنصر `label` يحمل `id` بقيمة `number-label`.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

يجب أن يحتوي `#name-label` على نص يصف المُدخل.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يحتوي `#email-label` على نص يصف المُدخل.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يحتوي `#number-label` على نص يصف المُدخل.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

يجب أن يكون `#name-label` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

يجب أن يكون `#email-label` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

يجب أن يكون `#number-label` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

يجب أن يحتوي `#name` على خاصية `placeholder` وقيمة.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

يجب أن يحتوي `#email` على خاصية `placeholder` وقيمة.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

يجب أن يحتوي `#number` على خاصية `placeholder` وقيمة.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

يجب أن يكون لديك حقل `select` يحمل `id` بقيمة `dropdown`.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

يجب أن يحتوي `#dropdown` على عنصرين `option` قابلين للاختيار على الأقل (غير معطلين).

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

يجب أن يكون `#dropdown` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

يجب أن يكون لديك على الأقل عنصران من نوع `input` مع `type` من نوع `radio` (أزرار اختيار).

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

يجب أن يكون لديك على الأقل زري اختيار تابعين لـ `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

يجب أن تحتوي جميع أزرار الاختيار على خاصية `value` وقيمة.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

يجب أن تحتوي جميع أزرار الاختيار على خاصية `name` وقيمة.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

يجب أن تحتوي كل مجموعة أزرار اختيار على زرين اختيار على الأقل.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
});

assert.isAbove(groupKeys.length, 0);
```

يجب أن يحتوي النموذج على عنصرين على الأقل من نوع `input` مع خاصية `type` تساوي `checkbox`، ويكون كل منهما عنصرًا تابعًا لـ `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

يجب أن تحتوي جميع مربعات الاختيار داخل `#survey-form` على خاصية `value` وقيمة.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

يجب أن يكون لديك على الأقل عنصر `textarea` تابع لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

يجب أن يكون لديك عنصر `input` أو `button` يحمل `id` بقيمة `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

يجب أن يكون لـ `#submit` خاصية `type` بقيمة `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

يجب أن يكون `#submit` عنصرًا تابعًا لـ `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Survey Form</title>
  </head>
  <body>
    <h1>Survey Form</h1>
    <p>The card below was built as a sample survey form for freeCodeCamp.</p>
    <main id="main">
      <h1 id="title">Join the Togepi Fan Club!</h1>
      <p id="description">
        Enter your information here to receive updates about club activities,
        our monthly newsletter, and other email communications.
      </p>
      <form id="survey-form" action="#">
        <label for="name" id="name-label"
          ><p>Name:</p>
          <input type="text" id="name" placeholder="e.g. John Smith" required />
        </label>
        <label for="email" id="email-label"
          ><p>Email:</p>
          <input
            type="email"
            id="email"
            placeholder="e.g. john.smith@email.com"
            required
          />
        </label>
        <label for="age" id="number-label"
          ><p>Age<em>(optional)</em>:</p>
          <input
            type="number"
            id="number"
            placeholder="e.g. 19"
            min="18"
            max="99"
          />
        </label>
        <label for="interest" id="interest-label"
          ><p>What are you most interested in?</p>
          <select id="dropdown">
            <option selected disabled hidden></option>
            <option id="battles">Battling</option>
            <option id="breeding">Breeding</option>
            <option id="catching">Completing my Pokedex</option>
            <option id="exploring">Exploring new regions</option>
          </select>
        </label>
        <p>Who is your favourite Pokemon?</p>
        <label for="togepi">
          <input
            id="togepi"
            type="radio"
            name="favorite"
            value="togepi"
          />Togepi!
        </label>
        <label for="pikachu">
          <input
            id="pikachu"
            type="radio"
            name="favorite"
            value="pikachu"
          />Pikachu
        </label>
        <label for="other">
          <input id="other" type="radio" name="favorite" value="other" />Other
        </label>
        <p>Which communications do you want to receive?</p>
        <label for="newsletter">
          <input
            id="newsletter"
            type="checkbox"
            name="communications"
            value="newsletter"
          />Newsletter
        </label>
        <label for="events">
          <input
            id="events"
            type="checkbox"
            name="communications"
            value="events"
          />Event updates
        </label>
        <label for="updates">
          <input
            id="updates"
            type="checkbox"
            name="communications"
            value="updates"
          />Club updates
        </label>
        <p>Any other information you would like to share?</p>
        <textarea id="additional-information" rows="4" cols="50">
You can provide comments, suggestions, or feedback here.</textarea
        >
        <p>
          <em
            >Please note: This form is a proof of concept. Submitting the form
            will not actually transmit your data.</em
          >
        </p>
        <button type="Submit" id="submit">Submit</button>
      </form>
    </main>
  </body>
  <footer>
    <a href="../">Return to Project List</a> |
    <a href="https://www.nhcarrigan.com">Return to HomePage</a>
  </footer>
</html>
```

```css
main {
  text-align: center;
  background-color: #92869c;
  background-blend-mode: lighten;
  max-width: 500px;
  margin: 20px auto;
  border-radius: 50px;
  box-shadow: 10px 10px rgba(0, 0, 0, 0.5);
  color: black;
}
body {
  text-align: center;
  background: #3a3240;
  color: white;
}
input,
textarea,
select,
button {
  background: #3a3240;
  color: white;
}
a {
  color: white;
}
```
