---
id: 587d7fb1367417b2b2512bf2
title: ".env 파일 사용하기"
challengeType: 2
forumTopicId: 301521
dashedName: use-the--env-file
---

# --description--

`.env` 파일은 애플리케이션에 환경 변수를 전달하는 데 사용되는 숨김 파일입니다. 이 파일은 비밀이며, 본인만 접근할 수 있고, 비공개로 유지하거나 숨기고 싶은 데이터를 저장하는 데 사용할 수 있습니다. 예를 들어, 외부 서비스의 API 키나 데이터베이스 URI를 저장할 수 있습니다. 또한 구성 옵션을 저장하는 데 사용할 수도 있습니다. 구성 옵션을 설정하면 일부 코드를 다시 작성하지 않고도 애플리케이션의 동작을 변경할 수 있습니다.

환경 변수는 `process.env.VAR_NAME`에서 접근할 수 있습니다. `process.env` 객체는 전역 Node 객체이며, 변수들은 문자열로 전달됩니다. 관례상 변수 이름은 모두 대문자이고, 단어는 밑줄로 구분합니다. `.env`는 셸 파일이므로 이름이나 값을 따옴표로 감쌀 필요가 없습니다. 변수에 값을 할당할 때 등호 양쪽에 공백이 있으면 안 된다는 점도 중요합니다. 예를 들어 `VAR_NAME=value`와 같습니다. 보통 각 변수 정의는 별도의 줄에 작성합니다.

# --instructions--

환경 변수를 구성 옵션으로 추가해 보세요.

프로젝트 디렉터리 루트에 `.env` 파일을 만들고, 그 안에 변수 `MESSAGE_STYLE=uppercase`를 저장하세요.

그런 다음, 이전 도전 과제에서 만든 `/json` GET 라우트 핸들러에서 `process.env.MESSAGE_STYLE`에 접근하고, 변수가 `message`과 같으면 응답 객체의 `uppercase`를 대문자로 변환하세요. 응답 객체는 `{"message": "Hello json"}` 값에 따라 `{"message": "HELLO JSON"}` 또는 `MESSAGE_STYLE`여야 합니다. 테스트가 실행되는 방식 때문에 `process.env.MESSAGE_STYLE`의 값을 라우트 핸들러 **내부에서** 읽어야 하며, 외부에서 읽으면 안 됩니다.

`dotenv` 패키지를 사용해야 합니다. 이 패키지는 `.env` 파일에서 환경 변수를 로드하여 `process.env`에 넣어 줍니다. `dotenv` 패키지는 이미 설치되어 있으며, 프로젝트의 `package.json` 파일에 있습니다. `myApp.js` 파일 상단에 `require('dotenv').config()`를 추가하여 환경 변수를 로드하세요.

# --hints--

엔드포인트 `/json`의 응답은 환경 변수 `MESSAGE_STYLE`에 따라 달라져야 합니다.

```js
  const response = await fetch(code + '/_api/use-env-vars');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(
    data.passed,
    'The response of "/json" does not change according to MESSAGE_STYLE'
  );
```
