---
id: 587d7fb2367417b2b2512bf6
title: 클라이언트에서 쿼리 매개변수 입력 받기
challengeType: 2
forumTopicId: 301512
dashedName: get-query-parameter-input-from-the-client
---

# --description--

클라이언트로부터 입력을 받는 또 다른 일반적인 방법은 라우트 경로 뒤에 데이터를 쿼리 문자열로 인코딩하는 것입니다. 쿼리 문자열은 물음표(?)로 구분되며, 필드=값 쌍을 포함합니다. 각 쌍은 앰퍼샌드(&)로 구분됩니다. Express는 쿼리 문자열에서 데이터를 파싱하여 `req.query` 객체를 채울 수 있습니다. 퍼센트(%) 같은 일부 문자는 URL에 포함될 수 없으며, 전송하기 전에 다른 형식으로 인코딩해야 합니다. JavaScript API를 사용하면 이러한 문자를 인코딩/디코딩하는 특정 메서드를 사용할 수 있습니다.

<blockquote>route_path: '/library'<br>actual_request_URL: '/library?userId=546&#x26;bookId=6754' <br>req.query: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

`GET /name`에 마운트된 API 엔드포인트를 만드세요. `{ name: 'firstname lastname'}` 구조를 가진 JSON 문서로 응답하세요. 이름과 성 매개변수는 쿼리 문자열로 인코딩되어야 합니다. 예를 들어 `?first=firstname&last=lastname`와 같습니다.

**참고:** 다음 연습에서는 같은 `/name` 라우트 경로에서 POST 요청으로 데이터를 받게 됩니다. 원한다면 `app.route(path).get(handler).post(handler)` 메서드를 사용할 수 있습니다. 이 구문은 같은 경로 라우트에서 여러 동사 핸들러를 체인으로 연결할 수 있게 해줍니다. 타이핑을 줄이고 코드를 더 깔끔하게 만들 수 있습니다.

# --hints--

Test 1 : `{ "name": "Mick Jagger" }` 엔드포인트가 `/name`와 함께 호출될 때 API 엔드포인트가 `?first=Mick&last=Jagger`로 응답해야 합니다.

```js
  const response = await fetch(code + '/name?first=Mick&last=Jagger');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "GET /name" route does not behave as expected'
  );
```

Test 2 : `{ "name": "Keith Richards" }` 엔드포인트가 `/name`와 함께 호출될 때 API 엔드포인트가 `?first=Keith&last=Richards`로 응답해야 합니다.

```js
  const response = await fetch(code + '/name?last=Richards&first=Keith');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "GET /name" route does not behave as expected'
  );
```
