---
id: 587d7fb2367417b2b2512bf8
title: POST 요청에서 데이터 가져오기
challengeType: 2
forumTopicId: 301511
dashedName: get-data-from-post-requests
---

# --description--

경로 `/name`에 POST 핸들러를 마운트하세요. 이전과 같은 경로입니다. html 프론트페이지에 폼이 준비되어 있습니다. 이 폼은 연습문제 10(Query string)과 동일한 데이터를 제출할 것입니다. 바디 파서가 올바르게 설정되어 있다면, 매개변수를 `req.body` 객체에서 찾을 수 있습니다. 일반적인 라이브러리 예제를 참고하세요:

<blockquote>라우트: POST '/library'<br>urlencoded_body: userId=546&#x26;bookId=6754 <br>req.body: {userId: '546', bookId: '6754'}</blockquote>

이전과 동일한 JSON 객체 `{name: 'firstname lastname'}`로 응답하세요. 제공된 앱 프론트페이지의 html 폼을 사용해 엔드포인트가 작동하는지 테스트하세요.

팁: GET과 POST 외에도 여러 http 메서드가 있습니다. 관례상 http 동사와 서버에서 실행할 작업 간에 대응 관계가 있습니다. 일반적인 매핑은 다음과 같습니다:

POST (때때로 PUT) - 요청과 함께 전송된 정보를 사용해 새 리소스를 생성합니다,

GET - 리소스를 수정하지 않고 읽습니다,

PUT 또는 PATCH (때때로 POST) - 전송된 데이터를 사용해 리소스를 업데이트합니다,

DELETE - 리소스를 삭제합니다.

서버와 연결을 협상하는 데 사용되는 몇 가지 다른 메서드도 있습니다. GET을 제외한 위에 나열된 모든 메서드는 페이로드(즉, 요청 본문에 포함된 데이터)를 가질 수 있습니다. 바디 파서 미들웨어는 이러한 메서드와도 작동합니다.

# --hints--

Test 1 : API 엔드포인트가 올바른 이름으로 응답하는지 확인하세요

```js
  const response = await fetch(code + '/name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/x-www-form-urlencoded'
    },
    body: new URLSearchParams({ first: 'Mick', last: 'Jagger' })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "POST /name" route does not behave as expected'
  );
```

Test 2 : API 엔드포인트가 올바른 이름으로 응답하는지 확인하세요

```js
  const response = await fetch(code + '/name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/x-www-form-urlencoded'
    },
    body: new URLSearchParams({
      first: 'Keith',
      last: 'Richards'
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "POST /name" route does not behave as expected'
  );
```
