---
id: 587d7fb0367417b2b2512bf0
title: 정적 자산 제공하기
challengeType: 2
forumTopicId: 301518
dashedName: serve-static-assets
---

# --description--

HTML 서버는 보통 사용자가 접근할 수 있는 하나 이상의 디렉터리를 갖고 있습니다. 애플리케이션에 필요한 정적 자산(스타일시트, 스크립트, 이미지)을 그곳에 배치할 수 있습니다.

Express에서는 미들웨어 `express.static(path)`를 사용해 이 기능을 구현할 수 있으며, `path` 매개변수는 자산이 들어 있는 폴더의 절대 경로입니다.

미들웨어가 무엇인지 모른다면 걱정하지 마세요. 나중에 자세히 다룰 예정입니다. 기본적으로 미들웨어는 라우트 핸들러를 가로채서 어떤 정보를 추가하는 함수입니다. 미들웨어는 `app.use(path, middlewareFunction)` 메서드를 사용해 마운트해야 합니다. 첫 번째 `path` 인수는 선택 사항입니다. 전달하지 않으면 모든 요청에 대해 미들웨어가 실행됩니다.

# --instructions--

`express.static()` 미들웨어를 `/public` 경로에 `app.use()`와 함께 마운트하세요. 자산 폴더의 절대 경로는 `__dirname + /public`입니다.

이제 앱이 CSS 스타일시트를 제공할 수 있어야 합니다. `/public/style.css` 파일이 프로젝트 보일러플레이트의 `/views/index.html`에서 참조된 점을 참고하세요. 프론트 페이지가 좀 더 보기 좋아질 것입니다!

# --hints--

앱은 `/public` 디렉터리의 자산 파일을 `/public` 경로로 제공해야 합니다

```js
  const response = await fetch(code + '/public/style.css');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /body\s*\{[^\}]*\}/,
    'Your app does not serve static assets'
  );
```

앱은 `/public` 디렉터리 외의 다른 폴더에서 파일을 제공하면 안 됩니다

```js
  const response = await fetch(code + '/server.js');
  if (response.ok) {
    const data = await response.text();
    assert.equal(
      response.status + '',
      404 + '',
      'Your app must serve files only from "public" directory'
    );
  } else {
    assert.equal(
      response.status + '',
      404 + '',
      'Your app must serve files only from "public" directory'
    );
  }
```
