---
id: 587d7fb1367417b2b2512bf3
title: 루트 레벨 요청 로거 미들웨어 구현하기
challengeType: 2
forumTopicId: 301514
dashedName: implement-a-root-level-request-logger-middleware
---

# --description--

앞서 `express.static()` 미들웨어 함수에 대해 배웠습니다. 이제 미들웨어가 무엇인지 좀 더 자세히 살펴볼 시간입니다. 미들웨어 함수는 3개의 인수를 받는 함수입니다: 요청 객체, 응답 객체, 그리고 애플리케이션의 요청-응답 사이클에서 다음 함수입니다. 이 함수들은 애플리케이션에 부수 효과를 일으킬 수 있는 코드를 실행하며, 보통 요청이나 응답 객체에 정보를 추가합니다. 또한 어떤 조건이 충족되면 응답을 보내면서 사이클을 종료할 수도 있습니다. 만약 응답을 보내지 않으면, 스택에 있는 다음 함수의 실행을 시작합니다. 이때 3번째 인수인 `next()`를 호출합니다.

다음 예제를 보세요:

```js
function(req, res, next) {
  console.log("I'm a middleware...");
  next();
}
```

이 함수를 라우트에 마운트했다고 가정해봅시다. 요청이 라우트와 일치하면, “I’m a middleware…”라는 문자열을 출력한 후 스택에 있는 다음 함수를 실행합니다. 이번 도전 과제에서는 루트 레벨 미들웨어를 만들 것입니다. 도전 과제 4에서 본 것처럼, 루트 레벨에 미들웨어 함수를 마운트하려면 `app.use(<mware-function>)` 메서드를 사용할 수 있습니다. 이 경우 모든 요청에 대해 함수가 실행되지만, 더 구체적인 조건도 설정할 수 있습니다. 예를 들어, POST 요청에만 함수를 실행하고 싶다면 `app.post(<mware-function>)`를 사용할 수 있습니다. 모든 HTTP 동사(GET, DELETE, PUT 등)에 대해 유사한 메서드가 존재합니다.

# --instructions--

간단한 로거를 만드세요. 모든 요청에 대해 다음 형식의 문자열을 콘솔에 기록해야 합니다: `method path - ip`. 예시는 다음과 같습니다: `GET /json - ::ffff:127.0.0.1`. `method`와 `path` 사이에 공백이 있고, `path`와 `ip`를 구분하는 대시는 양쪽에 공백이 있다는 점에 주의하세요. 요청 메서드(HTTP 동사), 상대 라우트 경로, 호출자의 IP는 요청 객체에서 `req.method`, `req.path`, `req.ip`를 사용해 얻을 수 있습니다. 작업이 끝나면 반드시 `next()`를 호출하세요. 그렇지 않으면 서버가 영원히 멈춥니다. ‘Logs’를 열어두고 요청이 도착할 때 어떤 일이 일어나는지 확인하세요.

**참고:** Express는 코드에 나타난 순서대로 함수를 평가합니다. 미들웨어도 마찬가지입니다. 모든 라우트에 대해 작동하게 하려면, 라우트보다 먼저 마운트해야 합니다.

# --hints--

루트 레벨 로거 미들웨어가 활성화되어야 합니다

```js
  const response = await fetch(code + '/_api/root-middleware-logger');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(
    data.passed,
    'root-level logger is not working as expected'
  );
```
