---
id: 587d7fb0367417b2b2512bee
title: 작동하는 Express 서버 시작하기
challengeType: 2
forumTopicId: 301519
dashedName: start-a-working-express-server
---

# --description--

파일 `myApp.js`의 처음 두 줄에서 Express 앱 객체를 얼마나 쉽게 생성할 수 있는지 볼 수 있습니다. 이 객체는 여러 메서드를 가지고 있으며, 이 도전 과제들에서 그중 많은 메서드를 배우게 됩니다. 기본적인 메서드 중 하나는 `app.listen(port)`입니다. 이 메서드는 서버가 특정 포트에서 대기하도록 하여 실행 상태로 만듭니다. 테스트를 위해 앱이 백그라운드에서 실행 중이어야 하므로 `server.js` 파일에 이 메서드를 추가했습니다.

첫 번째 문자열을 제공해 봅시다! Express에서 라우트는 다음과 같은 구조를 가집니다: `app.METHOD(PATH, HANDLER)`. METHOD는 소문자 HTTP 메서드입니다. PATH는 서버의 상대 경로로, 문자열이거나 정규 표현식일 수 있습니다. HANDLER는 라우트가 일치할 때 Express가 호출하는 함수입니다. 핸들러는 `function(req, res) {...}` 형태이며, 여기서 req는 요청 객체, res는 응답 객체입니다. 예를 들어, 핸들러

```js
function(req, res) {
  res.send('Response String');
}
```

는 문자열 'Response String'을 제공합니다.

# --instructions--

`app.get()` 메서드를 사용해 `/` (루트) 경로에 일치하는 GET 요청에 "Hello Express" 문자열을 제공하세요. 로그를 확인해 코드가 제대로 작동하는지 확인하세요.

**참고:** 이 수업의 모든 코드는 시작해 드린 몇 줄 코드 사이에 추가해야 합니다.

# --hints--

앱이 'Hello Express' 문자열을 제공해야 합니다.

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.equal(
    data,
    'Hello Express',
    'Your app does not serve the text "Hello Express"'
  );
```
