---
id: 587d7fb1367417b2b2512bf1
title: 특정 라우트에서 JSON 제공하기
challengeType: 2
forumTopicId: 301517
dashedName: serve-json-on-a-specific-route
---

# --description--

HTML 서버가 HTML을 제공하는 동안, API는 데이터를 제공합니다. <dfn>REST</dfn>(표현 상태 전송) API는 클라이언트가 서버에 대한 세부 정보를 알 필요 없이 간단한 방식으로 데이터 교환을 가능하게 합니다. 클라이언트는 리소스가 어디에 있는지(URL)와 그 리소스에 대해 수행하려는 동사만 알면 됩니다. GET 동사는 정보를 가져올 때, 아무것도 수정하지 않고 사용할 때 쓰입니다. 요즘 웹에서 정보를 주고받는 데 선호되는 데이터 형식은 JSON입니다. 간단히 말해, JSON은 JavaScript 객체를 문자열로 나타내어 쉽게 전송할 수 있게 하는 편리한 방법입니다.

경로 `/json`에서 JSON으로 응답하는 라우트를 만들어 간단한 API를 만들어 봅시다. 평소처럼 `app.get()` 메서드를 사용해 만들 수 있습니다. 라우트 핸들러 안에서는 `res.json()` 메서드를 사용하고, 인수로 객체를 전달하세요. 이 메서드는 요청-응답 루프를 닫아 데이터를 반환합니다. 내부적으로는 유효한 JavaScript 객체를 문자열로 변환하고, 브라우저에 JSON을 제공한다는 적절한 헤더를 설정한 뒤 데이터를 다시 보냅니다. 유효한 객체는 일반적으로 `{key: data}` 구조를 가집니다. `data`는 숫자, 문자열, 중첩된 객체 또는 배열일 수 있습니다. `data`는 변수나 함수 호출 결과일 수도 있는데, 이 경우 문자열로 변환되기 전에 평가됩니다.

# --instructions--

GET 요청에 대해 `{"message": "Hello json"}` 라우트에서 객체 `/json`를 JSON 형식으로 응답하세요. 그런 다음 브라우저에서 `your-app-url/json`로 접속하면 화면에 메시지가 표시됩니다.

# --hints--

엔드포인트 `/json`는 JSON 객체 `{"message": "Hello json"}`를 제공해야 합니다.

```js
  const response = await fetch(code + '/json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.message,
    'Hello json',
    "The '/json' endpoint does not serve the right data"
  );
```
