---
id: 587d7fb0367417b2b2512bef
title: HTML 파일 제공하기
challengeType: 2
forumTopicId: 301516
dashedName: serve-an-html-file
---

# --description--

`res.sendFile(path)` 메서드를 사용해 요청에 파일로 응답할 수 있습니다. `app.get('/', ...)` 라우트 핸들러 안에 넣을 수 있습니다. 내부적으로 이 메서드는 전송하려는 파일의 유형에 따라 브라우저가 파일을 처리하는 방법을 지시하는 적절한 헤더를 설정합니다. 그런 다음 파일을 읽어 전송합니다. 이 메서드는 절대 파일 경로가 필요합니다. Node 전역 변수 `__dirname` 를 사용해 다음과 같이 경로를 계산하는 것을 권장합니다:

```js
absolutePath = __dirname + '/relativePath/file.ext'
```

# --instructions--

GET 요청에 대해 `/views/index.html` 경로로 `/` 파일을 응답으로 보내세요. 라이브 앱을 보면 스타일이 적용되지 않은 큰 HTML 제목(그리고 나중에 사용할 폼…)이 보여야 합니다.

**참고:** 이전 도전 과제의 해답을 수정하거나 새로 만들 수 있습니다. 새 해답을 만들면 Express가 라우트를 위에서 아래로 평가하고 첫 번째 일치하는 핸들러를 실행한다는 점을 기억하세요. 이전 해답은 주석 처리해야 서버가 계속 문자열로 응답하지 않습니다.

# --hints--

앱이 views/index.html 파일을 제공해야 합니다

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /<h1>.*<\/h1>/,
    'Your app does not serve the expected HTML'
  );
```
