---
id: 587d7fb0367417b2b2512bed
title: 노드 콘솔 사용하기
challengeType: 2
forumTopicId: 301515
dashedName: meet-the-node-console
---

# --description--

이 도전 과제들을 진행하려면 다음 방법 중 하나를 사용해 코드를 작성해야 합니다:

- <a href="https://github.com/freeCodeCamp/boilerplate-express/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 도전 과제를 완료하세요.
- 원하는 사이트 빌더를 사용해 프로젝트를 완성하세요. GitHub 저장소의 모든 파일을 반드시 포함해야 합니다.

개발 과정에서는 코드에서 무슨 일이 일어나고 있는지 확인하는 것이 중요합니다.

노드는 단순한 JavaScript 환경입니다. 클라이언트 측 JavaScript처럼 콘솔을 사용해 유용한 디버그 정보를 표시할 수 있습니다. 로컬 컴퓨터에서는 터미널에서 콘솔 출력을 볼 수 있습니다.

이 도전 과제를 진행하는 동안 터미널을 열어 두는 것을 권장합니다. 터미널 출력을 읽으면 발생할 수 있는 오류를 확인할 수 있습니다.

서버 파일을 변경한 후에는 서버를 다시 시작해야 합니다.

터미널에서 `Ctrl + C`를 사용해 서버를 중지할 수 있으며, `node mainEntryFile.js`를 직접 사용하거나 `package.json` 파일의 실행 스크립트 `npm run`를 사용해 서버를 시작할 수 있습니다.

예를 들어, `"start": "node server.js"` 스크립트는 터미널에서 `npm run start` 명령어로 실행됩니다.

파일 저장 시 서버 자동 재시작을 구현하려면 노드가 제공하는 `--watch` 플래그를 시작 스크립트 `"start": "node --watch server.js"`에 추가하거나 `nodemon` 같은 npm 패키지를 설치할 수 있습니다. 이 부분은 직접 시도해 보세요.

# --instructions--

`myApp.js` 파일을 수정해 콘솔에 "Hello World"를 출력하세요.

# --hints--

`"Hello World"`가 콘솔에 있어야 합니다.

```js
  const response = await fetch(code + '/_api/hello-console');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.passed, '"Hello World" is not in the server console');
```
