---
id: 587d7fb2367417b2b2512bf5
title: 클라이언트로부터 라우트 매개변수 입력 받기
challengeType: 2
forumTopicId: 301513
dashedName: get-route-parameter-input-from-the-client
---

# --description--

API를 구축할 때, 사용자가 서비스에서 얻고자 하는 것을 전달할 수 있게 해야 합니다. 예를 들어, 클라이언트가 데이터베이스에 저장된 사용자에 관한 정보를 요청한다면, 어떤 사용자를 원하는지 알려줄 방법이 필요합니다. 이를 달성하는 한 가지 방법은 라우트 매개변수를 사용하는 것입니다. 라우트 매개변수는 URL의 이름이 지정된 구간으로, 슬래시(/)로 구분됩니다. 각 구간은 URL에서 해당 위치에 일치하는 부분의 값을 캡처합니다. 캡처된 값은 `req.params` 객체에서 찾을 수 있습니다.

<blockquote>route_path: '/user/:userId/book/:bookId'<br>actual_request_URL: '/user/546/book/6754' <br>req.params: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

라우트 `GET /:word/echo`에 마운트된 에코 서버를 만드세요. `{echo: word}` 구조를 가진 JSON 객체로 응답하세요. 반복할 단어는 `req.params.word`에서 찾을 수 있습니다. 브라우저 주소창에서 일치하는 라우트를 방문하여 예를 들어 `your-app-rootpath/freecodecamp/echo`에서 라우트를 테스트할 수 있습니다.

# --hints--

Test 1 : 에코 서버가 단어를 올바르게 반복해야 합니다

```js
  const response = await fetch(code + '/eChOtEsT/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'eChOtEsT',
    'Test 1: the echo server is not working as expected'
  );
```

Test 2 : 에코 서버가 단어를 올바르게 반복해야 합니다

```js
  const response = await fetch(code + '/ech0-t3st/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'ech0-t3st',
    'Test 2: the echo server is not working as expected'
  );
```
