---
id: 587d7fb1367417b2b2512bf4
title: 시간 서버를 만들기 위해 미들웨어 연결하기
challengeType: 2
forumTopicId: 301510
dashedName: chain-middleware-to-create-a-time-server
---

# --description--

미들웨어는 `app.METHOD(path, middlewareFunction)`를 사용해 특정 라우트에 장착할 수 있습니다. 미들웨어는 라우트 정의 내에서 연결할 수도 있습니다.

다음 예제를 살펴보세요:

```js
app.get('/user', function(req, res, next) {
  req.user = getTheUserSync();  // Hypothetical synchronous operation
  next();
}, function(req, res) {
  res.send(req.user);
});
```

이 방법은 서버 작업을 더 작은 단위로 나누는 데 유용합니다. 이는 더 나은 앱 구조와 여러 곳에서 코드를 재사용할 수 있는 가능성을 제공합니다. 이 방법은 데이터에 대한 유효성 검사를 수행하는 데도 사용할 수 있습니다. 미들웨어 스택의 각 지점에서 현재 체인의 실행을 중단하고 오류를 처리하도록 특별히 설계된 함수에 제어를 넘길 수 있습니다. 또는 특수한 경우를 처리하기 위해 다음 일치하는 라우트에 제어를 넘길 수도 있습니다. 고급 Express 섹션에서 그 방법을 살펴보겠습니다.

# --instructions--

라우트 `app.get('/now', ...)`에서 미들웨어 함수와 최종 핸들러를 연결하세요. 미들웨어 함수에서는 현재 시간을 `req.time` 키에 `new Date().toString()`를 사용해 요청 객체에 추가해야 합니다. 핸들러에서는 `{time: req.time}` 구조를 가진 JSON 객체로 응답하세요.

**참고:** 미들웨어를 연결하지 않으면 테스트를 통과하지 못합니다. 함수를 다른 곳에 장착하면 출력 결과가 맞더라도 테스트가 실패합니다.

# --hints--

/now 엔드포인트에는 미들웨어가 장착되어 있어야 합니다.

```js
  const response = await fetch(code + '/_api/chain-middleware-time');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.stackLength,
    2,
    '"/now" route has no mounted middleware'
  );
```

`/now` 엔드포인트는 현재 시간을 반환해야 합니다.

```js
  const response = await fetch(code + '/_api/chain-middleware-time');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  var now = new Date();
  assert.isAtMost(
    Math.abs(new Date(data.time) - now),
    20000,
    'the returned time is not between +- 20 secs from now'
  );
```
