---
id: 65c4f37ba6cbcfdc77e24165
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

ユーザーが`insert`メソッドを呼び出します。`self`パラメータに加えて、`key`パラメータも必要です。このパラメータは二分探索木に挿入するキーの値になります。

関数定義に`key`を2番目のパラメータとして追加してください。

# --hints--

`insert`メソッドには`self`パラメータを含める必要があります。

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

`insert`メソッドには`key`パラメータを含める必要があります。

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
