---
id: 65c4f09e074dd8d37830ea00
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

二分探索木を表す、もう一つの空のクラス`BinarySearchTree`を作成してください。

# --hints--

`BinarySearchTree`という名前のクラスを作成してください。`pass`キーワードを使うことを忘れないでください。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
