---
id: 65c4f0d32ca17ad4b1636b0e
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

`__init__` メソッド内で `pass` を削除し、`root` を値 `None` に初期化してください。

`root` 属性は二分探索木のルートノードを表します。これはコンストラクタなので、新しい `BinarySearchTree` オブジェクトが作成されるとき、空の木から始まるため、`root` 属性は `None` に設定されます。

# --hints--

`pass` メソッドから `__init__` 文を削除してください。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

`root` を使って `None` 属性を `self.root` に初期化してください。

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
