---
id: 65c4f300da28d8da361bfa93
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

`elif` 節の中で、`pass` を現在のノードの右の子を最初の引数、`_insert` を2番目の引数として `key` メソッドの呼び出しに置き換えてください。結果を現在のノードの `right` 属性に代入してください。

# --hints--

`pass` ブロックから `elif` キーワードを削除してください。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

`self._insert()` メソッドを呼び出し、`node.right` と `key` を引数として渡してください。

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

`self._insert()` の呼び出し結果を現在のノードの正しい属性（`node.right`）に代入してください。

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
