---
id: 65c4eb814cc977c95cd7df0e
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このプロジェクトでは、二分探索木を作成します。二分探索木は、各ノードが最大で2つの子を持ち、左の子は親ノードより小さい値を、右の子は親ノードより大きい値を持つデータ構造で、効率的な検索やソート操作が可能です。

二分探索木はこのような形をしています：

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

まず、空の`TreeNode`クラスを定義してください。`TreeNode`クラスは二分探索木のノードを表します。エラーを避けるために、クラス本体を埋めるには`pass`キーワードを使ってください。

# --hints--

空の`class`という名前のクラスを宣言するには、`TreeNode`キーワードを使う必要があります。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
