---
id: 65c4f33bc3c3d8db7f732787
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

実際の挿入を行うために、`insert` クラス内に空の `BinarySearchTree` メソッドを定義し、`self` パラメータを与えてください。

# --hints--

`insert` クラス内に `self` パラメータを持つ `BinarySearchTree` メソッドを定義する必要があります。`pass` キーワードを忘れないでください。

```js
({ test: () => assert.match(code, /^\s{4}def\s+insert\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node
--fcc-editable-region--    

--fcc-editable-region--

```
