---
id: 65c4f2851645e8d84d14f5d1
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

これからは再帰的に二分木をたどり、値を挿入する必要があります。二分木の原則に従ってください：

- キーより小さい値は左のサブツリーに配置します
- キーより大きい値は右のサブツリーに配置します

既存の条件文の後に、`if` 文をもう一つ書いて、`key` が `node.key` より小さいかどうかを確認してください。

# --hints--

`key` が `node.key` より小さいかどうかを確認するために、もう一つ if文 を書く必要があります。`pass` キーワードを使うことを忘れないでください。

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```
