---
id: 65c4f02dbd3557d23f12db62
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

`__init__` メソッド内で `pass` を削除し、`key` パラメータの値を `key` を使ってノードの `self.key` 属性に代入してください。

これは、`key` インスタンスの `TreeNode` 属性がオブジェクト作成時に渡された値に設定されることを意味します。

# --hints--

`pass` メソッドから `__init__` 文を削除してください。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

`key` パラメータの値を `key` を使ってノードの `self.key` 属性に代入してください。

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
