---
id: 65c4f013851cefd1a4fe4c96
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

`__init__` メソッドは2つのパラメータを取ります：`self`（作成されるクラスのインスタンスを表す）と `key`（ノードに格納される値）。これら2つのパラメータを `__init__()` メソッドに追加してください。

# --hints--

`self` パラメータをメソッドに追加してください。

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

`key` パラメータをメソッドに追加してください。

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
