---
id: 65c4f0b82db314d3fcc0b8e1
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`BinarySearchTree` クラス内で、`pass` を `__init__` メソッドに置き換え、このメソッドに `self` パラメータを追加してください。

# --hints--

`pass` クラスから `BinarySearchTree` キーワードを削除してください。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


`__init__` メソッドを定義し、このメソッドに `self` パラメータを追加してください。メソッド本体内で `pass` キーワードを使うことを忘れないでください。

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
