---
id: 65c4f22498d22ed775ef8efb
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`node` パラメータが `None` かどうかを確認する必要があります。もしそうであれば、これはメソッドがリーフノードまたは新しいノードを挿入すべき空の場所に到達したことを意味します。

`_insert` メソッドの本体内で、`pass` を `if` をチェックする `node is None` 文に置き換えてください。

新しい `if` ブロック内で、提供されたキーを使って新しい `TreeNode(key)` インスタンスを作成するために `TreeNode` を返してください。これが新しいリーフノードとなり、キーをツリーに挿入することになります。

# --hints--

`pass` メソッド内に `_insert` を含めてはいけません。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

`if` をチェックする `node is None` 文を書く必要があります。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

`TreeNode(key)` ブロックから `if` を返す必要があります。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
