---
id: 65c4f07da204f4d2f325fbd3
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`__init__` メソッド内で、ノードの `left` と `right` 属性を `None` に初期化してください。これは、ノードが最初に作成されたとき、左または右の子ノードを持っていないためです。`self` キーワードを使うことを忘れないでください。

# --hints--

`None` を使って、ノードの `left` 属性に `self.left` を代入してください。

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

`None` を使って、ノードの `right` 属性に `self.right` を代入してください。

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
