---
id: 65c4f195de7e2ad5932be717
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

次に、ツリーにノードを挿入する仕組みを定義する必要があります。そのために、`_insert` メソッドを定義してください。これはヘルパー関数で、後で実際の `insert` メソッドで使われます。

このメソッドは再帰的で、自身を呼び出してツリーをたどり、新しいノードを挿入する適切な場所を見つけます。

`_insert`、`self`、`node` のパラメータを持つ `key` メソッドを定義してください。

# --hints--

`_insert` クラス内に `BinarySearchTree` メソッドを定義してください。`pass` を使うことを忘れないでください。

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

`_insert` メソッドは3つのパラメータ、`self`、`node`、`key` を受け取る必要があります。

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
