---
id: 65c4f3258d2e4cdacc919dfd
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

`_insert` メソッドの最後で、挿入処理が完了した後、現在のノードを返して再帰呼び出しスタックの上位レベルでツリー構造を更新してください。


# --hints--

条件分岐の外側で現在のノードを返す必要があります。

```js
({ test: () => assert.match(code, /^\s{8}return\s+node/m) })
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)

--fcc-editable-region--

```
