---
id: 65c4f2b7178afed8e88f782f
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

もし `key` < `node.key` が `True` と評価されるなら、新しいノードは左のサブツリーに配置されるべきです。

`pass` を削除し、左の子ノードを最初の引数、`_insert` を2番目の引数として `key` メソッドを再帰的に呼び出してください。結果を現在のノードの `left` 属性に代入してください。

# --hints--

`pass` ブロックから `if` キーワードを削除してください。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

`self._insert` と `node.left` を引数として渡し、`key` メソッドを呼び出してください。

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

`self._insert()` の呼び出し結果を現在のノードの `left` 属性に代入してください。

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
