---
id: 65c4ef8463f869d0eaf87c67
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

`TreeNode` クラスの中で、`pass` を `__init__` メソッドに置き換えて、オブジェクトの属性を初期化できるようにしてください。今のところパラメータは追加しないでください。

# --hints--

`pass` クラスから `TreeNode` キーワードを削除し、`__init__` メソッドの中に移動してください。

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

`__init__` クラスの中に `TreeNode` メソッドを定義してください。中で `pass` キーワードを使うことを忘れないでください。

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
