---
id: 65c63fd9b387f0a6c5bf6a72
title: ステップ 24
challengeType: 20
dashedName: step-24
---

# --description--

`search` メソッド内で `pass` を削除し、次の引数を使ってヘルパーメソッド `_search` を呼び出してください。

- `self.root`: これは二分探索木のルートです。探索はルートから始まります。
- `key`: これはユーザーが二分探索木で見つけたい値です。

内部的に、`search` は実際の探索ロジックを `_search` ヘルパーメソッドに委譲し、二分探索木内での再帰的な探索を実行します。

# --hints--

`pass` メソッドから `search` キーワードを削除してください。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

`_search` メソッド内で `search` メソッドを呼び出してください。

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

`_search` メソッドを呼び出す際に `self.root` と `key` を引数として渡してください。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
