---
id: 65c63df529bd15a24c187c62
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

ターゲットキーが現在のノードのキーより小さいかどうかを確認する別の`if`文を書いてください。

`if`ブロック内で、現在のノードの左の子ノードと`_search`を引数として`key`メソッドを呼び出した結果を返してください。

# --hints--

`if` < `key`かどうかを確認する別の`node.key`文を書く必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

新しい`self._search(node.left, key)`ブロックから`if`を返す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_bodies()[0].is_equivalent("return self._search(node.left, key)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
```
