---
id: 65c63a4c4da62e9ae18e321a
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

検索機能に取り組む時間です。`_insert` メソッドのために作成したヘルパーメソッド `insert` と同様に、`_search` メソッドのためにヘルパーメソッド `search` を作成する必要があります。

`_search` メソッドを3つのパラメータ、すなわち `self`、`node`、`key` で定義してください。

# --hints--

`_search` メソッドを宣言してください。先頭のアンダースコアに注意し、`pass` を使うことを忘れないでください。

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

`_search` メソッドは3つのパラメータ、`self`、`node`、`key` を受け取る必要があります。

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
