---
id: 65c644829cfb63acf3479d09
title: ステップ 28
challengeType: 20
dashedName: step-28
---

# --description--

リスト `nodes` のノードを `bst` インスタンスに挿入してください。リストを繰り返し処理し、各ノードに対して `insert` メソッドを呼び出します。繰り返し変数には `node` を使ってください。

# --hints--

`for node in nodes` ループが必要です。

```js
({ test: () => (runPython(`assert _Node(_code).find_for_loops()[0].find_for_vars().is_equivalent("node") and _Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("nodes")`)) });
```

`bst.insert(node)` ループの中に `for` があるはずです。

```js
({ test: () => (runPython(`assert _Node(_code).find_for("node", "nodes").find_body().is_equivalent("bst.insert(node)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

--fcc-editable-region--
```
