---
id: 65c645b838e7deb080fc25e0
title: ステップ 29
challengeType: 20
dashedName: step-29
---

# --description--

ノードが挿入されました。正しく挿入されたか確認するために、ツリー内でノードを検索できます。

forループの外で、`80` インスタンス内のノード `bst` を検索し、それを `print` の呼び出しに追加してください。また、`print` 関数の最初の引数は `'Search for 80:'` にする必要があります。

# --hints--

`bst.search(80)` を呼び出した結果を出力し、`print` 文の最初の引数は `'Search for 80:'` にしてください。

```js
assert.match(code, /^print\(('|")Search for 80:('|"),\s*bst\.search\(80\)/gm);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

--fcc-editable-region--
```
