---
id: 65c6409418806da8d0636ffc
title: ステップ 26
challengeType: 20
dashedName: step-26
---

# --description--

`insert` と `search` の機能は完成しました。次はそれらをテストする時間です。

この時点では、ノードはソートされずにただ挿入されていることに注意してください。後で中間順走査を使ったソートに取り組みます。

Pythonでクラスのインスタンスを作成するには、次のようにします。

```python
object_name = ClassName()
```

クラス定義の外で、`BinarySearchTree` クラスのインスタンスを作成し、変数 `bst` に代入してください。


# --hints--

`BinarySearchTree` クラスのインスタンスを作成する必要があります。

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

新しい `BinarySearchTree` のインスタンスを変数 `bst` に代入してください。

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
