---
id: 65c646d4148ae3b2d1cbcac4
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

検索結果が `80: <__main__.TreeNode object at 0x108b3e0>` のようなものになることに注意してください。これはクラスのインスタンスを出力するときのデフォルトの文字列表現です。

それを有用な値を出力するように変更するには、`__str__` クラスに `TreeNode` という名前の別のメソッドを定義してください。このメソッドは単一の引数 `self` を取ります。

`__str__` を定義した後、`__str__` メソッドがまだ何も返していないため、コンソールに例外が発生します。次のステップでメソッドの本体を作成します。

# --hints--

単一の引数 `__str__` を取る `self` メソッドを定義してください。`pass` を使うことを忘れないでください。

```js
assert.match(code, /^(\s+)def\s+__init__.+?^\1def\s+__str__\(\s*self\s*\)\s*:\s*\n^\1\1pass/ms)
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
