---
id: 65c63da2ed6769a10e141341
title: ステップ 20
challengeType: 20
dashedName: step-20
---

# --description--

これから再帰的な探索の基本ケースを定義します。現在の`pass`を削除し、2つの条件をチェックする`if`文を書いてください。

- `node`が`None`の場合：これは探索が枝の終わりに達し、キーが見つからなかったことを示します。
- `node.key == key`の場合：これは現在のノードでキーが見つかったことを意味します。

2つの条件を`or`演算子で結合し、`if`ブロック内で現在のノードを返してください。

# --hints--

`pass`メソッドから`_search`キーワードを削除してください。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

`if`が`node`であるか、`None`が`node.key`と等しいかをチェックする`key`文を書く必要があります。

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

`node`ブロックから`if`を返してください。

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
