---
id: 65c63e6962159fa372ecdbec
title: ステップ 22
challengeType: 20
dashedName: step-22
---

# --description--

2つ目の`if`文が`True`でない場合、対象のキーは現在のノードのキー以上であることを意味します。

二分探索木では、対象のキーが現在のノードのキーより大きい場合、探索は右のサブツリーで続行されます。

`if`ブロックの後、現在のノードの右の子ノードとキーを引数にして`_search`メソッドを呼び出した結果を返してください。

# --hints--

`_search`を使って`self._search()`メソッドを再帰的に呼び出す必要があります。

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

`_search`メソッドの呼び出しには2つの引数、`node.right`と`key`を渡す必要があります。

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

`_search`メソッド呼び出しの結果を返してください。

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
