---
id: 65c9ddd336596e30a4266a50
title: ステップ 31
challengeType: 20
dashedName: step-31
---

# --description--

`__str__` メソッドの本体で、`pass` を削除し、`str()` を引数として `self.key` 関数を呼び出した結果を返してください。これは、ノードに関連付けられた値を格納している現在のノードオブジェクトの属性です。

結果を返した後、コンソールの例外が消え、出力にノードに関連付けられた `key` の値が表示されるはずです。

# --hints--

`pass` メソッドから `__str__` キーワードを削除してください。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__str__").has_pass()`
      )
    );
  },
});

```

`self.key` メソッドから `__str__` の文字列値を返してください。

```js
({ test: () => assert.match(code, /^\s{8}return\s+str\(\s*self\.key\s*\)/m) })
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--
    def __str__(self):
        pass

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
