---
id: 65ca06864129380054b55dc0
title: ステップ 36
challengeType: 20
dashedName: step-36
---

# --description--

`if` ブロック内で、`pass` を `_delete` メソッドの呼び出しに置き換え、現在のノードの左の子ノードと `key` を引数として渡してください。関数の呼び出し結果を左のノードに代入してください。

# --hints--

`pass` 文から `if` キーワードを削除してください。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_delete").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

`_delete` メソッドを `node.left` と `key` を引数にして呼び出してください。

```js
assert.match(code, /self\._delete\(\s*node\.left\s*,\s*key\s*\)/);
```

現在のノードの左の子ノード（`_delete()`）に `node.left` の呼び出し結果を代入してください。

```js
assert.match(code, /node\.left\s*=\s*self\._delete\(\s*node\.left\s*,\s*key\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            pass

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
