---
id: 65ca07dd6fa8840491b7a5cd
title: ステップ 41
challengeType: 20
dashedName: step-41
---

# --description--

前のどちらの条件も満たさない場合、そのノードは左と右の両方の子を持っていることを意味します。

後継ノードを選ぶには、右のサブツリーの最小値を見つける必要があります。
最小の値が現在のノードの中順後継になります。

最小の値を見つけるには、2つのパラメータ `_min_value` と `self` を受け取るヘルパー関数 `node` を作成してください。

# --hints--

`_min_value` メソッドを `self` と `node` をパラメータとして定義してください。`pass` キーワードを使うことを忘れないでください。

```js
({ test: () => assert.match(code, /def\s+_min_value\(\s*self\s*,\s*node\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
--fcc-editable-region--
    
--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)
    
# print('Search for 80:', bst.search(80))


```
