---
id: 65ca0d5adf39c410cd1177cc
title: ステップ 50
challengeType: 20
dashedName: step-50
---

# --description--

これから、中順による二分木の走査に取り組みます。中順走査は深さ優先の二分木走査アルゴリズムで、左のサブツリー、現在のノード、そして右のサブツリーの順に訪問します。

`_inorder_traversal` クラス内に `BinarySearchTree` メソッドを定義し、3つのパラメータ `self`、`node`、`result` を持たせてください。ここで `node` は走査中の現在のノード、`result` はキーがソートされた順に追加されるリストです。

# --hints--

3つのパラメータ `_inorder_traversal`、`self`、`node` を受け取る `result` メソッドを定義してください。`pass` を使うことを忘れないでください。

```js
assert.match(code, /def\s+_inorder_traversal\(\s*self\s*,\s*node\s*,\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
            
            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)   
        
        return node

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

--fcc-editable-region--

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
