---
id: 65d8a6fcb15a3a239ba35dfd
title: ステップ 46
challengeType: 20
dashedName: step-46
---

# --description--

最小値を見つけた後、右のサブツリーから最小値を持つノードを再帰的に削除する必要があります。

このステップは、二分探索木（BST）の性質を保ちながら、最小値を持つノードを木から取り除くことを保証します。

`_delete` メソッドを `node.right` と `node.key` を引数にして再帰的に呼び出してください。`_delete()` の呼び出しの戻り値を現在のノードの右の子に代入してください。

# --hints--

`_delete` メソッドを `node.right` と `node.key` を引数にして再帰的に呼び出す必要があります。

```js
assert.match(code, /self\._delete\(\s*node\.right\s*,\s*node\.key\s*/);
```

`_delete()` の呼び出しの戻り値を現在のノードの右の子に代入してください。

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(\s*node\.right\s*,\s*node\.key/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right) 
            
--fcc-editable-region--
    def _min_value(self, node):
        while node.left is not None:
            node = node.left

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
