---
id: 5d5a813321b9e3db6c106a46
title: Hatua ya 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript ni lugha yenye nguvu inayokuwezesha kujenga tovuti zinazoshirikiana.

*Kumbuka*: Kwa miradi yote iliyobaki katika mtaala huu, utahitaji kiwango cha msingi cha maarifa ya HTML na CSS. Ikiwa wewe ni mpya kwa HTML na CSS, tafadhali pitia [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

Ili kuanza, tengeneza muundo wako wa msingi wa HTML unaojumuisha `DOCTYPE`, `html`, `head`, na `body`, kisha ongeza lebo ya `meta` kwa ajili ya `charset`. Ongeza kipengele cha `title` na tumia maandishi ya `RPG - Dragon Repeller` kwa ajili yake. Jumuisha lebo ya `link` kwa ajili ya karatasi yako ya mtindo ili kuunganisha faili la `styles.css`.

Mwishowe, tengeneza kipengele cha `div` chenye `id` imewekwa kuwa `game` ndani ya `body` yako.

# --hints--

Msimbo wako unapaswa kuwa na rejea ya `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Unapaswa kujumuisha nafasi baada ya rejea ya `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Unapaswa kufafanua aina ya hati kuwa `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Unapaswa kufunga tamko la `DOCTYPE` kwa `>` baada ya aina.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufungua. Usisahau sifa ya `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufunga.

```js
assert(code.match(/<\/html\s*>/));
```

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Unapaswa kuwa na lebo ya kufungua ya `head`.

```js
assert(code.match(/<head\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Unapaswa kuwa na lebo ya kufungua ya `body`.

```js
assert(code.match(/<body\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Vipengele vya `head` na `body` vinapaswa kuwa ndugu.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

Kipengele cha `head` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

Kipengele cha `body` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

Msimbo wako unapaswa kuwa na kipengele cha `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Kipengele chako cha `meta` kinapaswa kuwa na sifa ya `charset` yenye thamani ya `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Msimbo wako unapaswa kuwa na kipengele cha `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Unapaswa kuwa na lebo ya kufunga ya `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Unapaswa kuwa na kipengele cha `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Kipengele chako cha `div` kinapaswa kuwa na sifa ya `id` yenye thamani ya `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Kipengele chako cha `div` kinapaswa kuwa ndani ya kipengele cha `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
