---
id: 62a1166ed9a56d439c0770e7
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Tengeneza vipengele vitatu vya `span` ndani ya kipengele chako cha `#stats`. Wape kila kimoja darasa la vitu la `stat`. Kisha mpe kipengele cha kwanza maandishi ya `XP: 0`, kipengele cha pili maandishi ya `Health: 100`, na kipengele cha tatu maandishi ya `Gold: 50`.

# --hints--

Unapaswa kuwa na vipengele vitatu vya `span` ndani ya kipengele cha `#stats`.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Unapaswa kuwapa vipengele vipya vitatu vya `span` darasa la vitu la `stat`.

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

Kipengele chako cha kwanza cha `.stat` kinapaswa kuwa na maandishi yaliyotolewa ya `XP: 0`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

Kipengele chako cha pili cha `.stat` kinapaswa kuwa na maandishi yaliyotolewa ya `Health: 100`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

Kipengele chako cha tatu cha `.stat` kinapaswa kuwa na maandishi yaliyotolewa ya `Gold: 50`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
