---
id: 62a115879a6d51422652cbfc
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Katika mchezo wako wa kuigiza jukumu, watumizi wataweza kufuatilia takwimu zao, kununua silaha, na kupigana na monsters. Kabla hujaendelea na sehemu ya mwingiliano ya JavaScript ya mchezo, unahitaji kwanza kuunda vipengele vya HTML vitakavyoonyesha taarifa za mchezo.

Unda vipengele vinne vya `div` ndani ya kipengele chako cha `#game`. Vipe thamani zifuatazo za `id` kwa mpangilio: `stats`, `controls`, `monsterStats`, na `text`.

# --hints--

Unapaswa kuunda vipengele vipya vinne vya `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

Unapaswa kumpa mojawapo ya vipengele vipya vya `div` `id` ya `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

Unapaswa kumpa mojawapo ya vipengele vipya vya `div` `id` ya `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

Unapaswa kumpa mojawapo ya vipengele vipya vya `div` `id` ya `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

Unapaswa kumpa mojawapo ya vipengele vipya vya `div` `id` ya `text`.

```js
assert.exists(document.querySelector('div#text'));
```

Unapaswa kuweka vipengele vipya vya `div` kwa mpangilio sahihi.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

Unapaswa kuweka vipengele vipya vya `div` ndani ya kipengele cha `#game`.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
