---
id: 62a23cb9bc467a147516b500
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Kwa kipengele chako cha `#controls`, tengeneza vipengele vitatu vya `button`. Kwanza iwe na `id` iliyowekwa kuwa `button1`, na maandishi ya `Go to store`. Pili iwe na `id` iliyowekwa kuwa `button2`, na maandishi ya `Go to cave`. Tatu iwe na `id` iliyowekwa kuwa `button3`, na maandishi ya `Fight dragon`.

# --hints--

Unapaswa kuongeza vipengele vitatu vya `button` kwenye kipengele chako cha `#controls`.

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

Kitufe chako cha kwanza kinapaswa kuwa na `id` iliyowekwa kuwa `button1`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

Kitufe chako cha kwanza kinapaswa kuwa na maandishi ya `Go to store`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

Kitufe chako cha pili kinapaswa kuwa na `id` iliyowekwa kuwa `button2`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

Kitufe chako cha pili kinapaswa kuwa na maandishi ya `Go to cave`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

Kitufe chako cha tatu kinapaswa kuwa na `id` iliyowekwa kuwa `button3`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

Kitufe chako cha tatu kinapaswa kuwa na maandishi ya `Fight dragon`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
--fcc-editable-region--
      <div id="controls">

      </div>
--fcc-editable-region--
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
