---
id: 62a257659d0d1e2456f24ba2
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Kabla hujaanza kuandika msimbo wa mradi wako, unapaswa kuuhamisha katika faili yake mwenyewe ili kuweka mambo yakiwa yamepangwa vizuri.

Ondoa mstari wako wa `console.log("Hello World");`. Kisha mpe kipengele chako cha `script` ambacho sasa ni tupu sifa ya kipengele ya `src` iliyowekwa kuwa `./script.js`.

# --hints--

Haupasi kuwa na mstari wa `console.log("Hello World");` katika msimbo wako.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Kipengele chako cha `script` kinapaswa kuwa na sifa ya kipengele ya `src` iliyowekwa kuwa `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
