---
id: 62a2509ba163e020bb9d84ea
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Sasa unaweza kuanza kuandika JavaScript yako. Anza kwa kuunda kipengele cha `script`. Kipengele hiki hutumika kupakia JavaScript katika faili yako ya HTML.

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

Unapaswa kuwa na kipengele cha `script`.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 1);
```

Kipengele chako cha `script` kinapaswa kuwa na lebo ya kufungua.

```js
assert.match(code, /<script\s*>/i);
```

Kipengele chako cha `script` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
