---
id: 62a23c1d505bfa13747c8a9b
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Panga nambari `0`, `100`, na `50` ndani ya vipengele vya `span`, kisha panga vipengele vipya vya `span` ndani ya vipengele vya `strong`. Kisha toa vipengele vipya vya `span` thamani za `id` za `xpText`, `healthText`, na `goldText`, kwa mpangilio huo huo.

Jibu lako linapaswa kufuata muundo huu wa msingi:

```html
<span class="stat">TEXT <strong><span id="VALUE">TEXT</span></strong></span>
```

# --hints--

Unapaswa kuongeza kipengele cha `strong` ndani ya kipengele chako cha kwanza cha `.stat`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

Kipengele chako kipya cha kwanza cha `strong` kinapaswa kuwa na kipengele cha `span`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

Kipengele chako cha kwanza kilichopangwa ndani cha `span` kinapaswa kuwa na `id` ya `xpText`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'xpText');
```

Kipengele chako cha kwanza cha `span` kinapaswa kuzunguka maandishi ya `0`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '0');
```

Maandishi ya kipengele chako cha kwanza cha `.stat` yanapaswa kubaki kuwa `XP: 0`.

```js
const stat = document.querySelectorAll('.stat')[0];
assert.equal(stat.innerText.trimEnd(), 'XP: 0');
```

Unapaswa kuongeza kipengele cha `strong` ndani ya kipengele chako cha pili cha `.stat`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

Kipengele chako kipya cha pili cha `strong` kinapaswa kuwa na kipengele cha `span`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

Kipengele chako cha pili kilichopangwa ndani cha `span` kinapaswa kuwa na `id` ya `healthText`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'healthText');
```

Kipengele chako cha pili cha `span` kinapaswa kuzunguka maandishi ya `100`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '100');
```

Maandishi ya kipengele chako cha pili cha `.stat` yanapaswa kubaki kuwa `Health: 100`.

```js
const stat = document.querySelectorAll('.stat')[1];
assert.equal(stat.innerText.trimEnd(), 'Health: 100');
```

Unapaswa kuongeza kipengele cha `strong` ndani ya kipengele chako cha tatu cha `.stat`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

Kipengele chako kipya cha tatu cha `strong` kinapaswa kuwa na kipengele cha `span`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

Kipengele chako cha tatu kilichopangwa ndani cha `span` kinapaswa kuwa na `id` ya `goldText`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'goldText');
```

Kipengele chako cha tatu cha `span` kinapaswa kuzunguka maandishi ya `50`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '50');
```

Maandishi ya kipengele chako cha tatu cha `.stat` yanapaswa kubaki kuwa `Gold: 50`.

```js
const stat = document.querySelectorAll('.stat')[2];
assert.equal(stat.innerText.trimEnd(), 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
--fcc-editable-region--
        <span class="stat">XP: 0</span>
        <span class="stat">Health: 100</span>
        <span class="stat">Gold: 50</span>
--fcc-editable-region--
      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
