---
id: 62a255dae245b52317da824a
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Moja ya zana zenye nguvu zaidi ni konsoli yako ya msanidi programu. Kulingana na kivinjari chako, hii inaweza kufunguliwa kwa kubonyeza `F12` au `Ctrl+Shift+I`. Kwenye Mac, unaweza kubonyeza `Option + ⌘ + C` na kuchagua "Console". Pia unaweza kubofya kitufe cha "Console" kilicho juu ya dirisha la onyesho kuona konsoli yetu iliyojengwa ndani.

Konsoli ya msanidi programu itajumuisha makosa yanayotokana na msimbo wako, lakini pia unaweza kuitumia kuona thamani za vigezo katika msimbo wako, jambo ambalo ni msaada kwa kurekebisha makosa.

Ongeza mstari wa `console.log("Hello World");` kati ya lebo zako za `script`. Kisha bofya kitufe cha "Console" kufungua konsoli. Unapaswa kuona maandishi ya `"Hello World"`.


# --hints--

Unapaswa kuongeza mstari wa `console.log("Hello World");` kwenye msimbo wako.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

Mstari wako wa `console.log("Hello World");` unapaswa kuwa kati ya lebo zako za `script`.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
