---
id: 62a240c67f3dbb1a1e6d95ee
title: Hatua 23
challengeType: 0
dashedName: step-23
---

# --description--

Toa `#game` yako upana wa juu kabisa wa `500px` na urefu wa juu kabisa wa `400px`. Weka `background-color` kuwa `#ffffff` na `color` kuwa `#ffffff`. 

Tumia vikingo kuikatia katikati kwa kuweka ukingo wa juu kuwa `30px`, ukingo wa chini kuwa `0px`, na vikingo vya kushoto na kulia kuwa `auto`. 

Mwishowe, mpe nafasi ya ndani ya `10px` pande zote nne.

# --hints--

Unapaswa kuwa na kichaguzi cha `#game`.

```js
const game = new __helpers.CSSHelp(document).getStyle('#game');
assert.exists(game);
```

Kichaguzi chako cha `#game` kinapaswa kuwa na `max-width` ya `500px`.

```js
const maxWidth = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-width');
assert.equal(maxWidth, '500px');
```

Kichaguzi chako cha `#game` kinapaswa kuwa na `max-height` ya `400px`.

```js
const maxHeight = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-height');
assert.equal(maxHeight, '400px');
```

Kichaguzi chako cha `#game` kinapaswa kuwa na `background-color` ya `#ffffff`.

```js
const background = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('background-color');
assert.equal(background, 'rgb(255, 255, 255)');
```

Kichaguzi chako cha `#game` kinapaswa kuwa na `color` ya `#ffffff`.

```js
const color = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('color');
assert.equal(color, 'rgb(255, 255, 255)');
```

Kichaguzi chako cha `#game` kinapaswa kuwa na `margin-top` iliyowekwa kuwa `30px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-top');
assert.equal(margin, '30px');
```

Kichaguzi chako cha `#game` kinapaswa kuwa na `margin-left` iliyowekwa kuwa `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-left');
assert.equal(margin, 'auto');
```

Kichaguzi chako cha `#game` kinapaswa kuwa na `margin-right` iliyowekwa kuwa `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-right');
assert.equal(margin, 'auto');
```


Kichaguzi chako cha `#game` kinapaswa kuwa na `margin-bottom` iliyowekwa kuwa `0px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-bottom');
assert.equal(margin, '0px');
```

Kichaguzi chako cha `#game` kinapaswa kuwa na nafasi ya ndani ya `10px` pande zote.

```js
const padding = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('padding');
assert.equal(padding, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
