---
id: 62a3a7e4f1060e2fc5ffb34b
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Tengeneza kigezo kingine kinachoitwa `currentWeaponIndex` na uweke thamani yake kuwa `0`.

# --hints--

Unapaswa kutumia `let` kutangaza kigezo kinachoitwa `currentWeaponIndex`.

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

Unapaswa kutumia camelCase kwa jina la kigezo chako.

```js
assert.match(code, /currentWeaponIndex/);
```

Kigezo chako cha `currentWeaponIndex` kinapaswa kuwekwa kuwa `0`.

```js
assert.equal(currentWeaponIndex, 0);
```

Unapaswa kuanzisha kigezo chako kwa `0`.

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
