---
id: 62a39f5a5790eb27c1e5d4bf
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Muonekano wako umebadilishwa kwenda kwenye faili yako mpya ya `script.js`. Kumbuka unaweza kutumia tabo zilizo juu kubadilisha kati ya mafaili.

Ongeza mstari wako wa `console.log("Hello World");` kwenye faili hii, na uone ukionekana kwenye konsoli yako.

# --hints--

Unapaswa kuwa na mstari wa `console.log("Hello World");` katika msimbo wako.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
