---
id: 62a3b79d520a7f3d0e25afd6
title: Hatua 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript huingiliana na HTML kwa kutumia <dfn>Document Object Model</dfn>, au DOM. DOM ni mti wa vitu unaowakilisha HTML. Unaweza kupata upatikanaji wa HTML kwa kutumia kitu cha `document`, ambacho kinakilisha hati yako yote ya HTML.

Njia moja ya kupata vipengele maalum katika HTML yako ni kutumia njia ya `querySelector()`. Njia ya `querySelector()` huchukua kichaguzi cha CSS kama hoja na hurudisha kipengele cha kwanza kinacholingana na kichaguzi hicho. Kwa mfano, kupata kipengele cha `<h1>` katika HTML yako, utaandika:

```js
let h1 = document.querySelector("h1");
```

Kumbuka kuwa `h1` ni mfuatano wa herufi na unalingana na kichaguzi cha CSS unachotumia.

Tengeneza kigezo cha `button1` na tumia `querySelector()` kuweka thamani ya kipengele chako chenye `id` ya `button1`. Kumbuka kuwa vichaguzi vya CSS vya `id` huanza na kiambishi awali cha `#`.

# --hints--

Unapaswa kutumia `let` kutangaza kigezo cha `button1`.

```js
assert.match(code, /let\s+button1/);
```

Unapaswa kutumia `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

Unapaswa kutumia kichaguzi cha `"#button1"`.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
