---
id: 62a3b3eab50e193608c19fc6
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Tangaza kigezo kinachoitwa `fighting` lakini usikizishe kwa thamani yoyote.

# --hints--

Unapaswa kutumia `let` kutangaza kigezo `fighting`.

```js
assert.match(code, /let\s+fighting/);
```

Kigezo chako cha `fighting` hakipaswi kuwa na thamani.

```js
assert.isUndefined(fighting);
```

Hupaswi kuweka thamani kwa kigezo chako cha `fighting`.

```js
assert.match(code, /let\s+fighting\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
--fcc-editable-region--
```
