---
id: 62a3a0a3c0a4b32915d26a6e
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Ondoa mstari wako wa `console.log("Hello World");` kuanza kuandika msimbo wa mradi wako.

Katika mradi uliopita, ulijifunza jinsi ya kufanya kazi na vigezo na neno kuu la `let` kama hii:

```js
let camperbot;
```

Pia ulijifunza jinsi ya kuanzisha kigezo na thamani kama hii:

```js
let age = 32;
```

Tumia neno kuu la `let` kutangaza kigezo kinachoitwa `xp` na kweke thamani ya `0`, nambari.

# --hints--

Haupaswi kuwa na mstari wa `console.log("Hello World");` katika msimbo wako.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Unapaswa kutumia neno kuu la `let` kutangaza kigezo chako.

```js
assert.match(code, /let/);
```

Unapaswa kutangaza kigezo kinachoitwa `xp`.

```js
assert.match(code, /let\s+xp/);
```

`xp` inapaswa kuwa na thamani ya `0`.

```js
assert.equal(xp, 0);
```

Unapaswa kuanzisha kigezo cha `xp` kwa `0`. 

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
