---
id: 62a3a75d8466a12e009eff76
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Anzisha kigezo kingine kinachoitwa `health` na thamani ya `100`, na kigezo kinachoitwa `gold` na thamani ya `50`.

# --hints--

Unapaswa kutumia `let` kutangaza kigezo kinachoitwa `health`.

```js
assert.match(code, /let\s+health/);
```

Unapaswa kuanzisha kigezo kinachoitwa `health` na thamani ya `100`.

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

Unapaswa kutumia `let` kutangaza kigezo kinachoitwa `gold`.

```js
assert.match(code, /let\s+gold/);
```

Unapaswa kuanzisha kigezo kinachoitwa `gold` na thamani ya `50`.

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` inapaswa kuwa na thamani ya `100`.

```js
assert.equal(health, 100);
```

`gold` inapaswa kuwa na thamani ya `50`.

```js
assert.equal(gold, 50);
```

`xp` bado inapaswa kuwa na thamani ya `0`.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
