---
id: 62a3b41c9494f937560640ab
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Tangaza vigezo viwili zaidi vinavyoitwa `monsterHealth` na `inventory`.

Kwa kigezo chako cha `inventory`, weka thamani yake kuwa safu ya data yenye mfuatano wa herufi `"stick"`.

Kumbuka kwamba ulifanya kazi na safu za data katika mradi uliopita hivi:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

Unapaswa kutumia `let` kutangaza kigezo chako cha `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Unapaswa kutumia camelCase kwa jina la kigezo chako cha `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` haipaswi kuwa na thamani.

```js
assert.isUndefined(monsterHealth);
```

Hupaswi kuweka thamani kwa kigezo chako cha `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

Unapaswa kutumia `let` kutangaza kigezo chako cha `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

Kigezo chako cha `inventory` kinapaswa bado kuwa safu ya data.

```js
assert.isArray(inventory);
```

Kigezo chako cha `inventory` kinapaswa kuwa na thamani moja tu.

```js
assert.lengthOf(inventory, 1);
```

Kigezo chako cha `inventory` kinapaswa kujumuisha mfuatano wa herufi `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
