---
id: 62b46e3a8d4be31be5af793d
title: Hatua 16
challengeType: 0
dashedName: step-16
---

# --description--

Tumekutana na tatizo dogo. Unajaribu kufanya hoja ya utafutaji kwenye ukurasa wako kwa ajili ya kipengele cha kitufe, lakini lebo yako ya `script` iko ndani ya `head` ya HTML yako. Hii inamaanisha msimbo wako unatekelezwa kabla ya kivinjari kumaliza kusoma HTML, na `document.querySelector()` yako haitagundua kitufe hicho - kwa sababu kivinjari hakijakamilisha kukichakata bado.

Ili kurekebisha hili, hamisha kipengele chako cha `script` kutoka ndani ya kipengele cha `head`, na uweke mwishoni mwa kipengele chako cha `body` (kabla tu ya lebo ya kufunga ya `</body>`).

# --hints--

Kipengele chako cha `script` hakipaswi kuwa ndani ya kipengele chako cha `head`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

Kipengele chako cha `script` kinapaswa kuja mwishoni mwa kipengele chako cha `body`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
