---
id: 587d8252367417b2b2512c67
title: 在链表中的特定索引处添加元素
challengeType: 1
forumTopicId: 301619
dashedName: add-elements-at-a-specific-index-in-a-linked-list
---

# --description--

让我们创建一个 addAt(index,element) 方法，在给定的索引处添加一个元素。 就像我们如何删除给定索引处的元素一样，我们需要在遍历链表时跟踪 currentIndex。 当 currentIndex 与给定索引匹配时，我们需要重新分配上一个节点的下一个属性以引用新添加的节点。 并且新节点应该引用 currentIndex 中的下一个节点。 回到康加线的例子，一个新人想加入这条线，但他想加入中间。 你处于中间位置，所以你要把手从你前面的人身上移开。 新人走过去，把手放在你曾经牵过手的那个人身上，现在你已经掌握了新人。

# --instructions--

创建一个 `addAt(index,element)` 方法，在指定索引处添加元素。如果无法添加元素，则返回 false。**注意：** 记得查看给定的索引是否为负数或超过链表的长度。

# --hints--

当给定的索引为 0 时，你的 `addAt` 方法应该将 `head` 重新分配给新节点。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'fish');
    return test.head().element === 'fish' && test.head().next.element === 'cat';
  })()
);
```

对于添加到链表的每个新节点，你的 `addAt` 方法应该将链表的长度增加 1。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'cat');
    return test.size() === 3;
  })()
);
```

如果无法添加节点，你的 `addAt` 方法应返回 `false`。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.addAt(4, 'cat') === false;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element) {
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function() {
    return head;
  };

  this.add = function(element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if (head === null){
        head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };
  this.addAt = function (index, element) {
    if (index > length || index < 0) {
      return false;
    }
    var newNode = new Node(element);
    var currentNode = head;
    if (index === 0) {
      head = newNode;
    } else {
      var previousNode = null;
      var i = 0;
      while (currentNode && i < index) {
        previousNode = currentNode;
        currentNode = currentNode.next;
        i++;
      }
      previousNode.next = newNode;
    }
    newNode.next = currentNode;
    length++;
  }
}
```
