---
id: 587d8254367417b2b2512c70
title: 在ES6中创建和添加集
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--

现在你已经完成了 ES5 的工作，你将在 ES6 中进行类似的工作。 这将相当容易。 ES6包含一个内置的数据结构 < code > Set </code > ，因此现在包含了许多手工编写的操作。 让我们来看以下的例子：

创建一个新的空集:

```js
var set = new Set();
```

您可以创建一个具有值的集合:

```js
var set = new Set(1);
```

您可以使用数组创建一个集合:

```js
var set = new Set([1, 2, 3]);
```

创建集合后，你可以使用 `add` 方法添加所需的值：

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

提醒一下，集合是一种不能包含重复值的数据结构：

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set contains [1, 2, 3] only
```

# --instructions--

在这个练习中，返回一个具有以下值的集合：`1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--

你的 `Set` 应该只包含值 `1, 2, 3, Taco, Cat, Awesome`。

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```
