---
id: 587d8255367417b2b2512c72
title: 在 ES6 集合上使用 .has 和 .size
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

Let's look at the .has and .size methods available on the ES6 Set object.

首先，创建一个 ES6 集合

```js
var set = new Set([1,2,3]);
```

.has 方法将检查该值是否包含在集合中。

```js
var hasTwo = set.has(2);
```

.size 方法将返回一个表示集合大小的整数

```js
var howBig = set.size;
```

# --instructions--

在本练习中，我们将向 checkSet() 函数传递一个数组和一个值。 您的函数应该从数组参数创建一个 ES6 集合。 查找该集合是否包含 value 参数。 求集合的大小。 并在数组中返回这两个值。

# --hints--

`checkSet([4, 5, 6], 3)` 应该返回 [ false, 3 ]

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
