---
id: 587d825b367417b2b2512c8d
title: 创建一个 ES6 JavaScript 映射
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

新版本的 JavaScript 为我们提供了一个内置的 Map 对象，它提供了我们在上次挑战中手工编写的许多功能。 这个 Map 对象虽然类似于普通的 JavaScript 对象，但是提供了一些普通对象所不具备的有用功能。 例如，ES6 Map 跟踪添加到其中的项的插入顺序。 以下是其方法的一个更完整的概述: 

- `.has(key)` 基于输入的key值返回 true 或 false
- `.get(key)` 返回与key关联的值
- `.set(key, value)` 设置一个新的键值对
- `.delete(key)` 移除一个键值对
- `.clear()` 移除所有键值对
- `.keys()` 返回一个包含所有密钥且按插入顺序排列的新建 map 迭代器对象`。
- `.values()` 返回一个包含所有值且按插入顺序排列的新建 map 迭代器对象
- `.entries()` 返回一个新的 map 迭代器对象，包含所有按插入顺序排列的密钥、值点对。

# --instructions--

定义一个 JavaScript 映射对象并将其分配给一个名为 myMap 的变量。 向其中添加键值对 `freeCodeCamp`、`Awesome!`。

# --hints--

`myMap` 对象应该存在。

```js
assert(typeof myMap === 'object');
```

< code > myMap </code > 应该包含键值对 < code > freeCodeCamp </code > ，< code > Awesome! </code > 。

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
