---
id: 587d825c367417b2b2512c90
title: 广度优先搜索
challengeType: 1
forumTopicId: 301622
dashedName: breadth-first-search
---

# --description--

到目前为止，我们已经学习了创建图形表示的不同方法。 现在要做什么？ 一个自然的问题是图中任何两个节点之间的距离是多少？ 进入 <dfn>图遍历算法</dfn> 。

<dfn>遍历算法</dfn> 是遍历或访问图中节点的算法。 一种类型的遍历算法是广度优先的搜索算法。

该算法从一个节点开始，并访问它的每个距离为 1 的邻居。 然后算法接着走访它们的每个邻居，如此，直到所有节点都已经到达为止。

一个有助于实现广度优先搜索算法的重要数据结构是队列。 这是一个数组，你可以添加元素到数组的一端并从另一端删除元素。 这也称为 <dfn>FIFO</dfn> 或 <dfn>先进先出</dfn> 数据结构。

从视觉上看，这就是算法正在做的事情：  
<img alt="animation showing the breadth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/breadth-first-search.gif'>

灰色阴影代表一个添加到队列中的节点，黑色阴影代表一个从队列中删除的节点。 注意每当某个节点从队列中被移除时（节点转为黑色），它们的所有邻居都会被添加到队列中（节点转为灰色）。

要实现此算法，你需要输入图形结构和你想要开始的节点。

首先，您需要了解距起始节点的距离。 在开始的时候，你想要给所有的距离一个很大的数字，例如 `Infinity`。 这将防止当某个节点无法从你的起始节点到达时出现计算问题。 接下来，您将要从开始节点转到其邻居。 这些邻居是一个边缘的距离，此时你应该添加一个距离单位到你要跟踪的距离。

# --instructions--

编写一个函数 `bfs()` ，它将一个邻接矩阵图（二维数组）和一个标记为根的节点标签作为参数。 节点标签只是 `0` 到 `n - 1` 之间节点的整数值，其中 `n` 是图中节点的总数。

你的函数将输出 JavaScript 对象键值对，即节点及其与根的距离。 如果无法到达节点，则其距离应为`Infinity` 。

# --hints--

输入图 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`，起始节点为 `1`，应该返回 `{0: 1, 1: 0, 2: 1, 3: 2}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 1);
assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: 2 });
```

输入图 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]`，起始节点为 `1`，应该返回 `{0: 1, 1: 0, 2: 1, 3: Infinity}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 0],
  [0, 0, 0, 0]
];
var results = bfs(graph, 1);
 assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: Infinity });
```

输入图 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`，起始节点为 `0`，应该返回 `{0: 0, 1: 1, 2: 2, 3: 3}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 0);
 assert.deepEqual(results, { 0: 0, 1: 1, 2: 2, 3: 3 });
```

输入图 `[[0, 1], [1, 0]]`，其实节点为 `0`，应该返回 `{0: 0, 1: 1}`

```js
var graph = [
  [0, 1],
  [1, 0]
];
var results = bfs(graph, 0);
assert.deepEqual(results, { 0: 0, 1: 1 });
```

# --seed--

## --seed-contents--

```js
function bfs(graph, root) {
  var nodesLen = {};

  return nodesLen;
};

var exBFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(bfs(exBFSGraph, 3));
```

# --solutions--

```js
function bfs(graph, root) {
  var nodesLen = {};
  // Set all distances to infinity
  for (var i = 0; i < graph.length; i++) {
    nodesLen[i] = Infinity;
  }
  nodesLen[root] = 0; // ...except root node
  var queue = [root]; // Keep track of nodes to visit
  var current; // Current node traversing
  // Keep on going until no more nodes to traverse
  while (queue.length !== 0) {
    current = queue.shift();
    // Get adjacent nodes from current node
    var curConnected = graph[current]; // Get layer of edges from current
    var neighborIdx = []; // List of nodes with edges
    var idx = curConnected.indexOf(1); // Get first edge connection
    while (idx !== -1) {
      neighborIdx.push(idx); // Add to list of neighbors
      idx = curConnected.indexOf(1, idx + 1); // Keep on searching
    }
    // Loop through neighbors and get lengths
    for (var j = 0; j < neighborIdx.length; j++) {
      // Increment distance for nodes traversed
      if (nodesLen[neighborIdx[j]] === Infinity) {
        nodesLen[neighborIdx[j]] = nodesLen[current] + 1;
        queue.push(neighborIdx[j]); // Add new neighbors to queue
      }
    }
  }
  return nodesLen;
}
```
