---
id: 69860157f7fce39f9df399a3
title: Реалізувати алгоритм N-ферзів
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

Задача N-ферзів полягає в тому, щоб розмістити N ферзів на шахівниці розміром N×N так, щоб жоден ферзь не атакував іншого (жодні два не знаходилися в одному рядку, стовпці або діагоналі).

Наприклад, якщо є дошка 4x4, одним із правильних розташувань є:

```md
[1, 3, 0, 2]
```

Це означає, що в рядку 0 ферзь розміщений у стовпці 1; у рядку 1 ферзь розміщений у стовпці 3; у рядку 2 ферзь розміщений у стовпці 0; а в рядку 3 ферзь розміщений у стовпці 2.

Візуально це розташування виглядає так:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Де `Q` позначає ферзя, а `.` — порожню клітинку.

У цій лабораторній роботі ви реалізуєте розв’язувач задачі N-ферзів, використовуючи підхід пошуку у глибину.


**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з назвою `dfsNQueens`.
2. Функція має приймати рівно один аргумент: ціле число `n`.
3. Якщо `n` менше за `1`, функція має повертати порожній масив (`[]`).
4. Функція має повертати список розв’язків; кожен розв’язок — це список довжини `n`, де елемент з індексом `i` — це індекс стовпця (починаючи з 0) ферзя в рядку `i`.

# --hints--

У вас має бути функція з назвою `dfsNQueens`, яка приймає один аргумент.

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

Якщо `n` менше за `1`, функція має повертати порожній масив.

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

Функція має повертати масив розв’язків, де кожен розв’язок — масив довжини `n`.

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` має повертати `[[0]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` має повертати `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` має повертати `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` має повертати `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` має повертати `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` має бути `10`.

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` має бути `92`.

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens` має повертати правильний результат.

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
