---
id: 684aaf9ec670c68d20efd0d0
title: 建置使用者設定管理員
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

在本實驗中，你將建置一個使用者設定管理員，允許使用者管理他們的設定，例如主題、語言和通知。你將實作函式來新增、更新、刪除和檢視使用者設定。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該定義一個名為 `add_setting` 的函式（程式），該函式帶有兩個參數，分別表述一個設定字典和一個包含鍵值配對的元組。

1. `add_setting` 函式（程式）應該：
    
   - 將 key 和 value 轉換為小寫。
   - 如果該鍵設定存在，傳回 `Setting '[key]' already exists! Cannot add a new setting with this name.`
   - 如果該鍵設定不存在，為給定的設定字典添加該鍵值配對，並傳回 `Setting '[key]' added with value '[value]' successfully!`。
   - 傳回的訊息應該具有小寫的鍵和值。

1. 你應該定義一個名為 `update_setting` 的函式（程式），該函式有兩個參數，分別表述一個設定字典和一個包含鍵值配對的元組。

1. `update_setting` 函式（程式）應該：
       
   - 將 key 和 value 轉換為小寫。
   - 如果該鍵設定存在，則在指定的設定字典中更新其值並傳回：`Setting '[key]' updated to '[value]' successfully!`
   - 如果鍵設定不存在，傳回 `Setting '[key]' does not exist! Cannot update a non-existing setting.`
   - 傳回的訊息應該具有小寫的鍵和值。


1. 你應該定義一個名為 `delete_setting` 的函式（程式），該函式帶有兩個參數，分別表述一個設定字典和一個鍵。

1. `delete_setting` 函式（程式）應該：
    
   - 將傳入的鍵轉換為小寫。
   - 如果該鍵設定存在，則從給定的設定字典中移除該鍵值配對，並傳回 `Setting '[key]' deleted successfully!`
   - 如果鍵設定不存在，傳回 `Setting not found!`
   - 傳回的訊息應該具有小寫的鍵。

1. 你應該定義一個名為 `view_settings` 的函式（程式），並帶有一個表述設定字典的參數。

1. `view_settings` 函式（程式）應該：
   
   - 如果給定的設定字典為空，傳回 `No settings available.`。
   - 如果字典包含任何設定，傳回一個顯示設定的字串。該字串應以 `Current User Settings:` 開頭，接著是鍵值對，每個鍵值對各佔一行且鍵字首字母大寫。為你的例子，`view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` 應該傳回：

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. 為測試程式碼，你應該創建一個名為 `test_settings` 的字典來儲存一些使用者設定偏好。

# --hints--

你應該創建一個名為 `test_settings` 的字典，並為它添加一些值。

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

你應該定義一個名為 `add_setting` 的函式（程式）。

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

`add_setting` 函式應該有兩個參數。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` 應該將鍵轉換為小寫。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting` 應該將值轉換為小寫。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` 應該傳回錯誤訊息 `Setting 'theme' already exists! Cannot add a new setting with this name.`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` 應該會新增一個新的鍵值對組並傳回成功的訊息 `Setting 'volume' added with value 'high' successfully!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting` 應該正確地將給定的鍵值配對添加到字典中。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

你應該定義一個名為 `update_setting` 的函式（程式）。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

`update_setting` 函式（程式）應該有兩個參數。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`update_setting` 函式（程式）應該將 key 轉換為小寫。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting` 函式（程式）應該將 value 轉換為小寫。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` 應該更新現有的鍵並傳回成功訊息 `Setting 'theme' updated to 'dark' successfully!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` 應該在鍵不存在時傳回錯誤訊息 `Setting 'volume' does not exist! Cannot update a non-existing setting.`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting` 應該正確更新字典中給定的鍵值配對。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

你應該定義一個名為 `delete_setting` 的函式（程式）。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

`delete_setting` 函式應該有兩個參數。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` 應該將鍵轉換為小寫。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` 應該移除現有的鍵並傳回成功的訊息 `Setting 'theme' deleted successfully!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

當鍵不存在時，`delete_setting` 應該傳回錯誤訊息 `Setting not found!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting` 應該正確地從字典中移除指定的鍵。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

你應該定義一個名為 `view_settings` 的函式（程式）。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

`view_settings` 函式（程式）應該有一個參數。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` 應該在給定的字典為空時傳回訊息 `No settings available.`。

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` 應該傳回格式化後的設定，針對非空字典。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` 應該將每個設定名稱的第一個字母大寫。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings` 應該顯示正確的結果並以換行字元結尾。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: English\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "English",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
