---
id: 68de7360e1832ea180577a8a
title: بناء حاسبة للعدد النوني في متتالية فيبوناتشي
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب إنشاء دالة باسم `fibonacci`.
2. يجب تعريف قائمة باسم `sequence` داخل دالة `fibonacci`، ويجب تهيئتها بالقيم `[0, 1]`.
3. يجب أن تقبل دالة `fibonacci` معلمة واحدة، وهي عدد صحيح غير سالب `n`.
4. يجب أن يؤدي استدعاء `fibonacci(n)` إلى استخدام أسلوب البرمجة الديناميكية لحساب وإرجاع العدد رقم `n` من متتالية فيبوناتشي، حيث كل عدد هو مجموع العددين السابقين له.
5. يجب إضافة كل عدد فيبوناتشي يتم حسابه إلى قائمة `sequence`.

# --hints--

يجب أن يكون لديك دالة باسم `fibonacci`.

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

يجب أن تأخذ دالة `fibonacci` معلمة واحدة فقط.

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

يجب أن يكون لديك قائمة باسم `sequence` داخل دالة `fibonacci` مهيأة إلى `[0, 1]`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

يجب أن تُرجع `fibonacci(0)` القيمة `0`.

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

يجب أن تُرجع `fibonacci(1)` القيمة `1`.

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

يجب أن تُرجع `fibonacci(2)` القيمة `1`.

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

يجب أن تُرجع `fibonacci(3)` القيمة `2`.

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

يجب أن تُرجع `fibonacci(5)` القيمة `5`.

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

يجب أن تُرجع `fibonacci(10)` القيمة `55`.

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

يجب أن تُرجع `fibonacci(15)` القيمة `610`.

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

يجب ألا تستخدم الاستدعاء الذاتي في شفرتك.

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```

