---
id: 66a9ca099e3ecb1be10a2696
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

现在是时候给表格再增加两名学生了。

按照与前面步骤相同的模式，添加第四个学生表格行。 表格数据元素使用以下数据：

- 姓氏：`Thompson`
- 名字：`Jane`
- 分数：`77`

对于第五个学生表格行，表格数据元素使用以下数据：

- 姓氏：`Williams`
- 名字：`Natalie`
- 分数：`83`

# --hints--

你应该在 `tbody` 元素中添加第四个 `tr`。

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

你应该在第四个 `tr` 元素中包含一个文本为 `Thompson` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

你应该在第四个 `tr` 元素中包含一个文本为 `Jane` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

你应该在第四个 `tr` 元素中包含一个文本为 `77` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

你应该在 `tbody` 元素中添加第五个 `tr`。

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

你应该在第五个 `tr` 元素中包含一个文本为 `Williams` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

你应该在第五个 `tr` 元素中包含一个文本为 `Natalie` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

你应该在第五个 `tr` 元素中包含一个文本为 `83` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
