---
id: 66a9c006e4f58819396605b3
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

现在你已经完成了标题部分，是时候添加<dfn>表体</dfn> `tbody` 了。 表体将显示所有学生姓名及其成绩。

为表格添加表体元素。

# --hints--

你应该有一个 `tbody` 开始标签。

```js
assert.match(code, /<tbody>/i);
```

你应该有一个 `tbody` 结束标签。

```js
assert.match(code, /<\/tbody>/i);
```

你的 `tbody` 开始标签应在 `table` 结束标签之前。

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
