---
id: 66a9cb26ec6bd41cf6c82bc5
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

最后要添加到表格中的部分是<dfn>表脚</dfn>元素 `tfoot`。 表脚元素将用于显示表格中所有学生的平均成绩。

在表格中添加 `tfoot` 元素。

在 `tfoot` 元素中，添加一个 `tr` 元素。

在 `tr` 元素中，添加两个 `td` 元素。

第一个 `td` 元素应包含文本 `Average Grade`。 第二个 `td` 元素应包含文本 `78.8`。

# --hints--

你应该有一个 `tfoot` 开始标签。

```js
assert.match(code, /<tfoot>/);
```

你应该有一个 `tfoot` 结束标签。

```js
assert.match(code, /<\/tfoot>/);
```

你应该在 `tfoot` 元素中包含一个 `tr` 元素。

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

你的第一个 `td` 元素应包含文本 `Average Grade`。 确保它位于 `tr` 元素内。

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

你的第二个 `td` 元素应包含文本 `78.8`。 确保它位于 `tr` 元素内。

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

你的 `tr` 元素应包含两个 `td` 元素。

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
