---
id: 66a9b48a211a73155621d0df
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

In previous lessons, you learned how to work with the `table` element to represent tabular data.

在 `body` 元素内，嵌套一个 `table` 元素。

# --hints--

你应该有一个 `table` 开始标签。

```js
assert.match(code, /<table>/i);
```

你应该有一个 `table` 结束标签。

```js
assert.match(code, /<\/table>/i);
```

你的 `table` 元素应该在 `body` 元素中。

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
