---
id: 66a9b8f14b963916a3baa732
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

要为 `table` 添加标题，你可以使用 <dfn>table caption</dfn> 元素。

下面是一个使用 `caption` 元素的示例：

```html
<table>
  <caption>Football Scores</caption>
</table>
```

在你的 `table` 元素中，嵌套一个 `caption` 元素，文本为 `Calculus Final Exam Grades`。

# --hints--

你应该有一个 `caption` 开始标签。

```js
assert.match(code, /<caption>/i);
```

你应该有一个 `caption` 结束标签。

```js
assert.match(code, /<\/caption>/i);
```

你的 `caption` 元素应该在 `table` 元素中。

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

你的 `caption` 应包含文本 `Calculus Final Exam Grades`。

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    --fcc-editable-region--
    <table>

    </table>
    --fcc-editable-region--
  </body>
</html>
```
