---
id: 66a98f42c7c06903e5f8dd07
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本工作坊，你将通过为一组学生制作期末考试表格来练习使用 HTML 表格。

开始项目，添加 `<!DOCTYPE html>`，添加一个 `html` 元素，其 `lang` 属性值为 `en`。

在 `html` 元素中，添加一个 `head` 元素。

# --hints--

你应该有 `<!DOCTYPE html>`。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你应该有一个 `html` 开始标签，语言设置为英语。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

你应该有一个 `html` 结束标签。

```js
assert.match(code, /<\/html>/i);
```

你的 `DOCTYPE` 应放在 `html` 元素之前。

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

你应该有一个 `head` 开始标签。

```js
assert.match(code, /<head>/i);
```

你应该有一个 `head` 结束标签。

```js
assert.match(code, /<\/head>/i);
```

你的 `head` 开始标签应在 `head` 结束标签之前。

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

你的 `head` 元素应该在 `html` 元素中。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
