---
id: 66a9bb8578459a175432e7d0
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

对于表格的第一部分，你需要对标题内容进行分组，标题内容代表学生的名、姓和期末考试成绩的列标签。

<dfn>表头</dfn>元素 `thead`，用于在表格中对标题内容进行分组。

下面是一个使用 `thead` 元素的示例：

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

在 `caption` 元素下方，添加一个表头元素。

# --hints--

你应该有一个 `thead` 开始标签。

```js
assert.match(code, /<thead>/i);
```

你应该有一个 `thead` 结束标签。

```js
assert.match(code, /<\/thead>/i);
```

你的 `thead` 元素应该在 `table` 元素中。

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
