---
id: 66a9bcf00f13a418368a272e
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

表头元素由<dfn>表格行</dfn>元素 `tr` 组成，其中包含<dfn>表头单元格</dfn>元素 `th`。

下面是一个使用 `tr` 和 `th` 元素的运动表格的示例：

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

在 `thead` 元素中，添加一个 `tr` 元素。

在你的 `tr` 元素中，添加三个 `th` 元素。

第一个 `th` 元素应包含文本 `Last Name`。 第二个 `th` 元素应包含文本 `First Name`。 第三个 `th` 元素应包含文本 `Grade`。

# --hints--

你应该有一个 `tr` 开始标签。

```js
assert.match(code, /<tr>/i);
```

你应该有一个 `tr` 结束标签。

```js
assert.match(code, /<\/tr>/i);
```

你的 `tr` 元素应该在 `thead` 元素中。

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

你的第一个 `th` 元素的文本应为 `Last Name`。

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

你的第二个 `th` 元素的文本应为 `First Name`。

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

你的第三个 `th` 元素的文本应为 `Grade`。

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

你所有的 `th` 元素都应该在 `tr` 元素中。

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      --fcc-editable-region--
      <thead>     

      </thead>
      --fcc-editable-region--
    </table>
  </body>
</html>
```
