---
id: 66a9b333487b9c14998539a5
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素内，嵌套一个 `meta` 元素，其 `charset` 属性的值设置为 `"UTF-8"`。

在该 `meta` 元素下方，添加一个 `title` 元素。

`title` 元素的文本应为 `Calculus Final Exams Table`。

完成 `head` 元素后，你可以添加 `body` 元素。

# --hints--

你应该有一个 `meta` 元素。

```js
assert.isNotNull(document.querySelector("meta"));
```

`meta` 元素是一个空元素，它不应该有一个结束标签 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 标签应该有一个 `charset` 属性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 属性的值应为 `"UTF-8"`。

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

你的 `meta` 元素应嵌套在 `head` 元素中。

```js
const meta = document.querySelector('head > meta');
assert.strictEqual(meta?.parentElement?.tagName, 'HEAD');
```

你应该有一个 `title` 开始标签。

```js
assert.match(code, /<title>/i);
```

你应该有一个 `title` 结束标签。

```js
assert.match(code, /<\/title>/i);
```

所有的 `title` 元素都应嵌套在 `head` 元素中。

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

你的 `title` 元素应包含文本 `Calculus Final Exams Table`。 你可能需要检查一下拼写。

```js
const titleText = document.querySelector('title')?.innerText.trim();
assert.strictEqual(titleText, "Calculus Final Exams Table");
```

你应该有一个 `body` 开始标签。

```js
assert.match(code, /<body>/i);
```

你应该有一个 `body` 结束标签。

```js
assert.match(code, /<\/body>/i);
```

你不应该改变你的 `head` 元素。 确认你没有删除结束标签。

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

你的 `body` 元素应该在你的 `head` 元素之后。

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
  --fcc-editable-region--
</html>
```
