---
id: 66a9c14e3b34c719e34bae20
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

要在表格中添加学生数据，需要使用表格行和表格数据元素。

<dfn>表格数据</dfn>元素 `td`，用于创建表格中的单元格。

下面是一个将 `td` 元素用于体育运动员表格的示例：

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

在表体元素中，添加一个表格行元素 `tr` 和三个表格数据元素 `td`。

第一个表格数据元素应包含姓氏 `Davis`。

第二个表格数据元素应包含名字 `Alex`。

第三个表格数据元素应包含分数 `54`。

# --hints--

你应该在 `tbody` 元素中包含一个 `tr` 元素。

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

你应该在 `tr` 元素中包含一个文本为 `Davis` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

你应该在 `tr` 元素中包含一个文本为 `Alex` 的 `td` 元素。

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

你应该在 `tr` 元素中包含一个文本为 `54` 的 `td` 元素。

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

你应将三个 `td` 元素都放在 `tr` 元素中。

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
