---
id: 66a9c7a4521d2b1b1ec6dcf0
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

现在是时候给表格再增加两名学生了。

按照与上一步相同的模式，添加第二行学生表。 表格数据元素使用以下数据：

- 姓氏：`Doe`
- 名字：`Samantha`
- 分数：`92`

对于第三个学生表格行，表格数据元素使用以下数据：

- 姓氏：`Rodriguez`
- 名字：`Marcus`
- 分数：`88`

# --hints--

你应该在 `tbody` 元素中添加第二个 `tr`。

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

你应该在第二个 `tr` 元素中包含一个文本为 `Doe` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

你应该在第二个 `tr` 元素中包含一个文本为 `Samantha` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

你应该在第二个 `tr` 元素中包含一个文本为 `92` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

你应该在 `tbody` 元素中添加第三个 `tr`。

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

你应该在第三个 `tr` 元素中包含一个文本为 `Rodriguez` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

你应该在第三个 `tr` 元素中包含一个文本为 `Marcus` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

你应该在第三个 `tr` 元素中包含一个文本为 `88` 的 `td` 元素。

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
