---
id: 63db7f4677d06d7500a13321
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 십진수와 이진수 변환기를 만들고 두 숫자 체계를 연습합니다. 또한 재귀를 사용해 변환을 수행합니다.

이 프로젝트에 필요한 모든 HTML과 CSS는 이미 제공되었습니다.

시작할 준비가 되면 `.getElementById()` 메서드를 사용해 id가 `"number-input"`인 `input` 요소를 가져와 `numberInput`이라는 `const` 변수에 저장하세요. 같은 메서드를 사용해 id가 `"convert-btn"`인 `button` 요소를 가져와 `convertBtn`이라는 `const` 변수에 저장하고, id가 `"result"`인 `output` 요소를 가져와 `result`라는 `const` 변수에 저장하세요.


**참고**: 이 프로젝트는 양수를 이진수로만 변환합니다.

# --hints--

`document.getElementById()` 메서드를 사용해 `#number-input` 요소를 가져와야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { numberInput } = explorer.variables;
assert.isTrue(numberInput?.value.matches('document.getElementById("number-input")'));
```

`const`를 사용해 `numberInput`라는 변수를 선언하고 `#number-input` 요소를 할당해야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { numberInput } = explorer.variables;
assert.isTrue(numberInput?.matches("const numberInput = document.getElementById('number-input')"));
```

`document.getElementById()` 메서드를 사용해 `#convert-btn` 요소를 가져와야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { convertBtn } = explorer.variables;
assert.isTrue(convertBtn?.value.matches('document.getElementById("convert-btn")'));
```

`const`를 사용해 `convertBtn`라는 변수를 선언하고 `#convert-btn` 요소를 할당해야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { convertBtn } = explorer.variables;
assert.isTrue(convertBtn?.matches("const convertBtn = document.getElementById('convert-btn')"));
```

`document.getElementById()` 메서드를 사용해 `#result` 요소를 가져와야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result?.value.matches('document.getElementById("result")'));
```

`const`를 사용해 `result`라는 변수를 선언하고 `#result` 요소를 할당해야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result?.matches("const result = document.getElementById('result')"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Decimal to Binary Converter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>
      Decimal to Binary <br />
      Converter
    </h1>
    <section class="input-container">
      <label for="number-input">Enter a decimal number:</label>
      <input
        value=""
        type="number"
        name="decimal number input"
        id="number-input"
        class="number-input"
      />
      <button class="convert-btn" id="convert-btn">Convert</button>
    </section>
    <section class="output-container">
      <output id="result" for="number-input"></output>
      <h2>Call stack</h2>
      <div id="animation-container"></div>
    </section>
    <script src="script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #f5f6f7;
  --dark-blue: #1b1b32;
  --orange: #f1be32;
}

body {
  display: flex;
  flex-direction: column;
  align-items: center;
  font-family: Menlo, Consolas, Monaco, Liberation Mono, Lucida Console,
    monospace;
  font-size: 1.125rem;
  color: var(--light-grey);
  background-color: var(--dark-blue);
  padding: 0 4px;
}

h1 {
  font-size: 2.125rem;
  text-align: center;
  margin: 20px 0;
}

h2 {
  font-size: 1.5rem;
  text-align: center;
  margin: 20px 0;
}

.input-container {
  display: flex;
  flex-direction: column;
  gap: 10px;
  justify-content: center;
  align-items: center;
  width: clamp(320px, 50vw, 460px);
  margin: 10px auto;
}

.input-container label {
  white-space: nowrap;
  word-spacing: -6px;
}

.convert-btn {
  font-size: inherit;
  font-family: inherit;
  background-color: var(--orange);
  width: 100%;
  height: 2rem;
  padding: 0 6px;
  border: none;
  border-radius: 2px;
  cursor: pointer;
}

.number-input {
  font-size: inherit;
  padding: 0.3rem;
  width: 100%;
}

.output-container {
  margin-inline: auto;
  width: clamp(320px, 50vw, 460px);
}

#result {
  display: flex;
  align-items: center;
  justify-content: center;
  font-size: 2rem;
  text-align: center;
  min-height: 80px;
  margin-block-start: 20px;
  padding: 15px;
  border: 2px solid var(--orange);
  border-radius: 2px;
}

#animation-container {
  display: flex;
  flex-direction: column-reverse;
  justify-content: end;
  gap: 1rem;
  margin-block-end: 1rem;
  min-height: 40vh;
  border: 2px dashed var(--orange);
  padding: 1rem;
}

.animation-frame {
  font-family: -apple-system, BlinkMacSystemFont, avenir next, avenir, segoe ui,
    helvetica neue, helvetica, Cantarell, Ubuntu, roboto, noto, arial,
    sans-serif;
  padding: 15px 10px;
  border: 5px solid var(--orange);
  font-size: 1.2rem;
  text-align: center;
}

@media screen and (min-width: 36em) {
  body {
    font-size: 1rem;
  }

  .input-container {
    flex-direction: row;
    width: unset;
  }

  .number-input {
    width: unset;
  }
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
