---
id: 68e6aff5325410118e76e7be
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você praticará trabalhar com ARIA construindo um controlador de áudio acessível.

Para começar, adicione um elemento `h1` com o texto `Audio Controls`.

# --hints--

Você deve ter um elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Seu elemento `h1` deve estar dentro do seu elemento `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Seu elemento `h1` deve conter o texto `Audio Controls`.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
