---
id: 68eac57fadcc662ce917e75e
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Em seguida, adicione um elemento `button` com o atributo `type` definido para o valor `button`.

O texto do elemento `button` deve ser `Play`.

# --hints--

Você deve ter um elemento `button`.

```js
assert.exists(document.querySelector('button'));
```

Seu elemento `button` deve ter um atributo `type`.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

Seu elemento `button` deve ter o atributo `type` definido como `button`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

Seu elemento `button` deve ter o texto `Play`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
