---
id: 68eb7f2397da251a40e04aa2
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Em seguida, adicione um elemento `div` para conter todos os elementos do controle de volume.

Dentro do elemento `div`, aninhe um elemento `span` com um atributo `id` definido como `volume-label`.  
Além disso, o texto para o elemento `span` deve ser `Volume`.

O `id` será posteriormente referenciado por atributos ARIA para que tecnologias assistivas possam identificar este texto como o rótulo do controle de volume.

# --hints--

Você deve ter um elemento `div` abaixo do seu elemento `button`.

```js
assert.exists(document.querySelector('button + div'));
```

Você deve ter um elemento `span` dentro do seu elemento `div`.

```js
assert.exists(document.querySelector('button + div > span'));
```

Seu elemento `span` deve ter um atributo `id`.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

Seu elemento `span` deve ter um atributo `id` com o valor definido como `volume-label`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

Seu elemento `span` deve conter o texto `Volume`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
