---
id: 68ec0a48c9d6db3594d768c3
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Agora é hora de adicionar outro elemento `span`.

Este elemento `span` deve ter um atributo `id` definido como `volume-description`.  
Além disso, o texto para este elemento `span` deve ser `Adjust the sound level`.

Semelhante ao elemento `span` anterior, o `id` deste `span` será posteriormente referenciado por atributos ARIA para que tecnologias assistivas possam identificar este texto como a descrição do controle de volume.

# --hints--

Você deve ter dois elementos `span` dentro do elemento `div`.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

Seu segundo elemento `span` deve ter um atributo `id`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

Seu segundo elemento `span` deve ter um atributo `id` com o valor `volume-description`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

Seu segundo elemento `span` deve ter o texto `Adjust the sound level`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
