---
id: 68ed488821845990c911a223
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Seu controlador de áudio acessível está quase completo mas há uma última coisa a adicionar.

Adicione um elemento `button` com o atributo `type` definido como `button` e com o texto `Mute` abaixo do elemento `div`.

E com essa mudança, você concluiu este workshop!

# --hints--

Você deve ter um elemento `button` abaixo do elemento `div`.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

Seu elemento `button` deve ter um atributo `type`.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

Seu elemento `button` deve ter um atributo `type` com o valor `button`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

Seu elemento `button` deve ter o texto `Mute`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
