---
id: 68ed4204ab7edc66b6cb2383
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Como você se lembra das lições anteriores, o atributo `aria-labelledby` é usado quando há texto existente na página que pode ser usado como rótulo. Neste caso, você adicionou o atributo `id` aos seus elementos `span` que agora serão usados aqui.

Adicione um atributo `aria-labelledby` ao elemento `input` definido como `volume-label volume-description`.

# --hints--

Seu elemento `input` deve ter um atributo `aria-labelledby`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

Seu elemento `input` deve ter um atributo `aria-labelledby` com o valor `volume-label volume-description`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
