---
id: 68ed3a96e0ebcc36565b39e5
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Em seguida, você precisa adicionar um elemento `input` com o atributo `type` definido como `range` para criar o controle deslizante de volume. Defina o atributo `min` como `0`, o atributo `max` como `100` e o atributo `value` como `50` para definir o nível de volume padrão.

# --hints--

Você deve ter um elemento `input` abaixo do seu segundo elemento `span`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

Seu elemento `input` deve ter um atributo `type`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

Seu elemento `input` deve ter um atributo `type` com o valor `range`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

Seu elemento `input` deve ter um atributo `min`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

Seu elemento `input` deve ter um atributo `min` com o valor `0`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

Seu elemento `input` deve ter um atributo `max`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

Seu elemento `input` deve ter um atributo `max` com o valor `100`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

Seu elemento `input` deve ter um atributo `value`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

Seu elemento `input` deve ter um atributo `value` com o valor `50`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
