---
id: 68e4c81577243e4db1676491
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з семантичним HTML, створюючи вебсторінку, яка містить деякі поради Квінсі Ларсона щодо працевлаштування розробником. Для вас підготовлено базовий шаблонний код HTML.

Почніть зі створення елемента `h1` з текстом `Quincy's Tips for Getting a Developer Job`.

# --hints--

У вас має бути елемент `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ваш елемент `h1` повинен містити текст `Quincy's Tips for Getting a Developer Job`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
