---
id: 68e65ec5348e40a739903e6d
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Тепер створіть абзац нижче елемента `h1`. Всередині цього абзацу додайте текст `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

# --hints--

У вас має бути елемент `p`.

```js
assert.exists(document.querySelector('p'));
```

Ваш абзац має містити текст `Learning to code is hard, but as Quincy Larson says, You can become a developer.`. Перевірте правильність написання.

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

Ваш абзац має бути розташований нижче елемента `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
