---
id: 68e90335bd53a901c41dfc13
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У попередньому уроці ви дізналися, що елемент `q` використовується для відокремлення цитованого тексту від навколишнього вмісту.

Ось приклад:

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

Більшість сучасних браузерів автоматично додадуть лапки навколо рядкової цитати, коли ви використовуєте елемент `q`.

Всередині абзацу обгорніть `You can become a developer.` у рядковий елемент цитати, залишаючи решту абзацу без змін.

# --hints--

У вас має бути елемент `q`.

```js
assert.exists(document.querySelector('q'));
```

Ваш елемент `q` має бути вкладений у абзац, який знаходиться під елементом `h1`.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

Ваш елемент `q` має містити текст `You can become a developer.`

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

Всередині вашого елемента `q` текст не повинен мати початкових або кінцевих пробілів (жодних пробілів, табуляцій або розривів рядка перед першим символом або після останнього).

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

Ви повинні додати лише теги `q` навколо відповідної частини речення, не змінюючи оригінальний текст. Елемент `p` має все ще читатися як: `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
