---
id: 68e964b84f954c322376e592
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Якщо джерелом цитати є вебсайт, ви можете посилатися на нього за допомогою атрибута `cite`. Значення цього атрибута має бути дійсною URL-адресою. Хоча цей атрибут не змінює вигляд блочного цитування, він дуже корисний для надання читачам екрана та пошуковим системам додаткової інформації про цитату.

Ось приклад рядкового елемента цитування з атрибутом `cite`:

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

Додайте атрибут `cite` до рядкового елемента цитування з цією URL-адресою:

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

Ваш елемент `q` має містити атрибут `cite`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

Атрибут `cite` вашого елемента `q` має мати значення `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
