---
id: 68ea7e7b916e8270483689b6
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Під елементом `h2` ви додасте ще одну цитату Квінсі. Цього разу цитата не буде частиною більшого абзацу. Натомість увесь абзац буде цитатою. Щоб відокремити такий цитований текст, слід використовувати елемент блочного цитування: `blockquote`. У браузері ви побачите, що текст трохи відступлений.

Ось приклад елемента блочного цитування з цитованим текстом:

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Тепер, всередині першого розділу, додайте елемент блочного цитування під елементом `h2` з текстом `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

# --hints--

Всередині першого елемента `section` має бути елемент `blockquote` під елементом `h2`.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

Ваш елемент `blockquote` має містити текст `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
