---
id: 66b62d0ad68488dd76228d6c
title: 调试一个随机背景颜色更改器
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot 正在学习 JavaScript，并尝试构建自己的随机背景颜色更改器。然而，他们在过程中犯了一些错误。   

**目标：** 完成以下用户故事并使所有测试通过，以确保实验室正常运行。

**用户故事：**

1. 你应该修复 `darkColorsArr` 数组中的语法错误。
1. 你应该修复在 `getRandomIndex` 函数内引用 `math` 对象时抛出的 `ReferenceError`。
1. 你应该将 `randomIndex` 向下取整到最接近的整数。
1. 你应该更新 `body` 变量的值以使用正确的方法名称并修复相关的 `TypeError`。
1. 你应该修复 `bgHexCodeSpanElement`，使其正确选择具有 `id` 为 `bg-hex-code` 的元素。
1. 你应该修改 `changeBackgroundColor` 内的 `color` 变量，使其设置为 `darkColorsArr` 数组中的随机颜色。
1. 你应该更新 `btn` 变量，使其指向按钮元素。打开 `index.html` 查看该按钮元素的正确 `id` 名称。
1. 你应该修复为 `btn` 添加 `click` 事件监听器的那一行，通过将对 `changeBackgroundColor` 函数的引用传递给该方法。
1. 每次点击 `#btn` 时，你的背景应随机更改为 `darkColorsArr` 中的某种颜色。

# --hints--

你应该解决 `darkColorsArr` 数组中的语法错误。 记住，数组元素应该用逗号分隔。

```js
assert.isArray(darkColorsArr);
```

您应该修复 `math.random()` 行中的大小写错误。 

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

你应该将 `darkColorsArr.length * Math.random()` 四舍五入为最接近的整数。

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

你应该更新你的 `body` 变量的值，以便 `queryselector` 使用驼峰命名法。

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

你应该修正 `"bg-hex-code"` 以便它正确引用 `id` 为 `bg-hex-code`。 请记住，CSS 中的 `id` 选择器以井号 (`#`) 开头。

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

你应该在 `darkColorsArr[getRandomIndex]` 行中调用 `getRandomIndex` 函数。

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

你应该修复 `"#click-btn"` 行中的 `id` 名称，使其与 `index.html` 文件中的正确 `id` 名称匹配。

```js
assert.match(btn?.id, /btn/);
```

你不应该调用 `changeBackgroundColor` 函数。相反，你应该将函数引用传递给 `addEventListener` 方法。

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
