---
id: 66e2e99caa0cb3570f4822b4
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

حان الوقت الآن لاختبار دالة `getConsonantCount` الخاصة بك.

أنشئ متغيرًا باسم `consonantCount` وقم بتعيين نتيجة استدعاء الدالة `getConsonantCount` مع المعلمة `"Coding is fun"` إليه.

بعد ذلك، سجّل في وحدة التحكم النص التالي: `"Consonant Count: [Consonant count goes here]"`. استبدل `[Consonant count goes here]` باسم المتغير الفعلي. يمكنك اختيار استخدام سلاسل القوالب أو ربط السلاسل النصية باستخدام العامل `+` هنا.

# --hints--

يجب أن تنشئ متغيرًا باسم `consonantCount`.

```js
assert.isNotNull(consonantCount)
```

يجب تعيين متغير `consonantCount` إلى نتيجة `getConsonantCount("Coding is fun")`.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

يجب أن تسجل متغير `consonantCount` في وحدة التحكم مع المعلمة `"Consonant Count: ${consonantCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
