---
id: 66e2e5d890ac8a54c9d703bd
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

حان الوقت لحساب الحروف الساكنة. أنشئ دالة `getConsonantCount` بمعلمة `sentence`.

داخل الدالة، استخدم حلقة لحساب عدد الحروف الساكنة في `sentence` التي ستُمرر إلى الدالة عند استدعائها. الحرف الساكن هو أي حرف ليس من الأحرف التالية: `"aeiou"`.

يجب أن تُرجع دالة `getConsonantCount` رقمًا.

# --hints--

يجب أن تنشئ دالة `getConsonantCount`.

```js
assert.isFunction(getConsonantCount);
```

يجب أن تحتوي دالة `getConsonantCount` على معلمة `sentence`.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

يجب أن تُرجع دالة `getConsonantCount` رقمًا.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

عندما تكون الجملة `"Coding is fun"`، يجب أن تُرجع دالة `getConsonantCount` القيمة `7`.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

عندما تكون الجملة `"Hello, World!"`، يجب أن تُرجع دالة `getConsonantCount` القيمة `7`.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

يجب أن يكون حساب الحروف الساكنة غير حساس لحالة الأحرف.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

يجب أن تُرجع دالة `getConsonantCount` العدد الصحيح للحروف الساكنة لأي جملة.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
