---
id: 66e2edb23013ca596f8d9568
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

حان الوقت الآن لاختبار دالة `getPunctuationCount` الخاصة بك.

أنشئ متغيرًا باسم `punctuationCount` وقم بتعيين نتيجة استدعاء الدالة `getPunctuationCount` مع المعلمة `"WHAT?!?!?!?!?"` له.

بعد ذلك، سجّل في وحدة التحكم النص التالي: `"Punctuation Count: [Punctuation count goes here]"`. استبدل `[Punctuation count goes here]` باسم المتغير الفعلي. يمكنك اختيار استخدام سلاسل القوالب أو دمج السلاسل النصية باستخدام العامل `+` هنا.

# --hints--

يجب أن تنشئ متغيرًا باسم `punctuationCount`.

```js
assert.isNotNull(punctuationCount)
```

يجب تعيين متغير `punctuationCount` إلى نتيجة `getPunctuationCount("WHAT?!?!?!?!?")`.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

يجب أن تسجل في وحدة التحكم النص التالي: `"Punctuation Count: [Punctuation count goes here]"`. استبدل `[Punctuation count goes here]` باسم المتغير الفعلي. تأكد من استخدام صيغة دمج السلاسل النصية الصحيحة هنا.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

