---
id: 66e2f06b191e305c00574e4d
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

أخيرًا، احسب عدد الكلمات بإنشاء دالة `getWordCount` بمعلمة `sentence`. يجب أن تُرجع الدالة العدد الإجمالي للكلمات في الجملة المُمررة عند استدعائها.

# --hints--

يجب إنشاء دالة `getWordCount`.

```js
assert.isFunction(getWordCount);
```

يجب أن تحتوي دالة `getWordCount` على معلمة `sentence`.

```js
assert.match(getWordCount.toString(), /sentence/);
```

يجب أن تُرجع دالة `getWordCount` رقمًا.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

عندما تكون الجملة هي `"When are you gonna start learning to code?"`، يجب أن تُرجع دالة `getWordCount` القيمة `8`.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

عندما تكون الجملة هي `"What's going on?"`، يجب أن تُرجع دالة `getWordCount` القيمة `3`.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

يجب أن يكون عد الكلمات غير حساس لحالة الأحرف.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

يجب أن تُرجع دالة `getWordCount` العدد الصحيح للكلمات لأي جملة.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both front-end and back-end development"), 7);
```

يجب أن تُرجع دالة `getWordCount` العدد الصحيح للكلمات في حالة السلسلة الفارغة أو السلسلة التي تحتوي فقط على فراغات.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
