---
id: 66e2d680e129e1423116a541
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستبني محلل جمل يأخذ جملة ويحسب عدد الكلمات، الحروف المتحركة، الحروف الساكنة، وعلامات الترقيم.

لتبدأ، أنشئ دالة `getVowelCount` بمعلمة تسمى `sentence`. يجب أن تُرجع دالتك العدد الإجمالي للحروف المتحركة في الجملة.

# --hints--

يجب أن تنشئ دالة `getVowelCount`.

```js
assert.isFunction(getVowelCount);
```

يجب أن تحتوي دالة `getVowelCount` على معلمة `sentence`.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

يجب أن تُرجع دالة `getVowelCount` رقمًا.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

عندما تكون الجملة `"Apples are tasty fruits"`، يجب أن تُرجع دالة `getVowelCount` القيمة `7`.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

عندما تكون الجملة `"Hello, World!"`، يجب أن تُرجع الدالة القيمة `3`.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

يجب أن يكون عد الحروف المتحركة غير حساس لحالة الأحرف.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

يجب أن تُرجع دالة `getVowelCount` العدد الصحيح للحروف المتحركة لأي جملة.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
