---
id: 66e2eab8a5638f57b637b7cc
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

يجب أن تحسب عدد علامات الترقيم الآن.

أنشئ دالة `getPunctuationCount` بمعلمة `sentence`.

داخل الدالة، أنشئ حلقة لحساب عدد حروف الترقيم في `sentence` التي سترسل إلى الدالة عند استدعائها. لغرضنا، حرف الترقيم هو أي حرف ليس فراغًا (`" "`) أو حرفًا أبجديًا.

يجب أن تُرجع دالة `getPunctuationCount` رقمًا.

# --hints--

يجب أن تنشئ دالة `getPunctuationCount`.

```js
assert.isFunction(getPunctuationCount);
```

يجب أن تحتوي دالة `getPunctuationCount` على معلمة `sentence`.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

يجب أن تُرجع دالة `getPunctuationCount` رقمًا.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

عندما تكون الجملة `"What's going on here?"`، يجب أن تُرجع دالة `getPunctuationCount` القيمة `2`.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

عندما تكون الجملة `"What????!"`، يجب أن تُرجع دالة `getPunctuationCount` القيمة `5`.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

يجب أن تُرجع دالة `getPunctuationCount` العدد الصحيح لعلامات الترقيم لأي جملة.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
