---
id: 66e2df8c5676fe51cead8ec8
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

حان الوقت الآن لاختبار دالة `getVowelCount` الخاصة بك.

عرّف متغيرًا باسم `vowelCount` وعيّن له نتيجة استدعاء الدالة `getVowelCount` مع المعلمة `"Apples are tasty fruits"`

بعد ذلك، سجّل في وحدة التحكم النص التالي: `"Vowel Count: [vowel count goes here]"`. استبدل `[vowel count goes here]` باسم المتغير الفعلي. يمكنك اختيار استخدام سلاسل القوالب أو دمج السلاسل النصية باستخدام العامل `+` هنا.

# --hints--

يجب أن تعرّف متغيرًا باسم `vowelCount`.

```js
assert.isNotNull(vowelCount)
```

يجب أن يُعيّن متغير `vowelCount` إلى نتيجة `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

يجب أن تسجل متغير `vowelCount` في وحدة التحكم مع المعلمة `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
