---
id: 66a8347f97a33e36ffc81b9f
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

To prepare to create some actual content, add a `body` element below the `head` element.

# --hints--

You should have an opening `body` tag.

```js
assert.match(code, /<body>/i);
```

You should have a closing `body` tag.

```js
assert.match(code, /<\/body>/i);
```

`head` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

`body` 要素を `head` 要素の後に置く必要があります。

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>

  --fcc-editable-region--
</html>
```
