---
id: 66ed41f912d0bb1dc62da5dd
title: 创建一个 JavaScript 趣味问答机器人
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该在控制台记录 `"Hello! I'm your coding fun fact guide!"` 作为对用户的问候消息。
1. 你应该创建三个变量：`botName`、`botLocation` 和 `favoriteLanguage`，分别保存机器人名称、它的所在地和它最喜欢的编程语言。
1. 你应该在控制台记录 `"My name is (botName) and I live on (botLocation)."`。
1. 你应该在控制台记录 `"My favorite programming language is (favoriteLanguage)."`。
1. 你应该使用 `let` 创建一个 `codingFact` 变量，并赋值一个关于机器人最喜欢的编程语言的有趣事实字串，使用 `favoriteLanguage` 变量。
1. 你应该在控制台记录 `codingFact`。
1. 你应该重新赋值 `codingFact` 变量为另一个关于机器人最喜欢的编程语言的新事实，仍然使用 `favoriteLanguage` 变量。
1. 你应该再次在控制台记录 `codingFact`。
1. 你应该第三次重新赋值 `codingFact` 变量为关于机器人最喜欢的编程语言的第三个事实，仍然使用 `favoriteLanguage` 变量。
1. 你应该第三次在控制台记录 `codingFact`。
1. 你应该在控制台记录 `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."`，作为机器人向用户的告别语。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该输出 `"Hello! I'm your coding fun fact guide!"` 到控制台。

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

你应该声明一个 `botName` 变量。 仔细检查是否有拼写或大小写错误。

```js
assert.exists(botName);
```

你的 `botName` 变量应该是一个字符串。

```js
assert.isString(botName);
```

你应该声明一个 `botLocation` 变量。 仔细检查是否有拼写或大小写错误。

```js
assert.exists(botLocation);
```

你的 `botLocation` 变量应该是一个字符串。

```js
assert.isString(botLocation);
```

你应该声明一个 `favoriteLanguage` 变量。 仔细检查是否有拼写或大小写错误。

```js
assert.exists(favoriteLanguage);
```

你 `favoriteLanguage` 变量应该是一个字符串。

```js
assert.isString(favoriteLanguage);
```

你应该向控制台日志 `"My name is (botName) and I live on (botLocation)."` 并将变量添加到字串中。

```js
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

你应该向控制台日志 `"My favorite programming language is (favoriteLanguage)."` 并将变量添加到字串中。

```js
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

你应该使用 `let` 声明一个名为 `codingFact` 的新变量。

```js
const explorer = await __helpers.Explorer(code);
const { codingFact } = explorer.variables;
assert.isTrue(codingFact?.matches(`let codingFact = ${codingFact?.value?.toString()}`));
```

你应该给 `codingFact` 一个包含 `favoriteLanguage` 的值。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.codingFact);
assert.include(String(codingFact), favoriteLanguage);
```

你应该将 `codingFact` 输出到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

你应该为 `codingFact` 重新赋值，其中也要包括 `favoriteLanguage`，并将其输出到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.exists(loggingCodingFacts);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
```

你应该第三次为 `codingFact` 赋值，其中还要包括 `favoriteLanguage`，并将其输出到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.lengthOf(loggingCodingFacts, 3);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.isNotEmpty(codingFact);
```

你应该向控制台记录 `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` 并添加变量的值。

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + " was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
