---
id: 587d7fb7367417b2b2512c0d
title: model.findById()를 사용해 _id로 데이터베이스를 검색하기
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

문서를 저장할 때 MongoDB는 자동으로 `_id` 필드를 추가하고 고유한 영숫자 키로 설정합니다. `_id`로 검색하는 작업은 매우 자주 발생하므로 Mongoose는 이를 위한 전용 메서드를 제공합니다.

# --instructions--

`findPersonById`를 수정해 `_id`를 사용하여 주어진 `Model.findById() -> Person`를 가진 유일한 사람을 찾으세요. 함수 인수 `personId`를 검색 키로 사용하세요.

# --hints--

Id로 항목을 찾는 데 성공해야 합니다

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
