---
id: 587d7fb7367417b2b2512c0a
title: model.create()로 여러 레코드 만들기
challengeType: 2
forumTopicId: 301537
dashedName: create-many-records-with-model-create
---

# --description--

가끔 모델 인스턴스를 여러 개 만들어야 할 때가 있습니다. 예를 들어 데이터베이스에 초기 데이터를 채울 때 그렇습니다. `Model.create()`는 `[{name: 'John', ...}, {...}, ...]` 같은 객체 배열을 첫 번째 인수로 받아 모두 데이터베이스에 저장합니다.

# --instructions--

`createManyPeople` 함수를 수정해서 `Model.create()`를 사용해 `arrayOfPeople` 인수로 여러 사람을 만드세요.

**참고:** 이전 도전 과제에서 생성한 모델을 재사용할 수 있습니다.

# --hints--

한 번에 여러 데이터베이스 항목을 만드는 데 성공해야 합니다

```js
  const response = await fetch(code + '/_api/create-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'John', age: 24, favoriteFoods: ['pizza', 'salad'] },
      { name: 'Mary', age: 21, favoriteFoods: ['onions', 'chicken'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an array');
  assert.equal(
    data.length,
    2,
    'the response does not contain the expected number of items'
  );
  assert.equal(data[0].name, 'John', 'The first item is not correct');
  assert.equal(
    data[0].__v,
    0,
    'The first item should be not previously edited'
  );
  assert.equal(data[1].name, 'Mary', 'The second item is not correct');
  assert.equal(
    data[1].__v,
    0,
    'The second item should be not previously edited'
  );
```
