---
id: 587d7fb6367417b2b2512c09
title: 모델의 레코드 생성 및 저장
challengeType: 2
forumTopicId: 301536
dashedName: create-and-save-a-record-of-a-model
---

# --description--

이 도전 과제에서는 모델의 레코드를 생성하고 저장해야 합니다.

# --instructions--

`createAndSavePerson` 함수 내에서, 이전에 만든 `Person` 모델 생성자를 사용해 문서 인스턴스를 생성하세요. 생성자에 `name`, `age`, `favoriteFoods` 필드를 가진 객체를 전달하세요. 이들의 타입은 `personSchema`에 있는 타입과 일치해야 합니다. 그런 다음, 반환된 문서 인스턴스에서 `document.save()` 메서드를 호출하세요. Node 관례에 따라 콜백 함수를 인수로 전달하세요. 이것은 일반적인 패턴이며, 다음의 모든 CRUD 메서드는 마지막 인수로 이와 같은 콜백 함수를 받습니다.

```js
/* Example */

// ...
person.save(function(err, data) {
  //   ...do your stuff here...
});
```

# --hints--

데이터베이스 항목 생성 및 저장이 성공해야 합니다

```js
  const response = await fetch(code + '/_api/create-and-save-person');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isString(data.name, '"item.name" should be a String');
  assert.isNumber(data.age, '28', '"item.age" should be a Number');
  assert.isArray(
    data.favoriteFoods,
    '"item.favoriteFoods" should be an Array'
  );
  assert.equal(data.__v, 0, 'The db item should be not previously edited');
```
