---
id: 587d7fb8367417b2b2512c0e
title: '찾기, 편집, 저장을 실행하여 고전적인 업데이트 수행하기'
challengeType: 2
forumTopicId: 301541
dashedName: perform-classic-updates-by-running-find-edit-then-save
---

# --description--

예전에는 문서를 편집하고 어떻게든 사용할 수 있게 하려면(예: 서버 응답으로 다시 보내기) 이렇게 해야 했습니다. Mongoose에는 전용 업데이트 메서드인 `Model.update()`가 있습니다. 이 메서드는 저수준 mongo 드라이버에 바인딩되어 있습니다. 특정 조건에 맞는 여러 문서를 일괄 편집할 수 있지만, 업데이트된 문서를 반환하지 않고 '상태' 메시지만 보냅니다. 게다가 mongo 드라이버를 직접 호출하기 때문에 모델 유효성 검사가 어렵습니다.

# --instructions--

`findEditThenSave` 함수를 수정하여 `_id` 매개변수를 검색 키로 사용해 `personId`로 사람을 찾으세요(위 방법 중 아무거나 사용 가능). 그런 다음 `"hamburger"`를 그 사람의 `favoriteFoods` 목록에 추가하세요(`Array.push()`를 사용할 수 있습니다). 그리고 찾기 콜백 안에서 업데이트된 `save()`를 `Person`하세요.

**참고:** 스키마에서 `favoriteFoods`를 타입을 지정하지 않고 배열로 선언했다면(i.e. `[String]`), `favoriteFoods`는 기본적으로 Mixed 타입이 되어 수동으로 `document.markModified('edited-field')`를 사용해 편집 표시를 해야 할 수 있습니다. 자세한 내용은 <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">Mongoose 기사</a>를 참조하세요.

# --hints--

찾기-편집-업데이트 작업이 성공해야 합니다

```js
  const response = await fetch(code + '/_api/find-edit-save', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Poldo',
      age: 40,
      favoriteFoods: ['spaghetti']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Poldo', 'item.name is not what is expected');
  assert.equal(data.age, 40, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['spaghetti', 'hamburger'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 1, 'The item should be previously edited');
```
