---
id: 587d7fb6367417b2b2512c07
title: 모델 만들기
challengeType: 2
forumTopicId: 301535
dashedName: create-a-model
---

# --description--

**C**RUD 1부 - 생성

먼저, 스키마가 필요합니다. 각 스키마는 MongoDB 컬렉션에 매핑됩니다. 스키마는 해당 컬렉션 내 문서의 구조를 정의합니다. 스키마는 모델의 구성 요소입니다. 복잡한 모델을 만들기 위해 중첩할 수 있지만, 여기서는 간단하게 유지하겠습니다. 모델은 문서라고 하는 객체 인스턴스를 생성할 수 있게 해줍니다.

서버에서는 데이터베이스와의 상호작용이 핸들러 함수에서 일어납니다. 이 함수들은 어떤 이벤트가 발생할 때 실행됩니다(예: 누군가 API의 엔드포인트를 호출할 때). 이 연습 문제들에서도 같은 방식을 따릅니다. `done()` 함수는 삽입, 검색, 업데이트, 삭제 같은 비동기 작업이 완료된 후 계속 진행할 수 있음을 알려주는 콜백입니다. Node 관례를 따르며, 성공 시 `done(null, data)`로, 오류 시 `done(err)`로 호출해야 합니다.

경고 - 원격 서비스와 상호작용할 때 오류가 발생할 수 있습니다!

```js
/* Example */

const someFunc = function(done) {
  //... do something (risky) ...
  if (error) return done(error);
  done(null, result);
};
```

# --instructions--

다음 구조를 가진 `personSchema`라는 사람 스키마를 만드세요:

* `name` 타입의 필수 `String` 필드
* `age` 타입의 `Number` 필드
* `favoriteFoods` 타입의 `[String]` 필드

Mongoose 기본 스키마 타입을 사용하세요. 원한다면 더 많은 필드를 추가하거나, required나 unique 같은 간단한 검증기를 사용하고 기본값을 설정할 수도 있습니다. 자세한 내용은 <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">Mongoose 기사</a>를 참조하세요.

이제 `personSchema`에서 모델을 만들고, 기존 변수 `Person`에 할당하세요.

# --hints--

Mongoose 스키마에서 인스턴스를 만드는 데 성공해야 합니다

```js
  const response = await fetch(code + '/_api/mongoose-model', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Mike',
      age: 28,
      favoriteFoods: ['pizza', 'cheese']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Mike', '"model.name" is not what expected');
  assert.equal(data.age, '28', '"model.age" is not what expected');
  assert.isArray(
    data.favoriteFoods,
    '"model.favoriteFoods" is not an Array'
  );
  assert.include(
    data.favoriteFoods,
    'pizza',
    '"model.favoriteFoods" does not include the expected items'
  );
  assert.include(
    data.favoriteFoods,
    'cheese',
    '"model.favoriteFoods" does not include the expected items'
  );
```
