---
id: 587d7fb6367417b2b2512c06
title: Mongoose 설치 및 설정
challengeType: 2
forumTopicId: 301540
dashedName: install-and-set-up-mongoose
---

# --description--

이 도전 과제들은 다음 방법 중 하나를 사용해 코드를 작성하는 작업을 포함합니다:

- <a href="https://github.com/freeCodeCamp/boilerplate-mongomongoose/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 도전 과제를 완료하세요.
- 원하는 사이트 빌더를 사용해 프로젝트를 완성하세요. GitHub 저장소의 모든 파일을 반드시 포함해야 합니다.

이번 도전 과제에서는 MongoDB Atlas 데이터베이스를 설정하고 연결에 필요한 패키지를 가져옵니다.

<a href='https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/' target="_blank" rel="noopener noreferrer nofollow">이 튜토리얼</a>을 따라 MongoDB Atlas에서 호스팅된 데이터베이스를 설정하세요.

# --instructions--

`mongoose@^5.11.15`가 프로젝트의 `package.json` 파일에 추가되었습니다. 먼저 `mongoose`에서 mongoose를 `myApp.js`로 require 하세요. 다음으로 `.env` 파일을 만들고 `MONGO_URI` 변수를 추가하세요. 이 변수의 값은 MongoDB Atlas 데이터베이스 URI여야 합니다. URI는 작은따옴표나 큰따옴표로 감싸야 하며, 환경 변수에서 `=` 주위에 공백을 넣으면 안 된다는 점을 기억하세요. 예를 들어 `MONGO_URI='VALUE'`와 같습니다.

완료하면 다음 구문을 사용해 `connect` 파일 내에서 `myApp.js` 메서드를 호출하여 데이터베이스에 연결하세요:

```js
mongoose.connect(<Your URI>, { useNewUrlParser: true, useUnifiedTopology: true });
```

# --hints--

"mongoose version ^5.11.15" 의존성이 package.json에 있어야 합니다

```js
  const response = await fetch(code + '/_api/file/package.json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  var packJson = JSON.parse(data);
  assert.property(packJson.dependencies, 'mongoose');
  assert.match(
    packJson.dependencies.mongoose,
    /^\^5\.11\.15/,
    'Wrong version of "mongoose". It should be ^5.11.15'
  );
```

"mongoose"가 데이터베이스에 연결되어야 합니다

```js
  const response = await fetch(code + '/_api/is-mongoose-ok');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.isMongooseOk, 'mongoose is not connected');
```
