---
id: 587d7fb8367417b2b2512c10
title: model.findByIdAndRemove로 문서 하나 삭제하기
challengeType: 2
forumTopicId: 301539
dashedName: delete-one-document-using-model-findbyidandremove
---

# --description--

`findByIdAndRemove`와 `findOneAndRemove`은 이전 업데이트 메서드와 비슷합니다. 삭제된 문서를 데이터베이스에 전달합니다. 평소처럼 함수 인수 `personId`를 검색 키로 사용하세요.

# --instructions--

`removeById` 함수를 수정하여 사람의 `_id`로 한 명을 삭제하세요. `findByIdAndRemove()` 또는 `findOneAndRemove()` 메서드 중 하나를 사용해야 합니다.

# --hints--

항목 삭제가 성공해야 합니다

```js
  const response = await fetch(code + '/_api/remove-one-person', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Jason Bourne',
      age: 36,
      favoriteFoods: ['apples']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Jason Bourne', 'item.name is not what expected');
  assert.equal(data.age, 36, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['apples'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0);
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
