---
id: 587d7fb8367417b2b2512c0f
title: model.findOneAndUpdate()를 사용해 문서에 새 업데이트 수행하기
challengeType: 2
forumTopicId: 301542
dashedName: perform-new-updates-on-a-document-using-model-findoneandupdate
---

# --description--

최근 Mongoose 버전에는 문서 업데이트를 간소화하는 메서드가 있습니다. 일부 고급 기능(예: pre/post 훅, 유효성 검사)은 이 방법과 다르게 동작하므로, 고전적인 방법이 여전히 많은 상황에서 유용합니다. `findByIdAndUpdate()`는 id로 검색할 때 사용할 수 있습니다.

# --instructions--

`findAndUpdate` 함수를 수정해 `Name`로 사람을 찾아 그 사람의 나이를 `20`으로 설정하세요. 함수 매개변수 `personName`를 검색 키로 사용하세요.

**참고:** 업데이트된 문서를 반환해야 합니다. 이를 위해 옵션 문서 `{ new: true }`를 `findOneAndUpdate()`의 세 번째 인수로 전달해야 합니다. 기본적으로 이 메서드들은 수정되지 않은 객체를 반환합니다.

# --hints--

findOneAndUpdate로 항목을 성공적으로 찾아 업데이트하세요

```js
  const response = await fetch(code + '/_api/find-one-update', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Dorian Gray',
      age: 35,
      favoriteFoods: ['unknown']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Dorian Gray', 'item.name is not what expected');
  assert.equal(data.age, 20, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['unknown'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(
    data.__v,
    0,
    'findOneAndUpdate does not increment version by design!'
  );
```
