---
id: 587d7fb7367417b2b2512c0b
title: model.find()로 데이터베이스 검색하기
challengeType: 2
forumTopicId: 301543
dashedName: use-model-find-to-search-your-database
---

# --description--

가장 간단한 사용법에서 `Model.find()`는 첫 번째 인수로 쿼리 문서(JSON 객체)를 받고, 그다음 콜백을 받습니다. 일치하는 항목들의 배열을 반환합니다. 매우 다양한 검색 옵션을 지원합니다. 자세한 내용은 문서를 참조하세요.

# --instructions--

`findPeopleByName` 함수를 수정하여 주어진 이름을 가진 모든 사람을 <code>Model.find() -&gt; [Person]</code>을 사용해 찾으세요.

함수 인수 `personName`를 검색 키로 사용하세요.

# --hints--

조건에 맞는 모든 항목을 찾는 데 성공해야 합니다.

```js
  const response = await fetch(code + '/_api/find-all-by-name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'r@nd0mN4m3',
      age: 24,
      favoriteFoods: ['pizza']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data[0].name,
    'r@nd0mN4m3',
    'item.name is not what expected'
  );
  assert.equal(data[0].__v, 0, 'The item should be not previously edited');
```
