---
id: 587d7fb9367417b2b2512c12
title: 검색 결과를 좁히기 위해 체인으로 검색 쿼리 도우미 사용하기
challengeType: 2
forumTopicId: 301533
dashedName: chain-search-query-helpers-to-narrow-search-results
---

# --description--

콜백을 `Model.find()` (또는 다른 검색 메서드)에 마지막 인수로 전달하지 않으면 쿼리가 실행되지 않습니다. 쿼리를 나중에 사용하기 위해 변수에 저장할 수 있습니다. 이런 종류의 객체는 체이닝 구문을 사용해 쿼리를 구성할 수 있게 해줍니다. 실제 데이터베이스 검색은 마지막으로 `.exec()` 메서드를 체이닝할 때 실행됩니다. 항상 이 마지막 메서드에 콜백을 전달해야 합니다. 쿼리 도우미가 여러 가지 있지만, 여기서는 가장 자주 사용하는 것들을 사용할 것입니다.

# --instructions--

`queryChain` 함수를 수정해 `foodToSearch` 변수에 지정된 음식을 좋아하는 사람들을 찾으세요. `name` 기준으로 정렬하고, 결과를 두 개의 문서로 제한하며, 나이는 숨기세요. `.find()`, `.sort()`, `.limit()`, `.select()`을 차례로 체이닝한 다음 `.exec()`을 체이닝하세요. `done(err, data)` 콜백을 `exec()`에 전달하세요.

# --hints--

쿼리 도우미 체이닝이 성공해야 합니다

```js
  const response = await fetch(code + '/_api/query-tools', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Pablo', age: 26, favoriteFoods: ['burrito', 'hot-dog'] },
      { name: 'Bob', age: 23, favoriteFoods: ['pizza', 'nachos'] },
      { name: 'Ashley', age: 32, favoriteFoods: ['steak', 'burrito'] },
      { name: 'Mario', age: 51, favoriteFoods: ['burrito', 'prosciutto'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data.length,
    2,
    'the data array length is not what expected'
  );
  assert.notProperty(
    data[0],
    'age',
    'The returned first item has too many properties'
  );
  assert.equal(
    data[0].name,
    'Ashley',
    'The returned first item name is not what expected'
  );
  assert.notProperty(
    data[1],
    'age',
    'The returned second item has too many properties'
  );
  assert.equal(
    data[1].name,
    'Mario',
    'The returned second item name is not what expected'
  );
```
