---
id: 587d7fb8367417b2b2512c11
title: model.remove()로 여러 문서 삭제하기
challengeType: 2
forumTopicId: 301538
dashedName: delete-many-documents-with-model-remove
---

# --description--

`Model.remove()`는 주어진 조건에 맞는 모든 문서를 삭제하는 데 유용합니다.

# --instructions--

`removeManyPeople` 함수를 수정하여 `nameToRemove` 변수에 포함된 이름을 가진 모든 사람을 `Model.remove()`를 사용해 삭제하세요. `name` 필드가 설정된 쿼리 문서와 콜백을 전달하세요.

**참고:** `Model.remove()`는 삭제된 문서를 반환하지 않고, 작업 결과와 영향을 받은 항목 수를 포함한 JSON 객체를 반환합니다. 테스트에 사용하므로 `done()` 콜백에 전달하는 것을 잊지 마세요.

# --hints--

한 번에 여러 항목을 삭제하는 데 성공해야 합니다

```js
  const response = await fetch(code + '/_api/remove-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Mary', age: 16, favoriteFoods: ['lollipop'] },
      { name: 'Mary', age: 21, favoriteFoods: ['steak'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(!!data.ok, 'The mongo stats are not what expected');
  assert.equal(
    data.n,
    2,
    'The number of items affected is not what expected'
  );
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
