---
id: 587d7fb7367417b2b2512c0c
title: 데이터베이스에서 단일 일치 문서를 반환하려면 model.findOne() 사용하기
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()`는 `Model.find()`처럼 동작하지만, 여러 항목이 있어도 배열이 아닌 단일 문서만 반환합니다. 특히 고유하게 선언한 프로프를 검색할 때 유용합니다.

# --instructions--

`findOneByFood` 함수를 수정하여 `Model.findOne() -> Person`를 사용해 특정 음식이 즐겨찾기에 있는 한 사람만 찾으세요. 함수 인수 `food`를 검색 키로 사용하세요.

# --hints--

한 항목 찾기가 성공해야 합니다

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
