---
id: ab6137d4e35944e21037b769
title: 构建一个标题大小写转换器
challengeType: 26
dashedName: build-a-title-case-converter
---

# --description--

在本实验中，你将创建一个函数，将字串转换为标题大小写。标题大小写意味着每个单词的第一个字母大写，其余字母小写。

`"Web Development Is Awesome"` 是一个标题式分支的字串示例。

**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `titleCase` 函数，它接受一个字串作为参数。
1. `titleCase` 函数应返回一个字串，其中每个单词的首字母大写，其余字母小写。
1. `titleCase("I like to code")` 应该返回 `"I Like To Code"`。
1. `titleCase("javaScript is fun")` 应该返回 `"Javascript Is Fun"`。

# --hints--

你应该创建一个名为 `titleCase` 的函数。

```js
assert.isFunction(titleCase);
```

`titleCase` 应该接受一个单一参数。

```js
assert.lengthOf(titleCase, 1);
```


`titleCase("I'm a little tea pot")` 应返回一个字符串。

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` 应返回 `I'm A Little Tea Pot`。

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` 应返回 `Short And Stout`。

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` 应返回 `Here Is My Handle Here Is My Spout`。

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}
```
