---
id: 564944c91be2204b269d51e3
title: 使用 jQuery 變更元素內的文字
challengeType: 6
forumTopicId: 16773
dashedName: change-text-inside-an-element-using-jquery
---

# --description--

使用 jQuery，你可以更改元素起始和結束 `tag` 之間的 `text`。你甚至可以更改 HTML 標記。

jQuery 有一個名為 `.html()` 的函式（程式），讓你可以在元素內添加 HTML 標記和文字。元素內先前的任何內容都會被你使用此函式提供的內容完全取代。

以下是你如何重寫並強調我們標題的文字：

```js
$("h3").html("<em>jQuery Playground</em>");
```

jQuery 也有一個類似的函式（程式）叫做 `.text()`，它只會更改文字而不會添加頁籤。換句話說，這個函式（程式）不會評估傳入的任何 HTML 頁籤，而是將其視為你想用來取代現有內容的文字。

透過強調其文字來更改 id 為 `target4` 的按鈕。

<a href="https://www.freecodecamp.org/news/html-elements-explained-what-are-html-tags/#em-element" target="_blank" rel="noopener noreferrer nofollow">查看我們關於 &lt;em&gt; 的新聞文章</a>，以了解 `<i>` 和 `<em>` 之間的差異及其用途。

請注意，雖然 `<i>` 頁籤傳統上用於強調文字，但它後來被用作圖示的頁籤。`<em>` 頁籤現在被廣泛接受為強調的頁籤。兩者在此挑戰中皆可使用。

# --hints--

你應該透過添加 HTML 標籤來強調你 `target4` 按鈕中的文字。

```js
assert.isTrue(
  /<em>|<i>\s*#target4\s*<\/em>|<\/i>/gi.test($('#target4').html())
);
```

文字應保持不變。

```js
assert($('#target4') && $('#target4').text().trim() === '#target4');
```

你不應該更改任何其他文字。

```js
assert.isFalse(/<em>|<i>/gi.test($('h3').html()));
```

你應該使用 `.html()` 而不是 `.text()`。

```js
assert(code.match(/\.html\(/g));
```

你應該使用 jQuery 選取 `button id="target4"`。

```js
assert(code.match(/\$\(\s*?(\"|\')#target4(\"|\')\s*?\)\.html\(/));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target4").html('<em>#target4</em>');
  });
</script>

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
