---
id: bad87fee1348bd9aed908626
title: 使用多個 jQuery 選擇器定位相同標的
challengeType: 6
forumTopicId: 18322
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-the-same-element-with-multiple-jquery-selectors
---

# --description--

現在你知道三種標的元素的方法：依型別：`$("button")`、依類別：`$(".btn")`，以及依 id：`$("#target1")`。

雖然可以在單一 `.addClass()` 呼叫中添加多個類別，但讓我們用*三種不同方式*將它們添加到同一個元素。

使用 `.addClass()`，一次只為同一個元素添加一個類別，有三種不同的方法：

為你的所有型別為 `button` 的元素添加 `animated` 類別。

為所有具有 `.btn` 類別的按鈕添加 `shake` 類別。

將 `btn-primary` 類別添加到 id 為 `#target1` 的按鈕。

**注意：**你應該一次只針對一個元素並且只添加一個類別。總共來說，你的三個獨立選擇器最終會將三個類別 `shake`、`animated` 和 `btn-primary` 添加到 `#target1`。

# --hints--

你的程式碼應該使用 `$("button")` 選擇器。

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?button\s*?(?:'|")/gi));
```

你的程式碼應該使用 `$(".btn")` 選擇器。

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?\.btn\s*?(?:'|")/gi));
```

你的程式碼應該使用 `$("#target1")` 選擇器。

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?#target1\s*?(?:'|")/gi));
```

你應該只為你的三個選擇器中的每一個添加一個類別。

```js
assert(
  code.match(/addClass/g) &&
    code.match(/addClass\s*?\(\s*?('|")\s*?[\w-]+\s*?\1\s*?\)/g).length > 2
);
```

你的 `#target1` 元素應該有 `animated`、`shake` 和 `btn-primary` 類別。

```js
assert(
  $('#target1').hasClass('animated') &&
    $('#target1').hasClass('shake') &&
    $('#target1').hasClass('btn-primary')
);
```

你應該只使用 jQuery 將這些類別添加到元素。

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated");
    $(".btn").addClass("shake");
    $("#target1").addClass("btn-primary");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
