---
id: bad87fee1348bd9aedc08826
title: 使用 jQuery 以類別選取標的元素
challengeType: 6
forumTopicId: 18316
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-class-using-jquery
---

# --description--

你看到我們如何讓你所有的 `button` 元素彈跳了嗎？我們用 `$("button")` 選取它們，然後用 `.addClass("animated bounce");` 為它們添加了一些 CSS 類別。

你剛剛使用了 jQuery 的 `.addClass()` 函式，這個函式允許你為元素添加類別。

首先，讓我們使用 `$(".well")` 選擇器來標的你的 `well` 類別的 `div` 元素。

請注意，就像 CSS 宣告式一樣，你在類別名稱前輸入 `.`。

然後使用 jQuery 的 `.addClass()` 函式為你的元素添加 `animated` 和 `shake` 類別。

例如，你可以為你的 `document ready function` 添加以下內容，使所有具有 `text-primary` 類別的元素搖動：

```js
$(".text-primary").addClass("animated shake");
```

# --hints--

你應該使用 jQuery 的 `addClass()` 函式（程式）為所有具有 `well` 類別的元素添加 `animated` 和 `shake` 類別。

```js
assert($('.well').hasClass('animated') && $('.well').hasClass('shake'));
```

你應該只使用 jQuery 將這些類別添加到元素。

```js
assert(!code.match(/class\.\*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
