---
id: bad87fee1348bd9acdd08826
title: 了解 `Script` 頁籤和 `Document Ready` 的運作方式
challengeType: 6
forumTopicId: 18224
dashedName: learn-how-script-tags-and-document-ready-work
---

# --description--

現在我們準備學習 jQuery，這是有史以來最受歡迎的 JavaScript 工具。

在我們開始使用 jQuery 之前，我們需要為你的 HTML 添加一些東西。

首先，為你的頁面頂部添加一個 `script` 元素。請確保在下一行將它關閉。

你的瀏覽器會執行 `script` 元素內的任何 JavaScript，包括 jQuery。

在你的 `script` 元素內，為你的 `script` 添加這段程式碼：`$(document).ready(function() {`。然後在下一行（仍在你的 `script` 元素內）以 `});` 關閉它。

我們稍後會進一步了解 `functions`。重要的是要知道，你放在這個 `function` 裡的程式碼會在你的瀏覽器載入頁面後立即執行。

這很重要，因為如果沒有你的 `document ready function`，你的程式碼可能會在 HTML 呈現之前執行，這會導致錯誤。

# --hints--

你應該創建一個 `script` 元素，確保它是有效的並且有閉合標記。

```js
assert(
  code.match(/<\/script\s*>/g) &&
    code.match(
      /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
    ) &&
    code.match(/<\/script\s*>/g).length ===
      code.match(
        /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
      ).length
);
```

你應該在你的 `script` 元素開頭添加 `$(document).ready(function() {`。

```js
assert(
  code.match(
    /\$\s*?\(\s*?document\s*?\)\.ready\s*?\(\s*?function\s*?\(\s*?\)\s*?\{/g
  )
);
```

你應該用 `});` 關閉你的 `$(document).ready(function() {` 函式（程式）。

```js
assert(code.match(/\n*?\s*?\}\s*?\);/g));
```

# --seed--

## --seed-contents--

```html
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
  });
</script>
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
