---
id: bad87fee1348bd9aed208826
title: 使用 jQuery 針對元素的子元素
challengeType: 6
forumTopicId: 18320
dashedName: target-the-children-of-an-element-using-jquery
---

# --description--

當 HTML 元素被放置在另一個元素下一層時，它們被稱為該元素的 <dfn>子元素</dfn>。例如，本挑戰中帶有文字 `#target1`、`#target2` 和 `#target3` 的按鈕元素，都是 `<div class="well" id="left-well">` 元素的子元素。

jQuery 有一個名為 `children()` 的函式（程式），允許你存取你所選取的任何元素的子元素。

以下是如何使用 `children()` 函式（程式）為你的 `left-well` 元素的子元素設定顏色為 `blue` 的範例：

```js
$("#left-well").children().css("color", "blue")
```

# --instructions--

將橘色指定給你的 `right-well` 元素的所有子元素。

# --hints--

`#right-well` 的所有子元素應該有橙色文字。

```js
assert($('#right-well').children().css('color') === 'rgb(255, 165, 0)');
```

你應該使用 `children()` 函式（程式）來修改這些元素。

```js
assert(code.match(/\.children\(\)\.css/g));
```

你應該只使用 jQuery 將這些類別添加到元素。

```js
assert(code.match(/<div class="well" id="right-well">/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
    $("#right-well").children().css("color", "orange");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
