---
id: bad87fee1348bd9bedc08826
title: 使用 jQuery 以選擇器定位標的 HTML 元素
challengeType: 6
forumTopicId: 18319
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-html-elements-with-selectors-using-jquery
---

# --description--

現在我們有一個 `document ready` 函式（程式）。

現在讓我們撰寫第一個 jQuery 述語。所有 jQuery 函式都以 `$` 開頭，通常稱為 dollar sign 運算子，或稱為 bling。

jQuery 通常會使用 <dfn>selector</dfn> 選取一個 HTML 元素，然後對該元素執行某些操作。

例如，讓你的所有 `button` 元素彈跳。只要將這段程式碼添加到你的文件 ready 函式中：

```js
$("button").addClass("animated bounce");
```

請注意，我們已經在背景中包含了 jQuery 函式庫和 Animate.css 函式庫，以便你可以在編輯器中使用它們。所以你正在使用 jQuery 將 Animate.css 的 `bounce` 類別套用到你的 `button` 元素。

# --hints--

你應該使用 jQuery 的 `addClass()` 函式（程式）為你的 `button` 元素添加 `animated` 和 `bounce` 這兩個類別。

```js
assert($('button').hasClass('animated') && $('button').hasClass('bounce'));
```

你應該只使用 jQuery 將這些類別添加到元素。

```js
assert(!code.match(/class.*animated/g));
```

你的 jQuery 程式碼應該放在 `$(document).ready();` 函式中。

```js
assert(
  code.replace(/\s/g, '').match(/\$\(document\)\.ready\(function\(\)\{\$/g)
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
