---
id: bad87fee1348bd9aed508826
title: 使用 jQuery 複製元素
challengeType: 6
forumTopicId: 16780
dashedName: clone-an-element-using-jquery
---

# --description--

除了移動元素之外，你也可以將它們從一個位置複製到另一個位置。

jQuery 有一個名為 `clone()` 的函式（程式），可以複製一個元素。

例如，如果我們想將 `target2` 從我們的 `left-well` 複製到我們的 `right-well`，我們會使用：

```js
$("#target2").clone().appendTo("#right-well");
```

你有注意到這涉及將兩個 jQuery 函式串接在一起嗎？這稱為 <dfn>函式串鏈</dfn>，是使用 jQuery 完成工作的方便方式。

複製你的 `target5` 元素並將它附加到你的 `left-well`。

# --hints--

你的 `target5` 元素應該在你的 `right-well` 裡面。

```js
assert($('#right-well').children('#target5').length > 0);
```

你的 `target5` 元素的複本也應該在你的 `left-well` 裡面。

```js
assert($('#left-well').children('#target5').length > 0);
```

你應該只使用 jQuery 來移動這些元素。

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
