---
id: 66b36358ed4f261d64840c24
title: 가용성 표 만들기
challengeType: 25
dashedName: build-an-availability-table
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 헤더를 포함해 최소 세 개의 열과 다섯 개의 행이 있는 표가 있어야 합니다.
1. 첫 번째 행의 셀은 요일을 포함하는 표 헤더여야 합니다.
1. 첫 번째 열의 셀은 `class`가 `time`인 표 헤더여야 하며, 텍스트로 시간 값을 포함해야 합니다.
1. 모든 데이터 셀은 `class`가 `available-#`인 속성을 가져야 하며, 여기서 `#`는 `0`부터 `5`까지의 숫자로 해당 시간에 가능한 인원 수를 나타냅니다.
1. `:root` 선택자에서 `--color#`라는 이름의 CSS 변수를 여섯 개 정의해야 하며, `#`는 `0`부터 `5`까지의 숫자입니다. 각 변수에 색상 값을 할당하고, 이 변수들을 사용해 해당 `.available-#` 요소들의 배경색을 설정하세요.
1. `:root` 선택자에서 `--solid-border`와 `--dashed-border`라는 CSS 변수를 정의해야 합니다. 이 변수들을 사용해 데이터 셀의 하단 테두리를 스타일링하며, 행에 따라 실선과 점선을 번갈아 적용하세요. 행에는 `sharp` 또는 `half` 클래스를 부여해 스타일을 지정하세요.
1. `div` 요소가 `id`가 `legend`인 속성을 가져야 합니다. 이 요소는 텍스트가 `span`인 `Availability` 요소와 `div`가 `id`인 속성을 가진 `legend-gradient` 요소를 포함해야 합니다.
1. `#legend-gradient` 요소에는 `--color0`부터 `--color5`까지 모든 색상이 경계가 뚜렷한 선형 그라디언트로 전환되는 배경 그라디언트를 지정해야 합니다.

**참고:** CSS를 적용하려면 HTML에 스타일시트를 반드시 연결하세요.

# --hints--

`table` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('table'));
```

`table` 요소는 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

표에는 최소 `5` 행이 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('tr')?.length, 5);
```

표에는 최소 `3` 열이 있어야 합니다.

```js
const rows = [...document.querySelectorAll('tr')];
assert.isNotEmpty(rows);
rows.forEach((row) => {
    assert.isAtLeast(row.children.length, 3);
})
```

표의 첫 번째 행에는 최소 세 개의 `th` 요소가 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('tr:first-child>th').length, 3)
```

표의 첫 번째 행에는 `td` 요소가 없어야 합니다.

```js
const firstRow = document.querySelector('tr');
assert.isNotNull(firstRow);
assert.isEmpty(firstRow.querySelectorAll('td'));
```

`sharp` 클래스를 가진 행이 최소 두 개 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('tr[class~="sharp"]').length, 2);
```

`half` 클래스를 가진 행이 최소 두 개 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('tr[class~="half"]').length, 2);
```

첫 번째 열의 `th` 요소는 `class`가 `time`인 속성을 가지며 시간 값을 포함해야 합니다.

```js
const firstColumnData = [...document.querySelectorAll('tr th:first-child')];
assert.isNotEmpty(firstColumnData);
for (let cell of firstColumnData) {
    if (cell.innerText.match(/\d/)?.length) {
        assert.isTrue(cell.classList.contains('time'));
    }
}
```

시간 값을 포함하는 `th` 클래스를 가진 `time` 요소가 최소 네 개 있어야 합니다.

```js
const timeClass = document.querySelectorAll('.time');
assert.isAtLeast(timeClass.length, 4);
for (let cell of timeClass) {
    assert.match(cell.innerText, /\d/)
}
```

모든 `td` 요소는 `class`가 `available-#`인 속성을 가져야 하며, `#`는 `0`부터 `5`까지의 숫자입니다.

```js
const cells = [...document.querySelectorAll('tr[class="sharp"]>td, tr[class="half"]>td')];
assert.isNotEmpty(cells);
for (let cell of cells) {
    const classString = [...cell.classList].join(' ');
    assert.lengthOf(classString.match(/(?<=\s|^)available-[0-5](?=\s|$)/gm), 1);
}
```

`.available-#` 요소가 최소 여덟 개 있어야 하며, `#`는 `0`과 `5` 사이의 숫자입니다.

```js
assert.isAtLeast(document.querySelectorAll('td[class*="available-"]').length, 8);
```

루트 선택자에 `--color0` 변수를 정의해야 합니다.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color0'))
```

루트 선택자에 `--color1` 변수를 정의해야 합니다.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color1'))
```

루트 선택자에 `--color2` 변수를 정의해야 합니다.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color2'))
```

루트 선택자에 `--color3` 변수를 정의해야 합니다.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color3'));
```

루트 선택자에 `--color4` 변수를 정의해야 합니다.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color4'));
```

루트 선택자에 `--color5` 변수를 정의해야 합니다.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color5'));
```

`--color0`를 사용해 `.available-0` 요소의 배경색을 설정해야 합니다.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-0');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color0)');
```

`--color1`를 사용해 `.available-1` 요소의 배경색을 설정해야 합니다.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-1');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color1)');
```

`--color2`를 사용해 `.available-2` 요소의 배경색을 설정해야 합니다.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-2');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color2)');
```

`--color3`를 사용해 `.available-3` 요소의 배경색을 설정해야 합니다.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-3');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color3)');
```

`--color4`를 사용해 `.available-4` 요소의 배경색을 설정해야 합니다.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-4');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color4)');
```

`--color5`를 사용해 `.available-5` 요소의 배경색을 설정해야 합니다.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-5');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color5)');
```

루트 선택자에 `--solid-border` 변수를 정의해야 합니다.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--solid-border'));
```

루트 선택자에 `--dashed-border` 변수를 정의해야 합니다.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--dashed-border'));
```

`td` 요소의 자식인 `.sharp` 요소를 선택해야 합니다.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.sharp td') || new __helpers.CSSHelp(document).getStyle('.sharp>td');
assert.exists(tds);
```

`--solid-border`를 사용해 `bottom-border` 요소의 자식인 `td` 요소의 `.sharp`를 설정해야 합니다.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.sharp td') || new __helpers.CSSHelp(document).getStyle('.sharp>td');
assert.exists(tds);
assert.strictEqual(tds.borderBottom, 'var(--solid-border)');
```

`td` 요소의 자식인 `.half` 요소를 선택해야 합니다.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.half td') || new __helpers.CSSHelp(document).getStyle('.half>td');
assert.exists(tds);
```

`--dashed-border`를 사용해 `bottom-border` 요소의 `td` 요소의 `.half`를 설정해야 합니다.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.half td') || new __helpers.CSSHelp(document).getStyle('.half>td');
assert.exists(tds);
assert.strictEqual(tds.borderBottom, 'var(--dashed-border)');
```

`div` 요소가 `id`가 `legend`인 속성을 가져야 합니다.

```js
assert.equal(document.querySelector('#legend')?.tagName, 'DIV');
```

`span` 안에 텍스트가 `Availability`인 `#legend` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector('#legend span')?.textContent.trim(), 'Availability');
```

`div` 안에 `id`가 `legend-gradient`인 속성을 가진 `#legend` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector('#legend #legend-gradient')?.tagName, 'DIV');
```

`#legend-gradient`의 배경 이미지를 선형 그라디언트로 설정해야 합니다.

```js
const legendGradient = new __helpers.CSSHelp(document).getStyle('#legend-gradient');
assert.exists(legendGradient); 
assert.include(legendGradient.backgroundImage, "linear-gradient(")
```

`#legend-gradient`에 대해 색상 간 전환이 경계가 뚜렷한 선으로 나타나야 하며, 백분율은 정수여야 합니다.

```js
const legendGradient = new __helpers.CSSHelp(document).getStyle('#legend-gradient');
assert.exists(legendGradient);
const gradientStr = legendGradient.backgroundImage;

const COLOR_COUNT = 6;
const EXPECTED_SHORTHAND_STOPS = 6;
const EXPECTED_EXPANDED_STOPS = 12;

for (let i = 0; i < COLOR_COUNT; i++) {
  assert.include(gradientStr, `--color${i}`, `gradient should include --color${i}`);
}

const GRADIENT_PATTERNS = {
  SHORTHAND: /var\(\s*--color[0-5]\s*\)\s+\d+%\s+\d+%/g,
  EXPANDED: /var\(\s*--color[0-5]\s*\)\s+\d+%(?!\s+\d+%)/g
};

const matchGradientFormat = (gradientStr, pattern) => gradientStr.match(pattern);

const shorthandMatches = matchGradientFormat(gradientStr, GRADIENT_PATTERNS.SHORTHAND);
const expandedMatches = matchGradientFormat(gradientStr, GRADIENT_PATTERNS.EXPANDED);

const colorPattern = /var\(\s*--color(\d)\s*\)/g;
const allColorMatches = [...gradientStr.matchAll(colorPattern)];
const colorNumbers = allColorMatches.map(match => parseInt(match[1]));

if (shorthandMatches?.length === EXPECTED_SHORTHAND_STOPS) {
  for (let i = 0; i < COLOR_COUNT; i++) {
    const count = colorNumbers.filter(num => num === i).length;
    assert.equal(count, 1, `In shorthand format, --color${i} should appear exactly once, but it appears ${count} times`);
  }
  
  for (let i = 0; i < COLOR_COUNT; i++) {
    const idx = colorNumbers.indexOf(i);
    assert.isAtLeast(idx, 0, `--color${i} should be in the gradient`);
    if (i > 0) {
      const prevIdx = colorNumbers.indexOf(i - 1);
      assert.isBelow(prevIdx, idx, `--color${i - 1} should appear before --color${i}`);
    }
  }
} else if (expandedMatches?.length === EXPECTED_EXPANDED_STOPS) {
  for (let i = 0; i < COLOR_COUNT; i++) {
    const count = colorNumbers.filter(num => num === i).length;
    assert.equal(count, 2, `In expanded format, --color${i} should appear exactly twice, but it appears ${count} times`);
  }
  
  for (let i = 0; i < COLOR_COUNT; i++) {
    const firstIdx = colorNumbers.indexOf(i);
    const lastIdx = colorNumbers.lastIndexOf(i);
    assert.equal(lastIdx - firstIdx, 1, `--color${i} should appear in consecutive positions for hard lines`);
    
    if (i > 0) {
      const prevLastIdx = colorNumbers.lastIndexOf(i - 1);
      assert.isBelow(prevLastIdx, firstIdx, `--color${i - 1} should complete before --color${i} starts`);
    }
  }
  
  const percentPattern = /var\(\s*--color\d\s*\)\s+(\d+)%/g;
  const percentMatches = [...gradientStr.matchAll(percentPattern)];
  const percentages = percentMatches.map(match => parseInt(match[1]));
  
  for (let i = 1; i < percentages.length - 1; i += 2) {
    assert.equal(percentages[i], percentages[i + 1], `For hard lines, percentage at position ${i} (${percentages[i]}%) should equal percentage at position ${i + 1} (${percentages[i + 1]}%) where colors meet`);
  }
} else {
  assert.fail('Gradient must use either shorthand format (6 stops with two percentages each) or expanded format (12 stops with one percentage each)');
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Availability Table</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Availability Table</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <table>
            <tbody>
                <tr class="sharp">
                    <th rowspan="2" class="time">9 AM</th>
                    <th>Monday</th>
                    <th>Tuesday</th>
                    <th>Wednesday</th>
                    <th>Thursday</th>
                    <th>Friday</th>
                </tr>
                <tr class="half">
                    <td class="available-0"></td>
                    <td class="available-1"></td>
                    <td class="available-2"></td>
                    <td class="available-2"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">10 AM</th>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>

                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">11 AM</th>
                    <td class="available-0"></td>
                    <td class="available-2"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-0"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-1"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">12 PM</th>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-3"></td>
                    <td class="available-2"></td>
                    <td class="available-0"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-0"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">1 PM</th>
                    <td class="available-2"></td>
                    <td class="available-3"></td>
                    <td class="available-3"></td>
                    <td class="available-2"></td>
                    <td class="available-0"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">2 PM</th>
                    <td class="available-0"></td>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">3 PM</th>
                    <td class="available-3"></td>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-4"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">4 PM</th>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-0"></td>
                    <td class="available-4"></td>
                    <td class="available-2"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">5 PM</th>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                </tr>
                <tr>
                    <td></td>
                    <td></td>
                    <td></td>
                    <td></td>
                    <td></td>
                </tr>
            </tbody>
        </table>
        <div id="legend">
            <span id="legend-title">Availability</span>
            <div id="legend-line"><span>0</span>
                <div id="legend-gradient"></div><span>5+</span>
            </div>
        </div>
    </main>
</body>

</html>
```

```css
:root {
    font-size: 16px;
    font-family: Arial, Helvetica, sans-serif;
    --color0: hsl(0, 0%, 100%);
    --color1: hsl(120, 95%, 90%);
    --color2: hsl(120, 95%, 80%);
    --color3: hsl(120, 95%, 65%);
    --color4: hsl(120, 95%, 40%);
    --color5: hsl(120, 95%, 25%);
    --solid-border: 0.1rem solid black;
    --dashed-border: 0.09rem dashed black;
}

* {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

body {
    background-color: hsl(150, 30%, 75%);
    width: 100vw;
    height: 100vh;
}

main {
    width: 100%;
    height: 100%;
    display: flex;
    flex-direction: column;
    align-items: center;
    justify-content: space-evenly;
}

table {
    border-collapse: collapse;
    table-layout: fixed;
    width: 32rem;
}

tr {
    height: 1.1rem;
}

tr :first-child {
    width: 4rem;
}

td {
    border-right: var(--solid-border);
    border-left: var(--solid-border);
}

.sharp td {
    border-bottom: var(--solid-border);
}

.sharp th:not([class*="time"]) {
    border-bottom: var(--solid-border);
}

.half td {
    border-bottom: var(--dashed-border);
}

tbody :last-child td {
    border: 0;
}

.time {
    text-align: right;
    border: 0;
    padding: 0.4rem;
}

.available-0 {
    background-color: var(--color0);
}

.available-1 {
    background-color: var(--color1);
}

.available-2 {
    background-color: var(--color2);
}

.available-3 {
    background-color: var(--color3);
}

.available-4 {
    background-color: var(--color4);
}

.available-5 {
    background-color: var(--color5);
}

#legend {
    display: flex;
    flex-direction: column;
    align-items: center;
    width: 18rem;
    height: 3.5rem;
    text-align: center;
}

#legend span {
    display: inline-block;
    width: 5rem;
    height: 2rem;
}

#legend-gradient {
    width: 100%;
    height: 60%;
    border: var(--solid-border);
    background-image: linear-gradient(90deg,
            var(--color0) 0% 17%,
            var(--color1) 17% 34%,
            var(--color2) 34% 50%,
            var(--color3) 50% 67%,
            var(--color4) 67% 83%,
            var(--color5) 83% 100%);
}

#legend-line {
    width: 60%;
    height: 100%;
    display: flex;
}
```
