---
id: 642dfb07e7fa6736251541c8
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Ihr Array muss die Größe des Bereichs haben. Sie können diese berechnen, indem Sie die Differenz zwischen `end` und `start` ermitteln und `1` zum Ergebnis hinzufügen.

Geben Sie diese Berechnung als Argument für Ihren `Array()`-Konstruktor weiter.

# --hints--

Sie sollten `start` von `end` subtrahieren.

```js
assert.match(code, /end\s*-\s*start/);
```

Sie sollten `1` zu Ihrer `end - start`-Berechnung hinzufügen.

```js
assert.match(code, /end\s*-\s*start\s*\+\s*1/);
```

Sie sollten Ihre Berechnung an den `Array()`-Konstruktor übergeben.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
const range = (start, end) => Array();
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
