---
id: 642df9df4b5216350de7b0d2
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Sie benötigen eine Funktion, um einen Zahlenbereich zu erzeugen.

Deklarieren Sie eine leere `range`-Funktion, die einen `start`- und einen `end`-Parameter entgegennimmt. Verwenden Sie den `Array()`-Konstruktor und geben Sie implizit ein leeres Array zurück.

# --hints--

Sie sollten eine `range`-Variable deklarieren.

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

Sie sollten `const` verwenden, um Ihre `range`-Variable zu deklarieren.

```js
assert.match(code, /const\s+range/);
```

Ihre `range`-Variable sollte eine Funktion sein.

```js
assert.isFunction(range);
```

Ihre `range`-Funktion sollte Pfeilsyntax verwenden.

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

Ihre `range`-Funktion sollte zuerst einen `start`-Parameter annehmen.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

Ihre `range`-Funktion sollte als zweiten Parameter einen `end` annehmen.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

Ihre `range`-Funktion sollte eine implizite Rückgabe verwenden. Denken Sie daran, dass dies bedeutet, dass Sie keine geschweiften Klammern verwenden.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

Ihre `range`-Funktion sollte den `Array()`-Konstruktor verwenden. Primitive Konstruktoren benötigen das Schlüsselwort `new` nicht.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

Sie sollten dem `Array()`-Konstruktor nichts übergeben.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
