---
id: 642db8c409d9991d0b3b2f0d
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Ihr Projekt beginnt mit einem einfachen HTML-Container und dem dazugehörigen CSS. Ihre erste Aufgabe wird es sein, die Zellen für Ihre Tabelle programmatisch zu erzeugen.

Das globale `window`-Objekt stellt das Browserfenster (oder den Tab) dar. Es hat eine `onload`-Eigenschaft, mit der Sie Verhalten definieren können, wenn das Fenster die gesamte Seite einschließlich Stylesheets und Skripten geladen hat.

Beginnen Sie damit, die `onload`-Eigenschaft von `window` auf eine Pfeilfunktion ohne Parameter zu setzen. Deklarieren Sie in der Funktion eine `container`-Variable und weisen Sie ihr den Wert zu, der durch das Abrufen des Elements mit der `id` von `"container"` erhalten wird.

# --hints--

Sie sollten auf die `onload`-Eigenschaft des `window`-Objekts zugreifen.

```js
assert.match(code, /window\.onload/);
```

Sie sollten die `onload`-Eigenschaft auf eine Funktion setzen.

```js
assert.isFunction(window.onload);
```

Sie sollten die Pfeilsyntax verwenden.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

Ihre `onload`-Funktion sollte keine Parameter haben.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Sie sollten eine `container`-Variable in Ihrer `onload`-Funktion deklarieren.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

Ihre `container`-Variable sollte mit `const` deklariert werden.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Sie sollten `document.getElementById()` verwenden.

```js
assert.match(code, /document\.getElementById\(/);
```

Sie sollten das Element mit der `id` von `"container"` abrufen.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Sie sollten das `#container`-Element `container` zuweisen.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
