---
id: 642dccb78549c9285835ebc2
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Funktionen sind ideal für wiederverwendbare Logik. Wenn eine Funktion selbst Logik wiederverwenden muss, können Sie eine verschachtelte Funktion deklarieren, die diese Logik übernimmt. Hier ist ein Beispiel für eine verschachtelte Funktion:

```js
const outer = () => {
  const inner = () => {

  };
};
```

Deklarieren Sie eine verschachtelte `createLabel`-Funktion mit Pfeilsyntax. Sie soll einen `name`-Parameter erhalten.

# --hints--

Sie sollten eine `createLabel`-Variable in Ihrer `onload`-Funktion deklarieren.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?)?\s*(?:let|var|const)\s+createLabel/);
```

Ihre `createLabel`-Variable sollte nach Ihrer `container`-Variable deklariert werden.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*(?:let|var|const)\s+createLabel/);
```

Ihre `createLabel`-Variable sollte mit `const` deklariert werden.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*const\s+createLabel/);
```

Ihre `createLabel`-Variable sollte eine Pfeilfunktion sein.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*const\s+createLabel\s*=\s*(\(.*\)|[^\s()]+)\s*=>/);
```

Ihre `createLabel`-Funktion sollte einen `name`-Parameter haben.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>/);
```

Ihre `createLabel`-Funktion sollte leer sein.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");

}
--fcc-editable-region--
```
