---
id: 642e02be7845f13b014cd2b0
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Jetzt, wo Sie eine `range`-Funktion haben, können Sie sie auch verwenden, um einen Buchstabenbereich zu erstellen.

Deklarieren Sie eine `charRange`-Funktion mit `const`- und Pfeilsyntax. Sie soll einen `start`- und einen `end`-Parameter annehmen. Die Funktion soll das Ergebnis des Aufrufs von `range()` mit `start` und `end` als Argumente implizit zurückgeben.

# --hints--

Sie sollten eine `charRange`-Variable deklarieren.

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

Ihre `charRange`-Variable sollte mit `const` deklariert werden.

```js
assert.match(code, /const\s+charRange/);
```

Ihre `charRange`-Variable sollte eine Funktion sein.

```js
assert.isFunction(charRange);
```

Ihre `charRange`-Funktion sollte Pfeilsyntax verwenden.

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

Ihre `charRange`-Funktion sollte `start` als ersten Parameter annehmen.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

Ihre `charRange`-Funktion sollte `end` als zweiten Parameter annehmen.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

Ihre `charRange`-Funktion sollte eine implizite Rückgabe verwenden.

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

Ihre `charRange`-Funktion sollte Ihre `range`-Funktion aufrufen.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

Sie sollten `start` und `end` als Argumente an Ihren `range`-Aufruf übergeben.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
