---
id: 643475e13dc727231acd0f72
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Deklarieren Sie nun in Ihrem verschachtelten `.forEach()`-Aufruf eine `input`-Variable. Verwenden Sie die `.createElement()`-Methode des `document`-Objekts, um ein `input`-Element zu erstellen. Setzen Sie das `type`-Attribut auf `"text"` und das `id`-Attribut auf `letter + number`.

# --hints--

Sie sollten eine `input`-Variable deklarieren.

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

Sie sollten `const` verwenden, um Ihre `input`-Variable zu deklarieren.

```js
assert.match(code, /const\s+input/)
```

Sie sollten die `.createElement()`-Methode des `document`-Objekts aufrufen.

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

Sie sollten den String `"input"` an die `.createElement()`-Methode übergeben.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Sie sollten Ihr neues `input`-Element Ihrer `input`-Variable zuweisen.

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Sie sollten auf die `type`-Eigenschaft Ihres `input`-Elements zugreifen.

```js
assert.match(code, /input\.type/);
```

Sie sollten das `type`-Attribut Ihres `input`-Elements auf `"text"` setzen.

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

Sie sollten auf die `id`-Eigenschaft Ihres `input`-Elements zugreifen.

```js
assert.match(code, /input\.id/);
```

Sie sollten das `id`-Attribut Ihres `input`-Elements auf `letter + number` setzen.

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
