---
id: 643715013330824ecaa70442
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Ihre nächste Funktion soll den Medianwert eines Arrays von Zahlen berechnen. Beginnen Sie damit, eine `median` Pfeilfunktion zu deklarieren, die einen `nums` Parameter entgegennimmt.

Deklarieren Sie in der Funktion eine `sorted` Variable und weisen Sie ihr den Wert zu, der durch das Sortieren einer Kopie des `nums` Arrays entsteht.

Sie sollten die `slice()` Methode verwenden, um eine flache Kopie des Arrays zu erstellen.

# --hints--

Sie sollten eine `median` Variable deklarieren.

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

Sie sollten `const` verwenden, um Ihre `median` Variable zu deklarieren.

```js
assert.match(code, /const\s+median/);
```

Ihre `median` Variable sollte eine Funktion sein.

```js
assert.isFunction(median);
```

Ihre `median` Funktion sollte die Pfeilsyntax verwenden.

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

Ihre `median` Funktion sollte einen `nums` Parameter haben.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

Ihre `median` Funktion sollte keine implizite Rückgabe verwenden.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

Ihre `median` Funktion sollte eine `sorted` Variable haben.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

Sie sollten `const` verwenden, um Ihre `sorted` Variable zu deklarieren.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

Sie sollten `.slice()` verwenden, um eine Kopie des `nums` Arrays an `sorted` zuzuweisen.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

Sie sollten die `.sort()` Methode an Ihre `.slice()` Methode anhängen.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

Sie sollten Ihrer `sort` Methode eine Callback-Funktion übergeben, um die Zahlen korrekt aufsteigend zu sortieren. Verwenden Sie eine implizite Rückgabe für bessere Übersicht.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
