---
id: 64496df724dd3716a71fe971
title: Schritt 26
challengeType: 0
dashedName: step-26
---

# --description--

Objekteigenschaften bestehen aus Schlüssel/Wert-Paaren. Sie können Kurzschreibweise-Eigenschaftsnamen verwenden, wenn Sie ein Objektliteral deklarieren. Bei Verwendung der Kurzschreibweise-Eigenschaftssyntax wird der Name der Variablen zum Eigenschaftsschlüssel und ihr Wert zum Eigenschaftswert.

Das folgende Beispiel deklariert ein `user`-Objekt mit den Eigenschaften `userId`, `firstName` und `loggedIn`.

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

Um alle Funktionen Ihrer Tabelle im Blick zu behalten, deklarieren Sie ein `spreadsheetFunctions`-Objekt. Verwenden Sie die Kurzschreibweise-Syntax, um `sum`, `average` und `median` als Eigenschaften des `spreadsheetFunctions`-Objekts zu setzen.

# --hints--

Sie sollten eine `spreadsheetFunctions`-Variable deklarieren.

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

Sie sollten `const` verwenden, um Ihre `spreadsheetFunctions`-Variable zu deklarieren.

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

Ihre `spreadsheetFunctions`-Variable sollte ein Objekt sein.

```js
assert.isObject(spreadsheetFunctions);
```

Ihr `spreadsheetFunctions`-Objekt sollte eine `sum`-Eigenschaft haben.

```js
assert.property(spreadsheetFunctions, "sum");
```

Ihre `sum`-Eigenschaft sollte Ihre `sum`-Funktion sein.

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

Ihr `spreadsheetFunctions`-Objekt sollte eine `average`-Eigenschaft haben.

```js
assert.property(spreadsheetFunctions, "average");
```

Ihre `average`-Eigenschaft sollte Ihre `average`-Funktion sein.

```js
assert.equal(spreadsheetFunctions?.average, average);
```

Ihr `spreadsheetFunctions`-Objekt sollte eine `median`-Eigenschaft haben.

```js
assert.property(spreadsheetFunctions, "median");
```

Ihre `median`-Eigenschaft sollte Ihre `median`-Funktion sein.

```js
assert.equal(spreadsheetFunctions?.median, median);
```

Sie sollten Destrukturierungssyntax verwenden, um Ihre Eigenschaften zuzuweisen.

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
