---
id: 65ae458e23954c3469e0c209
title: Schritt 18
challengeType: 0
dashedName: step-18
---


# --description--

In früheren Projekten haben Sie die `setAttribute`-Methode kennengelernt. Eine andere Möglichkeit, ein Attribut in JavaScript zu aktualisieren, ist die Verwendung der folgenden Syntax:

```js
el.attribute = value;
```

Die Eigenschaftsnamen für HTML-Attribute mit Bindestrich, wie `aria-label`, folgen der camelCase-Schreibweise und werden zu `ariaLabel`.

```js
el.ariaLabel = "Aria Label Value";
```

Setzen Sie das `aria-label`-Attribut für das `input`-Element auf denselben Wert wie das `id`-Attribut.


# --hints--

Sie sollten ein `input.ariaLabel` haben.

```js
assert.match(code, /input\.ariaLabel\s*=/);
```


Sie sollten `input.ariaLabel` denselben Wert wie `input.id` geben.

```js
assert.match(code, /input\.ariaLabel\s*=\s*((letter\s*\+\s*number)|(input.id))/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
     createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {
     const input = document.createElement("input");
     input.type = "text";
     input.id = letter + number;
 
    })   
--fcc-editable-region--
  })
}
```
