---
id: 67051431a73f1ca25d9a6f25
title: プロジェクトアイデアボードを作成する
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `projectStatus` という名前のオブジェクト定数を定義してください。このオブジェクトには `PENDING`、`SUCCESS`、`FAILURE` の3つのキーがあり、それぞれのステータスには `description` キーを持つオブジェクトが代入され、その値はそれぞれ `Pending Execution`、`Executed Successfully`、`Execution Failed` である必要があります。
2. `ProjectIdea` という名前のクラスを定義してください。このクラスには `constructor` があり、`title` と `description` をパラメータとして受け取ります。渡されたパラメータで `title` と `description` のプロパティを初期化してください。また、デフォルトでは `status` の値が代入された `projectStatus.PENDING` というプロパティも持つ必要があります。
3. `updateProjectStatus` クラスの中に `ProjectIdea` という名前のメソッドを定義してください。このメソッドは `newStatus` というパラメータを受け取り、`status` プロパティをその値に更新します。
4. `ProjectIdeaBoard` クラスを定義してください。このクラスには `constructor` があり、`title` を受け取り、`ideas` クラスのインスタンスを格納する空の配列 `ProjectIdea` を初期化します。
5. `pin` クラスの中に `ProjectIdeaBoard` という名前のメソッドを定義してください。このメソッドは `ProjectIdea` クラスのインスタンスを受け取り、そのインスタンスを `ideas` 配列に追加します。
6. `unpin` クラスの中に `ProjectIdeaBoard` という名前のメソッドを定義してください。このメソッドは `ProjectIdea` クラスのインスタンスを受け取り、それを `ideas` 配列から削除します。
7. `count` 配列にあるプロジェクトアイデアの数を返す `ProjectIdeaBoard` という名前のメソッドを定義してください。
8. プロジェクトアイデアの名前、説明、ステータスを以下の形式で返す `formatToString` という名前のメソッドを定義してください。

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

正しいプロジェクトステータスと説明を持つ `projectStatus` という名前のオブジェクト定数を定義してください。

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

`ProjectIdea` というクラスを用意してください。

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

`ProjectIdea` クラスは渡されたパラメータに基づいて `title` プロパティと `description` プロパティを初期化し、ステータスは `projectStatus.PENDING` に設定してください。

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

`ProjectIdea` クラスには `updateProjectStatus` というメソッドを持たせてください。

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

`ProjectIdeaBoard` というクラスを用意してください。

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

`ProjectIdeaBoard` は渡されたパラメータに基づいて `title` プロパティを初期化し、`ideas` クラスのインスタンスを格納する空の配列 `ProjectIdea` を初期化してください。

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

`ProjectIdeaBoard` クラスには `pin` というメソッドを持たせてください。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

`ProjectIdeaBoard` クラスには `unpin` というメソッドを持たせてください。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

`ProjectIdeaBoard` クラスには `count` というメソッドを持たせてください。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

`ProjectIdeaBoard` クラスには `formatToString` というメソッドを持たせてください。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")` は `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }` を返す必要があります。

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

`updateProjectStatus(projectStatus.SUCCESS)` で `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")` を呼び出すとステータスが `{ description: 'Executed Successfully' }` に更新されます。

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

`updateProjectStatus(projectStatus.FAILURE)` で `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")` を呼び出すとステータスが `{ description: 'Execution Failed' }` に更新されます。

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

`updateProjectStatus(projectStatus.SUCCESS)` で `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")` を呼び出すとステータスが `{ description: 'Executed Successfully' }` に更新されます。

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

`ProjectIdea` メソッドを使って `ProjectIdeaBoard` オブジェクトを `pin` にピン留めできる必要があります。

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

`ProjectIdea` メソッドを使って `ProjectIdeaBoard` オブジェクトを `unpin` から外せる必要があります。

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

`new ProjectIdeaBoard("Empty Board")` は `emptyBoard.formatToString()` が空のときに `Empty Board has 0 idea(s)\n` を返す必要があります。

```js

let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

// Create an empty board
let emptyBoard = new ProjectIdeaBoard("Empty Board");

// Attempt to unpin a project from an empty board
emptyBoard.unpin(smartHome);

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

`new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")` は `new ProjectIdeaBoard("Tech Projects Board")` を `techProjects.formatToString()` にピン留めしたときに `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n` を返す必要があります。

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
