---
id: 6852ef95ccc27660275692cd
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

يجب على صنف `Email` تخزين معلومات عن كل بريد إلكتروني عند إنشائه.

داخل صنف `Email`، احذف الكلمة المفتاحية `pass` الحالية وأنشئ الطريقة `__init__` التي تأخذ المعلمات `self`، و`sender`، و`receiver`، و`subject`، و`body`.

# --hints--

يجب تعريف طريقة `__init__` داخل صنف `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

يجب أن تحتوي طريقة `__init__` على المعلمات: `self`، و`sender`، و`receiver`، و`subject`، و`body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
--fcc-editable-region--
    pass
--fcc-editable-region--
```
