---
id: 6852f11d412c29665985a915
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

بما أن المستخدمين يمكنهم إرسال رسائل البريد الإلكتروني، فهم بحاجة إلى طريقة لتخزين الرسائل التي يستقبلونها. يجب أن يكون لكل مستخدم صندوق وارد لجمع رسائل البريد الخاصة به.

أضف خاصية `inbox` إلى طريقة `__init__` في صنف `User` وابدأها كقائمة فارغة `[]`. سيخزن هذا جميع الرسائل التي يستقبلها المستخدم.

# --hints--

يجب أن تضيف `self.inbox = []` إلى طريقة `__init__` في صنف `User`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
        
class User:
    def __init__(self, name):
        self.name = name
--fcc-editable-region--
        
--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
