---
id: 6852f0c3fcedb964ab959921
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

اختبر الآن أن الخاصية `read` أُضيفت بشكل صحيح إلى صنف `Email`. بما أن لديك كائن بريد إلكتروني من الخطوات السابقة، اطبع الخاصية `read` لترى أنها الآن `False` افتراضيًا.

# --hints--

يجب أن تطبع `email_obj.read` لاختبار الخاصية read.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
--fcc-editable-region--

--fcc-editable-region--
```
