---
id: 6852f06fdde1b4634a263f03
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

اختبر الآن صنف `Email` بإنشاء كائن بريد إلكتروني والتحقق من أنه يخزن المعلومات بشكل صحيح. اطبع بعض الخصائص لتتأكد من أن كل شيء يعمل.

أنشئ كائن بريد إلكتروني باسم `email_obj` مع `alice@example.com` كمرسل، و`bob@example.com` كمستقبل، و`Hello` كموضوع، و`Hi Bob!` كمحتوى. ثم اطبع خاصيتي `sender` و`subject` في عبارات طباعة منفصلة للتحقق من تخزينهما بشكل صحيح.

# --hints--

يجب أن تنشئ كائن بريد إلكتروني باسم `email_obj` بالمعلمات المحددة.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

يجب أن تطبع `email_obj.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

يجب أن تطبع `email_obj.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
--fcc-editable-region--

--fcc-editable-region--
```
