---
id: 6852ea0c3cadb54e897acdac
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستبني *محاكي بريد إلكتروني* يحاكي إرسال واستقبال وإدارة الرسائل الإلكترونية بين مستخدمين مختلفين. ستتعلم عن الأصناف، والكائنات، وكيفية تنظيم الكود بطريقة موجهة للكائنات.

ابدأ بإنشاء صنف اسمه `Email` باستخدام الكلمة المفتاحية `class`.

# --hints--

يجب أن تنشئ صنفًا اسمه `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
