---
id: 6852f01116912861c7d1b0e6
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

داخل طريقة `__init__`، تحتاج إلى تعيين قيم المعلمات إلى خاصيات الكائن حتى يتمكن كل بريد إلكتروني من تخزين معلوماته.

داخل طريقة `__init__` الخاصة بك، احذف الكلمة المفتاحية `pass` وقم بتعيين المعلمات إلى `self.sender`، و`self.receiver`، و`self.subject`، و`self.body`.

# --hints--

يجب تعيين المعلمة `sender` إلى `self.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

يجب تعيين المعلمة `receiver` إلى `self.receiver`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

يجب تعيين المعلمة `subject` إلى `self.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

يجب تعيين المعلمة `body` إلى `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
