---
id: 6853f4ded9f84896ff182dd8
title: الخطوة 31
challengeType: 20
dashedName: step-31
---

# --description--

بعد التحقق من أن صندوق الوارد فارغ، اطبع الرسالة `\nYour Emails:`. بعد ذلك، قم بالتكرار على جميع الرسائل في صندوق الوارد باستخدام حلقة `for` مع التعداد. هذه هي صيغة التعداد مع رقم بداية: `enumerate(iterable, start=start_number)`.

استخدم التعداد لترقيم الرسائل بدءًا من `1`. استخدم `i` (الفهرس) و`email` (كائن البريد الإلكتروني) كمتغيرات التكرار.

داخل الحلقة، اطبع سلسلة نصية بصيغة f-string تحتوي على فهرس التكرار متبوعًا بـ `.`, ثم مسافة وتمثيل السلسلة النصية لكائن البريد الإلكتروني بهذا الشكل:

```md
i. string_representation
```

حيث `i` هو الفهرس و`string_representation` هو تمثيل كائن البريد الإلكتروني.

# --hints--

يجب أن تطبع القسم العلوي `\nYour Emails:` خارج جملة `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

يجب أن تحتوي على حلقة `for` تتكرر على `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

يجب أن تستخدم حلقة `for` المتغيرين `i` و`email` للتكرار على `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

داخل الحلقة، يجب أن تطبع ملخص كل بريد إلكتروني مع بادئة مرقمة متبوعة بالسلسلة النصية المنسقة من طريقة `__str__`. لا تنسَ إضافة `.` بعد الرقم.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
        
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
--fcc-editable-region--
        
--fcc-editable-region--
```
