---
id: 6853f4584f508e95cc7a9f4d
title: الخطوة 30
challengeType: 20
dashedName: step-30
---

# --description--

الآن ستُنشئ طريقة لعرض جميع رسائل البريد الإلكتروني في صندوق الوارد. يجب أن تتعامل هذه الطريقة مع حالة فراغ صندوق الوارد وتعرض قائمة مرقمة بالرسائل عند وجود رسائل.

أضف طريقة تسمى `list_emails` إلى صنف `Inbox` الخاص بك تأخذ `self` كمعلمة. أولًا، أنشئ جملة `if` للتحقق مما إذا كان صندوق الوارد فارغًا عن طريق اختبار قائمة `emails`. إذا كانت فارغة، اطبع الرسالة `Your inbox is empty.\n`. أضف جملة `return` للخروج من الطريقة.

# --hints--

يجب أن تعرف طريقة باسم `list_emails` في صنف `Inbox` تأخذ `self` كمعلمة.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

يجب أن تتحقق مما إذا كانت `self.emails` فارغة باستخدام `if not self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

يجب أن تطبع الرسالة `Your inbox is empty.\n` إذا لم توجد رسائل.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

يجب أن تحتوي جملة `if` على جملة `return` في نهايتها.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
