---
id: 6859400f8f7f872e50e1238c
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

يحتاج صنف `User` إلى طريقة لإرسال رسائل بريد إلكتروني إلى مستخدمين آخرين. عند إرسال بريد إلكتروني، ستُنشئ كائن `Email` جديد وتضيفه إلى صندوق الوارد الخاص بالمستلم.

أنشئ طريقة باسم `send_email` في صنف `User` الخاص بك. يجب أن تأخذ هذه الطريقة معلمات لـ `receiver` و`subject` و`body`.

داخل الطريقة، أنشئ كائن `Email` جديد بالقيم التالية للمعلمات وعيّنه في متغير باسم `email`:

- `sender`: `self` (المستخدم الحالي)
- `receiver`: `receiver` (المستخدم الذي يستلم البريد الإلكتروني)
- `subject`: `subject` (موضوع البريد الإلكتروني)
- `body`: `body` (محتوى البريد الإلكتروني)

# --hints--

يجب تعريف طريقة `send_email` في صنف `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

يجب أن تحتوي طريقة `send_email` على المعلمات: `self` و`receiver` و`subject` و`body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

يجب أن تنشئ الطريقة كائن `Email` باسم `email` بالمعلمات الصحيحة.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        
--fcc-editable-region--
    
--fcc-editable-region--
```
