---
id: 68593ea22eac7e228ed585ce
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

يجب أن تتعقب الرسائل الإلكترونية ما إذا كانت قد قُرئت أم لا. أضف خاصية `read` إلى طريقة `__init__` وعيّنها إلى `False` افتراضيًا، لأن الرسائل الجديدة تبدأ كغير مقروءة.

# --hints--

يجب تعيين `False` إلى `self.read` في طريقة `__init__`.

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
--fcc-editable-region--
        
--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
