---
id: 685c1eba966c6396850736ea
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

بينما يمكن للمستخدمين إرسال رسائل إلكترونية وامتلاك صناديق واردة، هناك حاجة إلى صنف مخصص لإدارة عمليات الصندوق الوارد بكفاءة.

سيخزن صنف `Inbox` قائمة من الرسائل الإلكترونية ويوفر طرقًا لإضافة رسائل جديدة، وعرض جميع الرسائل، وإدارتها.

أنشئ صنفًا جديدًا باسم `Inbox` مع طريقة `__init__` تهيئ قائمة فارغة باسم `emails`.

# --hints--

يجب أن تنشئ صنف `Inbox`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

يجب أن يحتوي صنف `Inbox` على طريقة `__init__` تأخذ `self` كمعلمة.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

يجب أن تهيئ `self.emails = []` في طريقة `__init__` الخاصة بصنف `Inbox`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
