---
id: 685c1e32254e5993a5c8823e
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

الآن بعد أن أعددت صنف `Email` مبدئيًا، حان الوقت لإنشاء مستخدمين يمكنهم إرسال واستقبال الرسائل الإلكترونية. سيكون لكل مستخدم اسم وسيتمكن من تنفيذ عمليات البريد الإلكتروني.

أنشئ صنف `User` مع طريقة `__init__` تأخذ `self` و`name` كمعلمات. داخل الصنف، قم بتعيين المعلمة `name` إلى `self.name`.

# --hints--

يجب أن تنشئ صنف `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

يجب أن يحتوي صنف `User` على طريقة `__init__` تأخذ `self` و`name` كمعلمات.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

يجب أن تعيّن المعلمة `name` إلى `self.name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
