---
id: 688754d6d59578df1313cb3c
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

حان الوقت الآن لربط صنفي User و Inbox حتى تُرسل الرسائل الإلكترونية فعليًا!

يجب أن يمتلك المستخدمون كائنات Inbox صحيحة بدلًا من القوائم البسيطة، ويجب أن تُرسل الطريقة `send_email` الرسائل إلى صندوق الوارد الخاص بالمستلم.

حدّث دالة `__init__` في صنف `User` بحيث يحصل كل مستخدم على كائن Inbox فعلي بدلًا من القائمة الفارغة التي أنشئت سابقًا.

# --hints--

يجب تحديث صنف User لتهيئة `self.inbox = Inbox()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
--fcc-editable-region--
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
