---
id: 68875554c6c32de788afc64e
title: الخطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

يحتاج المستخدمون إلى طريقة لقراءة رسائل بريدهم الإلكتروني بشكل صحيح.

أضف طريقة `display_full_email` إلى صنف `Email` تأخذ فقط `self` كمعلمة. داخل هذه الطريقة، استدعِ طريقة `mark_as_read()` لوضع علامة على البريد الإلكتروني كمقروء عند عرضه.

تُستخدم هذه الطريقة لعرض المحتوى الكامل للبريد الإلكتروني بطريقة منسقة وجميلة.

# --hints--

يجب تعريف طريقة `display_full_email` في صنف `Email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

يجب أن تأخذ طريقة `display_full_email` فقط `self` كمعلمة.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

يجب أن تستدعي الطريقة `self.mark_as_read()` لوضع علامة على البريد الإلكتروني كمقروء.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
