---
id: 68b96d6a290d52d5e57a7485
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

احذف `email_obj` وعبارات الطباعة التالية.

# --hints--

يجب ألا تحتوي شفرتك على سطر إنشاء `email_obj`.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`));
    }
})
```

يجب ألا تحتوي شفرتك على `print(email_obj.sender)`.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.sender)")`));
    }
})
```

يجب ألا تحتوي شفرتك على `print(email_obj.subject)`.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.subject)")`));
    }
})
```

يجب ألا تحتوي شفرتك على `print(email_obj.read)`.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

يجب ألا تحتوي شفرتك على `email_obj.mark_as_read()`.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
email_obj.mark_as_read()
print(email_obj.read)
--fcc-editable-region--
```
