---
id: 68b930ee0e7b8fb81895e59a
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

الآن ستضيف طريقة لوضع علامة على البريد الإلكتروني كمقروء. عندما يفتح شخص ما بريدًا إلكترونيًا، تريد تغيير حالته من غير مقروء إلى مقروء. ستكون هذه الطريقة بسيطة - فقط تحتاج إلى تعيين الخاصية `read` إلى `True`.

أضف طريقة تسمى `mark_as_read` إلى صنف `Email` الخاص بك. يجب أن تأخذ هذه الطريقة `self` فقط كمعلمة وتعيّن الخاصية `read` إلى `True`.

# --hints--

يجب أن تعرف طريقة باسم `mark_as_read` في صنف `Email`.
يجب أن تأخذ فقط `self` كمعلمة.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

يجب أن تعيّن الطريقة `self.read` إلى `True`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False
--fcc-editable-region--
    
--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
