---
id: 68b9366fae8b13c1278761b1
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

تحتاج صندوق الوارد إلى طريقة لاستقبال الرسائل الجديدة. عندما يرسل شخص رسالة إلى مستخدم، يجب إضافتها إلى صندوق الوارد الخاص به.

أضف طريقة تسمى `receive_email` إلى صنف `Inbox` الخاص بك تأخذ `self` وكائن رسالة البريد الإلكتروني `email` كمعلمات. داخل جسم الطريقة، أضف `email` إلى قائمة `emails` باستخدام طريقة `append`.

# --hints--

يجب أن تعرف طريقة باسم `receive_email` في صنف `Inbox` تأخذ `self` و `email` كمعلمات.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

يجب أن تضيف `email` إلى `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []
--fcc-editable-region--
    
--fcc-editable-region--
```
