---
id: 68b936bf4f47d6c29a669f99
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

لا تزال داخل طريقة `send_email`، استدعِ طريقة `receive_email` الخاصة بـ `inbox` الخاص بـ `receiver` لإضافة البريد الإلكتروني إلى صندوق الوارد الخاص بالمستلم.

# --hints--

يجب أن تضيف `receiver.inbox.receive_email(email)` إلى طريقة `send_email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_stmt("receiver.inbox.receive_email(email)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()
        
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
--fcc-editable-region--
        
--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)
```
