---
id: 68b9350183ff05bd106c79e4
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

اختبر الآن الطريقة `mark_as_read` التي أنشأتها في صنف `Email`. استخدم الطريقة على كائن `Email` الموجود `email_obj` لتغيير الحالة. بعد تعليم البريد الإلكتروني كمقروء، اطبع الخاصية `read` لكائن `email_obj` لتأكيد التغيير.

# --hints--

يجب استدعاء الطريقة `mark_as_read` على `email_obj` لتعليم البريد الإلكتروني كمقروء.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

يجب طباعة الخاصية `read` لكائن `email_obj` بعد استدعاء الطريقة `mark_as_read`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
--fcc-editable-region--

--fcc-editable-region--
```
