---
id: 68be6af1f43cc44541debe41
title: الخطوة 37
challengeType: 20
dashedName: step-37
---

# --description--

يجب أن تتعامل مع الحالة التي يحاول فيها المستخدم حذف بريد إلكتروني باستخدام فهرس غير صالح تمامًا كما فعلت في طريقة `read_email`.

داخل طريقة `delete_email`:

- حوّل معلمة `index` التي تبدأ من 1 إلى فهرس يبدأ من 0 بطرح 1 وتخزينه في متغير يسمى `actual_index`.
- أنشئ جملة if الشرطية للتحقق مما إذا كان `actual_index` أقل من 0 أو أكبر من أو يساوي طول قائمة `self.emails`. إذا كان كذلك، اطبع الرسالة `Invalid email number.\n` واستخدم `return` للخروج من الطريقة.

# --hints--


يجب أن تطرح 1 من معلمة `index` وتخزنها في متغير يسمى `actual_index`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_body().find_variable("actual_index").is_equivalent("actual_index = index - 1")`))
    }
})
```

يجب أن تنشئ جملة if الشرطية للتحقق مما إذا كان `actual_index` أقل من 0 أو أكبر من أو يساوي طول قائمة `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_conditions()[0].is_equivalent("actual_index < 0 or actual_index >= len(self.emails)")`))
    }
})
```

داخل جملة `if`، يجب طباعة الرسالة `Invalid email number.\n` واستخدام `return` للخروج من الطريقة.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[0].is_equivalent("print('Invalid email number.\\\\n')")`))
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()
            
    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
--fcc-editable-region--
        
--fcc-editable-region--

```
