---
id: 66bb3e20d3dc5b6d0a21f5dd
title: بناء عداد في الوقت الحقيقي
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

في هذا المختبر، ستبني عدادًا لعدد الأحرف في الوقت الحقيقي باستخدام JavaScript. سيعرض عداد الأحرف عدد الأحرف المُدخلة في عنصر `textarea`. سيُحدَّث العداد في الوقت الحقيقي أثناء كتابة المستخدم في `textarea`.

**الهدف:** أتمم قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يحتوي على عنصر `textarea` مع خاصية `id` تساوي `text-input`.
2. يجب أن يوجد عنصر `p` مع خاصية `id` تساوي `char-count`. يجب أن يحتوي في البداية على النص `Character Count: 0/50`. يجب أن يُستبدل نص العنصر النائب هذا بعدد الأحرف الحالي في عنصر `textarea`.
3. عندما يحتوي عنصر `#text-input` على النص `hello coder` يجب أن يحتوي عنصر `#char-count` على النص `"Character Count: 11/50"`
4. عندما يحتوي عنصر `#text-input` على النص `i am a programmer` يجب أن يحتوي عنصر `#char-count` على النص `"Character Count: 17/50"`
5. عندما يحتوي عنصر `#text-input` على النص `hello world` يجب أن يحتوي عنصر `#char-count` على النص `"Character Count: 11/50"`
6. عندما يحتوي عنصر `#text-input` على النص `I am learning a new language and it is called c++.` يجب أن يحتوي عنصر `#char-count` على النص `"Character Count: 50/50"`. **ملاحظة**: رغم أن خاصية `maxlength` تحقق الوظيفة المطلوبة، هذا التمرين يتطلب حلًا باستخدام JavaScript فقط.
7. يجب ألا يتمكن المستخدم من إدخال أكثر من `50` حرفًا في عنصر `textarea`. عندما يصل عدد الأحرف إلى `50`، يجب أن يُقص أي مُدخل زائد تلقائيًا ويُعرض النص `Character Count: 50/50` باللون `red`.

**ملاحظة:** تأكد من ربط ورقة الأنماط وملف JavaScript في ملف HTML الخاص بك.

# --hints--

يجب أن يكون لديك عنصر `textarea` بالمعرف `id` الخاص به `text-input`.

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

يجب أن يكون لديك عنصر `p` بالمعرف `id` الخاص به `char-count`.

```js
assert.exists("p#char-count");
```

يجب أن يحتوي العنصر `#char-count` في البداية على النص `Character Count: 0/50`.

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

عندما يحتوي العنصر `#text-input` على النص `hello coder` يجب أن يحتوي العنصر `#char-count` على النص `"Character Count: 11/50"`

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

عندما يصل عدد الأحرف إلى `50`، يجب عرض النص باللون `red`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

إذا كان عداد الأحرف أكبر من أو يساوي `50`، يجب تقليم المُدخَل الإضافي.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
