---
id: 587d778b367417b2b2512aa7
title: 더 나은 접근성을 위해 라디오 버튼을 필드셋 요소로 감싸기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVefw'
forumTopicId: 301030
dashedName: wrap-radio-buttons-in-a-fieldset-element-for-better-accessibility
---

# --description--

다음 폼 주제는 라디오 버튼의 접근성에 관한 내용입니다. 각 선택지는 `label`와 `for` 속성으로 `id`에 연결되어 있는데, 이는 지난 도전 과제에서 다뤘습니다. 라디오 버튼은 사용자가 하나를 선택해야 하는 그룹으로 자주 나오기 때문에, 선택지가 집합의 일부임을 의미적으로 보여주는 방법이 있습니다.

이 목적을 위해 전체 라디오 버튼 그룹을 `fieldset` 태그로 감쌉니다. 보통 `legend` 태그를 사용해 그룹에 대한 설명을 제공하는데, 화면 읽기 프로그램이 `fieldset` 요소의 각 선택지에 대해 이를 읽어줍니다.

선택지가 성별 선택처럼 자명한 경우에는 `fieldset` 래퍼와 `legend` 태그가 필요하지 않습니다. 각 라디오 버튼에 `label`와 `for` 속성을 사용하는 것으로 충분합니다.

예시는 다음과 같습니다:

```html
<form>
  <fieldset>
    <legend>Choose one of these three items:</legend>
    <input id="one" type="radio" name="items" value="one">
    <label for="one">Choice One</label><br>
    <input id="two" type="radio" name="items" value="two">
    <label for="two">Choice Two</label><br>
    <input id="three" type="radio" name="items" value="three">
    <label for="three">Choice Three</label>
  </fieldset>
</form>
```

# --instructions--

Camper Cat은 이메일 목록에 가입할 때 사용자들의 닌자 레벨 정보를 원합니다. 그는 라디오 버튼 세트를 추가했고, 지난 수업에서 각 선택지에 `label` 태그와 `for` 속성을 사용하는 법을 배웠습니다. 잘했어요 Camper Cat! 하지만 그의 코드는 아직 개선이 필요합니다. 라디오 버튼을 감싸는 `div` 태그를 `fieldset` 태그로 바꾸고, 그 안의 `p` 태그를 `legend` 태그로 변경하세요.

# --hints--

라디오 버튼 세트 주위에 `fieldset` 태그가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('fieldset'),1);
```

`fieldset` 요소는 닫는 태그가 있어야 합니다.

```js
assert.isTrue(code.match(/<\/fieldset>/g)?.length === code.match(/<fieldset>/g)?.length);
```

사용자가 닌자 레벨이 무엇인지 묻는 텍스트 주위에 `legend` 태그가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('legend') ,1);
```

코드에 `div` 태그가 없어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div') ,0);
```

사용자가 닌자 레벨이 무엇인지 묻는 텍스트 주위에 더 이상 `p` 태그가 없어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('p') ,4);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <!-- Only change code below this line -->
      <div>
        <p>What level ninja are you?</p>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </div>
      <!-- Only change code above this line -->


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">

      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </fieldset>

      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
