---
id: 587d778f367417b2b2512aac
title: 충분한 명암 대비로 색맹 문제 피하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMEUw'
forumTopicId: 301012
dashedName: avoid-colorblindness-issues-by-using-sufficient-contrast
---

# --description--

색상은 시각 디자인에서 큰 부분을 차지하지만, 사용 시 두 가지 접근성 문제가 발생합니다. 첫째, 중요한 정보를 전달할 때 색상만 사용하는 것은 화면 읽기 프로그램 사용자가 볼 수 없기 때문에 적절하지 않습니다. 둘째, 전경과 배경 색상은 색맹인이 구분할 수 있도록 충분한 명암 대비가 필요합니다.

이전 도전 과제에서는 첫 번째 문제를 해결하기 위해 텍스트 대체 수단을 다뤘습니다. 마지막 도전 과제에서는 두 번째 문제를 돕기 위해 명암 대비 검사 도구를 소개했습니다. WCAG 권장 명암 대비율 4.5:1은 색상 사용뿐 아니라 그레이스케일 조합에도 적용됩니다.

색맹인은 보통 색조에서, 때로는 밝기에서 일부 색상을 구분하는 데 어려움을 겪습니다. 명암 대비율은 전경과 배경 색상의 상대 밝기(또는 밝기) 값을 사용해 계산된다는 점을 기억하세요.

실제로 4.5:1 명암 대비율은 더 어두운 색상에 검은색을 더하는 음영 처리와 더 밝은 색상에 흰색을 더하는 색조 처리로 달성할 수 있습니다. 색상환에서 더 어두운 음영은 파랑, 보라, 마젠타, 빨강 계열로 간주되며, 더 밝은 색조는 주황, 노랑, 초록, 청록 계열입니다.

# --instructions--

Camper Cat은 블로그 텍스트와 배경에 색상을 사용해 실험 중인데, 현재 녹색빛 `background-color`과 밤색 텍스트 `color` 조합의 명암 대비율이 2.5:1에 불과합니다. CSS `hsl()` 속성(색상, 채도, 밝기)을 사용해 선언했으므로 세 번째 인수를 변경해 색상의 밝기를 쉽게 조절할 수 있습니다. `background-color` 밝기 값을 35%에서 55%로 올리고, `color` 밝기 값을 20%에서 15%로 낮추세요. 이렇게 하면 명암 대비가 5.9:1로 향상됩니다.

# --hints--

텍스트 `color` 속성의 밝기 값만 15%로 변경하세요.

```js
assert.match(code ,/color:\s*?hsl\(0,\s*?55%,\s*?15%\)/gi);
```

`background-color` 속성의 밝기 값만 55%로 변경하세요.

```js
assert.match(code ,/background-color:\s*?hsl\(120,\s*?25%,\s*?55%\)/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 20%);
    background-color: hsl(120, 25%, 35%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 15%);
    background-color: hsl(120, 25%, 55%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
