---
id: 587d774e367417b2b2512a9f
title: main 요소를 사용해 콘텐츠로 바로 이동하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7zuE'
forumTopicId: 301018
dashedName: jump-straight-to-the-content-using-the-main-element
---

# --description--

HTML5는 개발자에게 더 많은 선택지를 제공하면서 접근성 기능도 포함하는 여러 새로운 요소를 도입했습니다. 이 태그에는 `main`, `header`, `footer`, `nav`, `article`, `section` 등이 있습니다.

기본적으로 브라우저는 이러한 요소를 겸손한 `div`와 비슷하게 렌더링합니다. 하지만 적절한 곳에 사용하면 마크업에 추가적인 의미를 부여합니다. 태그 이름만으로도 포함된 정보 유형을 나타내어 그 콘텐츠에 의미론적 의미를 더합니다. 보조 기술은 이 정보를 접근하여 사용자에게 더 나은 페이지 요약이나 탐색 옵션을 제공할 수 있습니다.

`main` 요소는 (말 그대로) 주요 콘텐츠를 감싸는 데 사용되며, 페이지당 하나만 있어야 합니다. 페이지의 중심 주제와 관련된 정보를 둘러싸도록 설계되었습니다. 내비게이션 링크나 배너처럼 페이지마다 반복되는 항목을 포함해서는 안 됩니다.

`main` 태그는 보조 기술이 주요 콘텐츠로 빠르게 이동할 수 있도록 하는 내장 랜드마크 기능도 갖추고 있습니다. 페이지 상단에 "주요 콘텐츠로 건너뛰기" 링크를 본 적이 있다면, `main` 태그를 사용하면 보조 장치에 자동으로 그 기능이 제공됩니다.

# --instructions--

Camper Cat가 닌자 무기 페이지에 대한 큰 아이디어를 가지고 있습니다. `header`과 `footer` 사이에 여는 `main` 태그와 닫는 `main` 태그를 추가해 마크업을 설정하도록 도와주세요 (다른 도전 과제에서 다룹니다). `main` 태그는 지금은 비워 두세요.

# --hints--

코드에는 `main` 태그가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`main` 태그는 닫는 `header` 태그와 여는 `footer` 태그 사이에 있어야 합니다.

```js
assert.match(code,/<\/header>\s*?<main>\s*?<\/main>/gi);
```

# --seed--

## --seed-contents--

```html
<header>
  <h1>Weapons of the Ninja</h1>
</header>



<footer></footer>
```

# --solutions--

```html
<header>
  <h1>Weapons of the Ninja</h1>
</header>
<main>

</main>
<footer></footer>
```
