---
id: 587d778d367417b2b2512aaa
title: 화면 읽기 프로그램에만 보이도록 사용자 지정 CSS로 요소 숨기기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJ8QGkhJ'
forumTopicId: 301020
dashedName: make-elements-only-visible-to-a-screen-reader-by-using-custom-css
---

# --description--

지금까지 적용한 모든 접근성 도전 과제에서 CSS를 사용하지 않은 것을 눈치채셨나요? 이는 시각적 디자인 요소를 도입하기 전에 논리적인 문서 개요와 의미론적으로 적절한 태그를 콘텐츠 주위에 사용하는 것이 얼마나 중요한지 보여줍니다.

하지만 CSS의 마법은 화면 읽기 프로그램 전용으로만 보이게 하려는 콘텐츠를 시각적으로 숨길 때 페이지의 접근성을 향상할 수도 있습니다. 예를 들어, 정보가 시각적 형식(차트 등)으로 제공되지만 화면 읽기 프로그램 사용자는 데이터를 접근하기 위해 대체 표현(표 등)이 필요할 때 그렇습니다. CSS는 화면 읽기 프로그램 전용 요소를 브라우저 창의 시각 영역 밖으로 위치시키는 데 사용됩니다.

다음은 이를 수행하는 CSS 규칙의 예입니다:

```css
.sr-only {
  position: absolute;
  left: -10000px;
  width: 1px;
  height: 1px;
  top: auto;
  overflow: hidden;
}
```

**참고:** 다음 CSS 방법들은 같은 효과를 내지 않습니다:

<ul>
<li><code>display: none;</code> 또는 <code>visibility: hidden;</code>는 화면 읽기 프로그램 사용자를 포함해 모두에게 콘텐츠를 숨깁니다</li>
<li><code>width: 0px; height: 0px;</code>와 같은 픽셀 크기 0 값은 해당 요소를 문서 흐름에서 제거해 화면 읽기 프로그램이 무시합니다</li>
</ul>

# --instructions--

Camper Cat이 자신의 교육 페이지에 정말 멋진 누적 막대 그래프를 만들고, 시각 장애가 있는 사용자를 위해 데이터를 표에 넣었습니다. 표에는 이미 `sr-only` 클래스가 있지만 CSS 규칙은 아직 채워지지 않았습니다. `position`에 `absolute` 값을, `left`에 `-10000px` 값을, 그리고 `width`와 `height`에는 모두 `1px` 값을 지정하세요.

# --hints--

코드는 `position` 클래스의 `sr-only` 속성을 `absolute` 값으로 설정해야 합니다.

```js
const srOnly = document.querySelector('.sr-only');
const position = window.getComputedStyle(srOnly).position; 
assert.equal(position, 'absolute');
```

코드는 `left` 클래스의 `sr-only` 속성을 `-10000px` 값으로 설정해야 합니다.

```js
const srOnly = document.querySelector('.sr-only');
const left = window.getComputedStyle(srOnly).left;
assert.equal(left, '-10000px');
```

코드는 `width` 클래스의 `sr-only` 속성을 `1` 픽셀 값으로 설정해야 합니다.

```js
assert.match(code , /width:\s*?1px/gi);
```

코드는 `height` 클래스의 `sr-only` 속성을 `1` 픽셀 값으로 설정해야 합니다.

```js
assert.match(code , /height:\s*?1px/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  .sr-only {
    position: ;
    left: ;
    width: ;
    height: ;
    top: auto;
    overflow: hidden;
  }
  </style>
</head>
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <section>
    <h2>Master Camper Cat's Beginner Three Week Training Program</h2>
    <figure>
      <!-- Stacked bar chart of weekly training -->
      <p>[Stacked bar chart]</p>
      <br />
      <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
    </figure>
    <table class="sr-only">
      <caption>Hours of Weekly Training in Stealth, Combat, and Weapons</caption>
      <thead>
        <tr>
          <th></th>
          <th scope="col">Stealth &amp; Agility</th>
          <th scope="col">Combat</th>
          <th scope="col">Weapons</th>
          <th scope="col">Total</th>
        </tr>
      </thead>
      <tbody>
        <tr>
          <th scope="row">Week One</th>
          <td>3</td>
          <td>5</td>
          <td>2</td>
          <td>10</td>
        </tr>
        <tr>
          <th scope="row">Week Two</th>
          <td>4</td>
          <td>5</td>
          <td>3</td>
          <td>12</td>
        </tr>
        <tr>
          <th scope="row">Week Three</th>
          <td>4</td>
          <td>6</td>
          <td>3</td>
          <td>13</td>
        </tr>
      </tbody>
    </table>
  </section>
  <section id="stealth">
    <h2>Stealth &amp; Agility Training</h2>
    <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
    <article><h3>No training is NP-complete without parkour</h3></article>
  </section>
  <section id="combat">
    <h2>Combat Training</h2>
    <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
    <article><h3>Goodbye, world: 5 proven ways to knock out an opponent</h3></article>
  </section>
  <section id="weapons">
    <h2>Weapons Training</h2>
    <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
    <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  .sr-only {
    position: absolute;
    left: -10000px;
    width: 1px;
    height: 1px;
    top: auto;
    overflow: hidden;
  }
  </style>
</head>
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <section>
    <h2>Master Camper Cat's Beginner Three Week Training Program</h2>
    <figure>
      <!-- Stacked bar chart of weekly training -->
      <p>[Stacked bar chart]</p>
      <br />
      <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
    </figure>
    <table class="sr-only">
      <caption>Hours of Weekly Training in Stealth, Combat, and Weapons</caption>
      <thead>
        <tr>
          <th></th>
          <th scope="col">Stealth &amp; Agility</th>
          <th scope="col">Combat</th>
          <th scope="col">Weapons</th>
          <th scope="col">Total</th>
        </tr>
      </thead>
      <tbody>
        <tr>
          <th scope="row">Week One</th>
          <td>3</td>
          <td>5</td>
          <td>2</td>
          <td>10</td>
        </tr>
        <tr>
          <th scope="row">Week Two</th>
          <td>4</td>
          <td>5</td>
          <td>3</td>
          <td>12</td>
        </tr>
        <tr>
          <th scope="row">Week Three</th>
          <td>4</td>
          <td>6</td>
          <td>3</td>
          <td>13</td>
        </tr>
      </tbody>
    </table>
  </section>
  <section id="stealth">
    <h2>Stealth &amp; Agility Training</h2>
    <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
    <article><h3>No training is NP-complete without parkour</h3></article>
  </section>
  <section id="combat">
    <h2>Combat Training</h2>
    <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
    <article><h3>Goodbye, world: 5 proven ways to knock out an opponent</h3></article>
  </section>
  <section id="weapons">
    <h2>Weapons Training</h2>
    <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
    <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
