---
id: 587d774c367417b2b2512a9d
title: Alt 텍스트를 비워둬야 할 때 알기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

지난 도전 과제에서 `alt` 태그를 사용할 때 `img` 속성을 포함하는 것이 필수임을 배웠습니다. 하지만 때로는 이미지가 이미 설명하는 캡션과 함께 그룹화되어 있거나 장식용으로만 사용되기도 합니다. 이런 경우 `alt` 텍스트는 중복되거나 불필요해 보일 수 있습니다.

이미지가 텍스트 내용으로 설명되어 있거나 페이지에 의미를 더하지 않는 경우에도 `img`에는 `alt` 속성이 필요하지만, 빈 문자열로 설정할 수 있습니다. 다음은 예시입니다:

```html
<img src="visualDecoration.jpeg" alt="">
```

배경 이미지는 보통 '장식용' 레이블에 해당합니다. 하지만 일반적으로 CSS 규칙으로 적용되므로 화면 읽기 프로그램이 처리하는 마크업의 일부가 아닙니다.

**참고:** 캡션이 있는 이미지의 경우, 이미지 내용을 검색 엔진이 분류하는 데 도움이 되므로 `alt` 텍스트를 포함하는 것이 좋습니다.

# --instructions--

Camper Cat이 자신의 웹사이트 블로그 부분을 위한 기본 페이지를 코딩했습니다. 두 개의 기사 사이에 장식용 사무라이 검 이미지를 넣어 시각적 구분을 추가하려고 합니다. `alt` 태그에 `img` 속성을 추가하고 빈 문자열로 설정하세요. (이미지 `src`가 실제 파일에 링크되어 있지 않으니, 화면에 검이 보이지 않는다고 걱정하지 마세요.)

# --hints--

`img` 태그에 `alt` 속성이 있어야 합니다.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

`alt` 속성은 빈 문자열로 설정하세요.

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```
