---
id: 587d778b367417b2b2512aa8
title: 접근 가능한 날짜 선택기 추가
challengeType: 0
videoUrl: 'https://scrimba.com/c/cR3bRbCV'
forumTopicId: 301008
dashedName: add-an-accessible-date-picker
---

# --description--

폼에는 종종 `input` 필드가 포함되며, 이를 사용해 여러 가지 폼 컨트롤을 만들 수 있습니다. 이 요소의 `type` 속성은 어떤 종류의 `input` 요소가 생성될지 나타냅니다.

이전 도전 과제에서 `text`와 `submit` 입력 유형을 본 적 있을 것이며, HTML5는 `date` 필드를 지정하는 옵션을 도입했습니다. 브라우저 지원에 따라, `input` 필드가 포커스될 때 날짜 선택기가 나타나서 모든 사용자가 폼을 더 쉽게 작성할 수 있게 합니다.

구형 브라우저에서는 타입이 기본적으로 `text`로 설정되므로, `label` 또는 `placeholder` 텍스트에 예상 날짜 형식을 표시하는 것이 도움이 됩니다.

예시는 다음과 같습니다:

```html
<label for="input1">Enter a date:</label>
<input type="date" id="input1" name="input1">
```

# --instructions--

Camper Cat가 Mortal Kombat 토너먼트를 준비 중이며, 참가자들에게 가장 적합한 날짜를 묻고 싶어 합니다. `input` 태그를 추가하고 `type` 속성에 `date`, `id` 속성에 `pickdate`, `name` 속성에 `date` 값을 넣으세요.

# --hints--

코드에는 날짜 선택 필드를 위한 `input` 태그가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('input'), 2);
```

`input` 태그에는 값이 `type`인 `date` 속성이 있어야 합니다.

```js
assert.equal(document.querySelector('input')?.getAttribute('type'), 'date');
```

`input` 태그에는 값이 `id`인 `pickdate` 속성이 있어야 합니다.

```js
assert.equal(document.querySelector('input')?.getAttribute('id'),'pickdate');
```

`input` 태그에는 값이 `name`인 `date` 속성이 있어야 합니다.

```js
assert.equal(document.querySelector('input')?.getAttribute('name'), 'date');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>

        <!-- Only change code below this line -->



        <!-- Only change code above this line -->

        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>
        <input type="date" id="pickdate" name="date">
        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
