---
id: 587d7787367417b2b2512aa1
title: 헤더 랜드마크로 화면 읽기 프로그램 내비게이션 쉽게 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB76vtv'
forumTopicId: 301023
dashedName: make-screen-reader-navigation-easier-with-the-header-landmark
---

# --description--

다음 HTML5 요소는 의미론적 의미를 추가하고 접근성을 향상시키는 `header` 태그입니다. 이 태그는 상위 태그의 소개 정보나 내비게이션 링크를 감싸는 데 사용되며, 여러 페이지 상단에 반복되는 콘텐츠 주위에 잘 어울립니다.

`header`은 `main`에서 본 임베드된 랜드마크 기능을 공유하여 보조 기술이 해당 콘텐츠로 빠르게 내비게이션할 수 있게 합니다.

**참고:** `header`는 HTML 문서의 `body` 태그 내에서 사용하도록 설계되었습니다. 페이지의 제목, 메타 정보 등을 포함하는 `head` 요소와는 다릅니다.

# --instructions--

Camper Cat이 닌자 훈련에 관한 훌륭한 글을 작성하고 있으며, 이를 위한 페이지를 사이트에 추가하려고 합니다. 현재 `div`를 포함하는 상단의 `h1`를 `header` 태그로 변경하세요.

# --hints--

코드에는 `header` 태그가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('header'),1);
```

`header` 태그는 `h1`를 감싸야 합니다.

```js
const header = document.querySelector('header');
const children = header?.querySelectorAll(`:scope ${'h1'}`);
assert.lengthOf(children , 1);
```

코드에 `div` 태그가 없어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

`header` 요소에는 닫는 태그가 있어야 합니다.

```js
assert.isTrue(code.match(/<\/header>/g)?.length === code.match(/<header>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>

  <div>
    <h1>Training with Camper Cat</h1>
  </div>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>

  <header>
    <h1>Training with Camper Cat</h1>
  </header>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
