---
id: 587d778a367417b2b2512aa6
title: label 요소로 폼 필드 접근성 향상하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMMAN'
forumTopicId: 301016
dashedName: improve-form-field-accessibility-with-the-label-element
---

# --description--

시맨틱 HTML 마크업으로 접근성을 향상하려면 적절한 태그 이름과 속성을 모두 사용하는 것이 중요합니다. 다음 몇 가지 도전 과제에서는 폼에서 속성을 사용하는 중요한 상황들을 다룹니다.

`label` 태그는 특정 폼 컨트롤 항목의 텍스트, 보통 선택지의 이름이나 레이블을 감쌉니다. 이는 항목에 의미를 부여하고 폼을 더 읽기 쉽게 만듭니다. `for` 속성은 `label` 태그에서 해당 `label`를 폼 컨트롤과 명시적으로 연결하며 화면 읽기 프로그램에서 사용됩니다.

기본 HTML 섹션의 수업에서 라디오 버튼과 그 레이블에 대해 배웠습니다. 그 수업에서는 라디오 버튼 입력 요소를 `label` 요소 안에 레이블 텍스트와 함께 감싸 텍스트를 클릭 가능하게 만들었습니다. 이 수업에서는 `for` 속성을 사용해 같은 효과를 내는 방법을 설명합니다.

`for` 속성의 값은 폼 컨트롤의 `id` 속성 값과 같아야 합니다. 예시는 다음과 같습니다:

```html
<form>
  <label for="name">Name:</label>
  <input type="text" id="name" name="name">
</form>
```

# --instructions--

Camper Cat은 자신의 생각이 담긴 블로그 게시물에 많은 관심이 있을 것으로 기대하며 이메일 가입 폼을 포함하려고 합니다. 이메일 `for`에 `label` 속성을 추가하고, 그 값이 `id` 필드의 `input`와 일치하도록 하세요.

# --hints--

코드에 비어 있지 않은 `for` 속성이 `label` 태그에 있어야 합니다.

```js
assert.isNotEmpty(document.querySelector('label')?.getAttribute('for'));
```

`for` 속성 값은 이메일 `id`의 `input` 값과 일치해야 합니다.

```js
assert.equal(document.querySelector('label')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label>Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
