---
id: 587d778a367417b2b2512aa5
title: figure 요소로 차트 접근성 향상하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMqtE'
forumTopicId: 301015
dashedName: improve-chart-accessibility-with-the-figure-element
---

# --description--

HTML5는 `figure` 요소와 관련된 `figcaption`을 도입했습니다. 이 두 가지를 함께 사용하면 이미지, 도표, 차트 같은 시각적 표현과 그에 대한 캡션을 함께 감쌀 수 있습니다. 이 요소들을 함께 감싸면 관련 내용을 의미론적으로 묶어주고 `figure`를 설명하는 텍스트 대체물을 제공하여 접근성을 두 배로 높입니다.

차트 같은 데이터 시각화의 경우, 캡션은 시각 장애가 있는 사용자를 위해 추세나 결론을 간략히 적는 데 사용할 수 있습니다. 또 다른 도전 과제에서는 화면 읽기 프로그램 사용자를 위해 CSS로 차트 데이터의 표 버전을 화면 밖으로 이동하는 방법을 다룹니다.

예를 들어, `figcaption`는 `figure` 태그 안에 들어가며 다른 요소와 함께 사용할 수 있습니다:

```html
<figure>
  <img src="roundhouseDestruction.jpeg" alt="Photo of Camper Cat executing a roundhouse kick">
  <br>
  <figcaption>
    Master Camper Cat demonstrates proper form of a roundhouse kick.
  </figcaption>
</figure>
```

# --instructions--

Camper Cat가 은신, 전투, 무기 훈련에 주당 투자하는 시간을 보여주는 누적 막대 그래프를 만들고 있습니다. 그가 사용한 `div` 태그를 `figure` 태그로, 캡션을 감싸는 `p` 태그를 `figcaption` 태그로 바꿔 페이지 구조를 더 잘 잡도록 도와주세요.

# --hints--

코드에 `figure` 태그가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('figure') , 1);
```

코드에 `figcaption` 태그가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('figcaption') , 1);
```

코드에 `div` 태그가 없어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

코드에 `p` 태그가 없어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('p') , 0);
```

`figcaption`는 `figure` 태그의 자식이어야 합니다.

```js
const figure = document.querySelector('figure');
const children = figure?.querySelectorAll(`:scope ${'figcaption'}`);
assert.lengthOf(children, 1);
```

`figure` 요소는 닫는 태그가 있어야 합니다.

```js
assert.isTrue(
    code.match(/<\/figure>/g)?.length === code.match(/<figure>/g)?.length
);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>

      <!-- Only change code below this line -->
      <div>
        <!-- Stacked bar chart will go here -->
        <br>
        <p>Breakdown per week of time to spend training in stealth, combat, and weapons.</p>
      </div>
      <!-- Only change code above this line -->

    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>
      <figure>
        <!-- Stacked bar chart will go here -->
        <br>
        <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
      </figure>
    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
