---
id: 587d7790367417b2b2512ab0
title: 키보드 포커스를 요소에 추가하기 위해 tabindex 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMDHW'
forumTopicId: 301027
dashedName: use-tabindex-to-add-keyboard-focus-to-an-element
---

# --description--

HTML의 `tabindex` 속성은 요소의 키보드 포커스와 관련해 세 가지 뚜렷한 기능을 가집니다. 태그에 이 속성이 있으면 해당 요소가 포커스될 수 있음을 나타냅니다. 값(양수, 음수 또는 0인 정수)은 동작 방식을 결정합니다.

링크나 폼 컨트롤 같은 특정 요소는 사용자가 페이지를 탭할 때 자동으로 키보드 포커스를 받습니다. 이 순서는 HTML 소스 마크업에 나타난 요소 순서와 같습니다. 이와 같은 기능을 `div`, `span`, `p` 같은 다른 요소에도 `tabindex="0"` 속성을 추가해 줄 수 있습니다. 예시는 다음과 같습니다:

```html
<div tabindex="0">I need keyboard focus!</div>
```

**참고:** 음수인 `tabindex` 값(보통 -1)은 요소가 포커스 가능하지만 키보드로는 접근할 수 없음을 나타냅니다. 이 방법은 일반적으로 팝업 창에 사용되는 `div`가 활성화될 때처럼 프로그래밍 방식으로 포커스를 이동할 때 사용하며, 이 도전 과제의 범위를 벗어납니다.

# --instructions--

Camper Cat이 사용자 정보를 수집하기 위해 새 설문조사를 만들었습니다. 입력 필드는 자동으로 키보드 포커스를 받는다는 것을 알지만, 탭으로 항목을 이동할 때 키보드 사용자가 안내 문구에서 잠시 멈추길 원합니다. `tabindex` 태그에 `p` 속성을 추가하고 값을 `0`로 설정하세요. 보너스 - `tabindex`을 사용하면 `:focus` 태그에서 CSS 가상 클래스 `p`도 작동합니다.

# --hints--

코드는 폼 안내 문구를 담고 있는 `tabindex` 태그에 `p` 속성을 추가해야 합니다.

```js
assert.isNotNull(document.querySelector('p')?.getAttribute('tabindex'));
```

코드는 `tabindex` 태그에 `p` 속성을 0으로 설정해야 합니다.

```js
assert.equal(document.querySelector('p')?.getAttribute('tabindex'), '0');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p>Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p tabindex="0">Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
