---
id: 587d7790367417b2b2512aaf
title: HTML 접근 키로 링크를 탐색 가능하게 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQvmaTp'
forumTopicId: 301021
dashedName: make-links-navigable-with-html-access-keys
---

# --description--

HTML은 요소를 활성화하거나 포커스를 맞추기 위한 단축키를 지정하는 `accesskey` 속성을 제공합니다. `accesskey` 속성을 추가하면 키보드만 사용하는 사용자가 더 효율적으로 탐색할 수 있습니다.

HTML5에서는 이 속성을 모든 요소에 사용할 수 있지만, 특히 링크, 버튼, 폼 컨트롤 같은 상호작용 요소에 사용할 때 유용합니다.

예시는 다음과 같습니다:

```html
<button accesskey="b">Important Button</button>
```

# --instructions--

Camper Cat은 두 블로그 기사 제목 주변의 링크에 키보드 단축키를 추가해 사이트 사용자가 전체 내용을 빠르게 탐색할 수 있게 하려고 합니다. 두 링크 모두에 `accesskey` 속성을 추가하고 첫 번째 링크에는 `g` (Garfield용), 두 번째 링크에는 `c` (Chuck Norris용)로 설정하세요.

# --hints--

코드는 `accesskey` 태그에 `a` 속성을 `id` 값 `first`로 추가해야 합니다.

```js
assert.isTrue(document.querySelector('#first')?.hasAttribute('accesskey'));
```

코드는 `accesskey` 태그에 `a` 속성을 `id` 값 `second`로 추가해야 합니다.

```js
assert.isTrue(document.querySelector('#second')?.hasAttribute('accesskey'));
```

코드는 `accesskey` 태그의 `a` 속성을 `id` 값 `first`에 `g`로 설정해야 합니다. 대소문자를 구분한다는 점에 유의하세요.

```js
assert.equal(document.querySelector('#first')?.getAttribute('accesskey'), 'g');
```

코드는 `accesskey` 태그의 `a` 속성을 `id` 값 `second`에 `c`로 설정해야 합니다. 대소문자를 구분한다는 점에 유의하세요.

```js
assert.equal(document.querySelector('#second')?.getAttribute('accesskey'), 'c');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" accesskey="g" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" accesskey="c" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
