---
id: 587d778c367417b2b2512aa9
title: HTML5 datetime 속성으로 시간 표준화하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMgtz'
forumTopicId: 301025
dashedName: standardize-times-with-the-html5-datetime-attribute
---

# --description--

날짜 주제를 이어서, HTML5는 시간을 표준화하기 위해 `time` 요소와 `datetime` 속성을 도입했습니다. `time` 요소는 페이지에서 날짜나 시간을 감쌀 수 있는 인라인 요소입니다. `datetime` 속성은 그 날짜의 유효한 형식을 담고 있습니다. 이 값은 보조 기기가 접근하는 값입니다. 텍스트에 비공식적이거나 구어체로 쓰여 있어도 표준화된 시간 버전을 명시해 혼란을 줄이는 데 도움이 됩니다.

예를 들어 다음과 같습니다:

```html
<p>Master Camper Cat officiated the cage match between Goro and Scorpion <time datetime="2013-02-13">last Wednesday</time>, which ended in a draw.</p>
```

# --instructions--

Camper Cat의 Mortal Kombat 설문조사 결과가 나왔습니다! 텍스트 `time` 주위에 `Thursday, September 15<sup>th</sup>` 태그를 감싸고, `datetime` 속성을 추가하여 `2016-09-15` 값으로 설정하세요.

# --before-all--

```js
const getTimeElement = () => {
  const pElement = [...document.querySelectorAll("article > p")]
  .filter(x => x?.textContent?.includes("Thank you to everyone for responding to Master Camper Cat's survey."));
  
  return pElement[0] ? pElement[0].querySelector("time") : null;
};

const getDatetimeAttr = () => {
  const timeElement = getTimeElement();
  return timeElement?.getAttribute("datetime");
};
```

# --hints--

코드에 텍스트 `p`를 포함하는 `Thank you to everyone for responding to Master Camper Cat's survey.` 요소와 `time` 요소가 있어야 합니다.

```js
assert.exists(getTimeElement());
```

추가한 `time` 태그는 텍스트 `Thursday, September 15<sup>th</sup>`를 감싸야 합니다.

```js
assert.strictEqual(getTimeElement()?.innerHTML?.trim(), 'Thursday, September 15<sup>th</sup>');
```

추가한 `time` 태그에는 비어 있지 않은 `datetime` 속성이 있어야 합니다.

```js
assert(getDatetimeAttr()?.length != 0);
```

추가한 `datetime` 속성은 `2016-09-15` 값으로 설정해야 합니다.

```js
assert.equal(getDatetimeAttr() , '2016-09-15');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <article>
    <h2>Mortal Kombat Tournament Survey Results</h2>

    <!-- Only change code below this line -->

    <p>Thank you to everyone for responding to Master Camper Cat's survey. The best day to host the vaunted Mortal Kombat tournament is Thursday, September 15<sup>th</sup>. May the best ninja win!</p>

    <!-- Only change code above this line -->

    <section>
      <h3>Comments:</h3>
      <article>
        <p>Posted by: Sub-Zero on <time datetime="2016-08-13T20:01Z">August 13<sup>th</sup></time></p>
        <p>Johnny Cage better be there, I'll finish him!</p>
      </article>
      <article>
        <p>Posted by: Doge on <time datetime="2016-08-15T08:12Z">August 15<sup>th</sup></time></p>
        <p>Wow, much combat, so mortal.</p>
      </article>
      <article>
        <p>Posted by: The Grim Reaper on <time datetime="2016-08-16T00:00Z">August 16<sup>th</sup></time></p>
        <p>Looks like I'll be busy that day.</p>
      </article>
    </section>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <article>
    <h2>Mortal Kombat Tournament Survey Results</h2>

    <p>Thank you to everyone for responding to Master Camper Cat's survey. The best day to host the vaunted Mortal Kombat tournament is <time datetime="2016-09-15">Thursday, September 15<sup>th</sup></time>. May the best ninja win!</p>

    <section>
      <h3>Comments:</h3>
      <article>
        <p>Posted by: Sub-Zero on <time datetime="2016-08-13T20:01Z">August 13<sup>th</sup></time></p>
        <p>Johnny Cage better be there, I'll finish him!</p>
      </article>
      <article>
        <p>Posted by: Doge on <time datetime="2016-08-15T08:12Z">August 15<sup>th</sup></time></p>
        <p>Wow, much combat, so mortal.</p>
      </article>
      <article>
        <p>Posted by: The Grim Reaper on <time datetime="2016-08-16T00:00Z">August 16<sup>th</sup></time></p>
        <p>Looks like I'll be busy that day.</p>
      </article>
    </section>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
