---
id: 587d774e367417b2b2512aa0
title: article 요소로 콘텐츠 감싸기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp79S3'
forumTopicId: 301029
dashedName: wrap-content-in-the-article-element
---

# --description--

`article`는 마크업에 의미론적 의미를 더하는 새로운 HTML5 요소 중 하나입니다. `article`는 구획 요소로 독립적이고 자립적인 콘텐츠를 감싸는 데 사용됩니다. 이 태그는 블로그 항목, 포럼 게시물 또는 뉴스 기사에 잘 어울립니다.

콘텐츠가 독립적으로 존재할 수 있는지 판단하는 것은 보통 주관적인 판단이지만, 몇 가지 간단한 테스트를 사용할 수 있습니다. 주변 문맥을 모두 제거해도 그 콘텐츠가 여전히 의미가 있을지 스스로에게 물어보세요. 마찬가지로, 텍스트의 경우 RSS 피드에 있어도 내용이 유지될지 생각해보세요.

보조 기술을 사용하는 분들은 조직적이고 의미론적으로 의미 있는 마크업에 의존해 작업을 더 잘 이해하므로 기억하세요.

**참고:** `section` 요소도 HTML5에서 새로 추가되었으며 `article`와는 약간 다른 의미론적 의미를 가집니다. `article`는 독립적인 콘텐츠용이고, `section`는 주제별로 관련된 콘텐츠를 그룹화하는 데 사용됩니다. 필요에 따라 서로 안에 사용할 수 있습니다. 예를 들어, 책이 `article`라면 각 장은 `section`입니다. 콘텐츠 그룹 간에 관계가 없으면 `div`를 사용하세요.

`<div>` - 콘텐츠를 그룹화함  
`<section>` - 관련 콘텐츠를 그룹화함  
`<article>` - 독립적이고 자립적인 콘텐츠를 그룹화함  

# --instructions--

Camper Cat은 블로그 페이지의 게시물을 감싸기 위해 `article` 태그를 사용했지만, 맨 위 게시물에는 이를 사용하지 않았습니다. `div` 태그를 `article` 태그로 변경하세요.

# --hints--

코드에는 `article` 태그가 세 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('article') ,3);
```

코드에는 `div` 태그가 없어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <div>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </div>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```
