---
id: 587d774d367417b2b2512a9e
title: 콘텐츠의 계층적 관계를 보여주기 위해 제목 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cqVEktm'
forumTopicId: 301026
dashedName: use-headings-to-show-hierarchical-relationships-of-content
---

# --description--

제목(`h1`부터 `h6` 요소까지)은 콘텐츠에 구조와 레이블을 제공하는 핵심 태그입니다. 화면 읽기 프로그램은 페이지의 제목만 읽도록 설정할 수 있어 사용자가 요약을 얻을 수 있습니다. 따라서 마크업의 제목 태그는 의미론적 의미를 가지고 서로 관련되어야 하며, 단순히 크기 값 때문에 선택되어서는 안 됩니다.

*의미론적 의미*란 콘텐츠를 감싸는 태그가 그 안에 담긴 정보 유형을 나타낸다는 뜻입니다.

소개, 본문, 결론이 있는 논문을 작성한다고 할 때, 결론을 본문의 하위 섹션으로 두는 것은 별로 의미가 없습니다. 결론은 독립된 섹션이어야 합니다. 마찬가지로 웹페이지의 제목 태그도 순서대로 배치되어 콘텐츠의 계층적 관계를 나타내야 합니다.

동일하거나 더 높은 등급의 제목은 새로운 암묵적 섹션을 시작하고, 더 낮은 등급의 제목은 이전 제목의 하위 섹션을 시작합니다.

예를 들어, `h2` 요소 다음에 `h4` 요소로 표시된 여러 하위 섹션이 있는 페이지는 화면 읽기 프로그램 사용자에게 혼란을 줄 수 있습니다. 6가지 선택지 중에서 태그를 사용하면 브라우저에서 더 보기 좋지만, CSS로 상대 크기를 조정할 수 있으므로 태그를 사용하는 것이 유혹적일 수 있습니다.

마지막으로, 각 페이지에는 항상 하나(그리고 단 하나)의 `h1` 요소가 있어야 하며, 이는 콘텐츠의 주요 주제를 나타냅니다. 이 제목과 다른 제목들은 검색 엔진이 페이지 주제를 이해하는 데 일부 사용됩니다.

# --instructions--

Camper Cat은 닌자가 되는 것에 관한 페이지를 사이트에 만들고 싶어 합니다. 제목을 수정하여 마크업이 콘텐츠에 의미론적 의미를 부여하고 섹션 간 올바른 부모-자식 관계를 보여주도록 도와주세요. 모든 `h5` 태그를 `h2` 태그의 하위 섹션임을 나타내도록 적절한 제목 수준으로 변경하세요. 이 목적을 위해 `h3` 태그를 사용하세요.

# --hints--

코드에는 6개의 `h3` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h3') , 6);
```

코드에는 6개의 `h3` 닫는 태그가 있어야 합니다.

```js
assert.lengthOf((code.match(/\/h3/g) || []) ,6);
```

코드에는 `h5` 요소가 없어야 합니다.

```js
assert.isEmpty(document.querySelectorAll('h5'));
```

코드에는 `h5` 닫는 태그가 없어야 합니다.

```js
assert.notMatch(code, /\/h5/);
```

# --seed--

## --seed-contents--

```html
<h1>How to Become a Ninja</h1>
<main>
  <h2>Learn the Art of Moving Stealthily</h2>
  <h5>How to Hide in Plain Sight</h5>
  <h5>How to Climb a Wall</h5>

  <h2>Learn the Art of Battle</h2>
  <h5>How to Strengthen your Body</h5>
  <h5>How to Fight like a Ninja</h5>

  <h2>Learn the Art of Living with Honor</h2>
  <h5>How to Breathe Properly</h5>
  <h5>How to Simplify your Life</h5>
</main>
```

# --solutions--

```html
<h1>How to Become a Ninja</h1>
<main>
  <h2>Learn the Art of Moving Stealthily</h2>
  <h3>How to Hide in Plain Sight</h3>
  <h3>How to Climb a Wall</h3>

  <h2>Learn the Art of Battle</h2>
  <h3>How to Strengthen your Body</h3>
  <h3>How to Fight like a Ninja</h3>

  <h2>Learn the Art of Living with Honor</h2>
  <h3>How to Breathe Properly</h3>
  <h3>How to Simplify your Life</h3>
</main>
```
