---
id: 587d7790367417b2b2512ab1
title: 여러 요소의 키보드 포커스 순서를 지정하기 위해 tabindex 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzRRcb'
forumTopicId: 301028
dashedName: use-tabindex-to-specify-the-order-of-keyboard-focus-for-several-elements
---

# --description--

`tabindex` 속성은 요소들의 정확한 탭 순서도 지정합니다. 이 속성 값이 1 이상의 양수로 설정될 때 이 기능이 작동합니다.

`tabindex="1"`을 설정하면 키보드 포커스가 해당 요소로 먼저 이동합니다. 그다음 지정된 `tabindex` 값들(2, 3 등)의 수열을 따라 순환한 후 기본값과 `tabindex="0"` 항목으로 이동합니다.

이렇게 탭 순서를 설정하면 기본 순서(HTML 소스 순서)를 재정의한다는 점을 기억하세요. 페이지 상단에서 탐색을 시작할 것으로 예상하는 사용자에게 혼란을 줄 수 있습니다. 이 기법은 특정 상황에서 필요할 수 있지만, 접근성 측면에서는 적용 전에 신중히 고려하세요.

예시는 다음과 같습니다:

```html
<div tabindex="1">I get keyboard focus, and I get it first!</div>
```

```html
<div tabindex="2">I get keyboard focus, and I get it second!</div>
```

# --instructions--

Camper Cat은 Inspirational Quotes 페이지에 검색 필드를 가지고 있으며, CSS로 오른쪽 상단에 배치할 계획입니다. 검색 `input`과 제출 `input` 폼 컨트롤이 탭 순서에서 처음 두 항목이 되길 원합니다. `tabindex` `1`에 `search` 속성을 `input` 값으로 추가하고, `tabindex` `2`에 `submit` 속성을 `input` 값으로 추가하세요.

또한 일부 브라우저는 요소를 클릭할 때 탭 순서 중간에서 시작할 수 있습니다. 페이지에 항상 탭 순서의 시작점에서 시작하도록 보장하는 요소가 추가되어 있습니다.

# --hints--

코드에 `tabindex` 속성을 `search` `input` 태그에 추가해야 합니다.

```js
assert.isNotNull(document.querySelector('#search')?.getAttribute('tabindex'));
```

코드에 `tabindex` 속성을 `submit` `input` 태그에 추가해야 합니다.

```js
assert.isNotNull(document.querySelector('#submit')?.getAttribute('tabindex'));
```

코드에 `tabindex` 속성을 `search` `input` 태그에 1 값으로 설정해야 합니다.

```js
assert.equal(document.querySelector('#search')?.getAttribute('tabindex'), '1');
```

코드에 `tabindex` 속성을 `submit` `input` 태그에 2 값으로 설정해야 합니다.

```js
assert.equal(document.querySelector('#submit')?.getAttribute('tabindex'), '2');
```

# --seed--

## --seed-contents--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input type="search" name="search" id="search">
    <input type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```

# --solutions--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input tabindex="1" type="search" name="search" id="search">
    <input tabindex="2" type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```
