---
id: 587d778f367417b2b2512aad
title: >-
  정보를 전달하는 색상을 신중하게 선택하여 색맹 문제를 피하세요
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437as3'
forumTopicId: 301011
dashedName: >-
  avoid-colorblindness-issues-by-carefully-choosing-colors-that-convey-information
---

# --description--

색맹에는 여러 가지 형태가 있습니다. 이는 특정 파장의 빛에 대한 민감도가 감소하는 것부터 색을 전혀 볼 수 없는 상태까지 다양합니다. 가장 흔한 형태는 녹색을 감지하는 민감도가 감소하는 경우입니다.

예를 들어, 두 가지 비슷한 녹색이 콘텐츠의 전경색과 배경색이라면, 색맹인 사용자는 이를 구분하지 못할 수 있습니다. 가까운 색상은 색상환에서 이웃한 색으로 생각할 수 있으며, 중요한 정보를 전달할 때 이러한 조합은 피해야 합니다.

**참고:** 일부 온라인 색상 선택 도구에는 다양한 색맹 유형에 따른 색상 시뮬레이션이 포함되어 있습니다. 이는 온라인 명암 대비율 계산기와 함께 매우 유용한 자료입니다.

# --instructions--

Camper Cat이 중요한 버튼에 대해 다양한 스타일을 테스트하고 있지만, 노란색(`#FFFF33`) `background-color`과 녹색(`#33FF33`) 텍스트 `color`는 색상환에서 이웃한 색조이며 일부 색맹인 사용자에게는 사실상 구분이 불가능합니다. (밝기가 비슷해 명암 대비율 검사도 통과하지 못합니다.) 텍스트 `color`를 어두운 파란색(`#003366`)으로 변경하여 두 문제를 모두 해결하세요.

# --hints--

코드는 `color`의 텍스트 `button`를 어두운 파란색으로 변경해야 합니다.

```js
const button = document.querySelector('button');
const buttonColor = window.getComputedStyle(button).color; 
assert.equal(buttonColor, 'rgb(0, 51, 102)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  button {
    color: #33FF33;
    background-color: #FFFF33;
    font-size: 14px;
    padding: 10px;
  }
  </style>
</head>
<body>
  <header>
    <h1>Danger!</h1>
  </header>
  <button>Delete Internet</button>
</body>
```

# --solutions--

```html
<head>
  <style>
    button {
      color: #003366;
      background-color: #FFFF33;
      font-size: 14px;
      padding: 10px;
    }
  </style>
</head>
<body>
  <header>
    <h1>Danger!</h1>
  </header>
  <button>Delete Internet</button>
</body>
```
