---
id: 587d7789367417b2b2512aa4
title: audio 요소로 오디오 콘텐츠 접근성 향상하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVkcZ'
forumTopicId: 301014
dashedName: improve-accessibility-of-audio-content-with-the-audio-element
---

# --description--

HTML5의 `audio` 요소는 소리나 오디오 스트림 콘텐츠를 마크업에 감쌀 때 의미론적 의미를 제공합니다. 청각 장애가 있거나 난청인 사용자가 접근할 수 있도록 오디오 콘텐츠에는 텍스트 대체물이 필요합니다. 이는 페이지 내 인근 텍스트나 대본 링크로 제공할 수 있습니다.

`audio` 태그는 `controls` 속성을 지원합니다. 이 속성은 브라우저 기본 재생, 일시정지 및 기타 컨트롤을 표시하며 키보드 기능도 지원합니다. 불리언 속성이므로 값이 필요 없고, 태그에 있으면 설정이 켜집니다.

예시는 다음과 같습니다:

```html
<audio id="meowClip" controls>
  <source src="audio/meow.mp3" type="audio/mpeg">
  <source src="audio/meow.ogg" type="audio/ogg">
</audio>
```

**참고:** 멀티미디어 콘텐츠는 보통 시각적 요소와 청각적 요소를 모두 포함합니다. 시각 또는 청각 장애가 있는 사용자가 접근할 수 있도록 동기화된 캡션과 대본이 필요합니다. 일반적으로 웹 개발자가 캡션이나 대본을 직접 만들지는 않지만, 포함해야 한다는 점은 알아야 합니다.

# --instructions--

Camper Cat과 잠시 쉬고, 접근성 챔피언이자 화면 읽기 프로그램 사용자 Zersiax(@zersiax)를 만나보세요. 그의 화면 읽기 프로그램 음성 클립을 듣기 위해 `audio` 요소 뒤에 `p` 요소를 추가하세요. `controls` 속성을 포함하세요. 그런 다음 `source` 태그 안에 `audio` 요소를 넣고 `src` 속성은 `https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3`로, `type` 속성은 `"audio/mpeg"`로 설정하세요.

**참고:** 오디오 클립이 빠르게 들리고 이해하기 어려울 수 있지만, 화면 읽기 프로그램 사용자에게는 정상 속도입니다.

# --hints--

코드에 `audio` 태그가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('audio'),1);
```

`audio` 요소에는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<audio.*>[\s\S]*<\/audio>/g);
assert.lengthOf(code.match(/<\/audio>/g),1);
```

`audio` 태그에는 `controls` 속성이 있어야 합니다.

```js
assert.exists(document.querySelector('audio')?.getAttribute('controls'));
```

코드에 `source` 태그가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('source'), 1);
```

`source` 태그는 `audio` 태그 안에 있어야 합니다.

```js
const audio = document.querySelector('audio');
const children = audio.querySelectorAll(`:scope ${'source'}`);
assert.lengthOf(children,1);
```

`src` 태그의 `source` 속성 값은 지침에 있는 링크와 정확히 일치해야 합니다.

```js
assert.equal(
  document.querySelector('source')?.getAttribute('src'),
  'https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3'
);
```

코드에 `type` 태그에 `source` 속성이 audio/mpeg 값으로 포함되어야 합니다.

```js
assert.equal(document.querySelector('source')?.getAttribute('type'), 'audio/mpeg');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>



  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>
    <audio controls>
      <source src="https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3" type="audio/mpeg"/>
    </audio>
  </main>
</body>
```
