---
id: 587d774c367417b2b2512a9c
title: 시각 장애 접근성을 위한 이미지 대체 텍스트 추가
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

다른 도전 과제에서 `alt` 속성이 `img` 태그에 있는 것을 본 적이 있을 것입니다. `alt` 텍스트는 이미지의 내용을 설명하며 텍스트 대체 수단을 제공합니다. `alt` 속성은 이미지가 로드되지 않거나 사용자가 볼 수 없을 때 도움이 됩니다. 검색 엔진도 이미지가 무엇을 포함하는지 이해하여 검색 결과에 포함시키기 위해 이 속성을 사용합니다. 예를 들어 다음과 같습니다:

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

시각 장애가 있는 사람들은 화면 읽기 프로그램을 사용해 웹 콘텐츠를 음성 인터페이스로 변환합니다. 시각적으로만 제공되는 정보는 얻을 수 없습니다. 이미지의 경우 화면 읽기 프로그램이 `alt` 속성에 접근해 그 내용을 읽어 주요 정보를 전달할 수 있습니다.

좋은 `alt` 텍스트는 독자에게 이미지에 대한 간단한 설명을 제공합니다. 항상 이미지에 `alt` 속성을 포함해야 합니다. HTML5 명세에 따르면 이제 이것은 필수로 간주됩니다.

# --instructions--

캠퍼 캣은 코딩 닌자이자 실제 닌자로, 자신의 지식을 공유할 웹사이트를 만들고 있습니다. 그가 사용하려는 프로필 사진은 그의 기술을 보여주며 모든 사이트 방문자가 감상할 수 있어야 합니다. `alt` 태그에 캠퍼 캣이 가라테를 하고 있음을 설명하는 `img` 속성을 추가하세요. (이미지 `src`는 실제 파일에 링크되어 있지 않으므로 화면에 `alt` 텍스트가 표시될 것입니다.)

# --hints--

`img` 태그에는 `alt` 속성이 있어야 하며 비어 있으면 안 됩니다.

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">
```
