---
id: 587d778e367417b2b2512aab
title: 명암 대비가 높은 텍스트로 가독성 향상
challengeType: 0
videoUrl: 'https://scrimba.com/c/cKb3nCq'
forumTopicId: 301017
dashedName: improve-readability-with-high-contrast-text
---

# --description--

전경과 배경 색상 간 명암 대비가 낮으면 텍스트를 읽기 어려울 수 있습니다. 충분한 명암 대비는 콘텐츠의 가독성을 높이지만, '충분한' 명암 대비가 정확히 무엇을 의미하는지 궁금할 수 있습니다.

웹 콘텐츠 접근성 지침(WCAG)은 일반 텍스트에 대해 최소 4.5:1 명암 대비율을 권장합니다. 이 비율은 두 색상의 상대 밝기 값을 비교하여 계산합니다. 동일한 색상일 경우 1:1로 명암 대비가 없으며, 흰색과 검은색의 경우 21:1로 가장 큰 명암 대비를 나타냅니다. 온라인에는 이 비율을 계산해 주는 다양한 명암 대비 검사 도구가 있습니다.

# --instructions--

Camper Cat가 최근 블로그 글에서 흰색 배경에 연한 회색 텍스트를 선택했는데, 명암 대비율이 1.5:1로 읽기 어렵습니다. 텍스트의 `color`를 현재 회색(`#D3D3D3`)에서 더 어두운 회색(`#636363`)으로 변경하여 명암 대비율을 6:1로 개선하세요.

# --hints--

코드에서 `color`의 텍스트 `body`를 더 어두운 회색으로 변경해야 합니다.

```js
const body = document.querySelector('body');
const bodyColor = window.getComputedStyle(body).color; 
assert(bodyColor == 'rgb(99, 99, 99)');
```

코드에서 `background-color`의 `body`는 변경하지 않아야 합니다.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body).backgroundColor; 
assert.equal(backgroundColor , 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: #D3D3D3;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: #636363;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
