---
id: 587d778f367417b2b2512aae
title: 설명적인 링크 텍스트로 링크에 의미 부여하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437DcV'
forumTopicId: 301013
dashedName: give-links-meaning-by-using-descriptive-link-text
---

# --description--

화면 읽기 프로그램 사용자는 기기가 읽을 콘텐츠 유형을 다양하게 선택할 수 있습니다. 이 옵션에는 랜드마크 요소로 건너뛰기(또는 건너뛰기), 주요 콘텐츠로 점프하기, 제목에서 페이지 요약 받기가 포함됩니다. 또 다른 옵션은 페이지에서 사용할 수 있는 링크만 듣는 것입니다.

화면 읽기 프로그램은 앵커 (`a`) 태그 사이에 있는 링크 텍스트를 읽음으로써 이를 수행합니다. "여기를 클릭" 또는 "더 읽기" 같은 링크 목록은 도움이 되지 않습니다. 대신 `a` 태그 안에 간결하지만 설명적인 텍스트를 사용해 이러한 사용자에게 더 많은 의미를 제공하세요.

# --instructions--

Camper Cat이 사용하는 링크 텍스트는 주변 문맥 없이는 그다지 설명적이지 않습니다. 앵커 (`a`) 태그를 `information about batteries` 대신 `Click here` 텍스트를 감싸도록 옮기세요.

# --hints--

코드에서 앵커 `a` 태그를 `Click here` 단어 주위에서 `information about batteries` 단어 주위로 옮겨야 합니다.

```js
assert.match(document.querySelector('a')?.textContent, /^(information about batteries)$/g);
```

`a` 요소는 빈 문자열 `href` 값을 가진 `""` 속성을 가져야 합니다.

```js
assert.isEmpty(document.querySelector('a')?.getAttribute('href'));
```

`a` 요소는 닫는 태그가 있어야 합니다.

```js
assert.isTrue(code.match(/<\/a>/g)?.length === code.match(/<a href=(''|"")>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. <a href="">Click here</a> for information about batteries</p>
  </article>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. Click here for <a href="">information about batteries</a></p>
  </article>
</body>
```
