---
id: 687d8e05a529a4838f1c961c
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

다음으로, `pass` 클래스 안에 있는 `Node` 키워드를 제거한 후, `Node` 클래스에 `__init__` 메서드를 두 개의 매개변수 `self`와 `element`로 추가하세요.

메서드 안에서는 `element`에 `self.element`를 할당하세요. 이렇게 하면 노드를 생성할 때 전달된 값을 저장합니다.

# --hints--

`__init__` 클래스 안에 `Node` 메서드를 정의해야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

`__init__` 메서드는 두 개의 매개변수 `self`과 `element`를 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

`element` 메서드 안에서 `self.element`에 `__init__`를 할당해야 합니다.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
