---
id: 687d8c6f066877768f7ee441
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

`LinkedList` 클래스 안에 `pass` 키워드를 `Node`라는 내부 클래스로 바꾸세요. 이 클래스는 리스트에서 데이터를 저장하는 개별 노드를 나타냅니다.

# --hints--

`Node` 클래스 안에 `LinkedList` 클래스를 만들어야 합니다. `pass` 클래스 안에 `Node` 키워드를 추가하는 것을 잊지 마세요.

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
