---
id: 687d7deea6a2a22fa8816a73
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이전 수업에서 몇 가지 일반적인 자료 구조에 대해 배웠습니다. 이번 워크숍에서는 연결 리스트를 구현할 것입니다.

먼저, 리스트 노드를 담을 구조를 나타내기 위해 `LinkedList`라는 이름의 클래스를 만드세요. 테스트를 통과하려면 클래스 안에 `pass` 키워드를 추가하세요.

# --hints--

`LinkedList`라는 이름의 클래스를 만들어야 합니다. `pass` 키워드를 추가하는 것도 잊지 마세요.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
