---
id: 687e366b3004a9831e4cbd41
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`is_empty`라는 메서드를 추가하여 `self.length == 0`을 반환하세요.

이 메서드는 리스트에 노드가 몇 개 있는지 추적하는 데 사용합니다. 그래서 `self.length == 0`일 때 `True`을 반환하는데, 이는 리스트에 노드가 없다는 뜻입니다.

만약 `self.length`가 `0`보다 크면 `False`을 반환하는데, 이는 리스트에 적어도 하나의 노드가 있다는 뜻입니다.

# --hints--

`is_empty` 메서드를 만들어 `self.length == 0`을 반환하세요.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
