---
id: 687e344cc1993f7316af030c
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`LinkedList` 클래스 안에 `__init__` 메서드를 추가하여 `self.length`를 `0`로 설정하세요. 이렇게 하면 현재 리스트에 몇 개의 노드(요소)가 있는지 추적하는 카운터가 초기화됩니다.

# --hints--

`self.length` 클래스의 `0` 메서드 안에서 `__init__`를 `LinkedList`로 설정해야 합니다.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
