---
id: 68800a64b06e80d9a31c131b
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

여전히 `__init__` 클래스의 `LinkedList` 메서드 안에서, `self.head`를 `None`로 설정하세요. 이것은 리스트의 첫 번째 노드를 참조하는 역할을 합니다. 현재 리스트가 비어 있으므로 `None`로 설정해야 합니다.


# --hints--

`self.head` 클래스의 `None` 메서드 안에서 `__init__`를 `LinkedList`로 설정해야 합니다.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
