---
id: 688003f589b875af60273dd9
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

여전히 `__init__` 메서드 안에서, `None`을 `self.next`에 할당하세요. 이렇게 하면 `next` 포인터가 `None`를 가리키게 되어, 새 노드가 아직 리스트 내 다른 노드를 가리키지 않음을 의미합니다.

나중에 `next`를 업데이트하여 다른 노드를 가리키게 하면 연결 리스트의 연결 고리를 형성할 수 있습니다.

# --hints--

`__init__` 메서드 안에서 `next` 속성을 `None`로 초기화해야 합니다.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
