---
id: 68857e0a49baf4751f737a2a
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

이제 `is_empty` 메서드 다음에 `add` 클래스 안에 `LinkedList`이라는 메서드를 만드세요. 새 메서드는 `self`와 `element` 두 개의 매개변수를 가져야 합니다.

새 메서드 안에서 `node`라는 변수를 만들고 `Node`을 사용해 `self.Node(element)`의 새 인스턴스를 할당하세요. 이렇게 하면 연결 리스트에 추가할 새 노드가 만들어집니다.

# --hints--

`add` 클래스 안에 `LinkedList` 메서드를 만들어야 합니다.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

`add` 메서드는 `self`와 `element` 두 개의 매개변수를 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

`self.Node(element)` 메서드에서 `node`에 `add`을 할당해야 합니다.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
