---
id: 68858a1b4e867acd87f6f7e8
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

목록이 비어 있으면 새 노드가 첫 번째 요소(`head`)가 됩니다. `add` 메서드 안에서 `if` 문을 만들어 `self.is_empty()`를 사용해 목록이 비어 있는지 확인하고 `node`를 `self.head`에 할당하세요.

# --hints--

목록이 비어 있는지 확인하려면 `is_empty()` 메서드를 사용하세요.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_conditions()[0]
  assert _cond.is_equivalent("self.is_empty()")
`)})
```

`node` 문 안에서 `self.head`를 `if`에 할당해야 합니다.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_body().is_equivalent("self.head = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
