---
id: 68868f502659139bed8e5c90
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

새 노드를 리스트에 추가한 후에는 길이를 업데이트해야 합니다. `add` 메서드 끝에서 `self.length`을 `1`만큼 증가시키세요.

# --hints--

`self.length` 메서드 끝에서 `1`을 `add`만큼 증가시켜야 합니다.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").has_stmt("self.length += 1")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
