---
id: 688649c8673f842f49228a5e
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

목록이 비어 있지 않으면, 메서드는 마지막 노드를 찾아 새 노드를 그 `next` 포인터에 연결해야 합니다.

이를 위해 `else` 블록 뒤에 `if` 블록을 만드세요. `else` 블록 안에 `current_node` 변수를 만들고 `self.head`를 할당하세요.

# --hints--

`self.head` 블록 안에서 `current_node`를 `else`에 할당해야 합니다.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
