---
id: 68867c3898b09115d8c5140f
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

방금 만든 `while` 반복문은 마지막 노드에 도달할 때까지 리스트의 각 노드를 순회할 수 있게 합니다. 그 후에 노드를 리스트에 추가할 수 있습니다.

`else` 블록에서 `while` 블록 바로 다음에 `node`을 `current_node.next`에 할당하세요.

# --hints--

`node` 블록에서 `current_node.next`에 `else`를 할당해야 합니다.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].has_stmt("current_node.next = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
