---
id: 6888d05ba074461b64f8eaca
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

다음으로, 연결 리스트에서 주어진 요소를 제거하는 데 사용할 또 다른 메서드를 만들겠습니다.

`remove`라는 이름의 메서드를 두 개의 매개변수 `self`과 `element`로 만드세요. 그 안에 `previous_node`라는 변수를 만들고 `None`를 할당하세요.

# --hints--

`remove`라는 메서드를 만들어야 합니다.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_class("LinkedList").has_function("remove")
    `)
  )
})
```

`remove` 메서드는 두 개의 매개변수 `self`과 `element`를 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("remove").has_args("self, element")`) 
})
```

`None` 메서드에서 `previous_node`에 `remove`를 할당해야 합니다.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").find_body().is_equivalent("previous_node = None")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
