---
id: 6888f9827637322d4f309192
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

`remove` 메서드 안에 `current_node`이라는 변수를 만들고 `self.head` 값을 할당하세요.

이렇게 하면 연결 리스트의 첫 번째 노드부터 순회가 시작됩니다. `current_node`는 지정한 요소를 제거할 노드를 찾는 동안 현재 검사 중인 노드를 추적하는 데 사용됩니다.

# --hints--

`self.head` 메서드 안에서 `current_node`에 `remove`를 할당해야 합니다.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").has_stmt("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
