---
id: 6888fe6c745b6852ed240229
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

다음으로, 지정된 요소를 제거할 노드를 찾기 위해 연결 리스트를 반복하세요.

`while` 루프를 만들어 `current_node`이 `None`가 아니고 `current_node.element`이 `element`와 같지 않을 때를 확인하세요. `while` 루프 안에서 `current_node`을 `previous_node`에 할당하세요.

# --hints--

`while` 메서드 안에 `remove` 루프가 있어야 합니다.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
`)})
```

`while` 루프는 `current_node`이 `None`가 아니고 `current_node.element`이 `element`와 같지 않을 때를 확인해야 합니다.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0].find_conditions()[0]
  assert _cond.is_equivalent("current_node is not None and current_node.element != element")
`)})
```

`current_node` 루프 안에서 `previous_node`을 `while`에 할당해야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
