---
id: 688ac1de985053329b4573d3
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

다음으로, `current_node.next` 반복문 안에서 `current_node`을 `while`에 할당하세요.

이렇게 하면 반복문이 실행될 때마다 `current_node` 포인터가 리스트의 다음 노드로 이동하여, 요소를 찾거나 리스트의 끝에 도달할 때까지 연결 리스트를 순회할 수 있습니다.

# --hints--

`current_node.next` 반복문 안에서 `current_node`를 `while`에 할당해야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node\\n current_node = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--
        while current_node is not None and current_node.element != element:
            previous_node = current_node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
