---
id: 688acb2053ac7978a06345f1
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

다음으로, `elif` 문을 만들어 `previous_node`이 `None`가 아님을 확인하세요. 그 안에서 `current_node.next`을 `previous_node.next`에 할당하세요.

이렇게 하면, 제거할 요소가 발견되고 그것이 헤드 노드에 없을 경우, 이전 노드의 `next` 포인터가 현재 노드(제거할 요소)를 건너뛰도록 업데이트되어 연결 리스트에서 제거됩니다.

# --hints--

`elif`이 `previous_node`가 아님을 확인하는 `None` 문을 만들어야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("previous_node is not None")
`) })
```

`current_node.next` 문 안에서 `previous_node.next`을 `elif`에 할당해야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return\\nelif previous_node is not None:\\n previous_node.next = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        if current_node is None:
            return
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
