---
id: 688ac806194cfc60ba1e1767
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

`while` 반복문 뒤에 `if` 문을 만들어 `current_node`가 `None`인지 확인하세요.

이 경우 연결 리스트를 끝까지 탐색했지만 제거할 요소를 찾지 못한 것이므로 함수 실행을 멈춰야 합니다. 따라서 `return` 본문 안에 `if` 키워드를 추가하세요.

# --hints--

`if`가 `current_node`인지 확인하는 `None` 문을 만들고, 그 안에 `return` 키워드를 추가하세요.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
