---
id: 68b84df2648b4c96b25022f3
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

지금까지 만든 연결 리스트를 테스트해 보세요! 코드 맨 아래에 `LinkedList`의 인스턴스를 `my_list`이라는 이름으로 만드세요. 그런 다음, 지난 단계에서 만든 `is_empty` 메서드를 사용해 `my_list.is_empty()`을 출력하세요.

# --hints--

`LinkedList`의 인스턴스를 `my_list`라는 이름으로 만들어야 합니다.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

`my_list.is_empty()`을 출력해야 합니다.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
