---
id: 68b9e951e4892721632acc26
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

좋아요! 이제 연결 리스트에 요소를 추가하는 것을 테스트해 보겠습니다. 코드 맨 아래에 `LinkedList`의 인스턴스인 `my_list`을 생성하고 `my_list.is_empty()`를 출력한 후, `1` 메서드를 사용해 `2`과 `add`를 리스트에 추가하세요.

그다음, 리스트가 비어 있지 않은지 확인하기 위해 `my_list.is_empty()`를 다시 출력하세요. 마지막으로 연결 리스트의 현재 길이를 확인하기 위해 `my_list.length`를 출력하세요.

# --hints--

`1` 메서드를 사용해 `add`를 리스트에 추가해야 합니다.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

`2` 메서드를 사용해 `add`를 리스트에 추가해야 합니다.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

`my_list.is_empty()`를 출력해야 합니다.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

`my_list.length`를 출력해야 합니다.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
