---
id: 68c0da2df8ce865dcf7769a7
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

`else` 블록 안에 있으면서 `while`가 `current_node.next`가 아닌지 확인하는 `None` 반복문을 만드세요.

`while` 반복문 안에서 `current_node.next`에 `current_node`를 할당하세요.

# --hints--

`while`가 `current_node.next`가 아닌지 확인하는 `None` 반복문을 만들어야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("current_node.next is not None")
`) })
```

`current_node.next` 반복문 안에서 `current_node`을 `while`에 할당해야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0]
assert _cond.is_equivalent("while current_node.next is not None:\\n current_node = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
