---
id: 587d7fa5367417b2b2512bbd
title: Розширення однієї множини стилів CSS для іншого елемента
challengeType: 0
forumTopicId: 301456
dashedName: extend-one-set-of-css-styles-to-another-element
---

# --description--

У Sass є функція `extend`, яка дозволяє легко запозичувати CSS-правила з одного елемента та розширювати їх в іншому.

Наприклад, наведений нижче блок CSS-правил стилізує клас `.panel`. Він має `background-color`, `height` і `border`.

```scss
.panel{
  background-color: red;
  height: 70px;
  border: 2px solid green;
}
```

Тепер ви хочете створити іншу панель під назвою `.big-panel`. Вона має ті ж базові властивості, що й `.panel`, але також потребує `width` і `font-size`. Можна скопіювати та вставити ініціальні CSS-правила з `.panel`, але код стає повторюваним, коли додаєте більше типів панелей. Директива `extend` — це простий спосіб повторно використовувати правила, написані для одного елемента, а потім додати більше для іншого:

```scss
.big-panel{
  @extend .panel;
  width: 150px;
  font-size: 2em;
}
```

`.big-panel` матиме ті ж властивості, що й `.panel`, а також нові стилі.

# --instructions--

Створіть клас `.info-important`, який розширює `.info` і також має множину `background-color`, встановлену в magenta.

# --hints--

Ваш клас `info-important` повинен мати множину `background-color`, встановлену в `magenta`.

```js
assert(
  code.match(
    /\.info-important\s*?{[\s\S]*background-color\s*?:\s*?magenta\s*?;[\s\S]*}/gi
  )
);
```

Ваш клас `info-important` повинен використовувати `@extend` для наслідування стилів від класу `info`.

```js
assert(
  code.match(/\.info-important\s*?{[\s\S]*@extend\s*?.info\s*?;[\s\S]*/gi)
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }




</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }
  .info-important{
    @extend .info;
    background-color: magenta;
  }



</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```
