---
id: 587d7dbf367417b2b2512bbb
title: "Застосувати стиль, доки не буде виконано умову з @while"
challengeType: 0
forumTopicId: 301454
dashedName: apply-a-style-until-a-condition-is-met-with-while
---

# --description--

Директива `@while` — це опція з подібною функціональністю до JavaScript циклу `while`. Вона створює CSS-правила, доки не буде виконано умову.

Завдання `@for` показало приклад створення простої системи сітки. Це також можна зробити з `@while`.

```scss
$x: 1;
@while $x < 13 {
  .col-#{$x} { width: 100%/12 * $x;}
  $x: $x + 1;
}
```

Спочатку визначте змінну `$x` і встановіть її в 1. Далі використайте директиву `@while`, щоб створити систему сітки *поки* `$x` менше за 13. Після встановлення CSS-правила для `width` змінна `$x` збільшується на 1, щоб уникнути нескінченного циклу.

# --instructions--

Використайте `@while`, щоб створити серію класів з різними `font-sizes`.

Повинно бути 5 різних класів від `text-1` до `text-5`. Потім встановіть `font-size` як `15px`, помножене на поточний індекс. Обов’язково уникайте нескінченного циклу!

# --hints--

Ваш код має використовувати директиву `@while`.

```js
assert(code.match(/@while /g));
```

Ваш код має використовувати змінну індексу, яка починається з 1.

```js
assert(code.match(/\$.*:\s*?1;/gi));
```

Ваш код має збільшувати змінну лічильника.

```js
assert(code.match(/\$(.*)\s*?:\s*\$\1\s*\+\s*1\s*;/gi));
```

Ваш клас `.text-1` має мати `font-size` зі значенням `15px`.

```js
assert($('.text-1').css('font-size') == '15px');
```

Ваш клас `.text-2` має мати `font-size` зі значенням `30px`.

```js
assert($('.text-2').css('font-size') == '30px');
```

Ваш клас `.text-3` має мати `font-size` зі значенням `45px`.

```js
assert($('.text-3').css('font-size') == '45px');
```

Ваш клас `.text-4` має мати `font-size` зі значенням `60px`.

```js
assert($('.text-4').css('font-size') == '60px');
```

Ваш клас `.text-5` має мати `font-size` зі значенням `75px`.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>
  $x: 1;
  @while $x < 6 {
    .text-#{$x}{
      font-size: 15px * $x;
    }
    $x: $x + 1;
  }
</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
