---
id: 587d7dbd367417b2b2512bb5
title: Вкладення CSS за допомогою Sass
challengeType: 0
forumTopicId: 301457
dashedName: nest-css-with-sass
---

# --description--

Sass дозволяє вкладати CSS-правила, що є корисним способом організації стилів.

Зазвичай кожен елемент стилізується на окремій лінії, наприклад:

```scss
article {
  height: 200px;
}

article p {
  color: white;
}

article ul {
  color: blue;
}
```

Для великого проєкту CSS-файл матиме багато ліній і правил. Саме тут вкладення допомагає організувати код, розміщуючи дочірні правила стилів всередині відповідних батьківських елементів:

```scss
article {
  height: 200px;

  p {
    color: white;
  }

  ul {
    color: blue;
  }
}

```

# --instructions--

Використайте техніку вкладення, показану вище, щоб реорганізувати CSS-правила для обох дочірніх елементів `.blog-post`. Для тестування `h1` має йти перед `p`.

# --hints--

Ваш код має реорганізувати CSS-правила так, щоб `h1` і `p` були вкладені в батьківський елемент `.blog-post`.

```js
assert(
  code.match(
    /\.blog-post\s*?{\s*?h1\s*?{\s*?text-align:\s*?center;\s*?color:\s*?blue;\s*?}\s*?p\s*?{\s*?font-size:\s*?20px;\s*?}\s*?}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  .blog-post {

  }
  h1 {
    text-align: center;
    color: blue;
  }
  p {
    font-size: 20px;
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  .blog-post {
    h1 {
      text-align: center;
      color: blue;
    }
    p {
      font-size: 20px;
    }
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```
