---
id: 587d7dbe367417b2b2512bb9
title: Використання @for для створення циклу в Sass
challengeType: 0
forumTopicId: 301462
dashedName: use-for-to-create-a-sass-loop
---

# --description--

Директива `@for` додає стилі в циклі, дуже схожому на цикл `for` у JavaScript.

`@for` використовується двома способами: «від початку до кінця» або «від початку через кінець». Головна різниця в тому, що «від початку **до** кінця» *не включає* кінцеве число у підрахунок, а «від початку **через** кінець» *включає* кінцеве число у підрахунок.

Ось приклад «від початку **через** кінець»:

```scss
@for $i from 1 through 12 {
  .col-#{$i} { width: 100%/12 * $i; }
}
```

Частина `#{$i}` — це синтаксис для поєднання змінної (`i`) з текстом, щоб створити рядок. Коли Sass-файл конвертується в CSS, це виглядає так:

```scss
.col-1 {
  width: 8.33333%;
}

.col-2 {
  width: 16.66667%;
}

...

.col-12 {
  width: 100%;
}
```

Це потужний спосіб створити макет сітки. Тепер у вас є дванадцять варіантів ширини колонок, доступних як CSS-класи.

# --instructions--

Напишіть директиву `@for`, яка приймає змінну `$j`, що йде від 1 **до** 6.

Вона має створити 5 класів з іменами `.text-1` до `.text-5`, де кожен має множину `font-size`, рівну 15px, помноженому на індекс.

# --hints--

Ваш код має використовувати директиву `@for`.

```js
assert(code.match(/@for /g));
```

Ваш клас `.text-1` має мати множину `font-size` 15px.

```js
assert($('.text-1').css('font-size') == '15px');
```

Ваш клас `.text-2` має мати множину `font-size` 30px.

```js
assert($('.text-2').css('font-size') == '30px');
```

Ваш клас `.text-3` має мати множину `font-size` 45px.

```js
assert($('.text-3').css('font-size') == '45px');
```

Ваш клас `.text-4` має мати множину `font-size` 60px.

```js
assert($('.text-4').css('font-size') == '60px');
```

Ваш клас `.text-5` має мати множину `font-size` 75px.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>

@for $i from 1 through 5 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

---

```html
<style type='text/scss'>

@for $i from 1 to 6 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
