---
id: 587d7dbf367417b2b2512bba
title: Використання @each для спроєктування елементів у списку
challengeType: 0
forumTopicId: 301461
dashedName: use-each-to-map-over-items-in-a-list
---

# --description--

Останнє завдання показало, як директива `@for` використовує початкове та кінцеве значення, щоб виконати цикл певну кількість разів. Sass також пропонує директиву `@each`, яка проходить по кожному елементу у списку або множині. На кожній ітерації змінна призначається поточному значенню зі списку або множини.

```scss
@each $color in blue, red, green {
  .#{$color}-text {color: $color;}
}
```

Множина має трохи інший синтаксис. Ось приклад:

```scss
$colors: (color1: blue, color2: red, color3: green);

@each $key, $color in $colors {
  .#{$color}-text {color: $color;}
}
```

Зверніть увагу, що змінна `$key` потрібна для посилання на ключі у множині. Інакше скомпільований CSS міститиме `color1`, `color2`... Обидва наведені приклади коду перетворюються у наступний CSS:

```scss
.blue-text {
  color: blue;
}

.red-text {
  color: red;
}

.green-text {
  color: green;
}
```

# --instructions--

Напишіть директиву `@each`, яка проходить по списку: `blue, black, red` і призначає кожну змінну відповідному класу `.color-bg`, де частина `color` змінюється для кожного елемента на відповідний колір. Кожен клас має встановлювати `background-color` у відповідний колір.

# --hints--

Ваш код має використовувати директиву `@each`.

```js
assert(code.match(/@each /g));
```

Ваш клас `.blue-bg` має мати `background-color` синього кольору.

```js
assert($('.blue-bg').css('background-color') == 'rgb(0, 0, 255)');
```

Ваш клас `.black-bg` має мати `background-color` чорного кольору.

```js
assert($('.black-bg').css('background-color') == 'rgb(0, 0, 0)');
```

Ваш клас `.red-bg` має мати `background-color` червоного кольору.

```js
assert($('.red-bg').css('background-color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

# --solutions--

```html
<style type='text/scss'>

  @each $color in blue, black, red {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

---

```html
<style type='text/scss'>

  $colors: (color1: blue, color2: black, color3: red);

  @each $key, $color in $colors {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```
