---
id: 587d7dbe367417b2b2512bb8
title: Використання @if та @else для додавання логіки до стилів
challengeType: 0
forumTopicId: 301463
dashedName: use-if-and-else-to-add-logic-to-your-styles
---

# --description--

Директива `@if` у Sass корисна для перевірки конкретного випадку — вона працює так само, як інструкція `if` у JavaScript.

```scss
@mixin make-bold($bool) {
  @if $bool == true {
    font-weight: bold;
  }
}
```

І так само, як у JavaScript, директиви `@else if` та `@else` перевіряють додаткові умови:

```scss
@mixin text-effect($val) {
  @if $val == danger {
    color: red;
  }
  @else if $val == alert {
    color: yellow;
  }
  @else if $val == success {
    color: green;
  }
  @else {
    color: black;
  }
}
```

# --instructions--

Створіть міксин з назвою `border-stroke`, який приймає параметр `$val`. Міксин має перевіряти такі умови, використовуючи директиви `@if`, `@else if` та `@else`:

```scss
light - 1px solid black
medium - 3px solid black
heavy - 6px solid black
```

Якщо значення параметра `$val` не є `light`, `medium` або `heavy`, тоді властивість `border` має бути встановлена в `none`.

# --hints--

Ваш код має оголосити міксин з ім’ям `border-stroke`, який має параметр з ім’ям `$val`.

```js
assert(code.match(/@mixin\s+?border-stroke\s*?\(\s*?\$val\s*?\)\s*?{/gi));
```

Ваш міксин має містити інструкцію `@if` для перевірки, чи `$val` дорівнює `light`, і встановлення `border` у `1px solid black`.

```js
assert(
  code.match(
    /@if\s+?\$val\s*?===?\s*?light\s*?{\s*?border\s*?:\s*?1px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Ваш міксин має містити інструкцію `@else if` для перевірки, чи `$val` дорівнює `medium`, і встановлення `border` у `3px solid black`.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?medium\s*?{\s*?border\s*?:\s*?3px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Ваш міксин має містити інструкцію `@else if` для перевірки, чи `$val` дорівнює `heavy`, і встановлення `border` у `6px solid black`.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?heavy\s*?{\s*?border\s*?:\s*?6px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Ваш міксин має містити інструкцію `@else` для встановлення `border` у `none`.

```js
assert(code.match(/@else\s*?{\s*?border\s*?:\s*?none\s*?;\s*?}/gi));
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```

# --solutions--

```html
<style type='text/scss'>
  @mixin border-stroke($val) {
    @if $val == light {
      border: 1px solid black;
    }
    @else if $val == medium {
      border: 3px solid black;
    }
    @else if $val == heavy {
      border: 6px solid black;
    }
    @else {
      border: none;
    }
  }


  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```
