---
id: 587d7dbd367417b2b2512bb4
title: Зберігання даних у змінних Sass
challengeType: 0
forumTopicId: 301460
dashedName: store-data-with-sass-variables
---

# --description--

Одна з функцій Sass, яка відрізняється від CSS, — це використання змінних. Їх оголошують і задають для зберігання даних, подібно до JavaScript.

У JavaScript змінні визначаються за допомогою ключових слів `let` та `const`. У Sass змінні починаються з `$`, за яким іде ім’я змінної.

Ось кілька прикладів:

```scss
$main-fonts: Arial, sans-serif;
$headings-color: green;
```

А щоб використати змінні:

```scss
h1 {
  font-family: $main-fonts;
  color: $headings-color;
}
```

Один із прикладів, де змінні корисні, — коли кілька елементів мають бути одного кольору. Якщо цей колір зміниться, редагувати код потрібно лише в значенні змінної.

# --instructions--

Створіть змінну `$text-color` і задайте їй значення `red`. Потім змініть значення властивості `color` для `.blog-post` і `h2` на змінну `$text-color`.

# --hints--

Ваш код має містити оголошену змінну Sass для `$text-color` зі значенням `red`.

```js
assert(code.match(/\$text-color\s*:\s*?red\s*;/g));
```

Ваш код має використовувати змінну `$text-color` для зміни властивості `color` для елементів `.blog-post` і `h2`.

```js
assert(code.match(/color\s*:\s*\$text-color\s*;?/g));
```

Ваш елемент `.blog-post` має мати властивість `color` зі значенням red.

```js
assert($('.blog-post').css('color') == 'rgb(255, 0, 0)');
```

Ваші елементи `h2` мають мати властивість `color` зі значенням red.

```js
assert($('h2').css('color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>


  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: red;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  $text-color: red;

  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: $text-color;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```
