---
id: 587d7dbf367417b2b2512bbc
title: Розбийте свої стилі на менші частини за допомогою partials
challengeType: 0
forumTopicId: 301459
dashedName: split-your-styles-into-smaller-chunks-with-partials
---

# --description--

<dfn>Partials</dfn> у Sass — це окремі файли, які містять сегменти CSS-коду. Їх імпортують і використовують в інших Sass-файлах. Це чудовий спосіб згрупувати схожий код у модуль, щоб підтримувати порядок.

Назви partials починаються з символу підкреслення (`_`), що повідомляє Sass, що це невеликий сегмент CSS і його не потрібно конвертувати у CSS-файл. Також Sass-файли мають розширення `.scss`. Щоб додати код із partial у інший Sass-файл, використовуйте директиву `@import`.

Наприклад, якщо всі ваші міксини збережені у partial з назвою "\_mixins.scss", і вони потрібні у файлі "main.scss", ось як їх використати у головному файлі:

```scss
@import 'mixins'
```

Зверніть увагу, що підкреслення та розширення файлу не потрібні у `import`-інструкції — Sass розуміє, що це partial. Після імпорту partial у файл усі змінні, міксини та інший код стають доступними для використання.

# --instructions--

Напишіть інструкцію `@import`, щоб імпортувати partial з назвою `_variables.scss` у файл main.scss.

# --hints--

Ваш код має використовувати директиву `@import` і не повинен містити підкреслення у назві файлу.

```js
assert(code.match(/@import\s+?('|")variables\1/gi));
```

# --seed--

## --seed-contents--

```html
<!-- The main.scss file -->
```

# --solutions--

```html
@import 'variables'
```
