---
id: 587d7dbd367417b2b2512bb6
title: Створення повторно використовуваного CSS за допомогою міксинів
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

У Sass <dfn>міксин</dfn> — це множина CSS оголошень, які можна повторно використовувати по всьому стилю. Визначення починається з директиви `@mixin`, за якою йде кастомна назва. Ви застосовуєте міксин за допомогою директиви `@include`.

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

Компіліюється у:

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

Ваші міксини також можуть приймати аргументи, що дозволяє налаштовувати їхню поведінку. Аргументи є обов’язковими при використанні міксина.

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

Ви можете зробити аргументи необов’язковими, задавши параметрам значення за замовчуванням.

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

Напишіть міксин з назвою `shape` і додайте йому 3 параметри: `$w`, `$h` та `$bg-color`.

Використайте міксин `shape`, щоб задати елементу `#square` ширину і висоту `50px`, а також колір фону `red`. Для елемента `#rect-a` додайте ширину `100px`, висоту `50px` і колір фону `blue`. Нарешті, для елемента `#rect-b` додайте ширину `50px`, висоту `100px` і колір фону `orange`.

# --hints--

Ви маєте оголосити міксин з назвою `shape` з 3 параметрами: `$w`, `$h` та `$bg-color`.

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

Ваш міксин повинен містити властивість `width`, яка використовує параметр `$w`.

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

Ваш міксин повинен містити властивість `height`, яка використовує параметр `$h`.

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

Ваш міксин повинен містити властивість `background-color`, яка використовує параметр `$bg-color`.

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

Ви маєте замінити стилі всередині селектора `#square` викликом міксина `shape` з ключовим словом `@include`. Задайте ширину і висоту `50px`, а також колір фону `red`.

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

Ви маєте замінити стилі всередині селектора `#rect-a` викликом міксина `shape` з ключовим словом `@include`. Задайте ширину `100px`, висоту `50px` і колір фону `blue`.

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

Ви маєте замінити стилі всередині селектора `#rect-b` викликом міксина `shape` з ключовим словом `@include`. Задайте ширину `50px`, висоту `100px` і колір фону `orange`.

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
