---
id: 616968c2c94c8071b349c146
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

`body` 안에 `h1` 요소를 중첩하고, 그 안에 `CSS Color Markers` 텍스트를 넣으세요.

# --hints--

코드에 `h1` 요소가 있어야 합니다.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

열린 `h1` 태그가 있어야 합니다.

```js
assert.match(code, /<h1\s*>/i);
```

닫힌 `h1` 태그가 있어야 합니다.

```js
assert.match(code, /<\/h1\s*>/i);
```

`h1` 요소는 `body` 요소 안에 있어야 합니다.

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

`h1` 요소에는 `CSS Color Markers` 텍스트가 있어야 합니다.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

대소문자와 철자가 `h1` 텍스트에 중요하다는 점을 기억하세요.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
