---
id: 616d4a84b756d9c4b8255093
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

마커에 색상을 추가할 시간입니다. 요소에 색상을 추가하는 한 가지 방법은 `black`, `cyan`, 또는 `yellow` 같은 <dfn>색상 키워드</dfn>를 사용하는 것입니다.

참고로, `freecodecamp` 클래스를 선택하는 방법은 다음과 같습니다:

```css
.freecodecamp {
  
}
```

`marker` 클래스를 선택하는 새로운 CSS 규칙을 만들고, 그 `background-color` 속성을 `red`으로 설정하세요.

**참고:** CSS를 추가해도 바로 변화가 보이지 않을 수 있습니다.

# --hints--

`marker` 클래스를 선택하기 위해 클래스 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.marker'));
```

`.marker` CSS 규칙에는 `background-color` 속성이 `red`으로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
