---
id: 616d595270d902f0e7086e18
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

세 개의 별도 마커 `div` 요소가 있지만, 하나의 큰 사각형처럼 보입니다. 각 요소를 더 쉽게 볼 수 있도록 사이에 약간의 공간을 추가하세요.

축약형 `margin` 속성에 두 개의 값이 있으면, 첫 번째 값이 `margin-top`와 `margin-bottom`에 설정되고, 두 번째 값이 `margin-left`와 `margin-right`에 설정됩니다.

`.marker` CSS 규칙에서 `margin` 속성을 `10px auto`로 설정하세요.

# --hints--

`.marker` CSS 규칙에 `margin` 속성을 `10px auto`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

```
