---
id: 616d3a67ccf800ad94ec89ae
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

이제 나중에 색상 표시기로 스타일을 지정할 요소들을 추가하세요.

먼저 `body` 요소 안에 `div` 요소를 추가하고 `class` 속성을 `container`으로 설정하세요. `div` 요소가 `h1` 요소 아래에 위치하도록 하세요.

# --hints--

`div` 여는 태그가 있어야 합니다.

```js
assert.match(code, /<div.*>/i);
```

`div` 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/div\s*>/i);
```

`div` 요소가 `body` 요소 안에 있어야 합니다.

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

`div` 요소가 `h1` 요소 뒤에 있어야 합니다.

```js
assert.exists(document.querySelector('h1 + div'));
```

`div` 요소에 `class` 속성이 `container`으로 설정되어 있어야 합니다.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
