---
id: 616d55bd160a95e22453a081
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

하나의 마커를 색상과 함께 가운데에 배치했으니, 이제 다른 마커들을 추가할 차례입니다.

`.container` 요소 안에 `div` 요소 두 개를 더 추가하고 각각에 `marker` 클래스를 지정하세요.

# --hints--

첫 번째 새 `div` 요소는 여는 태그가 있어야 합니다.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

첫 번째 새 `div` 요소는 닫는 태그가 있어야 합니다.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

두 번째 새 `div` 요소는 여는 태그가 있어야 합니다.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

두 번째 새 `div` 요소는 닫는 태그가 있어야 합니다.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

새 `div` 요소들은 `div` 안에 있고 `container` 클래스를 가져야 합니다.

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

새 `div` 요소들은 둘 다 `marker` 클래스를 가져야 합니다.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
