---
id: 61764fdda535587e1fefb3aa
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

두 번째 마커 `two`에 클래스 `div`를 추가하고, 세 번째 마커 `three`에 클래스 `div`를 추가하세요.

# --hints--

두 번째 마커 `two` 요소에 클래스 `div`를 `.container` 요소 안에 추가하세요.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('two'));
```

두 번째 마커 `div`에는 클래스 `marker`와 `two`가 있어야 합니다.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('marker'));
assert.isTrue(containerSecondChild.classList?.contains('two'));
```

세 번째 마커 `three` 요소에 클래스 `div`를 `.container` 요소 안에 추가하세요.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

세 번째 마커 `div`에는 클래스 `marker`와 `three`가 있어야 합니다.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('marker'));
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

```
