---
id: 61764c602bee6974e7790f35
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

마커에 서로 다른 색상을 주려면 각 마커에 고유한 클래스를 추가해야 합니다. 요소에 여러 클래스를 추가하려면 `class` 속성에 클래스를 나열하고 공백으로 구분하세요. 예를 들어, 다음 코드는 `animal` 요소에 `dog`과 `div` 클래스를 모두 추가합니다:

```html
<div class="animal dog">
```

HTML 요소에 여러 클래스를 추가하면 스타일시트에서 먼저 나열한 클래스의 스타일이 나중에 나열한 클래스에 의해 덮어씌워질 수 있습니다.

먼저, 첫 번째 마커 `one` 요소에 `div` 클래스를 추가하세요.


# --hints--

첫 번째 마커 `one` 요소에 `div` 클래스를 추가해야 합니다.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

첫 번째 마커 `div`에는 `marker`와 `one` 클래스가 있어야 합니다.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
--fcc-editable-region--
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
