---
id: 6176519636a76b810ab1219a
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

클래스 `two`를 선택하는 CSS 규칙을 만들고 그 `background-color` 속성을 `green`로 설정하세요.

또한, 클래스 `three`를 선택하는 별도의 CSS 규칙을 만들고 그 `background-color`를 `blue`로 설정하세요.

# --hints--

클래스 `two`를 선택하려면 클래스 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.two'));
```

`.two` CSS 규칙에 `background-color` 속성을 `green`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

클래스 `three`를 선택하려면 클래스 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.three'));
```

`.three` CSS 규칙에 `background-color` 속성을 `blue`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
