---
id: 617b65579ce424bf5f02ca73
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

수업에서 배운 것처럼, 두 가지 주요 색상 모델이 있습니다: 전자 기기에서 사용하는 <dfn>가산</dfn> RGB(빨강, 초록, 파랑) 모델과 인쇄에 사용하는 <dfn>감산</dfn> CMYK(청록, 자홍, 노랑, 검정) 모델입니다.

이 프로젝트에서는 RGB 모델을 다룹니다. 즉, 색상은 검정에서 시작해 빨강, 초록, 파랑의 다양한 수준이 추가되면서 변합니다. CSS `rgb` 함수를 사용하면 이를 쉽게 확인할 수 있습니다.

클래스 `container`를 대상으로 하는 새로운 CSS 규칙을 만들고 `background-color`를 사용해 `rgb(0, 0, 0)`를 검정으로 설정하세요.

# --hints--

클래스 선택자를 사용해 클래스 `container`를 대상으로 해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.container'));
```

`.container` CSS 규칙에 `background-color` 속성을 `rgb(0, 0, 0)`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
