---
id: 617bd6ec666b1da2587e4e37
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

세 가지 추가 3차 색상이 있습니다: chartreuse green(녹색 + 노란색), azure(파란색 + 청록색), 그리고 rose(빨간색 + 마젠타)입니다.

chartreuse green을 만들려면 `rgb` CSS 규칙의 `.one` 함수를 업데이트하여 빨간색을 `127`에 두고 녹색을 최대값으로 설정하세요.

azure의 경우 `rgb` CSS 규칙의 `.two` 함수를 업데이트하여 녹색을 `127`에 두고 파란색을 최대값으로 설정하세요.

그리고 때때로 밝은 분홍색이라고도 하는 rose는 `rgb` CSS 규칙의 `.three` 함수를 업데이트하여 파란색을 `127`에 두고 빨간색을 최대값으로 설정하세요.

# --hints--

`.one` CSS 규칙에 `background-color` 속성이 `rgb(127, 255, 0)`로 설정되어 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(127, 255, 0)');
```

`.two` CSS 규칙에 `background-color` 속성이 `rgb(0, 127, 255)`로 설정되어 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 255)');
```

`.three` CSS 규칙에 `background-color` 속성이 `rgb(255, 0, 127)`로 설정되어 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
