---
id: 617b954d9f4f6217a749380e
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

이제 마커와 그들이 속한 `container` 요소의 가장자리 사이에 수직 공간을 조금 더 추가하세요.

`.container` CSS 규칙에서 축약형 `padding` 속성을 사용해 위쪽과 아래쪽 패딩을 `10px`만큼 추가하고, 왼쪽과 오른쪽 패딩은 `0`로 설정하세요. 이는 이전에 사용한 축약형 `margin` 속성과 비슷하게 작동합니다.

# --hints--

`background-color` CSS 규칙에서 `.container` 속성과 그 값을 제거하지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

`.container` CSS 규칙에는 `padding` 속성이 `10px 0`으로 설정되어 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.padding, '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
--fcc-editable-region--

--fcc-editable-region--
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
