---
id: 617bc4824e233180553a8069
title: 28단계
challengeType: 0
dashedName: step-28
---

# --description--

삼차 색상인 보라색을 만들려면 마젠타와 파란색을 섞으세요. `rgb` CSS 규칙에서 `.three` 함수를 업데이트하여 파란색을 최대값으로 설정하고 빨간색은 `127`로 설정하세요.

# --hints--

`.three` CSS 규칙에 `background-color` 속성이 `rgb(127, 0, 255)`로 설정되어 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(127, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
--fcc-editable-region--
  background-color: rgb(255, 0, 255);
--fcc-editable-region--
}

```
