---
id: 617b92b9de349513466f6156
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

빨간색과 파란색 마커는 같아 보이지만, 녹색 마커는 이전보다 훨씬 더 연해졌습니다. 이는 `green` 색상 키워드가 실제로 더 어두운 색조이며, 검정과 녹색의 최대값 사이 중간쯤에 있기 때문입니다.

`.two` CSS 규칙에서 `rgb` 함수의 녹색 값을 `127`로 설정해 강도를 낮추세요.

# --hints--

`.two` CSS 규칙에 `background-color` 속성을 `rgb(0, 127, 0)`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
