---
id: 618a132676346ac9f7fd59dd
title: Step 34
challengeType: 0
dashedName: step-34
---

# --description--

가운데 빨간색에 자연스럽게 시선이 집중되는 것을 눈치채셨나요? 사이트를 디자인할 때 이 효과를 사용해 중요한 제목, 버튼, 링크에 주목을 끌 수 있습니다.

보색 외에도 중요한 색 조합이 여러 가지 있지만, 그것들은 나중에 배우게 됩니다.

지금은 `rgb` CSS 규칙에서 `.two` 함수를 사용해 `background-color`를 검은색으로 설정하세요.

# --hints--

`.two` CSS 규칙에 `background-color` 속성이 `rgb(0, 0, 0)`로 설정되어 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 0);
}

```
