---
id: 619b72a0db211f1c29afb906
title: Step 37
challengeType: 0
dashedName: step-37
---

# --description--

새로운 `.one` 클래스를 대상으로 하도록 `red` CSS 규칙을 업데이트하세요.

# --hints--

코드에 더 이상 `.one` 클래스 선택자가 없어야 합니다.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.one'));
```

클래스 `red`를 대상으로 하려면 클래스 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.red'));
```

`.red` CSS 규칙에 `background-color` 속성이 `rgb(0, 0, 0)`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
```
