---
id: 619b761916dac02643940022
title: Step 41
challengeType: 0
dashedName: step-41
---

# --description--

CSS로 요소에 색상을 적용하는 아주 일반적인 방법은 <dfn>16진수</dfn> 또는 헥스 값을 사용하는 것입니다. 헥스 값이 복잡하게 들릴 수 있지만, 사실은 RGB 값을 나타내는 또 다른 형태일 뿐입니다.

헥스 색상 값은 `#` 문자로 시작하며 0-9와 A-F 사이의 여섯 문자를 사용합니다. 첫 번째 두 글자는 빨간색을, 두 번째 두 글자는 초록색을, 세 번째 두 글자는 파란색을 나타냅니다. 예를 들어, `#4B5320`입니다.

`.green` 클래스 선택자에서 `background-color` 속성을 빨간색 `00`, 초록색 `FF`, 파란색 `00` 값을 가진 헥스 색상 코드로 설정하세요.

# --hints--

`.green` CSS 규칙에 `background-color` 속성이 `#00FF00`로 설정되어 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

.blue {
  background-color: rgb(0, 0, 0);
}

```
