---
id: 619b7fd56aa2253778bcf5f7
title: 43단계
challengeType: 0
dashedName: step-43
---

# --description--

<dfn>HSL</dfn> 색상 모델, 즉 색상, 채도, 밝기는 색상을 나타내는 또 다른 방법입니다.

CSS `hsl` 함수는 색상에 대해 0에서 360 사이의 숫자, 채도와 밝기에 대해 0에서 100 사이의 백분율 3개의 값을 받습니다.

색상환을 상상해 보면, 빨간색은 0도, 초록색은 120도, 파란색은 240도에 위치합니다.

채도는 0%일 때 회색, 100%일 때 순수한 색상으로 색상의 강도를 나타냅니다. 채도와 밝기 값에는 반드시 백분율 기호 `%`를 붙여야 합니다.

밝기는 색상이 얼마나 밝게 보이는지를 나타내며, 0%는 완전한 검정, 100%는 완전한 흰색이고, 50%는 중간값입니다.

`.blue` CSS 규칙에서 `hsl` 함수를 사용해 `background-color` 속성을 순수한 파란색으로 바꾸세요. 색상은 `240`, 채도는 `100%`, 밝기는 `50%`로 설정하세요.

# --hints--

`.blue` CSS 규칙에 `background-color` 속성이 `hsl(240, 100%, 50%)`으로 설정되어 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

.blue {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

```
