---
id: 61a489b8579e87364b2d2cdb
title: 44단계
challengeType: 0
dashedName: step-44
---

# --description--

CSS에서 단색을 설정하는 몇 가지 방법을 연습했지만, 요소에 색상 전환 또는 <dfn>그라디언트</dfn>를 사용할 수도 있습니다.

그라디언트는 한 색상이 다른 색상으로 전환되는 것입니다. CSS `linear-gradient` 함수는 전환 방향을 선을 따라 제어하고 어떤 색상을 사용할지 정할 수 있게 해줍니다.

기억할 점은 `linear-gradient` 함수가 실제로 `image` 요소를 생성하며, 보통 이미지를 값으로 받을 수 있는 `background` 속성과 함께 사용된다는 것입니다.

`.red` CSS 규칙에서 `background-color` 속성을 `background`으로 변경하세요.

# --hints--

`.red` CSS 규칙에 `background` 속성이 `rgb(255, 0, 0)` 값으로 있어야 합니다.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
