---
id: 61a49d15bdbb5e57cc6fd280
title: 49단계
challengeType: 0
dashedName: step-49
---

# --description--

색상 중단점은 그래디언트 선을 따라 색상이 배치되는 위치를 세밀하게 조정할 수 있게 해줍니다. `linear-gradient` 함수에서 색상이 끝나고 다음 색상으로 전환되는 위치를 지정하는 데 사용합니다. 색상 중단점 위치를 정의할 때 백분율이나 길이 값을 사용할 수 있습니다.

예를 들어, 이 빨간색-검은색 그래디언트에서는 빨간색에서 검은색으로 전환되는 지점이 그래디언트 선을 따라 90% 지점에 있으므로 빨간색이 대부분의 공간을 차지합니다:

```css
linear-gradient(90deg, red 90%, black);
```

`linear-gradient` 함수에서 첫 번째 빨간색 인수 뒤에 `75%` 색상 중단점을 추가하세요. 다른 색상 인수에는 색상 중단점을 추가하지 마세요.

# --hints--

`.red` CSS 규칙에 `background` 속성을 `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`로 설정해야 합니다.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
