---
id: 61a493ead935c148d2b76312
title: 45단계
challengeType: 0
dashedName: step-45
---

# --description--

`linear-gradient` 함수는 매우 유연합니다 — 이 튜토리얼에서 검토할 기본 구문은 다음과 같습니다:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection`은 전환에 사용되는 선의 방향입니다. `color1`와 `color2`은 전환 자체에 사용될 색상 인수로, 색상 키워드, 16진수, `rgb`, 또는 `hsl` 등 어떤 유형의 색상도 될 수 있습니다.

이제 첫 번째 마커에 90도 선을 따라 빨간색에서 초록색으로 그라디언트를 적용할 것입니다.

먼저, `.red` CSS 규칙에서 `background` 속성을 `linear-gradient()`로 설정하고, `90deg`으로 값 `gradientDirection`를 전달하세요.

# --hints--

`.red` CSS 규칙에 `background` 속성이 있고, 그 값이 `linear-gradient(90deg)`이어야 합니다.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
