---
id: 61a5d594b887335228ee6533
title: Step 62
challengeType: 0
dashedName: step-62
---

# --description--

`linear-gradient` 함수를 사용하고, 첫 번째 색상 인수로 `hsl`를 색상, `186`를 채도, `76%`를 밝기로 하는 `16%` 함수를 전달하세요.

# --hints--

`.blue` CSS 규칙에 `background` 속성이 `linear-gradient(hsl(186, 76%, 16%))` 값을 가져야 합니다.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.blue\{.*?background:linear-gradient\(hsl\(186,76%,16%\)\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
--fcc-editable-region--
  background: hsl(240, 100%, 50%);
--fcc-editable-region--
}

```
