---
id: 61adc91467b5d59328b9f781
title: Step 74
challengeType: 0
dashedName: step-74
---

# --description--

모든 HTML 요소에는 테두리가 있지만, 기본적으로는 `none`로 설정되어 있습니다. CSS를 사용하면 요소의 테두리 모든 측면을 제어할 수 있고, 모든 면에 테두리를 설정하거나 한 면씩 설정할 수도 있습니다. 테두리가 보이려면 너비와 스타일을 설정해야 합니다.

`.sleeve` CSS 규칙에 `border-left-width` 속성을 값 `10px`로 추가하세요.

# --hints--

`.sleeve` CSS 규칙에 `border-left-width` 속성이 있고 값은 `10px`이어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeftWidth, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
--fcc-editable-region--
  
--fcc-editable-region--
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
