---
id: 61b0a44a6b865738ba49b9fb
title: Step 81
challengeType: 0
dashedName: step-81
---

# --description--

마지막으로 각 마커에 약간의 그림자를 추가해 더욱 실감 나게 보이도록 하세요.

`box-shadow` 속성은 요소 주위에 하나 이상의 그림자를 적용할 수 있게 해줍니다. 기본 구문은 다음과 같습니다.

```css
box-shadow: offsetX offsetY color;
```

`offsetX`과 `offsetY` 값이 작동하는 방식은 다음과 같습니다:

* `offsetX`과 `offsetY` 모두 `px` 및 기타 CSS 단위로 숫자 값을 받습니다
* 양수 `offsetX` 값은 그림자를 오른쪽으로 이동시키고 음수 값은 왼쪽으로 이동시킵니다
* 양수 `offsetY` 값은 그림자를 아래로 이동시키고 음수 값은 위로 이동시킵니다
* `0`과 `offsetX` 중 하나 또는 둘 다에 대해 0(`offsetY`) 값을 원한다면 단위를 추가할 필요가 없습니다. 모든 브라우저는 0이 변함없음을 의미한다는 것을 이해합니다.

그림자의 높이와 너비는 그림자가 적용된 요소의 높이와 너비에 의해 결정됩니다. 선택적으로 `spreadRadius` 값을 사용해 그림자의 범위를 더 넓힐 수도 있습니다. 이에 대해서는 나중에 더 다룹니다.

빨간 마커에 간단한 그림자를 추가하는 것부터 시작하세요.

`.red` CSS 규칙에 `box-shadow` 속성을 추가하고 `5px`에 `offsetX`, `5px`에 `offsetY`, `red`에 `color` 값을 넣으세요.

# --hints--

`.red` CSS 규칙에는 `box-shadow` 단축 속성이 있어야 하며 값은 `5px 5px red`이어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow, 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
--fcc-editable-region--
  
--fcc-editable-region--
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
