---
id: 61b095a79e7fc020b43b1bba
title: Step 47
challengeType: 0
dashedName: step-47
---

# --description--

아직 그라디언트가 보이지 않을 것입니다. `linear-gradient` 함수는 작동하려면 최소 두 개의 색상 인수가 필요하기 때문입니다.

같은 `linear-gradient` 함수에서 `rgb` 함수를 사용해 두 번째 색상 인수를 순수한 녹색으로 설정하세요.

# --hints--

`.red` CSS 규칙에 `background` 속성이 `linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0))`로 설정되어 있어야 합니다.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0),rgb(0,255,0))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0),rgb(0,255,0))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(90deg, rgb(255, 0, 0));
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
