---
id: 61b0934c9e7fc020b43b1bb7
title: Step 72
challengeType: 0
dashedName: step-72
---

# --description--

클래스 `cap`를 대상으로 하는 새로운 CSS 규칙을 만드세요. 새 규칙에서 `width` 속성을 `60px`로 설정하고, `height`를 `25px`로 설정하세요.

# --hints--

클래스 `cap`를 대상으로 하려면 클래스 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cap'));
```

`.cap` CSS 규칙에 `width` 속성이 `60px`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cap')?.width, '60px');
```

`.cap` CSS 규칙에 `height` 속성이 `25px`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cap')?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
