---
id: 61b093329e7fc020b43b1bb5
title: 70단계
challengeType: 0
dashedName: step-70
---

# --description--

색상을 설정할 때 `rgb` 함수를 사용했습니다. `rgb` 색상에 알파 채널을 추가하려면 대신 `rgba` 함수를 사용하세요.

`rgba` 함수는 `rgb` 함수와 똑같이 작동하지만, 알파 채널을 위해 `0`부터 `1.0`까지 숫자 하나를 더 받습니다:

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

알파 채널은 `hsl`과 `hex` 색상에도 사용할 수 있습니다. 곧 그 방법을 배우게 될 것입니다.

`.sleeve` 규칙에서 `rgba` 함수를 사용해 `background-color` 속성을 불투명도 50%인 순수한 흰색으로 설정하세요.

# --hints--

`.sleeve` CSS 규칙에 `background-color` 속성이 `rgba(255, 255, 255, 0.5)`로 설정되어 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor, 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
--fcc-editable-region--
  background-color: white;
--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
