---
id: 61b095989e7fc020b43b1bb9
title: Step 46
challengeType: 0
dashedName: step-46
---

# --description--

이 그라디언트의 색상을 위해 `rgb` 함수를 사용하세요.

`linear-gradient` 함수 안에서 `rgb` 함수를 사용해 첫 번째 색상 인수를 순수한 빨간색으로 설정하세요.

# --hints--

`.red` CSS 규칙에 `background` 속성이 `linear-gradient(90deg, rgb(255, 0, 0))` 값을 가져야 합니다.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg,rgb\(255,0,0\)\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(90deg);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
