---
id: 61b0a1b2af494934b7ec1a72
title: Step 80
challengeType: 0
dashedName: step-80
---

# --description--

멋져요. 빨간 마커가 잘 보입니다. 이제 다른 마커들에도 캡과 소매를 추가하기만 하면 됩니다.

초록색과 파란색 마커 모두에 캡과 소매를 추가하세요. 빨간 마커에서 `div` 요소를 복사해서 다른 두 마커에 붙여넣으면 됩니다.

# --hints--

초록색 마커 `div`에는 `div` 요소가 두 개 있어야 합니다.

```js
const greenMarkerChildren = [...document.querySelector('.green')?.children];
greenMarkerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'))
assert.lengthOf(greenMarkerChildren, 2);
```

초록색 마커의 캡 `div` 요소는 소매 `div` 요소보다 앞에 있어야 합니다.

```js
const greenMarkerChildren = [...document.querySelector('.green')?.children];
const greenMarkerCap = document.querySelector('.green .cap');
const greenMarkerSleeve = document.querySelector('.green .sleeve');
assert.isBelow(greenMarkerChildren.indexOf(greenMarkerCap), greenMarkerChildren.indexOf(greenMarkerSleeve));
```

파란색 마커 `div`에는 `div` 요소가 두 개 있어야 합니다.

```js
const blueMarkerChildren = [...document.querySelector('.blue')?.children];
blueMarkerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(blueMarkerChildren, 2);
```

파란색 마커의 캡 `div` 요소는 소매 `div` 요소보다 앞에 있어야 합니다.

```js
const blueMarkerChildren = [...document.querySelector('.blue')?.children];
const blueMarkerCap = document.querySelector('.blue .cap');
const blueMarkerSleeve = document.querySelector('.blue .sleeve');
assert.isBelow(blueMarkerChildren.indexOf(blueMarkerCap), blueMarkerChildren.indexOf(blueMarkerSleeve));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      --fcc-editable-region--
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
