---
id: 61b09f739aa6572d2064f9b8
title: Step 79
challengeType: 0
dashedName: step-79
---

# --description--

테두리의 검은색이 더 투명한 소매와 대비되어 꽤 강해 보입니다. 알파 채널을 사용해 검은 테두리의 불투명도를 낮출 수 있습니다.

`border-left` 단축 속성에 `rgba` 함수를 사용해 색상 값을 75% 불투명도의 순수한 검은색으로 설정하세요.

# --hints--

`.sleeve` CSS 규칙에는 `border-left` 단축 속성이 있어야 하며 값은 `10px double rgba(0, 0, 0, 0.75)`이어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeft, '10px double rgba(0, 0, 0, 0.75)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
--fcc-editable-region--
  border-left: 10px double black;
--fcc-editable-region--
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
