---
id: 61b306305810f1c9040ce5a2
title: Step 82
challengeType: 0
dashedName: step-82
---

# --description--

보시다시피, 마커 주위에 오른쪽으로 5픽셀, 아래로 5픽셀 떨어진 간단한 빨간 그림자를 추가했습니다.

그런데 그림자를 반대쪽에 위치시키고 싶다면 어떻게 할까요? `offsetX`와 `offsetY`에 음수 값을 사용하면 됩니다.

`box-shadow` 속성의 값을 업데이트하고, `offsetX`를 `-5px`로, `offsetY`를 `-5px`로 설정하세요.

# --hints--

`.red` CSS 규칙에는 `box-shadow` 단축 속성이 있어야 하며, 값은 `-5px -5px red`이어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow, 'red -5px -5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
--fcc-editable-region--
  box-shadow: 5px 5px red;
--fcc-editable-region--
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
