---
id: 61b3183655ec10efd8c0bb07
title: Step 87
challengeType: 0
dashedName: step-87
---

# --description--

초록색과 파란색 마커의 그림자는 위치, 흐림, 확산이 동일합니다. 유일한 차이는 색상입니다.

`.green`와 `.blue` CSS 규칙에서 `box-shadow` 속성들의 값을 업데이트하여 `offsetX`가 `0`, `offsetY`가 `0`, `blurRadius`가 `20px`, `spreadRadius`가 `0`가 되도록 하세요. 색상은 당분간 `green`와 `blue`로 남겨두세요.

# --hints--

`.green` CSS 규칙에는 `box-shadow` 단축 속성이 있어야 하며, 값은 `0 0 20px 0 green`여야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.boxShadow, 'green 0px 0px 20px 0px');
```

`.blue` CSS 규칙에는 `box-shadow` 단축 속성이 있어야 하며, 값은 `0 0 20px 0 blue`여야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.boxShadow, 'blue 0px 0px 20px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: 0 0 20px 0 rgba(83, 14, 14, 0.8);
}

--fcc-editable-region--
.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
  box-shadow: 5px 5px 5px green;
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
  box-shadow: 0 0 0 5px blue;
}
--fcc-editable-region--

```
