---
id: 63ee3f71381756f9716727ef
title: CSS 基础练习 A
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**目标：**  
在本练习中，你将练习使用三种方法将 CSS 添加到超文本标记语言文件中：外部 CSS、内部 CSS 和内联 CSS。通过外部和内部方法添加样式时，你应仅使用类型选择器。你还应使用颜色关键字（例如 "blue"），而不是使用 `RGB` 或 `HEX` 值。

## 用户故事

- 你应该看到一个带有一些文本的 `div` 元素。
  - 它应该有一个 `red` 背景，`white` 文本，`32px` 的字体大小，文本居中对齐且为 `bold`。
  - `div` 元素的 CSS 应该外部添加，并使用类型选择器。
- 你应该看到一个带有一些文本的 `p` 元素。
  - 它应该有一个 `green` 背景，`white` 文本，和 `18px` 的字体大小。
  - `p` 元素的 CSS 应该内部添加，并使用类型选择器。
- 你应该看到一个带有一些文本的 `button` 元素。
  - `button` 元素应具有 `orange` 背景和 `18px` 字体大小。
  - `button` 元素的 CSS 应该使用内联样式添加。

# --hints--

你应该有一个包含一些文本的 `div` 元素。

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

你应该有一个包含 `div` 元素样式的外部样式表。

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

你的 `div` 元素的 CSS 不应使用内部或内联样式添加。

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

你的 `div` 元素应具有 `red` 的 `background-color` 和 `white` 的 `color`。

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

你的 `div` 元素的 `font-weight` 应设置为 `bold`，`font-size` 应设置为 `32px`，并且 `text-align` 应设置为 `center`。

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

你应该有一个 `p` 元素，并且它应包含一些文本。

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

你的 `p` 元素的样式应使用 `style` 元素在内部添加。

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

你的 `p` 元素的 `font-size` 应为 `18px`，并且 `color` 设置为 `white`。

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

你应该有一个包含一些文本的 `button` 元素。

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

你的 `button` 元素应具有内联样式。

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```
 
你的 `button` 元素的 `background-color` 应设置为 `orange`。

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

你的 `button` 元素的 `font-size` 应设置为 `18px`。

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
