---
id: 63ee3fe4381756f9716727f0
title: CSS 基础练习 B
challengeType: 25
dashedName: css-foundations-exercise-b
---

# --description--

**目标：** 提供的 HTML `file` 中有多个 `Element`，你需要为它们添加 `class` 或 `ID` `attribute`。然后你需要使用正确的选择器 `syntax` 在提供的 CSS `file` 中添加规则。

## 用户故事

1. 你应该看到列表中所有奇数编号元素的 `yellow` 背景。

1. 你应该为列表中所有奇数编号的元素使用 `class` 选择器。

1. 你应该看到列表中的第二个元素具有 `blue` 文本和 `font-size` 为 `36px`。

1. 第二个元素上的 `font-size` 和文本颜色应通过使用 `id` 属性来设置。

1. 你应该看到列表中的第三个元素的 `font-size` 是 `24px`。

1. 第三个元素上的 `font-size` 应该通过使用 `class` 属性来设置。

1. 你应该看到列表中的第四个元素有一个 `red` 背景，一个 `font-size` 为 `24px`，以及一个 `font-weight` 为 `bold`。

1. 第四个元素的 `font-size` 应该通过 `class` 属性设置。`font-weight` 和颜色应该通过 `id` 属性设置。

# --hints--

每个奇数 `Element` 应该有一个 `class` 属性。

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasClass = p?.every((paragraph) => paragraph.classList.length > 0);

assert(everyPHasClass);
```

你的奇数元素应具有 `background-color` 为 `yellow`。

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasBackgroundColor = p?.every((paragraph) => {
  const style = getComputedStyle(paragraph);
  
  return style?.backgroundColor === 'rgb(255, 255, 0)';
})

assert.equal(everyPHasBackgroundColor, true);
```

你的第二个元素应具有 `blue` 文本和 `font-size` 为 `36px`。

```js
const secondElementId = document.querySelectorAll('div')?.[0]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${secondElementId}`);

assert.equal(style?.color, 'blue');
assert.equal(style?.fontSize, '36px');
```

你的第三个元素应具有文本和 `font-size` 为 `24px`。

```js
const thirdElement = document.querySelectorAll('p')?.[1];

assert(thirdElement?.innerText?.length > 0);

const thirdElementClasses = Array.from(thirdElement?.classList?.values());

assert(thirdElementClasses.some(thirdElementClass => {

  const style = new __helpers.CSSHelp(document).getStyle(`.${thirdElementClass}`);

  return style?.fontSize === '24px';

}))

``` 

第四个元素的 `font-size` 应为 `24px`。

```js
const fourthElementClass = document.querySelectorAll('div')?.[1]?.classList[0];

const style = new __helpers.CSSHelp(document).getStyle(`.${fourthElementClass}`);

assert(style?.fontSize === '24px');
```

第四个元素应具有 `red` 的 `background-color`。

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.backgroundColor === 'red');
```

第四个元素的 `font-weight` 应为 `bold`。

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.fontWeight === 'bold');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p>Number 1 - I'm a class!</p>
    <div>Number 2 - I'm one ID.</div>
    <p>Number 3 - I'm a class, but cooler!</p>
    <div>Number 4 - I'm another ID.</div>
    <p>Number 5 - I'm a class!</p>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p class="odd">Number 1 - I'm a class!</p>
    <div id="two">Number 2 - I'm one ID.</div>
    <p class="odd adjust-font-size">Number 3 - I'm a class, but cooler!</p>
    <div id="four" class="adjust-font-size">Number 4 - I'm another ID.</div>
    <p class="odd">Number 5 - I'm a class!</p>
  </body>
</html>
```

```css
.odd {
  background-color: yellow;
  font-family: Verdana, "DejaVu Sans", sans-serif;
}

.adjust-font-size {
  font-size: 24px;
}

#two {
  color: blue;
  font-size: 36px;
}

#four {
  background-color: red;
  font-weight: bold;
}
```
