---
id: 66ea7fc048a6e26e33c598fe
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Ahora, es hora de generar un número aleatorio usando el método `Math.random()`.

Crea una variable llamada `randomNum` y asígnale el valor del resultado de llamar al método `Math.random()`.

Luego, registra la variable `randomNum` en la consola.

Intenta agregar un espacio en el código para reiniciar el bot y ver diferentes números aleatorios que se generan.

# --hints--

Deberías tener una variable llamada `randomNum`.

```js
assert.isNotNull(randomNum);
```

Deberías asignar el resultado de llamar al método `Math.random()` a la variable `randomNum`.

```js
assert.match(code, /randomNum\s*=\s*Math\.random\(\);?/);
```

No debes codificar manualmente un número aleatorio para la variable `randomNum`. Asegúrate de usar el método `Math.random()` para generar el número aleatorio.

```js
assert.notMatch(code, /randomNum\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum\s*=\s*\d+;?/);
```

Deberías registrar la variable `randomNum` en la consola.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
