---
id: 66eb3ec3b07dfe8f898646c0
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Como aprendiste en los videos de la conferencia anterior, el método `Math.floor()` redondea el valor hacia abajo al número entero más cercano.

```js
const price = 10.99;
Math.floor(price); // 10
```

Crea una variable llamada `numRoundedDown` y asígnale el resultado de redondear el número flotante `6.7` hacia abajo al número entero más cercano.

Después, registra la variable `numRoundedDown` en la consola para ver el resultado.

# --hints--

Deberías tener una variable `numRoundedDown`.

```js
assert.isNotNull(numRoundedDown);
```

La variable `numRoundedDown` debería ser un número.

```js
assert.isNumber(numRoundedDown);
```

Tu variable `numRoundedDown` debería ser igual al resultado de redondear el número flotante `6.7` hacia abajo al número entero más cercano.

```js
assert.equal(numRoundedDown, 6);
```

No deberías codificar de manera rígida el valor de `6` para la variable `numRoundedDown`.

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

Deberías registrar la variable `numRoundedDown` en la consola.

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
