---
id: 66ece147323c91f654bf8dd9
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

En los vídeos de la conferencia anterior, aprendiste cómo obtener los valores máximos y mínimos de un rango de números como este:

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

Crea una variable llamada `maxNum` y asígnale el resultado de encontrar el número máximo entre los números `3, 125, 55, 24`. Luego, registra el valor de `maxNum` en la consola.

Debajo de eso, crea una variable llamada `minNum` y asígnale el resultado de encontrar el número mínimo entre los números `6, 90, 14, 90, 2`. Luego, registra el valor de `minNum` en la consola.

# --hints--

Debes tener una variable llamada `maxNum`.

```js
assert.isNotNull(maxNum);
```

Tu variable `maxNum` debe ser asignada al resultado de encontrar el número máximo entre los números `3, 125, 55, 24`.

```js
assert.equal(maxNum, 125);
```

No debes codificar directamente el valor de `125` en la variable `maxNum`. Asegúrate de usar el método `Math.max()` en su lugar.

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

Debes registrar el valor de `maxNum` en la consola.

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

Debes tener una variable llamada `minNum`.

```js
assert.isNotNull(minNum);
```

Tu variable `minNum` debe ser asignada al resultado de encontrar el número mínimo entre los números `6, 90, 14, 90, 2`.

```js
assert.equal(minNum, 2);
```

No debes codificar directamente el valor de `2` en la variable `minNum`. Asegúrate de usar el método `Math.min()` en su lugar.

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

Debes registrar el valor de `minNum` en la consola.

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
