---
id: 685239f0cad43c7a909cbd98
title: 반려동물 입양 페이지 디버그하기
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

반려동물 입양 저장소 주인인 Sally가 첫 번째 웹 페이지를 만들었지만 몇 가지 문제가 있습니다.

Sally가 페이지를 계속 만들 수 있도록 모든 오류를 고치세요.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. Sally는 고양이 사진을 사용하고 싶지만 제대로 표시되지 않습니다. `img` 요소에서 다음을 고쳐야 합니다:
   - 이미지 소스를 위한 올바른 속성으로 `href` 속성을 교체하세요.
   - 이미지에 대한 짧고 설명적인 텍스트를 나타내는 올바른 속성으로 `att` 속성을 교체하세요.
   - `</img>` 요소는 닫는 태그가 없는 void 요소이므로 `img` 닫는 태그를 제거하세요.
2. Sally는 사용자들을 개와 고양이 페이지로 안내하는 링크를 사용하고 싶지만 링크가 제대로 작동하지 않습니다. `a` 요소에서 다음을 고쳐야 합니다:
   - URL을 지정하는 데 사용되는 올바른 속성으로 두 `src` 속성을 교체하세요.

# --hints--

`img` 요소는 `src` 속성 대신 `href` 속성을 가져야 합니다.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

`img` 요소는 존재하지 않는 `alt` 속성 대신 `att` 속성을 가져야 합니다.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

`img` 요소는 `</img>` 닫는 태그가 없어야 합니다.

```js
assert.notMatch(code, /<\/img>/);
```

텍스트가 `a`인 `Visit cats page` 요소는 `href` 속성 대신 `src` 속성을 가져야 합니다.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

텍스트가 `a`인 `Visit dogs page` 요소는 `href` 속성 대신 `src` 속성을 가져야 합니다.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
