---
id: 7a97b79c096ecb000d03bd36
title: 建置第 N 個費氏數列計算器
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成此實驗。

**使用者故事：**

1. 你應該創建一個名為 `fibonacci` 的函式。
2. 你應該在 `sequence` 函式內定義一個名為 `fibonacci` 的陣列，並將其初始化為 `[0, 1]` 的值。
3. `fibonacci` 函式應接受一個參數，該參數為正整數 `n`。
4. 呼叫 `fibonacci(n)` 時，應使用動態規劃方法計算並傳回費氏數列中第 `n` 個數字，每個數字都是前兩個數字的和。
5. 費氏數列中第 `n` 個位置計算出的每個數字，應存放在 `sequence` 陣列的索引 `n - 1`。

# --hints--

你應該有一個名為 `fibonacci` 的函式。

```js
assert.isFunction(fibonacci);
```

你的 `fibonacci` 函式應接受一個參數。

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

你應該在 `sequence` 函式內有一個名為 `fibonacci` 的陣列，並初始化為 `[0, 1]`。

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` 應傳回 `0`。

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` 應傳回 `1`。

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` 應傳回 `1`。

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` 應傳回 `2`。

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` 應傳回 `5`。

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` 應傳回 `55`。

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` 應傳回 `610`。

```js
assert.equal(fibonacci(15), 610);
```

你的程式碼中不應使用遞迴。

```js
const bodyMatch = __helpers.removeJSComments(code).match(/function\s+fibonacci\s*\([^)]*\)\s*\{([\s\S]*)\}/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
