---
id: 7a97b79c096ecb000d03bd36
title: 建置一個第 N 個費氏數列計算器
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `fibonacci` 的函式。
2. 你應該在 `sequence` 函式內定義一個名為 `fibonacci` 的陣列，並將其初始化為 `[0, 1]` 的值。
3. `fibonacci` 函式應接受一個參數，該參數為非負整數 `n`。
4. 呼叫 `fibonacci(n)` 時，應使用動態規劃方法計算並傳回費氏數列中第 `n` 個數字，每個數字都是前兩個數字的和。
5. 計算出的每個費氏數字應附加到 `sequence` 陣列中。

# --hints--

你應該有一個名為 `fibonacci` 的函式（程式）。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.fibonacci);
```

你的 `fibonacci` 函式應該接受一個單一參數。

```js
const explorer = await __helpers.Explorer(code);
const { fibonacci } = explorer.allFunctions;
assert.equal(fibonacci?.parameters.length, 1);
```

你應該在 `sequence` 函式內有一個名為 `fibonacci` 的陣列，並初始化為 `[0, 1]`。

```js
const explorer = await __helpers.Explorer(code);
const { fibonacci } = explorer.allFunctions;
const { sequence } = fibonacci?.variables ?? {};
assert.isTrue(sequence?.value.matches('[0, 1]'));
```

`fibonacci(0)` 應該傳回 `0`。

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` 應該傳回 `1`。

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` 應該傳回 `1`。

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` 應該傳回 `2`。

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` 應該傳回 `5`。

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` 應該傳回 `55`。

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` 應該傳回 `610`。

```js
assert.equal(fibonacci(15), 610);
```

你不應該在你的程式碼中使用遞迴。

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
