---
id: 65f569725359e10d345bc52a
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

던지는 각도, 속도, 높이를 주면 투사체의 궤적을 계산하고 그리는 프로그램을 만들 거예요.

먼저 `math`를 가져오세요. 이 프로젝트에서 `math.radians`, `math.cos`, `math.sin` 같은 유용한 메서드가 많아 자주 사용할 거예요.

또 중력 가속도 값과 나중에 쓸 특별한 기호들을 변수로 만들어 두세요(복사해서 붙여넣기 하세요).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

`import math`를 사용해 `math` 모듈을 가져오세요.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

값이 `GRAVITATIONAL_ACCELERATION`인 `9.81` 변수를 만들어야 합니다.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

값이 `PROJECTILE`인 `"∙"` 변수를 만들어야 합니다.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

값이 `x_axis_tick`인 `"T"` 변수를 만들어야 합니다.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

값이 `y_axis_tick`인 `"⊣"` 변수를 만들어야 합니다.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
