---
id: 65fd4bd84561a14d3e720061
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

클래스 변수 `__slots__`는 Python 클래스에서 특별한 용도로 사용됩니다. `__slots__`을 선언하고 문자열 수열을 할당하면, 그 수열에 포함된 속성만 생성할 수 있도록 제한합니다. 또한 `__dict__` 특별 속성의 생성을 막고, 속성 접근을 더 효율적으로 만듭니다.

클래스 내부에서 `__slots__` 변수를 사용해 클래스가 어떤 속성을 가지는지 정의해야 합니다: `__slots__`에 `__init__`에 정의된 속성 이름 중 하나와 같은 문자열 3개를 포함하는 튜플을 할당하세요.

# --hints--

`__slots__`에 `'__height'`, `'__speed'`, `'__angle'`를 포함하는 튜플을 할당해야 합니다.

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
