---
id: 65fd4b3bedee044b4b957d33
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

세 개의 인수, 즉 시작 속도, 시작 높이, 그리고 투사체 던지기의 시작 각도를 이 순서대로 받아 클래스를 초기화하는 `Projectile` 클래스를 `__init__` 메서드와 함께 만드세요.

`__init__` 메서드 안에서 이 인수들을 세 개의 비공개 속성인 `__speed`, `__height`, `__angle`에 할당하세요. 제공된 각도는 도 단위이지만 내부적으로는 라디안 단위로 저장해야 합니다. 이를 위해 각도를 도에서 라디안으로 변환할 때 `math.radians()` 함수를 사용하여 `__angle`에 할당하세요.

속성 이름 앞에 밑줄 두 개를 붙이면 Python에서 이름 맹글링이 발생합니다. 이는 해당 속성들이 클래스 외부에서 주어진 이름으로 직접 접근할 수 없으며, 외부에서 필요할 경우 `ball._Projectile__height` 같은 맹글링된 이름으로 접근해야 한다는 뜻으로, 내부 용도로만 사용됨을 나타냅니다.

# --hints--

클래스를 만들 때 `class Projectile:`를 사용하세요.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

`__init__` 클래스 안에 `Projectile` 메서드가 있어야 합니다.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

새로운 `Projectile` 클래스 인스턴스는 세 속성의 값이 제대로 초기화되어야 합니다.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
