---
id: 660400f02cf3a76795b82d4e
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

먼저 설정해야 할 것은 투사체의 변위를 계산하는 메서드입니다. 변위는 투사체가 던져진 지점부터 땅에 닿을 때까지 이동한 수평 거리입니다.

매개변수가 `__calculate_displacement` 하나뿐인 `self` 메서드를 만들고, 투사체의 변위를 반환하세요.

다음 공식을 사용해 투사체 변위를 계산하세요:
\\[ d = \frac{v \cdot \cos\theta \cdot \left(v\cdot\sin\theta + \sqrt{v^2 \cdot \sin^2\theta + 2 \cdot g \cdot h}\right)}{g} \\]

여기서 $d$는 변위, $v$는 초기 속도, $\theta$는 각도, $h$는 투사체의 초기 높이입니다.
$g$에는 `GRAVITATIONAL_ACCELERATION` 변수를 사용하세요.

삼각함수에는 `math.cos()`와 `math.sin()` 메서드를, 제곱근 계산에는 `math.sqrt()`를 사용하세요. 또한 파이썬에서 $x^y$는 `x ** y`로 표현합니다. $\sin^2\theta$는 사인 값에 제곱을 한 것을 의미합니다.

이름 맹글링을 사용할 경우, 메서드를 테스트하거나 클래스 외부에서 사용하려면 `_Projectile__calculate_displacement`로 호출해야 한다는 점을 기억하세요.

```py
ball = Projectile(10, 3, 45)
displacement_of_ball = ball._Projectile__calculate_displacement() # 12.6173996009878
```

# --hints--

`__calculate_displacement`라는 이름의 메서드를 `def __calculate_displacement(self):`와 함께 선언하세요.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").has_function("__calculate_displacement")`))})
```

`__calculate_displacement` 메서드는 `self` 인수만 가져야 합니다.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_args('self')`))})
```

`__calculate_displacement` 메서드는 올바른 값을 반환해야 합니다.

```js
({test: () => assert(runPython(`
p = Projectile(20, 21, 22)
disp = p._Projectile__calculate_displacement()
round(disp, 2) == 55.06 and round(disp, 2) != disp
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    
--fcc-editable-region--
```
