---
id: 660fcf3e1b9bb056b2edb567
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이제 문자열 표현을 만들 준비가 되었습니다.
`Projectile` 클래스를 인스턴스화하는 것부터 시작하세요. 클래스 정의 아래에 `ball` 변수를 만들고 `Projectile`를 인수로 사용하여 `10, 3, 45`를 호출한 결과를 할당하세요.

클래스 내에 `__str__` 메서드를 정의하여 인스턴스에 적절한 문자열 표현을 제공하고, 다음 형식의 문자열을 반환하게 하세요:

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

문자열은 반드시 줄 바꿈 문자로 시작하고 끝나야 하며, 각도는 정수로 도 단위로 작성하고 변위는 소수점 한 자리까지 출력해야 합니다.

라디안을 도 단위로 변환하는 데 유용한 `math.degrees`가 있습니다.

준비가 되면 `print(ball)`를 호출하여 함수를 테스트할 수 있습니다.

# --hints--

`ball` 변수가 있어야 합니다.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

`ball` 변수는 `Projectile(10, 3, 45)` 값을 가져야 합니다.

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

`Projectile(45, 45, 45)`의 문자열 표현이 올바른지 확인하세요.

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

다른 인스턴스에 대해서도 문자열 표현이 올바른지 확인하세요.

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

`print(ball)` 호출이 있어야 합니다.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
