---
id: 6616ec0bf9ee8548ce6b0f08
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

이제 궤적의 좌표를 계산할 차례입니다! `Projectile` 클래스에 `__calculate_y_coordinate`라는 메서드를 만들고, `self` 외에 `x` 매개변수를 하나 추가하세요.

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

위 식은 시작 각도 $\theta$, 속도 $v_0$, 높이 $y_0$가 주어졌을 때 임의의 수평 위치 $x$에 대한 수직 위치 $y$를 계산하는 공식입니다.

`math.tan()`와 `math.cos()`를 사용해야 하며, `x ** y`가 $x^y$를 표현하는 방법이라는 점과 $g$ 값이 `GRAVITATIONAL_ACCELERATION` 변수에 있다는 점을 기억하세요.

메서드를 구현하여 $y$ 좌표를 반환하도록 하세요.

메서드를 테스트하고 싶다면, 클래스 외부에서 `ball._Projectile__calculate_y_coordinate()`가 선언된 후 숫자를 인수로 하여 `ball`를 호출할 수 있습니다.

# --hints--

`Projectile` 클래스에 `__calculate_y_coordinate`라는 메서드가 있어야 합니다.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

`__calculate_y_coordinate` 메서드는 `self, x`를 인수로 가져야 합니다.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

메서드는 올바른 값을 반환해야 합니다. `Projectile(12, 12, 12)`에 대해, $x$ 좌표가 0, 1, 2, 3일 때 y 좌표는 대략 12, 12.18, 12.28, 12.32여야 합니다.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

메서드는 올바른 값을 반환해야 합니다. `Projectile(45, 12, 22)`에 대해, $x$ 좌표가 6, 9, 10, 12일 때 y 좌표는 대략 14.32, 15.41, 15.76, 16.44여야 합니다.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

`__calculate_y_coordinate` 메서드는 결과를 반올림하지 않아야 합니다.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
