---
id: 662fc35902038376c7a00b3c
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

getter를 작성한 후에는 setter를 작성할 수 있습니다. setter는 속성의 값을 간접적으로 설정할 수 있게 해줍니다.
지난 단계의 예를 따라 setter는 다음과 같이 작성합니다:

```py
class Nest:
    ...
    @number_of_eggs.setter
    def number_of_eggs(self, new_value):
        self.__number_of_eggs = new_value
```

getter와 마찬가지로 setter는 메서드처럼 호출하지 않고 속성처럼 사용합니다:

```py
nest = Nest()
nest.number_of_eggs = 12
```

이렇게 작성하면 setter가 호출되어 새 값을 설정합니다.

이번 단계에서는 세 개의 private 속성에 대한 setter를 작성해야 합니다. 각도 값은 도 단위로 받지만 내부적으로는 라디안 단위로 저장된다는 점을 기억하세요.

# --hints--

새로운 `speed` 메서드를 정의하세요.

```js
({test: () => assert(runPython(`
len(_Node(_code).find_class('Projectile').find_functions('speed')) == 2
`))})
```

새로운 `speed` 메서드에는 `@speed.setter` 데코레이터가 있어야 합니다.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_functions('speed')[1].has_decorators('speed.setter')
`))})
```

새로운 `speed` 메서드는 private `__speed` 속성의 값을 설정해야 합니다.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed = 33
a.speed == 33
`))})
```

새로운 `height` 메서드를 정의하세요.

```js
({test: () => assert(runPython(`
len(_Node(_code).find_class('Projectile').find_functions('height')) == 2
`))})
```

새로운 `height` 메서드에는 `@height.setter` 데코레이터가 있어야 합니다.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_functions('height')[1].has_decorators('height.setter')
`))})
```

새로운 `height` 메서드는 private `__height` 속성의 값을 설정해야 합니다.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.height = 44
a.height == 44
`))})
```

새로운 `angle` 메서드를 정의하세요.

```js
({test: () => assert(runPython(`
len(_Node(_code).find_class('Projectile').find_functions('angle')) == 2
`))})
```

새로운 `angle` 메서드에는 `@angle.setter` 데코레이터가 있어야 합니다.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_functions('angle')[1].has_decorators('angle.setter')
`))})
```

`angle` 메서드는 private `__angle` 속성의 값을 설정해야 합니다.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.angle = 12
a.angle == 12
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed


        
--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
