---
id: 662fa453ec8033322b3ab2e3
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

캡슐화와 게터에 대해 이야기할 시간입니다. 세 개의 인스턴스 속성을 앞에 이중 밑줄을 붙여 private으로 작성했습니다. 이 속성들은 관례상 private이라고 부르지만, 여전히 외부에서 접근할 수는 있으나 그렇게 하지 않기로 약속한 것입니다.

게터는 외부에서 값을 가져올 때 사용할 수 있습니다. 게터를 정의하려면, 원하는 속성의 값을 반환하는 메서드를 정의하고 `@property` 데코레이터를 붙이세요:

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

이 데코레이터는 메서드를 속성으로 바꾸어, 일반 메서드처럼 호출하지 않고 속성처럼 사용하게 만듭니다:

```py
n = Nest()
print(n.number_of_eggs)
```

위 예제에서 private 속성 `__number of eggs`은 `number_of_eggs`의 `n` 속성을 통해 접근됩니다.

`speed`라는 게터를 만들고 private 속성 `__speed`의 값을 반환하게 만드세요.

# --hints--

`speed` 메서드를 정의해야 합니다.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

`speed` 메서드는 `self`라는 단일 매개변수를 가져야 합니다.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

`speed` 메서드에는 `@property` 데코레이터가 있어야 합니다.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

`speed` 메서드는 private 속성 `__speed`의 값을 반환해야 합니다.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
