---
id: 662f9f65c50bd21d05ffdee9
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`calculate_all_coordinates`라는 이름의 메서드를 정의하세요. 이 메서드는 `0`부터 변위의 올림값(포함하지 않음)까지 모든 \\(x\\) 값에 대한 좌표를 계산한 후, 튜플 `(x, y)`의 리스트로 반환해야 합니다.

그런 다음 `calculate_all_coordinates` 인스턴스에서 `ball` 메서드를 호출하고, 결과를 `coordinates`라는 새 변수에 할당하세요.

숫자를 그 숫자보다 크거나 같은 가장 작은 정수로 올림하려면 `math.ceil()`를 사용할 수 있습니다.

# --hints--

인수가 하나인 `calculate_all_coordinates` 메서드를 정의해야 합니다, `self`.

```js
({test: () => assert(runPython(`
c = _Node(_code).find_class('Projectile')
c.has_function('calculate_all_coordinates') and c.find_function('calculate_all_coordinates').has_args('self')
`))})
```

`calculate_all_coordinates`에서 `ball` 메서드를 호출하고, 결과를 `coordinates`이라는 변수에 할당해야 합니다.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('coordinates'), "coordinates variable is missing"

        assert _Node(_code).find_variable('coordinates').is_equivalent('coordinates = ball.calculate_all_coordinates()'), "wrong value assigned to coordinates variable"
    `)
})
```

`calculate_all_coordinates` 메서드는 모든 `(x, y)` 좌표를 포함하는 리스트를 반환해야 합니다.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
ball = Projectile(12, 13, 14)
expected = [(x, ball._Projectile__calculate_y_coordinate(x)) for x in range(math.ceil(ball._Projectile__calculate_displacement()))]
actual = ball.calculate_all_coordinates()
expected == actual
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x**2 / (
                2 * self.__speed**2 * math.cos(self.__angle)**2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate

    --fcc-editable-region--
    
        
        
ball = Projectile(10, 3, 45)
print(ball)
   --fcc-editable-region--
```
