---
id: 6633b26c4f2d603072ef0818
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

클래스에 `__repr__` 특수 메서드를 사용해 표현을 제공하는 것이 좋은 습관입니다. `__str__` 메서드는 사용자 친화적인 읽기 쉬운 문자열 표현을 반환하는 반면, `__repr__`는 프로그래머를 위한 것입니다. 종종 `__repr__`는 객체를 다시 생성하는 데 사용할 수 있는 문자열을 제공합니다.

객체를 인스턴스화하는 데 필요한 문자열을 반환하는 `__repr__` 메서드를 작성하세요.

클래스의 인스턴스를 전달하여 `repr` 함수를 호출해 테스트할 수 있습니다.

# --hints--

`Projectile` 클래스는 `__repr__` 메서드를 가져야 합니다.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Projectile').has_function('__repr__')`
    ))
})
```

인스턴스 `Projectile(12, 12, 12)`의 표현은 `'Projectile(12, 12, 12)'`와 같아야 합니다.

```js
({
    test: () => assert(runPython(
        `repr(Projectile(12, 12, 12)) == 'Projectile(12, 12, 12)' or \\
        repr(Projectile(12, 12, 12)) == 'Projectile(12, 12, 12.0)'`
    ))
})
```

인스턴스 `Projectile(21, 22, 45)`의 표현은 `'Projectile(21, 22, 45)'`와 같아야 합니다.

```js
({
    test: () => assert(runPython(
        `repr(Projectile(21, 22, 45)) == 'Projectile(21, 22, 45)' or \\
        repr(Projectile(21, 22, 45)) == 'Projectile(21, 22, 45.0)'`
    ))
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.speed} m/s
height: {self.height} m
angle: {self.angle}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed
    
    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()

   
```
