---
id: 66335223f3e2a31d62d84367
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

투사체 궤적 계산기로 돌아와서, 이제 좌표 목록을 받아 궤적 그림을 생성하는 새 클래스를 만드세요.

새 클래스 `Graph`를 만들고, 좌표 목록이 저장되는 비공개 속성 `__coordinates`으로 인스턴스화하세요. `__slots__` 클래스 변수를 선언하고 `__init__` 및 `__repr__` 특수 메서드를 정의하세요.

또한 `Graph` 인스턴스를 `coordinates` 변수를 전달하여 생성하고, 이를 새 `graph` 변수에 할당하세요.

# --hints--

새 `Graph` 클래스를 만들어야 합니다.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_class('Graph')`
    ))
})
```

클래스 내부에 `__slots__` 변수가 있어야 합니다.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_variable('__slots__')`
    ))
})
```

`__slots__` 변수는 `('__coordinates')` 또는 `['__coordinates']` 값을 가져야 합니다.

```js
({
    test: () => (runPython(
        `
        slots = _Node(_code).find_class('Graph').find_variable('__slots__')

        assert any(slots.is_equivalent(str) for str in ['__slots__ = ("__coordinates")', '__slots__ = ["__coordinates"]'])`
    ))
})
```

`__init__` 메서드가 있어야 합니다.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('__init__')`
    ))
})
```

`__init__` 메서드는 전달된 값으로 `__coordinates` 비공개 속성을 인스턴스화해야 합니다.

```js
({
    test: () => assert(runPython(
        `
        l = [(3, 4), (4, 1), (3, 9)]
        g = Graph(l)
        g._Graph__coordinates == l
        `
    ))
})
```

`__repr__` 메서드가 있어야 합니다.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('__repr__')`
    ))
})
```

`__repr__` 메서드는 올바른 문자열을 반환해야 합니다.

```js
({
    test: () => assert(runPython(
        `
        ball = Projectile(10, 3, 45)
        coords = ball.calculate_all_coordinates()
        repr(Graph(coords)) == f'Graph({coords})'
        `
    ))
})
```

`Graph` 인스턴스를 `coordinates` 변수를 전달하여 생성하고 `graph` 변수에 할당해야 합니다.

```js
({
    test: () => {
        runPython(`

        assert _Node(_code).has_variable('graph'), "variable graph is missing"

        assert _Node(_code).find_variable('graph').is_equivalent('graph = Graph(coordinates)'), "assignment to graph is wrong"
        `)
    }
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.speed} m/s
height: {self.height} m
angle: {self.angle}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed

    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s

    def __repr__(self):
        return f"{self.__class__.__name__}({self.speed}, {self.height}, {self.angle})"

--fcc-editable-region--


ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()


--fcc-editable-region-- 
```
