---
id: 5d5a813321b9e3db6c106a46
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript はインタラクティブなウェブサイトを作るのに効果的な言語です。

*注*: 本カリキュラムのこれ以降のプロジェクトでは、HTML と CSS の基礎的な知識が必要です。 HTML と CSS に不慣れな場合は、[レスポンシブウェブデザイン認定講座](https://www.freecodecamp.org/learn/2022/responsive-web-design/)の学習をおすすめします。

まず、標準的な HTML のボイラープレートを作成しましょう。`DOCTYPE`、`html`、`head`、`body`、そして `charset` 用の `meta` タグを追加してください。 さらに `title` 要素を追加して、テキストには `RPG - Dragon Repeller` を使用してください。 `link` タグを追加して、スタイルシートとして `styles.css` ファイルをリンクしてください。

最後に、`body` 内に `div` 要素を作成して、`id` を `game` に設定してください。

# --hints--

コードに `DOCTYPE` 参照が含まれている必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプが `html` であると定義する必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

タイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 要素には開始タグが必要です。 `lang` 属性を忘れないようにしましょう。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 宣言は HTML の先頭に置く必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`head` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head\s*>/i));
```

`head` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

`body` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body\s*>/i));
```

`body` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 要素と `body` 要素は兄弟要素である必要があります。

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

`head` 要素は `html` 要素の中に入れる必要があります。

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

`body` 要素は `html` 要素の中に入れる必要があります。

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

コードには `meta` 要素が 1 つ必要です。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

`meta` 要素には値が `UTF-8` に設定されている `charset` 属性が必要です。

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/title\s*>/i));
```

コードには `link` 要素が 1 つ必要です。

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

`div` 要素が 1 つ必要です。

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div` 要素には値が `game` に設定されている `id` 属性が必要です。

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

`div` 要素は `body` 要素の中に入れる必要があります。

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
