---
id: 62a1166ed9a56d439c0770e7
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`#stats` 要素内に `span` 要素を 3 つ作成してください。 作成した各要素のクラスを `stat` に設定してください。 次に、作成した要素の 1 つ目にテキスト `XP: 0`、2 つ目にテキスト `Health: 100`、3 つ目にテキスト `Gold: 50` を設定してください。

# --hints--

`#stats` 要素内に 3 つの `span` 要素が必要です。

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

新しい 3 つの `span` 要素のクラスを `stat` を設定する必要があります。

```js
assert.exists(document.querySelectorAll('#stats > .stat')?.length, 3);
```

1 つ目の `.stat` 要素は `XP: 0` というテキストを持つ必要があります。

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

2 つ目の `.stat` 要素は `Health: 100` というテキストを持つ必要があります。

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

3 つ目の `.stat` 要素は `Gold: 50` というテキストを持つ必要があります。

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
