---
id: 62a115879a6d51422652cbfc
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

あなたのロールプレイングゲームでは、ユーザーが自分のステータスを追跡し、武器を購入し、モンスターと戦うことができます。 インタラクティブな JavaScript 部分のコードに進む前に、ゲーム情報を表示するための HTML 要素を最初に作成する必要があります。

`#game` の要素内に、`div` 要素を 4 つ作成してください。 それぞれに次の値の `id` を順番に付与してください: `stats`、`controls`、`monsterStats`、`text`

# --hints--

新しい `div` 要素を 4 つ作成する必要があります。

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

新しい `div` 要素の 1 つに `id` として `stats` を付与してください。

```js
assert.exists(document.querySelector('div#stats'));
```

新しい `div` 要素の 1 つに `id` として `controls` を付与してください。

```js
assert.exists(document.querySelector('div#controls'));
```

新しい `div` 要素の 1 つに `id` として `monsterStats` を付与してください。

```js
assert.exists(document.querySelector('div#monsterStats'));
```

新しい `div` 要素の 1 つに `id` として `text` を付与してください。

```js
assert.exists(document.querySelector('div#text'));
```

新しい `div` 要素を正しい順番で配置する必要があります。

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

新しい `div` 要素は `#game` の要素内に入れる必要があります。

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
