---
id: 62a2509ba163e020bb9d84ea
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

ここから JavaScript を書き始めます。 まずは `script` 要素から作成しましょう。 この要素を使って JavaScript を HTML ファイルに読み込ませます。

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

`script` 要素が 1 つ必要です。

```js
assert.isAtLeast(document.querySelectorAll('script').length, 2);
```

`script` 要素には開始タグが必要です。

```js
assert.match(code, /<script\s*>/i);
```

`script` 要素には終了タグが必要です。

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
