---
id: 62a257659d0d1e2456f24ba2
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

プロジェクトのコードを書き始める前に、JavaScript を個別のファイルに移して整理しましょう。

`console.log("Hello World");` という行を削除してください。 次に、空になった `script` 要素に `src` 属性を追加して、値を `./script.js` に設定してください。

# --hints--

コード内に `console.log("Hello World");` という行があってはなりません。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

`script` 要素の `src` 属性を `./script.js` に設定する必要があります。

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
