---
id: 62a23cb9bc467a147516b500
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

`#controls` 要素内に、`button` 要素を 3 つ作成してください。 1 つ目の `id` を `button1` に設定し、テキストは `Go to store` にしてください。 2 つ目の `id` を `button2` に設定し、テキストは `Go to cave` にしてください。 3 つ目の `id` を `button3` に設定し、テキストは `Fight dragon` にしてください。

# --hints--

`#controls` 要素に `button` 要素を 3 つ追加する必要があります。

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

1 つ目のボタンの `id` を `button1` に設定する必要があります。

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

1 つ目のボタンにはテキスト `Go to store` が必要です。

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

2 つ目のボタンの `id` を `button2` に設定する必要があります。

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

2 つ目のボタンにはテキスト `Go to cave` が必要です。

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

3 つ目のボタンの `id` を `button3` に設定する必要があります。

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

3 つ目のボタンにはテキスト `Fight dragon` が必要です。

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
--fcc-editable-region--
      <div id="controls">

      </div>
--fcc-editable-region--
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
