---
id: 62a255dae245b52317da824a
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

開発者コンソールは最も強力なツールのひとつです。 ブラウザーによりますが、 これは `F12` または `Ctrl+Shift+I` を押すと開くことができます。 On Mac, you can press `Option + ⌘ + C` and select "Console". または、プレビューウィンドウの上にある「Console」ボタンをクリックすることで組み込みのコンソールを表示できます。

The developer console will include errors that are produced by your code, but you can also use it to see values of variables in your code, which is helpful for debugging.

Add a `console.log("Hello World");` line between your `script` tags. Then click the "Console" button to open the console. You should see the text `"Hello World"`.


# --hints--

You should add a `console.log("Hello World");` line to your code.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

`console.log("Hello World");` という行を `script` タグの間に置く必要があります。

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
