---
id: 62a255dae245b52317da824a
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

開発者コンソールは最も強力なツールのひとつです。 ブラウザーによりますが、開発者ツールは `F12` または `Ctrl+Shift+I` を押すと開くことができます。 Macでは `Option + ⌘ + C` を押し、「コンソール」を選択します。 またはこの freeCodeCamp カリキュラムでは、プレビューウィンドウの上にある「コンソール」ボタンをクリックすることで組み込みコンソールを表示できます。

開発者コンソールにはコードによって生み出されたエラーが表示されることもありますが、デバッグに役立てるために、コード内の変数の値を表示させることもできます。

では、`script` タグの間に `console.log("Hello World");` という行を追加してください。 その後、「コンソール」ボタンをクリックしてコンソールを開きます。 `"Hello World"` という文字列が表示されるはずです。


# --hints--

`console.log("Hello World");` という行をコードに追加する必要があります。

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

`console.log("Hello World");` という行は `script` タグの間に置く必要があります。

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
