---
id: 62a23c1d505bfa13747c8a9b
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

数値 `0`、`100`、`50` をそれぞれ `span` 要素で囲み、その新しい `span` 要素をさらに `strong` 要素で囲んでください。 次に、新しい `span` 要素の `id` の値をそれぞれ `xpText`、`healthText`、`goldText` に設定してください。

回答は、次の構造に従った形になるはずです:

```html
<span class="stat">TEXT <strong><span id="VALUE">TEXT</span></strong></span>
```

# --hints--

1 つ目の `.stat` の要素内に、`strong` 要素を 追加する必要があります。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

1 つ目の新しい `strong` 要素内に、`span` 要素を入れる必要があります。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

1 つ目の、ネストの内側に位置する `span` 要素には、`xpText` という `id` を持たせてください。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'xpText');
```

1 つ目の、ネストの内側に位置する `span` 要素は、`0` というテキストを囲む必要があります。

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '0');
```

1 つ目の `.stat` の要素のテキストは変わらずに `XP: 0` でなければなりません。

```js
const stat = document.querySelectorAll('.stat')[0];
assert.equal(stat.innerText.trimEnd(), 'XP: 0');
```

2 つ目の `.stat` の要素内に、`strong` 要素を 追加する必要があります。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

2 つ目の新しい `strong` 要素内に、`span` 要素を入れる必要があります。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

2 つ目の、ネストの内側に位置する `span` 要素には、`healthText` という `id` を持たせてください。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'healthText');
```

2 つ目の、ネストの内側に位置する `span` 要素は、`100` というテキストを囲む必要があります。

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '100');
```

2 つ目の `.stat` の要素のテキストは変わらずに `Health: 100` でなければなりません。

```js
const stat = document.querySelectorAll('.stat')[1];
assert.equal(stat.innerText.trimEnd(), 'Health: 100');
```

3 つ目の `.stat` の要素内に、`strong` 要素を 追加する必要があります。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

3 つ目の新しい `strong` 要素内に、`span` 要素を入れる必要があります。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

3 つ目の、ネストの内側に位置する `span` 要素には、`goldText` という `id` を持たせてください。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'goldText');
```

3 つ目の、ネストの内側に位置する `span` 要素は、`50` というテキストを囲む必要があります。

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '50');
```

3 つ目の `.stat` の要素のテキストは変わらずに `Gold: 50` でなければなりません。

```js
const stat = document.querySelectorAll('.stat')[2];
assert.equal(stat.innerText.trimEnd(), 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
--fcc-editable-region--
        <span class="stat">XP: 0</span>
        <span class="stat">Health: 100</span>
        <span class="stat">Gold: 50</span>
--fcc-editable-region--
      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
