---
id: 62a3b79d520a7f3d0e25afd6
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript は<dfn>ドキュメントオブジェクトモデル (Document Object Model)</dfn>、または DOM を使用して HTML とやりとりします。 DOM は HTML を表すオブジェクトのツリーです。 HTML ドキュメント全体を表す `document` オブジェクトを使用することで、HTML にアクセスできます。

HTML 内の特定の要素を見つけるための方法の 1 つは、`querySelector()` メソッドを使用することです。 `querySelector()` メソッドは CSS セレクターを引数として受け取り、そのセレクターと一致する最初の要素を返します。 例えば、HTML 内の `<h1>` 要素を見つけるには、次のように書きます。

```js
let h1 = document.querySelector("h1");
```

引数の `h1` は文字列であり、CSS セレクターとしてマッチすることに注意してください。

では変数 `button1` を作成し、`querySelector()` を使って `id` が `button1` の要素を見つけてこの変数に割り当ててください。 CSS の `id` セレクターは、先頭が `#` で始まることを思い出してください。

# --hints--

`let` を使って `button1` という変数を宣言してください。

```js
assert.match(code, /let\s+button1/);
```

`document.querySelector()` を使用してください。

```js
assert.match(code, /document\.querySelector/);
```

`"#button1"` セレクターを使用してください。

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
