---
id: 62a3a0a3c0a4b32915d26a6e
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

プロジェクトのコードを書き始めるために、まず `console.log("Hello World");` という行を削除してください。

以前のプロジェクトで、次のような `let` キーワードと変数の使い方を学びました。

```js
let camperbot;
```

また、次のようにして変数に値を与えて初期化する方法も学びました:

```js
let age = 32;
```

では、`let` キーワードを使用して `xp` という名前の変数を宣言し、数値 `0` の値を割り当ててください。

# --hints--

コード内に `console.log("Hello World");` という行があってはなりません。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

`let` キーワードを使用して変数を宣言してください。

```js
assert.match(code, /let/);
```

`xp` という名前の変数を宣言してください。

```js
assert.match(code, /let\s+xp/);
```

`xp` は `0` の値を持つ必要があります。

```js
assert.equal(xp, 0);
```

変数 `xp` を `0` で初期化してください。

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
