---
id: 62a39f5a5790eb27c1e5d4bf
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

表示が新しい `script.js` ファイルに切り替えられています。 上のタブを使用してファイルの切り替えができることを忘れないようにしましょう。

`console.log("Hello World");` という行をこのファイルに追加し、コンソール内にそれが表示されるのを確認してください。

# --hints--

コード内に `console.log("Hello World");` という行が必要です。

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
