---
id: 62a3a7e4f1060e2fc5ffb34b
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

もう 1 つ、`currentWeaponIndex` という名前の変数を作成し、その値を `0` にしてください。

# --hints--

`currentWeaponIndex` という変数の宣言には `let` を使用してください。

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

変数名には camelCase を使用してください。

```js
assert.match(code, /currentWeaponIndex/);
```

変数 `currentWeaponIndex` の値は `0` にしてください。

```js
assert.equal(currentWeaponIndex, 0);
```

変数を `0` で初期化してください。

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
