---
id: 62a3b41c9494f937560640ab
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

さらに 2 つ、`monsterHealth` と `inventory` という名前の変数を宣言してください。

変数 `inventory` には、値として文字列 `"stick"` を含む配列を代入してください。

以前のプロジェクトで、次のように配列を使用したことを思い出しましょう:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

変数 `monsterHealth` の宣言には `let` を使用してください。

```js
assert.match(code, /let\s+monsterHealth/i);
```

変数 `monsterHealth` の命名にはキャメルケースを使用してください。

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` には値を持たせないでください。

```js
assert.isUndefined(monsterHealth);
```

変数 `monsterHealth` に値を代入してはいけません。

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

変数 `inventory` の宣言には `let` を使用してください。

```js
assert.match(code, /let\s+inventory/i);
```

変数 `inventory` は配列にしてください。

```js
assert.isArray(inventory);
```

変数 `inventory` は値を 1 つだけ持つ必要があります。

```js
assert.lengthOf(inventory, 1);
```

変数 `inventory` には文字列 `"stick"` が含まれている必要があります。

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
