---
id: 62a3c4a0e52767482c5202d4
title: ステップ 37
challengeType: 0
dashedName: step-37
---

# --description--

同じ構文を使用して、`button2` と `button3` の `onclick` プロパティをそれぞれ `goCave` および `fightDragon` に設定してください。

完了したらコンソールを開いて、ボタンをクリックしてみてください。

# --hints--

関数 `goStore` への参照を `button1.onclick` に代入してください。 関数を呼び出さないようにしてください。

```js
assert.match(code, /button1\.onclick\s*=\s*goStore\s*;?\s*$/m);
```

ドット記法を使用して、`button2` の `onclick` プロパティにアクセスしてください。

```js
assert.match(code, /button2\.onclick/);
```

`button2.onclick` に代入する際に `let` や `const` は使用しないでください。

```js
assert.notMatch(code, /(let|const)\s+button2\.onclick/);
```

関数 `goCave` への参照を `button2.onclick` に代入してください。 関数を呼び出さないようにしてください。

```js
assert.match(code, /button2\.onclick\s*=\s*goCave\s*;?\s*$/m);
```

ドット記法を使用して、`button3` の `onclick` プロパティにアクセスしてください。

```js
assert.match(code, /button3\.onclick/);
```

`button3.onclick` に代入する際に `let` や `const` は使用しないでください。

```js
assert.notMatch(code, /(let|const)\s+button3\.onclick/);
```

関数 `fightDragon` への参照を `button3.onclick` に代入してください。 関数を呼び出さないようにしてください。

```js
assert.match(code, /button3\.onclick\s*=\s*fightDragon\s*;?\s*$/m);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");

--fcc-editable-region--
// initialize buttons
button1.onclick = goStore;

--fcc-editable-region--
function goStore() {
  console.log("Going to store.");
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
