---
id: 62a8c2bbbd8aa82052f47c53
title: ステップ 71
challengeType: 0
dashedName: step-71
---

# --description--

配列 `locations` に 3 つ目のオブジェクトを追加しましょう。 他のオブジェクトと同じプロパティを持たせます。

`name` を `cave` に設定してください。 `button text` は、文字列 `"Fight slime"`、`"Fight fanged beast"`、`"Go to town square"` の配列に設定してください。 `"button functions"` は、変数 `fightSlime`、`fightBeast`、`goTown` の配列に設定してください。 `text` プロパティは `"You enter the cave. You see some monsters."` に設定してください。

# --hints--

配列 `locations` に 3 つの値を持たせる必要があります。

```js
assert.lengthOf(locations, 3);
```

`locations` の 3 番目の値はオブジェクトでなければなりません。

```js
assert.isObject(locations[2]);
```

`locations` の 3 番目のオブジェクトは `name` プロパティを持ち、その値は `"cave"` でなければなりません。

```js
assert.equal(locations[2].name, "cave");
```

`locations` の 3 番目のオブジェクトには、配列である `"button text"` プロパティが必要です。

```js
assert.isArray(locations[2]["button text"]);
```

`"button text"` プロパティには、値として文字列の `"Fight slime"`、`"Fight fanged beast"`、`"Go to town square"` が必要です。

```js
assert.equal(locations[2]["button text"][0], "Fight slime");
assert.equal(locations[2]["button text"][1], "Fight fanged beast");
assert.equal(locations[2]["button text"][2], "Go to town square");
```

`locations` の 3 番目のオブジェクトには、配列である `"button functions"` プロパティが必要です。

```js
assert.isArray(locations[2]["button functions"]);
```

`"button functions"` プロパティには、値として関数 `fightSlime`、`fightBeast`、`goTown` が必要です。

```js
assert.equal(locations[2]["button functions"][0], fightSlime);
assert.equal(locations[2]["button functions"][1], fightBeast);
assert.equal(locations[2]["button functions"][2], goTown);
```

`locations` の 3 番目のオブジェクトには、文字列である `text` プロパティが必要です。

```js
assert.isString(locations[2].text);
```

`locations` の 3 番目のオブジェクトは `text` プロパティを持ち、その値は `"You enter the cave. You see some monsters."` でなければなりません。

```js
assert.equal(locations[2].text, "You enter the cave. You see some monsters.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
--fcc-editable-region--

--fcc-editable-region--
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}

function buyHealth() {

}

function buyWeapon() {

}

function fightSlime() {

}

function fightBeast() {

}
```
