---
id: 62a8c7322e42962ad53ad204
title: ステップ 81
challengeType: 0
dashedName: step-81
---

# --description--

配列 `locations` と同様に、配列 `weapons` にもオブジェクトを格納します。 配列 `weapons` に 4 つのオブジェクトを追加してください。それぞれのオブジェクトには `name` と `power` という 2 つのプロパティを持たせてください。 最初のオブジェクトは、`name` を `"stick"` に、`power` を `5` に設定してください。 2 番目のオブジェクトは `"dagger"` と `30` にしましょう。 3 番目は `"claw hammer"` と `50` にしましょう。 4 番目は `"sword"` と `100` にしましょう。

# --hints--

配列 `weapons` には 4 つの要素が必要です。

```js
assert.lengthOf(weapons, 4);
```

配列 `weapons` には 4 つのオブジェクトが必要です。

```js
assert.isObject(weapons[0]);
assert.isObject(weapons[1]);
assert.isObject(weapons[2]);
assert.isObject(weapons[3]);
```

`weapons` の最初のオブジェクトは、`name` を `"stick"` に、`power` を `5` に設定してください。

```js
assert.equal(weapons[0].name, 'stick');
assert.equal(weapons[0].power, 5);
```

`weapons` の 2 番目のオブジェクトは、`name` を `"dagger"` に、`power` を `30` に設定してください。

```js
assert.equal(weapons[1].name, 'dagger');
assert.equal(weapons[1].power, 30);
```

`weapons` の 3 番目のオブジェクトは、`name` を `"claw hammer"` に、`power` を `50` に設定してください。

```js
assert.equal(weapons[2].name, 'claw hammer');
assert.equal(weapons[2].power, 50);
```

`weapons` の 4 番目のオブジェクトは、`name` を `"sword"` に、`power` を `100` に設定してください。

```js
assert.equal(weapons[3].name, 'sword');
assert.equal(weapons[3].power, 100);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
--fcc-editable-region--
const weapons = [];
--fcc-editable-region--
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
  {
    name: "cave",
    "button text": ["Fight slime", "Fight fanged beast", "Go to town square"],
    "button functions": [fightSlime, fightBeast, goTown],
    text: "You enter the cave. You see some monsters."
  }
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  update(locations[2]);
}

function fightDragon() {
  console.log("Fighting dragon.");
}

function buyHealth() {
  if (gold >= 10) {
    gold -= 10;
    health += 10;
    goldText.innerText = gold;
    healthText.innerText = health;
  } else {
    text.innerText = "You do not have enough gold to buy health.";
  }
}

function buyWeapon() {

}

function fightSlime() {

}

function fightBeast() {

}
```
