---
id: 67fe626cc464f1238514cd8b
title: 第 20 步
challengeType: 0
dashedName: step-20
---

# --description--

在 `filteredItems.map()` 呼叫中 `return` 述語之前，創建一個名為 `isChecked` 的變數，並使用 `includes()` 方法檢查目前的 `item` 是否在 `selectedItems` 陣列中。

然後，為 `li` 元素添加一個 `style` 屬性，並使用三元運算子將 `textDecoration` 屬性設為 `line-through`（如果 `isChecked` 為 true），或設為 `none`（如果為 false）。請記得你需要用大括號包住三元運算子，因為它是一個 JavaScript 表達式。

例如，這裡示範如何使用三元運算子在一個簡單的組件中有條件地設定段落的顏色：

```jsx
const MyComponent = () => {
  const isActive = true;

  return (
    <p style={{ color: isActive ? "green" : "black"}}>This text is conditionally styled.</p>
  );
};
```

# --hints--

你應該創建一個名為 `isChecked` 的變數，該變數使用 `selectedItems` 陣列上的 `includes()` 方法來檢查目前的 `item` 是否在陣列中。

```js
const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
const exports = {};
const a = eval(script);
const shoppingListString = exports.ShoppingList.toString();

assert.match(shoppingListString, /var\s*isChecked\s*=\s*selectedItems\.includes\s*\(\s*item\s*\)/);
```

你應該為你的 `li` 元素添加一個 `style` 屬性。

```js
  const testElem = await __helpers.prepTestComponent(window.index.ShoppingList);
  const li = testElem.querySelector('li');
  const key = li && Object.keys(li)?.find(key => key.startsWith("__reactProps"));
  const reactLi = li && li[key];

  assert.exists(reactLi.style);
```

你的 `style` 屬性應該包含一個三元運算子，當 `isChecked` 為 true 時將 `textDecoration` 屬性設為 `line-through`，當為 false 時設為 `none`。

```js
const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
const exports = {};
const a = eval(script);
const shoppingListString = exports.ShoppingList.toString();

assert.match(shoppingListString, /"li"\s*,\s*{(.|\n)*style:\s*{\s*textDecoration\s*:\s*isChecked\s*\?\s*("|')line-through("|')\s*:\s*("|')none("|')/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Shopping List</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main id="app-container"></main>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ShoppingList } from "./index.jsx";
      const appContainer = document.getElementById("app-container");
      const root = ReactDOM.createRoot(appContainer);
      root.render(<ShoppingList />);
    </script>
  </body>
</html>

```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  box-sizing: border-box;
}

body {
  font-family: Arial, sans-serif;
  line-height: 1.5;
  color: var(--dark-grey);
  background-color: var(--dark-grey);
}

main,
.container {
  display: flex;
  flex-direction: column;
  align-items: center;
}

.container {
  background-color: var(--light-grey);
  width: 90%;
  margin: 20px;
  padding: 10px;
}

h1 {
  color: var(--dark-grey);
}

form {
  text-align: center;
}

button {
  cursor: pointer;
}

button {
  cursor: pointer;
  width: 100px;
  margin: 3px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

input {
  color: var(--dark-grey);
  margin-left: 5px;
  padding: 3px;
}

li {
  text-align: left;
  margin: 10px 0;
}

.selected-item {
  font-weight: bold;
}

@media (min-width: 425px) {
  .container {
    width: 400px;
  }
}

```

```jsx
const { useState } = React;

const items = [
  "Apples",
  "Bananas",
  "Strawberries",
  "Blueberries",
  "Mangoes",
  "Pineapple",
  "Lettuce",
  "Broccoli",
  "Paper Towels",
  "Dish Soap",
];

export const ShoppingList = () => {
  const [query, setQuery] = useState("");
  const [selectedItems, setSelectedItems] = useState([]);

  const filteredItems = items.filter((item) => 
    item.toLowerCase().includes(query.toLowerCase())
  );

  const toggleItem = (item) => {
    setSelectedItems((prev) =>
      prev.includes(item) ? prev.filter((i) => i !== item) : [...prev, item]
    );
  };

  return (
    <div className="container">
      <h1>Shopping List</h1>
      <form>
        <label htmlFor="search">Search for an item:</label>
        <input
          id="search"
          type="search"
          placeholder="Search..."
          aria-describedby="search-description"
          value={query}
          onChange={(e) => setQuery(e.target.value)}
        /> 
        <p id="search-description">Type to filter the list below:</p>
        <ul>
          {filteredItems.map((item) => {
            --fcc-editable-region--

            return (
              <li
                key={item}
              >
                <label>
                  <input
                    type="checkbox"
                    onChange={() => toggleItem(item)}
                  />
                  {item}
                </label>
              </li>
            );
            --fcc-editable-region--
          })}
        </ul>
      </form>
    </div>
  );
};

```
