---
id: 587d78b1367417b2b2512b0a
title: Tumia picha ya retina kwa azimio la juu zaidi
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cVZ4Rfp'
forumTopicId: 301142
dashedName: use-a-retina-image-for-higher-resolution-displays
---

# --description--

Kwa kuongezeka kwa vifaa vinavyounganishwa na intaneti, ukubwa na sifa zao hutofautiana, na skrini wanazotumia zinaweza kuwa tofauti nje na ndani. Msongamano wa pixel ni kipengele ambacho kinaweza kutofautiana kati ya kifaa kimoja na kingine na msongamano huu unajulikana kama Pixel Per Inch (PPI) au Dots Per Inch (DPI). Skrini maarufu zaidi ya aina hii ni ile inayojulikana kama "Retina Display" kwenye kompyuta za Apple MacBook Pro za kisasa, na hivi karibuni kwenye kompyuta za iMac. Kutokana na tofauti ya msongamano wa pixel kati ya skrini za "Retina" na "Non-Retina", baadhi ya picha ambazo hazijatengenezwa kwa kuzingatia skrini za Azimio la Juu zinaweza kuonekana "zikiwa na pixel" wakati zinaonyeshwa kwenye skrini za Azimio la Juu.

Njia rahisi zaidi ya kufanya picha zako zionekane ipasavyo kwenye skrini za Azimio la Juu, kama vile skrini za "retina display" za MacBook Pro, ni kuweka thamani za `width` na `height` kuwa nusu tu ya ile ya faili asili. Hapa kuna mfano wa picha inayotumia nusu tu ya urefu na upana wa asili:

```html
<style>
  img {
    height: 250px;
    width: 250px;
  }
</style>
<img src="coolPic500x500" alt="A most excellent picture" />
```

# --instructions--

Weka `width` na `height` za lebo ya `img` kuwa nusu ya thamani zao za asili. Katika kesi hii, `height` na `width` za asili zote ni `200px`.

# --hints--

Lebo yako ya `img` inapaswa kuwa na `width` ya pikseli 100.

```js
assert.strictEqual(document.querySelector('img').width, 100);
```

Lebo yako ya `img` inapaswa kuwa na `height` ya pikseli 100.

```js
assert.strictEqual(document.querySelector('img').height, 100);
```

# --seed--

## --seed-contents--

```html
<style></style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```

# --solutions--

```html
<style>
  img {
    height: 100px;
    width: 100px;
  }
</style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```
