---
id: 587d78b1367417b2b2512b0c
title: Fanya Typography Iwe Inayojibadilisha Kulingana na Kifaa
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/crzN7T8'
forumTopicId: 301141
dashedName: make-typography-responsive
---

# --description--

Badala ya kutumia `em` au `px` kupima ukubwa wa maandishi, unaweza kutumia vitengo vya eneo la kuona kwa typography inayojibadilisha kulingana na kifaa. Vitengo vya eneo la kuona, kama asilimia, ni vitengo vinavyohusiana, lakini vinategemea vitu tofauti. Vitengo vya eneo la kuona vinahusiana na vipimo vya eneo la kuona (upana au urefu) wa kifaa, na asilimia zinahusiana na ukubwa wa kipengele cha kontena cha mzazi.

Vitengo vinne tofauti vya eneo la kuona ni:

<ul><li><code>vw</code> (upana wa eneo la kuona): <code>10vw</code> itakuwa asilimia 10 ya upana wa eneo la kuona.</li><li><code>vh</code> (urefu wa eneo la kuona): <code>3vh</code> itakuwa asilimia 3 ya urefu wa eneo la kuona.</li><li><code>vmin</code> (kiwango cha chini cha eneo la kuona): <code>70vmin</code> itakuwa asilimia 70 ya kipimo kidogo cha eneo la kuona (urefu au upana).</li><li><code>vmax</code> (kiwango cha juu cha eneo la kuona): <code>100vmax</code> itakuwa asilimia 100 ya kipimo kikubwa cha eneo la kuona (urefu au upana).</li></ul>

Hapa kuna mfano unaoweka lebo ya `body` kuwa asilimia 30 ya upana wa eneo la kuona.

```css
body {
  width: 30vw;
}
```

# --instructions--

Weka `width` ya lebo ya `h2` kuwa asilimia 80 ya upana wa eneo la kuona na `width` ya aya kuwa asilimia 75 ya kipimo kidogo cha eneo la kuona.

# --hints--

Leboa yako ya `h2` inapaswa kuwa na `width` ya 80vw.

```js
assert.match(
  __helpers.removeCssComments(code),
  /h2\s*?{\s*?width:\s*?80vw;\s*?}/g
);
```

Leboa yako ya `p` inapaswa kuwa na `width` ya 75vmin.

```js
assert.match(
  __helpers.removeCssComments(code),
  /p\s*?{\s*?width:\s*?75vmin;\s*?}/g
);
```

# --seed--

## --seed-contents--

```html
<style>

</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  h2 {
    width: 80vw;
  }
  p {
    width: 75vmin;
  }
</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
