---
id: 587d78b1367417b2b2512b09
title: Fanya picha i-jibadilishe kulingana na kifaa
challengeType: 0
forumTopicId: 301140
dashedName: make-an-image-responsive
---

# --description--

Kufanya picha zi-jibadilishe kulingana na kifaa kwa kutumia CSS ni rahisi sana. Unahitaji tu kuongeza vigezo hivi kwenye picha:

```css
img {
  max-width: 100%;
  height: auto;
}
```

`max-width` ya `100%` itahakikisha picha haizidi upana wa kontena iliyo ndani yake, na `height` ya `auto` itafanya picha ihifadhi uwiano wake wa asili.

# --instructions--

Ongeza sheria za mtindo kwenye darasa la `responsive-img` ili lifanye picha i-jibadilishe kulingana na kifaa. Haipaswi kamwe kuwa pana zaidi ya kontena lake (katika kesi hii, ni dirisha la onyesho) na inapaswa kuhifadhi uwiano wake wa asili. Baada ya kuongeza msimbo wako, badilisha ukubwa wa dirisha la onyesho kuona jinsi picha zako zinavyo-jibadilisha.

# --hints--

Darasa lako la `responsive-img` linapaswa kuwa na `max-width` iliyowekwa kuwa `100%`.

```js
assert.strictEqual(getComputedStyle(document.querySelector('.responsive-img')).maxWidth, '100%');
```

Darasa lako la `responsive-img` linapaswa kuwa na `height` iliyowekwa kuwa `auto`.

```js
assert.match(code, /height:\s*?auto;/g);
```

# --seed--

## --seed-contents--

```html
<style>
.responsive-img {


}

img {
  width: 600px;
}
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```

# --solutions--

```html
<style>
  .responsive-img {
    max-width: 100%;
    height: auto;
  }

  img {
    width: 600px;
  }
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```
