---
id: 587d78b0367417b2b2512b08
title: Unda media query
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cqwKrtm'
forumTopicId: 301139
dashedName: create-a-media-query
---

# --description--

Media Queries ni mbinu mpya iliyotangazwa katika CSS3 inayobadilisha muonekano wa maudhui kulingana na ukubwa tofauti wa eneo la kuona. Eneo la kuona ni eneo linaloonekana na mtumizi wa ukurasa wa mtandao, na hubadilika kulingana na kifaa kinachotumika kufikia tovuti.

Media Queries zinajumuisha aina ya media, na ikiwa aina hiyo ya media inalingana na aina ya kifaa ambacho hati inaonyeshwa, mitindo hutumika. Unaweza kuwa na vichaguzi na mitindo mingi ndani ya media query yako kama unavyotaka.

Hapa kuna mfano wa media query inayorejesha maudhui wakati upana wa kifaa ni chini au sawa na `100px`:

```css
@media (max-width: 100px) {
  /* CSS Rules */
}
```

na media query ifuatayo inarejesha maudhui wakati urefu wa kifaa ni zaidi au sawa na `350px`:

```css
@media (min-height: 350px) {
  /* CSS Rules */
}
```

Kumbuka, CSS ndani ya media query hutumika tu ikiwa aina ya media inalingana na ile ya kifaa kinachotumika.

# --instructions--

Ongeza media query, ili lebo ya `p` iwe na `font-size` ya `10px` wakati urefu wa kifaa ni chini au sawa na `800px`.

# --hints--

Unapaswa kutangaza hoja ya utafutaji ya `@media` kwa vifaa vyenye `height` chini au sawa na `800px`.

```js
const media = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(
  media.some(x => x.media?.mediaText?.includes('(max-height: 800px)'))
);
```

Kipengele chako cha `p` kinapaswa kuwa na `font-size` ya `10px` wakati `height` ya kifaa ni chini au sawa na `800px`.

```js
const rules = new __helpers.CSSHelp(document).getRuleListsWithinMedia(
  '(max-height: 800px)'
);
assert.strictEqual(
  rules?.find(x => x.selectorText === 'p')?.style?.fontSize,
  '10px'
);
```

Kipengele chako cha `p` kinapaswa kuwa na `font-size` ya awali ya `20px` wakati `height` ya kifaa ni zaidi ya `800px`.

```js
const ifPFirst = new __helpers.CSSHelp(document)
  .getCSSRules()
  ?.find(x => x?.selectorText === 'p' || x?.media);
assert.strictEqual(ifPFirst?.style?.fontSize, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  p {
    font-size: 20px;
  }

  /* Only change code below this line */

  /* Only change code above this line */
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  p {
    font-size: 20px;
  }

  @media (max-height: 800px) {
    p {
      font-size: 10px;
    }
  }
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
