---
id: 587d8249367417b2b2512c3e
title: 使用 helmet.noCache() 停用客戶端快取
challengeType: 2
forumTopicId: 301576
dashedName: disable-client-side-caching-with-helmet-nocache
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

如果你正在釋出網站的更新，並且你希望使用者總是下載較新的版本，你可以嘗試停用客戶端瀏覽器的快取。這在開發時也很有用。快取具有效率利益，但你會失去這些利益，因此只有在真正需要時才使用此選項。

# --instructions--

在你的伺服端使用 `helmet.noCache()` 方法。

# --hints--

應正確掛載 helmet.noCache() 中介層

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nocache');
assert.equal(
  data.headers['cache-control'],
  'no-store, no-cache, must-revalidate, proxy-revalidate'
);
```

