---
id: 587d8248367417b2b2512c3a
title: 避免使用 helmet.noSniff() 推斷回應的 MIME 型別
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

瀏覽器可以使用內容或 MIME 嗅探來改寫回應的 `Content-Type` 標頭，以隱喻的內容型別猜測並處理資料。雖然這在某些情況下很方便，但也可能導致一些危險的攻擊。此中介層會將 `X-Content-Type-Options` 標頭設為 `nosniff`，指示瀏覽器不要繞過所提供的 `Content-Type`。

# --instructions--

在你的伺服器上使用 `helmet.noSniff()` 方法。

# --hints--

應正確掛載 helmet.noSniff() 中介層

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```

