---
id: 587d8247367417b2b2512c39
title: >-
  Mitigate the Risk of Cross Site Scripting (XSS) Attacks with
  helmet.xssFilter()
challengeType: 2
forumTopicId: 301583
dashedName: mitigate-the-risk-of-cross-site-scripting-xss-attacks-with-helmet-xssfilter
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

跨站腳本攻擊（XSS）是一種常見的攻擊型別，攻擊者會將惡意腳本注入易受攻擊的頁面中，目的是竊取像是會話 Cookie 或密碼等敏感資料。

降低 XSS 攻擊風險的基本規則很簡單：「永遠不要相信使用者的輸入」。作為開發者，你應該始終對所有來自外部的輸入進行淨化。這包括來自表單、GET 查詢網址，甚至是 POST 主體的資料。淨化意味著你應該尋找並編碼可能危險的字元，例如 &lt;、>。

現代瀏覽器可以透過採用更好的軟體策略來降低風險。這些策略通常可以透過 http 標頭進行設定。

X-XSS-Protection HTTP 標頭是一種基本的保護。瀏覽器使用啟發式過濾器偵測潛在的注入腳本。如果啟用該標頭，瀏覽器會更改腳本程式碼，使其失效。它仍然只有有限的支援。

# --instructions--

使用 `helmet.xssFilter()` 來清理傳送到你的伺服器的輸入。

# --hints--

應正確掛載 helmet.xssFilter() 中介層

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'xXssProtection');
assert.property(data.headers, 'x-xss-protection');
```

