---
id: 587d8247367417b2b2512c37
title: 使用 helmet.hidePoweredBy() 隱藏潛在危險的資訊
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

駭客如果看到你的網站是由 Express 威力本源，可能會利用 Express/Node 中已知的漏洞。預設情況下，每個來自 Express 的請求都會帶有 `X-Powered-By: Express`。使用 `helmet.hidePoweredBy()` 中介層來移除 X-Powered-By 標頭。

# --hints--

應正確掛載 helmet.hidePoweredBy() 中介層

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```

