---
id: 587d8247367417b2b2512c36
title: 安裝並引入 Helmet
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

處理這些挑戰將涉及你使用以下其中一種方法來撰寫你的程式碼：

- 複製 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">這個 GitHub 倉庫</a> 並在區域端完成這些挑戰。
- 使用你選擇的網站建置工具來完成專案。務必將我們 GitHub 倉庫中的所有檔案納入其中。

Helmet 幫助你透過設定各種 HTTP 標頭來保護你的 Express 應用程式。

# --instructions--

你所有這些課程的程式碼都放在 `myApp.js` 檔案中，位於我們為你開始的程式碼行之間。請勿更改或刪除我們為你添加的程式碼。

Helmet 版本 `3.21.3` 已經安裝完成，所以在 `myApp.js` 中以 `helmet` 引用它。

# --hints--

`helmet` 版本 `3.21.3` 應該在 `package.json` 中。

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```

