---
id: 587d8247367417b2b2512c38
title: 使用 helmet.frameguard() 降低點擊劫持的風險
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

你的頁面可能會在未經你同意的情況下被放入 `<frame>` 或 `<iframe>`。這可能導致點擊劫持攻擊等問題。點擊劫持是一種誘騙使用者與與他們認為不同的頁面互動的技術。這可以透過在惡意背景關係中執行你的頁面（透過 iframing）來達成。在該背景關係中，駭客可以在你的頁面上放置隱藏層。隱藏的按鈕可以用來執行惡意腳本。這個中介層會設定 X-Frame-Options 標頭。它限制誰可以將你的網站放入影格。它有三種模式：DENY、SAMEORIGIN 和 ALLOW-FROM。

我們不需要讓我們的應用程式被框住。

# --instructions--

使用 `helmet.frameguard()` 並傳入設定物件 `{action: 'deny'}`。

# --hints--

應正確掛載 helmet.frameguard() 中介層

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

helmet.frameguard() 的 'action' 應該設為 'DENY'

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```

