---
id: 587d8248367417b2b2512c3d
title: 使用 helmet.dnsPrefetchControl() 停用 DNS 預取。
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

為了提升效率，大多數瀏覽器會預先擷取頁面中聯結的 DNS 記錄。如此一來，當使用者點擊聯結時，目的地 IP 已經知道。這可能導致 DNS 服務過度使用（如果你擁有一個由數百萬人造訪的大型網站……）、隱私問題（某個竊聽者可能推斷你正在瀏覽某個特定頁面），或頁面統計數據異動（某些聯結即使未被點擊也可能顯示為已造訪）。如果你有高度安全需求，可以停用 DNS 預擷取，但會以效率損失為代價。

# --instructions--

在你的伺服器上使用 `helmet.dnsPrefetchControl()` 方法。

# --hints--

應正確掛載 helmet.dnsPrefetchControl() 中介層

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```

