---
id: 587d8248367417b2b2512c3b
title: 防止 IE 使用 helmet.ieNoOpen() 開啟不受信任的 HTML
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

某些網頁應用程式會提供不受信任的 HTML 供下載。某些版本的 Internet Explorer 預設會在你的網站背景關係中開啟這些 HTML 檔案。這表示不受信任的 HTML 頁面可能會開始在你的頁面背景關係中執行不良行為。這個中介層會將 X-Download-Options 標頭設定為 noopen。這將防止 IE 使用者在受信任的網站背景關係中執行下載。

# --instructions--

在你的伺服器上使用 `helmet.ieNoOpen()` 方法。

# --hints--

應正確掛載 helmet.ieNoOpen() 中介層

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```

