---
id: 587d8249367417b2b2512c40
title: 使用「parent」`helmet()` 中介層組態 Helmet
challengeType: 2
forumTopicId: 301575
dashedName: configure-helmet-using-the-parent-helmet-middleware
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

`app.use(helmet())` 會自動包含上述所有介紹的中介層，除了 `noCache()` 和 `contentSecurityPolicy()`，但如果有需要可以啟用它們。你也可以使用組態物件個別停用或組態其他任何中介層。

**範例：**

```js
app.use(helmet({
  frameguard: {         // configure
    action: 'deny'
  },
  contentSecurityPolicy: {    // enable and configure
    directives: {
      defaultSrc: ["'self'"],
      styleSrc: ['style.com'],
    }
  },
  dnsPrefetchControl: false     // disable
}))
```

我們為教學目的及方便測試而分別介紹每個中介層。在真實專案中使用「父」`helmet()`中介層很容易實作。

# --hints--

沒有測試－這是一個描述性挑戰

```js
assert(true);
```

