---
id: 587d8248367417b2b2512c3c
title: 使用 helmet.hsts() 要求瀏覽器僅透過 HTTPS 存取你的網站
challengeType: 2
forumTopicId: 301573
dashedName: ask-browsers-to-access-your-site-via-https-only-with-helmet-hsts
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

HTTP 嚴格傳輸安全性（HSTS）是一種網路安全政策，有助於保護網站免受通訊協定降級攻擊和 Cookie 劫持。如果你的網站可以透過 HTTPS 存取，你可以要求使用者的瀏覽器避免使用不安全的 HTTP。透過設定標頭 `Strict-Transport-Security`，你告訴瀏覽器在指定的時間內對未來的請求使用 HTTPS。這將適用於初始請求之後的請求。

# --instructions--

組態 `helmet.hsts()` 以在接下來的 90 天內使用 HTTPS。傳遞組態物件 `{maxAge: timeInSeconds, force: true}`。你可以創建一個變數 `ninetyDaysInSeconds = 90*24*60*60;` 來用於 `timeInSeconds`。

注意：在訂製網站上設定 HTTPS 需要取得網域和 SSL/TLS 憑證。

# --hints--

應正確掛載 helmet.hsts() 中介層

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hsts');
assert.property(data.headers, 'strict-transport-security');
```

maxAge 應該相等於 7776000 秒（90 天）

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.match(
  data.headers['strict-transport-security'],
  /^max-age=7776000;?/
);
```

