---
id: 58a25bcff9fc0f352b528e7e
title: 同步雜湊並比較密碼
challengeType: 2
forumTopicId: 301579
dashedName: hash-and-compare-passwords-synchronously
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

同步進行雜湊同樣很簡單，但如果在伺服端使用高成本或頻繁進行雜湊，可能會導致延遲。使用此方法進行雜湊就像呼叫一樣簡單。

```js
var hash = bcrypt.hashSync(myPlaintextPassword, saltRounds);
```

將這種雜湊方法添加到你的程式碼中，然後將結果記錄到主控台。再次強調，所使用的變數已經在伺服端定義，因此你不需要調整它們。你可能會注意到，儘管你正在對與非同步函式中相同的密碼進行雜湊，但主控台中的結果不同－這是因為每次都會隨機產生 salt，如雜湊字串第三個字串的前 22 個字元所示。現在，若要將密碼輸入與新的同步雜湊進行比較，你會使用 compareSync 方法：

```js
var result = bcrypt.compareSync(myPlaintextPassword, hash);
```

結果為布林值 true 或 false。

# --instructions--

將函式添加到其中並將結果記錄到主控台以查看其運作。

當你認為已經正確時，提交你的頁面。

# --hints--

應該產生同步雜湊並正確比較。

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_SYNC[^]*hash.*=.*bcrypt.hashSync.*myPlaintextPassword( |),( |)saltRounds[^]*END_SYNC/gi,
  'You should call bcrypt.hashSync on myPlaintextPassword with saltRounds'
);
assert.match(
  data,
  /START_SYNC[^]*result.*=.*bcrypt.compareSync.*myPlaintextPassword( |),( |)hash[^]*END_SYNC/gi,
  'You should call bcrypt.compareSync on myPlaintextPassword with the hash generated in the last line'
);
```

