---
id: a202eed8fc186c8434cb6d61
title: بناء معكوس سلسلة نصية
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

في هذا المختبر، ستبني معكوسًا بسيطًا لسلسلة نصية يعكس حروف السلسلة المعطاة.

على سبيل المثال، يجب أن تتحول `"hello"` إلى `"olleh"`.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `reverseString` تأخذ سلسلة نصية كمعلمة.
2. يجب أن تُرجع الدالة السلسلة النصية المعكوسة.

# --hints--

يجب أن يكون لديك دالة باسم `reverseString`.

```js
assert.isFunction(reverseString);
```

يجب أن تأخذ `reverseString` سلسلة نصية كمعلمة.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

يجب أن تُرجع `reverseString("hello")` سلسلة نصية.

```js
assert.isString(reverseString('hello'));
```

يجب أن تُرجع `reverseString("hello")` السلسلة النصية `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

يجب أن تُرجع `reverseString("Howdy")` السلسلة النصية `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

يجب أن تُرجع `reverseString("Greetings from Earth")` السلسلة النصية `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
