---
id: 66ad9c633fa26d3c1475eae3
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Ora è il momento di aggiungere un altro messaggio dal bot.

Aggiungi un'altra istruzione `console` al codice che registra il messaggio `"Allow me to introduce myself."`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Dovresti avere una terza istruzione `console.log()` nel tuo codice.

```js
assert.lengthOf(logSpy.calls, 3);
```

La tua istruzione `console` dovrebbe contenere il messaggio `"Allow me to introduce myself."`. Non dimenticare le virgolette.

```js
assert.isTrue(
  logSpy.calls.some(
    (call) => /Allow\s+me\s+to\s+introduce\s+myself\s*\./.test(call[0])
  )
);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
