---
id: 66ada3f46945763dd97f43f8
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

In precedenza, hai creato le variabili `bot` e `botLocation`. Ora, le userai per mostrare nuovi messaggi nella console.

Nelle lezioni precedenti, hai imparato a lavorare con la concatenazione di stringhe usando l'operatore `+` per unire stringhe insieme in questo modo:

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

Ricorda che devi fare attenzione agli spazi quando concateni stringhe con variabili.

Crea una variabile chiamata `botIntroduction`.

Poi usa la concatenazione di stringhe con l'operatore `+` per unire la stringa `"My name is "` seguita dalla variabile `bot` seguita da un punto (`.`).

Assegna questo valore alla variabile `botIntroduction`.

Infine, mostra la variabile `botIntroduction` nella console.

# --hints--

Dovresti avere una variabile chiamata `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

La tua variabile `botIntroduction` dovrebbe contenere il valore di una stringa.

```js
assert.isString(botIntroduction);
```

Dovresti usare la concatenazione di stringhe con l'operatore `+` per unire la stringa `"My name is "` con la variabile `bot` seguita da un punto (`.`). Fai attenzione agli spazi. Assegna questa stringa concatenata alla variabile `botIntroduction`.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

Dovresti avere una quarta istruzione `console.log()` nel tuo codice.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

La tua istruzione `console` dovrebbe mostrare la variabile `botIntroduction`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
