---
id: 66adc82170acb1464bd348eb
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Gli ultimi messaggi del bot si concentreranno sull'argomento preferito del bot.

Inizia creando una nuova variabile chiamata `favoriteSubject` e assegnale la stringa `"Computer Science"`.

# --hints--

Dovresti avere una variabile chiamata `favoriteSubject`.

```js
assert.isNotNull(favoriteSubject);
```

La tua variabile `favoriteSubject` dovrebbe essere una stringa.

```js
assert.isString(favoriteSubject);
```

Dovresti assegnare la stringa `"Computer Science"` alla variabile `favoriteSubject`.

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
