---
id: 66ad8294a0ad902f1b31b612
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop imparerai a lavorare con i fondamenti di JavaScript costruendo un bot di saluto.

In questo primo passaggio, dovrai mostrare un messaggio nella console dal bot di saluto.

Ricorda che nelle lezioni precedenti hai imparato a usare `console.log()` e le stringhe.

Ecco un promemoria su come usare `console.log()` con le stringhe:

```js
console.log("Hello, World!");
```

Aggiungi un’istruzione `console.log()` che mostri nella console la stringa `"Hi there!"`. Non dimenticare le virgolette intorno al messaggio!

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Dovresti avere un’istruzione `console.log()` nel tuo codice.

```js
assert.isNotEmpty(logSpy.calls);
```

La tua istruzione `console` dovrebbe contenere il messaggio `"Hi there!"`. Non dimenticare le virgolette.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
