---
id: 66adb844118ba74107ce771f
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Il prossimo messaggio del bot riguarderà la posizione del bot.

Crea una variabile chiamata `botLocationSentence`.

Poi usa la concatenazione di stringhe con l'operatore `+` per unire la stringa `"I live in "` con la variabile `botLocation` seguita da un punto (`.`).

Assegna questo valore alla variabile `botLocationSentence`.

Infine, stampa il valore di `botLocationSentence` nella console.

# --hints--

Dovresti avere una variabile `botLocationSentence`.

```js
assert.isNotNull(botLocationSentence);
```

La tua variabile `botLocationSentence` dovrebbe contenere il valore di una stringa.

```js
assert.isString(botLocationSentence);
```

Dovresti usare la concatenazione di stringhe con l'operatore `+` per unire la stringa `"I live in "` con la variabile `botLocation` seguita da un punto (`.`). Fai attenzione agli spazi.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

Dovresti avere un'istruzione `console` nel tuo codice.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

La tua istruzione `console` dovrebbe stampare la variabile `botLocationSentence`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
