---
id: 66adcf383276814776aba3ca
title: Passaggio 14
challengeType: 1
dashedName: step-14
---

# --description--

Crea una variabile chiamata `favoriteSubjectSentence`.

Usa la concatenazione di stringhe per unire la stringa `"My favorite subject is "` con la variabile `favoriteSubject` seguita da un punto.

Assegna il risultato alla variabile `favoriteSubjectSentence`.

Poi, stampa il valore di `favoriteSubjectSentence` nella console.

# --hints--

Dovresti avere una variabile chiamata `favoriteSubjectSentence`.

```js
assert.isNotNull(favoriteSubjectSentence);
```

La tua variabile `favoriteSubjectSentence` dovrebbe essere una stringa.

```js
assert.isString(favoriteSubjectSentence);
```

Dovresti usare la concatenazione di stringhe per unire la stringa `"My favorite subject is "` con la variabile `favoriteSubject` seguita da un punto. Fai attenzione agli spazi.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

Dovresti avere un'istruzione `console` nel tuo codice.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

Dovresti stampare il valore di `favoriteSubjectSentence` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
