---
id: 66ad8294a0ad902f1b31b612
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop imparerai a lavorare con i fondamenti di JavaScript costruendo un bot di saluto.

In questo primo passaggio, dovrai mostrare un messaggio nella console dal bot di saluto.

Ricorda che nelle lezioni precedenti hai imparato a usare `console.log()` e le stringhe.

Ecco un promemoria su come usare `console.log()` con le stringhe:

```js
console.log("Hello, World!");
```

Aggiungi un’istruzione `console.log()` che mostri nella console la stringa `"Hi there!"`. Non dimenticare le virgolette intorno al messaggio!

# --hints--

Dovresti avere un’istruzione `console.log()` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

La tua istruzione `console` dovrebbe contenere il messaggio `"Hi there!"`. Non dimenticare le virgolette.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
