---
id: 66ad998285df023b23dacdd3
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora è il momento di inizializzare la variabile `botLocation`.

Quando devi dichiarare variabili con più parole, puoi usare la convenzione di denominazione <dfn>camelCase</dfn>.

Quando usi `camelCase`, la prima parola è tutta in minuscolo e la prima lettera di ogni parola successiva è maiuscola.

Ecco un esempio:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

Dichiara e assegna la stringa `"the universe"` alla variabile `botLocation` sulla stessa riga usando la parola chiave `let`.

# --hints--

Usa la parola chiave `let` per dichiarare una variabile.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

La tua variabile dovrebbe chiamarsi `botLocation`.

```js
assert.isNotNull(botLocation);
```

Dovresti assegnare una stringa alla tua variabile `botLocation`.

```js
assert.isString(botLocation);
```

Alla tua variabile `botLocation` dovrebbe essere assegnata la stringa `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
