---
id: 66adc42868cab843ccee87d9
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Ora è il momento di vedere il nuovo valore `bot`.

Inizia creando una nuova variabile chiamata `nicknameIntroduction`.

Usa la concatenazione di stringhe per unire la stringa `"My nickname is "` con la variabile `bot` seguita da un punto (`.`).

Assegna la stringa risultante alla variabile `nicknameIntroduction`.

Poi, stampa il valore di `nicknameIntroduction` nella console.

# --hints--

Dovresti avere una variabile chiamata `nicknameIntroduction`.

```js
assert.isNotNull(nicknameIntroduction);
```

La tua variabile `nicknameIntroduction` dovrebbe essere una stringa.

```js
assert.isString(nicknameIntroduction);
```

Dovresti usare la concatenazione di stringhe per unire la stringa `"My nickname is "` con la variabile `bot` seguita da un punto (`.`). Fai attenzione agli spazi.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

Dovresti avere un'istruzione `console` nel tuo codice.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

Dovresti stampare il valore di `nicknameIntroduction` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
