---
id: 66ad984e92ce213a66f7e51c
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora è il momento di assegnare un valore alla tua variabile `bot`.

Nelle lezioni precedenti, hai imparato come assegnare valori alle variabili in questo modo:

```js
variableName = "Here is the value";
```

Ricorda che ciò che si trova a destra dell'operatore di assegnazione `=` è il valore che stai assegnando alla variabile a sinistra.

Assegna la stringa `"teacherBot"` alla variabile `bot`.

# --hints--

Dovresti assegnare una stringa alla tua variabile `bot`.

```js
assert.isString(bot);
```

Alla tua variabile `bot` dovrebbe essere assegnata la stringa `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
