---
id: 66adc5919c1853448f119d11
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Ora sembra che il bot voglia cambiare il suo soprannome.

Usando la riassegnazione, assegna la stringa `"awesomeTeacherBot"` alla variabile `bot`.

# --hints--

Dovresti assegnare la stringa `"awesomeTeacherBot"` alla variabile `bot`.

```js
assert.strictEqual(bot, "awesomeTeacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
