---
id: 66ad8ab945d266383f318cbf
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Ora dovresti vedere il primo messaggio del bot nella console.

È il momento di aggiungere un secondo messaggio del bot.

Aggiungi un'altra istruzione `console.log` per mostrare il messaggio `"I am excited to talk to you."` nella console.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Dovresti avere una seconda istruzione `console.log()` nel tuo codice.

```js
assert.lengthOf(logSpy.calls, 2);
```

La tua istruzione `console` dovrebbe contenere il messaggio `"I am excited to talk to you."`. Non dimenticare le virgolette.

```js
assert.isTrue(
  logSpy.calls.some(
    (call) => /I\s+am\s+excited\s+to\s+talk\s+to\s+you\s*\./.test(call[0])
  )
);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
