---
id: 66adc6046acc18453decf577
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Per vedere il nuovo soprannome del bot, devi registrare un nuovo messaggio nella console.

Crea una nuova variabile chiamata `newNicknameGreeting`.

Poi usa la concatenazione di stringhe per unire la stringa `"I love my nickname but I wish people would call me "` con la variabile `bot` seguita da un punto.

Assegna il risultato alla variabile `newNicknameGreeting`.

Infine, registra il valore di `newNicknameGreeting` nella console.

# --hints--

Dovresti avere una variabile chiamata `newNicknameGreeting`.

```js
assert.isNotNull(newNicknameGreeting);
```

La tua variabile `newNicknameGreeting` dovrebbe essere una stringa.

```js
assert.isString(newNicknameGreeting);
```

Dovresti usare la concatenazione di stringhe per unire la stringa `"I love my nickname but I wish people would call me "` con la variabile `bot` seguita da un punto.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

Dovresti avere un'istruzione `console` nel tuo codice.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

Dovresti registrare il valore di `newNicknameGreeting` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
