---
id: 66ad8d150264db3926eccfeb
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Nelle lezioni precedenti, hai imparato la parola chiave `let` e come dichiarare variabili in JavaScript.

Ecco un promemoria su come dichiarare una variabile usando la parola chiave `let`:

```js
let greeting;
```

Usa la parola chiave `let` per dichiarare una variabile chiamata `bot`.

**NOTA**: Qui usi `let` perché più avanti nel workshop cambierai il valore della variabile `bot`.

# --hints--

Dovresti avere la parola chiave `let` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

Dovresti dichiarare una variabile chiamata `bot`.

```js
assert.isNotNull(bot);
```

Dovresti usare la parola chiave `let` per dichiarare la variabile.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
