---
id: 66fbcf750a62784cf11f5630
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Sasa unapaswa kusukuma vitu vitatu ndani ya safu ya data `recipes`. Ili kufanya hivyo, unaweza kutumia njia ya `push()`.

Tumia njia ya `push()` kusukuma vitu vyote vya mapishi ndani ya safu ya data `recipes`. Hakikisha kusukuma `recipe1`, `recipe2`, na `recipe3` kwa mpangilio huo huo.

Pia futa vigezo `recipe1Name`, `recipe2CookingTime` na `recipe3Ingredients`, pamoja na taarifa za `console.log` zinazochapisha vigezo hivyo.

# --hints--

Unapaswa kusukuma vitu vyote vya mapishi ndani ya safu ya data `recipes`. Hakikisha kusukuma kwa mpangilio walipotangazwa.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

Unapaswa kufuta kigezo cha `recipe1Name`.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

Unapaswa kufuta taarifa ya `console.log(recipe1Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

Unapaswa kufuta kigezo cha `recipe2CookingTime`.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

Unapaswa kufuta taarifa ya `console.log(recipe2CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

Unapaswa kufuta kigezo cha `recipe3Ingredients`.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

Unapaswa kufuta taarifa ya `console.log(recipe3Ingredients);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
