---
id: 66fbcf750a62784cf11f5633
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Unda kitendakazi cha `getDifficultyLevel` kinachopokea nambari inayosema muda wa kupika kama kigezo.

Ikiwa muda wa kupika ni chini au sawa na `30`, kitendakazi kinapaswa kurudisha `"easy"`. Ikiwa ni chini au sawa na `60`, kitendakazi kinapaswa kurudisha `"medium"`. Vinginevyo, kitendakazi kinapaswa kurudisha `"hard"`.

# --hints--

Unapaswa kuunda kitendakazi cha `getDifficultyLevel`.

```js
assert.isFunction(getDifficultyLevel);
```

Kitendakazi chako cha `getDifficultyLevel` kinapaswa kuwa na kigezo kimoja.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

Kitendakazi chako cha `getDifficultyLevel` kinapaswa kurudisha `"easy"` wakati muda wa kupika ni chini au sawa na `30`.

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

Kitendakazi chako cha `getDifficultyLevel` kinapaswa kurudisha `"medium"` wakati muda wa kupika ni mkubwa kuliko `31` na chini au sawa na `60`.

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

Kitendakazi chako cha `getDifficultyLevel` kinapaswa kurudisha `"hard"` wakati muda wa kupika ni mkubwa kuliko `60`.

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
