---
id: 66fbcf750a62784cf11f562d
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Tengeneza kitu kinachoitwa `recipe1`. Ndani ya kitu cha `recipe1`, tengeneza sifa ya `name` yenye thamani `"Spaghetti Carbonara"`.

Pia ndani ya kitu cha `recipe1`, tengeneza sifa ya `ingredients` yenye thamani ya safu ya data. Safu ya data hiyo inapaswa kuwa na `spaghetti`, `Parmesan cheese`, `pancetta`, na `black pepper` ndani yake.

# --hints--

Unapaswa kutengeneza kitu kinachoitwa `recipe1`.

```js
assert.isObject(recipe1);
```

Kitu chako cha `recipe1` kinapaswa kuwa na sifa ya `name`.

```js
assert.property(recipe1, 'name');
```

Sifa yako ya `name` inapaswa kuwekwa kuwa `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

Kitu chako cha `recipe1` kinapaswa kuwa na sifa ya `ingredients`.

```js
assert.property(recipe1, 'ingredients');
```

Sifa yako ya `ingredients` inapaswa kuwa na thamani ya safu ya data.

```js
assert.isArray(recipe1.ingredients);
```

Thamani ya safu ya data ya sifa yako ya `ingredients` inapaswa kuwa na `spaghetti`, `Parmesan cheese`, `pancetta`, na `black pepper` ndani yake.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

--fcc-editable-region--

--fcc-editable-region--
```
