---
id: 66fbcf750a62784cf11f5635
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Sasa unaweza kujaza kila kipengee cha safu ya data ya `recipes` na thamani za sifa za `totalIngredients` na `difficultyLevel`.

Kwa sasa, pata upatikanaji wa `totalIngredients` na `difficultyLevel` za `recipe1` na weka thamani zao kwa matokeo yanayofaa ya wito za kitendakazi na hoja.

# --hints--

Unapaswa kupata upatikanaji wa sifa ya `totalIngredients` ya `recipe1`.

```js
assert.isNotNull(recipe1.totalIngredients);
```

Unapaswa kuweka thamani ya matokeo ya wito wa `getTotalIngredients` ukiwa na `recipe1.ingredients` kwa sifa ya `totalIngredients` ya `recipe1`.

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

Unapaswa kupata upatikanaji wa sifa ya `difficultyLevel` ya `recipe1`.

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

Unapaswa kuweka thamani ya matokeo ya wito wa `getDifficultyLevel` ukiwa na `recipe1.cookingTime` kwa sifa ya `cookingTime` ya `recipe1`.

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
