---
id: 66fbcf750a62784cf11f5634
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Ni wakati wa kujaribu kila kitendakazi. Unaweza kutumia mapishi yoyote kwa hili, lakini mafunzo haya yataanza na `recipe1`.

Tengeneza vigezo viwili vipya: `recipe1TotalIngredients` na `recipe1DifficultyLevel`. Weka thamani zao kwa kuita kitendakazi kinachofaa kwa kila kigezo na kupitisha sifa ya `recipe1` inayohitajika.

Mwishowe, andika kila kigezo kwenye konsoli ili kuona matokeo.

# --hints--

Unapaswa kutengeneza kigezo cha `recipe1TotalIngredients`.

```js
assert.isNotNull(recipe1TotalIngredients);
```

Kigezo chako cha `recipe1TotalIngredients` kinapaswa kuwekwa thamani ya `getTotalIngredients(recipe1.ingredients)`.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

Unapaswa kuandika `recipe1TotalIngredients` kwenye konsoli.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

Unapaswa kutengeneza kigezo cha `recipe1DifficultyLevel`.

```js
assert.isNotNull(recipe1DifficultyLevel);
```

Kigezo chako cha `recipe1DifficultyLevel` kinapaswa kuwekwa thamani ya `getDifficultyLevel(recipe1.cookingTime)`.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

Unapaswa kuandika `recipe1DifficultyLevel` kwenye konsoli.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
