---
id: 66fbcf750a62784cf11f562e
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Unda kitu cha `recipe2` chenye sifa na thamani zifuatazo:

| Funguo | Thamani   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

Unapaswa kuunda kitu kinachoitwa `recipe2`.

```js
assert.isObject(recipe2);
```

Kitu chako cha `recipe2` kinapaswa kuwa na sifa ya `name`.

```js
assert.property(recipe2, "name");
```

Sifa yako ya `name` inapaswa kuwekwa kuwa `Chicken Curry`.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

Kitu chako cha `recipe2` kinapaswa kuwa na sifa ya `ingredients`.

```js
assert.property(recipe2, "ingredients");
```

Sifa yako ya `ingredients` inapaswa kuwa na thamani ya safu ya data.

```js
assert.isArray(recipe2.ingredients);
```

Thamani ya safu ya data ya sifa yako ya `ingredients` inapaswa kuwa na `chicken breast`, `coconut milk`, `curry powder`, `onion` na `garlic` ndani yake.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

Kitu chako cha `recipe2` kinapaswa kuwa na sifa ya `cookingTime`.

```js
assert.property(recipe2, "cookingTime");
```

Thamani ya sifa yako ya `cookingTime` inapaswa kuwa nambari.

```js
assert.isNumber(recipe2.cookingTime);
```

Sifa yako ya `cookingTime` inapaswa kuwekwa kuwa `42`.

```js
assert.equal(recipe2.cookingTime, 42);
```

Kitu chako cha `recipe2` kinapaswa kuwa na sifa ya `totalIngredients`.

```js
assert.property(recipe2, "totalIngredients");
```

Sifa yako ya `totalIngredients` inapaswa kuwekwa kuwa `null`.

```js
assert.isNull(recipe2.totalIngredients);
```

Kitu chako cha `recipe2` kinapaswa kuwa na sifa ya `difficultyLevel`.

```js
assert.property(recipe2, "difficultyLevel");
```

Sifa yako ya `difficultyLevel` inapaswa kuwekwa kuwa mfuatano wa herufi tupu.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
