---
id: 66fbcdfed6848e48947dbbc9
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda kifuatiliaji cha mapishi kwa kutumia vitu vya JavaScript.

Anza kwa kuunda safu ya data tupu iitwayo `recipes`. Hii ndiyo safu ya data utakayoweka vitu vya mapishi baadaye.

# --hints--

Unapaswa kuunda safu ya data iitwayo `recipes`.

```js
assert.isArray(recipes);
```

Safu ya data yako ya `recipes` inapaswa kuwa tupu.

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
