---
id: 66fbcf750a62784cf11f5632
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Unda kitendakazi cha `getTotalIngredients` kinachopokea hoja moja, inayo wakilisha safu ya data yenye viungo, na kurudisha idadi ya viungo kutoka kwenye safu ya data iliyotumwa kwa kitendakazi.

# --hints--

Unapaswa kuunda kitendakazi cha `getTotalIngredients`.

```js
assert.isFunction(getTotalIngredients);
```

Kitendakazi chako cha `getTotalIngredients` kinapaswa kuwa na kigezo kimoja.

```js
assert.lengthOf(getTotalIngredients, 1);
```

Kitendakazi chako cha `getTotalIngredients` kinapaswa kurudisha nambari.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)` inapaswa kurudisha `4`.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` inapaswa kurudisha `5`.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` inapaswa kurudisha `3`.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
