---
id: 66ff02792b3042a2b811f475
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Ongeza sifa zifuatazo kwenye kitu cha `recipe1`:

| Funguo | Thamani   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

Sifa zilizo na `null` na thamani za mfuatano wa herufi tupu zitasasishwa baadaye baada ya kuzihesabu.

# --hints--

Kitu chako cha `recipe1` kinapaswa kuwa na sifa ya `cookingTime`.

```js
assert.property(recipe1, "cookingTime");
```

Sifa yako ya `cookingTime` inapaswa kuwekwa kuwa `22`.

```js
assert.equal(recipe1.cookingTime, 22);
```

Kitu chako cha `recipe1` kinapaswa kuwa na sifa ya `totalIngredients`.

```js
assert.property(recipe1, "totalIngredients");
```

Sifa yako ya `totalIngredients` inapaswa kuwekwa kuwa `null`.

```js
assert.isNull(recipe1.totalIngredients);
```

Kitu chako cha `recipe1` kinapaswa kuwa na sifa ya `difficultyLevel`.

```js
assert.property(recipe1, "difficultyLevel");
```

Sifa yako ya `difficultyLevel` inapaswa kuwekwa kuwa mfuatano wa herufi tupu.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
