---
id: 670e4f45f7116c0f216a5177
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Kabla hujaendelea, unapaswa kufanya mazoezi ya jinsi ya kupata sifa kutoka kwa kitu.

Unaweza kutumia alama ya nukta (`.`) au alama ya mabano (`[]`) kufanya hivi. Hapa kuna mfano:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Pata sifa ya `name` ya `recipe1`, na weka thamani yake katika mabadiliko `recipe1Name`.

Kisha, pata sifa ya `cookingTime` ya `recipe2` na weka thamani yake katika mabadiliko `recipe2CookingTime`.

Mwishowe, pata sifa ya `ingredients` ya `recipe3` na weka thamani yake katika mabadiliko `recipe3Ingredients`.

Hakikisha mabadiliko yote uliyounda yameandikwa kwenye konsoli.

# --hints--

Unapaswa kuunda mabadiliko `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

Unapaswa weka thamani ya sifa ya `name` ya `recipe1` katika kigezo chako cha `recipe1Name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Unapaswa kuandika `recipe1Name` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Unapaswa kuunda mabadiliko `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

Unapaswa kuweka thamani ya sifa ya `cookingTime` ya `recipe2` katika mabadiliko `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Unapaswa kuandika `recipe2CookingTime` kwenye konsoli.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

Unapaswa kuunda mabadiliko `recipe3Ingredients`.

```js
assert.isDefined(recipe3Ingredients);
```

Unapaswa kuweka thamani ya sifa ya `ingredients` ya `recipe3` katika mabadiliko `recipe3Ingredients`.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

Unapaswa kuandika `recipe3Ingredients` kwenye konsoli.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
