---
id: 69161f4c0c9c8c48b3c843c5
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Unda kitu cha `recipe3` chenye sifa na thamani zifuatazo:

|Key|Value|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

Unapaswa kuunda kitu kinachoitwa `recipe3`.

```js
assert.isObject(recipe3);
```

Kitu chako cha `recipe3` kinapaswa kuwa na sifa ya `name`.

```js
assert.property(recipe3, "name");
```

Sifa yako ya `name` inapaswa kuwekwa kuwa `Vegetable Stir Fry`.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

Kitu chako cha `recipe3` kinapaswa kuwa na sifa ya `ingredients`.

```js
assert.property(recipe3, "ingredients");
```

Sifa yako ya `ingredients` inapaswa kuwa na thamani ya safu ya data.

```js
assert.isArray(recipe3.ingredients);
```

Thamani ya safu ya data ya sifa yako ya `ingredients` inapaswa kuwa na `broccoli`, `carrot` na `bell pepper` ndani yake.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

Kitu chako cha `recipe3` kinapaswa kuwa na sifa ya `cookingTime`.

```js
assert.property(recipe3, "cookingTime");
```

Thamani ya sifa yako ya `cookingTime` inapaswa kuwa nambari.

```js
assert.isNumber(recipe3.cookingTime);
```

Sifa yako ya `cookingTime` inapaswa kuwekwa kuwa `15`.

```js
assert.equal(recipe3.cookingTime, 15);
```

Kitu chako cha `recipe3` kinapaswa kuwa na sifa ya `totalIngredients`.

```js
assert.property(recipe3, "totalIngredients");
```

Sifa yako ya `totalIngredients` inapaswa kuwekwa kuwa `null`.

```js
assert.isNull(recipe3.totalIngredients);
```

Kitu chako cha `recipe3` kinapaswa kuwa na sifa ya `difficultyLevel`.

```js
assert.property(recipe3, "difficultyLevel");
```

Sifa yako ya `difficultyLevel` inapaswa kuwekwa kuwa mfuatano wa herufi tupu.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
