---
id: 619665c9abd72906f3ad30f9
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ハッピーなFlappy Penguinを作りながら、CSSの変形とアニメーションについてさらに学んでいきます。

基本的なHTMLのボイラープレートから始めてください。`DOCTYPE`宣言、言語が英語に設定された`html`要素、適切な`meta`タグ、`head`、`body`、`title`要素を含めてください。また、スタイルシートをページにリンクしてください。

# --hints--

コードには`<!DOCTYPE html>`宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

コードには`html`要素が必要です。

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

`html`要素の開始タグには`lang`属性が`en`に設定されている必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html`要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/gi));
```

コードには`head`要素が`html`要素の中に必要です。

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

コードには`body`要素が`html`要素の中に必要です。

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

`head`要素は`body`要素の前に配置してください。

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

`meta`要素が2つ必要です。

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

1つの`meta`要素には`name`が`viewport`に、`content`が`width=device-width, initial-scale=1.0`に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

もう1つの`meta`要素には`charset`属性が`UTF-8`に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

コードには`title`要素が必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title`にはテキストが含まれている必要があります。

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

コードには`link`要素が必要です。

```js
assert.exists(document.querySelector('link'));
```

`link`要素は`head`要素の中にある必要があります。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

`link`要素には`rel`属性が`stylesheet`の値で設定されている必要があります。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link`要素には`href`属性が`styles.css`の値で設定されている必要があります。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
