---
id: 0095d0567b6817114989a1dd
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 JavaScript로 최단 경로 알고리즘을 구현할 것입니다. 가중치 그래프에서 노드 간 최단 경로를 계산하고, 거리와 경로를 모두 반환하는 함수를 작성하세요.

예를 들어, 도시들이 서로 다른 거리의 도로로 연결된 그래프가 있을 때, 알고리즘은 한 도시에서 다른 도시로 가는 최단 라우트를 찾습니다. 만약 A 도시에서 D 도시로 이동하려면, 알고리즘은 A ⇨ B ⇨ C ⇨ D (총 15km)가 직접 A ⇨ D (20km)로 가는 것보다 더 짧다는 것을 찾아낼 수 있습니다.

시작하려면, `INF`라는 변수를 정의하고 `Infinity` 값을 할당하세요. 나중에 두 노드 사이에 직접 연결이 없음을 나타내는 데 사용할 것입니다.

# --hints--

`INF`라는 변수를 정의하세요.

```js
assert.isDefined(INF);
```

`INF` 변수는 `Infinity` 값을 가져야 합니다.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
