---
id: 69c918841987bf16c1d1baca
title: 19단계
challengeType: 1
dashedName: step-19
---

# --description--

이웃 노드로 가는 더 짧은 경로를 찾으면, 그 경로를 기록한 경로도 업데이트해야 합니다.

JavaScript에서는 전개 연산자를 사용해 기존 배열과 새 요소를 결합한 새 배열을 만들 수 있습니다:

```js
const newPath = [...existingArray, newElement];
```

중첩된 `if` 블록 안에서, `paths[nodeNo]`을 현재 노드까지의 경로에 `nodeNo`를 끝에 추가한 경로로 전개 구문을 사용해 다음과 같이 업데이트하세요: `[...paths[current], nodeNo]`.

# --hints--

중첩된 `paths[nodeNo]` 블록 안에서 `if`를 업데이트해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=/
);
```

`paths[nodeNo]`는 `[...paths[current], nodeNo]`로 설정해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=\s*\[\s*\.\.\.\s*paths\s*\[\s*current\s*\]\s*,\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          --fcc-editable-region--

          --fcc-editable-region--
        }
      }
    }
  }
}
```
