---
id: 69c937c73f23b1140bcd297a
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

`targetNode` 매개변수는 선택 사항입니다. 제공하지 않으면, 함수가 `startNode`에서 그래프의 다른 모든 노드까지 최단 경로를 계산합니다.

`targetNode` 매개변수에 `null` 기본값을 지정하세요.

# --hints--

`targetNode` 함수의 `shortestPath` 매개변수는 `null` 기본값이어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {

}
--fcc-editable-region--
```
