---
id: 69c918841987bf16c1d1babb
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이제 모든 작업을 수행하는 메인 함수를 만드세요.

`shortestPath`라는 이름의 함수를 만들고 세 개의 매개변수 `matrix`, `startNode`, `targetNode`를 받으세요.

다음 몇 단계에 걸쳐 이 함수의 로직을 완성할 것입니다.

# --hints--

`shortestPath`라는 이름의 함수를 만들어야 합니다.

```js
assert.isFunction(shortestPath);
```

`shortestPath` 함수는 `matrix`, `startNode`, `targetNode` 매개변수를 가져야 합니다.

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
