---
id: 69c918841987bf16c1d1bac8
title: 17단계
challengeType: 1
dashedName: step-17
---

# --description--

거리를 업데이트하기 전에 이웃이 고려할 가치가 있는지 확인해야 합니다. 실제 간선이 있어야 하고(`distance !== INF`) 아직 방문하지 않은 상태여야 합니다.

이웃 `for` 반복문 안에 `if` 문을 추가하여 `distance`가 `INF`와 엄격히 같지 않고 **그리고** `nodeNo`가 방문되지 않았는지 확인하세요.

`if` 블록 안에 `newDistance` 변수를 선언하고 `distances[current] + distance`를 할당하세요. 이것은 현재 노드를 통해 이웃에 도달하는 총 비용입니다.

# --hints--

조건 `if`을 가진 `distance !== INF && !visited[nodeNo]` 문을 추가해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

`newDistance` 블록 안에 값 `distances[current] + distance`를 가진 `if`을 선언해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
