---
id: 69c918841987bf16c1d1babf
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

각 노드에 도달하기 위해 실제로 거친 경로도 추적해야 합니다. 각 항목이 거친 라우트를 나타내는 노드 인덱스 배열인 경로 목록을 저장할 것입니다.

초기에는 각 노드의 경로가 자기 자신만 포함합니다. `Array.from()`를 사용해 이 구조를 만들 수 있습니다.

```js
Array.from({ length: n }, (_, i) => [i]);
```

이것은 `n`개의 요소를 가진 배열을 만드는데, 인덱스 `i`에 있는 요소는 `[i]`입니다.

`shortestPath` 함수 내에서 `paths`라는 변수를 만들고 `Array.from({ length: n }, (_, i) => [i])`를 할당하세요.

# --hints--

`paths` 함수 안에 `shortestPath` 변수를 만들어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

`paths` 변수는 `Array.from`를 사용해 초기화해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

`paths` 변수는 `{ length: n }`의 첫 번째 인수로 `Array.from`를 사용해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

`Array.from`의 두 번째 인수는 콜백 `(_, i) => [i]`여야 하며, `_`는 사용하지 않는 요소 자리 표시자이고 `i`는 인덱스입니다. 이 콜백은 `[i]`를 반환해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
