---
id: 69c918841987bf16c1d1baba
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

그래프의 인접 행렬을 나타내려면 2차원 배열이 필요합니다. 각 값은 두 노드 사이의 간선 가중치(거리)를 나타냅니다. `INF` 값은 직접 연결된 간선이 없음을 의미합니다.

`adjMatrix`이라는 변수를 새로 만들고 다음 값을 가진 2차원 배열을 할당하세요:

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

`adjMatrix`라는 변수를 만들어야 합니다.

```js
assert.isDefined(adjMatrix);
```

`adjMatrix`은 6행 6열의 2차원 배열이어야 합니다.

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix`에는 올바른 값이 들어가야 합니다.

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
