---
id: 69c918841987bf16c1d1bac2
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

반복문 안에서, 아직 방문하지 않은 가장 작은 알려진 거리를 가진 노드를 찾기 위해 모든 노드를 검사해야 합니다.

메인 반복문 안에 `for` 반복문을 추가하여 각 노드를 반복하세요. `nodeNo`을 `0`에서 `n`까지 반복 변수로 사용하세요.

현재는 본문을 비워 두세요.

# --hints--

메인 반복문 안에 `for` 두 번째(내부) 반복문이 있어야 하며, `nodeNo`가 반복 변수여야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+nodeNo\s*=\s*0\s*;\s*nodeNo\s*<\s*n\s*;\s*nodeNo\+\+\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

--fcc-editable-region--
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;

  }
--fcc-editable-region--
}
```
