---
id: 69c918841987bf16c1d1bac5
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

가장 가까운 방문하지 않은 노드를 찾기 위해 모든 노드를 스캔한 후, 유효한 노드를 찾지 못한 경우를 처리해야 합니다. 이 상황은 남은 모든 노드가 도달 불가능할 때 발생합니다.

내부 `for` 반복문이 끝난 후, 그러나 외부 반복문 안에 있을 때, `if`가 `current`와 엄격히 같은지 확인하는 `-1` 문을 추가하세요. 만약 참이라면, `break` 블록 안에서 `if`를 사용해 외부 반복문을 조기에 종료하세요.

# --hints--

`if`인지 확인하는 `current === -1` 문을 추가해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

`break` 블록 안에서 `if`를 사용해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
