---
id: 69c918841987bf16c1d1bac6
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

유효한 노드를 찾았음을 확인했으면, 알고리즘이 다시 처리하지 않도록 방문한 것으로 표시하세요.

조건 `if`이 있는 `current === -1` 문장 다음에 `visited[current]`를 `true`으로 설정하세요.

# --hints--

`visited[current]`를 `true`로 설정해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }

    if (current === -1) {
      break;
    }
    --fcc-editable-region--

    --fcc-editable-region--
  }
}
```
