---
id: 69c918841987bf16c1d1bacb
title: 20단계
challengeType: 1
dashedName: step-20
---

# --description--

메인 루프가 끝나면 어떤 노드의 결과를 표시할지 결정해야 합니다.

특정 `targetNode`가 제공되었다면, 그 노드의 결과만 보여주세요. 그렇지 않으면 모든 노드의 결과를 보여주세요.

JavaScript에서는 `[...Array(n).keys()]`를 사용해 배열의 모든 인덱스를 반복 가능한 객체로 얻을 수 있습니다.

외부 `for` 루프 후에 `targets`라는 변수를 만드세요. 삼항 표현식을 사용해 `targetNode`가 엄격하게 `null`가 아니면 `[targetNode]`를 할당하고, 그렇지 않으면 `[...Array(n).keys()]`를 할당하세요.

# --hints--

메인 루프 후에 `targets` 변수를 만드세요.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+targets\s*=/
);
```

`targets` 변수는 `targetNode !== null`을 확인하는 삼항식을 사용해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /targets\s*=\s*targetNode\s*!==\s*null\s*\?/
);
```

대상이 제공되면 `targets`는 `[targetNode]`여야 하고, 그렇지 않으면 `[...Array(n).keys()]`여야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /\?\s*\[\s*targetNode\s*\]\s*:\s*\[\s*\.\.\.\s*Array\s*\(\s*n\s*\)\s*\.keys\s*\(\s*\)\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
