---
id: 69c918841987bf16c1d1bacf
title: 24단계
challengeType: 1
dashedName: step-24
---

# --description--

이제 도달 가능한 각 노드에 대한 결과를 출력하세요.

`path` 변수 뒤에 `console.log()` 호출을 추가하고, 템플릿 리터럴로 다음을 출력하세요:

```md
\n{startNode}-{nodeNo} distance: {distances[nodeNo]}\nPath: {path}
```

예를 들어, `startNode`가 `0`이고 `nodeNo`가 `3`라면 출력은 다음과 같아야 합니다:

```md
0-3 distance: 4
Path: 0 -> 2 -> 3
```

# --hints--

`console.log()` 반복문 안에 `for...of`가 있어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*(const|let)\s+nodeNo\s+of\s+targets\s*\)[\s\S]*?console\s*\.\s*log\s*\(/
);
```

`console.log()`는 `${startNode}-${nodeNo} distance: ${distances[nodeNo]}`를 포함하는 템플릿 리터럴을 사용해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*`[\s\S]*?\$\{\s*startNode\s*\}[\s\S]*?\$\{\s*nodeNo\s*\}[\s\S]*?\$\{\s*distances\s*\[\s*nodeNo\s*\]\s*\}[\s\S]*?`\s*\)/
);
```

템플릿 리터럴에는 `${path}`도 포함되어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*`[\s\S]*?\$\{\s*path\s*\}[\s\S]*?`\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }

  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
    if (nodeNo === startNode || distances[nodeNo] === INF) {
      continue;
    }
    const path = paths[nodeNo].join(' -> ');
    --fcc-editable-region--

    --fcc-editable-region--
  }
}
```
