---
id: 69c918841987bf16c1d1bace
title: 23단계
challengeType: 1
dashedName: step-23
---

# --description--

노드에 도달할 수 있고 시작 노드가 아니라는 것을 알았으니, 그 경로를 읽기 쉬운 문자열로 형식화하세요.

`paths[nodeNo]` 배열은 `nodeNo`에 도달하기 위해 방문한 노드 인덱스 수열을 담고 있습니다. `0 -> 2 -> 3` 메서드를 사용해 `.join()`처럼 사람이 읽기 쉬운 문자열로 바꿀 수 있습니다.

`for...of` 반복문 안에서, `if`를 사용한 `continue` 문장 뒤에 `path`라는 변수를 선언하고 `paths[nodeNo].join(' -> ')` 값을 할당하세요.

# --hints--

`path` 반복문 안에 `for...of`라는 변수가 있어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*(const|let)\s+nodeNo\s+of\s+targets\s*\)[\s\S]*?(const|let)\s+path\s*=/
);
```

`path` 변수는 `paths[nodeNo].join(' -> ')` 값을 가져야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /path\s*=\s*paths\s*\[\s*\w+\s*\]\s*\.\s*join\s*\(\s*['"]\s*->\s*['"]\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }

  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
    if (nodeNo === startNode || distances[nodeNo] === INF) {
      continue;
    }
    --fcc-editable-region--

    --fcc-editable-region--
  }
}
```
