---
id: 69c918841987bf16c1d1bac3
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

내부 `for` 반복문 안에서 현재 노드가 방문하지 않았고 지금까지 찾은 것보다 더 가까운지 확인해야 합니다.

`if`가 방문하지 않았고 `nodeNo`이 `distances[nodeNo]`보다 작은지 확인하는 `minDistance` 문을 추가하세요.

본문은 일단 비워 두세요.

# --hints--

내부 `if` 반복문 안에 조건 `!visited[nodeNo] && distances[nodeNo] < minDistance`을 가진 `for` 문이 있어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*!visited\s*\[\s*nodeNo\s*\]\s*&&\s*distances\s*\[\s*nodeNo\s*\]\s*<\s*minDistance\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
