---
id: 69c918841987bf16c1d1bac4
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

조건이 참이면 현재 노드는 지금까지 찾은 최상의 방문하지 않은 후보입니다. 이를 반영하도록 추적 변수를 업데이트하세요.

`if` 문장 안에서 `minDistance`을 `distances[nodeNo]`로, `current`을 `nodeNo`로 업데이트하세요.

# --hints--

`minDistance` 블록 안에서 `distances[nodeNo]`를 `if`으로 설정해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /minDistance\s*=\s*distances\s*\[\s*nodeNo\s*\]/
);
```

`current` 블록 안에서 `nodeNo`을 `if`로 설정해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /current\s*=\s*nodeNo/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
