---
id: 69c918841987bf16c1d1bac0
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

알고리즘이 실행되는 동안 이미 처리한 노드를 추적하여 다시 방문하지 않도록 해야 합니다.

`visited`라는 변수를 만들고 `new Array(n).fill(false)`을 할당하세요.

# --hints--

`visited`라는 변수를 만들어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+visited\s*=/
);
```

`visited` 변수는 `new Array(n).fill(false)`로 설정해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*false\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
--fcc-editable-region--

--fcc-editable-region--
}
```
