---
id: 69c918841987bf16c1d1bad1
title: 26단계
challengeType: 1
dashedName: step-26
---

# --description--

함수가 완성되었습니다. 이제 호출하여 결과를 확인하세요.

파일 하단(함수 외부)에서 `shortestPath`를 `adjMatrix`을 인접 행렬로, `0`를 시작 노드로, `5`를 목표 노드로 호출하세요.

이렇게 하면 인접 행렬에서 노드 `0`에서 노드 `5`까지의 최단 경로가 출력됩니다.

이로써 최단 경로 알고리즘 워크숍이 완료됩니다!

# --hints--

`shortestPath` 함수를 `adjMatrix, 0, 5`와 함께 호출해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /shortestPath\s*\(\s*adjMatrix\s*,\s*0\s*,\s*5\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }

  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
    if (nodeNo === startNode || distances[nodeNo] === INF) {
      continue;
    }
    const path = paths[nodeNo].join(' -> ');
    console.log(`\n${startNode}-${nodeNo} distance: ${distances[nodeNo]}\nPath: ${path}`);
  }
  return [distances, paths];
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }

  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
    if (nodeNo === startNode || distances[nodeNo] === INF) {
      continue;
    }
    const path = paths[nodeNo].join(' -> ');
    console.log(`\n${startNode}-${nodeNo} distance: ${distances[nodeNo]}\nPath: ${path}`);
  }
  return [distances, paths];
}

shortestPath(adjMatrix, 0, 5);
```
