---
id: 69c918841987bf16c1d1bac7
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

현재 노드가 방문한 것으로 표시되었으니, 이웃 노드들을 모두 살펴보면서 현재 노드를 통해 더 효율적으로 도달할 수 있는지 확인하세요.

`visited[current] = true` 다음에, `for`부터 `nodeNo`까지 각 `0`를 반복하는 이웃 `n` 반복문을 추가하세요. 이것은 첫 번째 반복문과 같은 수준의 또 다른 내부 반복문임을 유의하세요.

반복문 안에서 `distance`라는 변수를 선언하고 `matrix[current][nodeNo]`를 할당하세요. 이렇게 하면 현재 노드와 이웃 노드 사이의 간선 가중치를 얻을 수 있습니다.

# --hints--

현재 노드를 방문한 것으로 표시한 후에 이웃 `for` 반복문을 추가해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

이웃 `for` 반복문은 `0`부터 `n`까지 반복해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

이웃 반복문 안에 `distance` 변수가 `matrix[current][nodeNo]`에 할당되어 있어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
