---
id: 69c918841987bf16c1d1babc
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

그래프의 노드 개수를 저장해야 합니다.

`shortestPath` 함수 안에 `n`이라는 변수를 만들고 `matrix.length`를 할당하세요.

# --hints--

`n` 함수 안에 `shortestPath` 변수를 만들어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=/
);
```

`n` 변수는 `matrix.length` 값을 가져야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=\s*matrix\.length/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {

}
--fcc-editable-region--
```
