---
id: 69c918841987bf16c1d1babe
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

시작 노드에서 자기 자신까지의 거리는 항상 `0`입니다. 이를 반영하도록 `distances` 배열을 업데이트해야 합니다.

`distances` 배열 다음에 `distances[startNode]`을 `0`로 설정하세요.

# --hints--

`distances[startNode]`를 `0`로 설정해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*startNode\s*\]\s*=\s*0/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);

}
--fcc-editable-region--
```
