---
id: 69c918841987bf16c1d1bac9
title: 18단계
challengeType: 1
dashedName: step-18
---

# --description--

이제 현재 노드를 거쳐 이웃으로 가는 경로가 더 짧은지 확인하세요. `newDistance`가 이미 저장된 값보다 더 좋으면, 그것을 업데이트해야 합니다.

기존 `if` 블록 안에 `if`가 `newDistance`보다 작은지 확인하는 중첩된 `distances[nodeNo]` 문을 추가하세요. 그 중첩된 `if` 블록 안에서 `distances[nodeNo]`를 `newDistance`로 업데이트하세요.

# --hints--

조건 `if`를 가진 중첩된 `newDistance < distances[nodeNo]` 문을 추가해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*newDistance\s*<\s*distances\s*\[\s*\w+\s*\]\s*\)/
);
```

중첩된 `distances[nodeNo]` 블록 안에서 `newDistance`를 `if`로 업데이트해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*\w+\s*\]\s*=\s*newDistance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
