---
id: 69c918841987bf16c1d1bac1
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

이제 알고리즘을 구동하는 메인 루프를 추가하세요. 이 루프는 그래프의 각 노드마다 한 번씩 실행되며, 매번 가장 가까운 방문하지 않은 노드를 선택합니다.

`shortestPath` 함수 안에 `for`를 루프 변수로 사용하여 `n`번 실행되는 `i` 루프를 추가하세요.

루프 본문 시작 부분에서 `let`를 사용해 `minDistance`와 `current`로 각각 설정된 두 변수 `INF`와 `-1`를 선언하세요. 이 변수들은 각 반복에서 발견된 가장 가까운 방문하지 않은 노드를 추적합니다.

# --hints--

`for` 안에 `shortestPath`번 실행되는 `n` 루프가 있어야 합니다. 그 안에 `i`로 `let`를 설정하는 것을 잊지 마세요.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

루프 안에 `minDistance`로 설정된 `INF` 변수가 있어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

루프 안에 `current`로 설정된 `-1` 변수가 있어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
