---
id: 69c918841987bf16c1d1babd
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

시작 노드에서 다른 모든 노드까지 알려진 가장 짧은 거리를 추적해야 합니다. 처음에는 모든 노드가 무한히 멀리 있다고 가정하세요.

JavaScript에서는 `new Array(n).fill(value)`를 사용해 값이 미리 채워진 배열을 만들 수 있습니다. 예를 들어, `new Array(3).fill(0)`는 `[0, 0, 0]`를 만듭니다.

`shortestPath` 함수 안에 `distances`라는 변수를 만들고 `new Array(n).fill(INF)`를 할당하세요.

# --hints--

`distances`라는 변수를 `shortestPath` 안에 만들어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

`distances` 변수는 `new Array(n).fill(INF)`로 설정해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
