---
id: 6909ed80e1c195f49564edac
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会中，你将学习如何使用条件式语句和比较操作符。

首先声明一个名为 `truthyOrFalsy` 的变量，并赋值为布尔值 `true`。

# --hints--

你应该有一个名为 `truthyOrFalsy` 的变量。

```js
assert.exists(truthyOrFalsy);
```

你应该将 `true` 赋值给你的 `truthyOrFalsy` 变量。

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
