---
id: 690a58cc4f828a49eab1b53a
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

声明一个名为 `hasDeveloperJob` 的变量，并赋值为布尔值 `true`。

# --hints--

你应该有一个名为 `hasDeveloperJob` 的变量。

```js
assert.exists(hasDeveloperJob);
```

你应该为你的 `hasDeveloperJob` 变量赋值 `true`。

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
