---
id: 690a3f884e083a6a33b81fcf
title: 步骤 13
challengeType: 1
dashedName: step-13
---

# --description--

最后，更改 `timmyAge` 的声明并将其设置为小于 `16` 的数字。之后你会看到 `"Timmy is not old enough to drive."` 被记录到控制台。

# --hints--

你的 `timmyAge` 变量的值应小于 `16`。

```js
assert.isAtMost(timmyAge, 15);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--
const timmyAge = 18;
--fcc-editable-region--

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```

# --solutions--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 15;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```
