---
id: 690a0df27ee9ee36d935d530
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

声明另一个名为 `isTimmyAGamer` 的变量并赋值为布尔值 `false`。

# --hints--

你应该声明一个名为 `isTimmyAGamer` 的变量。

```js
assert.exists(isTimmyAGamer);
```

你应该为你的 `isTimmyAGamer` 变量赋值布尔值 `false`。

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
