---
id: 690a57095c4afb27bbe751b3
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

假值是在布尔上下文中计算为 `false` 的值。假值的示例有：

- `""`（空字符串）
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

你将在完成本课程时遇到所有这些内容。

现在，通过将 `truthyOrFalsy` 的值设置为空字串来更改赋值。之后，你会看到 `false` 被记录到控制台。

# --hints--

你应该将空字串赋值给 `truthyOrFalsy`。

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
