---
id: 690a562fd611e8155599ce98
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

通过将 `truthyOrFalsy` 的值赋值为 `"freeCodeCamp"` 来更改赋值。你仍然会看到 `true` 被记录到控制台，因为非空字串是一个真值。

# --hints--

你的 `truthyOrFalsy` 变量应具有值 `"freeCodeCamp"`。

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
