---
id: 690a06833d3e83b5226ade78
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

条件式用于根据条件运行代码。它由 `if` 关键字组成，后跟用括号括起来的条件，以及用花括弧括起来的体。下面是一个 `if` 语句的示例：

```js
if (condition) {
  console.log("condition is truthy");
}
```

当 `condition` 是真值（非空字串、非零数字、`true` 等）时，将执行 `if` 语句的体内的代码。

在你的变量下方，创建一个使用 `hasDeveloperJob` 作为条件的 `if` 语句。在你的 `if` 语句体内，将 `"Timmy is employed as a developer."` 记录到控制台。

# --hints--

你应该有一个 `if` 语句。

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

你的 `if` 语句应使用 `hasDeveloperJob` 作为条件。

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

你应该在你的 `if` 语句中将 `"Timmy is employed as a developer."` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
