---
id: 690a110a76d0496d9e7c28eb
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

声明另一个名为 `timmyAge` 的变量并将其设置为 `18`。

# --hints--

你应该有一个名为 `timmyAge` 的变量。

```js
assert.exists(timmyAge);
```

你应该将 `18` 赋值给你的 `timmyAge` 变量。

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
