---
id: 690a3ad44874a928a2fe3077
title: 步骤11
challengeType: 1
dashedName: step-11
---

# --description--

比较操作符允许你比较值并编写更复杂的条件式语句。

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

例如，上述代码仅在变量 `today` 和 `birthday` 的值相等时将 `"Happy Birthday!"` 记录到控制台。

创建第三个 `if` 语句。对于它的条件，使用 `>=` 操作符来查看 `timmyAge` 是否大于或相等于 `16`。

在你的新建 `if` 语句体内，将 `"Timmy is old enough to drive."` 记录到控制台。

# --hints--

你的代码中应该有第三个 `if` 语句。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

你的第三个 `if` 语句应使用 `timmyAge >= 16` 作为其条件。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

你应该在第三个 `if` 语句内将 `"Timmy is old enough to drive."` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
