---
id: 690a57fef66aac39dffc61a2
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

既然你知道什么是真值和假值，就从你的代码中删除变量声明和 `console.log` 语句。

# --hints--

你不应该有 `truthyOrFalsy` 变量。

```js
const explorer = await __helpers.Explorer(code);
assert.notExists(explorer.variables.truthyOrFalsy);
```

你的代码中不应有 `console.log(Boolean(truthyOrFalsy))`。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
