---
id: 690a0f7a5c22c15222a6aa5c
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

现在创建另一个使用 `isTimmyAGamer` 作为条件的 `if` 语句。在你的 `if` 语句体内，将 `"Timmy loves to play World of Warcraft."` 记录到控制台。

请注意，由于你的新建 `if` 语句的条件为假，你不会看到任何内容被记录到控制台。

# --hints--

你的代码中应该有两个 `if` 语句。

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

你的第二个 `if` 语句应使用 `isTimmyAGamer` 作为条件。

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

你应该在你的第二个 `if` 语句内将 `"Timmy loves to play World of Warcraft."` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
