---
id: 690a51806fb743bc896c3af7
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

正如你在之前的课程中学到的，布尔是一种只有 `true` 和 `false` 值的数据类型。

一个真值是指在布尔上下文中计算结果为 `true` 的值。真值的示例有：

- 非空字串
- 非零数字
- 布尔 `true`

以及你稍后会遇到的其他值。

你可以使用 `Boolean()` 函数来查看一个值的真值。例如，`Boolean("Hello World!")` 会返回 `true`，因为 `"Hello World!"` 是真值。

现在，调用 `Boolean(truthyOrFalsy)` 并将结果记录到控制台日志。

# --hints--

你应该将 `Boolean(truthyOrFalsy)` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
