---
id: 698c4f4e965bd853ad5084e1
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Abaixo do `h1`, adicione um elemento `h2` com o texto `Can't Stay Down`, este é o título da primeira música.

Abaixo do `h2` elemento adicione um `p` elemento com o texto `Artist: Quincy Larson`.

# --hints--

Você deve ter um elemento `h2`.

```js
assert.exists(document.querySelector('h2'));
```

O elemento `h2` deve estar abaixo do elemento `h1`.

```js
assert.exists(document.querySelector('h1 + h2'));
```

Seu `h2` elemento deve conter o texto `Can't Stay Down`.

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

Você deve ter um elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Seu `p` elemento deve conter o texto `Artist: Quincy Larson`.

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
