---
id: 698c58745b73289bb6c89ece
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

A última música se chama `Scratching the Surface`, o artista é Quincy Larson e o arquivo para usar é `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

Adicione a última música para completar o player de música seguindo o mesmo padrão que você usou para as duas primeiras músicas.

# --hints--

Você deve ter um terceiro `h2` elemento.

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

Seu terceiro `h2` elemento deve conter o texto `Scratching the Surface`.

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

Você deve ter um elemento `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

Seu `p` elemento deve conter o texto `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

Você deve ter um terceiro `audio` elemento.

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

O novo `audio` elemento deve estar abaixo do `p` elemento.

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

Seu `audio` elemento deve ter um `src` atributo.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

Seu elemento `audio` deve ter um atributo `src` com o valor `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

Seu `audio` elemento deve ter o atributo `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Seu `audio` elemento deve ter o atributo `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
