---
id: 698c4fe24f1f1a17be74a23e
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Em seguida, crie um `audio` elemento abaixo do `p` elemento. Nos próximos passos, você vai adicionar os atributos necessários para fazer o elemento de áudio tocar música. 

# --hints--

Você deve ter um `audio` elemento.

```js
assert.exists(document.querySelector('audio'));
```

O novo `audio` elemento deve estar abaixo do `p` elemento.

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
