---
id: 698c4e69abe14a9746c72312
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Nesta oficina, você construirá um player de música em HTML. O boilerplate HTML foi fornecido para você.

Crie um elemento `h1` e dê a ele o texto `freeCodeCamp Tunes`.

# --hints--

Você deve ter um elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Seu `h1` elemento deve conter o texto `freeCodeCamp Tunes`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'freeCodeCamp Tunes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
