---
id: 698c5229e47d57c6dfaa91bd
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

O `controls` atributo fornece controles de reprodução incluindo pausar/retomar a reprodução, busca e controle de volume para o `audio` elemento. 

O atributo `controls` é um atributo booleano e não precisa de um valor.

Adicione o atributo `controls` ao elemento `audio`.

Agora você deve ver o `audio` elemento exibido na página.

# --hints--

Seu `audio` elemento deve ter o atributo `controls`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
