---
id: 698c55e7aaacfae48a227459
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Adicione um novo `h2` elemento abaixo do `audio` elemento com o texto `Cruising for a Musing`.

Abaixo do `h2`, adicione um elemento `p` com o texto `Artist: Quincy Larson`.

# --hints--

Você deve ter um segundo `h2` elemento.

```js
assert.exists(document.querySelector('h2:nth-of-type(2)'));
```

Seu segundo `h2` elemento deve conter o texto `Cruising for a Musing`.

```js
const h2 = document.querySelector('h2:nth-of-type(2)');
assert.strictEqual(h2?.textContent.trim(), 'Cruising for a Musing');
```

Você deve ter um segundo elemento `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(2)'));
```

Seu segundo `p` elemento deve conter o texto `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(2)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
