---
id: 698c5710136539797f296098
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Abaixo do elemento `p`, adicione um elemento `audio`, dê a ele um atributo `src` com um valor de `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3` e os atributos `loop` e `controls`.

# --hints--

Você deve ter um segundo `audio` elemento.

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

O novo `audio` elemento deve estar abaixo do `p` elemento.

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

Seu `audio` elemento deve ter um `src` atributo.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

Seu elemento `audio` deve ter um atributo `src` com o valor `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

Seu `audio` elemento deve ter o atributo `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Seu `audio` elemento deve ter o atributo `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
