---
id: 698c50535b54b2d75fa985fb
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

O atributo `loop` reiniciará o áudio assim que a reprodução for concluída. Pense em um meme da internet que repete a reprodução. Omitir o atributo `loop` fará com que o áudio seja reproduzido uma vez.

O atributo `loop` é um atributo booleano e não precisa de um valor.

Adicione o atributo `loop` ao elemento `audio`.

# --hints--

Seu `audio` elemento deve ter o atributo `loop`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
