---
id: 69447054012a22f38c944ff0
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Aggiungi un'istruzione `if` all'interno della tua funzione `countup`.

Controlla se `number` è minore di `1` e restituisci un array vuoto.

Questo servirà come caso base per la tua ricorsione.

# --hints--

Quando `number` è minore di `1` la funzione dovrebbe restituire un array vuoto.

```js
assert.deepStrictEqual(countup(0), []);
```

Quando `number` **non** è minore di `1` la funzione **non dovrebbe** restituire nulla.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
