---
id: 69446f18fdcb11aa516f7c5c
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop creerai una `countup` funzione che genera un array con numeri da 1 a un numero dato usando la ricorsione.

Inizia creando una funzione vuota chiamata `countup` che accetta un solo parametro.

# --hints--

Dovresti avere una funzione chiamata `countup`.

```js
assert.isFunction(countup);
```

La tua funzione `countup` dovrebbe accettare un solo parametro.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
