---
id: 6948705160b66cf93a916e29
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Dichiara un array vuoto chiamato `countArray` all'interno della funzione `countup` appena sopra la tua istruzione `if`.

Questo array archivierà il valore di `number` a ogni chiamata della funzione.

# --hints--

Dovresti dichiarare un array vuoto chiamato `countArray`.

```js
const explorer = await __helpers.Explorer(code);
const { countup } = explorer.functions;
assert.exists(countup);
const { countArray } = countup.variables;
assert.exists(countArray);
assert.isTrue(countArray.value.matches('[]'));
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
