---
id: 694870bcf3c6874112b6f951
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Per tutti i numeri maggiori di e uguali a `1`, crea un blocco `else` e al suo interno chiama la funzione `countup` con `number - 1` e assegna il risultato a `countArray`. 

Poi usa `.push()` per aggiungere `number` a `countArray`, e infine restituisci `countArray`. 

Questo garantisce che la ricorsione raggiunga infine il caso base e termini.

# --hints--

Dovresti assegnare il risultato della chiamata a `countup` con `number - 1` a `countArray`.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

Dovresti inserire `number` in `countArray`.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

Dovresti restituire `countArray`.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
