---
id: 6948705160b66cf93a916e29
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Dichiara un array vuoto chiamato `countArray` all'interno della funzione `countup` appena sopra la tua istruzione `if`.

Questo array archivierà il valore di `number` a ogni chiamata della funzione.

# --hints--

Dovresti dichiarare un array vuoto chiamato `countArray`.

```js
assert.match(countup.toString(), /(let|const|var)\s+countArray\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
