---
id: 60a3e3396c7b40068ad6996c
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Katika mafundisho yaliyopita kuhusu mfano wa boksi wa CSS, ulijifunza kwamba kila kipengele cha HTML hutendewa kama boksi lenye maeneo manne.

Fikiria unapata boksi kutoka kwa muuzaji wako mtandaoni unayempenda — yaliyomo ni bidhaa ndani ya boksi, au katika kesi yako, kichwa, aya, au kipengele cha picha.

Badilisha sifa ya `src` katika `<img>` kutoka `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` hadi `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

Kipengele cha `img` kinapaswa kuwa na sifa ya `src` ya `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```
