---
id: 60a3e3396c7b40068ad69979
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Ongeza kipengele kipya cha `div` ndani ya kipengele chako cha `.canvas`.

Toa kipengele kipya cha `div` sifa ya kipengele `class` yenye thamani ya `one`. Hii itakuwa mstatili wako wa kwanza.

# --hints--

Unapaswa kuunda kipengele kipya cha `div`.

```js
assert(document.querySelectorAll('div').length === 3);
```

Unapaswa kupanga ndani kipengele kipya cha `div` ndani ya kipengele chako cha `.canvas`.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

`div` mpya inapaswa kuwa na sifa ya kipengele `class` yenye thamani `one`.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
