---
id: 60a3e3396c7b40068ad6996b
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Anza mradi wako kwa kuongeza kipengele cha `img` chenye `src` ya `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` ndani ya kipengele cha `body`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

Kipengele cha `img` kinapaswa kuwa na `src` ya `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
