---
id: 60a3e3396c7b40068ad69977
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Tumia `padding` kurekebisha nafasi ndani ya kipengele.

Katika `.frame`, tumia sifa ya ufupisho ya `padding` kuongeza nafasi kati ya vipengele vya `.frame` na `.canvas` kwa `50px`. Sifa ya ufupisho itaongeza nafasi juu, chini, kushoto, na kulia ya mpaka na turubai ya kuchora ndani ya kipengele.

# --hints--

Unapaswa kuweka sifa ya `padding` kuwa `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

Kipengele chako cha `.frame` kinapaswa kuwa na thamani ya `padding` ya `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
