---
id: 60a3e3396c7b40068ad69980
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Ongeza `div` nyingine yenye thamani ya `class` ya `two` chini kabisa ya kipengele chako cha `one`. Hii itakuwa mstatili wako wa pili.

# --hints--

Haupaswi kubadilisha kipengele kilichopo cha `.one`.

```js
assert(document.querySelectorAll('.one').length === 1);
```

Unapaswa kuwa na kipengele cha `div` cha pili ndani ya kipengele chako cha `.canvas`.

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

Kipengele chako cha `div` cha pili kinapaswa kuwa na thamani ya `class` ya `two`.

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

Kipengele chako cha `.two` kinapaswa kuja baada ya kipengele chako cha `.one`.

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
