---
id: 646ceb843412c74edee27a79
title: Крок 22
challengeType: 0
dashedName: step-22
---

# --description--

Зараз ви навчитеся трюку CSS, щоб малювати трикутники.

Його буде використано, щоб намалювати вуха кота.

Використовуючи селектор класу, надайте елементу `.cat-right-ear` властивості `height` та `width` зі значенням `100px`. Встановіть `background-color` на `white`. Встановіть лівий та правий кордони на `35px solid blue`. Встановіть верхній та нижній кордони на `35px solid red`.

# --hints--

Ви повинні мати селектор `.cat-right-ear`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Селектор `.cat-right-ear` повинен мати властивість `height` зі значенням `100px`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height === '100px')
```

Селектор `.cat-right-ear` повинен мати властивість `width` зі значенням `100px`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '100px')
```

Селектор `.cat-right-ear` повинен мати властивість `background-color` зі значенням `white`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor === 'white')
```

Селектор `.cat-right-ear` повинен мати властивість `border-left` зі значенням `35px solid blue`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid blue')
```

Селектор `.cat-right-ear` повинен мати властивість `border-right` зі значенням `35px solid blue`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid blue')
```

Селектор `.cat-right-ear` повинен мати властивість `border-top` зі значенням `35px solid red`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid red')
```

Селектор `.cat-right-ear` повинен мати властивість `border-bottom` зі значенням `35px solid red`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
