---
id: 6676a8a8d9861319ee901a90
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Тепер ви можете почати працювати над вухами кота.

Приберіть код огляду, видаливши селектор `.cat-right-ear` та всі його властивості.

Надайте елементу `.cat-left-ear` лівий та правий кордон зі значенням `35px solid transparent` за допомогою селектора класу. Також встановіть нижній кордон на `70px solid #5e5e5e`.

# --hints--


Селектор `.cat-right-ear` не повинен мати властивості `height`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
```

Селектор `.cat-right-ear` не повинен мати властивості `width`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
```


Селектор `.cat-right-ear` не повинен мати властивості `border-left`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
```

Селектор `.cat-right-ear` не повинен мати властивості `border-right`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
```

Селектор `.cat-right-ear` не повинен мати властивості `border-top`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
```

Селектор `.cat-right-ear` не повинен мати властивості `border-bottom`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
```

Ви не повинні мати селектор `.cat-right-ear`.

```js
assert(!new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Ви повинні мати селектор `.cat-left-ear`.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

Селектор `.cat-left-ear` повинен мати властивість `border-left` зі значенням `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

Селектор `.cat-left-ear` повинен мати властивість `border-right` зі значенням `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

Селектор `.cat-left-ear` повинен мати властивість `border-bottom` зі значенням `70px solid #5e5e5e`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
