---
id: 6823f9df49cc206af5471a30
title: Depurar la página de perfil de Camperbot
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot está tratando de construir una página de perfil. Pidieron a un amigo que revisara su código y dijeron que tiene algunos errores.

Su trabajo es corregir todos los errores de Camperbot para que puedan continuar construyendo su página de perfil. Complete los elementos en las historias de usuario a continuación y haga clic en "Ejecutar las Pruebas" para ver si se corrigieron todos los errores.

**Historias de usuario:**

1. Camperbot está tratando de usar un elemento `heading2`, pero ese elemento no existe. Corrija esas etiquetas para que utilice un elemento de encabezado de segundo nivel correcto.
2. Camperbot está tratando de agregar dos párrafos con `pp`, pero esos tampoco existen. Corríjalos para que usen etiquetas de párrafo correctas.
3. Camperbot está usando un elemento `h3` para el subencabezado `Background and Interests`, pero tiene un error de sintaxis. Identifique el problema y resuélvalo.

# --hints--

Debe tener exactamente un elemento `h1` en la página.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

No debes eliminar el elemento `<h1>Hello from Camperbot!</h1>` de la página.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

Debe tener exactamente un elemento `h2` en la página.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

Debes tener un elemento `h2` con el texto `About`. Aquí tienes un ejemplo: `<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

Debes tener un elemento de párrafo con el texto `My name is Camperbot and I love learning new things.` Aquí tienes un ejemplo: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

Debes tener un elemento de párrafo con el texto `I enjoy solving puzzles.` Aquí tienes un ejemplo: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

Debe tener exactamente dos elementos de párrafo en la página.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Debe tener exactamente un elemento `h3` en la página.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

Debes tener un elemento `h3` con el texto `Background and Interests`. Aquí tienes un ejemplo: `<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

Debe tener exactamente cinco elementos en total en la página.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
