---
id: 68720c179cd34be6793bfdf6
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

애플리케이션이 예상대로 작동하지만, 보통 이렇게 여러 `Card` 컴포넌트를 수동으로 추가하지는 않습니다.

`<Card />` 요소 안에 있는 세 개의 `<div className="flex-container">` 컴포넌트를 제거하세요. 다음 몇 단계에서 데이터 배열을 만들고 `map()` 함수를 사용해 `Card` 컴포넌트를 렌더링할 것입니다. 이것이 React에서 더 일반적인 방법입니다.

# --hints--

코드에서 세 개의 `<Card />` 컴포넌트를 제거해야 합니다.

```js
const testElem = await __helpers.prepTestComponent(window.index.App);
const cards = testElem.querySelectorAll(".card");
assert.lengthOf(cards, 0);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  return (
    <div className="flex-container">
--fcc-editable-region--
      <Card
        name="Mark"
        title="Front-End developer"
        bio="I like to work with different front-end technologies and play video games."
      />
      <Card
        name="Tiffany"
        title="Engineering manager"
        bio="I have worked in tech for 15 years and love to help people grow in this industry."
      />
      <Card
        name="Doug"
        title="Back-End developer"
        bio="I have been a software developer for over 20 years and I love working with Go and Rust."
      />
--fcc-editable-region--
    </div>
  );
}
```
