---
id: 68778af56eb5d980a4e2daed
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

각 개별 `Card` 컴포넌트에 프로프스를 설정하는 대신, 프로필 데이터를 담은 배열을 만들고 그 데이터를 사용해 `Card` 컴포넌트를 동적으로 렌더링하세요.

`profiles` 컴포넌트 안에 `App` 배열을 만드는 것부터 시작하세요.

다음 값을 가진 객체 하나를 `Mark` 배열에 추가하세요:

- `id`: `1`
- `name`: `"Mark"`
- `title`: `"Front-End developer"`
- `bio`: `"I like to work with different front-end technologies and play video games."`

# --hints--

`profiles` 컴포넌트 안에 `App` 배열을 정의해야 합니다.

```js
assert.match(code, /(const|let)\s+profiles\s*=\s*\[/);
```

프로필 배열에는 `id`, `name`, `title`, `bio` 모든 필수 프로퍼티를 포함하는 객체가 있어야 합니다.

```js
assert.match(
  code,
  /(const|let)\s+profiles\s*=\s*\[\s*{[^}]*id\s*:\s*\d+[^}]*name\s*:\s*["'`][^"'`]+["'`][^}]*title\s*:\s*["'`][^"'`]+["'`][^}]*bio\s*:\s*["'`][^"'`]+["'`][^}]*}\s*\]/s);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  --fcc-editable-region--
     
  --fcc-editable-region--
  return (
    <div className="flex-container"></div>
  );
}
```
