---
id: 687799aefd1f2af1704065d0
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

이제 페이지에 카드 하나가 표시될 것입니다. 하지만 `Card` 컴포넌트 안에 빠진 부분이 있습니다.

이전 수업에서 기억하듯이, 이런 식으로 컴포넌트 목록을 렌더링할 때는 `key` 프로프를 포함하는 것이 모범 사례입니다.

```jsx
<div>
  {fruits.map(fruit => (
    <Card
      key={fruit.id}
      name={fruit.name}
      color={fruit.color}
    />
  ))}
</div>
```

`key` 프로프는 React가 어떤 항목이 변경되었고, 추가되었으며, 제거되었는지 식별하는 데 도움을 줍니다. 이는 성능을 향상시키고 더 효율적인 업데이트를 가능하게 합니다.

`Card` 컴포넌트를 업데이트하여 `key` 값을 가진 `profile.id` 프로프를 포함하세요.

# --hints--

`key` 컴포넌트에 `Card` 프로프를 추가해야 합니다. 값은 `profile.id` 이어야 합니다.

```js
assert.match(code, /<Card[^>]*key=\{\s*profile\.id\s*\}[^>]*\/?>/);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  const profiles = [
    {
      id: 1,
      name: "Mark",
      title: "Front-End developer",
      bio: "I like to work with different front-end technologies and play video games."
    }
  ];
  return (
    <div className="flex-container">
      --fcc-editable-region--
      {profiles.map((profile) => (
        <Card
          name={profile.name}
          title={profile.title}
          bio={profile.bio}
        />
      ))}
      --fcc-editable-region--
    </div>
  );
}
```
