---
id: 67d83df6f82eda3868dd2a84
title: Erstellen Sie einen RPG-Charakter
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

In diesem Labor üben Sie die Grundlagen von Python, indem Sie eine kleine App erstellen, die einen Charakter für ein RPG-Abenteuer anlegt.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `create_character` haben.
1. Die Funktion sollte nacheinander einen Charakternamen gefolgt von drei Werten annehmen: Stärke, Intelligenz und Charisma.
1. Der Charakternamen sollte validiert werden:
   - Wenn der Charakternamen kein String ist, sollte die Funktion `The character name should be a string` zurückgeben.
   - Wenn der Charakternamen ein leerer String ist, sollte die Funktion `The character should have a name` zurückgeben.
   - Wenn der Charakternamen länger als 10 Zeichen ist, sollte die Funktion `The character name is too long` zurückgeben.
   - Wenn der Charakternamen Leerzeichen enthält, sollte die Funktion `The character name should not contain spaces` zurückgeben.
1. Die Werte sollten ebenfalls validiert werden:
   - Wenn einer oder mehrere Werte keine ganzen Zahlen sind, sollte die Funktion `All stats should be integers` zurückgeben.
   - Wenn einer oder mehrere Werte kleiner als 1 sind, sollte die Funktion `All stats should be no less than 1` zurückgeben.
   - Wenn einer oder mehrere Werte größer als 4 sind, sollte die Funktion `All stats should be no more than 4` zurückgeben.
   - Wenn die Summe aller Werte nicht 7 ergibt, sollte die Funktion `The character should start with 7 points` zurückgeben.
1. Wenn alle Werte die Überprüfung bestehen, sollte die Funktion einen String mit vier Zeilen zurückgeben:
   - Die erste Zeile sollte den Charakternamen enthalten
   - Die Zeilen 2–4 sollten mit der Abkürzung der Werte `STR`, `INT` oder `CHA` (in dieser Reihenfolge) beginnen, dann ein Leerzeichen, gefolgt von einer Anzahl voller Punkte (`●`) entsprechend dem Wert, und einer Anzahl leerer Punkte (`○`), um auf 10 zu kommen. Beispiel: Wenn der Wert der Stärke 3 ist, müssen 3 volle Punkte gefolgt von 7 leeren Punkten angezeigt werden. Die Punkte sind im Editor vorgegeben.

Hier ist der String, der von `create_character('ren', 4, 2, 1)` zurückgegeben werden sollte:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

HINWEIS: Während `str` und `int` gängige Abkürzungen für die Werte sind, denken Sie daran, dass diese reservierte Schlüsselwörter in Python sind und nicht als Variablennamen verwendet werden sollten.

# --hints--

Sie sollten eine Funktion mit dem Namen `create_character` haben.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

Wenn `create_character` mit einem ersten Argument aufgerufen wird, das **kein** String ist, sollte es `The character name should be a string` zurückgeben.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

Wenn `create_character` mit einem ersten Argument aufgerufen wird, das ein String ist, sollte es **nicht** `The character name should be a string` zurückgeben.

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

Wenn `create_character` mit einem ersten Argument aufgerufen wird, das ein leerer String ist, sollte es `The character should have a name` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

Wenn `create_character` mit einem ersten Argument aufgerufen wird, das **kein** leerer String ist, sollte es **nicht** `The character should have a name` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

Wenn `create_character` mit einem ersten Argument aufgerufen wird, das länger als 10 Zeichen ist, sollte es `The character name is too long` zurückgeben.

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

Die Funktion `create_character` sollte nicht sagen, dass der Charakter zu lang ist, wenn er nicht länger als 10 Zeichen ist.

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

Wenn `create_character` mit einem ersten Argument aufgerufen wird, das ein Leerzeichen enthält, sollte es `The character name should not contain spaces` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

Wenn `create_character` mit einem ersten Argument aufgerufen wird, das *kein* Leerzeichen enthält, sollte es *nicht* `The character name should not contain spaces` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

Wenn `create_character` mit einem zweiten, dritten oder vierten Argument aufgerufen wird, das **keine** ganze Zahl ist, sollte es `All stats should be integers` zurückgeben.

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

Wenn `create_character` mit einem zweiten, dritten und vierten Argument aufgerufen wird, die alle ganze Zahlen sind, sollte es **nicht** `All stats should be integers` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

Wenn `create_character` mit einem zweiten, dritten oder vierten Argument aufgerufen wird, das kleiner als `1` ist, sollte es `All stats should be no less than 1` zurückgeben.

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

Wenn `create_character` mit einem zweiten, dritten und vierten Argument aufgerufen wird, die alle mindestens `1` sind, sollte es **nicht** `All stats should be no less than 1` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

Wenn `create_character` mit einem zweiten, dritten oder vierten Argument aufgerufen wird, das größer als `4` ist, sollte es `All stats should be no more than 4` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

Wenn `create_character` mit einem zweiten, dritten und vierten Argument aufgerufen wird, die alle höchstens `4` sind, sollte es **nicht** `All stats should be no more than 4` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

Wenn `create_character` mit einem zweiten, dritten oder vierten Argument aufgerufen wird, die **nicht** zu `7` summieren, sollte es `The character should start with 7 points` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

Wenn `create_character` mit einem zweiten, dritten und vierten Argument aufgerufen wird, die zu `7` summieren, sollte es **nicht** `The character should start with 7 points` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` sollte `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○` zurückgeben.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

Wenn `create_character` mit gültigen Werten aufgerufen wird, sollte es die Charakterwerte wie gefordert ausgeben.

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
