---
id: 66e45c8140f9fda5c105ae26
title: Construir una Página de Volantes de Evento
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

Cumpla con las historias de usuario a continuación y haga que todas las pruebas pasen para completar el laboratorio.

**Instrucciones:**

1. Debería tener un elemento `header` dentro del cuerpo.
1. Su elemento `header` debería tener una imagen para su evento y un `h1` con el título del evento, en ese orden. Puede utilizar esta imagen, si lo desea: `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. Debería tener un elemento `main` dentro del cuerpo.
1. Su elemento `main` debería tener al menos dos elementos `section` dentro que muestren las características del evento.
1. Cada uno de sus elementos `section` debería tener un `h2` dentro de ellos.
1. Su cuerpo debería tener un relleno superior e inferior de `50px`.
1. Su cuerpo debería tener un margen superior e inferior de `0` y un margen izquierdo y derecho centrados.
1. Su cuerpo debería tener un ancho establecido en relación con el ancho de la vista.
1. Su cuerpo debería usar la función `calc` para establecer su propiedad `min-height` al 100% de la altura de la vista menos todo el relleno aplicado a la parte superior e inferior del cuerpo.
1. Debería tener al menos un `hr` dentro de su volante.
1. Debería establecer el `ancho` de todos los elementos `hr` y `section` a un valor en porcentaje relativo a su elemento padre.

**Nota:** Asegúrate de enlazar tu hoja de estilos en tu HTML y aplicar tu CSS.

# --hints--

Debería tener un elemento `header` dentro del elemento `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Debería tener un elemento `img` como primer hijo en el `header`.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

Debería tener un elemento `h1` como segundo hijo en el `header`.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

Solo Debes tener un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Su `h1` debería tener el título de su evento.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

Debería tener un elemento `main` dentro del elemento `body`.

```js
assert.exists(document.querySelector('body > main'));
```

Debería tener al menos dos elementos `section` dentro de su elemento `main`.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

Sus elementos `section` deberían tener cada uno un `h2` dentro de ellos.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

Sus elementos `h2` no deberían estar vacíos.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

Su elemento `body` debería tener un relleno superior e inferior de `50px`.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

Su elemento `body` debería tener un margen superior e inferior de `0`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

Su elemento `body` debería tener un margen izquierdo y derecho de `auto`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

Su elemento `body` debería establecer su ancho utilizando `vw`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

Su `body` debería utilizar `calc` para establecer su `min-height` a `100vh - 100px`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

Debería tener al menos un elemento `hr`.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

El ancho de tus elementos `hr` debería establecerse utilizando un valor en porcentaje.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

El ancho de tus elementos `section` debería establecerse utilizando un valor en porcentaje.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
