---
id: 6733b5632fb854bc6c58dc0d
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Erstellen Sie nun ein weiteres Objekt mit denselben drei Eigenschaften, um eine andere Sorte spanischer Lavendel namens `Pretty Polly` darzustellen, und weisen Sie es einer Variablen namens `prettyPolly` zu.

# --hints--

Sie sollten eine Variable namens `prettyPolly` haben.

```js
assert.exists(prettyPolly);
```

Sie sollten ein Objekt Ihrer `prettyPolly`-Variablen zuweisen.

```js
assert.isObject(prettyPolly);
```

Ihr `prettyPolly`-Objekt sollte eine `commonName`-Eigenschaft mit dem Wert `"Spanish lavender"` haben.

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

Ihr `prettyPolly`-Objekt sollte eine `scientificName`-Eigenschaft mit dem Wert `"Lavandula stoechas"` haben.

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

Ihr `prettyPolly`-Objekt sollte eine `cultivar`-Eigenschaft mit dem Wert `"Pretty Polly"` haben.

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

Ihr `prettyPolly`-Objekt sollte drei SchlĂĽssel haben.

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

--fcc-editable-region--

--fcc-editable-region--
```
