---
id: 6733b5e6f210c1bd1afdf8b3
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Die restlichen Pflanzenobjekte wurden für Sie hinzugefügt. Jetzt ist es Zeit, Ihren Katalog zu erstellen.

Wie reguläre JavaScript-Objekte enthalten `Map`-Objekte eine Sammlung von Schlüssel-Wert-Paaren. Um ein `Map`-Objekt zu erstellen, können Sie den `Map`-Konstruktor verwenden:

```js
const myMap = new Map();
```

Deklarieren Sie eine `catalog`-Variable und weisen Sie ihr ein `Map`-Objekt zu.

# --hints--

Sie sollten eine Variable mit dem Namen `catalog` haben.

```js
assert.exists(catalog);
```

Sie sollten ein `Map`-Objekt Ihrer `catalog`-Variable zuweisen.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

--fcc-editable-region--

--fcc-editable-region--
```
