---
id: 673339d96c56d36e963e888e
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop erstellen Sie einen Katalog fĂĽr eine Pflanzenzucht und ĂĽben dabei die JavaScript-Objekte Map und Set.

Um jede Pflanze in Ihrem Katalog darzustellen, verwenden Sie Objekte, die jeweils drei Eigenschaften haben:

- `commonName`, der gebrĂ¤uchliche Name der Pflanze
- `scientificName`, der wissenschaftliche Name (oder lateinische Name) der Pflanze
- `cultivar`, die Sorte der Pflanze

Beginnen Sie damit, ein Objekt mit den drei oben genannten Eigenschaften fĂĽr die `Spanish lavender` (wissenschaftlicher Name `Lavandula stoechas`) einer Sorte namens `Ballerina` zu erstellen und weisen Sie es einer Variablen namens `ballerina` zu.

# --hints--

Sie sollten eine Variable namens `ballerina` haben.

```js
assert.exists(ballerina);
```

Sie sollten Ihrer Variablen `ballerina` ein Objekt zuweisen.

```js
assert.isObject(ballerina);
```

Ihr `ballerina`-Objekt sollte eine `commonName`-Eigenschaft mit dem Wert `"Spanish lavender"` haben.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

Ihr `ballerina`-Objekt sollte eine `scientificName`-Eigenschaft mit dem Wert `"Lavandula stoechas"` haben.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

Ihr `ballerina`-Objekt sollte eine `cultivar`-Eigenschaft mit dem Wert `"Ballerina"` haben.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

Ihr `ballerina`-Objekt sollte drei SchlĂĽssel haben.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
