---
id: 6734a606e5968a5a19b77171
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

FĂĽgen Sie die restlichen Pflanzen zu Ihrem Katalog hinzu, wie Sie es in den vorherigen Schritten getan haben.

# --hints--

Sie sollten Ihr `willowVale`-Objekt als SchlĂĽssel Ihrer `catalog`-Map festlegen.

```js
assert.isTrue(catalog.has(willowVale));
```

Der Wert des `willowVale`-SchlĂĽssels in der `catalog`-Map sollte ein Objekt mit den Eigenschaften `small`, `medium` und `large` sein, jeweils mit dem Wert einer positiven ganzen Zahl.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Sie sollten Ihr `hidcote`-Objekt als SchlĂĽssel Ihrer `catalog`-Map festlegen.

```js
assert.isTrue(catalog.has(hidcote));
```

Der Wert des `hidcote`-SchlĂĽssels in der `catalog`-Map sollte ein Objekt mit den Eigenschaften `small`, `medium` und `large` sein, jeweils mit dem Wert einer positiven ganzen Zahl.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Sie sollten Ihr `imperialGem`-Objekt als SchlĂĽssel Ihrer `catalog`-Map festlegen.

```js
assert.isTrue(catalog.has(imperialGem));
```

Der Wert des `imperialGem`-SchlĂĽssels in der `catalog`-Map sollte ein Objekt mit den Eigenschaften `small`, `medium` und `large` sein, jeweils mit dem Wert einer positiven ganzen Zahl.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Sie sollten Ihr `royalCrown`-Objekt als SchlĂĽssel Ihrer `catalog`-Map festlegen.

```js
assert.isTrue(catalog.has(royalCrown));
```

Der Wert des `royalCrown`-SchlĂĽssels in der `catalog`-Map sollte ein Objekt mit den Eigenschaften `small`, `medium` und `large` sein, jeweils mit dem Wert einer positiven ganzen Zahl.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Sie sollten weiterhin den `ballerina`-SchlĂĽssel in Ihrer `catalog`-Map haben.

```js
assert.isTrue(catalog.has(ballerina));
```

Sie sollten weiterhin den `prettyPolly`-SchlĂĽssel in Ihrer `catalog`-Map haben.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Ihre `catalog`-Map sollte sechs Elemente enthalten.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
