---
id: 6734a2385c52c54d916a8b86
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Fügen Sie Ihrem Katalog ein weiteres Schlüssel-Wert-Paar hinzu. Verwenden Sie Ihre `prettyPolly`-Variable als Schlüssel und machen Sie den Wert zu einem Objekt mit den Eigenschaften `small`, `medium` und `large`, deren Werte wieder ganze Zahlen sind.

# --hints--

Sie sollten Ihr `prettyPolly`-Objekt als Schlüssel Ihrer `catalog`-Map setzen.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Der Wert des `prettyPolly`-Schlüssels in der `catalog`-Map sollte ein Objekt mit den Eigenschaften `small`, `medium` und `large` sein.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

Die `small`-Eigenschaft des Objekts, das mit dem `prettyPolly`-Schlüssel in der `catalog`-Map verknüpft ist, sollte eine positive ganze Zahl sein.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

Die `medium`-Eigenschaft des Objekts, das mit dem `prettyPolly`-Schlüssel in der `catalog`-Map verknüpft ist, sollte eine positive ganze Zahl sein.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

Die `large`-Eigenschaft des Objekts, das mit dem `prettyPolly`-Schlüssel in der `catalog`-Map verknüpft ist, sollte eine positive ganze Zahl sein.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Sie sollten weiterhin den `ballerina`-Schlüssel in Ihrer `catalog`-Map haben.

```js
assert.isTrue(catalog.has(ballerina));
```

Ihre `catalog`-Map sollte zwei Elemente enthalten.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
