---
id: 67349ba230cff237c0f2bdf2
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Map-Objekte merken sich die Reihenfolge, in der die Schlüssel in die Map eingefügt werden. Um ein Schlüssel-Wert-Paar in Ihrem Map-Objekt zu setzen, können Sie die `set`-Methode verwenden. Das Folgende setzt den Schlüssel `"name"` auf den Wert `"Jerry"` im `myMap`-Map-Objekt.

```js
myMap.set("name", "Jerry");
```

Setzen Sie den Schlüssel `"ballerina"` auf den Wert Ihres `ballerina`-Objekts in Ihrer `catalog`-Map. Geben Sie dann Ihren Katalog in der Konsole aus.

Beachten Sie, dass Sie die Browser-Konsole öffnen müssen, um Ihr Map-Objekt richtig zu sehen.

# --hints--

Sie sollten den Schlüssel `"ballerina"` auf den Wert Ihres `ballerina`-Objekts in Ihrer `catalog`-Map setzen.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

Ihre `catalog`-Map sollte genau ein Schlüssel-Wert-Paar enthalten.

```js
assert.lengthOf(catalog, 1);
```

Sie sollten Ihre `catalog`-Map in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
--fcc-editable-region--

--fcc-editable-region--
```
