---
id: 6734c10a327760a665b7d5b9
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

In den nächsten Schritten arbeiten Sie an einer Funktion, mit der Sie den Katalog aktualisieren können, wenn Sie Pflanzen verkaufen.

Erstellen Sie eine Funktion namens `sellPlants`, die drei Argumente entgegennimmt: ein Pflanzenobjekt, die Topfgröße und die Anzahl der zu verkaufenden Töpfe, in dieser Reihenfolge.

Wenn die angegebene Anzahl der Töpfe größer ist als die verfügbaren Töpfe, soll Ihre Funktion `Not enough <pot-size> size pots for <item-name>. Only <pot-number> left.` zurückgeben. Ersetzen Sie `<pot-size>` durch die Topfgröße, `<pot-number>` durch die verbleibenden Töpfe dieser Größe und `<item-name>` durch den wissenschaftlichen Namen der Pflanze, gefolgt von einem Leerzeichen und dem Namen der Sorte in einfachen Anführungszeichen.

Um zu testen, ob alles funktioniert, geben Sie das Ergebnis des Aufrufs Ihrer `sellPlants`-Funktion mit dem `ballerina`-Objekt, `"small"` und `25` in der Konsole aus.

# --hints--

Sie sollten Ihre `catalog`-Map nicht verändern.

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

Sie sollten eine Funktion namens `sellPlants` haben.

```js
assert.isFunction(sellPlants);
```

Ihre `sellPlants`-Funktion sollte drei Argumente entgegennehmen.

```js
assert.lengthOf(sellPlants, 3);
```

`sellPlants(ballerina, "small", 25)` sollte `Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.` zurückgeben.

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

Sie sollten `sellPlants(ballerina, "small", 25)` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)` sollte `undefined` zurückgeben.

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog)
```
