---
id: 6734f6290ebd5cfbc75996af
title: Schritt 16
challengeType: 1
dashedName: step-16
---

# --description--

Ändern Sie jetzt Ihren `removePlant`-Aufruf, indem Sie ihm das `ballerina`-Objekt übergeben, und schauen Sie in die Konsole. Sie werden einen Fehler sehen, den Sie im nächsten Schritt beheben werden.

# --hints--

Sie sollten `removePlant(ballerina)` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*removePlant\s*\(\s*ballerina\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const removePlant = plant => catalog.delete(plant);
--fcc-editable-region--
console.log(removePlant("ballerina"));
--fcc-editable-region--
const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```
