---
id: 6734fa225c6667121e589f7e
title: Schritt 17
challengeType: 1
dashedName: step-17
---

# --description--

Ihre `sellPlants`-Funktion löst einen `TypeError` aus, weil sie versucht, eine Eigenschaft eines Objekts zu lesen, das aus dem Katalog gelöscht wurde. Daher müssen Sie den Fall berücksichtigen, in dem ein Schlüssel in Ihrer `catalog`-Map nicht gefunden wird. Dafür ist die `has`-Methode sehr nützlich. Sie gibt einen Boolean zurück, der angibt, ob der angegebene Schlüssel in der Map vorhanden ist oder nicht.

Ändern Sie Ihre `sellPlants`-Funktion so, dass sie `Item not found.` zurückgibt, wenn die Pflanze, die an die Funktion übergeben wird, nicht im Katalog enthalten ist.

# --hints--

Ihre `sellPlants`-Funktion sollte die `has`-Methode auf `catalog` aufrufen.

```js
let flag = false;
const temp = catalog.has;
try {
    catalog.has = (p) => {flag = true; return false;}
    sellPlants("test", "", 0);
    sellPlants(imperialGem, "medium", 200);
    assert.isTrue(flag);
} finally {
    catalog.has = temp;
}
```

Sie sollten die in Ihrer `catalog`-Map eingefügten Schlüssel-Wert-Paare nicht direkt ändern.

```js
const testCatalog = new Map();
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

`sellPlants(ballerina, "small", 10)` sollte `Item not found.` zurückgeben

```js
assert.equal(sellPlants(ballerina, "small", 10), "Item not found.")
```

`sellPlants(hidcote, "large", 95)` sollte `Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.` zurückgeben

```js
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")
assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

`sellPlants(hidcote, "large", 5)` sollte `Catalog successfully updated.` zurückgeben und `5` von der `large`-Eigenschaft des Objekts subtrahieren, das mit dem `hidcote`-Schlüssel im Katalog verknüpft ist.

```js
try {
    assert.equal(sellPlants(hidcote, "large", 5), "Catalog successfully updated.");
    assert.strictEqual(catalog.get(hidcote).large, 13);

    assert.equal(sellPlants(royalCrown, "medium", 5), "Catalog successfully updated.");
    assert.strictEqual(catalog.get(royalCrown).medium, 17);
} finally {
    catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
}
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const sellPlants = (plant, size, potsNo) => {
--fcc-editable-region--
    
--fcc-editable-region--
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

const removePlant = plant => catalog.delete(plant);
console.log(removePlant(ballerina));

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```
