---
id: 6734a13a29acdf47f2b8e161
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Im Gegensatz zu regulären Objekten, bei denen Schlüssel immer Strings sein müssen, erlauben Map-Objekte die Verwendung beliebiger Datentypen als Schlüssel. In Ihrem Katalog werden Sie Objekt-Objekt-Schlüssel-Wert-Paare verwenden, wobei jedes Ihrer Pflanzenobjekte mit einem Objekt verknüpft wird, das die Anzahl verfügbarer Töpfe für verschiedene Größen enthält.

Ändern Sie Ihren `set`-Aufruf so, dass Ihr `ballerina`-Objekt als Schlüssel verwendet wird. Für den Wert verwenden Sie ein Objekt mit drei Eigenschaften, `small`, `medium` und `large`, die jeweils einen positiven ganzzahligen Wert haben.

# --hints--

Sie sollten Ihr `ballerina`-Objekt als Schlüssel Ihrer `catalog`-Map festlegen.

```js
assert.isTrue(catalog.has(ballerina));
```

Der Wert des `ballerina`-Schlüssels in der `catalog`-Map sollte ein Objekt mit den Eigenschaften `small`, `medium` und `large` sein.

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

Die `small`-Eigenschaft des Objekts, das mit dem `ballerina`-Schlüssel in der `catalog`-Map verknüpft ist, sollte eine positive ganze Zahl sein.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

Die `medium`-Eigenschaft des Objekts, das mit dem `ballerina`-Schlüssel in der `catalog`-Map verknüpft ist, sollte eine positive ganze Zahl sein.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

Die `large`-Eigenschaft des Objekts, das mit dem `ballerina`-Schlüssel in der `catalog`-Map verknüpft ist, sollte eine positive ganze Zahl sein.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Ihre `catalog`-Map sollte genau ein Schlüssel-Wert-Paar enthalten.

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
--fcc-editable-region--
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
