---
id: 6734f0ae9d1cf4dfd43a7805
title: Schritt 15
challengeType: 1
dashedName: step-15
---

# --description--

Um ein Element aus einem Map-Objekt zu entfernen, können Sie die `delete`-Methode verwenden, die einen Schlüssel als Argument nimmt und den Eintrag entfernt, der mit diesem Schlüssel verknüpft ist. Sie gibt einen Boolean zurück, der angibt, ob das Entfernen erfolgreich war.

Erstellen Sie eine Funktion namens `removePlant`, die ein Pflanzenobjekt als Argument nimmt und es aus dem Katalog entfernt. Lassen Sie sie entweder `true` oder `false` zurückgeben, je nachdem, ob die angegebene Pflanze aus dem Katalog gelöscht wurde.

Protokollieren Sie danach das Ergebnis des Aufrufs Ihrer neuen Funktion mit dem String `"ballerina"` als Argument.

# --hints--

Sie sollten eine Funktion namens `removePlant` haben.

```js
assert.isFunction(removePlant);
```

Ihre `removePlant`-Funktion sollte ein einzelnes Argument annehmen.

```js
assert.lengthOf(removePlant, 1);
```

Ihre `removePlant`-Funktion sollte die `delete`-Methode auf `catalog` aufrufen.

```js
let flag = false;
const temp = catalog.delete;
try {
    catalog.delete = (p) => {flag = true}
    removePlant("");
    assert.isTrue(flag);
} finally {
    catalog.delete = temp;
}
```

`removePlant(item)` sollte `item` aus Ihrem Katalog löschen.

```js
const testArr = [];
const temp = catalog.delete;
try {
    catalog.delete = (p) => {testArr.push(p)}
    removePlant("test");
    assert.deepEqual(testArr, ["test"]);
} finally {
    catalog.delete = temp;
}
```

Ihre `removePlant`-Funktion sollte einen Boolean zurückgeben, der angibt, ob das Entfernen erfolgreich war.

```js
try {
    assert.isFalse(removePlant("test"));
    assert.isTrue(removePlant(royalCrown));
} finally {
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 })
}
```

Sie sollten `removePlant("ballerina")` in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*removePlant\s*\(\s*("|')ballerina\1\s*\)\s*\)/);
```



# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

--fcc-editable-region--

--fcc-editable-region--

const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```
