---
id: 6734a772196245601e4b80e0
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Wenn Sie ein Element aus einer Map abrufen müssen, können Sie die `get`-Methode verwenden. Sie nimmt einen Schlüssel als Argument und gibt den zugehörigen Wert zurück.

Verwenden Sie die `get`-Methode, um den Wert des Schlüssels `ballerina` aus Ihrem Katalog abzurufen und weisen Sie ihn einer Variablen namens `getBallerina` zu. Geben Sie dann Ihre neue Variable in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `getBallerina` haben.

```js
assert.exists(getBallerina);
```

Ihre `getBallerina`-Variable sollte den Wert von `catalog.get(ballerina)` haben.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+getBallerina\s*=\s*catalog\s*\.\s*get\s*\(\s*ballerina\s*\)/);
```

Sie sollten Ihre `getBallerina`-Variable in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*getBallerina\s*\)/);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
```
