---
id: 67350acfd2006f2777295b96
title: Schritt 19
challengeType: 1
dashedName: step-19
---

# --description--

Die `keys`-Methode ist direkt auf dem `Map`-Objekt verfügbar und gibt ein `MapIterator` zurück, mit dem Sie die Schlüssel Ihrer Map durchlaufen können.

Erstellen Sie eine Funktion namens `displayCatalog`. Erstellen Sie innerhalb Ihrer Funktion eine `for...of`-Schleife, die die Schlüssel in `catalog` durchläuft und jeden davon in der Konsole ausgibt.

Rufen Sie dann `displayCatalog` auf und sehen Sie sich die Ausgabe in der Konsole an.

# --hints--

Sie sollten eine Funktion namens `displayCatalog` haben.

```js
assert.isFunction(displayCatalog);
```

Sie sollten Ihre `displayCatalog`-Funktion aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /displayCatalog\s*\(\s*\)/);
```

Ihre `displayCatalog`-Funktion sollte die `keys`-Methode auf `catalog` aufrufen.

```js
const catalogCopy = new Map(catalog);
let flag = false;
const temp = catalog.keys;
try {
    catalog.keys = () => {
        flag = true;
        return catalogCopy.keys()
    };
    displayCatalog();    
    assert.isTrue(flag);
} finally {
    catalog.keys = temp;
}
```

Ihre `displayCatalog`-Funktion sollte jeden Schlüssel, der in `catalog` enthalten ist, in der Konsole ausgeben.

```js
const testArr = [];
const temp = console.log;
try {
    console.log = obj => {testArr.push(obj)};
    displayCatalog();    
    assert.isNotEmpty(testArr);
    const k = catalog.keys();
    for (let i = 0; i < catalog.size; i++) {
        assert.strictEqual(testArr[i], k.next().value)
    }
} finally {
    console.log = temp;
}
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

const removePlant = plant => catalog.delete(plant);

--fcc-editable-region--

--fcc-editable-region--
```
