---
id: a8e512fbe388ac2f9198f0fa
title: 实现匹配对象过滤器
challengeType: 26
dashedName: implement-a-matching-object-filter
---

# --description--

在本实验中，你将创建一个函数，该函数过滤一个对象数组，并且只返回那些匹配给定源对象中所有密钥-值点对的对象。

**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `whatIsInAName` 函数，它接受两个参数，一个对象数组和一个源对象。
1. `whatIsInAName` 函数应返回一个新建数组，该数组仅包含集合中所有存在于源对象中的密钥–值点对的对象。
1. 如果没有对象匹配源中的所有密钥–值点对，函数应返回一个空数组。例如：

```js
whatIsInAName(
  [
    { first: "Romeo", last: "Montague" },
    { first: "Mercutio", last: null },
    { first: "Tybalt", last: "Capulet" }
  ],
  { last: "Capulet" }
);
// should return [{ first: "Tybalt", last: "Capulet" }]
```

# --hints--

你应该有一个 `whatIsInAName` 函数。

```js
assert.isFunction(whatIsInAName);
```

`whatIsInAName([{ first: "Romeo", last: "Montague" }, { first: "Mercutio", last: null }, { first: "Tybalt", last: "Capulet" }], { last: "Capulet" })` 应返回 `[{ first: "Tybalt", last: "Capulet" }]`。

```js
assert.deepEqual(
  whatIsInAName(
    [{ first: "Romeo", last: "Montague" }, { first: "Mercutio", last: null }, { first: "Tybalt", last: "Capulet" }],
    { last: "Capulet" }
  ),
  [{ first: "Tybalt", last: "Capulet" }]
);
```

`whatIsInAName([{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }], { "apple": 1 })` 应该返回 `[{"apple": 1}, {"apple": 1}, {"apple": 1, "bat": 2}]`。

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }],
    { "apple": 1 }
  ),
  [{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }], { "apple": 1, "bat": 2 })` 应该返回 `[{"apple": 1, "bat": 2}, {"apple": 1, "bat": 2, "cookie": 2}]`。

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }],
    { "apple": 1, "bat": 2 }
  ),
  [{ "apple": 1, "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }], { "apple": 1, "cookie": 2 })` 应该返回 `[{"apple": 1, "bat": 2, "cookie": 2}]`。

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }],
    { "apple": 1, "cookie": 2 }
  ),
  [{ "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }, { "bat": 2 }], { "apple": 1, "bat": 2 })` 应该返回 `[{"apple": 1, "bat": 2}, {"apple": 1, "bat": 2, "cookie": 2}]`

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }, { "bat": 2 }],
    { "apple": 1, "bat": 2 }
  ),
  [{ "apple": 1, "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{"a": 1, "b": 2, "c": 3}], {"a": 1, "b": 9999, "c": 3})` 应该返回 `[]`。

```js
assert.deepEqual(
  whatIsInAName([{ "a": 1, "b": 2, "c": 3 }], { "a": 1, "b": 9999, "c": 3 }),
  []
);
```

`whatIsInAName([{"a": 1, "b": 2, "c": 3, "d": 9999}], {"a": 1, "b": 9999, "c": 3})` 应该返回 `[]`。

```js
assert.deepEqual(
  whatIsInAName([{ "a": 1, "b": 2, "c": 3, "d": 9999 }], { "a": 1, "b": 9999, "c": 3 }),
  []
);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function whatIsInAName(collection, source) {
  const sourceKeys = Object.keys(source);

  return collection.filter(obj =>
    sourceKeys.every(key => obj.hasOwnProperty(key) && obj[key] === source[key])
  );
}
```
