---
id: 6976db7ecfa770bf21307b20
title: 플레이리스트 리믹스 엔진 만들기
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

이 실습에서는 청취자가 제출한 여러 플레이리스트에서 하나의 리믹스 플레이리스트를 만드는 프로그램을 만듭니다.

각 청취자는 듣고 싶은 노래 목록을 제공합니다. 어떤 노래는 여러 번 나타날 수 있고, 어떤 아티스트는 너무 자주 등장할 수 있습니다. 이 플레이리스트들을 단계별로 처리하세요: 하나의 목록으로 합치고, 각 노래에 점수를 매기고, 중복 노래를 제거하고, 같은 아티스트가 등장하는 횟수를 제한한 다음 최종 재생 순서를 만드세요.

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `flattenPlaylists`라는 이름의 함수를 만들어야 하며, 이 함수는 각 플레이리스트가 `trackId`, `artist`, `title`, `votes`, `bpm` 속성을 가진 객체 배열인 플레이리스트 배열을 입력으로 받습니다. 입력이 배열이 아니면 `flattenPlaylists`는 빈 배열을 반환해야 합니다. 예시 플레이리스트가 제공되어 있으니 이 예시로 함수를 테스트할 수 있습니다.

2. `flattenPlaylists`는 트랙 객체의 평탄화된 배열을 반환해야 하며, 각 객체는 원래 트랙 속성 모두와 함께 트랙이 어디에서 왔는지 나타내는 플레이리스트 인덱스와 트랙 인덱스를 포함하는 배열인 `source` 속성을 포함해야 합니다.

3. `scoreTracks`라는 이름의 함수를 만들어야 하며, 이 함수는 `flattenPlaylists`에서 반환된 트랙 객체 배열을 입력으로 받아야 합니다(각각 `trackId`, `artist`, `title`, `votes`, `bpm`, `source` 속성을 가짐). 그리고 `score` 공식을 사용해 `votes * 10 - Math.abs(bpm - 120)` 속성이 추가된 새로운 트랙 객체 배열을 반환해야 합니다.

4. `dedupeTracks`라는 이름의 함수를 만들어야 하며, 이 함수는 `scoreTracks`에서 반환된 트랙 객체 배열을 입력으로 받아 중복된 `trackId` 항목을 제거한 새 배열을 반환해야 하며, 각 항목의 첫 번째 등장만 유지해야 합니다.

5. `enforceArtistQuota`라는 이름의 함수를 만들어야 하며, 이 함수는 `dedupeTracks`에서 반환된 트랙 객체 배열과 아티스트별 최대 허용 등장 횟수를 나타내는 숫자를 입력으로 받습니다. 함수는 주어진 횟수보다 아티스트가 더 많이 등장하지 않도록 가장 이른 등장만 유지하는 새 배열을 반환해야 합니다.

6. `buildSchedule`라는 이름의 함수를 만들어야 하며, 이 함수는 `enforceArtistQuota`에서 반환된 트랙 객체 배열을 입력으로 받아 `{ slot, trackId }` 객체 배열을 반환해야 하며, 여기서 `slot`은 방송 순서에서 각 트랙의 위치를 나타내는 1부터 시작하는 인덱스입니다.

7. `remixPlaylist`라는 이름의 함수를 만들어야 하며, 이 함수는 플레이리스트 배열과 아티스트별 최대 허용 등장 횟수를 입력으로 받아 `{ slot, trackId }`, `flattenPlaylists`, `scoreTracks`, `dedupeTracks`, `enforceArtistQuota`를 순서대로 호출하여 최종 방송 스케줄을 `buildSchedule` 객체 배열로 반환해야 합니다.

# --hints--

`flattenPlaylists`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(flattenPlaylists);
```

입력이 배열이 아닐 경우 `flattenPlaylists`에서 빈 배열을 반환해야 합니다.

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

`flattenPlaylists`에서 반환하는 각 트랙은 플레이리스트 인덱스와 트랙 인덱스를 포함하는 배열인 `source` 필드를 포함해야 합니다.

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

`scoreTracks`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(scoreTracks)
```

`scoreTracks`에서 반환하는 각 트랙은 숫자형 `score` 필드를 포함해야 합니다.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

목표 BPM을 `score`으로 사용하고 다음 공식을 사용해 `120`를 계산해야 합니다: `votes * 10 - Math.abs(bpm - 120)`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

`dedupeTracks`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(dedupeTracks)
```

중복된 `trackId` 값이 있을 때 `dedupeTracks`는 각 트랙의 첫 번째 등장만 유지해야 합니다.

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

`enforceArtistQuota`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota`는 아티스트가 `maxPerArtist`번 이상 등장하지 않도록 하며, 가장 이른 트랙을 유지하면서 초과 트랙을 제거해야 합니다.

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

`buildSchedule`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(buildSchedule);
```

`buildSchedule`는 `{ slot, trackId }` 형태의 객체 배열을 반환해야 하며, `slot`는 `1`부터 시작합니다.

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

`remixPlaylist`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist`는 최종 스케줄을 만들기 위해 헬퍼 함수를 순서대로 호출해야 합니다.

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
