---
id: 69757cc0faae0152c1418aad
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви потренуєтеся працювати з числами та математичними операціями, щоб створити калькулятор для розподілу рахунку. Цей інструмент обчислить, скільки кожна людина має заплатити після додавання вартості їжі та чайових.

Для початку вам потрібен спосіб відстежувати загальну суму, коли додаються витрати. У Python ви можете використати змінну, щоб зберігати ціле число, яке змінюється з часом.

Наприклад, ви можете написати:

```py
my_number = 2
```

Створіть змінну з іменем `running_total` і призначте їй значення `0`.

# --hints--

У вас має бути змінна з іменем `running_total`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

Ви маєте призначити ціле число `0` вашій змінній `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
