---
id: 6976cc0b0c135686a4620b94
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Далі потрібно врахувати кількість людей, які ділять рахунок. Збережіть це значення у змінній, як ви робили на попередньому кроці.

Створіть змінну з назвою `num_of_friends` і призначте їй значення `4`. Це буде використано пізніше на практичному занятті для обчислення фінального розбиття.

# --hints--

У вас має бути змінна з назвою `num_of_friends`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

Ви повинні призначити ціле число `4` вашій змінній `num_of_friends`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
