---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Кожен курс має свою вартість. Вам потрібно зберегти ці суми у змінних, щоб використовувати їх пізніше. Оскільки ці суми включають копійки, ви будете використовувати тип float, який застосовується для представлення десяткових чисел. Ось приклад змінної зі значенням типу float:

```py
change = 2.35
```

Створіть чотири змінні: `appetizers` призначте значення `37.89`, `main_courses` призначте значення `57.34`, `desserts` призначте значення `39.39`, а `drinks` призначте значення `64.21`. 

# --hints--

Ви повинні визначити змінну з ім’ям `appetizers` і призначити їй значення `37.89`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

Ви повинні визначити змінну з ім’ям `main_courses` і призначити їй значення `57.34`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

Ви повинні визначити змінну з ім’ям `desserts` і призначити їй значення `39.39`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

Ви повинні визначити змінну з ім’ям `drinks` і призначити їй значення `64.21`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
