---
id: 69779dcb44016eccc05c15a4
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Обслуговування було відмінним, тому група вирішує залишити чайові у розмірі 25%. Щоб обчислити відсоток у Python, ви можете помножити загальну суму на десятковий еквівалент відсотка.

Наприклад, щоб знайти 10% від значення, ви помножите його на `0.10` за допомогою оператора `*`:

```py
tax = total * 0.10
```

Створіть змінну з ім’ям `tip` і призначте їй результат множення `running_total` на `0.25`.

Нарешті, використайте `print()` щоб вивести рядок `Tip amount:`, за яким іде пробіл і значення вашої змінної `tip`.

# --hints--

У вас має бути змінна з ім’ям `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

Ваша змінна `tip` має бути результатом `running_total * 0.25`.

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

Ви маєте вивести рядок `Tip amount:`, за яким іде пробіл і змінна `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
