---
id: 6977a4306bb3e856690a4890
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер, коли ви порахували чайові, потрібно додати їх до вашого `running_total`, щоб знайти фінальну суму рахунку.

У Python ви можете використовувати оператор розширеного присвоєння `+=`, щоб додати значення до змінної та одночасно оновити цю змінну. Наприклад, `total += 5` — це скорочений спосіб запису `total = total + 5`.

Використайте оператор `+=`, щоб додати значення `tip` до вашого `running_total`. Нарешті, використайте `print()`, щоб вивести рядок `Total with tip:`, за яким іде пробіл і значення `running_total`.

# --hints--

Ви повинні використати оператор `+=`, щоб додати значення `tip` до вашого `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

Ви повинні вивести рядок `Total with tip:`, за яким іде пробіл і змінна `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total with tip:', running_total)") or _Node(_code).has_call("print(f'Total with tip: {running_total}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
