---
id: 6977a8bbcbdd557886bd0718
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Тепер, коли чайові включено, у вас є фінальна сума для всієї групи. Ви маєте визначити, скільки кожна людина повинна заплатити, поділивши загальний рахунок на кількість друзів.

У Python для виконання ділення використовується коса риска `/`. Наприклад:

```py
half = 10 / 2
```

Створіть змінну з ім’ям `final_bill` і призначте їй результат ділення `running_total` на `num_of_friends`.

Нарешті, використайте функцію `print()`, щоб вивести рядок `Bill per person:`, за яким іде пробіл і значення `final_bill`.

# --hints--

Ви маєте визначити змінну з ім’ям `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

Ви маєте використати оператор `/`, щоб поділити `running_total` на `num_of_friends` і призначити результат змінній `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

Ви маєте вивести рядок `Bill per person:`, за яким іде пробіл і змінна `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
