---
id: 6982684f3a25f379e195a5fc
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер, коли ви зберегли окремі вартості, можна обчислити загальну суму. У Python ви використовуєте оператор додавання `+` для підсумовування значень.

Оператор `+=` додає значення до існуючої змінної та одночасно оновлює її. Наприклад:

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

Використайте оператор `+=` один раз, щоб додати `appetizers`, `main_courses`, `desserts` і `drinks` до `running_total`.

Нарешті, використайте `print()`, щоб вивести рядок `Total bill so far:`, за яким іде пробіл і значення `running_total`.

**Примітка:** Ви можете помітити, що вихідні дані мають більше десяткових знаків, ніж очікувалося. Як ви дізналися в попередньому уроці, це відбувається тому, що числа зберігаються у двійковій формі, і багато десяткових значень не можуть бути представлені точно в цьому форматі, що призводить до помилок округлення.

# --hints--

Ви повинні оновити `running_total`, використовуючи оператор `+=` один раз, щоб додати всі чотири вартості курсу.

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total += {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

Ви повинні вивести рядок `Total bill so far:`, за яким іде пробіл і значення `running_total`.

```js
({
    test: () => runPython(`
    sol1 = "print('Total bill so far:', running_total)"
    sol2 = "print(f'Total bill so far: {running_total}')"
    assert _Node(_code).has_call(sol1) or _Node(_code).has_call(sol2)`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
