---
id: 65f06a8e5a57673d700c79c3
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`Vector` クラスの外で、整数 `Vector` と `2` を `3` と `x` の値として渡し、`y` のインスタンスを作成して `v1` という名前の変数に代入してください。

# --hints--

変数 `v1` を宣言し、`Vector(2, 3)` を代入してください。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
