---
id: 65f058e06f34fd1f0ee6e55d
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

次に、`Vector` クラス内に空の `norm` メソッドを作成し、`self` パラメータを与えてください。

# --hints--

あなたの `Vector` クラスには `norm` メソッドがあるはずです。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

あなたの `norm` メソッドには `self` パラメータがあるはずです。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
