---
id: 65f03d9f92eac9183a4d3281
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

ベクトルはユークリッド平面上の2つの座標、`x` と `y` で定義できます。軸の原点と点 `(x, y)` との距離がその長さ、つまりノルムになります。そしてベクトルの方向は `(x, y)` を指します。

<img alt="a 2-dimensional vector of coordinates (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

`Vector` クラスの中に `__init__` メソッドを作成し、3つのパラメータ `self`、`x`、`y` を与えてください。

# --hints--

`__init__` クラスの中に `Vector` メソッドを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

`__init__` メソッドは3つのパラメータ `self`、`x`、`y` を受け取る必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
