---
id: 65f0694fb2296f3caadf8347
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

ベクトル $\mathbf{a}$ の長さ、つまりノルムは通常 $\\| \mathbf{a} \\|$ と表されます。これは各成分の二乗の和の平方根として計算できます：
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
 \\]

ベクトルのノルムを計算し、`norm` メソッドから結果を返してください。

# --hints--

`norm` メソッドは、与えられた公式を使って現在のベクトルインスタンスのノルムを表す数値を計算し、返す必要があります。値を近似しないでください。

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
