---
id: 65f055b9190fc41ca35549b8
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

Python には、名前の先頭と末尾にダブルアンダースコアが付くさまざまなメソッドがあります。`__init__` や `__str__` のようなものにすでに馴染みがあるかもしれません。これらのメソッドは `__<name>__` という命名パターンに従っており、特別なメソッド、マジックメソッド、または *dunder*（ダブルアンダースコアの略）メソッドと呼ばれます。

クラス内で特別なメソッドを定義すると、そのクラスの動作に影響します。これらは特定の状況で裏側で呼び出されます（例：インスタンス化時の `__init__`、オブジェクトが出力されたり `__str__` に渡されたときの `str()`）。このプロジェクトでは、よく使われる特別なメソッドのいくつかを学びます。

まずは、`x` を `x` オブジェクトの `Vector` 属性に代入してください。

# --hints--

`x` を `self.x` に代入する必要があります。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
