---
id: 65f01969115f933073b6be03
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

ベクトルとは、長さ（または大きさ）と方向を持ち、単一の数値では表せないオブジェクトです。物理学では、力、速度、加速度などの量を表すためによく使われます。

このプロジェクトでは、ベクトル空間を構築します。これは、その集合の要素間で一連の演算が定義されている集合です。詳細はすぐに学びます。

まずは、`Vector` という名前の空のクラスを宣言してプロジェクトを始めてください。

# --hints--

`Vector` という名前の空のクラスを宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
