---
id: 65f3fe07cc763212efe91285
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

ベクトルは `n` 個の次元（コンポーネント）を持つことができます。こちらは3次元ベクトルの表現です：

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

これまでに2次元ベクトルを作成しました。異なる次元数のベクトルを、それぞれの場合に必要なコードを書き直すことなく表現できるようにしたいです。そのために継承を使います。

まず、`Vector` クラスの名前を `R2Vector` に変更して、このクラスが2次元ベクトルを扱うことを明示してください。`v1` のインスタンス化も忘れずに修正してください。

# --hints--

`Vector` クラスの名前を `R2Vector` に変更してください。`v1` のインスタンス化も修正することを忘れないでください。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
