---
id: 65f3fa097f26b510db6c710b
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

素晴らしいです、`norm` は期待通りに動作しています。次に、`v1` を出力しようとすると、オブジェクトのデフォルトの文字列表現（`<__main__.Vector object at 0x11eb778>` のようなもの）が表示されます。

`Vector` クラスの中で、読みやすい文字列表現を実装するために空の `__str__` メソッドを宣言してください。パラメータとして `self` を必ず付けてください。

`v1` が文字列を返すまで `__str__` を出力しないように注意してください。そうしないと `TypeError` が発生します。なぜなら `__str__` は常に文字列を返す必要があるからです。

# --hints--

`__str__` クラス内に `Vector` メソッドを定義してください。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

`__str__` メソッドには `self` パラメータが必要です。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
