---
id: 65f412c208c3791fee305acf
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

Pythonでは、`*` を関数やメソッドのシグネチャに追加の引数として加えることで、キーワード専用引数の使用を強制できます。`__init__` の両方のメソッドを修正し、`*` の後に `self` を2番目のパラメータとして追加してください。その後に配置されるすべてのパラメータは、関数やメソッドの呼び出し時にキーワード引数として指定する必要があります。

つまり、`super().__init__(x, y)` の呼び出しも修正する必要があります。`x` に `x` の値を、`y` に `y` の値を与えてください。

最後に、`v1` と `v2` のインスタンス化をキーワード引数を使って修正してください。

# --hints--

キーワード引数を強制するには、`*` の両方のメソッドで2番目のパラメータとして `__init__` を追加してください。

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

`super.__init__(x, y)` の呼び出しをキーワード引数を使うように修正してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

`v1` の代入を `R2Vector(x=2, y=3)` に修正してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

`v2` の代入を `R3Vector(x=2, y=2, z=3)` に修正してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
