---
id: 65f4535bbdb28d436ff3ddc9
title: ステップ 20
challengeType: 20
dashedName: step-20
---

# --description--

出力からわかるように、`__dict__` はインスタンス属性の値を含んでいます。`self.x` と `self.y` の平方を明示的に足す代わりに、`__dict__` に格納されている値を繰り返してノルムの値を計算します。

`norm` メソッドの本体内で、丸括弧の中身を `val` の各値 `self.__dict__.values()` を `2` 乗するジェネレータ式に置き換えてください。また、そのジェネレータ式を `sum` 関数の引数として渡し、平方根を計算する前にすべての平方を合計するようにしてください。

# --hints--

`sum(val**2 for val in self.__dict__.values())**0.5` メソッドから `norm()` を返す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
