---
id: 65f40051d6b09a139f253e8e
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

継承を使うと、既存のクラスから新しいクラスを定義できます。新しいクラスは子クラスと呼ばれ、既存のクラス（親クラス）のすべてのメソッドとプロパティを継承します。

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

上の例では、子クラス`Oak`が`Tree`から継承し、親クラス`sprout`の`Tree`メソッドを継承しています。

`R3Vector`という名前の新しいクラスを作成し、上の例に従って`R2Vector`クラスから継承してください。

# --hints--

`R3Vector`という新しいクラスを定義してください。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

新しいクラス`R3Vector`は`R2Vector`から継承する必要があります。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
