---
id: 65f4148dea0f802040225e0c
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

既存の `print` の呼び出しを削除してください。それから、以前 `v1` に対して行ったように、`v2` と `v2.norm()` が返す値を出力してください。

# --hints--

コード内に `print(v1)` と `print(v1.norm())` があってはいけません。

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v1)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v1.norm())")`));
    }
})
```

`v2` と `v2.norm()` を出力してください。

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_call("print(v2)")`));
      assert(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
print(v1.norm())
print(v1)

v2 = R3Vector(x=2, y=2, z=3)
--fcc-editable-region--
```
