---
id: 65f407ea37ad6e181b90462e
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

前と同じように、各`i`パラメータを`self.i`に代入することもできます。繰り返す行数が少ない場合でも、この繰り返しを避ける方法の一つは`super()`関数を使うことです。`super()`は親クラスを暗黙的に参照できるようにします：`super().__init__(x, y)`は親クラスの`__init__`メソッドを呼び出します。

`super().__init__(x, y)`メソッドに`__init__`の呼び出しを追加してください。

# --hints--

`super().__init__(x, y)`メソッドに`__init__`の呼び出しがあるはずです。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
