---
id: 65f40401e6ef53173c04e27d
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

新しいクラスの中で、`__init__` メソッドを宣言してください。パラメータとして `self`、`x`、`y`、`z` を与えてください。

# --hints--

`__init__` クラスの中に `R3Vector` メソッドを定義する必要があります。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").has_function("__init__")
    `))
})
```

`__init__` メソッドは、4つのパラメータ `self`、`x`、`y`、`z` を持つべきです。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, x, y, z")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
