---
id: 65f4520e363e2642f8112e33
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

ご覧のとおり、何かがうまくいっていません。`norm` と `__str__` メソッドは `R2Vector` から継承されていますが、3次元ベクトルには対応できません。これらの実装はもっと柔軟である必要があります。

Pythonのすべてのオブジェクトには `__dict__` という特別な属性があり、これはオブジェクトの属性を格納する辞書です。

既存の `print` の呼び出しを削除してください。その後、`__dict__` と `v1` のベクトルの `v2` 属性を出力して、それらがどのようなものか確認してください。

# --hints--

コードに `print(v2)` と `print(v2.norm())` が含まれていないことを確認してください。

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

`__dict__` と `v1` の `v2` 属性を出力してください。

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
