---
id: 65f876d17832001e8e1abb05
title: ステップ 28
challengeType: 20
dashedName: step-28
---

# --description--

変数`arg_list`を宣言してください。`key`の各キーと値のペアに対して`val`と`vars(self).items()`を繰り返し、各反復でf文字列`f'{key}={val}'`を計算するリスト内包表記の値を代入してください。

# --hints--

変数`arg_list`を宣言する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_variable("arg_list")`)) })
```

`vars(self).items()`を繰り返すリスト内包表記を`arg_list`に代入してください。

```js
({ test: () => runPython(`
import ast
var = _Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list")
assert isinstance(var.tree.value, ast.ListComp), "Expected arg_list to be a list comprehension"
assert (actual := var.find_comp_iters()[0].is_equivalent("vars(self).items()")), f"Expected vars(self).items(), got {actual}"
`) })
```

`arg_list`に代入されたリスト内包表記は、`key`と`val`を使って`vars(self).items()`を繰り返す必要があります。

```js
({ test: () => assert(runPython(`
_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_targets()[0].is_equivalent("key, val")
`)) })
```

`arg_list`に代入されたリスト内包表記は、`f'{key}={val}'`の各キーと値のペアに対して文字列`vars(self).items()`を計算する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_expr().is_equivalent("f'{key}={val}'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
