---
id: 65f8749b13774b1d2e4a7fba
title: ステップ 27
challengeType: 20
dashedName: step-27
---

# --description--

`__str__` メソッドはオブジェクトの人間が読める文字列の表現を返しますが、`__repr__` メソッドはオブジェクトを生成するために必要な文字列を返すことになっています。

`__repr__` クラス内に `self` パラメータを持つ `R2Vector` メソッドを追加してください。

# --hints--

`R2Vector` クラスには `__repr__` メソッドが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__repr__")`)) })
```

`__repr__` メソッドには `self` パラメータが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
