---
id: 65f8709620b2ce1a62608f5a
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

文字列の入力から動的に属性にアクセスする必要がある場合、組み込みの `getattr()` 関数が役立ちます。この関数は、最初の引数にオブジェクトを取り、2番目の引数に属性名を含む文字列を取ります。

`__str__` メソッドを修正し始めて、`__str__()` が返す文字列を、オブジェクトの属性を反復し、各属性に対して `getattr()` を呼び出すジェネレータ式に置き換えてください `i`。

# --hints--

オブジェクトの属性を反復するジェネレータ式を返す必要があります `vars(self)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

`i` を反復変数として使うジェネレータ式を返す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

`getattr(self, i)` の各アイテム `i` に対して `vars(self)` を呼び出すジェネレータ式を返す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
