---
id: 65f9dd6e5a08af19c196c2df
title: ステップ 45
challengeType: 20
dashedName: step-45
---

# --description--

あるベクトルから別のベクトルを引いた結果のベクトルは、それぞれのコンポーネントの差を計算することで得られます。例えば、`(2, 4)` から `(7, 3)` を引くと、3つ目のベクトル `(5, -1)` が生成されます。ここで `5` は `7` と `2` の差であり、`-1` は `3` と `4` の差です。

以前 `__add__` メソッド内で行ったように、`kwargs` 文の後に `if` 変数を宣言してください。これは、オブジェクトの属性を繰り返し処理し、各キー `i` に対してキーと値のペアを作成する辞書内包表記を代入します。ここでキーは現在の繰り返しのキーと同じで、その値は2つのオペランドの `i` 属性の値の差です。

# --hints--

`kwargs` という変数を宣言し、オブジェクトの属性を繰り返す辞書内包表記を代入してください。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

`kwargs` に代入する辞書内包表記は、`i` という変数を使ってオブジェクトの属性を繰り返す必要があります。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

`kwargs` に代入する辞書内包表記は、`i` をキーとして使う必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

`kwargs` に代入する辞書内包表記は、各 `i` に対して `self` の `i` 属性と `other` の `i` 属性の差を `i` キーに代入する必要があります。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
