---
id: 65f93b67169a9c703264458a
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

メソッドは継承されたときに `R2Vector` と `R3Vector` のオブジェクトを生成するための文字列を返す必要があるため、クラス名を指定して文字列を作成することはできません。

クラスの名前は `__class__.__name__` でアクセスできます。`return` メソッドに `__repr__` 文を追加し、オブジェクトを生成するために必要な文字列を返してください。

# --hints--

ベクトルオブジェクトを生成するために必要な f-string を作成し、`__class__.__name__` メソッドから返すには `__repr__` を使うべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_return("f'{self.__class__.__name__}({args})'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
