---
id: 65f9bc10a9fb1612e066e9e8
title: ステップ 42
challengeType: 20
dashedName: step-42
---

# --description--

2つのベクトルを加算するときは、一方のベクトルの各コンポーネントに対して、もう一方のベクトルの同じコンポーネントを加算します。たとえば、`(1, 2)` と `(2, 4)` を加算すると、3つ目のベクトル `(3, 6)` が生成されます。ここで `3` は `1` と `2` の和であり、`6` は `2` と `4` の和です。

`kwargs` 辞書には、加算された2つのベクトルと同じクラスの新しいベクトルをインスタンス化するために必要なキーと値のペアが含まれます。

空の辞書を、`vars(self)` を繰り返し処理する辞書内包表記に変え、各キー（`i`）に対してキーと値のペアを作成してください。ここでキーは現在の繰り返しのキーと同じで、その値は2つのオペランドの `i` 属性の値の和です。

# --hints--

`kwargs` に代入された空の辞書を、`vars(self)` を繰り返す辞書内包表記に変える必要があります。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

`kwargs` に代入された辞書内包表記は、`i` を繰り返すために変数 `vars(self)` を使うべきです。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

`kwargs` に代入された辞書内包表記は、キーとして `i` を使うべきです。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

`kwargs` に代入された辞書内包表記は、`i` の各 `self` に対して、`i` の `other` 属性と `i` の `i` 属性の和を `vars(self)` キーに代入するべきです。そのために `getattr()` 関数を使ってください。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
