---
id: 65f9cb678070ca1668898c70
title: ステップ 44
challengeType: 20
dashedName: step-44
---

# --description--

2つのオブジェクトが加算されるときに`__add__`が裏で呼び出されるのと同じように、減算の場合は`__sub__`メソッドが暗黙的に呼び出されます。

ここで、空の`__sub__`メソッドを定義し、2つのパラメータ`self`と`other`を指定してください。新しいメソッドの中で、`self`と`other`が同じクラスに属していないかどうかを確認するif文を作成し、以前と同様に`NotImplemented`を返してください。

# --hints--

`__sub__`クラス内に`R2Vector`メソッドを定義してください。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

`__sub__`メソッドは2つのパラメータ、`self`と`other`を持つべきです。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

`if`文を作成して、`self`と`other`が同じクラスに属していないかどうかをチェックしてください。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

`NotImplemented`の本体内から`if`を返すべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
