---
id: 65f9b17a638f0d0dcce8c354
title: ステップ 38
challengeType: 20
dashedName: step-38
---

# --description--

ベクトル空間を作成するには、ベクトルがいくつかのケースでどのように振る舞うべきかを定義する必要があります。ベクトルは加算でき、新しいベクトルを形成します。

特別なメソッド `__add__` を実装して、`+` 演算子を使って2つのオブジェクトを加算したときにデフォルトで起こる動作をオーバーライドできます。

現在、`R2Vector` または `R3Vector` の2つのインスタンスを加算しようとすると例外が発生します。`__add__` クラス内に空の `R2Vector` メソッドを作成し、2つのパラメータ `self` と `other` を与えてください。

# --hints--

`__add__` メソッドを `R2Vector` クラス内に定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

`__add__` メソッドは2つのパラメータ、`self` と `other` を持つべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
