---
id: 66682150151af29efec9727d
title: ステップ 33
challengeType: 20
dashedName: step-33
---

# --description--

`__getattribute__` メソッドは、インスタンス属性にアクセスしようとするたびに内部で呼び出されます。属性がインスタンスレベルで見つからない場合、このメソッドはクラスレベルで検索し、最終的には親クラスでも探します。

クラス内に、2つのパラメータ `__getattribute__` と `self` を持つ `attr` メソッドを定義し、文字列 `'calling __getattribute__'` を返すようにしてください。属性アクセスがどのように動作するかを見るために、デフォルトの実装をすぐにオーバーライドします。

# --hints--

`__getattribute__` クラス内に、2つのパラメータ `self` と `attr` を持つ `R2Vector` という名前のメソッドを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

`__getattribute__` メソッドは文字列 `'calling __getattribute__'` を返すべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
