---
id: 666826f258fda1ab3396a509
title: ステップ 34
challengeType: 20
dashedName: step-34
---

# --description--

今度は、`x`の`v1`属性にドット演算子と`getattr()`組み込み関数の両方でアクセスした結果を出力してください。

どちらの場合も`__getattribute__`が呼び出されることがわかります。

# --hints--

`x`の`v1`属性をドット演算子と`getattr()`関数の両方で出力してください。

```js
({ test: () => assert(runPython(`
n = _Node(_code)
(n.has_call("print(v1.x)") and n.has_call("print(getattr(v1, 'x'))")) or n.has_call("print(v1.x, getattr(v1, 'x'))") or n.has_call("print(getattr(v1, 'x'), v1.x)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
    
    def __getattribute__(self, attr):                
        return 'calling __getattribute__'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')
--fcc-editable-region--

--fcc-editable-region--
```
