---
id: 66680ddfd0f8c76782923cb0
title: ステップ 22
challengeType: 20
dashedName: step-22
---

# --description--

`norm()` メソッドは正しい値を返していますが、まだ改善できる点があります：可読性です。

`vars()` 組み込み関数はオブジェクトを引数に取り、そのオブジェクトの `__dict__` 属性を返します。

`__dict__` の `self` 属性に直接アクセスする代わりに、`norm` メソッドを修正して `vars()` 関数を使うようにしてください。

# --hints--

`norm` メソッドを修正して、`vars(self)` の代わりに `self.__dict__` を使うようにしましょう。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in vars(self).values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
