---
id: 666832e427d70bc5219dc62a
title: ステップ 36
challengeType: 20
dashedName: step-36
---

# --description--

出力からわかるように、クラスで `__getattr__` を定義していても、このメソッドはまだ呼び出されていません。これは、デフォルトの属性アクセスが `__getattribute__` を通じて行われるためです。そのため、ターミナルに属性の値が出力されています。

では、コードの最後の2行を修正して、両方の `z` 呼び出しで `v1` の `print()` 属性にアクセスしてください。今回は、`v1` が持っていない属性にアクセスします。その結果、`__getattribute__` がエラーを発生させ、`__getattr__` が呼び出されます。

# --hints--

コードを修正して、ドット演算子と `z` 関数の両方を使って `v1` の `getattr()` 属性を出力してください。

```js
({ test: () => assert(runPython(`
n = _Node(_code)
(n.has_call("print(v1.z)") and n.has_call("print(getattr(v1, 'z'))")) or n.has_call("print(v1.z, getattr(v1, 'z'))") or n.has_call("print(getattr(v1, 'z'), v1.z)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __getattr__(self, attr):                
        return 'calling __getattr__'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')
--fcc-editable-region--
print(v1.x)
print(getattr(v1, 'x'))
--fcc-editable-region--
```
