---
id: 666827a6fd0dbaafe8330ea6
title: ステップ 35
challengeType: 20
dashedName: step-35
---

# --description--

`__getattr__` は、属性にアクセスする際に役割を果たすもう一つの特別なメソッドです。

`__getattribute__` のデフォルトの実装は、要求された属性がインスタンス属性でないか、クラスツリーに存在しない場合に `AttributeError` を発生させることです。

その場合、クラスで定義されていれば `__getattr__` が呼び出されます。これは通常の属性アクセス手順が失敗したときのフォールバックのようなものと考えられます。

`__getattribute__` メソッドを `__getattr__` に変更し、そのメソッドが返す文字列を `'calling __getattr__'` に書き換えてください。

# --hints--

`__getattr__` クラスの中に、`self` と `attr` の2つのパラメータを持つ `R2Vector` という名前のメソッドを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_args("self, attr")`)) })
```

`__getattr__` メソッドは文字列 `'calling __getattr__'` を返す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_return("'calling __getattr__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __getattribute__(self, attr):                
        return 'calling __getattribute__'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

print(v1.x)
print(getattr(v1, 'x'))

```
