---
id: 5d822fd413a79914d39e98cb
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Далі додайте відкриваючий і закриваючий теги `head` та `body` всередині елемента `html`.

# --hints--

У вас має бути відкриваючий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

У вас має бути закриваючий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

У вас має бути відкриваючий тег `body`.

```js
assert(code.match(/<body\s*>/i));
```

У вас має бути закриваючий тег `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` та `body` мають бути братськими.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

Елемент `head` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

Елемент `body` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
  
</html>
--fcc-editable-region--

```
