---
id: 5d822fd413a79914d39e98cc
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Всередині `head` вкладіть елемент `meta` з атрибутом `charset` зі значенням `UTF-8`, елемент `title` із заголовком `City Skyline` та елемент `link`, який посилається на ваш файл `styles.css`.

# --hints--

Ви повинні створити елемент `meta` всередині елемента `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('meta'));
```

Ви повинні надати тегу `meta` атрибут `charset` зі значенням `UTF-8`.

```js
const linkElement = document.querySelector('meta');
const charset = linkElement?.getAttribute("charset").toLowerCase();
assert.strictEqual(charset, 'utf-8');
```

У вашому коді має бути елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Елемент `title` має бути всередині елемента `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

Ваш проєкт має мати заголовок `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Пам’ятайте, що важливі регістр і правопис заголовка.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

У вашому коді має бути елемент `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Ви не повинні змінювати існуючі теги `head`. Переконайтеся, що ви не видалили кінцевий тег.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

У вас має бути один порожній елемент `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Ваш елемент `link` має бути всередині елемента `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Елемент `link` має мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Елемент `link` має мати атрибут `href` зі значенням `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
