---
id: 5d822fd413a79914d39e98ca
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте початковий та кінцевий теги `html` під `DOCTYPE`, щоб ви мали місце для коду. Переконайтесь, що встановили англійську мову.

# --hints--

Оголошення `DOCTYPE` має бути на початку HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Елемент `html` повинен мати початковий тег та атрибут `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Елемент `html` повинен мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Теги `html` повинні бути в правильному порядку.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Ви повинні мати тільки один елемент `html`.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
