---
id: 5d822fd413a79914d39e98c9
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Ласкаво просимо до проєкту CSS Variables Skyline! Почніть із додавання оголошення `!DOCTYPE html` на початку документа, щоб браузер знав, який тип документа він читає.

# --hints--

Ваш код має містити референс `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Після референса `DOCTYPE` слід додати пробіл.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Ви маєте визначити тип документа як `html`.

```js
assert(code.match(/html/gi));
```

Оголошення `DOCTYPE` слід закрити `>` після типу.

```js
assert(code.match(/html\s*>/gi));
```

Ваше оголошення `DOCTYPE` має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
