---
id: 5d822fd413a79914d39e98cf
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Ви бачите `body` (це найвнутрішніший блок на вашій сторінці); блок навколо нього — це елемент `html`. Зробіть так, щоб `body` займав усе вікно перегляду, задавши йому властивість `height` зі значенням `100vh`. Видаліть стандартний відступ `margin` у `body`, встановивши `margin` у `0`. Нарешті, задайте властивість `overflow` значення `hidden`, щоб приховати будь-які смуги прокрутки, які з’являються, коли щось виходить за межі вікна перегляду.

# --hints--

Ви повинні використовувати селектор `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Ваш `body` має мати властивість `height` зі значенням `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

Ваш `body` має мати властивість `margin` зі значенням `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Ваш `body` має мати властивість `overflow`, встановлену в `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```
