---
id: 5d822fd413a79914d39e98d4
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Задайте частинам вашої будівлі `width` і `height` властивості зі значеннями: `70%` і `10%` до `.bb1a`, `80%` і `10%` до `.bb1b`, `90%` і `10%` до `.bb1c`, а також `100%` і `70%` до `.bb1d`. Пам’ятайте, що ці відсотки відносні до батьківського елемента, і зверніть увагу, що висоти в сумі становитимуть 100% — вертикально заповнюючи контейнер.

# --hints--

Ви повинні використати селектор класу, щоб стилізувати `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Ви повинні задати `.bb1a` властивість `width` зі значенням `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Ви повинні задати `.bb1a` властивість `height` зі значенням `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Ви повинні використати селектор класу, щоб стилізувати `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Ви повинні задати `.bb1b` властивість `width` зі значенням `80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Ви повинні задати `.bb1b` властивість `height` зі значенням `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Ви повинні використати селектор класу, щоб стилізувати `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Ви повинні задати `.bb1c` властивість `width` зі значенням `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Ви повинні задати `.bb1c` властивість `height` зі значенням `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Ви повинні використати селектор класу, щоб стилізувати `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Ви повинні задати `.bb1d` властивість `width` зі значенням `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Ви повинні задати `.bb1d` властивість `height` зі значенням `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```
