---
id: 5d822fd413a79914d39e98ca
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте початковий і кінцевий `html` теги під `DOCTYPE`, щоб мати місце для початку написання коду. Обов’язково встановіть мову на англійську.

# --hints--

Ваше оголошення `DOCTYPE` має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ваш елемент `html` повинен мати початковий тег і атрибут `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш елемент `html` повинен мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Ваші теги `html` мають бути у правильному порядку.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

У вас має бути лише один елемент `html`.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
