---
id: 5d822fd413a79914d39e98d1
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Надайте вашому `.background-buildings` елементу `width` і `height` зі значенням `100%`, щоб він займав повну ширину і висоту свого батьківського елемента, `body`.

# --hints--

Ви повинні використати клас `background-buildings` для вибору правильного елемента.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

Вашому `.background-buildings` елементу слід надати `width` зі значенням `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

Вашому `.background-buildings` елементу слід надати `height` зі значенням `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--
    
```
