---
id: 69a56b5069ca99f7317e6e19
title: Costruire un validatore di manifesti di carico
challengeType: 26
dashedName: lab-cargo-manifest-validator
---

# --description--

In questo laboratorio userai JavaScript per normalizzare e validare i manifesti di carico. Un manifesto di carico è un documento che solitamente elenca le merci trasportate (ad esempio, via nave o treno) e include dettagli su tali merci.

Ogni manifesto di carico sarà rappresentato come un oggetto con le seguenti proprietà:

- `containerId`: un numero intero positivo che identifica il contenitore associato.
- `destination`: una stringa non vuota (dopo aver rimosso gli spazi) che indica la destinazione finale del carico.
- `weight`: un numero positivo che rappresenta il peso del carico.
- `unit`: una stringa che descrive l'unità della proprietà `weight` del carico (o `"kg"` per chilogrammi o `"lb"` per libbre).
- `hazmat`: un valore booleano che indica se è necessario gestire materiali pericolosi.

Esempio di oggetto manifesto di carico:

```js
{
  containerId: 1,
  destination: "Monterey, California, USA",
  weight: 831,
  unit: "lb",
  hazmat: false
}
```

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi implementare una funzione chiamata `normalizeUnits` con un parametro `manifest`.
    - La funzione non deve modificare l'oggetto manifest originale e deve sempre restituire un nuovo oggetto in cui `weight` è normalizzato in chilogrammi e `unit` è impostato su `"kg"`.
    - Se il `weight` dell'oggetto manifest è espresso in libbre (`unit: "lb"`), la funzione deve convertire il `weight` in chilogrammi usando la conversione approssimativa `1 lb = 0.45 kg` e aggiornare di conseguenza la `unit`.
    - Se il `weight` è già espresso in chilogrammi (`unit: "kg"`), `weight` e `unit` devono rimanere invariati.
2. Devi implementare una funzione chiamata `validateManifest` con un parametro `manifest`.
    - La funzione non deve modificare l'oggetto manifest originale e deve sempre restituire un nuovo oggetto.
    - Se il manifest in input è valido (nessuna proprietà mancante o non valida), la funzione deve restituire un oggetto vuoto.
    - Se il manifest in input non è valido, la funzione deve restituire un oggetto contenente elementi per ogni proprietà mancante o non valida. Le proprietà mancanti devono avere il valore `"Missing"` e quelle non valide il valore `"Invalid"`.

Esempio di valore di ritorno quando l'oggetto in input manca della proprietà `destination` e ha una proprietà `weight` non valida:

    ```js
    {
      destination: "Missing",
      weight: "Invalid"
    }
    ```

3. Devi implementare una funzione chiamata `processManifest` con un parametro `manifest`. La funzione deve stampare:
    - Se l'oggetto manifest è valido, `Validation success: ${containerId}` e poi il `weight` del manifest in chilogrammi, nel formato `Total weight: ${weight} kg`. Usa `normalizeUnits()` per questa conversione.
    - Se l'oggetto manifest non è valido, `Validation error: ${containerId}` e poi l'oggetto restituito chiamando `validateManifest()` con l'oggetto manifest.

**Nota:** Ciascuno di questi due casi deve prevedere due chiamate a `console.log()`.

**Nota:** Non dichiarare `normalizeUnits`, `validateManifest` o `processManifest` usando `const`, perché i test devono poterle riassegnare.

# --before-each--

```js
const _validManifests = [
  {
    containerId: 4,
    destination: "Monterey, California, USA",
    weight: 831,
    unit: "lb",
    hazmat: false
  },
  {
    containerId: 5,
    destination: "Montreal, Quebec, Canada",
    weight: 151,
    unit: "kg",
    hazmat: false
  }
];

const _invalidManifests = [
  {
    containerId: -6,
    destination: 123,
    weight: 0,
    unit: "pounds",
    hazmat: true
  },
  {
    containerId: 0,
    weight: -21,
    unit: "KG"
  },
  {
    containerId: "eight",
    destination: "Guadalajara, Jalisco, Mexico",
    weight: 9001,
    hazmat: "no"
  },
  {
    destination: " ",
    weight: NaN,
    unit: "kg",
    hazmat: 1
  },
  {
    containerId: 10,
    destination: " ",
  },
];

const _testObj = {
  containerId: 1.51,
  destination: "Lavender Town",
  weight: NaN,
};

console.log = () => {};
```

# --hints--

Non devi dichiarare `normalizeUnits`, `validateManifest` o `processManifest` usando `const`, perché i test devono poterle riassegnare.

```js
assert.doesNotThrow(() => {
  normalizeUnits = normalizeUnits;
  validateManifest = validateManifest;
  processManifest = processManifest;
});
```

Devi avere una funzione chiamata `normalizeUnits` con un parametro `manifest`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.normalizeUnits);
const params = explorer.allFunctions.normalizeUnits.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), 'manifest');
```

La tua funzione `normalizeUnits` deve restituire il nuovo oggetto `{ containerId: 68, destination: "Salinas", weight: 45.45, unit: "kg", hazmat: true }` quando viene chiamata con `{ containerId: 68, destination: "Salinas", weight: 101, unit: "lb", hazmat: true }`, senza modificare l'originale.

```js
const testObj = {
  containerId: 68,
  destination: "Salinas",
  weight: 101,
  unit: "lb",
  hazmat: true
};

const expected = {
  containerId: 68,
  destination: "Salinas",
  weight: 45.45,
  unit: "kg",
  hazmat: true
};

const copy = normalizeUnits(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "normalizeUnits() did not return the expected normalized object.");
assert.notStrictEqual(copy, testObj, "normalizeUnits() should return a new object.");
assert.strictEqual(testObj.containerId, 68, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, "Salinas", "Original object should not be mutated.");
assert.strictEqual(testObj.weight, 101, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "lb", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, true, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 5, "Original object should not be mutated.");
```

La tua funzione `normalizeUnits` deve restituire una copia dell'oggetto manifest in input con il suo `weight` normalizzato in chilogrammi e la sua `unit` impostata su `"kg"`. Usa la conversione approssimativa `1 lb = 0.45 kg` per la conversione del `weight`.

```js
_validManifests.forEach((obj, i) => {
  const result = normalizeUnits(obj);
  assert.isObject(result, "normalizeUnits() should return an object.");
  assert.notStrictEqual(result, obj, "normalizeUnits() should return a new object, not mutate the original.");
  const expectedWeight = obj.unit === "lb" ? obj.weight * 0.45 : obj.weight;
  const expected = {
    ...obj,
    weight: expectedWeight,
    unit: "kg"
  };

  assert.deepEqual(
    result,
    expected,
    `normalizeUnits() did not return the expected normalized object for manifest at index ${i}.`
  );
});
```

La tua funzione `normalizeUnits` deve restituire una nuova copia dell'oggetto manifest in input senza modificare l'originale.

```js
const original = {
  containerId: 151,
  destination: "Saffron City",
  weight: 151,
  unit: "lb",
  hazmat: false
};

const expected = {
  containerId: 151,
  destination: "Saffron City",
  weight: 67.95,
  unit: "kg",
  hazmat: false
};

const copy = normalizeUnits(original);
assert.isObject(copy);
assert.deepEqual(copy, expected);
assert.notStrictEqual(copy, original, "normalizeUnits() should return a new object.");
assert.strictEqual(original.weight, 151, "Original object should not be mutated.");
assert.strictEqual(original.unit, "lb", "Original object should not be mutated.");
assert.equal(Object.keys(original).length, 5, "Original object should not be mutated.");
```

Devi avere una funzione chiamata `validateManifest` con un parametro `manifest`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.validateManifest);
const params = explorer.allFunctions.validateManifest.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), 'manifest');
```

La tua funzione `validateManifest` deve restituire un nuovo oggetto vuoto `{}` quando viene chiamata con `{ containerId: 1, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }`.

```js
const testObj = {
  containerId: 1,
  destination: "Santa Cruz",
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for a valid manifest.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, 1, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, "Santa Cruz", "Original object should not be mutated.");
assert.strictEqual(testObj.weight, 304, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, false, "Original object should not be mutated.");
```

La tua funzione `validateManifest` deve restituire un oggetto vuoto `{}` se l'oggetto manifest in input è valido.

```js
const solution = {};
for (const obj of _validManifests) {
  assert.deepEqual(
    validateManifest(obj),
    solution,
    "validateManifest() should return an empty object for valid manifests."
  );
}
```

La tua funzione `validateManifest` deve restituire il nuovo oggetto `{ containerId: "Missing", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` quando viene chiamata con `{}`, senza modificare l'originale.

```js
const testObj = {};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.equal(Object.keys(testObj).length, 0, "validateManifest() should return a new object, not mutate the original.");
```

La tua funzione `validateManifest` deve restituire il nuovo oggetto `{ containerId: "Invalid" }` quando viene chiamata con `{ containerId: null, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }`, senza modificare l'originale.

```js
const testObj = {
  containerId: null,
  destination: "Santa Cruz",
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {
  containerId: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for null containerId.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
```

La tua funzione `validateManifest` deve restituire il nuovo oggetto `{ containerId: "Invalid", destination: "Invalid", weight: "Invalid", unit: "Invalid", hazmat: "Invalid" }` quando viene chiamata con `{ containerId: 0, destination: 405, weight: -84, unit: "pounds", hazmat: "no" }`, senza modificare l'originale.

```js
const testObj = {
  containerId: 0,
  destination: 405,
  weight: -84,
  unit: "pounds",
  hazmat: "no"
};

const expected = {
  containerId: "Invalid",
  destination: "Invalid",
  weight: "Invalid",
  unit: "Invalid",
  hazmat: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, 0, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, 405, "Original object should not be mutated.");
assert.strictEqual(testObj.weight, -84, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "pounds", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, "no", "Original object should not be mutated.");
```

La tua funzione `validateManifest` deve restituire il nuovo oggetto `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` quando viene chiamata con `{ containerId: -2 }`, senza modificare l'originale.

```js
const testObj = {
  containerId: -2,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, -2, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

La tua funzione `validateManifest` deve restituire il nuovo oggetto `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` quando viene chiamata con `{ containerId: 3.50 }`, senza modificare l'originale. Puoi usare `Number.isInteger()` per validare i valori interi.

```js
const testObj = {
  containerId: 3.50,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, 3.50, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

La tua funzione `validateManifest` deve restituire il nuovo oggetto `{ containerId: "Missing", destination: "Invalid", weight: "Missing", unit: "Missing", hazmat: "Missing" }` quando viene chiamata con `{ destination: "  " }`, senza modificare l'originale. Puoi usare `.trim()` per rimuovere gli spazi da una stringa.

```js
const testObj = {
  destination: "  ",
};

const expected = {
  containerId: "Missing",
  destination: "Invalid",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.destination, "  ", "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

La tua funzione `validateManifest` deve restituire il nuovo oggetto `{ containerId: "Missing", destination: "Missing", weight: "Invalid", unit: "Missing", hazmat: "Missing" }` quando viene chiamata con `{ weight: NaN }`, senza modificare l'originale. Puoi usare `Number.isNaN()` per validare i valori `NaN`.

```js
const testObj = {
  weight: NaN,
};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Invalid",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.ok(Number.isNaN(testObj.weight), "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

La tua funzione `validateManifest` deve restituire un oggetto che descrive le proprietà mancanti e/o non valide se l'oggetto manifest in input non è valido.

```js
const solutions = [
  { containerId: 'Invalid', destination: 'Invalid', weight: 'Invalid', unit: 'Invalid' },
  { containerId: 'Invalid', destination: 'Missing', weight: 'Invalid', unit: 'Invalid', hazmat: 'Missing' },
  { containerId: 'Invalid', unit: 'Missing', hazmat: 'Invalid' },
  { containerId: 'Missing', destination: 'Invalid', weight: 'Invalid', hazmat: 'Invalid' },
  { destination: 'Invalid', weight: 'Missing', unit: 'Missing', hazmat: 'Missing' }
];

let i = 0;
for (const obj of _invalidManifests) {
  const result = validateManifest(obj);
  assert.deepEqual(
    result,
    solutions[i],
    "validateManifest() should return an object with errors for invalid manifests."
  );
  i += 1;
}
```

La tua funzione `validateManifest` deve restituire un nuovo oggetto senza modificare l'originale.

```js
const invalidOriginal = {
  containerId: -151,
  destination: "",
  weight: -151,
  unit: "kilograms",
  hazmat: "idk"
};

const validOriginal = {
  containerId: 151,
  destination: "Pallet Town",
  weight: 151,
  unit: "kg",
  hazmat: false
};

const res0 = validateManifest(invalidOriginal);
assert.isObject(res0);
assert.notStrictEqual(res0, invalidOriginal, "validateManifest() should return a new object.");
assert.strictEqual(invalidOriginal.containerId, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.destination, "", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.weight, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.unit, "kilograms", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.hazmat, "idk", "Original object should not be mutated.");
assert.equal(Object.keys(invalidOriginal).length, 5, "Original object should not be mutated.");

const res1 = validateManifest(validOriginal);
assert.isObject(res1);
assert.notStrictEqual(res1, validOriginal, "validateManifest() should return a new object.");
assert.strictEqual(validOriginal.containerId, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.destination, "Pallet Town", "Original object should not be mutated.");
assert.strictEqual(validOriginal.weight, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(validOriginal.hazmat, false, "Original object should not be mutated.");
assert.equal(Object.keys(validOriginal).length, 5, "Original object should not be mutated.");
```

Devi avere una funzione chiamata `processManifest` con un parametro `manifest`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.processManifest);
const params = explorer.allFunctions.processManifest.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), 'manifest');
```

La tua funzione `processManifest` deve prima stampare `Validation success: 55` e poi stampare `Total weight: 180 kg` quando viene chiamata con `{ containerId: 55, destination: "Carmel", weight: 400, unit: "lb", hazmat: false }`.

```js
const testObj = {
  containerId: 55,
  destination: "Carmel",
  weight: 400,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);

  const expectedCalls = [
    ["Validation success: 55"],
    ["Total weight: 180 kg"]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation success message.");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the normalized weight in kg.");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

La tua funzione `processManifest` deve prima stampare il messaggio di successo, `Validation success: ${containerId}`, se l'oggetto manifest in input è valido.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

La tua funzione `processManifest` deve normalizzare un oggetto manifest valido in input in chilogrammi usando `normalizeUnits()` e poi stampare `Total weight: ${weight} kg`.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ];
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

La tua funzione `processManifest` deve stampare un messaggio di successo con il `containerId` dell'oggetto, e poi stampare il `weight` dell'oggetto in chilogrammi, se l'oggetto manifest in input è valido. Devi usare `normalizeUnits()` per la conversione, e questo deve prevedere due chiamate a `console.log()`.

```js
const originalNormalize = normalizeUnits;
let called = false;
try {
  normalizeUnits = function (...args) {
    called = true;
    return originalNormalize(...args);
  };

  processManifest(_validManifests[0]);

  assert.isTrue(
    called,
    "processManifest() should call normalizeUnits() when the manifest is valid."
  );

} catch (err) {
  assert.fail(err);
} finally {
  normalizeUnits = originalNormalize;
}

const spy = __helpers.spyOn(console, "log");

try {
  for (const obj of _validManifests) {
    processManifest(obj);
  }
  const expectedCalls = [
    ["Validation success: 4"],
    ["Total weight: 373.95 kg"],
    ["Validation success: 5"],
    ["Total weight: 151 kg"],
  ];
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

La tua funzione `processManifest` deve prima stampare `Validation error: -88` e poi stampare l'oggetto `{ containerId: "Invalid", weight: "Invalid", unit: "Missing", hazmat: "Missing" }` quando viene chiamata con `{ containerId: -88, destination: "Soledad", weight: NaN }`.

```js
const testObj = {
  containerId: -88,
  destination: "Soledad",
  weight: NaN
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: -88"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message.");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object.");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

La tua funzione `processManifest` deve prima stampare `Validation error: undefined` e poi stampare l'oggetto `{ containerId: "Missing", weight: "Missing", unit: "Missing" }` quando viene chiamata con `{ destination: "Watsonville", hazmat: true }`.

```js
const testObj = {
  destination: "Watsonville",
  hazmat: true
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: undefined"],
    [{
        containerId: "Missing",
        weight: "Missing",
        unit: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message.");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object.");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

La tua funzione `processManifest` deve prima stampare il messaggio di errore, `Validation error: ${containerId}`, se l'oggetto manifest in input non è valido.

```js
const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

La tua funzione `processManifest` deve anche stampare l'oggetto restituito chiamando `validateManifest()` con l'oggetto manifest originale se l'oggetto manifest in input non è valido. Chiama direttamente `console.log()` con l'oggetto restituito.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest() should call validateManifest() when the manifest is invalid."
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

La tua funzione `processManifest` deve stampare un messaggio di errore con il `containerId` dell'oggetto, e poi stampare l'oggetto restituito chiamando `validateManifest()` con l'oggetto in input, se l'oggetto manifest in input non è valido. Questo deve prevedere due chiamate a `console.log()`.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest() should call validateManifest() when the manifest is invalid."
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function normalizeUnits(manifest) {
  const normalized = {
    ...manifest
  }
  if (normalized.unit === "lb") {
    normalized.weight = normalized.weight * 0.45;
    normalized.unit = "kg";
  }
  return normalized;
}

function validateManifest(manifest) {
  const result = {}

  if (manifest.containerId === undefined) {
    result.containerId = "Missing";
  } else if (
    typeof manifest.containerId !== "number" ||
    !Number.isInteger(manifest.containerId) ||
    manifest.containerId <= 0
  ) {
    result.containerId = "Invalid";
  }
  if (manifest.destination === undefined) {
    result.destination = "Missing";
  } else if (typeof manifest.destination !== "string" || manifest.destination.trim() === "") {
    result.destination = "Invalid";
  }
  if (manifest.weight === undefined) {
    result.weight = "Missing";
  } else if (typeof manifest.weight !== "number" || manifest.weight <= 0 || Number.isNaN(manifest.weight)) {
    result.weight = "Invalid";
  }
  if (manifest.unit === undefined) {
    result.unit = "Missing";
  } else if (manifest.unit !== "kg" && manifest.unit !== "lb") {
    result.unit = "Invalid";
  }
  if (manifest.hazmat === undefined) {
    result.hazmat = "Missing";
  } else if (typeof manifest.hazmat !== "boolean") {
    result.hazmat = "Invalid";
  }

  return result;
}

function processManifest(manifest) {
  const validation = validateManifest(manifest);

  if (Object.keys(validation).length > 0) {
    console.log(`Validation error: ${manifest.containerId}`);
    console.log(validation);
  } else {
    const normalized = normalizeUnits(manifest);
    console.log(`Validation success: ${manifest.containerId}`);
    console.log(`Total weight: ${normalized.weight} kg`);
  }
}
```
