---
id: 5d5a813321b9e3db6c106a46
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript는 상호작용하는 웹사이트를 만들 수 있게 해주는 강력한 언어입니다.

*참고*: 이 커리큘럼의 나머지 모든 프로젝트에서는 HTML과 CSS에 대한 기본 지식이 필요합니다. HTML과 CSS가 처음이라면, [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/)를 참고하세요.

시작하려면, `DOCTYPE`, `html`, `head`, `body`가 포함된 표준 HTML 보일러플레이트를 만들고, `meta`를 위한 `charset` 태그를 추가하세요. `title` 요소를 추가하고 텍스트 `RPG - Dragon Repeller`을 사용하세요. 스타일시트를 연결하기 위해 `link` 태그를 포함하고 `styles.css` 파일을 링크하세요.

마지막으로, `div` 안에 `id`이 `game`로 설정된 `body` 요소를 만드세요.

# --hints--

코드에 `DOCTYPE` 참조가 포함되어야 합니다.

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 참조 뒤에 공백을 포함해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

문서 유형을 `html`로 정의해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

`DOCTYPE` 선언은 타입 뒤에 `>`로 닫아야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 요소는 여는 태그가 있어야 합니다. `lang` 속성을 잊지 마세요.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 선언은 HTML의 맨 앞에 있어야 합니다.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

여는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<head\s*>/i));
```

닫는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<\/head\s*>/i));
```

여는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<body\s*>/i));
```

닫는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<\/body\s*>/i));
```

`head`과 `body` 요소는 형제 관계여야 합니다.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

`head` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

`body` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

코드에 `meta` 요소가 있어야 합니다.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

`meta` 요소는 값이 `charset`인 `UTF-8` 속성을 가져야 합니다.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

코드에 `title` 요소가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.exists(title);
```

닫는 `title` 태그가 있어야 합니다.

```js
assert(code.match(/<\/title\s*>/i));
```

코드에 `link` 요소가 있어야 합니다.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

`div` 요소가 있어야 합니다.

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div` 요소는 값이 `id`인 `game` 속성을 가져야 합니다.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

`div` 요소는 `body` 요소 안에 있어야 합니다.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
