---
id: 62a115879a6d51422652cbfc
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

역할 수행 게임에서 사용자는 자신의 능력치를 추적하고, 무기를 구매하며, 몬스터와 싸울 수 있습니다. 게임의 대화형 JavaScript 부분을 진행하기 전에, 게임 정보를 표시할 HTML 요소를 먼저 만들어야 합니다.

`div` 요소 안에 네 개의 `#game` 요소를 만드세요. 각각에 다음 순서대로 `id` 값을 부여하세요: `stats`, `controls`, `monsterStats`, `text`.

# --hints--

네 개의 새로운 `div` 요소를 만들어야 합니다.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

새로운 `div` 요소 중 하나에 `id` 속성으로 `stats` 값을 주어야 합니다.

```js
assert.exists(document.querySelector('div#stats'));
```

새로운 `div` 요소 중 하나에 `id` 속성으로 `controls` 값을 주어야 합니다.

```js
assert.exists(document.querySelector('div#controls'));
```

새로운 `div` 요소 중 하나에 `id` 속성으로 `monsterStats` 값을 주어야 합니다.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

새로운 `div` 요소 중 하나에 `id` 속성으로 `text` 값을 주어야 합니다.

```js
assert.exists(document.querySelector('div#text'));
```

새로운 `div` 요소들을 올바른 순서로 배치해야 합니다.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

새로운 `div` 요소들을 `#game` 요소 안에 배치해야 합니다.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
