---
id: 62a1166ed9a56d439c0770e7
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

`span` 요소 안에 `#stats` 요소 세 개를 만드세요. 각각에 `stat` 클래스를 지정하세요. 그런 다음 첫 번째 요소에는 `XP: 0` 텍스트를, 두 번째 요소에는 `Health: 100` 텍스트를, 세 번째 요소에는 `Gold: 50` 텍스트를 넣으세요.

# --hints--

`span` 요소 안에 `#stats` 요소 세 개가 있어야 합니다.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

새로 만든 `span` 요소 세 개에 `stat` 클래스를 지정해야 합니다.

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

첫 번째 `.stat` 요소에는 제공된 `XP: 0` 텍스트가 있어야 합니다.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

두 번째 `.stat` 요소에는 제공된 `Health: 100` 텍스트가 있어야 합니다.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

세 번째 `.stat` 요소에는 제공된 `Gold: 50` 텍스트가 있어야 합니다.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
