---
id: 62a2509ba163e020bb9d84ea
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

이제 JavaScript를 작성할 수 있습니다. `script` 요소를 만들어 시작하세요. 이 요소는 JavaScript를 HTML 파일에 불러오는 데 사용됩니다.

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

`script` 요소가 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 1);
```

`script` 요소에는 여는 태그가 있어야 합니다.

```js
assert.match(code, /<script\s*>/i);
```

`script` 요소에는 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
