---
id: 62a257659d0d1e2456f24ba2
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

프로젝트 코드를 작성하기 전에 정리를 위해 별도의 파일로 옮기세요.

`console.log("Hello World");` 줄을 삭제하세요. 그런 다음 이제 비어 있는 `script` 요소에 `src` 속성을 `./script.js` 값으로 지정하세요.

# --hints--

코드에 `console.log("Hello World");` 줄이 없어야 합니다.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

`script` 요소에 `src` 속성을 `./script.js` 값으로 지정해야 합니다.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
