---
id: 62a23c1d505bfa13747c8a9b
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

숫자 `0`, `100`, `50`를 `span` 요소로 감싸고, 그 새 `span` 요소들을 `strong` 요소로 감싸세요. 그런 다음 새 `span` 요소에 각각 `id`, `xpText`, `healthText` 값을 `goldText` 속성으로 부여하세요.

답안은 다음 기본 구조를 따라야 합니다:

```html
<span class="stat">TEXT <strong><span id="VALUE">TEXT</span></strong></span>
```

# --hints--

첫 번째 `strong` 요소 안에 `.stat` 요소를 추가해야 합니다.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

첫 번째 새 `strong` 요소는 `span` 요소를 가져야 합니다.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

첫 번째 중첩된 `span` 요소는 `id`의 `xpText`을 가져야 합니다.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'xpText');
```

첫 번째 `span` 요소는 텍스트 `0`를 감싸야 합니다.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '0');
```

첫 번째 `.stat` 요소의 텍스트는 여전히 `XP: 0`이어야 합니다.

```js
const stat = document.querySelectorAll('.stat')[0];
assert.equal(stat.innerText.trimEnd(), 'XP: 0');
```

두 번째 `strong` 요소 안에 `.stat` 요소를 추가해야 합니다.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

두 번째 새 `strong` 요소는 `span` 요소를 가져야 합니다.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

두 번째 중첩된 `span` 요소는 `id`의 `healthText`을 가져야 합니다.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'healthText');
```

두 번째 `span` 요소는 텍스트 `100`를 감싸야 합니다.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '100');
```

두 번째 `.stat` 요소의 텍스트는 여전히 `Health: 100`이어야 합니다.

```js
const stat = document.querySelectorAll('.stat')[1];
assert.equal(stat.innerText.trimEnd(), 'Health: 100');
```

세 번째 `strong` 요소 안에 `.stat` 요소를 추가해야 합니다.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

세 번째 새 `strong` 요소는 `span` 요소를 가져야 합니다.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

세 번째 중첩된 `span` 요소는 `id`의 `goldText`을 가져야 합니다.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'goldText');
```

세 번째 `span` 요소는 텍스트 `50`를 감싸야 합니다.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '50');
```

세 번째 `.stat` 요소의 텍스트는 여전히 `Gold: 50`이어야 합니다.

```js
const stat = document.querySelectorAll('.stat')[2];
assert.equal(stat.innerText.trimEnd(), 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
--fcc-editable-region--
        <span class="stat">XP: 0</span>
        <span class="stat">Health: 100</span>
        <span class="stat">Gold: 50</span>
--fcc-editable-region--
      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
