---
id: 62a23cb9bc467a147516b500
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

`#controls` 요소에 `button` 요소 세 개를 만드세요. 첫 번째는 `id`가 `button1`로 설정되고, 텍스트는 `Go to store`여야 합니다. 두 번째는 `id`가 `button2`로 설정되고, 텍스트는 `Go to cave`여야 합니다. 세 번째는 `id`가 `button3`로 설정되고, 텍스트는 `Fight dragon`여야 합니다.

# --hints--

`button` 요소에 `#controls` 요소 세 개를 추가하세요.

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

첫 번째 버튼은 `id`가 `button1`로 설정되어야 합니다.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

첫 번째 버튼의 텍스트는 `Go to store`여야 합니다.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

두 번째 버튼은 `id`가 `button2`로 설정되어야 합니다.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

두 번째 버튼의 텍스트는 `Go to cave`여야 합니다.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

세 번째 버튼은 `id`가 `button3`로 설정되어야 합니다.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

세 번째 버튼의 텍스트는 `Fight dragon`여야 합니다.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
--fcc-editable-region--
      <div id="controls">

      </div>
--fcc-editable-region--
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
