---
id: 62a23d1c5f1c93161f3582ae
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

`#stats` 요소와 비슷하게, `#monsterStats` 요소에도 두 개의 `span` 요소가 필요합니다. 이들에게 `stat` 클래스를 주고 첫 번째 요소에는 `Monster Name: ` 텍스트를, 두 번째 요소에는 `Health: ` 텍스트를 넣으세요. 각 텍스트 뒤에는 빈 중첩 `strong` 요소가 있는 `span` 요소를 추가하세요. 첫 번째 내부 `span` 요소에는 `id`의 `monsterName`를, 두 번째 내부 `span` 요소에는 `id`의 `monsterHealth`를 주세요.

# --hints--

`#monsterStats` 요소에는 두 개의 `span` 요소가 있어야 합니다.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

새로 만든 `span` 요소 두 개 모두 `class` 값을 가진 `stat`을 가져야 합니다.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

첫 번째 `span` 요소에는 빈 중첩 `strong` 요소가 있는 `span` 요소가 있어야 합니다.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

두 번째 `span` 요소에는 빈 중첩 `strong` 요소가 있는 `span` 요소가 있어야 합니다.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

첫 번째 `span` 요소에는 `Monster Name: ` 텍스트가 있어야 합니다. 공백이 올바른지 확인하세요.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+/);
```

두 번째 `span` 요소에는 `Health: ` 텍스트가 있어야 합니다. 공백이 올바른지 확인하세요.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[1].innerHTML, /Health:\s+/);
```

`strong`과 `span` 요소는 텍스트 뒤에 와야 합니다.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+<strong>/);
assert.match(spans[1].innerHTML, /Health:\s+<strong>/);
```

첫 번째 내부 `span` 요소에는 `id`의 `monsterName`가 있어야 합니다.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[0]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterName');
```

두 번째 내부 `span` 요소에는 `id`의 `monsterHealth`가 있어야 합니다.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[1]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterHealth');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
--fcc-editable-region--
      <div id="monsterStats">

      </div>
--fcc-editable-region--
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
