---
id: 62a255dae245b52317da824a
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

가장 강력한 도구 중 하나는 개발자 콘솔입니다. 사용하는 브라우저에 따라 `F12` 또는 `Ctrl+Shift+I` 키를 눌러 열 수 있습니다. Mac에서는 `Option + ⌘ + C` 키를 누르고 "Console"을 선택할 수 있습니다. 미리보기 창 위의 "Console" 버튼을 클릭해 내장 콘솔을 볼 수도 있습니다.

개발자 콘솔에는 코드에서 발생한 오류가 표시되지만, 디버깅에 도움이 되도록 코드 내 변수의 값을 확인하는 데도 사용할 수 있습니다.

`console.log("Hello World");` 태그 사이에 `script` 줄을 추가하세요. 그런 다음 "Console" 버튼을 클릭해 콘솔을 열면 `"Hello World"` 텍스트가 보여야 합니다.


# --hints--

코드에 `console.log("Hello World");` 줄을 추가해야 합니다.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

`console.log("Hello World");` 태그 사이에 `script` 줄이 있어야 합니다.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
