---
id: 62a3b79d520a7f3d0e25afd6
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript는 <dfn>문서 객체 모델</dfn>(DOM)을 사용해 HTML과 상호작용합니다. DOM은 HTML을 나타내는 객체들의 트리입니다. 전체 HTML 문서를 나타내는 `document` 객체를 사용해 HTML에 접근할 수 있습니다.

HTML에서 특정 요소를 찾는 한 가지 방법은 `querySelector()` 메서드를 사용하는 것입니다. `querySelector()` 메서드는 CSS 선택자를 인수로 받아 그 선택자와 일치하는 첫 번째 요소를 반환합니다. 예를 들어, HTML에서 `<h1>` 요소를 찾으려면 다음과 같이 작성합니다:

```js
let h1 = document.querySelector("h1");
```

`h1`는 문자열이며, 사용하려는 CSS 선택자와 일치한다는 점을 기억하세요.

`button1` 변수를 선언하고 `querySelector()`를 사용해 `id`가 `button1`인 요소를 할당하세요. CSS `id` 선택자는 `#`가 접두사로 붙는다는 점을 잊지 마세요.

# --hints--

`let`를 사용해 `button1` 변수를 선언하세요.

```js
assert.match(code, /let\s+button1/);
```

`document.querySelector()`를 사용하세요.

```js
assert.match(code, /document\.querySelector/);
```

`"#button1"` 선택자를 사용하세요.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
