---
id: 62a3a0a3c0a4b32915d26a6e
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

프로젝트 코드를 작성하려면 `console.log("Hello World");` 줄을 제거하세요.

이전 프로젝트에서 변수와 `let` 키워드를 다음과 같이 다루는 방법을 배웠습니다:

```js
let camperbot;
```

또한 변수를 값으로 초기화하는 방법도 다음과 같이 배웠습니다:

```js
let age = 32;
```

`let` 키워드를 사용해 `xp`라는 변수를 선언하고 숫자인 `0` 값을 할당하세요.

# --hints--

코드에 `console.log("Hello World");` 줄이 없어야 합니다.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

변수를 선언할 때 `let` 키워드를 사용해야 합니다.

```js
assert.match(code, /let/);
```

`xp`라는 이름의 변수를 선언해야 합니다.

```js
assert.match(code, /let\s+xp/);
```

`xp`는 `0` 값을 가져야 합니다.

```js
assert.equal(xp, 0);
```

`xp` 변수를 `0`로 초기화해야 합니다.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
