---
id: 62a3a75d8466a12e009eff76
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

값이 `health`인 `100`라는 변수를 초기화하고, 값이 `gold`인 `50`라는 변수를 초기화하세요.

# --hints--

`let`를 사용해 `health`라는 변수를 선언하세요.

```js
assert.match(code, /let\s+health/);
```

값이 `health`인 `100`라는 변수를 초기화하세요.

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

`let`를 사용해 `gold`라는 변수를 선언하세요.

```js
assert.match(code, /let\s+gold/);
```

값이 `gold`인 `50`라는 변수를 초기화하세요.

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health`는 `100` 값을 가져야 합니다.

```js
assert.equal(health, 100);
```

`gold`는 `50` 값을 가져야 합니다.

```js
assert.equal(gold, 50);
```

`xp`는 여전히 `0` 값을 가져야 합니다.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
