---
id: 62a3b41c9494f937560640ab
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

`monsterHealth`와 `inventory`라는 이름의 변수를 두 개 더 선언하세요.

`inventory` 변수에는 문자열 `"stick"`를 포함하는 배열 값을 할당하세요.

이전 프로젝트에서 배열을 다뤘던 것을 기억하세요:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

`let`를 사용해 `monsterHealth` 변수를 선언하세요.

```js
assert.match(code, /let\s+monsterHealth/i);
```

`monsterHealth` 변수 이름은 카멜 케이스를 사용하세요.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` 변수에는 값을 할당하지 마세요.

```js
assert.isUndefined(monsterHealth);
```

`monsterHealth` 변수에 값을 할당하지 마세요.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

`let`를 사용해 `inventory` 변수를 선언하세요.

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` 변수는 여전히 배열이어야 합니다.

```js
assert.isArray(inventory);
```

`inventory` 변수는 값이 하나만 있어야 합니다.

```js
assert.lengthOf(inventory, 1);
```

`inventory` 변수에는 문자열 `"stick"`가 포함되어야 합니다.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
