---
id: 62a3a7e4f1060e2fc5ffb34b
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`currentWeaponIndex`라는 변수를 하나 더 만들고 `0`로 설정하세요.

# --hints--

`let`를 사용해 `currentWeaponIndex`라는 변수를 선언하세요.

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

변수 이름은 카멜 케이스를 사용하세요.

```js
assert.match(code, /currentWeaponIndex/);
```

`currentWeaponIndex` 변수는 `0`로 설정해야 합니다.

```js
assert.equal(currentWeaponIndex, 0);
```

변수를 `0`로 초기화하세요.

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
