---
id: 62a3b3eab50e193608c19fc6
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

`fighting`라는 변수를 선언하되 값을 초기화하지 마세요. 

# --hints--

변수 `let`를 선언하려면 `fighting`을 사용하세요.

```js
assert.match(code, /let\s+fighting/);
```

`fighting` 변수에는 값이 없어야 합니다.

```js
assert.isUndefined(fighting);
```

`fighting` 변수에 값을 할당하지 마세요.

```js
assert.match(code, /let\s+fighting\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
--fcc-editable-region--
```
