---
id: 62a39f5a5790eb27c1e5d4bf
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

뷰가 새 `script.js` 파일로 전환되었습니다. 위의 탭을 사용해 파일 간 전환할 수 있다는 점을 기억하세요.

이 파일에 `console.log("Hello World");` 줄을 추가하고 콘솔에 출력되는지 확인하세요.

# --hints--

코드에 `console.log("Hello World");` 줄이 있어야 합니다.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
