---
id: 62a3c2fccf186146b59c6e96
title: 36단계
challengeType: 0
dashedName: step-36
---

# --description--

`button1`는 첫 번째 `button` 요소를 나타냅니다. 이 요소들은 `onclick`라는 특별한 속성을 가지고 있는데, 이 속성을 사용해 누군가가 그 버튼을 클릭했을 때 어떤 일이 일어날지 결정할 수 있습니다.

JavaScript에서 속성에 접근하는 방법은 여러 가지가 있습니다. 첫 번째는 <dfn>점 표기법</dfn>입니다. 다음은 점 표기법을 사용해 버튼의 `onclick` 속성을 함수 참조로 설정하는 예시입니다.

```js
button.onclick = myFunction;
```

이 예시에서 `button`는 버튼 요소이고, `myFunction`는 함수에 대한 참조입니다. 버튼이 클릭되면 `myFunction`가 호출됩니다.

점 표기법을 사용해 `onclick`의 `button1` 속성을 `goStore` 함수 참조로 설정하세요. `button1`은 이미 선언되어 있으니 `let`이나 `const`를 사용할 필요가 없습니다.

# --hints--

점 표기법을 사용해 `onclick`의 `button1` 속성에 접근해야 합니다.

```js
assert.match(code, /button1\.onclick/);
```

`let`나 `const`를 사용하면 안 됩니다.

```js
assert.notMatch(code, /(let|const)\s+button1\.onclick/);
```

`goStore` 함수 참조를 `button1.onclick`에 할당하세요. 함수를 호출하지 않도록 주의하세요.

```js
assert.match(code, /button1\.onclick\s*=\s*goStore\s*;?\s*$/m);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");

--fcc-editable-region--
// initialize buttons

--fcc-editable-region--
function goStore() {
  console.log("Going to store.");
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
