---
id: 62a7bf06d2ad9d1c5024e833
title: 40단계
challengeType: 0
dashedName: step-40
---

# --description--

버튼을 다시 클릭할 때 실행되는 함수도 업데이트해야 합니다.

`goStore()` 함수에서 각 버튼의 `onclick` 속성을 업데이트하여 각각 `buyHealth`, `buyWeapon`, `goTown`를 실행하게 하세요.

# --hints--

점 표기법을 사용해 `onclick`의 `button1` 속성에 접근하세요.

```js
assert.match(goStore.toString(), /button1\.onclick/);
```

`let`의 `const` 속성에 접근할 때 `onclick`나 `button1`를 사용하지 마세요.

```js
assert.notMatch(code, /(let|const)\s+button1\.onclick/);
```

`onclick`의 `button1` 속성을 `buyHealth`로 설정하세요.

```js
assert.match(code, /button1\.onclick\s*=\s*buyHealth/);
```

`onclick` 함수에서 `button1`의 `goStore` 속성을 설정하세요.

```js
assert.match(goStore.toString(), /button1\.onclick\s*=\s*buyHealth/);
```

점 표기법을 사용해 `onclick`의 `button2` 속성에 접근하세요.

```js
assert.match(goStore.toString(), /button2\.onclick/);
```

`let`의 `const` 속성에 접근할 때 `onclick`나 `button2`를 사용하지 마세요.

```js
assert.notMatch(code, /(let|const)\s+button2\.onclick/);
```

`onclick`의 `button2` 속성을 `buyWeapon`로 설정하세요.

```js
assert.match(code, /button2\.onclick\s*=\s*buyWeapon/);
```

`onclick` 함수에서 `button2`의 `goStore` 속성을 설정하세요.

```js
assert.match(goStore.toString(), /button2\.onclick\s*=\s*buyWeapon/);
```

점 표기법을 사용해 `onclick`의 `button3` 속성에 접근하세요.

```js
assert.match(goStore.toString(), /button3\.onclick/);
```

`let`의 `const` 속성에 접근할 때 `onclick`나 `button3`를 사용하지 마세요.

```js
assert.notMatch(code, /(let|const)\s+button3\.onclick/);
```

`onclick`의 `button3` 속성을 `goTown`로 설정하세요.

```js
assert.match(code, /button3\.onclick\s*=\s*goTown/);
```

`onclick` 함수에서 `button3`의 `goStore` 속성을 설정하세요.

```js
assert.match(goStore.toString(), /button3\.onclick\s*=\s*goTown/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function goStore() {
  button1.innerText = "Buy 10 health (10 gold)";
  button2.innerText = "Buy weapon (30 gold)";
  button3.innerText = "Go to town square";
}
--fcc-editable-region--

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
