---
id: 62a8c2bbbd8aa82052f47c53
title: 71단계
challengeType: 0
dashedName: step-71
---

# --description--

`locations` 배열에 세 번째 객체를 추가하세요. 다른 두 객체와 같은 속성을 가지게 하세요.

`name`을 `cave`로 설정하세요. `button text`을 `"Fight slime"`, `"Fight fanged beast"`, `"Go to town square"` 문자열을 가진 배열로 설정하세요. `"button functions"`을 `fightSlime`, `fightBeast`, `goTown` 변수를 가진 배열로 설정하세요. `text` 속성을 `"You enter the cave. You see some monsters."`로 설정하세요.

# --hints--

`locations` 배열에 세 개의 값이 있어야 합니다.

```js
assert.lengthOf(locations, 3);
```

세 번째 `locations` 값은 객체여야 합니다.

```js
assert.isObject(locations[2]);
```

세 번째 `locations` 객체는 `name` 속성을 가지고, 그 값은 `"cave"`이어야 합니다.

```js
assert.equal(locations[2].name, "cave");
```

세 번째 `locations` 객체는 배열인 `"button text"` 속성을 가져야 합니다.

```js
assert.isArray(locations[2]["button text"]);
```

`"button text"` 속성은 `"Fight slime"`, `"Fight fanged beast"`, `"Go to town square"` 문자열 값을 가져야 합니다.

```js
assert.equal(locations[2]["button text"][0], "Fight slime");
assert.equal(locations[2]["button text"][1], "Fight fanged beast");
assert.equal(locations[2]["button text"][2], "Go to town square");
```

세 번째 `locations` 객체는 배열인 `"button functions"` 속성을 가져야 합니다.

```js
assert.isArray(locations[2]["button functions"]);
```

`"button functions"` 속성은 `fightSlime`, `fightBeast`, `goTown` 함수 값을 가져야 합니다.

```js
assert.equal(locations[2]["button functions"][0], fightSlime);
assert.equal(locations[2]["button functions"][1], fightBeast);
assert.equal(locations[2]["button functions"][2], goTown);
```

세 번째 `locations` 객체는 문자열인 `text` 속성을 가져야 합니다.

```js
assert.isString(locations[2].text);
```

세 번째 `locations` 객체는 `text` 속성을 가지고, 그 값은 `"You enter the cave. You see some monsters."`이어야 합니다.

```js
assert.equal(locations[2].text, "You enter the cave. You see some monsters.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
--fcc-editable-region--

--fcc-editable-region--
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}

function buyHealth() {

}

function buyWeapon() {

}

function fightSlime() {

}

function fightBeast() {

}
```
