---
id: 62a8d81f539f004776dd9b1e
title: 110단계
challengeType: 0
dashedName: step-110
---

# --description--

이제 몬스터와 싸우는 코드를 시작할 수 있습니다. 코드를 깔끔하게 정리하기 위해 `fightDragon` 함수가 다른 `fight` 함수들 근처로 이동되었습니다.

`weapons` 배열 아래에 `monsters` 변수를 선언하고 배열을 할당하세요. 그 배열에는 `name`, `level`, `health` 속성을 가진 객체 세 개가 있어야 합니다. 첫 번째 객체의 값은 순서대로 `"slime"`, `2`, `15`이어야 합니다. 두 번째 객체는 `"fanged beast"`, `8`, `60`이어야 합니다. 세 번째 객체는 `"dragon"`, `20`, `300`이어야 합니다.

# --hints--

`const`를 사용해 `monsters` 변수를 선언하세요.

```js
assert.match(code, /const\s+monsters/);
```

`monsters` 변수는 배열이어야 합니다.

```js
assert.isArray(monsters);
```

`monsters` 변수에는 3개의 값이 있어야 합니다.

```js
assert.lengthOf(monsters, 3);
```

`monsters` 배열에는 3개의 객체가 있어야 합니다.

```js
assert(monsters.every(val => typeof val === "object"));
```

`monsters` 배열의 첫 번째 값은 `name` 속성이 `"slime"`으로 설정된 객체여야 합니다.

```js
assert.equal(monsters[0].name, "slime");
```

`monsters` 배열의 첫 번째 값은 `level` 속성이 `2`으로 설정된 객체여야 합니다.

```js
assert.equal(monsters[0].level, 2);
```

`monsters` 배열의 첫 번째 값은 `health` 속성이 `15`으로 설정된 객체여야 합니다.

```js
assert.equal(monsters[0].health, 15);
```

`monsters` 배열의 두 번째 값은 `name` 속성이 `"fanged beast"`으로 설정된 객체여야 합니다.

```js
assert.equal(monsters[1].name, "fanged beast");
```

`monsters` 배열의 두 번째 값은 `level` 속성이 `8`으로 설정된 객체여야 합니다.

```js
assert.equal(monsters[1].level, 8);
```

`monsters` 배열의 두 번째 값은 `health` 속성이 `60`으로 설정된 객체여야 합니다.

```js
assert.equal(monsters[1].health, 60);
```

`monsters` 배열의 세 번째 값은 `name` 속성이 `"dragon"`으로 설정된 객체여야 합니다.

```js
assert.equal(monsters[2].name, "dragon");
```

`monsters` 배열의 세 번째 값은 `level` 속성이 `20`으로 설정된 객체여야 합니다.

```js
assert.equal(monsters[2].level, 20);
```

`monsters` 배열의 세 번째 값은 `health` 속성이 `300`으로 설정된 객체여야 합니다.

```js
assert.equal(monsters[2].health, 300);
```

객체에 추가 속성을 넣지 마세요.

```js
assert.deepEqual(monsters, [
  {name: "slime", level: 2, health: 15},
  {name: "fanged beast", level: 8, health: 60},
  {name: "dragon", level: 20, health: 300}
])
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const weapons = [
  { name: 'stick', power: 5 },
  { name: 'dagger', power: 30 },
  { name: 'claw hammer', power: 50 },
  { name: 'sword', power: 100 }
];
--fcc-editable-region--

--fcc-editable-region--
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
  {
    name: "cave",
    "button text": ["Fight slime", "Fight fanged beast", "Go to town square"],
    "button functions": [fightSlime, fightBeast, goTown],
    text: "You enter the cave. You see some monsters."
  }
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  update(locations[2]);
}

function buyHealth() {
  if (gold >= 10) {
    gold -= 10;
    health += 10;
    goldText.innerText = gold;
    healthText.innerText = health;
  } else {
    text.innerText = "You do not have enough gold to buy health.";
  }
}

function buyWeapon() {
  if (currentWeaponIndex < weapons.length - 1) {
    if (gold >= 30) {
      gold -= 30;
      currentWeaponIndex++;
      goldText.innerText = gold;
      let newWeapon = weapons[currentWeaponIndex].name;
      text.innerText = "You now have a " + newWeapon + ".";
      inventory.push(newWeapon);
      text.innerText += " In your inventory you have: " + inventory;
    } else {
      text.innerText = "You do not have enough gold to buy a weapon.";
    }
  } else {
    text.innerText = "You already have the most powerful weapon!";
    button2.innerText = "Sell weapon for 15 gold";
    button2.onclick = sellWeapon;
  }
}

function sellWeapon() {
  if (inventory.length > 1) {
    gold += 15;
    goldText.innerText = gold;
    let currentWeapon = inventory.shift();
    text.innerText = "You sold a " + currentWeapon + ".";
    text.innerText += " In your inventory you have: " + inventory;
  } else {
    text.innerText = "Don't sell your only weapon!";
  }
}

function fightSlime() {

}

function fightBeast() {

}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
