---
id: 62b46e3a8d4be31be5af793d
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

약간 문제가 생겼습니다. 페이지에서 버튼 요소를 쿼리하려고 하는데 `script` 태그가 HTML의 `head` 안에 있습니다. 이는 브라우저가 HTML을 모두 읽기 전에 코드가 실행되어 `document.querySelector()`가 버튼을 찾지 못한다는 뜻입니다 — 브라우저가 아직 버튼을 처리하지 않았기 때문입니다.

이 문제를 고치려면 `script` 요소를 `head` 요소 밖으로 옮기고 `body` 요소의 끝(닫는 `</body>` 태그 바로 앞)에 배치하세요.

# --hints--

`script` 요소는 `head` 요소 안에 있으면 안 됩니다.

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

`script` 요소는 `body` 요소의 끝에 와야 합니다.

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
