---
id: 62fc20387ef88d1d1998aac5
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

`querySelector()`를 사용해 `button`인 `id`와 `button2` 두 `button3` 요소를 가져오세요. `button2`와 `button3`라는 변수에 저장하세요. `const`를 사용하는 것을 잊지 마세요.

# --hints--

`button2`로 `const` 변수를 선언하세요.

```js
assert.match(code, /const\s+button2/);
```

`button2` 변수는 `"#button2"` 요소의 값을 가져야 합니다.

```js
assert.match(code, /const\s+button2\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#button2\1\s*\)/);
```

`button3`로 `const` 변수를 선언하세요.

```js
assert.match(code, /const\s+button3/);
```

`button3` 변수는 `"#button3"` 요소의 값을 가져야 합니다.

```js
assert.match(code, /const\s+button3\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#button3\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");

--fcc-editable-region--
```
