---
id: 660880e67dfed9eb6adb7178
title: 48단계
challengeType: 0
dashedName: step-48
---

# --description--

`locations` 배열을 만들기 시작하기 전에 <dfn>객체</dfn>에 대해 먼저 배워야 합니다. 객체는 JavaScript에서 중요한 데이터 유형입니다. 다음 몇 단계는 객체에 대해 배우는 데 할애되어, 프로젝트에 어떻게 적용할지 더 잘 이해할 수 있게 됩니다.

객체는 키-값 쌍을 저장하는 비원시 데이터 유형입니다. 비원시 데이터 유형은 변경 가능한 데이터 유형으로, `undefined`, `null`, `boolean`, `number`, `string`, 또는 `symbol`가 아닙니다. 변경 가능하다는 것은 데이터가 생성된 후에도 변경될 수 있음을 의미합니다.

다음은 객체의 기본 구문입니다:

```js
{
  key: value
}
```

키와 값에 대해서는 다음 몇 단계에서 배우게 됩니다.

지금은 `const` 변수를 만들고 `cat`라는 이름으로 빈 객체 `{}`를 할당하세요.

그 아래 `cat` 변수에 대해 콘솔에서 객체를 확인하는 `console.log(cat)` 문을 추가하세요.

# --hints--

`const` 변수 `cat`가 있어야 합니다.

```js
assert.match(code, /const\s+cat\s*=/);
```

빈 객체를 `cat` 변수에 할당해야 합니다.

```js
assert.deepEqual(cat, {});
```

`console.log` 객체를 기록하는 `cat` 문이 있어야 합니다.

```js
assert.match(code, /console\.log\(cat\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const locations = [];
--fcc-editable-region--

--fcc-editable-region--

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {

}

function goTown() {
  button1.innerText = "Go to store";
  button2.innerText = "Go to cave";
  button3.innerText = "Fight dragon";
  button1.onclick = goStore;
  button2.onclick = goCave;
  button3.onclick = fightDragon;
  text.innerText = "You are in the town square. You see a sign that says \"Store\".";
}

function goStore() {
  button1.innerText = "Buy 10 health (10 gold)";
  button2.innerText = "Buy weapon (30 gold)";
  button3.innerText = "Go to town square";
  button1.onclick = buyHealth;
  button2.onclick = buyWeapon;
  button3.onclick = goTown;
  text.innerText = "You enter the store.";
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}

function buyHealth() {

}

function buyWeapon() {

}
```
