---
id: 641d9a19bff38d34d5a5edb8
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop imparerai a lavorare con la Fetch API per caricare dati da una fonte esterna e mostrarli su una pagina piena di autori di freeCodeCamp.

Tutto l'HTML e il CSS per questo workshop sono già forniti. Puoi dare un'occhiata ai due file per familiarizzare con essi.

Quando sei pronto, inizia consultando gli elementi `#author-container` e `#load-more-btn`. Ricorda che puoi usare i metodi `getElementById` o `querySelector` per farlo.

Poi assegna questi elementi a variabili chiamate rispettivamente `authorContainer` e `loadMoreBtn`.

# --hints--

Dovresti avere una variabile chiamata `authorContainer`.

```js
assert.isDefined(authorContainer);
```

Dovresti assegnare l'elemento `#author-container` alla variabile `authorContainer`.

```js
assert.strictEqual(authorContainer, document.querySelector('#author-container'));
```

Dovresti avere una variabile chiamata `loadMoreBtn`.

```js
assert.isDefined(loadMoreBtn);
```

Dovresti assegnare l'elemento `#load-more-btn` alla variabile `loadMoreBtn`.


```js
assert.strictEqual(loadMoreBtn, document.querySelector('#load-more-btn'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
