---
id: 66ebd4ae2812430bb883c787
title: 建置一個 Event Hub
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

在本實驗中，你將利用語意化 HTML 元素來創建網頁的結構。你會添加內容和圖片，使其看起來像一個真實的事件中心。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `header` 元素。
1. 在 `header` 元素內，你應該有一個包含文字 `Event Hub` 的 `h1` 元素，以及一個 `nav` 元素。
1. 在 `nav` 元素內，你應該有一個包含兩個項目的無序列表，這些項目包含指向頁面不同區段的連結。第一個項目應該有文字 `Upcoming Events`，第二個項目應該有文字 `Past Events`。
1. 每個聯結都應該由一個帶有 `href` 屬性的 `a` 元素表述，該屬性分別連結到頁面中對應的區段 `#upcoming-events` 和 `#past-events`。
1. 你應該有一個包含頁面不同區段的 `main` 元素。
1. 在 `main` 元素內，你應該有兩個 `section` 元素。
1. 第一個 `section` 元素應該有一個 `id` 屬性，值為 `upcoming-events`。
1. 在 `#upcoming-events` 區段內，你應該有：
   
   - 一個帶有文字 `Upcoming Events` 的 `h2` 元素。
   - 兩個 `article` 元素。每個 article 應該表述一個事件，並且應該有：
      - 用於事件標題的 `h3` 元素。
      - 一個用於事件描述的 `p` 元素。你可以在底部添加日期（如果你願意）。

1. 第二個 `section` 元素應該有一個 `id` 屬性，值為 `past-events`。
1. 在 `#past-events` 區段內，你應該有：

   - 一個帶有文字 `Past Events` 的 `h2` 元素。
   - 兩個 `article` 元素。每個 `article` 元素應該表述一個過去的事件，並且應該具有：
     - 用於事件標題的 `h3` 元素，
     - 一個用於事件描述的 `p` 元素。你可以在底部添加日期（如果你願意）。
     - 一個影像元素，其 `src` 屬性指向一個影像檔案，且 `alt` 屬性包含影像的描述。

**注意：**你可以使用任何文字作為事件描述和日期。你可以使用以下影像 URL 作為影像，如果你願意：

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`。
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`。

# --hints--

你應該有一個 `header` 元素。

```js
assert.isNotNull(document.querySelector("header"));
```

你的 `header` 元素應該位於開啟的 `body` 標籤之後。

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

在 `header` 元素內，你應該有一個包含文字 `Event Hub` 的 `h1` 元素。

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

在 `header` 元素內，`h1` 元素之後，你應該有一個 `nav` 元素。

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

你的 `nav` 元素應該包含一個有兩個項目的無序列表。

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

無序列表中的第一個項目應該是一個聯結。

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

無序列表中的第二個項目應該是一個聯結。

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

無序列表中第一個項目的文字應該是 `Upcoming Events`。

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

無序列表中的第一個項目應將 `href` 設定為 `#upcoming-events`。

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

無序列表中第二個條款的文字應該是 `Past Events`。

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

無序列表中的第二個項目應將 `href` 設定為 `#past-events`。

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

你應該在 `header` 元素結束標籤之後有一個 `main` 元素。

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

在 `main` 元素內，你應該有兩個 `section` 元素。

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

你的第一個 `section` 元素應該有一個 `id` 屬性，值為 `upcoming-events`。

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

你的第二個 `section` 元素應該有一個 `id` 屬性，值為 `past-events`。

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

在 `#upcoming-events` 區段內，你應該有一個帶有文字 `Upcoming Events` 的 `h2` 元素。

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

在 `#upcoming-events` 區段內，`h2` 元素下方應該有兩個 `article` 元素。

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`#upcoming-events` 區段內的兩個 `article` 元素都應該有一個用於事件標題的 `h3` 元素。

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`#upcoming-events` 區段內的兩個 `article` 元素都應該有一個段落元素來描述該事件。

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

在 `#past-events` 區段內，你應該有一個 `h2` 元素，其文字為 `Past Events`。

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

在 `#past-events` 區段內，`h2` 元素下方應該有兩個 `article` 元素。

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`#past-events` 區段內的兩個 `article` 元素都應該有一個用於事件標題的 `h3` 元素。

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`#past-events` 區段內的兩個 `article` 元素都應該有一個段落元素來描述該事件。

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

`#past-events` 區段內的兩個 `article` 元素都應該有一個影像元素。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

`#past-events` 區段內的兩個影像元素都應該有指向影像檔案的 `src` 屬性。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

`#past-events` 區段內的兩個影像元素都應該有帶有影像描述的 `alt` 屬性。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

每個 `h3` 元素都應該有事件標題。

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

每個 `p` 元素都應該有事件描述。

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```

