---
id: 66ebd4ae2812430bb883c787
title: 創建一個活動中心
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

在本實驗室中，你將利用語義化 HTML 元素創建網頁結構。 你將添加內容和圖片，使其看起來像一個真正的活動中心。

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶故事：**

1. 你應該有一個 `header` 元素。
1. `header` 元素內部，應該有一個包含文本 `Event Hub` 的 `h1` 元素和一個 `nav` 元素。
1. 在 `nav` 元素中，應該有一個由兩個項目組成的無序列表，其中包含指向頁面不同部分的鏈接。 第一個列表項應包含文本 `Upcoming Events`，第二個列表項應包含文本 `Past Events`。
1. 每個鏈接都應由一個帶有 `href` 屬性的 `a` 元素表示，該元素鏈接到頁面的相應部分，分別是 `#upcoming-events` 和 `#past-events`。
1. 你應該有一個 `main` 元素，其中包含頁面的不同部分。
1. 在 `main` 元素內部，應該有兩個 `section` 元素。
1. 第一個 `section` 元素應具有 `id` 屬性，其值爲 `upcoming-events`。
1. 在 `#upcoming-events` 部分中，應該有：

   - 包含文本 `Upcoming Events` 的 `h2` 元素。
   - 兩個 `article` 元素。 每篇文章都應代表一個活動，並應包含 ：
      - 一個用於活動標題的 `h3` 元素。
      - 一個用於活動描述的 `p` 元素。 你還可以在底部添加日期。

1. 第二個 `section` 元素應具有 `id` 屬性，其值爲 `past-events`。
1. 在 `#past-events` 部分中，應該有：

   - 包含文本 `Past Events` 的 `h2` 元素。
   - 兩個 `article` 元素。 每個文章元素都應代表一個已結束的活動，而且應該有：
     - 一個用於活動標題的 `h3` 元素，
     - 一個用於活動描述的 `p` 元素。 你還可以在底部添加日期。
     - 一個圖像元素，其 `src` 屬性指向圖像文件，`alt` 屬性包含圖像描述。

**注意：**你可以使用任何文本來描述活動和日期。 你可以使用以下 URL 獲取圖像：

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`。
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`。

# --hints--

你應該有一個 `header` 元素。

```js
assert.isNotNull(document.querySelector("header"));
```

你的 `header` 元素應放在 `body` 開始標籤之後。

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

在 `header` 元素內部，應該有一個 `h1` 元素，其中包含文本 `Event Hub`。

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

在 `header` 元素內部，在 `h1` 元素之後，你應該有一個 `nav` 元素。

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

你的 `nav` 元素應包含一個由兩個項目組成的無序列表。

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

無序列表中的第一個項目應該是一個鏈接。

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

無序列表中的第二個項目應該是一個鏈接。

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

無序列表中第一個項目的文本應爲 `Upcoming Events`。

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

無序列表中的第一個項目應將 `href` 設置爲 `#upcoming-events`。

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

無序列表中第二個項目的文本應爲 `Past Events`。

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

無序列表中的第二個項目應將 `href` 設置爲 `#past-events`。

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

你應該在 `header` 元素結束標籤之後添加一個 `main` 元素。

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

在 `main` 元素內部，應該有兩個 `section` 元素。

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

你的第一個 `section` 元素應包含一個 `id` 屬性，值爲 `upcoming-events`。

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

你的第二個 `section` 元素應包含一個 `id` 屬性，值爲 `past-events`。

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

在 `#upcoming-events` 部分中，應該有一個 `h2` 元素，其文本爲 `Upcoming Events`。

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

在 `#upcoming-events` 部分中，在 `h2` 元素下面應該有兩個 `article` 元素。

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`#upcoming-events` 部分中的兩個 `article` 元素都應有一個 `h3` 元素用於活動標題。

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`#upcoming-events` 部分中的兩個 `article` 元素都應有一個段落元素用於活動描述。

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

在 `#past-events` 部分中，應該有一個 `h2` 元素，其文本爲 `Past Events`。

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

在 `#past-events` 部分中，在 `h2` 元素下面應該有兩個 `article` 元素。

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`#past-events` 部分中的兩個 `article` 元素都應有一個 `h3` 元素用於活動標題。

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`#past-events` 部分中的兩個 `article` 元素都應有一個段落元素用於活動描述。

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

`#past-events` 部分中的兩個 `article` 元素都應包含一個圖像元素。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

`#past-events` 部分中的兩個圖像元素的 `src` 屬性都應指向一個圖像文件。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

`#past-events` 部分中的兩個圖像元素都應具有 `alt` 屬性，其中包含圖像描述。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

每個 `h3` 元素都應包含活動標題。

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

每個 `p` 元素都應包含活動描述。

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```

