---
id: 66b205e6eacba4c4e54ea434
title: Jenga kificha barua pepe
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

Katika maabara hii, utaficha sehemu ya jina la mtumiaji katika anwani ya barua pepe kwa nyota. Kuficha ni neno linalotumika kuficha au kubadilisha taarifa nyeti kwa nyota au herufi nyingine.

Kwa mfano, ikiwa anwani ya barua pepe ilikuwa `myEmail@email.com`, basi anwani ya barua pepe iliyofichwa itakuwa `m*****l@email.com`.

**Lengo:** Timiza hadithi za mtumiaji zilizo hapa chini na upite mitihani yote ili kumaliza maabara hii.

**Hadithi za mtumiaji:**

1. Tengeneza kitendakazi kinachoitwa `maskEmail` kinachopokea `email` kama hoja.
2. Ndani ya kitendakazi, unapaswa kuficha `email` na kuambatanisha jina la kikoa nalo. Kumbuka unaweza kutumia njia kama `slice`, `repeat`, `indexOf` au hata `replace` kusaidia.
3. Nje ya kitendakazi, tangaza kigezo kinachoitwa `email` kuhifadhi anwani ya barua pepe unayotaka kuficha.
4. Fanya wito wa kitendakazi cha `maskEmail` kwa kigezo cha `email` na toa matokeo kwenye konsoli.
5. `maskEmail("apple.pie@example.com")` inapaswa kurudisha `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` inapaswa kurudisha `"f**********p@example.com"`.
7. `maskEmail("info@test.dev")` inapaswa kurudisha `"i**o@test.dev"`.
8. `maskEmail("user@domain.org")` inapaswa kurudisha `"u**r@domain.org"`.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `maskEmail`.

```js
assert.isFunction(maskEmail);
```

Kitendakazi cha `maskEmail` kinapaswa kupokea mfuatano wa herufi, `email`, kama hoja.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

Nje ya kitendakazi, unapaswa kuwa na kigezo cha `email`.

```js
assert.isDefined(email);
```

Unapaswa kuweka thamani ya anwani halali ya barua pepe kwenye kigezo chako cha `email`.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` inapaswa kurudisha `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` inapaswa kurudisha `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` inapaswa kurudisha `"i**o@test.dev"`.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` inapaswa kurudisha `"u**r@domain.org"`.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

`maskEmail` yako inapaswa kutoa matokeo sahihi.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

Unapaswa kuandika matokeo ya wito wa `maskEmail` ukiwa na `email` kama hoja.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
