---
id: 66faac4139dbbd5dd632c860
title: Construir un formulario de solicitud de trabajo
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**Objetivo:** Cumplir las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Instrucciones:**

1. Debe tener un `div` elemento con la clase `container`.
1. Dentro del elemento `div`, debe tener un elemento `form`.
1. El formulario debe contener un elemento `input` con el tipo `text` y el id `name` para ingresar el nombre completo del usuario.
1. Debe tener otro elemento `input` con el tipo `email` y el id `email` para ingresar la dirección de correo electrónico del usuario.
1. El formulario debe incluir un elemento `select` con el id `position` que permita a los usuarios seleccionar un puesto de trabajo.
1. Debe tener un elemento `fieldset` con clase de `radio-group`.
1. Dentro de `.radio-group` debe tener un conjunto de elementos `input` con el tipo `radio` y etiquetas relevantes para seleccionar opciones de disponibilidad (por ejemplo, tiempo completo, tiempo parcial). El nombre del grupo debe ser `availability`.
1. Debe tener un elemento `textarea` con el id `message` para ingresar un mensaje.
1. Debe asociar cada elemento `input` con un elemento `label`.
1. Debe tener un elemento `button` con el tipo `submit` para enviar el formulario.
1. Agregue una pseudo-clase `:focus` a los elementos `input` y `textarea` para cambiar el color del borde cuando estén enfocados.
1. Los elementos `input`, `select` y `textarea` deben tener una pseudo-clase `:invalid` que cambie el color del borde a rojo cuando se detecte una entrada inválida.
1. Los elementos `input`, `select` y `textarea` deben tener una pseudo-clase `:valid` que cambie el color del borde a verde cuando se ingrese una entrada válida.
1. El elemento `button` debe tener una pseudo-clase `:hover` que cambie el color de fondo cuando se pase el ratón sobre él.
1. Use la pseudo-clase `:checked` en `.radio-group input[type="radio"]` para agregar un color de borde, color de fondo y una sombra de caja cuando se seleccione el botón de radio.
1. Use la pseudo-clase `:checked` en botones de radio para cambiar el color de texto de la etiqueta asociada cuando se seleccione la opción.
1. Agregue una pseudo-clase `:first-of-type` al elemento `input` para estilizar el primer campo de entrada de manera diferente. (por ejemplo, bordes redondeados).

**Nota:** Asegúrese de vincular su hoja de estilos en su HTML y aplicar su CSS.

# --hints--

Debe tener un `div` elemento con la clase `container`.

```js
assert.exists(document.querySelector("div.container"));
```

Dentro del elemento `div`, debe tener un elemento `form`.

```js
assert.exists(document.querySelector("div.container > form"));
```

El formulario debe contener un elemento `input` con el tipo `text` y el id `name` para ingresar el nombre completo del usuario.

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

Debe tener otro elemento `input` con el tipo `email` y el id `email` para ingresar la dirección de correo electrónico del usuario.

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

El formulario debe incluir un elemento `select` con el id `position` con algunos elementos `option`.

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

Debe tener un elemento `fieldset` o `section` con la clase `.radio-group`.

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

Dentro de `.radio-group` debe tener un grupo de elementos `input` con el tipo `radio` para seleccionar opciones de disponibilidad. El nombre del grupo debe ser `availability`.

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

Debe tener un elemento `textarea` con el id `message` para ingresar un mensaje.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

Debe asociar cada elemento `input` con un elemento `label`.

```js
const els = document.querySelectorAll('input');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

Debe tener un elemento `button` con el tipo `submit` para enviar el formulario.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

Debe agregar una pseudo-clase `:focus` a los elementos `input` y `textarea` para cambiar su color de borde cuando estén enfocados. Use un selector de lista en el orden dado.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'))
```

Los elementos `input`, `select` y `textarea` deben tener una pseudo-clase `:invalid` que cambie el color del borde a rojo cuando se detecte una entrada inválida. Use un selector de lista en el orden dado.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

Los elementos `input`, `select` y `textarea` deben tener una pseudo-clase `:valid` que cambie el color del borde a verde cuando se ingrese una entrada válida. Use un selector de lista en el orden dado.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

El elemento `button` debe tener una pseudo-clase `:hover` que cambie el color de fondo cuando se pase el ratón sobre él.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

Debe usar la pseudo-clase `:checked` en `.radio-group input[type="radio"]` para agregar un color de borde cuando se seleccione el botón de radio.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

Debe usar la pseudo-clase `:checked` en `.radio-group input[type="radio"]` para agregar un color de fondo cuando se seleccione el botón de radio.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

Debe usar la pseudo-clase `:checked` en `.radio-group input[type="radio"]` para agregar una sombra de caja cuando se seleccione el botón de radio.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

Debe usar la pseudo-clase `:checked` en botones de radio para cambiar el color de texto de la etiqueta asociada cuando se seleccione la opción.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

Agregue una pseudo-clase `:first-of-type` a los elementos `input` para estilizar el primer campo de entrada de manera diferente.

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
