---
id: bad87fee1348bd9aec908850
title: applica lo stile predefinito bootstrap per i pulsanti
challengeType: 0
forumTopicId: 16657
dashedName: apply-the-default-bootstrap-button-style
---

# --description--

Bootstrap ha un'altra classe per i pulsanti chiamata `btn-default`.

Applica sia la classe `btn` sia la classe `btn-default` a ciascuno dei tuoi elementi `button`.

# --hints--

Dovresti applicare la classe `btn` a ciascuno dei tuoi elementi `button`.

```js
assert.lengthOf(document.querySelectorAll('.btn'),6);
```

Dovresti applicare la classe `btn-default` a ciascuno dei tuoi elementi `button`.

```js
assert.lengthOf(document.querySelectorAll('.btn-default'), 6);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
      </div>
    </div>
  </div>
</div>
```
