---
id: bad87fee1348bd9aec908846
title: creare un titolo con bootstrap
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

Ora costruiamo qualcosa da zero per esercitarci con HTML, CSS e bootstrap.

Creeremo un'area di prova jQuery, che presto useremo nelle nostre sfide jQuery.

Per cominciare, crea un elemento `h3` con il testo `jQuery Playground`.

Colora il tuo elemento `h3` con la classe bootstrap `text-primary` e centrala con la classe bootstrap `text-center`.

# --hints--

Dovresti aggiungere un elemento `h3` alla tua pagina.

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

Il tuo elemento `h3` dovrebbe avere un tag di chiusura.

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

Il tuo elemento `h3` dovrebbe essere colorato applicando la classe `text-primary`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

Il tuo elemento `h3` dovrebbe essere centrato applicando la classe `text-center`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

Il tuo elemento `h3` dovrebbe contenere il testo `jQuery Playground`.

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```
