---
id: 60eebd07ea685b0e777b5583
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the Registration Form project! Start by adding the `!DOCTYPE html` declaration at the top of the document so the browser knows what type of document it's reading.

# --hints--

Il codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti aggiungere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire il tipo di documento come `html`.

```js
assert(code.match(/html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un carattere `>`.

```js
assert(code.match(/html\s*>/gi));
```

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio dell'HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
