---
id: 68e4276483b0c71658aab1ce
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

有時你會只想要修剪字串開頭的空白字元。這時候 `trimStart()` 方法就非常有用。

這裡有一個範例：

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

請記住，字串是不可變的。此方法將創建一個包含結果的新字串，而不會修改原本的字串。

創建一個變數 `trimmedStart` 並將使用 `trimStart()` 方法於 `userInput` 變數的結果指定給它。

# --hints--

你應該有一個 `trimmedStart` 變數。

```js
assert.exists(trimmedStart);
```

你的 `trimmedStart` 變數應該是一個字串。

```js
assert.isString(trimmedStart);
```

你應該將 `userInput.trimStart()` 指定到你的 `trimmedStart` 變數。

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
