---
id: 68e4276483b0c71658aab1d2
title: 步驟 13
challengeType: 1
dashedName: step-13
---

# --description--

有時你需要使用多個方法來創建新的字串。在接下來的幾個步驟中，你將使用迄今為止學到的一些方法來創建一個使用 camelCase 命名慣例的新字串。

首先建立一個名為 `lowercaseWord` 的變數，並指定其值為字串 `"camelcase"`。 

在 `lowercaseWord` 變數下方，創建另一個名為 `camelCasedVersion` 的變數，並暫時將它指定為空字串。

在 `camelCasedVersion` 變數下方，為你的主控台添加一個帶有字串 `"Camel cased version:"` 的 `console.log()`，接著再為 `camelCasedVersion` 變數添加另一個 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該有一個名為 `lowercaseWord` 的變數。

```js
assert.exists(lowercaseWord);
```

你的 `lowercaseWord` 應該是一個字串。

```js
assert.isString(lowercaseWord);
```

你應該將字串 `"camelcase"` 指定給你的 `lowercaseWord` 變數。

```js
assert.equal(lowercaseWord, "camelcase");
```

你應該有一個 `camelCasedVersion` 變數。

```js
assert.exists(camelCasedVersion);
```

你的 `camelCasedVersion` 變數應該是一個字串。

```js
assert.isString(camelCasedVersion);
```

你應該將空字串指定給你的 `camelCasedVersion` 變數。

```js
assert.isEmpty(camelCasedVersion);
```

你應該將字串 `"Camel cased version:"` 記錄到主控台。

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

你應該將 `camelCasedVersion` 變數記錄到主控台。

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
