---
id: 68e4276483b0c71658aab1ca
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將練習使用各種用於格式化字串的 `string methods`。

首先建立一個名為 `userInput` 的變數，並指定其值為 `"   Hello World!   "`。

# --hints--

你應該有一個 `userInput` 變數。

```js
assert.exists(userInput);
```

你的 `userInput` 變數應該是一個字串。

```js
assert.isString(userInput);
```

你應該將 `"   Hello World!   "` 指定到你的 `userInput` 變數。請確保在 `H` 前有三個空格，在 `!` 後也有三個空格。

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
