---
id: 68e4276483b0c71658aab1d3
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

如你從先前的課程和工作坊中回想，你可以使用 `slice()` 方法從字串中擷取部分內容。

這裡有一個提醒：

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

將 `camelCasedVersion` 的值從空字串改為對 `lowercaseWord` 變數使用 `slice()` 方法的結果。傳入數字 `0` 和 `5`，它們分別表述 `slice()` 方法的起始和結束索引。

現在你應該會在主控台看到字 `"camel"`。

# --hints--

你的 `camelCasedVersion` 變數不應該是空字串。

```js
assert.isNotEmpty(camelCasedVersion);
```

你應該將 `lowercaseWord.slice(0, 5)` 指定給 `camelCasedVersion` 變數。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

你不應該將字面字串 `"camel"` 指定給 `camelCasedVersion`。請重新閱讀指示以了解正確的方法。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
