---
id: 68e4276483b0c71658aab1d1
title: 步驟 12
challengeType: 1
dashedName: step-12
---

# --description--

如果你需要創建一個所有字母皆為小寫的字串，你可以使用 `toLowerCase()` 方法。  

首先建立一個名為 `lowerCaseInput` 的變數，並將對 `cleanedInput` 變數使用 `toLowerCase()` 方法的結果指定給它。 

然後記錄字串 `"Result of using the toLowerCase() method:"`。

最後，為你的 `lowerCaseInput` 變數添加另一個 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該有一個 `lowerCaseInput` 變數。

```js
assert.exists(lowerCaseInput);
```

你的 `lowerCaseInput` 變數應該是一個字串。

```js
assert.isString(lowerCaseInput);
```

你應該將 `cleanedInput.toLowerCase()` 指定到你的 `lowerCaseInput` 變數。

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

你應該將字串 `"Result of using the toLowerCase() method:"` 記錄到主控台。

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

你應該將 `lowerCaseInput` 變數記錄到主控台。

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
