---
id: 68e4276483b0c71658aab1d0
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

有時你需要創建一個不同大小寫的字串。 

要創建一個所有字母都轉換為大寫的新字串，你可以使用 `toUpperCase()` 方法。

這裡有一個範例：

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

創建一個名為 `upperCaseInput` 的變數，並將使用 `toUpperCase()` 方法於 `cleanedInput` 變數上的結果指定給它。

# --hints--

你應該有一個 `upperCaseInput` 變數。

```js
assert.exists(upperCaseInput);
```

你的 `upperCaseInput` 變數應該是一個字串。

```js
assert.isString(upperCaseInput);
```

你應該將 `cleanedInput.toUpperCase()` 指定到你的 `upperCaseInput` 變數。

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
