---
id: 68e4276483b0c71658aab1cd
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

現在是時候查看結果。

首先為你的程式碼添加一個帶有字串 `"Result of trimming whitespace from both ends:"` 的 `console.log()`。

然後在那個 `console.log()` 下面，為你的 `cleanedInput` 變數添加另一個 `console.log()`。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將字串 `"Result of trimming whitespace from both ends:"` 記錄到主控台。

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

你應該將 `cleanedInput` 變數記錄到主控台。

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
