---
id: 68e4276483b0c71658aab1cc
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

你的 `userInput` 值在字串的兩端有一些多餘的空格。若要創建一個移除該空白字元的新字串，你可以使用 `trim()` 方法。  

這裡有一個範例：

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

創建一個名為 `cleanedInput` 的變數，並將使用 `trim()` 方法於 `userInput` 變數的結果指定給它。

# --hints--

你應該有一個 `cleanedInput` 變數。

```js
assert.exists(cleanedInput);
```

你的 `cleanedInput` 變數應該是一個字串。

```js
assert.isString(cleanedInput);
```

你應該將 `userInput.trim()` 指定到你的 `cleanedInput` 變數。

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
