---
id: 68f1179eaea3bb59ba434386
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

要去除字串末端的空白字元，你可以使用 `trimEnd()` 方法。

這裡有一個範例：

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

創建一個名為 `trimmedEnd` 的變數，並將對 `userInput` 變數使用 `trimEnd()` 方法的結果指定給它。 

# --hints--

你應該有一個 `trimmedEnd` 變數。

```js
assert.exists(trimmedEnd);
```

你的 `trimmedEnd` 變數應該是字串。

```js
assert.isString(trimmedEnd);
```

你應該將 `userInput.trimEnd()` 指定到你的 `trimmedEnd` 變數。

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
