---
id: 68f3f3d65d51d80ac0b412fc
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

現在是時候查看在字串開頭修剪空白的結果。

在你的 `trimmedStart` 變數下方，添加一個帶有字串 `"After using the trimStart() method, leading spaces removed:"` 的 `console.log()`。 

然後在 `console.log()` 下面，為你的 `trimmedStart` 變數添加另一個 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將字串 `"After using the trimStart() method, leading spaces removed:"` 記錄到主控台。

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

你應該將 `trimmedStart` 變數記錄到主控台。

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
