---
id: 68f3f5ca1d2a490c97f14f03
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

現在是時候查看在字串末端修剪空白的結果。

在你的 `trimmedEnd` 變數下方，添加一個帶有字串 `"After using the trimEnd() method, trailing spaces removed:"` 的 `console.log()`。 

然後在 `console.log()` 下面，為你的 `trimmedEnd` 變數添加另一個 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將字串 `"After using the trimEnd() method, trailing spaces removed:"` 記錄到主控台。

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

你應該將 `trimmedEnd` 變數記錄到主控台。

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
