---
id: 68f3eeaa65614107ecd5f1f2
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

在你的第一個 `console.log()` 下面，為你的主控台添加另一個 `console.log()` 以記錄 `userInput` 變數。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將 `userInput` 變數記錄到主控台。

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
