---
id: 68f3f890010c870e17bbe754
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

現在是時候查看結果。

在你的 `upperCaseInput` 變數下方，為你的 `console.log()` 添加字串 `"Result of using the toUpperCase() method:"`。 

然後在 `console.log()` 下面，為你的 `console.log()` 添加另一個帶有 `upperCaseInput` 變數的 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將字串 `"Result of using the toUpperCase() method:"` 記錄到主控台。

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

你應該將 `upperCaseInput` 變數記錄到主控台。

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
