---
id: 68f405c8c9df44132c8fe1bb
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

如你從先前課程回想，camelCase 命名慣例以第一個字的英文字母小寫開始，接著每個後續的字以大寫字母開頭，後面跟著小寫字母。

以下是一些提醒範例： 

```md
freeCodeCamp
camelCase
learningIsFun
```

`lowercaseWord` 變數中的第二個字是 `"case"`。若要存取該字中的 `c`，你可以使用 `lowercaseWord[5]`。 

使用 `+` 運算子將 `lowercaseWord.slice(0, 5)` 與對 `lowercaseWord[5]` 使用正確方法將字串轉換為大寫的結果串接起來。

現在你應該會在主控台看到 `camelC`。

# --hints--

你應該將 `lowercaseWord.slice(0, 5)` 與 `lowercaseWord[5].toUpperCase()` 串接起來。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

你不應該將字面字串 `"camelC"` 指定給 `camelCasedVersion`。請重新閱讀指示以取得正確的方法。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

