---
id: 68f40925d810ae14b2d5a169
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

目前你的 `camelCasedVersion` 變數持有 `camelC` 的值。但期望的結果是 `camelCase`。 

要取得剩餘的 `ase` 字母，你可以使用帶有負索引值的 `slice()` 方法。請記住，使用負索引值會告訴 `slice()` 方法從結尾開始擷取。 

這裡是一個提醒範例：

```js
const str = "abcd";
str.slice(-2); // cd
```

使用 `+` 運算子將 `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` 與對 `lowercaseWord` 使用 `slice()` 方法的結果串接。為 `slice()` 方法傳入數字 `-3`。

隨著最後一次變更，你的工作坊已完成！

# --hints--

你應該將 `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` 與 `lowercaseWord.slice(-3)` 串接起來。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

你不應該將字面字串 `"camelCase"` 指定給 `camelCasedVersion`。請重新閱讀指示以了解正確的方法。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
