---
id: 646c59e942f35541923104bf
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Torna nel tuo HTML e crea un elemento `main`. All'interno di quell'elemento `main`, aggiungi un elemento `div` con la classe `cat-head`.

# --hints--

Dovresti avere un elemento `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Dovresti avere un elemento `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

Il tuo elemento `div` dovrebbe avere la classe `cat-head`.

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

Il tuo elemento `div` dovrebbe essere all'interno del tuo tag `main`.

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
