---
id: 646c48df8674cf2b91020ecc
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Ti ├© stato fornito un codice boilerplate di base. Inizia aggiungendo un elemento `link` all'interno del tuo elemento `head`. Per quell'elemento `link`, imposta l'attributo `rel` su `stylesheet` e l'attributo `href` su `./styles.css`.

# --hints--

Dovresti avere un elemento `link`.

```js
assert.exists(document.querySelector("head>link"));
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` e `href` impostati su `stylesheet` e `./styles.css`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
