---
id: 6476f5c17f99146071ee884c
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Potresti aver notato che non succede nulla. L'elemento `.cat-head` non si è mosso nonostante l'impostazione di un `top` e `left` di `100px` ciascuno. Ma non è così con il posizionamento `relative`.

Quando usi il valore `relative`, l'elemento viene ancora posizionato secondo il flusso normale del documento, ma diventano attivi i valori `top`, `left`, `bottom` e `right`.

Invece di `static`, imposta la proprietà position del tuo `.cat-head` su `relative` e lascia inalterate sia la proprietà `top` sia la proprietà `left`.

# --hints--

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `position` impostata su `relative`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `top` impostata su `100px`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `left` impostata su `100px`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
