---
id: 6476f7a4827bcc61682f2347
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Quando usi il valore `absolute` per la proprietà `position`, l'elemento viene rimosso dal flusso normale del documento, e la sua posizione viene poi determinata dalle proprietà `top`, `right`, `bottom` e `left`.

Imposta la proprietà `position` del tuo elemento `.cat-head` su `absolute`, poi imposta le proprietà `top` e `left` su un qualsiasi valore positivo in pixel.

<!-- **Nota**: Puoi sperimentare con le proprietà `top`, `left`, `bottom` e `right` qui, ma il test passerà solo per `top` di `300px` e left di `400px`. -->

# --hints--

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `position` impostata su `absolute`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `top` impostata su un qualsiasi numero in `px`. Assicurati di aggiungere un punto e virgola.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `left` impostata su un qualsiasi numero in `px`. Assicurati di aggiungere un punto e virgola.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
