---
id: 6476fb7cbaafa36d65e9cf35
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

`fixed` è un valore della proprietà `position` che ti permette di fissare un elemento alla pagina indipendentemente da dove l’utente scorre.

Dovrai fare qualche markup in più per vedere come funziona il posizionamento `fixed`. Nel tuo HTML, crea un elemento `div` con la classe `box`.

# --hints--

Dovresti avere un elemento `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Il tuo elemento `div` dovrebbe avere la classe `box`.

```js
assert.equal(document.querySelectorAll('div')[1]?.getAttribute('class'), 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
