---
id: 64770351e8586671ec0911f0
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Il posizionamento `sticky` è un ibrido tra il posizionamento `relative` e `fixed`. Permette a un elemento di **attaccarsi** a una posizione specifica all'interno del suo elemento contenitore o della viewport, in base alla posizione dello scroll.

Cambia il valore della proprietà `position` di `.cat-head` in `sticky`, imposta `top` su `0`, poi rimuovi `left` e il suo valore.

**Nota**: Per vedere come funziona `sticky`, devi posizionare un paio di testi prima e dopo il tuo elemento `.cat-head`. Se scorri verso il basso dopo, vedrai che `.cat-head` si attacca in cima e rimane lì.

# --hints--

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `position` impostata su `sticky`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'sticky');
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `top` impostata su `0`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

Non dovresti avere la proprietà `left` e il suo valore nel tuo codice.

```js
assert.notMatch(code, /left:\s*100px;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
