---
id: 6477062778c85972eb648030
title: Passaggio 14
challengeType: 0
dashedName: step-14
---

# --description--

Ora dovresti centrare la testa del gatto.

Dai all'elemento `.cat-head` una proprietà `position` impostata su `absolute`. Imposta un valore `0` per le proprietà `right`, `left`, `top`, `bottom`, poi imposta la sua proprietà `margin` su tutti i lati a `auto`. Questo è un modo per centrare un elemento verticalmente e orizzontalmente usando il posizionamento CSS.

# --hints--

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `position` impostata su `absolute`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `top` impostata su `0`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `left` impostata su `0`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '0px');
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `right` impostata su `0`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right, '0px');
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `bottom` impostata su `0`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom, '0px');
```

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `margin` impostata su `auto`. Assicurati di aggiungere un punto e virgola.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
