---
id: 647d855ac12cd436059acd39
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Usa il selettore universale per aggiungere `box-sizing: border-box;` al tuo CSS. Questo assicura che gli elementi includano padding e bordo nella larghezza e altezza specificate.

# --hints--

Dovresti avere il selettore universale (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Il selettore universale dovrebbe avere una proprietà `box-sizing` impostata su `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
