---
id: 6676a8b01e56ec1a1e07c202
title: Passaggio 21
challengeType: 0
dashedName: step-21
---

# --description--

Nota che ora hai un quadrato bianco con bordi spessi. I bordi hanno angoli diagonali che possono essere usati per creare triangoli. Ti serviranno per creare le orecchie del tuo gatto.

All'interno dello stesso selettore di classe, regola `height` e `width` in `0`. Imposta il bordo sinistro, destro e superiore su `transparent`.

Rimuovi la proprietà `background-color` e dovresti vedere un triangolo.

# --hints--


Dovresti avere un selettore `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `height` impostata su `0`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '0px')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `width` impostata su `0`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '0px')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `border-left` impostata su `35px solid transparent`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid transparent')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `border-right` impostata su `35px solid transparent`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid transparent')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `border-top` impostata su `35px solid transparent`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid transparent')
```


Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `border-bottom` impostata su `35px solid red`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

Il tuo selettore `.cat-right-ear` non dovrebbe avere una proprietà `background-color`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
