---
id: 6676a8a8d9861319ee901a90
title: Passaggio 22
challengeType: 0
dashedName: step-22
---

# --description--

Ora puoi iniziare a lavorare sulle orecchie del tuo gatto.

Pulisci il tuo codice di revisione rimuovendo il selettore `.cat-right-ear` e tutte le sue proprietà.

Usando un selettore di classe, dai all'elemento `.cat-left-ear` un bordo sinistro e destro di `35px solid transparent` ciascuno. Imposta anche il bordo inferiore a `70px solid #5e5e5e`.

# --hints--


Il tuo selettore `.cat-right-ear` non dovrebbe avere una proprietà `height`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
}
```

Il tuo selettore `.cat-right-ear` non dovrebbe avere una proprietà `width`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
}
```


Il tuo selettore `.cat-right-ear` non dovrebbe avere una proprietà `border-left`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
}
```

Il tuo selettore `.cat-right-ear` non dovrebbe avere una proprietà `border-right`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
}
```

Il tuo selettore `.cat-right-ear` non dovrebbe avere una proprietà `border-top`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
}
```

Il tuo selettore `.cat-right-ear` non dovrebbe avere una proprietà `border-bottom`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
}
```

Non dovresti avere un selettore `.cat-right-ear`.

```js
assert.notExists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Dovresti avere un selettore `.cat-left-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-left-ear'))
```

Il tuo selettore `.cat-left-ear` dovrebbe avere una proprietà `border-left` impostata su `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft, '35px solid transparent')
```

Il tuo selettore `.cat-left-ear` dovrebbe avere una proprietà `border-right` impostata su `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight, '35px solid transparent')
```

Il tuo selettore `.cat-left-ear` dovrebbe avere una proprietà `border-bottom` impostata su `70px solid #5e5e5e`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom, '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
