---
id: af7588ade1100bde429baf20
title: Побудувати детектор відсутньої літери
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з назвою `fearNotLetter`.
1. Функція `fearNotLetter` повинна приймати один аргумент: рядок, який представляє діапазон літер в алфавітному порядку, у якому може бути відсутня одна літера.
1. Функція має знайти відсутню літеру в переданому діапазоні літер і повернути її.
1. Якщо всі літери присутні в діапазоні, функція повинна повернути `undefined`.

# --hints--

У вас має бути функція `fearNotLetter`.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` має повертати рядок `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` має повертати рядок `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` має повертати рядок `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` має повертати рядок `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` має повертати `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
