---
id: 67172b43f84bcd2dec238a3d
title: Construir um Organizador de Livros
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório. 

**Histórias de Usuário:**

1. Você deve ter um array de objetos chamado `books` onde cada objeto no array deve ter uma string `title`, outra string `authorName` e um número `releaseYear`.

2. Seu array `books` deve conter no mínimo três objetos.

3. Você deve ter uma função de callback chamada `sortByYear` que aceita dois livros como parâmetro para ordenar o array.

4. A função `sortByYear` deve retornar `-1` se o `releaseYear` do primeiro livro for menor que o do segundo livro.

5. A função `sortByYear` deve retornar `1` se o `releaseYear` do primeiro livro for maior que o do segundo livro.

6. A função `sortByYear` deve retornar `0` se ambos os valores de `releaseYear` forem iguais.

7. Você deve filtrar os livros escritos após um determinado ano, como 1950, do array `books` e salvar o array filtrado em um novo array chamado `filteredBooks`.

8. Você deve ordenar os livros no array `filteredBooks` de acordo com seu `releaseYear` em ordem crescente.  You learned in a prior lesson that the `sort()` method will sort the array in place. Isso significa que o array `filteredBooks` será modificado.

# --hints--

Você deve ter uma função `sortByYear` no seu código.

```js
assert.isFunction(sortByYear);
```

Sua função `sortByYear` deve receber dois parâmetros.

```js
assert.lengthOf(sortByYear, 2);
```

Sua função `sortByYear` deve retornar `-1` se o `releaseYear` do objeto `book1` for menor que o do objeto `book2`, `1` se o `releaseYear` do objeto `book1` for maior que o do objeto `book2`, e `0` em todos os outros cenários.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

Você deve ter um array `books` no seu código.

```js
assert.isArray(books)
```

Seu array `books` deve conter pelo menos três objetos de livro.

```js
assert.isAtLeast(books.length, 3)
```

Seu array `books` deve conter objetos, cada um com uma string `title`, outra string `authorName` e um número `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

Você deve ter um array `filteredBooks` no seu código.

```js
assert.isArray(filteredBooks)
```

O array `filteredBooks` deve conter alguns dos livros filtrados do array `books` e não deve estar vazio.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

O array `filteredBooks` deve conter apenas livros do array original `books`.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

O array `filteredBooks` deve incluir apenas livros lançados em ou antes de um ano especificado.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

Você deve chamar a função de ordem superior `sort` passando a função de callback `sortByYear` no array `filteredBooks`.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
