---
id: 68af274560e009597dd861e3
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Sasa, tengeneza mfano wa `Movie` ukitumia hoja unazotaka na weka thamani yake katika kigezo kinachoitwa `movie1`. Kisha, chapisha `movie1` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `movie1`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

Unapaswa kuweka thamani ya mfano wa `Movie` katika kigezo cha `movie1`.

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

Unapaswa kuchapisha `movie1` kwenye konsoli.

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

--fcc-editable-region--

--fcc-editable-region--
```
