---
id: 68af274e60e009597dd861e4
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Katika hatua zilizopita, uliunda kipengele cha sinema cha msingi. Sasa, utaongeza uthibitisho ili kuhakikisha kwamba sinema haiwezi kuanzishwa ikiwa na sifa za kipengele zisizo halali.

Mwanzoni mwa njia ya `__init__`, unda tamko la `if` linalokagua kama `title` ni tupu au lina nafasi tu za blanki (unaweza kutumia njia ya `strip` kuunda hali hiyo).

Ikiwa ni hivyo, onyesha `ValueError` yenye ujumbe wa `Title cannot be empty`.

# --hints--

Unapaswa kuonyesha `ValueError` yenye ujumbe wa `Title cannot be empty` wakati `title` ni tupu au lina nafasi tu za blanki.

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Haupasi kuonyesha `ValueError` wakati `title` ni mfuatano wa herufi halali.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
