---
id: 68af272f60e009597dd861e1
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda orodha ya vyombo vya habari. Anza kwa kufafanua darasa la vitu linaloitwa `Movie`.

Toa darasa lako njia ya `__init__` yenye vigezo vifuatavyo: `self`, `title`, `year`, `director`, na `duration`. Ndani ya njia ya `__init__`, weka thamani za vigezo `title`, `year`, `director`, na `duration` kwa sifa za kipengele za mfano zenye majina sawa.

# --hints--

Unapaswa kuunda darasa la vitu linaloitwa `Movie`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

Darasa lako la `Movie` linapaswa kuwa na njia ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

Njia yako ya `__init__` inapaswa kuwa na vigezo vifuatavyo: `self`, `title`, `year`, `director`, na `duration`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

Unapaswa kuweka thamani ya `title` kwa `self.title` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

Unapaswa kuweka thamani ya `year` kwa `self.year` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

Unapaswa kuweka thamani ya `director` kwa `self.director` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

Unapaswa kuweka thamani ya `duration` kwa `self.duration` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
