---
id: 68af273960e009597dd861e2
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Unataka onyesho linaloweza kusomwa kwa ajili ya kurekebisha makosa na kuchapisha. Unda njia ya `__str__` inayorudisha mfuatano wa herufi wenye muundo wa `title (year) - duration min, director` (ambapo `title`, `year`, `duration`, na `director` zinapaswa kubadilishwa na thamani za sifa za kipengele zinazolingana) ili kuchapisha filamu ipate muhtasari mzuri.

# --hints--

Darasa lako la `Movie` linapaswa kuwa na njia ya `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

Njia yako ya `__str__` inapaswa kurudisha mfuatano wa herufi wenye muundo wa `title (year) - duration min, director`, ambapo `title`, `year`, `duration`, na `director` zinapaswa kubadilishwa na thamani za sifa za kipengele zinazolingana.

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
--fcc-editable-region--
    
--fcc-editable-region--
```
