---
id: 68af276360e009597dd861e6
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Tengeneza tamko jingine la `if` ili kuangalia kama `director` ni tupu au lina herufi za nafasi tu na onyesha `ValueError` yenye ujumbe wa `Director cannot be empty`.

# --hints--

Unapaswa kuonyesha `ValueError` yenye ujumbe wa `Director cannot be empty` wakati `director` ni tupu au lina herufi za nafasi tu.

```js
({
  test: () => runPython(`
  _msg = 'Director cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1999, ' ', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 1999, '  ', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Haupasi kuonyesha `ValueError` wakati `director` ni mfuatano wa herufi halali.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid director'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
--fcc-editable-region--
        
--fcc-editable-region--
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
