---
id: 68b704862dd6559e890c2a86
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Baada ya usemi wa `if`, tengeneza kigezo kinachoitwa `result`. Hii itakuwa mfuatano wa herufi unaorudishwa kutoka kwa njia.

Kwa hiyo weka thamani yake kuwa f-string yenye muundo wa `Media Catalogue (N items):`. Badilisha `N` na idadi ya vitu katika orodha ya `items` na fanya mfuatano huo mwisho na herufi mbili za kuingia mstari mpya.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `result` katika njia yako ya `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_variable("result")`)
})
```

Unapaswa kuweka `f'Media Catalogue ({len(self.items)} items):\n\n'` kwa kigezo cha `result`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_stmt("result = f'Media Catalogue ({len(self.items)} items):\\\\n\\\\n'")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []
    
    def add(self, media_item):
        self.items.append(media_item)
    
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'
        
--fcc-editable-region--
        
--fcc-editable-region--

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
