---
id: 68b804f3be853a355f055fef
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Kitendakazi cha `super` kinatumika kurejelea darasa la mzazi bila kuliita kwa jina moja kwa moja. Kwa mfano, `super().__init__(arg1, arg2, ...)` hufanya wito wa njia ya `__init__` iliyobainishwa katika darasa la mzazi.

Kwa sasa, umebadilisha kwa ufanisi njia ya `__init__` iliyorithiwa kutoka `Movie`. Hii inamaanisha kuwa kitu cha `TVSeries` kitakuwa na sifa ya `seasons` pekee na sifa ya `total_episodes`. Hii siyo unachotaka.

Ili kuhifadhi msimbo wote uliobainishwa katika njia ya `Movie` ya darasa la `__init__`, tumia `super` kufanya wito wa `__init__` wa mzazi. Kumbuka kuwa utahitaji kupitisha hoja zinazohitajika na njia hiyo. Katika kesi hii, `title`, `year`, `director`, na `duration`.

# --hints--

Unapaswa kuita `super().__init__(title, year, director, duration)` ndani ya njia yako ya `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("super().__init__(title, year, director, duration)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
--fcc-editable-region--
        
--fcc-editable-region--
        
        self.seasons = seasons
        self.total_episodes = total_episodes

class MediaCatalogue:
    def __init__(self):
        self.items = []
    
    def add(self, media_item):
        self.items.append(media_item)
    
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'
        
        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
