---
id: 660f23b53db70af0f2620e78
title: 第 43 步
challengeType: 1
dashedName: step-43
---

# --description--

現在，你的所有數字都出現在同一行上。 這對於創建金字塔不起作用。

你需要在每一行中添加一個新行。 但是，在 JavaScript 中按回車鍵在引號之間插入換行符將導致解析錯誤。 相反，你需要使用特殊的 <dfn>轉義序列</dfn> `\n`，它在記錄字符串時被解釋爲新行。 例如：

```js
lineOne = lineOne + "\n" + lineTwo;
```

Use a second addition operator to append a new line after the existing `result` value and the added `row` value.

# --hints--

你應該使用 `\n` 轉義序列。 請記住，它需要是一個字符串，因此用引號括起來。

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

You should concatenate your `row` variable to your `result` variable.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

You should concatenate the `\n` escape sequence to your `row` variable.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

你應該將整個連接分配回你的 `result` 變量。 不要忘記分號。

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*result\s*\+\s*row\s*\+\s*('|")\\n\1;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
