---
id: 68e44c305f1f935d949d76bb
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Kama ilivyotajwa katika hatua iliyopita, kuna nafasi ya ziada baada ya `"love"` ya mwisho katika sentensi `"I love love love  learning."`. Ili kuondoa nafasi hiyo ya ziada, unaweza kutumia njia ya `trimEnd()`.

Mwisho wa njia ya `"love ".repeat(3)`, ungana na njia ya `trimEnd()`. Unaweza kuunganisha njia kama hii:

```js
.firstMethod().secondMethod()
```

Sasa unapoangalia konsoli, unapaswa kuona kuwa nafasi ya ziada imeondolewa.

Na kwa mabadiliko hayo ya mwisho, warsha yako imekamilika!

# --hints--

Unapaswa kuunganisha njia ya `trimEnd()` na njia ya `"love ".repeat(3)`.

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
