---
id: 68e4465fba0bd256868b7deb
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Sasa ni wakati wa kufanya kazi na njia mpya ya mfuatano wa herufi.

Anza kwa kuunda kigezo kinachoitwa `learningSentence` na weka thamani ya mfuatano wa herufi `"I love learning!"`.

Kisha chini ya kigezo hicho, andika kwenye konsoli `"Original learning sentence:"`. Chini ya `console.log()`, ongeza `console.log()` nyingine kwa kigezo cha `learningSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na kigezo cha `learningSentence`.

```js
assert.exists(learningSentence);
```

Kigezo chako cha `learningSentence` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(learningSentence);
```

Unapaswa kuweka thamani ya mfuatano wa herufi `"I love learning!"` kwa kigezo cha `learningSentence`.

```js
assert.equal(learningSentence, "I love learning!");
```

Unapaswa kuandika `"Original learning sentence:"` kwenye konsoli.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

Unapaswa kuandika kigezo cha `learningSentence` kwenye konsoli.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
