---
id: 68e41084bb275539f72de9fe
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utazoea kufanya kazi na njia za `replace()`, `replaceAll()` na `repeat()` kwa kujenga programu ya kubadilisha mfuatano wa herufi.

Ili kuanza, tengeneza kigezo kinachoitwa `originalString` na weka thamani ya `"I love cats."`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `originalString`.

```js
assert.exists(originalString);
```

Kigezo chako cha `originalString` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(originalString);
```

Unapaswa kuweka mfuatano wa herufi `"I love cats."` katika kigezo cha `originalString`.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
