---
id: 68e44abea678a65b877864dc
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Itakuwa vizuri kuona neno hilo likirudiwa katika sentensi. Unda kigezo kipya kinachoitwa `newSentence` na kipa thamani ya matokeo ya kuunganisha neno `"I "` na kigezo `repeatedLove` kisha kuunganisha mfuatano wa herufi `" learning."`. Kisha andika thamani ya kigezo `newSentence` kwenye konsoli. Unaweza kuchagua kutumia au mwekaji `+` au herufi za kiolezo kwa ajili ya kuunganisha mfuatano wa herufi.

Sasa unapaswa kuona maandishi `"I love love love  learning."`.

*TAFADHALI*: Kuna nafasi ya ziada baada ya `"love"` ya mwisho ambayo ni sahihi. Katika hatua ya mwisho, utakata nafasi hiyo ya ziada.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `newSentence`.

```js
assert.exists(newSentence);
```

Kigezo chako cha `newSentence` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(newSentence);
```

Unapaswa kuunganisha neno `"I "` na kigezo `repeatedLove` kisha kuunganisha mfuatano wa herufi `" learning."`. Matokeo haya yanapaswa kuwekwa thamani ya kigezo `newSentence`. Kuwa makini na nafasi.

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

HUPASWI kuweka mfuatano wa herufi `"I love love love  learning."` kama thamani ya kigezo `newSentence`. Tafadhali soma tena maelekezo kwa njia sahihi.

```js
const explorer = await __helpers.Explorer(code);
const { newSentence } = explorer.variables;
assert.notInclude(newSentence?.value?.toString(), "I love love love  learning.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
