---
id: 68e44519c769bf545bb3f22a
title: Hatua ya 7
challengeType: 1
dashedName: step-7
---

# --description--

Ili kupata matokeo unayotaka ya `"I love dogs and dogs are so much fun!"`, utahitaji kutumia njia ya `replaceAll()` badala ya njia ya `replace()`. Njia hii inarudisha mfuatano mpya wa herufi kwa kila mechi ya sehemu ndogo ya mfuatano.

Sahihisha `exampleSentence.replace("cats", "dogs")` ili itumie njia ya `replaceAll()` badala yake. Sasa unapaswa kuona maandishi sahihi kwenye konsoli.

# --hints--

Unapaswa kusahihisha mstari wa `exampleSentence.replace("cats", "dogs")` wa msimbo ili utumie njia ya `replaceAll()`.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
