---
id: 68e43e884ef09d4d2ab05631
title: Hatua ya 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa ni wakati wa kuona matokeo ya kutumia njia ya `replace()`.

Anza kwa kuongeza `console.log()` yenye thamani ya `"After using the replace() method:"`.

Kisha chini ya `console.log()` hiyo, ongeza `console.log()` nyingine yenye variable ya `replacedString`.

Tazama kwenye konsoli na utapaswa kuona kwamba sentensi mpya inasema `"I love dogs."` badala ya `"I love cats."` ya awali.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuandika `"After using the replace() method:"` kwenye konsoli.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

Unapaswa kuandika variable ya `replacedString` kwenye konsoli.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
