---
id: 68e42b62d7093345c6975f22
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Chini ya variable yako ya `originalString`, ongeza `console.log()` yenye mfuatano wa herufi `"Original string:"`. 

Kisha chini ya `console.log()` hiyo, ongeza `console.log()` nyingine yenye variable ya `originalString`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuandika mfuatano wa herufi `"Original string:"` kwenye konsoli.

```js
assert.equal(getLogs()[0], "Original string:");
```

Unapaswa kuandika variable ya `originalString` kwenye konsoli. 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
