---
id: 68e43f94365a964f26c88f78
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Sasa ni wakati wa kufanya kazi na njia mpya ya mfuatano wa herufi.

Anza kwa kuunda kigezo kipya kinachoitwa `exampleSentence` na ukaweke thamani ya mfuatano wa herufi `"I love cats and cats are so much fun!"`.

Kisha chini ya kigezo hicho, andika kwenye konsoli `"Original sentence:"`. Kisha chini ya hicho `console.log()`, ongeza `console.log()` nyingine yenye kigezo `exampleSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na kigezo `exampleSentence`.

```js
assert.exists(exampleSentence);
```

Kigezo chako cha `exampleSentence` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(exampleSentence);
```

Unapaswa kuweka thamani ya mfuatano wa herufi `"I love cats and cats are so much fun!"` kwa kigezo cha `exampleSentence`.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

Unapaswa kuandika `"Original sentence:"` kwenye konsoli.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

Unapaswa kuandika kigezo cha `exampleSentence` kwenye konsoli.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
