---
id: 68e448ec4ac4205968c0a805
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Itakuwa vizuri kama neno `"love"` linarudiwa mara tatu katika sentensi `"I love learning!"`.

Kama unavyokumbuka katika funzo la awali, unaweza kurudia mfuatano wa herufi idadi maalum ya mara kwa kutumia njia ya `repeat()`. Hapa kuna mfano:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

Kwa kuwa mfuatano wa herufi hauwezi kubadilika, njia hii haitabadilisha mfuatano wa herufi wa awali. Itarudisha mfuatano mpya wa herufi wenye maudhui yaliyorudiwa.

Tengeneza kigezo kinachoitwa `repeatedLove` na weka thamani `"love ".repeat(3)`. Kisha andika kwenye konsoli kigezo `repeatedLove` ili uone matokeo.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `repeatedLove`.

```js
assert.exists(repeatedLove);
```

Kigezo chako cha `repeatedLove` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(repeatedLove);
```

Unapaswa kuweka thamani ya `"love ".repeat(3)` kwenye kigezo cha `repeatedLove`.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

Unapaswa kuandika kigezo cha `repeatedLove` kwenye konsoli.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
