---
id: 68e42de6396f5c49d90d1951
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Katika funzo lililopita, ulijifunza kuhusu njia ya `replace()` ambayo hutumika kutafuta sehemu ndogo ya mfuatano wa herufi ndani ya mfuatano wa herufi na kuibadilisha na thamani nyingine.

Hapa kuna mfano:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

Kumbuka kuwa mfuatano wa herufi hauwezi kubadilika, ambayo ina maana mfuatano wa herufi wa asili `text` haukubadilishwa katika mfano huo. Badala yake, mfuatano mpya wa herufi huundwa.

Unda kigezo kipya kinachoitwa `replacedString` na uweke thamani `originalString.replace("cats", "dogs")` ndani yake.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `replacedString`.

```js
assert.exists(replacedString);
```

Kigezo chako cha `replacedString` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(replacedString);
```

Unapaswa kuweka thamani `originalString.replace("cats", "dogs")` katika kigezo cha `replacedString`.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
