---
id: 68e4417b7317835208145323
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Ifuatayo, tengeneza kigezo kipya kinachoitwa `dogsOnlySentence` na weka thamani `exampleSentence.replace("cats", "dogs")`.

Kisha chini ya kigezo hicho, andika kwenye konsoli `"Replacing all occurrences of cats with dogs:"`. Chini ya `console.log()`, ongeza `console.log()` nyingine kwa kigezo cha `dogsOnlySentence`.

Sasa unapaswa kuona maandishi `"I love dogs and cats are so much fun!"` yameandikwa kwenye konsoli. Lakini hiyo si matokeo unayotaka. Katika hatua inayofuata, utajifunza jinsi ya kurekebisha ili maandishi yasomeke `"I love dogs and dogs are so much fun!"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na kigezo cha `dogsOnlySentence`.

```js
assert.exists(dogsOnlySentence);
```

Kigezo chako cha `dogsOnlySentence` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(dogsOnlySentence);
```

Unapaswa kuweka thamani `exampleSentence.replace("cats", "dogs")` kwa kigezo chako cha `dogsOnlySentence`.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

Unapaswa kuandika `"Replacing all occurrences of cats with dogs:"` kwenye konsoli.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

Unapaswa kuandika kigezo cha `dogsOnlySentence` kwenye konsoli.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
