---
id: 6842a6cd9836f0114a5b7a8a
title: crea un generatore di pattern numerici
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

In questo laboratorio eserciterai le basi di Python creando una piccola app che genera un pattern numerico.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti definire una funzione chiamata `number_pattern` che prende un singolo parametro `n` (che rappresenta un numero intero positivo).
1. `number_pattern` dovrebbe usare un loop `for`.
1. `number_pattern(n)` dovrebbe restituire una stringa con tutti i numeri interi da 1 fino a `n` (incluso) separati da uno spazio. Per esempio, `number_pattern(4)` dovrebbe restituire la stringa `1 2 3 4`.
1. Se l'argomento passato alla funzione non è un numero intero, la funzione dovrebbe restituire `Argument must be an integer value.`
1. Se l'argomento passato alla funzione è minore di 1, la funzione dovrebbe restituire `Argument must be an integer greater than 0.`

# --hints--

Dovresti avere una funzione `number_pattern`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

La funzione `number_pattern` dovrebbe avere un parametro chiamato `n`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` dovrebbe restituire `1 2 3 4`.

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` dovrebbe restituire `1 2 3 4 5 6 7 8 9 10 11 12`.

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` dovrebbe restituire una lista di numeri separati da spazi per ogni numero intero positivo.

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

`number_pattern` dovrebbe restituire `Argument must be an integer value.` quando riceve un valore che non è un numero intero.

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` dovrebbe restituire `Argument must be an integer greater than 0.` quando riceve un numero intero non positivo.

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
