---
id: a6b0bb188d873cb2c8729495
title: تنفيذ محول كيان HTML
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

هذا المختبر يختص بتحويل الأحرف الخاصة في سلسلة نصية إلى كيانات HTML المقابلة لها.

**الهدف**: نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة `convertHTML` تقبل سلسلة نصية كمعلمة.
1. يجب أن تُرجع دالة `convertHTML` سلسلة نصية جديدة بتحويل الأحرف الخاصة في سلسلة المعلمة إلى كيانات HTML المقابلة لها.

- يجب تحويل `&` إلى `&amp;`.
   - يجب تحويل `<` إلى `&lt;`.
   - يجب تحويل `>` إلى `&gt;`.
   - يجب تحويل `"` إلى `&quot;`.
   - يجب تحويل `'` إلى `&apos;`.

# --hints--

يجب أن تملك دالة `convertHTML`.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` يجب أن تُرجع السلسلة النصية `Dolce &amp; Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` يجب أن تُرجع السلسلة النصية `Hamburgers &lt; Pizza &lt; Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` يجب أن تُرجع السلسلة النصية `Sixty &gt; twelve`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` يجب أن تُرجع السلسلة النصية `Stuff in &quot;quotation marks&quot;`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` يجب أن تُرجع السلسلة النصية `Schindler&apos;s List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` يجب أن تُرجع السلسلة النصية `&lt;&gt;`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` يجب أن تُرجع السلسلة النصية `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
