---
id: ab306dbdcc907c7ddfc30830
title: Ein Tool zum tiefen Flattening erstellen
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

In diesem Lab implementieren Sie einen Algorithmus zum Flattening von Arrays.

Ein Array zu flatten bedeutet, ein verschachteltes Array beliebiger Tiefe in ein einziges eindimensionales Array umzuwandeln. Dabei werden alle Elemente der Reihe nach extrahiert, wobei nur Arrays aufgelöst werden. Andere Typen bleiben unverändert.

Zum Beispiel:

| Original | Geflattet |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `steamrollArray` haben.
1. Die `steamrollArray`-Funktion sollte ein Argument akzeptieren: ein verschachteltes Array.
1. Die Funktion sollte das verschachtelte Array flatten, wobei unterschiedliche Verschachtelungsebenen berücksichtigt werden.
1. Ihre Lösung sollte die Methoden `Array.prototype.flat()` oder `Array.prototype.flatMap()` nicht verwenden.
1. Globale Variablen sollten nicht verwendet werden.

# --hints--

Sie sollten eine `steamrollArray`-Funktion haben.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` sollte `["a", "b"]` zurückgeben.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` sollte `[1, 2, 3, 4]` zurückgeben.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` sollte `[1, 3, 4]` zurückgeben.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` sollte `[1, {}, 3, 4]` zurückgeben.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

Ihre Lösung sollte die Methoden `Array.prototype.flat()` oder `Array.prototype.flatMap()` nicht verwenden.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

Sie sollten keine globalen Variablen verwenden.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
