---
id: 61537bb9b1a29430ac15ad38
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

All’interno del tuo elemento `head`, aggiungi un elemento `title` con il testo impostato su `Photo Gallery` e un elemento `link` per aggiungere il tuo file `styles.css` alla pagina.

# --hints--

Il tuo elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Il tuo codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Il tuo progetto dovrebbe avere un titolo `Photo Gallery`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

Ricorda, la maiuscola e la grafia sono importanti per il titolo.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
