---
id: 61537485c4f2a624f18d7794
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Inizia con il tuo codice boilerplate HTML standard. Aggiungi una dichiarazione `DOCTYPE`, un elemento `html`, un elemento `head` e un elemento `body`.

Aggiungi l'attributo `lang` al tag di apertura `<html>` impostando `en` come valore.

# --hints--

Il tuo codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti includere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire `html` come tipo di documento.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un tag di chiusura `>` dopo il tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di apertura.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

Il tuo tag di apertura `<html>` dovrebbe avere l'attributo `lang` con `en` come valore.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

La tua dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dovresti avere un tag di apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag di chiusura `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag di apertura `body`.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag di chiusura `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere fratelli.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

L'elemento `head` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

L'elemento `body` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
