---
id: 61537d86bdc3dd343688fceb
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Successivamente, assegna a ogni `img` un attributo `src` in base al suo ordine nel documento. Il primo `img` dovrebbe avere un `src` di `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`. Gli altri dovrebbero essere uguali, tranne che sostituire il `1` con il numero che il `img` ha nel documento.

# --hints--

Tutti e nove i tuoi elementi `img` dovrebbero avere un attributo `src`.

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

Il tuo primo elemento `img` dovrebbe avere `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` impostato come valore dell'attributo `src`.

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Il tuo secondo elemento `img` dovrebbe avere `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` impostato come valore dell'attributo `src`.

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

Il tuo terzo elemento `img` dovrebbe avere `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` impostato come valore dell'attributo `src`.

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

Il tuo quarto elemento `img` dovrebbe avere `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` impostato come valore dell'attributo `src`.

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

Il tuo quinto elemento `img` dovrebbe avere `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` impostato come valore dell'attributo `src`.

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

Il tuo sesto elemento `img` dovrebbe avere `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` impostato come valore dell'attributo `src`.

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

Il tuo settimo elemento `img` dovrebbe avere `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` impostato come valore dell'attributo `src`.

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

L'ottavo elemento `img` dovrebbe avere `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` impostato come valore dell'attributo `src`.

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

Il nono elemento `img` dovrebbe avere `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` impostato come valore dell'attributo `src`.

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
