---
id: 61537a8054753e2f1f2a1574
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

All'interno del tuo elemento `head`, aggiungi un tag `meta` con l'attributo `name` impostato su `viewport` e l'attributo `content` impostato su `width=device-width, initial-scale=1`.

Aggiungi anche un tag `meta` con l'attributo `charset` impostato su `UTF-8`.

# --hints--

Dovresti avere due elementi `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Uno degli elementi `meta` dovrebbe avere un attributo `name` impostato su `viewport` e un attributo `content` impostato su `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

L'altro elemento `meta` dovrebbe avere l'attributo `charset` impostato su `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
