---
id: 61537c9eecea6a335db6da79
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Sotto il tuo elemento `.header`, crea un nuovo elemento `div` e assegna a esso un `class` con valore `gallery`. Questo `div` fungerà da contenitore per le immagini della galleria.

All’interno di quell’elemento `.gallery`, crea nove elementi `img`.

# --hints--

Dovresti creare un elemento `div` nel tuo elemento `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

Il tuo nuovo elemento `div` dovrebbe avere un `class` con `gallery` impostato come valore.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

Il tuo nuovo elemento `div` dovrebbe venire dopo il tuo elemento `.header`.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Il tuo elemento `.gallery` dovrebbe contenere nove elementi `img`.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
