---
id: 6153893900438b4643590367
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Allinea il testo `.header` al centro. Rendi il testo maiuscolo usando la proprietà `text-transform` con `uppercase` come valore.

Dagli una spaziatura interna di `32px` su tutti i lati. Imposta lo sfondo su `#0a0a23` e il testo su `#fff` come valori colore.

Aggiungi una `border-bottom` con `4px solid #fdb347` come valore.

# --hints--

Dovresti avere un selettore `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Il tuo selettore `.header` dovrebbe avere una proprietà `text-align` impostata su `center` come valore.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

Il tuo selettore `.header` dovrebbe avere una proprietà `text-transform` impostata su `uppercase` come valore.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

Il tuo selettore `.header` dovrebbe avere una proprietà `padding` impostata su `32px` come valore.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

Il tuo selettore `.header` dovrebbe avere una proprietà `background-color` impostata su `#0a0a23` come valore.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

Il tuo selettore `.header` dovrebbe avere una proprietà `color` impostata su `#fff` come valore.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

Il tuo selettore `.header` dovrebbe avere una proprietà `border-bottom` impostata su `4px solid #fdb347` come valore.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
