---
id: 61537c5f81f0cf325b4a854c
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Aggiungi un elemento `header` all'interno dell'elemento `body` e assegna a esso una classe `header`.

All'interno del `header`, crea un `h1` con `css flexbox photo gallery` come testo.

# --hints--

Dovresti avere un elemento `header` all'interno del tuo elemento `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Il tuo elemento `header` dovrebbe avere un `class` con `header` come valore.

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Il tuo elemento `header` dovrebbe avere un elemento `h1` al suo interno.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Il tuo elemento `h1` dovrebbe avere il testo `css flexbox photo gallery` al suo interno.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
