---
id: 6494da0daf5df5197963671d
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Nota come il bordo blu dell'immagine si estende oltre il bordo rosso della galleria. Questo succede a causa del modo in cui i browser calcolano la dimensione degli elementi contenitore.

La proprietà `box-sizing` serve a impostare questo comportamento. Per impostazione predefinita, si usa il modello `content-box`. Con questo modello, quando un elemento ha una larghezza specifica, quella larghezza viene calcolata solo in base al contenuto dell'elemento. I valori di spaziatura interna e bordo vengono aggiunti alla larghezza totale, quindi l'elemento cresce per accogliere questi valori.

Prova a impostare esplicitamente `box-sizing` su `content-box` con il selettore globale `*`. A questo punto non vedrai cambiamenti, perché stai usando il valore predefinito.

# --hints--

Dovresti avere un selettore `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Il tuo selettore `*` dovrebbe avere una proprietà `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

Il tuo selettore `*` dovrebbe avere la sua proprietà `box-sizing` impostata su `content-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
