---
id: 67cd525b0e575b1e5abf4c03
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Como se mencionó anteriormente, la función set es responsable de actualizar el estado. Un enfoque común al alternar entre valores de estado booleanos es actualizar el estado de esta manera:

```js
const [isUserLoggedIn, setIsUserLoggedIn] = useState(false);

// updating state
setIsUserLoggedIn(!isUserLoggedIn);
```

En el siguiente ejemplo se usa la función set llamada `setIsUserLoggedIn` para actualizar el estado de inicio de sesión para el usuario. Es una práctica común usar el operador lógico NO (`!`) para alternar entre `true` y `false`. Dado que el valor inicial es `false`, aplicar el operador `!` lo cambiará a `true`.

Dentro de la función `handleToggleVisibility`, use `setIsVisible` para actualizar el estado de visibilidad usando el operador lógico NO (`!`).

# --hints--

Debe usar `setIsVisible` para actualizar la visibilidad del texto. Consulte el ejemplo si necesita ayuda.

```js
assert.match(code, /setIsVisible\s*\(\s*!isVisible\s*\)/);  
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(false);

  --fcc-editable-region--
  const handleToggleVisibility = () => {

  }
  --fcc-editable-region--

  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      {isVisible && <p id="message">I love freeCodeCamp!</p>}
    </div>
  );
};
```
