---
id: 587d7b84367417b2b2512b37
title: 捕捉單引號和雙引號的混合使用
challengeType: 1
forumTopicId: 301188
dashedName: catch-mixed-usage-of-single-and-double-quotes
---

# --description--

JavaScript 允許使用單引號（`'`）和雙引號（`"`）來宣告字串。決定使用哪一種通常取決於個人偏好，但有些例外情況。

當字串有縮寫或另一段被引號包住的文字時，有兩種選擇是很棒的。只要小心不要太早關閉字串，否則會導致語法錯誤。

以下是混合引號的一些範例：

```js
const grouchoContraction = "I've had a perfectly wonderful evening, but this wasn't it.";
const quoteInString = "Groucho Marx once said 'Quote me as saying I was mis-quoted.'";
const uhOhGroucho = 'I've had a perfectly wonderful evening, but this wasn't it.';
```

前兩個是正確的，但第三個是不正確的。

當然，只使用一種引號樣式也是可以的。你可以使用反斜線（`\`）跳脫字元來跳脫字串內的引號：

```js
const allSameQuotes = 'I\'ve had a perfectly wonderful evening, but this wasn\'t it.';
```

# --instructions--

修正字串，使其對 `href` 值使用不同的引號，或跳脫現有的引號。保持整個字串周圍的雙引號標記。

# --hints--

你的程式碼應該透過更改或跳脫，修正 `href` 值 `#Home` 周圍的引號。

```js
assert(__helpers.removeJSComments(code).match(/<a href=\s*?('|\\")#Home\1\s*?>/g));
```

你的程式碼應該保留整個字串周圍的雙引號。

```js
assert(__helpers.removeJSComments(code).match(/"<p>.*?<\/p>";/g));
```

# --seed--

## --seed-contents--

```js
let innerHtml = "<p>Click here to <a href="#Home">return home</a></p>";
console.log(innerHtml);
```

# --solutions--

```js
let innerHtml = "<p>Click here to <a href=\"#Home\">return home</a></p>";
console.log(innerHtml);
```
