---
id: 587d7b85367417b2b2512b38
title: 捕捉使用指定運算子而非等號運算子的情況
challengeType: 1
forumTopicId: 301191
dashedName: catch-use-of-assignment-operator-instead-of-equality-operator
---

# --description--

分支程式，也就是在符合特定條件時執行不同操作的程式，依賴 JavaScript 中的 `if`、`else if` 和 `else` 述語。條件有時採用測試結果是否相等於某個值的形式。

這個邏輯（至少用英文來說）是「如果 x 相等於 y，則……」，這可以直接用 `=`，或指定運算子，來翻譯成程式碼。這會導致你的程式中出現意外的控制流程。

如先前挑戰中所涵蓋，JavaScript 中的指定運算子（`=`）會將值指定給變數名稱。而 `==` 和 `===` 運算子則用來檢查相等性（三重 `===` 測試嚴格相等性，意即值和型別皆相同）。

以下程式碼將 `x` 指定為 2，該值評估為 `true`。在 JavaScript 中，幾乎所有單獨的值都評估為 `true`，除了所謂的「falsy」值：`false`、`0`、`""`（空字串）、`NaN`、`undefined` 和 `null`。

```js
let x = 1;
let y = 2;
if (x = y) {

} else {

}
```

在此範例中，`if` 述語內的程式碼區塊會在任何 `y` 的值下執行，除非 `y` 是 falsy。`else` 區塊，我們預期會在此執行，實際上不會執行。

# --instructions--

修正條件，使程式執行正確的分支，並將適當的值指定給 `result`。

# --hints--

你的程式碼應該修正條件，使其檢查相等性，而不是使用指派。

```js
assert(result == 'Not equal!');
```

條件應該使用 `==` 或 `===` 來測試相等性。

```js
assert(__helpers.removeJSComments(code).match(/x\s*?===?\s*?y/g));
```

# --seed--

## --seed-contents--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x = y) {
  result = "Equal!";
} else {
  result = "Not equal!";
}

console.log(result);
```

# --solutions--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x === y) {
 result = "Equal!";
} else {
 result = "Not equal!";
}

console.log(result);
```
