---
id: 587d7b84367417b2b2512b35
title: 捕捉拼寫錯誤的變數和函式名稱
challengeType: 1
forumTopicId: 301186
dashedName: catch-misspelled-variable-and-function-names
---

# --description--

`console.log()` 和 `typeof` 行為是檢查程式輸出中間值和型別的兩種主要方法。現在是時候進入錯誤常見的形式。一個快速打字者能夠同感的語法層級問題是簡單的拼寫錯誤。

在變數或函式（程式）名稱中，字元顛倒、遺漏或大小寫錯誤，瀏覽器將尋找不存在的物件，並以引用錯誤的形式提出警告。JavaScript 變數和函式（程式）名稱區分大小寫。

# --instructions--

修正程式碼中的兩個拼寫錯誤，讓 `netWorkingCapital` 計算能正常運作。

# --hints--

檢查用於 `netWorkingCapital` 計算的兩個變數拼寫，主控台輸出應顯示「Net working capital is: 2」。

```js
assert(netWorkingCapital === 2);
```

程式碼中不應該有拼錯的變數實體。

```js
assert(!__helpers.removeJSComments(code).match(/recievables/g));
```

`receivables` 變數應該在程式碼中正確宣告和使用。

```js
assert(__helpers.removeJSComments(code).match(/receivables/g).length == 2);
```

程式碼中不應該有拼錯的變數實體。

```js
assert(!__helpers.removeJSComments(code).match(/payable;/g));
```

`payables` 變數應該在程式碼中正確宣告和使用。

```js
assert(__helpers.removeJSComments(code).match(/payables/g).length == 2);
```

# --seed--

## --seed-contents--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = recievables - payable;
console.log(`Net working capital is: ${netWorkingCapital}`);
```

# --solutions--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = receivables - payables;
console.log(`Net working capital is: ${netWorkingCapital}`);
```
