---
id: 587d7b85367417b2b2512b3a
title: 捕捉在呼叫函式時以錯誤順序傳遞的引數
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--

繼續討論呼叫函式，下一個要注意的錯誤是當函式的引數以錯誤的順序提供時。如果引數是不同型別，例如函式預期一個陣列和一個整數，這很可能會發出執行期錯誤。如果引數是相同型別（例如全部是整數），那麼程式碼的邏輯將不合理。請確保以正確的順序提供所有必要的引數，以避免這些問題。

# --instructions--

函式（程式）`raiseToPower` 將底數提升到指數。不幸地，它沒有被正確呼叫－修正程式碼，使 `power` 的值為預期的 8。

# --hints--

你的程式碼應該修正變數 `power`，使其等於 2 引發到第 3 次方，而不是 3 引發到第 2 次方。

```js
assert(power == 8);
```

你的程式碼應該使用 `raiseToPower` 函式呼叫的正確引數順序。

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```
