---
id: 587d7b86367417b2b2512b3d
title: 使用有效的終端機條件防止無窮迴圈
challengeType: 1
forumTopicId: 301192
dashedName: prevent-infinite-loops-with-a-valid-terminal-condition
---

# --description--

最後一個主題是令人畏懼的無窮迴圈。當你需要你的程式執行某個程式碼區塊特定次數或直到條件達成時，迴圈是很棒的工具，但它們需要一個終止條件來結束迴圈。無窮迴圈很可能會凍結或當掉瀏覽器，並造成整體程式執行混亂，這是沒有人想要的。

在本節介紹中有一個無窮迴圈的範例－它沒有終止條件來跳出 `loopy()` 裡的 `while` 迴圈。請勿呼叫這個函式！

```js
function loopy() {
  while(true) {
    console.log("Hello, world!");
  }
}
```

確保終端條件（告訴程式何時跳出迴圈程式碼）最終被達成是程式設計師的工作。一個錯誤是將計數變數朝錯誤方向遞增或遞減，與終端條件相反。另一個錯誤是在迴圈程式碼中意外重設計數器或索引變數，而非遞增或遞減它。

# --instructions--

`myFunc()` 函式（程式）包含無窮迴圈，因為終端機條件 `i != 4` 永遠不會計算為 `false`（並中斷迴圈）－`i` 每次通過會增加 2，且因為 `i` 起始為奇數，會直接跳過 4。修正終端機條件中的比較運算子，使迴圈只在 `i` 小於或相等於 4 時執行。

# --hints--

你的程式碼應該更改 `for` 迴圈終端條件（中間部分）中的比較運算子。

```js
assert(__helpers.removeJSComments(code).match(/i\s*?<=\s*?4;/g).length == 1);
```

你的程式碼應該修正迴圈終端機條件中的比較運算子。

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?!=\s*?4;/g));
```

# --seed--

## --seed-contents--

```js
function myFunc() {
  for (let i = 1; i != 4; i += 2) {
    console.log("Still going!");
  }
}
```

# --solutions--

```js
function myFunc() {
 for (let i = 1; i <= 4; i += 2) {
   console.log("Still going!");
 }
}
```
