---
id: 587d7b83367417b2b2512b37
title: 了解 freeCodeCamp 與瀏覽器主控台之間的差異
challengeType: 1
forumTopicId: 301193
dashedName: understanding-the-differences-between-the-freecodecamp-and-browser-console
---

# --description--

你可能已經注意到某些 freeCodeCamp 挑戰包含他們自己的主控台。這個主控台的行為與瀏覽器主控台略有不同。

有許多方法可以與 `console` 一起使用來輸出訊息。舉例來說，`log`、`warn` 和 `clear`。freeCodeCamp 主控台只會輸出 `log` 訊息，而瀏覽器主控台會輸出所有訊息。當你對程式碼進行更改時，它會自動執行並顯示日誌。每次你的程式碼執行時，freeCodeCamp 主控台都會被清除。

# --instructions--

首先，打開你的瀏覽器主控台，以便你可以看到日誌。為此，你可以在大多數瀏覽器中，於頂部的 freeCodeCamp 導覽列上按右鍵，然後點擊 `inspect`。接著，在開啟的視窗中找到 `console` 分頁。

接著，使用 `console.log` 記錄 `output` 變數。查看兩個主控台以檢視記錄。最後，在你的記錄後使用 `console.clear` 清除瀏覽器主控台。查看兩個主控台的差異。

# --hints--

你應該使用 `console.log()` 來列印 `output` 變數。

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/console\.log\(output\)/));
```

你應該使用 `console.clear()` 來清除瀏覽器主控台。

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/console.clear\(\)/)
);
```

你應該在記錄後清除主控台。

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/console\.log\(output\)[\s\S]*console.clear\(\)/)
);
```

# --seed--

## --seed-contents--

```js
let output = "Get this to show once in the freeCodeCamp console and not at all in the browser console";

```

# --solutions--

```js
let output = "Get this to show once in the freeCodeCamp console and not at all in the browser console";

console.log(output);
console.clear();
```
