---
id: 587d7b85367417b2b2512b39
title: 捕捉函式呼叫後遺失的開啟與關閉括號
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--

當函式或方法不帶任何引數時，你可能會忘記在呼叫時包含（空的）開啟和關閉小括弧。通常函式呼叫的結果會存檔在變數中，以便在你的程式碼中其他地方使用。這個錯誤可以透過將變數值（或其型別）記錄到主控台來偵測，並發現其中一個被設定為函式引用，而非函式傳回的預期值。

以下範例中的變數是不同的：

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

這裡 `varOne` 是函式（程式）`myFunction`，而 `varTwo` 是字串 `You rock!`。

# --instructions--

修正程式碼，使變數 `result` 設定為從呼叫函式 `getNine` 傳回的值。

# --hints--

你的程式碼應該修正變數 `result`，使其設定為函式 `getNine` 傳回的數字。

```js
assert(result == 9);
```

你的程式碼應該呼叫 `getNine` 函式（程式）。

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```
