---
id: 587d7b84367417b2b2512b34
title: 使用 typeof 來檢查變數的型別
challengeType: 1
forumTopicId: 18374
dashedName: use-typeof-to-check-the-type-of-a-variable
---

# --description--

你可以使用 `typeof` 來檢查變數的資料結構或型別。當你在處理多種資料型別時，這在偵錯中非常有用。如果你以為你正在將兩個數字相加，但其中一個實際上是字串，結果可能會出乎意料。型別錯誤可能潛藏在計算或函式呼叫中。當你存取並處理以 JavaScript Object Notation（JSON）物件形式存在的外部資料時，尤其要小心。

以下是使用 `typeof` 的一些範例：

```js
console.log(typeof "");
console.log(typeof 0);
console.log(typeof []);
console.log(typeof {});
```

主控台將依序顯示字串 `string`、數字 `number`、物件 `object` 和物件 `object`。

JavaScript 識別七種原始（不可變的）資料型別：`Boolean`、`Null`、`Undefined`、`Number`、`String`、`Symbol`（ES6 新增）和 `BigInt`（ES2020 新增），以及一種可變的項目型別：`Object`。請注意，在 JavaScript 中，陣列在技術上屬於物件的一種型別。

# --instructions--

為你的程式碼中的兩個變數 `seven` 和 `three` 添加兩個 `console.log()` 述語，以檢查它們的 `typeof`。

# --hints--

你的程式碼應該在兩個 `console.log()` 述語中使用 `typeof` 來檢查變數的型別。

```js
assert(__helpers.removeJSComments(code).match(/console\.log\s*\(typeof[\( ].*\)?\)/g).length == 2);
```

你的程式碼應該使用 `typeof` 來檢查變數 `seven` 的型別。

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]seven\)?/g));
```

你的程式碼應該使用 `typeof` 來檢查變數 `three` 的型別。

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]three\)?/g));
```

# --seed--

## --seed-contents--

```js
let seven = 7;
let three = "3";
console.log(seven + three);
// Only change code below this line
```

# --solutions--

```js
let seven = 7;let three = "3";console.log(typeof seven);
console.log(typeof three);
```
