---
id: 587d7b84367417b2b2512b36
title: '捕捉未封閉的小括號、中括號、大括號和引號'
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

另一個需要注意的語法錯誤是所有開啟的小括弧、方括弧、大括弧和引號都必須有對應的關閉配對。當你編輯現有程式碼並插入其中一種配對型別的項目時，常常會忘記其中一部分。此外，當將程式碼區段巢狀放入其他區段時，也要特別小心，例如將回呼函式作為引數添加到方法中。

避免此錯誤的一種方法是當輸入開啟字元時，立即包含對應的關閉字元，然後將游標移回它們之間繼續編寫程式碼。幸運的是，大多數現代程式碼編輯器會自動產生配對的第二半部。

# --instructions--

修正程式碼中的兩個配對錯誤。

# --hints--

你的程式碼應該修正陣列中遺失的部分。

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

你的程式碼應該修正 `.reduce()` 方法中遺失的部分。主控台輸出應該顯示 `Sum of array values is: 6`。

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
