---
id: 587d7b86367417b2b2512b3b
title: 使用索引時捕捉越界錯誤
challengeType: 1
forumTopicId: 301189
dashedName: catch-off-by-one-errors-when-using-indexing
---

# --description--

<dfn>Off by one errors</dfn>（有時稱為 OBOE）會在你嘗試標的字串或陣列的特定索引（用於切片或存取區段）時發生，或在對它們的索引進行迴圈時發生。JavaScript 的索引從零開始，而非一，這表示最後一個索引總是比項目的長度少一。如果你嘗試存取等於長度的索引，程式可能會丟出「index out of range」引用錯誤或列印 `undefined`。

當你使用接受索引範圍作為引數的字串或陣列方法時，閱讀文件並了解它們是否包含邊界（給定索引的項目是否為傳回內容的一部分）會很有幫助。以下是一些常見的錯位一個錯誤範例：

```js
let alphabet = "abcdefghijklmnopqrstuvwxyz";
let len = alphabet.length;
for (let i = 0; i <= len; i++) {
  console.log(alphabet[i]);
}
for (let j = 1; j < len; j++) {
  console.log(alphabet[j]);
}
for (let k = 0; k < len; k++) {
  console.log(alphabet[k]);
}
```

第一個範例這裡迴圈多執行了一次，第二個迴圈少執行了一次（缺少第一個索引，0）。第三個範例是正確的。

# --instructions--

修正以下函式中的兩個索引錯誤，使所有數字 1 到 5 都能列印到主控台。

# --hints--

你的程式碼應該設定迴圈的初始條件，使其從第一個索引開始。

```js
assert(__helpers.removeJSComments(code).match(/i\s*?=\s*?0\s*?;/g).length == 1);
```

你的程式碼應該修正迴圈的初始條件，使索引從 0 開始。

```js
assert(!__helpers.removeJSComments(code).match(/i\s?=\s*?1\s*?;/g));
```

你的程式碼應該設定迴圈的終端機條件，使其在最後一個索引停止。

```js
assert(__helpers.removeJSComments(code).match(/i\s*<\s*len\s*;|i\s*<=\s*len\s*-\s*1\s*;/g).length == 1);
```

你的程式碼應該修正迴圈的終端機條件，使其在長度前停止於 1。

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?<=\s*?len;/g));
```

# --seed--

## --seed-contents--

```js
function countToFive() {
  let firstFive = "12345";
  let len = firstFive.length;
  // Only change code below this line
  for (let i = 1; i <= len; i++) {
  // Only change code above this line
    console.log(firstFive[i]);
  }
}

countToFive();
```

# --solutions--

```js
function countToFive() {
 let firstFive = "12345";
 let len = firstFive.length;
 // Only change code below this line
 for (let i = 0; i < len; i++) {
 // Only change code above this line
   console.log(firstFive[i]);
 }
}

countToFive();
```
