---
id: 587d7b83367417b2b2512b33
title: 使用 JavaScript 主控台檢查變數的值
challengeType: 1
forumTopicId: 18372
dashedName: use-the-javascript-console-to-check-the-value-of-a-variable
---

# --description--

Chrome 和 Firefox 都有優秀的 JavaScript 主控台，也稱為 DevTools，用於偵錯你的 JavaScript。

你可以在 Chrome 的選單中找到 Developer tools，或在 Firefox 的選單中找到 Web Console。如果你使用不同的瀏覽器，或是手機，我們強烈建議切換到桌面版 Firefox 或 Chrome。

`console.log()` 方法會將其小括弧內的輸出「列印」到主控台，這很可能是最有用的偵錯工具。在你的程式碼中於策略性位置放置它，可以讓你看到變數的中間值。最好在查看輸出結果之前，先有個預期的想法。在你的程式碼中設置檢查點以查看計算狀態，有助於縮小問題所在的位置。

以下是一個將字串 `Hello world!` 列印到主控台的範例：

```js
console.log('Hello world!');
```

# --instructions--

在程式碼中標註的地方，使用 `console.log()` 方法列印變數 `a` 的值。

# --hints--

你的程式碼應該使用 `console.log()` 來檢查變數 `a` 的值。

```js
assert(__helpers.removeJSComments(code).match(/console\.log\(a\)/g));
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 1;
a++;
// Only change code below this line


let sumAB = a + b;
console.log(sumAB);
```

# --solutions--

```js
var a = 5; console.log(a);
```
