---
id: 587d7b86367417b2b2512b3c
title: 在迴圈內重新初始化變數時請小心
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--

有時候在迴圈中需要存檔資訊、遞增計數器，或重新設定變數。潛在的問題是變數應該被重新初始化卻沒有，或反之亦然。如果你不小心重設用於終端條件的變數，這尤其危險，會導致無窮迴圈。

透過在你的迴圈的每個週期使用 `console.log()` 列印變數值，可以揭露與重設或未能重設變數相關的錯誤行為。

# --instructions--

以下函式應該要創建一個具有 `m` 列和 `n` 行的二維陣列，內容為零。不幸的是，它沒有產生預期的輸出，因為 `row` 變數沒有在外層迴圈中重新初始化（設回空陣列）。修正程式碼，使其傳回正確的 3×2 零陣列，形式為 `[[0, 0], [0, 0], [0, 0]]`。

# --hints--

你的程式碼應該將 `matrix` 變數設定為一個陣列，該陣列包含 3 列 2 行的零。

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

`matrix` 變數應該有 3 列。

```js
assert(matrix.length == 3);
```

`matrix` 變數的每一列應該有 2 行。

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

`zeroArray(4,3)` 應該傳回一個陣列，該陣列包含 4 列，每列有 3 行的零。

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```
