---
id: 694648acde178bb8202d9514
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

现在你需要为成绩单添加学生的年龄。为此你将使用整数，这是 Python 中的数字数据类型之一。

声明一个名为 `age` 的变量并赋值为整数值 `20`。

然后，打印 `age` 的值和数据类型，用逗号分隔。查看终端中的输出，显示 `age` 的值及其类型为 `int`（整数）。

# --hints--

你应该有一个 `age` 变量。

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

变量 `age` 应该保存值 `20`。不要用引号包围该值。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

你应该使用逗号分隔符打印 `age` 和 `type(age)`。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
