---
id: 694648acde178bb8202d9517
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

输出是 `True`，确认 `score` 是一个 `float`。

通过使用单个 `print()` 语句打印 `score` 值及其数据类型来完成成绩单。

# --hints--

你应该使用逗号分隔符打印 `score` 和 `type(score)`。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
