---
id: 69461e4f98594012939bc1c6
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次工作坊中，你将通过创建一个简单的成绩单打印机来练习 Python 数据类型。

正如你在之前的课程中学到的，变量是以这种方式赋值的：

```py
greeting = 'Hello World'
```

在上面的例子中，变量 `greeting` 的值是一个字串，即被引用包围的字符序列。

创建一个名为 `name` 的变量，并将字串 `Alice` 赋值给它。记住要像示例中那样用单引号或双引号将值括起来。

**重要：** Python 使用缩排（行首的空格）来组织代码。对于本次研讨会，请确保每行代码开头没有多余的空格。添加多余的空格会导致 `IndentationError` 并阻止你的代码运行。

# --hints--

你应该有一个 `name` 变量。

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

变量 `name` 应该保存值 `'Alice'`。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
