---
id: 694648acde178bb8202d9515
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

现在，添加学生的分数。

声明一个名为 `score` 的变量并赋值为 `80.5`。

虽然 `age` 和 `score` 都是数字，但它们可能不是同一种类型。Python 提供了一个名为 `isinstance()` 的函数来查看这一点。

```py
x = 10
print(isinstance(x, int)) # Output: True
```

使用 `isinstance()` 查看 `score` 是否为 `int`，并将结果打印到终端，如上例所示。

# --hints--

你应该有一个 `score` 变量。

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

变量 `score` 应该保存值 `80.5`。不要用引号包围该值。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

你应该打印 `isinstance(score, int)` 的结果。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
