---
id: 694648abde178bb8202d9511
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

你现在应该能在终端看到学生姓名被打印。

Python 提供了一个名为 `type()` 的函数，你可以用它来查看一个值的类型。

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

在上面的示例中，输出 `<class 'str'>` 表示传递给 `type()` 函数的变量是一个字串。

使用 `type()` 函数，将 `name` 作为其参数，并像示例一样打印输出。查看终端中的输出，显示 `name` 的类型是 `str`（字串）。

# --hints--

你应该打印调用 `type(name)` 的结果。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
