---
id: 694648acde178bb8202d9513
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

学生姓名应遵循与其他详细信息相同的形式。

删除之前 `name` 变量的输出。然后，将 `name` 和 `type(name)` 一起打印在一行，用逗号分隔，类似上一步。

# --hints--

你的代码中不应有 `print(name)`。

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

你的代码中不应有 `print(type(name))`。

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

你应该在同一个 `print()` 调用中使用逗号分隔符打印 `name` 和 `type(name)`。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
