---
id: 694648acde178bb8202d9516
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

输出是 `False`，这表明 `score` 不是一个 `int`。

Python 中另一种常见的数字类型是 `float`，它表现为带小数的数字。在现有的 `isinstance()` 调用中将 `int` 替换为 `float` 以确认这一点。

# --hints--

你应该在现有的 `isinstance(score, int)` 调用中将 `int` 替换为 `float`。

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
