---
id: 694648acde178bb8202d9512
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

成绩单还应显示学生当前是否在读。这可以用布尔值来表现。

布尔值表现一个是或否的条件，它们经常用于在代码中做出决策。布尔值只有两个：`True` 和 `False`。

声明一个名为 `is_student` 的变量并赋值为 `True`。

# --hints--

你应该有一个 `is_student` 变量。

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

变量 `is_student` 应该保存值 `True`。不要用引号包围该值。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
