---
id: 6947a3e516bb8a4a681274a7
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

`print()` 函数可以一次显示多个值。用逗号（`,`）分隔值以在同一行打印它们。

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

在同一行使用逗号 `,` 打印 `is_student` 和 `type(is_student)`，如示例所示。然后，在终端查看输出，显示 `is_student` 的值及其类型为 `bool`（布尔）。

# --hints--

你应该使用逗号分隔符打印 `is_student` 和 `type(is_student)`。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
