---
id: 6947a0c2b6de035dd8d3eea1
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你可以使用 `print()` 函数打印变量的值。

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

你将在后续课程中学习更多关于函数的内容。目前，了解函数是一个可重用的代码块，可以被调用或调用以运行其代码，并且可以向其传递参数。

在上面的例子中，`print(greeting)` 是一个函数调用，`greeting` 是该函数的参数。

参考示例并打印 `name` 变量。查看终端中的输出。

# --hints--

你应该调用 `print()`，并将 `name` 作为它的参数。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
