---
id: 66fbcf750a62784cf11f5633
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

調理時間を示す数値をパラメータとして受け取る `getDifficultyLevel` 関数を作成してください。

調理時間が `30` 以下の場合、関数は `"easy"` を返す必要があります。調理時間が `60` 以下の場合、関数は `"medium"` を返す必要があります。それ以外の場合、関数は `"hard"` を返す必要があります。

# --hints--

`getDifficultyLevel` 関数を作成してください。

```js
assert.isFunction(getDifficultyLevel);
```

`getDifficultyLevel` 関数は1つのパラメータを持つ必要があります。

```js
assert.lengthOf(getDifficultyLevel, 1);
```

調理時間が `getDifficultyLevel` 以下の場合、`"easy"` 関数は `30` を返す必要があります。

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

調理時間が `getDifficultyLevel` より大きく `"medium"` 以下の場合、`31` 関数は `60` を返す必要があります。

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

調理時間が `getDifficultyLevel` より大きい場合、`"hard"` 関数は `60` を返す必要があります。

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
