---
id: 66fbcf750a62784cf11f5636
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

`totalIngredients` と `difficultyLevel` の `recipe2` と `recipe3` プロパティについても同じ手順を繰り返してください。

# --hints--

`totalIngredients` の `recipe2` プロパティにアクセスしてください。

```js
assert.isNotNull(recipe2.totalIngredients);
```

`getTotalIngredients` の `recipe2.ingredients` プロパティに、`totalIngredients` を引数にして `recipe2` を呼び出した結果を代入してください。

```js
assert.strictEqual(recipe2.totalIngredients, getTotalIngredients(recipe2.ingredients));
```

`difficultyLevel` の `recipe2` プロパティにアクセスしてください。

```js
assert.isNotEmpty(recipe2.difficultyLevel);
```

`getDifficultyLevel` の `recipe2.cookingTime` プロパティに、`difficultyLevel` を引数にして `recipe2` を呼び出した結果を代入してください。

```js
assert.strictEqual(recipe2.difficultyLevel, getDifficultyLevel(recipe2.cookingTime));
```

`totalIngredients` の `recipe3` プロパティにアクセスしてください。

```js
assert.isNotNull(recipe3.totalIngredients);
```

`getTotalIngredients` の `recipe3.ingredients` プロパティに、`totalIngredients` を引数にして `recipe3` を呼び出した結果を代入してください。

```js
assert.strictEqual(recipe3.totalIngredients, getTotalIngredients(recipe3.ingredients));
```

`difficultyLevel` の `recipe3` プロパティにアクセスしてください。

```js
assert.isNotEmpty(recipe3.difficultyLevel);
```

`getDifficultyLevel` の `recipe3.cookingTime` プロパティに、`difficultyLevel` を引数にして `recipe3` を呼び出した結果を代入してください。

```js
assert.strictEqual(recipe3.difficultyLevel, getDifficultyLevel(recipe3.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

recipe1.totalIngredients = getTotalIngredients(recipe1.ingredients);
recipe1.difficultyLevel = getDifficultyLevel(recipe1.cookingTime);

--fcc-editable-region--

--fcc-editable-region--
```
