---
id: 66fbcf750a62784cf11f5634
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

各関数をテストする時間です。どのレシピでも使えますが、このチュートリアルでは `recipe1` から始めます。

2つの新しい変数を作成してください：`recipe1TotalIngredients` と `recipe1DifficultyLevel`。それぞれの変数には、対応する関数を呼び出して適切な `recipe1` プロパティを渡し、その戻り値を代入してください。

最後に、結果を確認するためにそれぞれの変数をコンソールにログ出力してください。

# --hints--

`recipe1TotalIngredients` という変数を作成してください。

```js
assert.isNotNull(recipe1TotalIngredients);
```

`recipe1TotalIngredients` 変数には `getTotalIngredients(recipe1.ingredients)` を代入してください。

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

`recipe1TotalIngredients` をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

`recipe1DifficultyLevel` という変数を作成してください。

```js
assert.isNotNull(recipe1DifficultyLevel);
```

`recipe1DifficultyLevel` 変数には `getDifficultyLevel(recipe1.cookingTime)` を代入してください。

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

`recipe1DifficultyLevel` をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
