---
id: 66fbcf750a62784cf11f5630
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

今、3つのオブジェクトを`recipes`配列に追加する必要があります。これを行うには、`push()`メソッドを使えます。

`push()`メソッドを使って、すべてのレシピオブジェクトを`recipes`配列に追加してください。`recipe1`、`recipe2`、`recipe3`の順番で追加することを確認してください。

また、`recipe1Name`、`recipe2CookingTime`、`recipe3Ingredients`の変数と、それらの変数をログに出力する`console.log`の文を削除してください。

# --hints--

すべてのレシピオブジェクトを`recipes`配列に追加する必要があります。宣言された順番で追加してください。

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

`recipe1Name`変数を削除してください。

```js
assert.throws(() => recipe1Name, ReferenceError);
```

`console.log(recipe1Name);`の文を削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

`recipe2CookingTime`変数を削除してください。

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

`console.log(recipe2CookingTime);`の文を削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

`recipe3Ingredients`変数を削除してください。

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

`console.log(recipe3Ingredients);`の文を削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
