---
id: 66fbcf750a62784cf11f5632
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

1つの引数を受け取る`getTotalIngredients`関数を作成してください。この引数は材料の配列を表し、関数に渡された配列の材料の数を返します。

# --hints--

`getTotalIngredients`関数を作成してください。

```js
assert.isFunction(getTotalIngredients);
```

`getTotalIngredients`関数は1つのパラメータを持つ必要があります。

```js
assert.lengthOf(getTotalIngredients, 1);
```

`getTotalIngredients`関数は数値を返す必要があります。

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)`は`4`を返すべきです。

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)`は`5`を返すべきです。

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)`は`3`を返すべきです。

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
