---
id: 66fbcdfed6848e48947dbbc9
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、JavaScriptオブジェクトを使ってレシピトラッカーを作成します。

まず、`recipes`という名前の空の配列を作成してください。これは後でレシピオブジェクトを追加する配列です。

# --hints--

`recipes`という名前の配列を作成してください。

```js
assert.isArray(recipes);
```

`recipes`配列は空である必要があります。

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
