---
id: 66fbcf750a62784cf11f5635
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

`recipes` 配列の各アイテムに、`totalIngredients` と `difficultyLevel` プロパティの値を入力できるようになりました。

まずは `totalIngredients` の `difficultyLevel` と `recipe1` にアクセスし、それらを関数呼び出しと引数の適切な結果に設定してください。

# --hints--

`totalIngredients` の `recipe1` プロパティにアクセスしてください。

```js
assert.isNotNull(recipe1.totalIngredients);
```

`getTotalIngredients` を `recipe1.ingredients` で呼び出した結果を `totalIngredients` の `recipe1` プロパティに代入してください。

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

`difficultyLevel` の `recipe1` プロパティにアクセスしてください。

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

`getDifficultyLevel` を `recipe1.cookingTime` で呼び出した結果を `cookingTime` の `recipe1` プロパティに代入してください。

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
