---
id: 66fbcf750a62784cf11f562e
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

次のプロパティと値を持つ `recipe2` オブジェクトを作成してください。

| Key | Value   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

`recipe2` という名前のオブジェクトを作成してください。

```js
assert.isObject(recipe2);
```

`recipe2` オブジェクトには `name` プロパティが必要です。

```js
assert.property(recipe2, "name");
```

`name` プロパティは `Chicken Curry` に設定してください。

```js
assert.equal(recipe2.name, "Chicken Curry");
```

`recipe2` オブジェクトには `ingredients` プロパティが必要です。

```js
assert.property(recipe2, "ingredients");
```

`ingredients` プロパティは配列の値にしてください。

```js
assert.isArray(recipe2.ingredients);
```

`ingredients` プロパティの配列値には `chicken breast`、`coconut milk`、`curry powder`、`onion`、`garlic` が含まれている必要があります。

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

`recipe2` オブジェクトには `cookingTime` プロパティが必要です。

```js
assert.property(recipe2, "cookingTime");
```

`cookingTime` プロパティの値は数値にしてください。

```js
assert.isNumber(recipe2.cookingTime);
```

`cookingTime` プロパティは `42` に設定してください。

```js
assert.equal(recipe2.cookingTime, 42);
```

`recipe2` オブジェクトには `totalIngredients` プロパティが必要です。

```js
assert.property(recipe2, "totalIngredients");
```

`totalIngredients` プロパティは `null` に設定してください。

```js
assert.isNull(recipe2.totalIngredients);
```

`recipe2` オブジェクトには `difficultyLevel` プロパティが必要です。

```js
assert.property(recipe2, "difficultyLevel");
```

`difficultyLevel` プロパティは空の文字列に設定してください。

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
