---
id: 66fbcf750a62784cf11f562d
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

`recipe1` という名前のオブジェクトを作成してください。`recipe1` オブジェクトの中に、値が `name` の `"Spaghetti Carbonara"` プロパティを作成してください。

また、`recipe1` オブジェクトの中に、配列を値とする `ingredients` プロパティを作成してください。その配列には `spaghetti`、`Parmesan cheese`、`pancetta`、`black pepper` が含まれている必要があります。

# --hints--

`recipe1` という名前のオブジェクトを作成してください。

```js
assert.isObject(recipe1);
```

`recipe1` オブジェクトには `name` プロパティが必要です。

```js
assert.property(recipe1, 'name');
```

`name` プロパティは `Spaghetti Carbonara` に設定してください。

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

`recipe1` オブジェクトには `ingredients` プロパティが必要です。

```js
assert.property(recipe1, 'ingredients');
```

`ingredients` プロパティは配列の値を持つ必要があります。

```js
assert.isArray(recipe1.ingredients);
```

`ingredients` プロパティの配列の値には `spaghetti`、`Parmesan cheese`、`pancetta`、`black pepper` が含まれている必要があります。

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
