---
id: 66ff02792b3042a2b811f475
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

次のプロパティを `recipe1` オブジェクトに追加してください：

| キー | 値   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

`null` のプロパティで空文字列の値を持つものは、後で計算して更新します。

# --hints--

`recipe1` オブジェクトには `cookingTime` プロパティが必要です。

```js
assert.property(recipe1, "cookingTime");
```

`cookingTime` プロパティは `22` に設定してください。

```js
assert.equal(recipe1.cookingTime, 22);
```

`recipe1` オブジェクトには `totalIngredients` プロパティが必要です。

```js
assert.property(recipe1, "totalIngredients");
```

`totalIngredients` プロパティは `null` に設定してください。

```js
assert.isNull(recipe1.totalIngredients);
```

`recipe1` オブジェクトには `difficultyLevel` プロパティが必要です。

```js
assert.property(recipe1, "difficultyLevel");
```

`difficultyLevel` プロパティは空文字列に設定してください。

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
