---
id: 670e4f45f7116c0f216a5177
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

次に進む前に、オブジェクトからプロパティにアクセスする方法を練習しましょう。

ドット記法（`.`）かブラケット記法（`[]`）のどちらかを使えます。例を示します。

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

`name`の`recipe1`プロパティにアクセスし、変数`recipe1Name`に代入してください。

次に、`cookingTime`の`recipe2`プロパティにアクセスし、変数`recipe2CookingTime`に代入してください。

最後に、`ingredients`の`recipe3`プロパティにアクセスし、変数`recipe3Ingredients`に代入してください。

作成したすべての変数をコンソールにログ出力することを忘れないでください。

# --hints--

`recipe1Name`変数を作成してください。

```js
assert.isDefined(recipe1Name);
```

`name`の`recipe1`プロパティの値を`recipe1name`変数に代入してください。

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

`recipe1Name`をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

`recipe2CookingTime`変数を作成してください。

```js
assert.isDefined(recipe2CookingTime);
```

`cookingTime`の`recipe2`プロパティの値を`recipe2CookingTime`変数に代入してください。

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

`recipe2CookingTime`をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

`recipe3Ingredients`変数を作成してください。

```js
assert.isDefined(recipe3Ingredients);
```

`ingredients`の`recipe3`プロパティの値を`recipe3Ingredients`変数に代入してください。

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

`recipe3Ingredients`をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
