---
id: 69161f4c0c9c8c48b3c843c5
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

次のプロパティと値を持つ `recipe3` オブジェクトを作成してください。

|キー|値|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

`recipe3` という名前のオブジェクトを作成してください。

```js
assert.isObject(recipe3);
```

`recipe3` オブジェクトには `name` プロパティが必要です。

```js
assert.property(recipe3, "name");
```

`name` プロパティは `Vegetable Stir Fry` に設定してください。

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

`recipe3` オブジェクトには `ingredients` プロパティが必要です。

```js
assert.property(recipe3, "ingredients");
```

`ingredients` プロパティは配列の値を持つ必要があります。

```js
assert.isArray(recipe3.ingredients);
```

`ingredients` プロパティの配列値には `broccoli`、`carrot`、`bell pepper` が含まれている必要があります。

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

`recipe3` オブジェクトには `cookingTime` プロパティが必要です。

```js
assert.property(recipe3, "cookingTime");
```

`cookingTime` プロパティの値は数値である必要があります。

```js
assert.isNumber(recipe3.cookingTime);
```

`cookingTime` プロパティは `15` に設定してください。

```js
assert.equal(recipe3.cookingTime, 15);
```

`recipe3` オブジェクトには `totalIngredients` プロパティが必要です。

```js
assert.property(recipe3, "totalIngredients");
```

`totalIngredients` プロパティは `null` に設定してください。

```js
assert.isNull(recipe3.totalIngredients);
```

`recipe3` オブジェクトには `difficultyLevel` プロパティが必要です。

```js
assert.property(recipe3, "difficultyLevel");
```

`difficultyLevel` プロパティは空の文字列に設定してください。

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
