---
id: 685976a92ffab8b2dc744b72
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie innerhalb Ihres `.tabs`-Elements ein `div`-Element mit einem `role`-Attribut hinzu, das auf `"tablist"` gesetzt ist, und einem `aria-labelledby`-Attribut, das auf `"tabs-title"` gesetzt ist.

Das `aria-labelledby`-Attribut wird hier benötigt, um die Tabliste mit einer sichtbaren Bezeichnung zu verknüpfen, damit unterstützende Technologien wie Screenreader den Zweck der Tabliste den Nutzenden ansagen können.

# --hints--

Sie sollten ein `div`-Element innerhalb Ihres `.tabs`-Elements haben.

```js
assert.exists(document.querySelector(".tabs div"));
```

Ihr `div`-Element sollte ein `role`-Attribut mit dem Wert `"tablist"` besitzen.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

Ihr `div`-Element sollte ein `aria-labelledby`-Attribut mit dem Wert `"tabs-title"` besitzen.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
