---
id: 68599af28ead38c22590a89d
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Es muss eine Möglichkeit geben, die Buttons mit dem Inhalt zu verbinden, den sie darstellen. Dies geschieht durch die Verwendung des `aria-controls`-Attributs bei jedem Button. Später fügen Sie den Inhalt und die entsprechenden `id`-Attribute zu den Inhaltselementen hinzu.

Beginnen Sie damit, jedem der `aria-controls`-Elemente das `button`-Attribut hinzuzufügen. Der Wert für das `aria-controls`-Attribut des ersten Buttons sollte auf `"panel-earth"` gesetzt werden. Das `aria-controls`-Attribut des zweiten Buttons sollte auf `"panel-saturn"` gesetzt werden. Das `aria-controls`-Attribut des dritten Buttons sollte auf `"panel-mars"` gesetzt werden.

# --hints--

Jedes Ihrer `button`-Elemente sollte ein `aria-controls`-Attribut haben.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

Das erste `button`-Element sollte ein `aria-controls`-Attribut mit dem Wert `"panel-earth"` haben.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

Das zweite `button`-Element sollte ein `aria-controls`-Attribut mit dem Wert `"panel-saturn"` haben.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

Das dritte `button`-Element sollte ein `aria-controls`-Attribut mit dem Wert `"panel-mars"` haben.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
