---
id: 68599e9259c0eec3e78c52fd
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Standardmäßig sollten die anderen Tabs nicht ausgewählt sein. Sie können dies erreichen, indem Sie das `aria-selected`-Attribut zu den zweiten und dritten `button`-Elementen hinzufügen und deren Werte auf `"false"` setzen.

# --hints--

Ihr zweites `button`-Element sollte ein `aria-selected`-Attribut haben.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[1].hasAttribute("aria-selected")
);
```

Ihr zweites `button`-Element sollte ein `aria-selected`-Attribut mit dem Wert `"false"` haben.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-selected"), "false");
``` 

Ihr drittes `button`-Element sollte ein `aria-selected`-Attribut haben.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[2].hasAttribute("aria-selected")
);
```

Ihr drittes `button`-Element sollte ein `aria-selected`-Attribut mit dem Wert `"false"` haben.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-selected"), "false");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true">🌍 Earth</button>
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
