---
id: 68599f9df77791c4a311cc5f
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Das letzte Element, das Sie zu jedem der `button` Elemente hinzufügen, sind die `id` Attribute. Diese werden verwendet, um die `button` Elemente mit ihren entsprechenden Inhaltsbereichen zu verknüpfen.

Das erste `button` Element sollte ein `id` von `"tab-earth"` haben, das zweite sollte ein `id` von `"tab-saturn"` haben, und das dritte sollte ein `id` von `"tab-mars"` haben.

# --hints--

Ihr erstes `button` Element sollte ein `id` Attribut haben.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[0].hasAttribute("id")
);
```

Ihr erstes `button` Element sollte ein `id` Attribut mit dem Wert `"tab-earth"` haben.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("id"), "tab-earth");
```

Ihr zweites `button` Element sollte ein `id` Attribut haben.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[1].hasAttribute("id")
);
```

Ihr zweites `button` Element sollte ein `id` Attribut mit dem Wert `"tab-saturn"` haben.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("id"), "tab-saturn");
```

Ihr drittes `button` Element sollte ein `id` Attribut haben.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[2].hasAttribute("id")
);
```

Ihr drittes `button` Element sollte ein `id` Attribut mit dem Wert `"tab-mars"` haben.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("id"), "tab-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-earth" aria-selected="true">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
