---
id: 68599633407374c0d574587b
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Der nächste Schritt besteht darin, die Buttons hinzuzufügen, die für das Wechseln zwischen den verschiedenen Planeten und ihren Fakten verantwortlich sind.

Fügen Sie innerhalb des `tablist`-Elements drei `button`-Elemente hinzu, von denen jedes ein `role`-Attribut mit dem Wert `"tab"` besitzt. Das erste `button`-Element sollte den Text `🌍 Earth` enthalten. Das zweite `button`-Element sollte den Text `🪐 Saturn` enthalten. Das dritte `button`-Element sollte den Text `🔴 Mars` enthalten.

# --hints--

Sie sollten drei `button`-Elemente innerhalb des `tablist`-Elements haben.

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

Jedes Ihrer `button`-Elemente sollte ein `role`-Attribut mit dem Wert `"tab"` besitzen.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

Ihr erstes `button`-Element sollte den Text `🌍 Earth` enthalten.

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

Ihr zweites `button`-Element sollte den Text `🪐 Saturn` enthalten.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

Ihr drittes `button`-Element sollte den Text `🔴 Mars` enthalten.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
