---
id: 6859a484a55387c6543ab2c0
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Jetzt ist es Zeit, die Fakten für jeden Planeten hinzuzufügen. Fügen Sie im ersten `p`-Element den folgenden Text ein:
`Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.`

Fügen Sie im zweiten `p`-Element den folgenden Text ein:
`Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.`

Fügen Sie im dritten `p`-Element den folgenden Text ein:
`Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.`

# --hints--

Ihr erstes `p`-Element sollte die Fakten über die Erde enthalten. Stellen Sie sicher, dass Ihr Text der bereitgestellten Beschreibung entspricht.

```js
assert.equal(
  document.querySelector('div[role="tabpanel"] p').textContent.trim(),
  'Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It\'s the only planet in the solar system known to harbor life.'
);
```

Ihr zweites `p`-Element sollte die Fakten über den Saturn enthalten. Stellen Sie sicher, dass Ihr Text der bereitgestellten Beschreibung entspricht.

```js
assert.equal(
  document.querySelectorAll('div[role="tabpanel"] p')[1].textContent.trim(),
  'Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It\'s a gas giant with dozens of moons orbiting it.'
);
```

Ihr drittes `p`-Element sollte die Fakten über den Mars enthalten. Stellen Sie sicher, dass Ihr Text der bereitgestellten Beschreibung entspricht.

```js
assert.equal(
  document.querySelectorAll('div[role="tabpanel"] p')[2].textContent.trim(),
  'Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It\'s a key focus for exploration in the search for past or present life.'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div role="tabpanel">
        --fcc-editable-region--
        <p></p>
      </div>
      <div role="tabpanel">
        <p></p>
      </div>
      <div role="tabpanel">
        <p></p>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
