---
id: 6859ae3ef795c4c9d4badf9c
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Jetzt ist es an der Zeit, die Funktionalität zur Tabliste hinzuzufügen. Der erste Schritt besteht darin, auf alle Tab- und Panel-Elemente in der Tabliste zuzugreifen.

Erstellen Sie eine Variable namens `tabs` und weisen Sie ihr das Ergebnis einer Abfrage zu, die alle Elemente mit dem `role`-Attribut, das auf `tab` gesetzt ist, auswählt.

Erstellen Sie dann eine Variable namens `panels` und weisen Sie ihr das Ergebnis einer Abfrage zu, die alle Elemente mit dem `role`-Attribut, das auf `tabpanel` gesetzt ist, auswählt.

Denken Sie daran, dass Sie die `querySelectorAll`-Methode verwenden können, um mehrere Elemente gleichzeitig auszuwählen.

# --hints--

Sie sollten eine Variable namens `tabs` haben.

```js
assert.isDefined(tabs);
```

Sie sollten `document.querySelectorAll('[role="tab"]')` der `tabs`-Variablen zuweisen.

```js
assert.deepEqual(
  Array.from(tabs),
  Array.from(document.querySelectorAll('[role="tab"]'))
);
```

Sie sollten eine Variable namens `panels` haben.

```js
assert.isDefined(panels);
```

Sie sollten `document.querySelectorAll('[role="tabpanel"]')` der `panels`-Variablen zuweisen.

```js
assert.deepEqual(
  Array.from(panels),
  Array.from(document.querySelectorAll('[role="tabpanel"]'))
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div id="panel-earth" role="tabpanel" aria-labelledby="tab-earth">
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.
        </p>
      </div>
      <div id="panel-saturn" role="tabpanel" aria-labelledby="tab-saturn" hidden>
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div id="panel-mars" role="tabpanel" aria-labelledby="tab-mars" hidden>
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.
        </p>
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
