---
id: 685db0bd1280154ae85bcf9f
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop bauen Sie eine barrierefreie Planeten-Tabliste mit HTML, CSS und JavaScript. Der Boilerplate-Code für HTML und CSS wurde für Sie bereitgestellt.

In diesen ersten Schritten erstellen Sie den restlichen HTML-Code für die Tabliste.

Beginnen Sie damit, ein `div`-Element mit einer Klasse `"tabs"` zu erstellen. Fügen Sie innerhalb dieses `div`-Elements ein `h2`-Element mit einem `id`-Attribut, das auf `"tabs-title"` gesetzt ist, und dem Textinhalt `Planets` hinzu.

# --hints--

Sie sollten ein `div`-Element haben.

```js
assert.exists(document.querySelector("div"));
```

Ihr `div`-Element sollte eine Klasse `"tabs"` haben.

```js
assert.exists(document.querySelector("div.tabs"));
```

Ihr `div`-Element sollte ein `h2`-Element enthalten.

```js
assert.exists(document.querySelector(".tabs h2"));
```

Ihr `h2`-Element sollte ein `id`-Attribut mit dem Wert `"tabs-title"` haben.

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

Ihr `h2`-Element sollte den Textinhalt `Planets` haben.

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
