---
id: 69860157f7fce39f9df399a3
title: 實作 N-Queens 演算法
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

N 皇后問題要求你將 N 個皇后放置在一個 N×N 的棋盤上，使得沒有兩個皇后互相攻擊（沒有兩個皇后共用同一列、行或對角線）。 

例如，如果有一個 4x4 棋盤，其中一個有效的排列是：

```md
[1, 3, 0, 2]
```

這表示在第 0 列，皇后放置在第 1 行；在第 1 列，皇后放置在第 3 行；在第 2 列，皇后放置在第 0 行；在第 3 列，皇后放置在第 2 行。

從視覺上來看，這種排列方式如下：

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

其中 `Q` 表示一個皇后，`．` 表示一個空格。

在本實驗中，你將使用深度優先搜尋方法實作 N-Queens 問題解決器。


**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `dfsNQueens` 的函式（程式）。
2. 該函式（程式）應該只接受一個引數：整數 `n`。
3. 如果 `n` 小於 `1`，該函式應該傳回一個空的陣列（`[]`）。
4. 該函式應傳回一個解決方案列表；每個解決方案本身是一個長度為 `n` 的列表，其中索引為 `i` 的元素是第 `i` 列中皇后的行索引（從 0 開始）。

# --hints--

你應該有一個名為 `dfsNQueens` 的函式，該函式接受一個引數。

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

如果 `n` 小於 `1`，該函式應該傳回一個空的陣列。

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

該函式應該傳回一個解決方案的陣列，其中每個解決方案都是長度為 `n` 的陣列。

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` 應該傳回 `[[0]]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` 應該傳回 `[]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` 應該傳回 `[]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` 應該傳回 `[[1, 3, 0, 2], [2, 0, 3, 1]]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` 應該傳回 `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`。

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` 應該是 `10`。

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` 應該是 `92`。

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens` 應該傳回正確的結果。

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
