---
id: 6a89c69c99298438d6fbe60b
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Ora aggiungi il caso ricorsivo. Dopo l'istruzione `if`, crea una variabile chiamata `count_list`. Assegna a essa il risultato della chiamata a `countup` con `number - 1`.

Ogni chiamata usa un numero più piccolo, quindi la ricorsione si avvicina al caso base.

# --hints--

Dovresti assegnare `countup(number - 1)` a una variabile chiamata `count_list`.

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
