---
id: 6a89c69c6af682856558d78b
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

La chiamata ricorsiva restituisce una lista contenente i numeri più piccoli. Usa `append()` per aggiungere `number` a `count_list`, poi restituisci `count_list`.

# --hints--

Dovresti aggiungere `number` a `count_list`.

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

Dovresti restituire `count_list` dalla tua funzione `countup`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

La tua funzione `countup` dovrebbe restituire `[1, 2, 3]` quando viene chiamata con `3`.

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
