---
id: 6a89c69c91bc7365fbd4b733
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Aggiungi il caso base alla tua `countup` function. Sostituisci `pass` con un'istruzione `if` che verifica se `number` è minore di `1`. All'interno dell'istruzione `if`, restituisci una lista vuota.

# --hints--

La tua `countup` function dovrebbe restituire una lista vuota quando `number` è minore di `1`.

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

La tua `countup` function non dovrebbe restituire una lista vuota quando `number` è maggiore o uguale a `1`.

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
