---
id: 6a89c69c8dbfda74151fa0b6
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, creerai una `countup` function che usa la ricorsione per creare una lista di numeri da `1` a un numero dato.

Inizia definendo una function chiamata `countup` con un parametro chiamato `number`. Aggiungi la parola chiave `pass` al corpo della function in modo che sia valida mentre è vuota.

# --hints--

Dovresti avere una function chiamata `countup` con un parametro chiamato `number`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

La tua `countup` function dovrebbe contenere la parola chiave `pass`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
