---
id: 61fd71d596e8f253b9408b39
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Das `tr`-Element wird verwendet, um eine Tabellenzeile anzuzeigen. Fügen Sie ein `tr`-Element innerhalb Ihres `thead`-Elements hinzu. Fügen Sie in Ihrem neuen `tr`-Element ein `td`-Element hinzu, gefolgt von drei `th`-Elementen.

Das `td`-Element kennzeichnet eine Datenzelle, während das `th`-Element eine Kopfzeile anzeigt.

# --hints--

Ihr `thead`-Element sollte ein `tr`-Element enthalten.

```js
assert(document.querySelector('thead')?.children?.[0]?.localName === 'tr');
```

Ihr `tr`-Element sollte als erstes Kind ein `td`-Element haben.

```js
assert(document.querySelector('tr')?.children?.[0]?.localName === 'td');
```

Ihr `tr`-Element sollte nach dem `th`-Element drei `td`-Elemente enthalten.

```js
assert(document.querySelector('tr')?.children?.[1]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[2]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[3]?.localName === 'th');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
