---
id: 61fd719788899952e67692b9
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Die `thead`- und `tbody`-Elemente werden verwendet, um anzugeben, welcher Teil Ihrer Tabelle die Kopfzeile ist und welcher Teil die Hauptdaten oder den Hauptinhalt enthält.

Fügen Sie Ihrer ersten `thead` unterhalb des `tbody`-Elements ein `table` und `caption` hinzu.

# --hints--

Ihr erstes `table`-Element sollte ein `thead`-Element enthalten.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

Ihr erstes `table`-Element sollte ein `tbody`-Element enthalten.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

Ihr `thead`-Element sollte sich direkt unter Ihrem `caption`-Element befinden.

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

Ihr `tbody`-Element sollte sich direkt unter Ihrem `thead`-Element befinden.

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
