---
id: 61fd6cc9475a784b7776233e
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Fügen Sie unter Ihrem bestehenden `div`-Element ein neues `div`-Element mit einem `class` hinzu, das auf `table-wrap` gesetzt ist. Dies wird der Container für Ihre Tabellen sein.

Fügen Sie darin drei `table`-Elemente hinzu. Sie werden CSS verwenden, um diese zu einer einzigen Tabelle zu gestalten, aber die Verwendung separater Tabellen hilft Screenreadern, den Dokumentenfluss besser zu verstehen.

# --hints--

Sie sollten ein neues `div`-Element erstellen.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Ihr neues `div`-Element sollte das `class` auf `table-wrap` gesetzt haben.

```js
assert(document.querySelector('.table-wrap')?.localName === 'div');
```

Ihr `.table-wrap`-Element sollte nach Ihrem bestehenden `div` kommen.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap'));
```

Ihr `.table-wrap`-Element sollte drei `table`-Elemente enthalten.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert(children?.length === 3);
children.forEach(child => assert(child?.localName === 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
