---
id: 61fd6b7c83dbf54a08cf0498
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Fügen Sie innerhalb Ihres `div`-Elements drei `span`-Elemente hinzu. Geben Sie jedem von ihnen ein `class`-Attribut mit dem Wert `year` und fügen Sie den folgenden Text in der Reihenfolge hinzu: `2019`, `2020` und `2021`.

# --hints--

Ihr `div`-Element sollte drei `span`-Elemente enthalten.

```js
assert(document.querySelector('div')?.children?.length === 3);
```

Jedes `span`-Element sollte ein `class`-Attribut mit dem Wert `year` haben.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert(span?.classList?.contains('year')));
```

Ihr erstes `span` sollte den Text `2019` enthalten.

```js
assert(document.querySelector('div')?.children?.[0]?.textContent === '2019');
```

Ihr zweites `span` sollte den Text `2020` enthalten.

```js
assert(document.querySelector('div')?.children?.[1]?.textContent === '2020');
```

Ihr drittes `span` sollte den Text `2021` enthalten.

```js
assert(document.querySelector('div')?.children?.[2]?.textContent === '2021');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
