---
id: 61fd77f7ad2aeb5ae34d07d6
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Fügen Sie in Ihrem zweiten `tr`-Element ein `th`-Element mit dem Text `Checking Our primary transactional account.` hinzu. Wickeln Sie diesen Text, außer `Checking `, in ein `span`-Element mit dem `class`-Attribut, das auf `description` gesetzt ist.

Fügen Sie anschließend drei `td`-Elemente mit folgendem Text (in dieser Reihenfolge) hinzu: `$54`, `$56`, `$53`. Geben Sie dem dritten `td`-Element ein `class`-Attribut, das auf `current` gesetzt ist.

# --hints--

Ihr zweites `tr` sollte ein `th`-Element enthalten.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th'));
```

Ihr `th`-Element sollte den Text `Checking Our primary transactional account.` haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th')?.innerText === 'Checking Our primary transactional account.');
```

Sie sollten den Text `Our primary transactional account.` in ein `span`-Element einwickeln.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.textContent === 'Our primary transactional account.');
```

Ihr `span`-Element sollte das `class`-Attribut auf `description` gesetzt haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.classList?.contains('description'));
```

Sie sollten drei `td`-Elemente haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td').length === 3);
```

Ihr erstes `td`-Element sollte den Text `$54` haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[0]?.textContent === '$54');
```

Ihr zweites `td`-Element sollte den Text `$56` haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[1]?.textContent === '$56');
```

Ihr drittes `td`-Element sollte den Text `$53` haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.textContent === '$53');
```

Ihr drittes `td`-Element sollte das `class` auf `current` gesetzt haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
