---
id: 61fd5a93fd62bb35968adeab
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Richten Sie Ihr HTML mit den `DOCTYPE`, `html` ein, die anzeigen, dass dieses Dokument auf Englisch ist, `head` und `body` Elemente.

Geben Sie Ihrem `head` Element die passenden `meta` Elemente für das `charset` und `viewport`, ein `title` Element mit einem passenden Titel und ein `link` Element für Ihr Stylesheet.

# --hints--

Ihr Code sollte eine `<!DOCTYPE html>` Deklaration enthalten.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Ihr Code sollte ein `html` Element enthalten.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Ihr `html` Element sollte ein öffnendes Tag mit einem `lang` Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr `html` Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Ihr Code sollte ein `head` Element innerhalb des `html` Elements enthalten.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Ihr Code sollte ein `body` Element innerhalb des `html` Elements enthalten.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Ihr `head` Element sollte vor Ihrem `body` Element stehen.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Sie sollten zwei `meta` Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Ein `meta` Element sollte ein `name` mit dem Wert `viewport` und `content` mit dem Wert `width=device-width, initial-scale=1.0` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Das andere `meta` Element sollte das `charset` Attribut mit dem Wert `UTF-8` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Ihr Code sollte ein `title` Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ihr `title` sollte etwas Text enthalten.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Ihr Code sollte ein `link` Element enthalten.

```js
assert(/<link/.test(code))
```

Ihr `link` Element sollte innerhalb Ihres `head` Elements sein.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

Ihr `link` Element sollte ein `rel` Attribut mit dem Wert `stylesheet` haben.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

Ihr `link` Element sollte ein `href` Attribut mit dem Wert `styles.css` haben.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
