---
id: 61fd67a656743144844941cb
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Screenreader geben HTML-Elemente basierend auf dem Dokumentenfluss aus. Wir möchten schließlich, dass die Bilanz eine Überschrift „Balance Sheet“ und eine Unterüberschrift „AcmeWidgetCorp“ hat. Diese Reihenfolge ergibt jedoch keinen Sinn, wenn sie von einem Screenreader vorgelesen wird.

Geben Sie Ihrem bestehenden `span` das `class`-Attribut mit dem Wert `flex` und fügen Sie zwei `span`-Elemente darin ein. Geben Sie dem ersten den Text `AcmeWidgetCorp`. Geben Sie dem zweiten den Text `Balance Sheet`. Sie werden CSS verwenden, um die Reihenfolge des Texts auf der Seite umzukehren, aber die HTML-Reihenfolge ist für einen Screenreader sinnvoller.

# --hints--

Ihr bestehendes `span`-Element sollte das `class`-Attribut mit dem Wert `flex` haben.

```js
assert(document.querySelector('h1')?.children?.[0]?.classList?.contains('flex'));
```

Ihr bestehendes `span`-Element sollte zwei neue `span`-Elemente darin enthalten.

```js
assert(document.querySelector('.flex')?.children?.[0]?.localName === 'span');
assert(document.querySelector('.flex')?.children?.[1]?.localName === 'span');
```

Ihre neuen `span`-Elemente sollten kein `class`-Attribut haben.

```js
assert(!document.querySelector('.flex')?.children?.[0]?.classList?.length);
assert(!document.querySelector('.flex')?.children?.[1]?.classList?.length);
```

Ihr erstes neues `span`-Element sollte den Text `AcmeWidgetCorp` haben.

```js
assert(document.querySelector('.flex')?.children?.[0]?.textContent === 'AcmeWidgetCorp');
```

Ihr zweites neues `span`-Element sollte den Text `Balance Sheet` haben.

```js
assert(document.querySelector('.flex')?.children?.[1]?.textContent === 'Balance Sheet');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
