---
id: 587d7dab367417b2b2512b70
title: "introduzione al currying e all'applicazione parziale"
challengeType: 1
forumTopicId: 301232
dashedName: introduction-to-currying-and-partial-application
---

# --description--

L'<dfn>arità</dfn> di una funzione è il numero di argomenti che richiede. Il <dfn>currying</dfn> di una funzione significa trasformare una funzione di arità N in N funzioni di arità 1.

In altre parole, ristruttura una funzione in modo che prenda un argomento, poi restituisca un'altra funzione che prende il prossimo argomento, e così via.

Ecco un esempio:

```js
function unCurried(x, y) {
  return x + y;
}

function curried(x) {
  return function(y) {
    return x + y;
  }
}

const curried = x => y => x + y

curried(1)(2)
```

`curried(1)(2)` restituirebbe `3`.

Questo è utile nel tuo programma se non puoi fornire tutti gli argomenti a una funzione in una sola volta. Puoi salvare ogni chiamata di funzione in una variabile, che conterrà il riferimento alla funzione restituita che prende il prossimo argomento quando sarà disponibile. Ecco un esempio che usa la funzione curried nell'esempio sopra:

```js
const funcForY = curried(1);
console.log(funcForY(2)); // 3
```

Allo stesso modo, l'<dfn>applicazione parziale</dfn> può essere descritta come l'applicare alcuni argomenti a una funzione alla volta e restituire un'altra funzione che viene applicata a più argomenti. Ecco un esempio:

```js
function impartial(x, y, z) {
  return x + y + z;
}

const partialFn = impartial.bind(this, 1, 2);
partialFn(10); // 13
```

# --instructions--

Completa il corpo della funzione `add` in modo che usi il currying per aggiungere i parametri `x`, `y` e `z`.

# --hints--

`add(10)(20)(30)` dovrebbe restituire `60`.

```js
assert(add(10)(20)(30) === 60);
```

`add(1)(2)(3)` dovrebbe restituire `6`.

```js
assert(add(1)(2)(3) === 6);
```

`add(11)(22)(33)` dovrebbe restituire `66`.

```js
assert(add(11)(22)(33) === 66);
```

Il tuo codice dovrebbe includere un'istruzione finale che restituisce `x + y + z`.

```js
assert(__helpers.removeJSComments(code).match(/[xyz]\s*?\+\s*?[xyz]\s*?\+\s*?[xyz]/g));
```

# --seed--

## --seed-contents--

```js
function add(x) {
  // Only change code below this line


  // Only change code above this line
}

add(10)(20)(30);
```

# --solutions--

```js
const add = x => y => z => x + y + z
```
