---
id: 587d7daa367417b2b2512b6b
title: Dividere una stringa in un array usando il metodo split
challengeType: 1
forumTopicId: 18305
dashedName: split-a-string-into-an-array-using-the-split-method
---

# --description--

Il metodo `split` divide una stringa in un array di stringhe. Prende un argomento per il delimitatore, che può essere un carattere usato per spezzare la stringa o un'espressione regolare. Per esempio, se il delimitatore è uno spazio, ottieni un array di parole, e se il delimitatore è una stringa vuota, ottieni un array di ogni carattere nella stringa.

Ecco due esempi che dividono una stringa per spazi, poi un'altra per cifre usando un'espressione regolare:

```js
const str = "Hello World";
const bySpace = str.split(" ");

const otherString = "How9are7you2today";
const byDigits = otherString.split(/\d/);
```

`bySpace` avrebbe il valore `["Hello", "World"]` e `byDigits` avrebbe il valore `["How", "are", "you", "today"]`.

Poiché le stringhe sono immutabili, il metodo `split` rende più facile lavorarci.

# --instructions--

Usa il metodo `split` all'interno della funzione `splitify` per dividere `str` in un array di parole. La funzione dovrebbe restituire l'array. Nota che le parole non sono sempre separate da spazi, e l'array non dovrebbe contenere punteggiatura.

# --hints--

Il tuo codice dovrebbe usare il metodo `split`.

```js
assert(__helpers.removeJSComments(code).match(/\.split/g));
```

`splitify("Hello World,I-am code")` dovrebbe restituire `["Hello", "World", "I", "am", "code"]`.

```js
assert(
  JSON.stringify(splitify('Hello World,I-am code')) ===
    JSON.stringify(['Hello', 'World', 'I', 'am', 'code'])
);
```

`splitify("Earth-is-our home")` dovrebbe restituire `["Earth", "is", "our", "home"]`.

```js
assert(
  JSON.stringify(splitify('Earth-is-our home')) ===
    JSON.stringify(['Earth', 'is', 'our', 'home'])
);
```

`splitify("This.is.a-sentence")` dovrebbe restituire `["This", "is", "a", "sentence"]`.

```js
assert(
  JSON.stringify(splitify('This.is.a-sentence')) ===
    JSON.stringify(['This', 'is', 'a', 'sentence'])
);
```

# --seed--

## --seed-contents--

```js
function splitify(str) {
  // Only change code below this line


  // Only change code above this line
}

splitify("Hello World,I-am code");
```

# --solutions--

```js
function splitify(str) {
  return str.split(/\W/);
}
```
