---
id: 587d7b90367417b2b2512b65
title: Restituire una parte di un array usando il metodo slice
challengeType: 1
forumTopicId: 301239
dashedName: return-part-of-an-array-using-the-slice-method
---

# --description--

Il metodo `slice` restituisce una copia di alcuni elementi di un array. Può prendere due argomenti, il primo indica l'indice da cui iniziare il taglio, il secondo è l'indice dove terminare il taglio (e non è incluso). Se gli argomenti non sono forniti, il valore predefinito è iniziare dall'inizio dell'array fino alla fine, il che è un modo semplice per creare una copia dell'intero array. Il metodo `slice` non modifica l'array originale, ma ne restituisce uno nuovo.

Ecco un esempio:

```js
const arr = ["Cat", "Dog", "Tiger", "Zebra"];
const newArray = arr.slice(1, 3);
```

`newArray` avrebbe il valore `["Dog", "Tiger"]`.

# --instructions--

Usa il metodo `slice` nella funzione `sliceArray` per restituire una parte dell'array `anim` dati gli indici `beginSlice` e `endSlice` forniti. La funzione dovrebbe restituire un array.

# --hints--

Il tuo codice dovrebbe usare il metodo `slice`.

```js
assert(__helpers.removeJSComments(code).match(/\.slice/g));
```

La variabile `inputAnim` non dovrebbe cambiare.

```js
assert(
  JSON.stringify(inputAnim) ===
    JSON.stringify(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 3)` dovrebbe restituire `["Dog", "Tiger"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 3)) ===
    JSON.stringify(['Dog', 'Tiger'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 0, 1)` dovrebbe restituire `["Cat"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 0, 1)) ===
    JSON.stringify(['Cat'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 4)` dovrebbe restituire `["Dog", "Tiger", "Zebra"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 4)) ===
    JSON.stringify(['Dog', 'Tiger', 'Zebra'])
);
```

# --seed--

## --seed-contents--

```js
function sliceArray(anim, beginSlice, endSlice) {
  // Only change code below this line


  // Only change code above this line
}

const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
sliceArray(inputAnim, 1, 3);
```

# --solutions--

```js
function sliceArray(anim, beginSlice, endSlice) {
  return anim.slice(beginSlice, endSlice);
}
const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
```
