---
id: 587d7da9367417b2b2512b67
title: aggiungere elementi alla fine di un array usando concat invece di push
challengeType: 1
forumTopicId: 301226
dashedName: add-elements-to-the-end-of-an-array-using-concat-instead-of-push
---

# --description--

La programmazione funzionale riguarda la creazione e l’uso di funzioni non mutanti.

L’ultima sfida ha introdotto il metodo `concat` come modo per unire array in un nuovo array senza mutare gli array originali. Confronta `concat` con il metodo `push`. `push` aggiunge elementi alla fine dello stesso array su cui viene chiamato, il che muta quell’array. Ecco un esempio:

```js
const arr = [1, 2, 3];
arr.push(4, 5, 6);
```

`arr` avrebbe un valore modificato di `[1, 2, 3, 4, 5, 6]`, che non è il modo della programmazione funzionale.

`concat` offre un modo per unire nuovi elementi alla fine di un array senza effetti collaterali mutanti.

# --instructions--

Modifica la funzione `nonMutatingPush` in modo che usi `concat` per unire `newItem` alla fine di `original` senza mutare gli array `original` o `newItem`. La funzione deve restituire un array.

# --hints--

Il tuo codice dovrebbe usare il metodo `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

Il tuo codice non dovrebbe usare il metodo `push`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?push/g));
```

L’array `first` non dovrebbe cambiare.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

L’array `second` non dovrebbe cambiare.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingPush([1, 2, 3], [4, 5])` dovrebbe restituire `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingPush([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingPush(original, newItem) {
  // Only change code below this line
  return original.push(newItem);

  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingPush(first, second);
```

# --solutions--

```js
function nonMutatingPush(original, newItem) {
  return original.concat(newItem);
}
const first = [1, 2, 3];
const second = [4, 5];
```
