---
id: 587d7dab367417b2b2512b6f
title: Usa il metodo some per verificare se qualche elemento in un array soddisfa un criterio
challengeType: 1
forumTopicId: 301314
dashedName: use-the-some-method-to-check-that-any-elements-in-an-array-meet-a-criteria
---

# --description--

Il metodo `some` lavora con gli array per verificare se *qualche* elemento supera un test specifico. Restituisce un valore booleano: `true` se almeno uno dei valori soddisfa il criterio, `false` altrimenti.

Ad esempio, il seguente codice controllerebbe se qualche elemento nell'array `numbers` è minore di 10:

```js
const numbers = [10, 50, 8, 220, 110, 11];

numbers.some(function(currentValue) {
  return currentValue < 10;
});
```

Il metodo `some` restituirebbe `true`.

# --instructions--

Usa il metodo `some` all'interno della funzione `checkPositive` per verificare se qualche elemento in `arr` è positivo. La funzione dovrebbe restituire un valore booleano.

# --hints--

Il tuo codice dovrebbe usare il metodo `some`.

```js
assert(__helpers.removeJSComments(code).match(/\.some/g));
```

`checkPositive([1, 2, 3, -4, 5])` dovrebbe restituire `true`.

```js
assert(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` dovrebbe restituire `true`.

```js
assert(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([-1, -2, -3, -4, -5])` dovrebbe restituire `false`.

```js
assert(!checkPositive([-1, -2, -3, -4, -5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.some(elem => elem > 0);
}
```
