---
id: 587d7daa367417b2b2512b6c
title: combinare un array in una stringa usando il metodo join
challengeType: 1
forumTopicId: 18221
dashedName: combine-an-array-into-a-string-using-the-join-method
---

# --description--

Il metodo `join` serve a unire gli elementi di un array per creare una stringa. Accetta un argomento che indica il delimitatore usato per separare gli elementi dell'array nella stringa.

Ecco un esempio:

```js
const arr = ["Hello", "World"];
const str = arr.join(" ");
```

`str` avrebbe il valore della stringa `Hello World`.
# --instructions--

Usa il metodo `join` (tra gli altri) all'interno della funzione `sentensify` per creare una frase dalle parole nella stringa `str`. La funzione deve restituire una stringa. Per esempio, `I-like-Star-Wars` verrebbe convertito in `I like Star Wars`. Per questa sfida, non usare il metodo `replace`.

# --hints--

Il tuo codice dovrebbe usare il metodo `join`.

```js
assert(__helpers.removeJSComments(code).match(/\.join/g));
```

Il tuo codice non dovrebbe usare il metodo `replace`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?replace/g));
```

`sentensify("May-the-force-be-with-you")` dovrebbe restituire una stringa.

```js
assert(typeof sentensify('May-the-force-be-with-you') === 'string');
```

`sentensify("May-the-force-be-with-you")` dovrebbe restituire la stringa `May the force be with you`.

```js
assert(sentensify('May-the-force-be-with-you') === 'May the force be with you');
```

`sentensify("The.force.is.strong.with.this.one")` dovrebbe restituire la stringa `The force is strong with this one`.

```js
assert(
  sentensify('The.force.is.strong.with.this.one') ===
    'The force is strong with this one'
);
```

`sentensify("There,has,been,an,awakening")` dovrebbe restituire la stringa `There has been an awakening`.

```js
assert(
  sentensify('There,has,been,an,awakening') === 'There has been an awakening'
);
```

# --seed--

## --seed-contents--

```js
function sentensify(str) {
  // Only change code below this line


  // Only change code above this line
}

sentensify("May-the-force-be-with-you");
```

# --solutions--

```js
function sentensify(str) {
  return str.split(/\W/).join(' ');
}
```
