---
id: 587d7da9367417b2b2512b66
title: Unire due array usando il metodo concat
challengeType: 1
forumTopicId: 301229
dashedName: combine-two-arrays-using-the-concat-method
---

# --description--

<dfn>Concatenazione</dfn> significa unire elementi uno dopo l'altro. JavaScript offre il metodo `concat` sia per le stringhe sia per gli array, che funziona allo stesso modo. Per gli array, il metodo viene chiamato su uno, quindi un altro array viene fornito come argomento a `concat`, che viene aggiunto alla fine del primo array. Restituisce un nuovo array e non modifica nessuno dei due array originali. Ecco un esempio:

```js
[1, 2, 3].concat([4, 5, 6]);
```

L'array restituito sarebbe `[1, 2, 3, 4, 5, 6]`.

# --instructions--

Usa il metodo `concat` nella funzione `nonMutatingConcat` per concatenare `attach` alla fine di `original`. La funzione dovrebbe restituire l'array concatenato.

# --hints--

Il tuo codice dovrebbe usare il metodo `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

L'array `first` non dovrebbe cambiare.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

L'array `second` non dovrebbe cambiare.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingConcat([1, 2, 3], [4, 5])` dovrebbe restituire `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingConcat([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingConcat(original, attach) {
  // Only change code below this line


  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingConcat(first, second);
```

# --solutions--

```js
function nonMutatingConcat(original, attach) {
  return original.concat(attach);
}
const first = [1, 2, 3];
const second = [4, 5];
```
