---
id: 6853e58d52ec64349a5760ca
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次研讨会中，你将实现最短路径算法。你将编写一个 Python 函数，用于计算图中节点之间的最短路径，并且还返回所经过的路径。

例如，给定一个城市通过不同距离的道路连接的图，算法将查找从一个城市到另一个城市的最短路线。如果你想从 City A 旅行到 City D，算法可能会查找出走 A ⇨ B ⇨ C ⇨ D（总计：15km）比直接走 A ⇨ D（20km）更短。

要开始，定义一个名为 `INF` 的变量并赋值为 `float('inf')`，它表现正无穷大。稍后，你将用它来表示两个节点之间的无限距离。

# --hints--

你应该在你的 `file` 顶部定义变量 `INF`。

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

你应该将 `float('inf')` 赋值给变量 `INF`。

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
