---
id: 686253f1b051998ad4904e3e
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

你现在将创建主函数，该函数接受三个参数：邻接矩阵、起始节点和一个可选的目标节点。

创建一个名为 `shortest_path` 的函数，接收三个参数：`matrix`、`start_node` 和 `target_node`。为 `target_node` 赋值为默认的 `None`。

`target_node` 参数默认设置为 `None`，表示如果未指定目标节点，函数应计算从起始节点到图中所有其他节点的最短路径。

暂时在函数体内添加一个 `pass` 语句。

# --hints--

你应该创建一个名为 `shortest_path` 的函数。

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

该函数应接受三个参数：`matrix`、`start_node` 和 `target_node`，其中 `target_node` 的默认值为 `None`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
