---
id: 68625ad34aa3f58d7f3b9b07
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

你需要保存图中节点的数量。为此你需要一个与邻接矩阵长度匹配的变量。

在 `shortest_path` 函数内，创建一个变量 `n` 并将其设置为 `matrix` 的长度。

# --hints--

你应该将 `n` 初始化为 `matrix` 的长度。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("n").is_equivalent("n = len(matrix)")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    pass

--fcc-editable-region--
```
