---
id: 686251fb49ebc489ba60ccb5
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你需要创建一个二维列表来表现图的邻接矩阵。该矩阵将用于表现图中节点之间边的权值。

创建一个名为 `adj_matrix` 的变量，并赋值一个二维列表，表现具有以下权值的图：

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

你应该创建一个名为 `adj_matrix` 的变量。

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` 应该是一个包含所提供值的二维列表。

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
