---
id: 686b73007fc582db0132c361
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

你需要跟踪从起始节点到图中每个其他节点的已知最短距离。

为此，创建一个名为 `distances` 的变量，并将其初始化为包含单个元素 `INF` 的列表。  


# --hints--

你应该创建一个名为 `distances` 的变量。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("distances")`
        )
      );
    }
});
```

`distances` 列表应初始化为 `[INF]`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF]")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)

--fcc-editable-region--
```
