---
id: 686b82af323098e7a5203bf4
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

当算法运行时，你需要跟踪你已经访问过的节点，以免多次处理它们。

为此，创建一个名为 `visited` 的列表，并为每个节点初始化为 `False`。

# --hints--

你应该在 `shortest_path` 函数内有一个名为 `visited` 的列表。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("visited")`
        )
      );
    }
});
```

你应该使用 `[False] * n` 为每个节点初始化 `visited` 列表为 `False`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("visited").is_equivalent("visited = [False] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]

--fcc-editable-region--
```
