---
id: 686b78bc13d518e1267448e4
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

除了跟踪距离，你还需要跟踪到达每个节点所经过的实际路径。

你将创建一个列表，其中每个条目保存到达该节点所经过的路径。最初，每个节点的路径只包含它自己。  

列表推导式提供了一种简洁的方式来创建列表。例如：

```py
[x * 2 for x in range(3)]
```

创建一个名为 `paths` 的变量，并使用列表推导式初始化它，该推导式为 `range(n)` 中的每个 `node_no` 创建一个包含 `[node_no]` 的列表。

# --hints--

你应该有一个名为 `paths` 的变量。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

你应该使用列表推导式初始化 `paths`。

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

你的列表推导应使用 `node_no` 来迭代 `range(n)`。

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

你的列表推导式应对 `range(n)` 中的每个 `node_no` 计算 `[node_no]`。  

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0

--fcc-editable-region--
```
