---
id: 68760c4c2648e7832c5eef16
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

在查找最近未访问节点的循环之后，你需要查看是否实际找到了有效节点。

如果不存在这样的节点，则意味着剩余节点无法从起始节点到达，算法应提前停止。

在与嵌套的 `for` 循环同一级别，添加一个 `if` 语句，查看 `current == -1`，如果为真则跳出循环。

# --hints--

你应该有一个 `if` 语句，用来查看在内部 `for` 循环之后 `current` 是否仍然是 `-1`。

```js
({
    test: () => {
      runPython(`
      cond = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
      assert cond.is_equivalent("current == -1") or cond.is_equivalent("-1 == current")
      `);
    }
});
```

你应该使用 `break` 语句跳出循环。

```js
({
    test: () => {
      runPython(`
      assert _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("break")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

--fcc-editable-region--

--fcc-editable-region--
```
