---
id: 6876089f7a500b7c782e5b83
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

你需要决定当前节点是否比你已经找到的节点（如果有的话）更合适。为此，你将在循环内添加一个条件式语句。

该条件应执行两件事：

- 查看节点是否尚未被访问。
- 将已知距离与当前的 `min_distance` 进行比较。

在内部的 `for` 循环中，添加一个 `if` 语句，查看 `node_no` 是否未被访问 **且** `distances[node_no]` 是否小于 `min_distance`。

为你的条件式块添加 `pass` 作为占位符。

# --hints--

你应该在你的内层 `for` 循环中添加一个 `if` 语句。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

你的 `if` 语句应查看节点是否未访问（`visited[node_no]` 为假值）且 `distances[node_no]` 小于 `min_distance`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass

--fcc-editable-region--
```
