---
id: 68760464eb9c2e79cc913fac
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

在此步骤中，你将添加一个循环，该循环将针对图中的每个节点运行一次。此循环将允许算法在多次遍历中更新距离和路径。

创建一个运行 `n` 次的 `for` 循环。使用 `_` 作为循环变量，因为你不需要使用迭代值。  

在 `loop` 内，你需要通过创建两个 `variable` 来准备选择下一个要处理的节点。 

- 一个用于保存当前迭代中找到的最小距离
- 还有一个用于保存具有此最小距离的节点索引。

创建变量 `min_distance` 和 `current`，并分别将它们设置为 `INF` 和 `-1`。 

# --hints--

你应该在 `shortest_path` 函数中有一个 `for` 循环。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

你的 `for` 循环应迭代 `range(n)`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

你的 for 循环应使用 `_` 作为迭代变量。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

你应该在 `for` 循环内创建一个名为 `min_distance` 的变量并将其设置为 `INF`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

你应该在 for 循环内创建一个名为 `current` 的变量并将其设置为 `-1`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
