---
id: 687775bed85144353d298665
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

现在你已经选择并标记了一个节点为已访问，是时候查看它的所有邻居，看看你是否能找到通往它们的更短路径。

在 `visited[current] = True` 这一行之后，添加一个 `for` 循环，迭代 `range(n)` 中的 `node_no`。

在这个循环内，创建一个变量 `distance` 并将其设置为 `matrix[current][node_no]`。这将给你当前节点到邻居节点的距离。

# --hints--

你应该在将当前节点标记为已访问后添加一个 `for` 循环。

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 2
      `)
      );
    }
});
```

循环应迭代 `range(n)` 中的 `node_no`。

```js
({
    test: () => {
      runPython(`  
loop = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1]  
assert loop.find_for_vars().is_equivalent("node_no")  
assert loop.find_for_iter().is_equivalent("range(n)")  
`);
    }
});
```

你应该在 `loop` 内创建一个变量 `distance`。

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
neighbor_loop = outer_loop.find_bodies()[0].find_for_loops()[1]
neighbor_loop.find_bodies()[0].has_variable("distance")
      `)
      );
    }
});
```

变量 `distance` 应设置为 `matrix[current][node_no]`。

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].is_equivalent("distance = matrix[current][node_no]")
      `)
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]  
                current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--

--fcc-editable-region--
```
