---
id: 68777816e95d4936f129819f
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在尝试更新到邻居的距离之前，你需要验证该邻居既可达又未被访问。然后你将计算通过当前节点到达该邻居的总距离。

在 `for` 循环内，添加一个 `if` 语句来查看：

- `distance` 不等于 `INF`（意味着节点之间存在一条边）
- 邻居 `node_no` 尚未被访问

在 `if` 块内，创建一个名为 `new_distance` 的变量，并赋值为 `distances[current]`（到当前节点的最短距离）和 `distance`（从当前节点到邻居的距离）之和。

# --hints--

你应该有一个 `if` 语句，用来查看 `distance` 是否不相等于 `INF`，并且邻居节点尚未被访问。

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

你应该在 `if` 块内创建一个名为 `new_distance` 的变量。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

你应该将 `new_distance` 赋值为 `distances[current]` 和 `distance` 的和。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
