---
id: 687b45e16f38ad5f4e4bd799
title: 步骤-23
challengeType: 20
dashedName: step-23
---

# --description--

既然你已经确定某个节点应该被显示，你需要格式化它的 `path`，以便能够清晰地打印。为此你将使用一个生成器表达式。

生成器表达式类似于列表推导式，但它不是创建列表，而是一次生成一个值。它使用括号 `()` 而不是方括弧 `[]`。例如：

```py
numbers = [1, 2, 3]
squared = (x**2 for x in numbers)  # Generator expression
```

在 `if` 语句后的循环内，创建一个名为 `string_path` 的变量。为它赋值一个生成器表达式，该表达式使用 `str()` 将 `paths[node_no]` 中的每个节点编号转换为字串。 

生成器表达式应迭代 `paths[node_no]` 中的每个节点编号 `n`。

# --hints--

你应该创建一个名为 `string_path` 的变量。

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].has_variable("string_path")
      `)
      );
    }
});
```

你应该将生成器表达式赋值给 `string_path`。

```js
({
  test: () => {
      runPython(`
      import ast
      assert isinstance(_Node(_code).find_function("shortest_path").find_for_loops()[1].find_bodies()[0].find_variable("string_path").tree.value, ast.GeneratorExp)
      `)
 }
})
```

你的生成器表达式应迭代 `paths[node_no]`。

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert path_var.find_comp_iters()[0].is_equivalent("paths[node_no]")
      `);
    }
});
```

你的生成器表达式应使用 `n` 作为迭代变量来迭代 `paths[node_no]`。

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_targets()[0].is_equivalent("n")
      `);
    }
});
```

你的生成器表达式应对 `paths[node_no]` 中的每个 `n` 计算 `str(n)`。

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_expr().is_equivalent("str(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no

        if current == -1:
            break
        visited[current] = True

        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
        if node_no == start_node or distances[node_no] == INF:
            continue

--fcc-editable-region--
```
