---
id: 69f0a37c3bf194a7c9c62492
title: Construir um livro razão de empréstimo de dispositivo
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

Você é um desenvolvedor de software encarregado de criar uma aplicação de livro razão para gerenciar dispositivos de hardware de TI fornecidos pelo Service Desk da sua empresa. O livro razão é um objeto cujas chaves são tags de asset e cujos valores são objetos:

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve modelar o livro razão como um objeto onde cada dispositivo é identificado pela sua tag de asset e possui as propriedades `type`, `status`, `borrower` (com `name` e `email`) e `dueDate`.

1. Você deve implementar uma função `checkoutDevice(ledger, assetTag, borrower)` que recebe um objeto ledger e o clona sem modificar o original. No clone, deve definir o `name` e o `email` do objeto `borrower`, atualizar o `status` para `"CheckedOut"` e retornar um objeto no formato `{ ledger: updatedLedger, message: confirmationString }`, onde a `message` inclui a tag de asset e o `name` do objeto `borrower`.

   - Se o `assetTag` não existir no ledger, a função deve retornar o ledger inalterado junto com um `message` que inclua o asset tag e explique que ele não foi encontrado.
   - Se o dispositivo já estiver com o status `"CheckedOut"`, a função deve retornar o ledger inalterado e manter o borrower atual, junto com um `message` que inclua o asset tag e explique que o dispositivo já está emprestado.

1. Você deve implementar uma função `checkinDevice(ledger, assetTag)` que recebe um objeto ledger e o clona sem modificar o original. No clone, deve limpar o `name` e o `email` do objeto `borrower`, resetar o `dueDate` para `""`, atualizar o `status` para `"CheckedIn"` e retornar um objeto no formato `{ ledger: updatedLedger, message: confirmationString }`, onde a `message` inclui a tag de asset.

   - Se a `assetTag` não existir no ledger, a função deve retornar o ledger inalterado junto com uma `message` que inclui a tag de asset e explica que ela não foi encontrada.

1. Você deve implementar uma função `listOverdueDevices(ledger, today)` que recebe um objeto ledger e retorna um array dos objetos de dispositivo que estão atrasados, ordenados pelo `dueDate` em ordem crescente. Um dispositivo está atrasado quando seu `status` é `"CheckedOut"` e seu `dueDate` é estritamente anterior a `today`. Tanto `today` quanto `dueDate` são strings no formato mês/dia/ano, e o mês e o dia podem ou não ter zero à esquerda (por exemplo, `"9/5/2025"` e `"09/05/2025"` são válidos).

1. Você não deve usar o objeto `Date` do JavaScript em nenhuma parte do seu código.

1. Você deve implementar uma função `serializeLedger(ledger)` que converte o objeto ledger em uma string JSON, e uma função `loadLedger(json)` que converte uma string JSON de volta em um objeto JavaScript.

# --hints--

Você deve criar uma função chamada `checkoutDevice`.

```js
assert.isFunction(checkoutDevice);
```

Você deve criar uma função chamada `checkinDevice`.

```js
assert.isFunction(checkinDevice);
```

Você deve criar uma função chamada `listOverdueDevices`.

```js
assert.isFunction(listOverdueDevices);
```

Você deve criar uma função chamada `serializeLedger`.

```js
assert.isFunction(serializeLedger);
```

Você deve criar uma função chamada `loadLedger`.

```js
assert.isFunction(loadLedger);
```

Sua função `checkoutDevice` não deve modificar o ledger original.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

Sua função `checkoutDevice` deve definir o `name` e o `email` do borrower e definir o `status` para `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

Sua função `checkoutDevice` deve retornar uma mensagem de confirmação que inclui a tag de asset e o `name` do borrower.

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

Sua função `checkoutDevice` deve preservar dispositivos não relacionados e campos não alterados no ledger retornado.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Sua função `checkoutDevice` deve retornar o ledger inalterado e uma mensagem que inclui a tag de asset quando a tag de asset não for encontrada.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Sua função `checkoutDevice` deve retornar o ledger inalterado, manter o borrower atual e retornar uma mensagem que inclui a tag de asset quando o dispositivo já estiver emprestado.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

Sua função `checkinDevice` não deve modificar o ledger original.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

Sua função `checkinDevice` deve limpar o `name` e o `email` do borrower, resetar o `dueDate` para `""` e definir o `status` para `"CheckedIn"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

Sua função `checkinDevice` deve retornar uma mensagem de confirmação que inclui a tag de asset.

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

Sua função `checkinDevice` deve preservar dispositivos não relacionados e campos não alterados no ledger retornado.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Sua função `checkinDevice` deve retornar o ledger inalterado e uma mensagem que inclui a tag de asset quando a tag de asset não for encontrada.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Sua função `listOverdueDevices` deve retornar um array.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

Sua função `listOverdueDevices` deve retornar apenas dispositivos atrasados com `status` igual a `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

Sua função `listOverdueDevices` deve retornar os dispositivos atrasados ordenados pelo `dueDate` em ordem crescente.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

Sua função `listOverdueDevices` deve ordenar corretamente mesmo quando o mês ou o dia não tiver zero à esquerda.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

Você não deve usar o objeto `Date` do JavaScript em nenhuma parte do seu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

Sua função `serializeLedger` deve retornar uma string JSON que representa o ledger passado para ela.

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

Sua função `loadLedger` deve retornar o objeto JavaScript que a string JSON representa.

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
