---
id: 660ef5105b8ba095307a0e50
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

もう一度 `character` 変数をコンソールにログしてください。コンソールには文字列 `"Hello"` と文字列 `"World"` が表示されるはずです。

# --hints--

`console.log` をもう一度使う必要があります。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

`character` をコンソールにもう一度ログしてください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


新しい `console.log()` は再代入の後に来るべきです。

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
