---
id: 660ef55dd468079679ee0092
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

次の変数を宣言する準備ができました。`console.log` の両方の文と、`character` の再代入を削除してください。

また、`secondCharacter` の変数も削除しますが、`character` の初期化はそのままにしてください。

# --hints--

コードに `console.log()` の文が含まれていないことを確認してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

`character` の値を再代入してはいけません。

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

`character` の変数は `"Hello"` の値を持っている必要があります。

```js
assert.equal(character, "Hello");
```

`secondCharacter` の変数を持ってはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
