---
id: 660ef0f7c4b8e68ccd1f0786
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript には 7 種類のプリミティブデータ型があります。その 1 つが文字列 (String) です。 JavaScript では、<dfn>文字列</dfn>はひと続きの文字 (characters) を表し、シングルクォート (`'`) またはダブルクォート (`"`) で囲んで表記します。

ここで、文字列は<dfn>イミュータブル</dfn>であることに注意してください。つまり、一度作成された後は変更できません。 文字列を割り当てた変数に別の値を再割り当てすることは可能です。

では、`"Hello"` という文字列を、シングルクォートを使う形に書き直してください。

# --hints--

コードにダブルクォーテーションが入らないようにしてください。

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

`"Hello"` の文字列にはシングルクォートを使用してください。

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

コード内では引き続き `let` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

変数 `character` の宣言は引き続き必要です。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

変数 `character` には引き続き、文字列 `"Hello"` を値として持たせてください。

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
