---
id: 660ef6355e8f5a9e67fe5f46
title: ステップ 16
challengeType: 1
dashedName: step-16
---

# --description--

`number` データ型を使うと、加算のような数学的な操作ができます。`count + 1` をコンソールに出力してみてください。

# --hints--

コード内で `console` にアクセスしてください。

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

`.log()` の `console` にアクセスしてください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

`count + 1` をコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
