---
id: 660ef857f2806aa626d29d17
title: ステップ 17
challengeType: 1
dashedName: step-17
---

# --description--

引き算 (`-`)、掛け算 (`*`)、割り算 (`/`) も実行できます。 `console.log` の括弧内の演算子や数値をいろいろ変更して実験してみましょう。 次に進む準備ができたら、`console.log` を削除してくださしてください。

# --hints--

コード内に `console.log()` ステートメントが無いようにしてください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

変数 `character` の値はまだ `"Hello"` である必要があります。

```js
assert.equal(character, "Hello");
```

変数 `count` の値はまだ `8` である必要があります。

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
