---
id: 660ef31a5be625914a0102cd
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

`let` キーワードで変数を宣言すると、その変数の値を後で<dfn>再代入</dfn>（または変更）できます。この例では、`programmer` の値が `"Naomi"` から `"CamperChan"` に変わっています。

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

変数を再代入するときは、`let` キーワードを再度使わないことに注意してください。

`console.log` の後に、`"World"` 変数に値 `character` を代入してください。

# --hints--

コード内で `character` を2回使う必要があります。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

`character` を再代入するには代入演算子を使ってください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

コード内で文字列 `"World"` を使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

`character` 変数の値は `"World"` であるべきです。

```js
assert.equal(character, "World");
```

再代入の際に `let` を使わないでください。

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
