---
id: 660ef857f2806aa626d29d17
title: ステップ 17
challengeType: 1
dashedName: step-17
---

# --description--

引き算 (`-`)、掛け算 (`*`)、割り算 (`/`) も行えます。`console.log` の中の演算子や数字を自由に試してみてください。準備ができたら、`console.log` を削除してください。

# --hints--

コードに `console.log()` が含まれていないことを確認してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

`character` 変数はまだ `"Hello"` の値を持っているはずです。

```js
assert.equal(character, "Hello");
```

`count` 変数はまだ `8` の値を持っているはずです。

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
