---
id: 660ef0f7c4b8e68ccd1f0786
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScriptには7つのプリミティブデータ型があり、その中の1つがStringです。JavaScriptでは、<dfn>文字列</dfn>は文字のシーケンスを表し、シングルクォート（`'`）または二重引用符（`"`）のどちらかで囲むことができます。

文字列は<dfn>不変</dfn>であるため、一度作成されると変更できません。ただし、変数には別の値を代入することはできます。

`"Hello"`の文字列をシングルクォートに変更してください。

# --hints--

コード内に二重引用符があってはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

`"Hello"`の文字列にはシングルクォートを使ってください。

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

コード内では引き続き`let`を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

`character`の変数は引き続き宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

`character`の変数には文字列`"Hello"`が値として代入されている必要があります。

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
