---
id: 660eebd83100d37862268781
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

あなたの `character` 変数には現在値がありません。<dfn>代入</dfn>演算子 `=` を使って値を代入できます。例えば：

```js
let hello = "Hello";
```

変数を宣言するときに値を代入することは <dfn>初期化</dfn> と呼ばれます。

`character` 変数を宣言時に値 `"Hello"` を代入して初期化してください。

# --hints--

代入演算子 `=` を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

文字列 `"Hello"` を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

`"Hello"` 文字列には二重引用符を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

`"Hello"` を `character` 変数に代入してください。

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
