---
id: 660ef19b95d3308e7dd31bb6
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>コンソール</dfn>はJavaScriptの出力を出力して表示することができます。`console.log()`を使ってコンソールに情報を送ることができます。例えば、このコードは`"Naomi"`をコンソールに出力します。

```js
let developer = "Naomi";
console.log(developer);
```

上のコードは`developer`の中の名前で`console.log()`変数にアクセスしています。丸括弧の中の値が出力される値であることに注意してください。

`character`変数の値をコンソールに出力してください。その後、「Console」ボタンをクリックしてJavaScriptコンソールを表示してください。

# --hints--

`character`変数の初期化は変更しないでください。

```js
assert.strictEqual(character, 'Hello');
```

コード内で`console`にアクセスしてください。

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

`log`の`console`プロパティにアクセスしてください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

`.log()`メソッドを呼び出すには丸括弧を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

`character`変数をコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
