---
id: 660ef5c1904955978a986a5c
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

`let` を使って `count` 変数を宣言してください。<dfn>number</dfn> 型の `8` を代入してください。数値を使う場合は、引用符は使いません。例えば：

```js
let money = 100;
```

# --hints--

コードには2つ目の `let` キーワードが必要です。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

`let` を使って `count` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

`8` 変数に数値 `count` を代入してください。

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
