---
id: 660ee6e3a242da6bd579de69
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

プログラミングで最も重要な概念の一つが変数です。<dfn>変数</dfn>は値を格納する特定のメモリアドレスを指します。変数には名前が付けられ、その名前を使ってコード全体でその値にアクセスできます。

変数を宣言するとは、名前を付けることを意味します。JavaScriptでは、これは多くの場合`let`キーワードで行います。例えば、`hello`変数を宣言する方法は次の通りです。

```js
let hello;
```

変数名には特定のルールがあります：名前には文字、数字、ドル記号、アンダースコアを含めることができますが、スペースは含められず、数字で始めてはいけません。

`let`キーワードを使って`character`という名前の変数を宣言してください。

_注意_: JavaScriptでは文の末尾にセミコロンを付けるのが一般的な慣習です。`;`

# --hints--

コード内で`let`を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

コード内で`character`を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

`let`を使って`character`変数を宣言するべきです。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

宣言はセミコロンで終わるべきです。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
