---
id: 660f07d231941bc11719f664
title: ステップ 21
challengeType: 1
dashedName: step-21
---

# --description--

配列は<dfn>ミュータブル</dfn>と見なされる点で特別です。これは、インデックスの値を直接変更できることを意味します。

たとえば、このコードは配列の2番目の要素に数値`25`を代入します。

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

`rows`配列の**3番目**の要素を数値`10`に更新してください。その後、`rows`配列をコンソールに出力してください。

# --hints--

`rows`配列に対して再びブラケット記法を使うべきです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

`rows`配列の3番目の要素にアクセスするべきです。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

`rows`配列の3番目の要素に代入演算子を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

`10`配列の3番目の要素に値`rows`を代入するべきです。

```js
assert.equal(rows[2], 10);
```

コード内に2つ目の`console.log`文があるべきです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

`rows`配列をログに出力するべきです。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
