---
id: 660f4377a359972c521d3f4b
title: ステップ 79
challengeType: 1
dashedName: step-79
---

# --description--

`true` は文字通り「真」であるため、コンソールに文字列が表示されます。

では、`if` 文の条件をブール値の `false` に変更してください。

# --hints--

`if` 文の条件として `false` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (true) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
