---
id: 660f1e3f047bf4e403268713
title: ステップ 39
challengeType: 1
dashedName: step-39
---

# --description--

残念ながら、今はループが何をしているのか見えません。

`let` を使って `result` 変数を宣言し、空の文字列を代入してください。空の文字列は、`""` のように、間に何もない引用符で表されます。

# --hints--

`result` 変数を宣言する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

`let` を使って `result` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

`result` 変数は空の文字列である必要があります。

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
