---
id: 660f0c34aad72dc712b97624
title: ステップ 27
challengeType: 1
dashedName: step-27
---

# --description--

このプロジェクトに欠かせないもう一つのメソッドは `.pop()` メソッドです。これは配列の最後の要素を取り除き、その要素を<dfn>返します</dfn>。

メソッドが値を返すとき、それはその値をあなたに返していると考えられ、コードの他の部分で使えるようになります。

`popped` という新しい変数を作成し、`rows.pop()` の結果を代入してください。そして、`popped` をコンソールに出力してください。

# --hints--

`popped` という変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

`let` を使って `popped` という変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

`.pop()` メソッドを呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

`.pop()` メソッドを `rows` 配列に対して呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

`popped` 変数をコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
