---
id: 660f061d259bbebc37461080
title: ステップ 20
challengeType: 1
dashedName: step-20
---

# --description--

配列の中の値には、その値の<dfn>インデックス</dfn>を使ってアクセスできます。インデックスは配列内の値の位置を表す数字で、最初の値は`0`から始まります。

`array[0]`のように、<dfn>ブラケット記法</dfn>を使って値にアクセスできます。

`console.log`とブラケット記法を使って、`rows`配列の最初の値を出力してください。

# --hints--

コード内に`console.log()`文があるはずです。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

`rows`配列にアクセスしてください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

`rows`配列でブラケット記法を使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

`rows`配列の最初の要素にブラケット記法でアクセスしてください。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

`rows`配列の最初の要素を出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
