---
id: 660f280dda5040f707c76b4a
title: ステップ 45
challengeType: 1
dashedName: step-45
---

# --description--

今、`#` 文字が並んでいますが、ピラミッドの形はまだできていません。幸いなことに、`i` 変数はループ内の現在の「行」番号を表しているので、これを使ってピラミッドのような構造を作ることができます。

これを実現するために、文字列に使える `.repeat()` メソッドを使います。このメソッドは引数として数字を受け取り、その数字の回数だけ対象の文字列を繰り返します。例えば、`.repeat()` を使って `"Code! Code! Code!"` という文字列を生成します。

```js
const activity = "Code! ";
activity.repeat(3);
```

`.repeat()` メソッドを `character` に使い、引数に `i` を渡してください。

# --hints--

`.repeat()` メソッドを使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

`.repeat()` メソッドを `character` 変数に使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

`i` を `.repeat()` メソッドに渡すべきです。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

`.repeat()` メソッドを `.push()` メソッドの中で使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
