---
id: 660f4f79e2a82a4e92290f44
title: ステップ 110
challengeType: 1
dashedName: step-110
---

# --description--

時々、コメントアウトした以前のコードを復活させたい場合があります。 そのコードの周りにある `/*` と `*/` を削除すれば復活できます。 これを<dfn>コメント解除</dfn>といいます。

では、最初の `for` ループだけをコメント解除してください。 一行コメントと、他の 2 つの複数行コメントはそのままにしてください。

# --hints--

最初の `for` ループをコメント解除してください。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

一行コメントは削除しないでください。

```js
assert.match(code, /\/\//);
```

`while` ループはコメント解除しないでください。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

2 番目の `for` ループはコメント解除しないでください。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
