---
id: 660f1a00ac619ddc1e259a66
title: ステップ 35
challengeType: 1
dashedName: step-35
---

# --description--

`for` ループの<dfn>条件</dfn>は、ループが何回反復するかを指定します。`condition` が false になると、ループは停止します。

JavaScript では、ブール値は `true` または `false` のどちらかになります。これらは文字列ではありません — 違いについては後で詳しく学びます。

今回は、<dfn>小なり演算子</dfn>（`<`）を使います。これは左側の値が右側の値より小さいかどうかをチェックできます。例えば、`count < 3` は `true` が `count` の場合 `2` に評価され、`false` が `count` の場合は `4` に評価されます。

`"condition"` の文字列を、`i` が `count` より小さいかどうかをチェックする条件に置き換えてください。

# --hints--

小なり演算子を使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /</);
```

`i` が `count` より小さいかどうかをチェックするために小なり演算子を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

`for` ループは `i < count` を条件として使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
