---
id: 660f505d02b2bd513a1c3468
title: ステップ 111
challengeType: 1
dashedName: step-111
---

# --description--

ピラミッドが逆さまではなくなりました。これは、配列の**末尾**に新しい行を追加しているためです。

ループの本体を更新して、新しい行を配列の先頭に追加してください。

# --hints--

`unshift`の`rows`メソッドを使うべきです。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

`padRow()`メソッドの引数として`.unshift()`の呼び出しを渡すべきです。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

`i`の呼び出しの最初の引数として`padRow()`を渡すべきです。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

`count`の呼び出しの2番目の引数として`padRow()`を渡すべきです。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
