---
id: 660f374d532dc41189cc9cc2
title: ステップ 69
challengeType: 1
dashedName: step-69
---

# --description--

ここで少し数学の時間です。3段のピラミッドを考えてみましょう。中央揃えにしたい場合、次のような形になります。

```md
··#··
·###·
#####
```

空白は読みやすくするために中黒（・）に置き換えられています。このパターンを推測すると、各行の先頭と末尾の空白は一定の規則に従っていることがわかります。

空白の文字列を `rowCount - rowNumber` 回繰り返すように更新してください。

結果を見るためにコンソールを開いてください。

# --hints--

`.repeat()` の文字列に対して `" "` を呼び出し、`rowCount - rowNumber` 回繰り返すようにしてください。

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
