---
id: 660f48a419b40238e2b8b4d5
title: ステップ 93
challengeType: 1
dashedName: step-93
---

# --description--

ループはもう`continueLoop`変数に依存していません。これにより、その変数は<dfn>未使用の宣言</dfn>となります。一般的に、将来の混乱を避けるために未使用の宣言は避けるべきです。

`continueLoop`変数を削除してください。

# --hints--

もう`continueLoop`変数は存在しないはずです。

```js
assert.notMatch(__helpers.removeJSComments(code), /continueLoop/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
