---
id: 660f44f10ea40f300b896a5e
title: ステップ 84
challengeType: 1
dashedName: step-84
---

# --description--

`if...else if...else` 文を使う練習ができたので、コードからそれらを削除してください。

それが終わったら、`let` を使って `continueLoop` 変数を宣言し、ブール値の `false` を代入してください。次に `let` を使って `done` 変数を宣言し、値 `0` を代入してください。

# --hints--

`if` 文は使わないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

`else if` 文は使わないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

`else` 文は使わないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

`let` を使って `continueLoop` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

`continueLoop` 変数には値 `false` を代入してください。

```js
assert.isFalse(continueLoop);
```

`let` を使って `done` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

`done` 変数には値 `0` を代入してください。

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
