---
id: 660f0ee51d7460ce88cd248d
title: ステップ 29
challengeType: 1
dashedName: step-29
---

# --description--

コンソールに`4`が表示されると思っていましたか？`.push()`は、指定した値を追加した後の配列の新しい長さを返します。

メソッドがどんな値を返すかを理解することは重要です。`.push()`や`.pop()`を使って少し試してみてください。準備ができたら、すべての`.push()`と`.pop()`の呼び出し、および`console.log`の文を削除してください。

# --hints--

`.push()`の呼び出しはしてはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

`.pop()`の呼び出しはしてはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

ログ文は一切あってはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

`popped`の変数はあってはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

`pushed`の変数はあってはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
