---
id: 660f415b76859a2736771607
title: ステップ 78
challengeType: 1
dashedName: step-78
---

# --description--

ピラミッドがまた消えてしまいましたが、大丈夫です。 これは予想通りです。

新しいループを作成する前に、`if` 文について学びましょう。 <dfn>`if` 文</dfn>を使用すると、条件 (condition) が満たされた場合にのみコードブロックを実行できます。 次の構文を使用します。

```js
if (condition) {
  logic
}
```

では、条件にブール値 `true` を使用して、`if` 文を作成してください。 本文では文字列 `"Condition is true"` を出力してください。

# --hints--

`if` 文を作成してください。

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

`if` 文の条件として `true` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

`if` 文の本文では `"Condition is true"` をログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
