---
id: 660f455b044d3230ed971e98
title: ステップ 85
challengeType: 1
dashedName: step-85
---

# --description--

<dfn>`while`</dfn> ループは、指定された `condition` がもはや真でなくなるまで繰り返し実行されます。構文は次の通りです。

```js
while (condition) {
  logic;
}
```

この構文を使って、条件に `while` を持つ `continueLoop` ループを宣言してください。ループの本体は空にしてください。

# --hints--

`while` ループを使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

`while` ループは条件に `continueLoop` を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
