---
id: 660f1b6e60bd9edf902c81fd
title: ステップ 36
challengeType: 1
dashedName: step-36
---

# --description--

<dfn>反復</dfn>文は、ループの各回の後にイテレーターをどう扱うかを指示します。

変数に再代入するとき、再代入前の値を参照するためにその変数を使うことができます。これにより、既存の数値に3を加えるなどの操作が可能です。例えば、`bees = bees + 3` は `bees` の値を3増やします。

その構文を使って、`"iteration"` という文字列を `i` を1増やす再代入文に置き換えてください。

# --hints--

`i` 変数に1を加える必要があります。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

`i + 1` を `i` 変数に再代入してください。

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

`for` ループは各反復で `i` を `1` だけ増やす必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
