---
id: 660f09a2694b59c3a10ee304
title: ステップ 23
challengeType: 1
dashedName: step-23
---

# --description--

まずは最初のコンソールログと`rows[rows.length - 1]`の代入を削除してください。2つ目の`rows`のログ文は後で使うので残しておきましょう。

# --hints--

`console.log(rows[0])`の文は削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

`rows[rows.length - 1]`の再代入は削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

`console.log(rows)`の文は削除しないでください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
