---
id: 660f229d2dbe09ef2954a4a1
title: ステップ 42
challengeType: 1
dashedName: step-42
---

# --description--

前のループで、`+` の値を `i` だけ増やすために加算演算子 `1` を使ったことを覚えていますか。

文字列の値でも同じようなことができ、既存の文字列に新しい文字列を追加できます。例えば、`hello = hello + " World";` は `" World"` 変数に格納されている既存の文字列に `hello` という文字列を追加します。これを<dfn>連結</dfn>と呼びます。

あなたの `for...of` ループでは、加算演算子を使って `row` の値に `result` の値を連結してください。

# --hints--

`result` 変数に連結演算子を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

`row` 変数に `result` を連結するべきです。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

連結の結果を `result` 変数に代入してください。

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
