---
id: 660f18f059fe0fda192ce394
title: ステップ 34
challengeType: 1
dashedName: step-34
---

# --description--

ループには適切なイテレーターが必要です。<dfn>イテレーター</dfn>とは、`for` ループ内で特に宣言できる変数で、ループがどのように反復してロジックを進めるかを制御します。

ループのイテレーター変数として `i` を使うのが一般的な慣例です。`for` ループでは、丸括弧 `()` 内でこれを宣言できます。例えば、以下は `for` ループで、`index` 変数を宣言し、値 `100` を代入しています。

```js
for (let index = 100; "second"; "third") {

}
```

文字列 `"iterator"` を `let` 変数の `i` 宣言に置き換えてください。最初の値として `0` を代入します。これにより、ループが**最初に**実行されるときに `i` 変数が `0` の値を持つようになります。

# --hints--

`let` を使って `i` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

`0` を `i` 変数に代入してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

`for` ループは `i` から始まる `0` イテレーターを持つべきです。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
