---
id: 660f4e74f7fd3f4a99ac2e50
title: ステップ 108
challengeType: 1
dashedName: step-108
---

# --description--

配列には `.shift()` メソッドもあります。これは配列の**最初の**要素を削除します。最後の要素を削除する `.pop()` とは異なります。以下は `.shift()` メソッドの例です。

```js
const numbers = [1, 2, 3];
numbers.shift();
```

`numbers` 配列は `[2, 3]` になります。

`numbers` 配列のすぐ下に、`shifted` 変数を宣言し、`.shift()` 配列に対して `numbers` を呼び出した結果を代入してください。次の行で、`shifted` 変数をコンソールに出力してください。

# --hints--

`const` を使って `shifted` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

`.shift()` 配列に対して `numbers` を呼び出してください。

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

`.shift()` の呼び出し結果を `shifted` 変数に代入してください。

```js
assert.equal(shifted, 1);
```

`shifted` 変数をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
