---
id: 660f207334fabaeac3269c38
title: ステップ 41
challengeType: 1
dashedName: step-41
---

# --description--

`result` 文字列を操作するには、別の種類のループを使います。具体的には、イテラブルオブジェクトの各アイテムを繰り返し処理し、一時的に変数に代入する `for...of` ループです。

`for...of` ループの構文は次のようになります。

```js
for (const value of iterable) {

}
```

変数はループ全体ではなく、単一の繰り返し処理の間だけ存在するため、`const` を使えます。

`for...of` 配列を繰り返し処理し、各値を `rows` 変数に代入する `row` ループを作成してください。

# --hints--

別の `for` キーワードを使う必要があります。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

`row` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

`const` を使って `row` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

`for...of` ループは `row` 変数を宣言する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

`row` 変数は `rows` から `of` キーワードを使って取り出す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

`for...of` ループの本体は空にしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
