---
id: 660f415b76859a2736771607
title: ステップ 78
challengeType: 1
dashedName: step-78
---

# --description--

あなたのピラミッドがまた消えてしまいましたね。大丈夫です、それは予想されることです。

新しいループを作成する前に、`if` 文について学ぶ必要があります。<dfn>`if` 文</dfn>は、条件が満たされたときだけコードのブロックを実行できます。構文は次の通りです。

```js
if (condition) {
  logic
}
```

ブール値 `if` を条件にして `true` 文を作成してください。本文では文字列 `"Condition is true"` を出力してください。

# --hints--

`if` 文を作成してください。

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

`if` 文の条件は `true` にしてください。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

`if` の本文で `"Condition is true"` を出力してください。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
