---
id: 660f3ce51f70571e1c5227c8
title: ステップ 76
challengeType: 1
dashedName: step-76
---

# --description--

コメントは、なぜコードが特定の方法を取っているのか説明したり、将来の自分へのやることメモを残したりするのに役立ちます。

JavaScriptでは、`//` を使ってコード内に一行コメントを残せます。

`for` ループの上に一行コメントを追加して、コードを別の種類のループに変更することを自分に思い出させてください。

# --hints--

一行コメントは `//` で始める必要があります。

```js
assert.match(code, /\/\//);
```

一行コメントは少なくとも5文字以上にしてください。

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
