---
id: 660f359af3e32e0f1a6880b7
title: ステップ 68
challengeType: 1
dashedName: step-68
---

# --description--

コンソールに同じ文字の塊が表示されているはずです。あなたの`padRow`関数は、先ほど行っていたこととまったく同じことをしていますが、今は再利用可能な独立したセクションになっています。

繰り返した`" "`文字列の前後に、1つのスペース`character`を連結するために加算演算子を使ってください。

文字列を連結するには、`+`演算子をこのように使えることを覚えておいてください。

```js
" " + "string"
```

# --hints--

返される値の先頭に1つのスペースを連結してください。

```js
assert.match(padRow(1, 1), /^\s/);
```

返される値の末尾に1つのスペースを連結してください。

```js
assert.match(padRow(1, 1), /\s$/);
```

あなたの`padRow()`関数は、スペースが前後に付いた繰り返し`character`の文字列を返す必要があります。

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
