---
id: 660f039ff313dbb696b007ca
title: ステップ 19
challengeType: 1
dashedName: step-19
---

# --description--

配列が値、つまり<dfn>要素</dfn>を保持するとき、それらの値はカンマで区切られます。こちらは2つの文字列を保持する配列です：

```js
let array = ["first", "second"];
```

`rows` の宣言を、文字列 `"Naomi"`、`"Quincy"`、`"CamperChan"` を持つ配列に変更してください。配列内の値の順序は重要なので、その順序に従ってください。文字列は大文字と小文字が区別されることを忘れないでください。

# --hints--

配列の最初の要素は文字列 `"Naomi"` であるべきです。

```js
assert.equal(rows[0], "Naomi");
```

配列の2番目の要素は文字列 `"Quincy"` であるべきです。

```js
assert.equal(rows[1], "Quincy");
```

配列の3番目の要素は文字列 `"CamperChan"` であるべきです。

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
