---
id: 660f447efc0e722f016c1be0
title: ステップ 81
challengeType: 1
dashedName: step-81
---

# --description--

テキストが再び表示されました！これは `"false"` が文字列であり、ブール値として評価されると `true` になるためです。つまり、`"false"` は真値です。

<dfn>真値</dfn>とは、ブール値として評価したときに真とみなされる値のことです。JavaScriptで出会うほとんどの値は真値です。

<dfn>偽値</dfn>はその逆で、ブール値として評価したときに偽とみなされる値です。JavaScriptには定義された偽値のリストがあります。その中には `false`、`0`、`""`、`null`、`undefined`、および `NaN` が含まれます。

`if` の条件を空文字列の `""` に変更してみてください。これは偽値です。

# --hints--

`if` 文の条件には `""` を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
