---
id: 660f383d4c772c12ff59904b
title: ステップ 70
challengeType: 1
dashedName: step-70
---

# --description--

式全体を引数として渡すことができます。関数はその式を評価した結果を受け取ります。例えば、以下の2つの関数呼び出しは同じ結果になります。（演算の順序ルールPEMDAS（括弧、指数、乗算、除算、加算、減算）が適用されることに注意してください）：

```js
test(2 * 3 + 1);
test(7);
```

もう一度パターンを見てみましょう：

```md
··#··
·###·
#####
```

`character` の値を `2 * rowNumber - 1` 回繰り返すように更新してください。

コンソールを再度開いて、更新された結果を確認してください。

# --hints--

関数は `2 * rowNumber - 1` を `.repeat()` の呼び出しに渡す必要があります。

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
