---
id: 660f3915b41a441537ec9f5e
title: ステップ 72
challengeType: 1
dashedName: step-72
---

# --description--

`i` を `1` だけ増やしているので、<dfn>インクリメント演算子</dfn> `++` を使うことができます。この演算子は変数の値を1だけ増やし、その変数への代入を更新します。例えば、ここで `test` は `8` になります。

```js
let test = 7;
test++;
```

ループの繰り返しでの加算代入をインクリメント演算子に置き換えてください。

# --hints--

`for` のループでは `i` を使った加算代入を使わないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

`for` のループでは `i` にインクリメント演算子を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
