---
id: 660f1cedf3676fe26122ebf6
title: ステップ 38
challengeType: 1
dashedName: step-38
---

# --description--

コンソールに0から7までの数字が1行ずつ出力されているはずです。これはピラミッドを作る基礎になります。

ログ出力の文を、`i`を`rows`配列に追加する文に置き換えてください。

# --hints--

`console.log`の呼び出しは不要です。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

`.push()`配列に対して`rows`を呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

`i`配列に`rows`を追加してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

`.push()`は`for`ループの中で行う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
