---
id: 660f2fbd45b520046cac68e8
title: ステップ 65
challengeType: 1
dashedName: step-65
---

# --description--

以前のステップで、戻り値について学んだことを覚えていますか。関数はアプリケーションが別々に利用できる値を<dfn>戻す</dfn>ことができます。

関数内では、`return` キーワードを使って戻り値を指定します。例えば、この関数は最初のパラメータに渡された値を戻します。

```js
function name(parameter) {
  return parameter;
}
```

`return` キーワードを使って、`character` 変数の値を `rowNumber` 回繰り返して戻してください。

# --hints--

`.repeat()` メソッドを使うべきです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

`.repeat()` 変数に対して `character` メソッドを使うべきです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

`rowNumber` 呼び出しに `.repeat()` を渡すべきです。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

`return` キーワードを使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

`.repeat()` 呼び出しの結果を戻すべきです。

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
