---
id: 660f34e99571070d56d2f231
title: ステップ 67
challengeType: 1
dashedName: step-67
---

# --description--

あなたの `padRow` 関数には2つのパラメータがあります。関数が呼び出されるときに、それらのパラメータに値が渡されます。

関数呼び出しに渡す値は<dfn>引数</dfn>と呼ばれ、関数呼び出しに引数を<dfn>渡す</dfn>と言います。以下は `"Hello"` を引数として渡した関数呼び出しの例です。

```js
test("Hello");
```

`i + 1` と `count` を `padRow` の呼び出しに引数として渡してください。パラメータと同様に、引数はカンマで区切ります。

# --hints--

`i + 1` を `padRow()` の呼び出しに渡す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

`i + 1` の引数の後にカンマがあるはずです。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

`count` を2番目の引数として渡す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
