---
id: 660f0ee51d7460ce88cd248d
title: ステップ 29
challengeType: 1
dashedName: step-29
---

# --description--

`4` が表示されると予想できましたか？ `.push()` は、指定された値を追加した後の配列の新しい長さを返すのです。

メソッドがどのような値を返すかを知っておくことは重要です。 `.push()` と `.pop()` を使って、色々実験してみてください。 次に進む準備ができたら、すべての `.push()` および `.pop()` の呼び出しと、`console.log` ステートメントを削除してください。

# --hints--

`.push()` の呼び出しが無いようにしてください。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

`.pop()` の呼び出しが無いようにしてください。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

ログステートメントが無いようにしてください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

変数 `popped` が無いようにしてください。

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

変数 `pushed` が無いようにしてください。

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
