---
id: 660f17294346b7d69e79db3d
title: ステップ 32
challengeType: 1
dashedName: step-32
---

# --description--

これでピラミッドジェネレーターの作成を始める準備ができました。`character` 変数はピラミッドの構成要素として使います。

`"Hello"` はあまりうまく機能しないかもしれません。`character` の値をハッシュ記号の `"#"` に変更してください。

# --hints--

`character` 変数はハッシュ記号にしてください。

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
