---
id: 660f1e3f047bf4e403268713
title: ステップ 39
challengeType: 1
dashedName: step-39
---

# --description--

今はループが何をしているか見えません。

`let` を使って変数 `result` を宣言し、空の文字列を代入してください。 空の文字列は、`""` のように、間に何もない引用符のペアで表されます。

# --hints--

変数 `result` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

変数 `result` の宣言には `let` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

変数 `result` の値は空の文字列でなければなりません。

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
