---
id: 660f23b53db70af0f2620e78
title: ステップ 43
challengeType: 1
dashedName: step-43
---

# --description--

今、すべての数字が同じ行に表示されています。これではピラミッドを作成できません。

各行に改行を追加する必要があります。ただし、JavaScriptで引用符の間に改行キーを押して改行を入れると、構文エラーになります。代わりに、文字列がログに出力されるときに改行として解釈される特別な<dfn>エスケープシーケンス</dfn> `\n` を使う必要があります。例えば：

```js
lineOne = lineOne + "\n" + lineTwo;
```

`result` と `row` の値の後に改行を追加するために、2つ目の加算演算子を使ってください。

# --hints--

`\n` エスケープシーケンスを使うべきです。文字列である必要があるので、引用符で囲まれていることを覚えておいてください。

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

`row` 変数を `result` 変数に連結するべきです。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

`\n` エスケープシーケンスを `row` 変数に連結するべきです。

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

連結した全体を `result` 変数に代入してください。

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
