---
id: 660f2eccfe3f820304af1b39
title: ステップ 64
challengeType: 1
dashedName: step-64
---

# --description--

行のフォーマット方法を知るために、`padRow` 関数は現在の行番号と、生成される行の総数を知る必要があります。

これを行う最良の方法は、それらのために関数のパラメータを作成することです。`padRow` 関数に `rowNumber` と `rowCount` のパラメータを与えてください。複数のパラメータはカンマで区切ります。

```js
function name(first, second) {

}
```

# --hints--

`padRow` 関数には `rowNumber` のパラメータが必要です。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

`rowNumber` のパラメータの後にカンマを追加してください。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

`padRow` 関数には `rowCount` のパラメータが必要です。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
