---
id: 660f0a55847d6cc485f29ba5
title: ステップ 26
challengeType: 1
dashedName: step-26
---

# --description--

JavaScriptの<dfn>メソッド</dfn>とは、特定の値やオブジェクトに関連付けられた関数のことです。すでに出会った例としては、`.log()`メソッドがあり、これは`console`オブジェクトの一部です。

配列には独自のメソッドがあり、最初に探求するのは`.push()`メソッドです。これは配列の末尾に値を「プッシュ」することができます。以下は、配列の末尾に数値`12`を追加する例です。

```js
array.push(12);
```

`.push()`を使って、文字列`"freeCodeCamp"`を`rows`配列の末尾に追加してください。`console.log`の前にこのコードを追加して、配列に加えた変更を確認できるようにしましょう。

# --hints--

コード内で`.push()`を使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

`.push()`配列の`rows`メソッドを使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

`"freeCodeCamp"`メソッドに文字列`.push()`を渡す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
