---
id: 660f2a70ad6225fa503e71c3
title: ステップ 46
challengeType: 1
dashedName: step-46
---

# --description--

もう少しです！ここで遭遇しているのは、JavaScriptのようなゼロベースのインデックス言語でよくある<dfn>オフバイワンエラー</dfn>です。

`rows` 配列の最初のインデックスは `0` なので、`for` ループは `i = 0` から始まります。しかし、文字列をゼロ回繰り返すと何も出力されません。

これを直すには、`1` 呼び出しの `i` の値に `.repeat()` を加えてください。ループ条件のように `i` に代入しないでください。

# --hints--

`1` メソッドで `i` に `.repeat()` を加えるべきです。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
