---
id: 660f4efcb8068e4cb470dca1
title: ステップ 109
challengeType: 1
dashedName: step-109
---

# --description--

これらのメソッドの使い方を試したので、逆さピラミッドを作成するもう一つの方法に取り掛かりましょう。 まずは実験用のコードを整理します。

配列 `numbers` と、そのメソッドの呼び出し、ログの呼び出しを削除してください。

# --hints--

配列 `numbers` は必要ありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /numbers/);
```

変数 `unshifted` は必要ありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /unshifted/);
```

変数 `shifted` は必要ありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /shifted/);
```

`console.log` 文は必要ありません。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];
const shifted = numbers.shift();
console.log(shifted);
const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
