---
id: 660f2b6fd54ac1fc142804dd
title: ステップ 47
challengeType: 1
dashedName: step-47
---

# --description--

このピラミッドをフォーマットするロジックは複雑になりそうなので、そのコードを関数に切り出す絶好のタイミングです。

<dfn>関数</dfn>とは、アプリケーション全体で再利用できるコードのブロックです。関数は次の構文で宣言します。

```js
function name(parameter) {

}
```

`function` キーワードは、`name` 変数が関数になることをJavaScriptに伝えます。`parameter` は、関数が使われるときに渡される値を表す変数です。関数は好きなだけ多くても少なくても <dfn>パラメータ</dfn> を持てます。`for` ループのように、波括弧の間が <dfn>関数本体</dfn> です。

`padRow` 関数を宣言してください。まだパラメータ変数は作らないでください。関数本体は空にしてください。命名規則にはキャメルケースを使う必要があることを忘れないでください。

# --hints--

`function` キーワードを使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

`padRow` 関数を宣言するべきです。

```js
assert.isFunction(padRow);
```

`padRow()` 関数にはパラメータを持たせないでください。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

`padRow()` 関数本体は空にしてください。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
