---
id: 660f4cffb1459d45e34902d1
title: ステップ 106
challengeType: 1
dashedName: step-106
---

# --description--

逆ピラミッドは、前のように「逆方向に」ループしなくても実際に作れます。

これを行うには、新しい配列メソッドをいくつか覚える必要があります。まず `const` を使って `numbers` 変数を宣言してください。`1`、`2`、`3` の要素を持つ配列を代入します。次に `numbers` 配列をログに出力してください。

# --hints--

`const` を使って `numbers` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

`numbers` 変数は配列である必要があります。

```js
assert.isArray(numbers);
```

`numbers` 配列は `1`、`2`、`3` の順に要素を持つ必要があります。

```js
assert.deepEqual(numbers, [1,2,3]);
```

`numbers` 配列をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
