---
id: 660f0da9bf1035c9097af20a
title: ステップ 28
challengeType: 1
dashedName: step-28
---

# --description--

`"freeCodeCamp"` がコンソールに出力されているはずです。これは `.pop()` が配列から削除された値を返し、以前に `"freeCodeCamp"` を配列の末尾に追加したためです。

では、`.push()` は何を返すでしょうか？ 既存の `rows.push()` を新しい `pushed` 変数に代入し、ログに出力してください。

# --hints--

`pushed` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

`let` を使って `pushed` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

`rows.push("freeCodeCamp")` を `pushed` 変数に代入してください。

```js
assert.equal(pushed, 4);
```

`pushed` 変数をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
