---
id: 660f3ba3cceef11b6ba08b59
title: ステップ 75
challengeType: 1
dashedName: step-75
---

# --description--

残念ながら、今度はピラミッドの底辺が消えてしまいました。これは別の<dfn>オフバイワンエラー</dfn>を作ってしまったためです。

元のループは`i`の値を`0`から`7`まで回っていました。なぜなら`count`は`8`であり、条件は`i`が`count`より小さいことを要求していたからです。今のループは`i`の値を`1`から`7`まで回っています。

ループは`i`が`8`のときも回るように更新する必要があります。ロジックを見ると、これはループが`i`が`count`以下のときに回るべきだということです。これには以下演算子`<=`を使えます。

ループ条件を、`i`が以下演算子で`count`以下の間回るように更新してください。

# --hints--

`for`のループは`i`が`count`より小さいかどうかをチェックすべきではありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

`for`のループは`i`が以下演算子で`count`以下かどうかをチェックすべきです。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
