---
id: 660f1cedf3676fe26122ebf6
title: ステップ 38
challengeType: 1
dashedName: step-38
---

# --description--

コンソールに数字の 0 から 7 までが 1 行ずつ表示されたはずです。 これを基礎にして、ピラミッドを組み立てていきましょう。

ではログ文を削除し、代わりに `i` を配列 `rows` にプッシュする文に置き換えてください。

# --hints--

`console.log` 呼び出しは必要ありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

配列 `rows` に対して`.push()`を呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

配列 `rows` に `i` を追加 (プッシュ) してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

`.push()` は `for` ループの中で行いましょう。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
