---
id: 660f0f980e98e8cf77f1ce31
title: ステップ 30
challengeType: 1
dashedName: step-30
---

# --description--

`rows` の宣言を再び空の配列に代入してください。

また、`'Hello'` の文字列を再び二重引用符で囲むように変更してください。一般的にどちらを使っても問題ありませんが、プロジェクト全体で一貫性を持たせることが重要です。

# --hints--

`rows` の配列は空である必要があります。

```js
assert.empty(rows);
```

`"Hello"` の文字列は二重引用符を使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
