---
id: 660f165270622fd4ec0da3f7
title: ステップ 31
challengeType: 1
dashedName: step-31
---

# --description--

`let` キーワードで変数を宣言すると、その変数は再代入が可能になります。つまり、後で `character` をまったく異なる値に変更することができます。

このプロジェクトでは、これらの変数の値を変更したくありません。ですので、代わりに `const` を使って宣言してください。`const` 変数は特別です。

まず、`const` 変数は `let` 変数のように再代入できません。このコードはエラーになります。

```js
const firstName = "Naomi";
firstName = "Jessica";
```

また、`const` 変数は初期化されていない状態にできません。このコードはエラーになります。

```js
const firstName;
```

`let` キーワードを `const` に置き換えてください。

# --hints--

`const` を使って `character` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

`const` を使って `count` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

`const` を使って `rows` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

コード内で `let` を使わないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
