---
id: 660f033cf051ebb50ea3bf48
title: ステップ 18
challengeType: 1
dashedName: step-18
---

# --description--

プログラミングでは、多くのデータを扱う必要がよくあります。データを整理し管理するのに役立つ多くのデータ構造があります。最も基本的なデータ構造の一つが<dfn>配列</dfn>です。

<dfn>配列</dfn>は、一連の値を保持できる非プリミティブ型のデータ型です。非プリミティブ型は、より複雑なデータを保持できる点でプリミティブ型と異なります。文字列や数値のようなプリミティブ型は、一度に一つの値しか保持できません。

配列は大括弧（`[]`）を使って表されます。以下は空の配列の値を持つ変数の例です。

```js
let array = [];
```

`rows` 変数を宣言し、空の配列を代入してください。

# --hints--

`rows` 変数を持っているはずです。

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

`let` を使って `rows` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

`rows` 変数に空の配列を代入してください。

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
