---
id: 660f4934fb48f63abd5ae371
title: ステップ 95
challengeType: 1
dashedName: step-95
---

# --description--

`done` を使って生成された行数を追跡するのは機能しますが、実際にはロジックをもう少し整理できます。

配列には特別な `length` プロパティがあり、配列にいくつの値、つまり<dfn>要素</dfn>があるかを確認できます。このプロパティには `myArray.length` のような構文でアクセスします。

`rows.length` の `padRow` はオフバイワンエラーになることに注意してください。なぜなら `done` は呼び出しの*前に*インクリメントされるからです。

条件を `rows.length` が `count` より小さいかどうかをチェックするように更新してください。

# --hints--

`while` ループは `rows.length` が `count` より小さいかどうかをチェックする必要があります。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
