---
id: 660f45ccf4ca5c31f253005a
title: ステップ 86
challengeType: 1
dashedName: step-86
---

# --description--

今、`continueLoop` を true に変更すると、`while` ループが永遠に実行されます。これは <dfn>無限ループ</dfn> と呼ばれ、避けるように注意が必要です。無限ループはシステムをフリーズさせ、完全な再起動が必要になる場合があります。

これを避けるために、ループ内で `done` 変数の値をインクリメント演算子で増やすことから始めてください。

# --hints--

`while` ループは `done` 変数をインクリメントする必要があります。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
