---
id: 660f1bf673487ae0bb25b900
title: ステップ 37
challengeType: 1
dashedName: step-37
---

# --description--

ループは今、8回実行されるはずです。ループの本体内で、`i` イテレーターの値を出力して、何が起きるか確認してください。

# --hints--

`console.log()` を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

`i` の値を出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

`i` ループ内で `for` の値を出力してください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
