---
id: 660f46460f9c36330ebc07d8
title: ステップ 88
challengeType: 1
dashedName: step-88
---

# --description--

等値演算子は JavaScript で奇妙な動作を引き起こす場合があります。例えば、`"0" == 0` は、片方が文字列で片方が数値であっても true になります。

<dfn>厳密等値演算子</dfn> `===` は、2つの値が等しく、かつ同じ型であるかどうかをチェックするために使われます。一般的なルールとして、常にこの等値演算子を使うことをお勧めします。厳密等値演算子を使うと、`"0" === 0` は false になります。なぜなら、値はどちらもゼロかもしれませんが、型が同じではないからです。

`done == count` の条件を厳密等値演算子を使うように更新してください。

# --hints--

`if` の条件は厳密等値演算子を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
