---
id: 660f39b444fd6f16d1e49c1f
title: ステップ 73
challengeType: 1
dashedName: step-73
---

# --description--

`i + 1` を `padRow` の呼び出しに渡す代わりに、ループを `1` から始めることができます。これにより、1から始まるインデックスのループを作成できます。

イテレーターを `1` ではなく `0` から始めるように更新してください。

# --hints--

`for` ループは `i` を `1` で初期化する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
