---
id: 660f4f79e2a82a4e92290f44
title: ステップ 110
challengeType: 1
dashedName: step-110
---

# --description--

以前コメントアウトしたコードを戻したい場合があります。その場合は、そのコードの周りにある`/*`と`*/`を取り除いてください。これを<dfn>コメント解除</dfn>と呼びます。

最初の`for`ループだけをコメント解除してください。1行コメントと他の2つの複数行コメントはそのままにしてください。

# --hints--

最初の`for`ループをコメント解除してください。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

1行コメントは削除しないでください。

```js
assert.match(code, /\/\//);
```

`while`ループはコメント解除しないでください。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

2番目の`for`ループもコメント解除しないでください。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
