---
id: 660f38c34a4de6141c0c369f
title: ステップ 71
challengeType: 1
dashedName: step-71
---

# --description--

あなたのピラミッドジェネレーターは期待通りに動作しています。しかし、これは書いたコードをさらに詳しく探求する絶好の機会です。

加算演算子は変数に値を加える唯一の方法ではありません。<dfn>加算代入</dfn>演算子は「変数の元の値を取り、この値を加え、その結果を変数に代入する」という意味の省略形として使えます。例えば、これら2つの文は同じ結果になります。

```js
test = test + 1;
test += 1;
```

`for`ループの繰り返し文を加算代入を使うように更新してください。

# --hints--

あなたの`for`ループは`i = i + 1`を使うべきではありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

あなたの`for`ループは`i`で加算代入を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
