---
id: 660f34626216270c682e2f7b
title: ステップ 66
challengeType: 1
dashedName: step-66
---

# --description--

<dfn>関数呼び出し</dfn>は、関数を実際に使うことを可能にします。気づいていなかったかもしれませんが、これまで使ってきた `.push()` のようなメソッドは関数呼び出しです。

関数は関数名を参照し、`()` を付けることで呼び出されます。`test` 関数を呼び出す方法は次の通りです。

```js
test();
```

`character.repeat(i + 1)` 呼び出しの `.push()` を、`padRow` 関数の関数呼び出しに置き換えてください。まだ引数は追加しないでください。

# --hints--

`i + 1` 呼び出しで `push` を使わないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

`character.repeat` 呼び出しで `.push()` を使わないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

`padRow` 呼び出しで `.push()` を呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

`padRow` 呼び出しに引数を入れないでください。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
