---
id: 660f47afe4c98536715d5fa4
title: ステップ 91
challengeType: 1
dashedName: step-91
---

# --description--

<dfn>厳密不等価</dfn>演算子 `!==` は、2つの値が等しくないか、または同じ型でないかを確認できます。構文は等値演算子と似ています：`value !== 4`。

現在、`while` ループは `continueLoop` が true の場合にのみ実行されます。`while` ループの条件を更新して、`done` が `count` と等しくないかどうかをチェックしてください。

# --hints--

`while` ループは `done` と `count` が等しくないかどうかをチェックする必要があります。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
