---
id: 660f17d4e9f227d86e834abd
title: ステップ 33
challengeType: 1
dashedName: step-33
---

# --description--

ピラミッドを作成するには、複数の行を作る必要があります。条件が満たされるまで繰り返し処理を行う必要がある場合、<dfn>ループ</dfn>を使います。ループを書く方法はいくつもあります。

まずは基本的な`for`ループから始めましょう。`for`ループは次の構文を使います。

```js
for (iterator; condition; iteration) {
  logic;
}
```

今後のステップで、ループの各コンポーネントを詳しく見ていきます。ここでは、3つのコンポーネントとして`for`、`"iterator"`、`"condition"`を含む`"iteration"`ループを作成してください。ループの<dfn>本体</dfn>、つまり波括弧 `{}` 内の部分は空にしてください。


# --hints--

`for`ループを作成してください。

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

`for`ループの最初のコンポーネントは文字列`"iterator"`にしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

`for`ループの2番目のコンポーネントは文字列`"condition"`にしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

`for`ループの3番目のコンポーネントは文字列`"iteration"`にしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

`for`ループの本体は空にしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
