---
id: 660f255022991ef34ed0ee88
title: ステップ 44
challengeType: 1
dashedName: step-44
---

# --description--

数字を出力しても見た目に美しいピラミッドにはなりません。`rows` 配列のフォーマット済みの内容を出力しているので、元のループを更新する時です。

配列に `i` を追加する代わりに、`character` 変数の値を追加してください。

# --hints--

もう `i` 変数を追加する必要はありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

`character` 変数を追加してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
