---
id: 660f4b641290da41b2cf0dd9
title: ステップ 103
challengeType: 1
dashedName: step-103
---

# --description--

加算と同様に、減算に使えるさまざまな演算子があります。<dfn>減算代入</dfn>演算子 `-=` は、指定された値を現在の変数の値から引き、その結果を変数に代入します。

繰り返し文を、減算代入演算子を使った正しい文に置き換えてください。

# --hints--

あなたの `for` ループは `i = i - 1` を使わないようにしてください。

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /for\(leti=count;i>0;i=i-1\)\{rows\.push\(padRow\(i,count\)\);/);
```

あなたの `for` ループは減算代入を使って `i` を `1` ずつ減らすようにしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*-=\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i = i - 1) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
