---
id: 660f34e99571070d56d2f231
title: ステップ 67
challengeType: 1
dashedName: step-67
---

# --description--

`padRow` 関数には、先ほど定義した 2 つの仮引数があります。 関数が呼び出されるとき、これらの仮引数に値が提供されます。

関数呼び出しに提供される値は<dfn>実引数 (argument)</dfn> と呼ばれ、実引数を関数呼び出しに<dfn>渡す</dfn>と言います。 こちらは、`"Hello"` を実引数として渡す関数呼び出しの例です。

```js
test("Hello");
```

`i + 1` と `count` を、`padRow` 関数呼び出しの実引数として渡してください。 仮引数と同じように、実引数もコンマで区切ります。

# --hints--

`i + 1` を `padRow()` 呼び出しに渡してください。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

実引数 `i + 1` の後にコンマを入れてください。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

`count` を 2 番目の実引数として渡してください。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
