---
id: 660f4455f457ef2e3ec6920f
title: ステップ 80
challengeType: 1
dashedName: step-80
---

# --description--

今、文字列は出力されなくなりました。なぜなら `false` は `true` ではないからです。でも、他の値はどうでしょうか？

条件を文字列 `"false"` に変更してみてください。

# --hints--

`if` 文の条件には文字列 `"false"` を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*('|")false\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (false) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
