---
id: 660f3b664421471aa595170f
title: ステップ 74
challengeType: 1
dashedName: step-74
---

# --description--

ピラミッドが少し変に見えています。`1`ではなく`0`からループを始めているので、`i`に渡すときに`padRow`に1を足す必要はありません。

`padRow`の呼び出しの最初の引数を`i`に更新してください。

# --hints--

`i + 1`に`padRow`を渡してはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

`i`に`padRow`を渡す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
