---
id: 6610bbed59bc2a0194d85533
title: ステップ 22
challengeType: 1
dashedName: step-22
---

# --description--

配列 `rows` の中の値が直接変更されたことに気づきましたか？ これを「<dfn>ミューテーション</dfn>」といいます。 配列についてもっと学ぶと、いつ配列を変更 (mutate) すべきか、いつすべきでないかがわかります。

次に進む前に、絶好の機会なので一般的な配列の使い方を学びましょう。 現在、あなたのコードは `rows[2]` を使用して配列の最後の要素にアクセスしています。 しかし、最後の element を取得したいけれど、配列の中に何個の要素があるのか分からないこともあります。

そのような場合、配列の `.length` プロパティを使うことができます。これは配列内の要素の数を返します。 任意の配列の最後の要素を取得するには、次の構文が使用できます:

```js
array[array.length - 1]
```

`array.length` は配列内の要素の数を返します。 そこから `1` を引くことで、配列の最後の要素のインデックスを取得できます。 この考え方を `rows` 配列にも適用することができます。

`rows[2]` というコードを、配列 `rows` の最後の要素に動的にアクセスするように変更しましょう。 上記の例を参考にしてください。

コンソールの出力内容は何も変わらないはずです。

# --hints--

配列 `rows` の `.length` プロパティを使用してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

ブラケット記法で `rows.length` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

ブラケット記法で、配列の長さから `1` を引いてください。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
