---
id: 6610c77d50636722e5b6be17
title: ステップ 60
challengeType: 1
dashedName: step-60
---

# --description--

`return` 文の下に、文字列 `"This works!"`をコンソールにログ出力するコードを書いてみましょう。

そうすると、文字列 `"This works!"` がコンソールに表示されず、`console.log("This works!")` の行がグレーアウトされているのがわかります。

そのコンソールログ文をコピーして、`return` 文の上に貼り付けてください。 今度は、文字列 `"This works!"` がコンソールに表示されるはずです。

`return` キーワードに関して知っておくべき重要なことは、それが関数から返される値を定義するだけでなく、関数またはブロック文内のコードの実行を中止するということです。 つまり、`return` 文より後のコードは実行されません。

# --hints--

`padRow` 関数は変数 `test` を返す必要があります。

```js
assert.equal(padRow("Naomi"), "Testing");
```

最初の `console.log`は、`return` キーワードの後に置く必要があります。

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

2 つ目の `console.log` は、`return` キーワードの前に置く必要があります。

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";

  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
