---
id: 6610c48c4ea0891afa7c4696
title: ステップ 54
challengeType: 1
dashedName: step-54
---

# --description--

もう一度コンソールを開くと、`padRow` 関数が戻り値を定義しているにもかかわらず `undefined` を返しているのがわかります。これは、関数を呼び出すときにパラメータに値を渡す必要があるためです。

関数呼び出しに値を渡すと、その値は <dfn>引数</dfn> と呼ばれます。以下は `demo` 関数を呼び出し、`"Naomi"` を `name` パラメータの引数として渡す例です。

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

`name` の呼び出しで、`padRow` パラメータに自分の名前を引数として渡してください。名前は文字列なので、引用符で囲む必要があります。

# --hints--

`padRow()` の呼び出しには文字列を渡してください。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
