---
id: 6610b741b54b90f0c0fb3d58
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

変数名が複数の単語で構成される場合、単語の大文字・小文字の付け方には特定の慣例があります。JavaScriptでは、使う慣例は<dfn>キャメルケース</dfn>です。

キャメルケースとは、名前の最初の単語はすべて小文字にし、続く単語はすべてタイトルケースにすることを意味します。キャメルケースの例をいくつか示します。

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

キャメルケースを使って新しい`secondCharacter`変数を宣言してください。

# --hints--

`secondCharacter`変数を宣言する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

`let`変数を宣言するには`secondCharacter`を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

`secondCharacter`変数に値を代入してはいけません。セミコロンを忘れないでください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
