---
id: 6610b8017d1671f2814e8c77
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

変数を初期化せずに宣言すると、それは<dfn>未初期化</dfn>とみなされます。現在、あなたの`secondCharacter`変数は未初期化です。

`console.log()`を追加して、`secondCharacter`変数の値を確認してください。

# --hints--

2つ目の`console.log`を追加してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

`secondCharacter`変数をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
