---
id: 6610b8f6a98d25f4d485a94d
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

初期化されていない変数のデフォルト値は `undefined` です。これはまだ定義されていない値を表す特別なデータ型です。

初期化されていない変数にも値を代入することができます。以下はその例です。

```js
let uninitialized;
uninitialized = "assigned";
```

宣言の下にある `"Test"` 変数に文字列 `secondCharacter` を代入してください。コンソールを開いてログがどのように変わったか確認しましょう。

# --hints--

`secondCharacter` を初期化してはいけません。初期化とは変数を宣言するときに値を代入することを意味します。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

`secondCharacter` に代入演算子を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

`"Test"` 変数に文字列 `secondCharacter` を代入してください。

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
