---
id: 6610c87eac0f0b256d7b037e
title: ステップ 62
challengeType: 1
dashedName: step-62
---

# --description--

関数がもうパラメータを使っていなかったので、引数を変えても影響はありませんでした。

`test` 関数から `return` 宣言と `padRow` 文を削除して、関数を再び空にしてください。

# --hints--

`padRow` 関数には `test` 変数があってはいけません。

```js
assert.notMatch(padRow.toString(), /test/);
```

`padRow` 関数は値を返してはいけません。

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
