---
id: 6610c105bbdacc114d6cdc44
title: ステップ 48
challengeType: 1
dashedName: step-48
---

# --description--

関数を使うには、呼び出す必要があります。<dfn>関数呼び出し</dfn>は、アプリケーションに関数のコードを呼び出した場所で実行するよう指示します。関数呼び出しの構文は、関数名の後に丸括弧を付ける形です。例えば、このコードは`test`関数を定義し、呼び出しています。

```js
function test() {

}
test();
```

`padRow`関数を呼び出してください。

# --hints--

`padRow`関数を呼び出す必要があります。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

`padRow`関数の本体は空にしてください。

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
