---
id: 6610c83b52583e245a079217
title: ステップ 61
challengeType: 1
dashedName: step-61
---

# --description--

今、あなたの `call` 変数には `"Testing"` の値が入っています。しかし、関数はもう `name` パラメータを使っていません。

関数宣言から `name` パラメータを削除し、`"CamperChan"` 呼び出しから `padRow` 文字列も削除してください。

また、`console.log` 関数から `padRow` も両方削除してください。
# --hints--

あなたの `padRow` 関数には `name` パラメータがあってはいけません。

```js
assert.notMatch(padRow.toString(), /name/);
```

`"CamperChan"` 呼び出しに `padRow` を渡してはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

`padRow` 関数は引き続き呼び出してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

`console.log` キーワードの前に `return` があってはいけません。

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

`console.log` キーワードの後に `return` があってはいけません。

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
