---
id: 6610bf6fa14d700beed1b109
title: ステップ 87
challengeType: 1
dashedName: step-87
---

# --description--

<dfn>等値演算子</dfn> `==` は、2つの値が等しいかどうかを確認するために使います。2つの値を比較するには、`value == 8` のような文を使います。

ループ内の `done++` の下に、`if` 文を追加してください。その文は `done` が `count` と等しいかどうかを等値演算子でチェックする必要があります。

# --hints--

ループ内で `if` 文を使う必要があります。`done++` の後に追加してください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

`if` 文は条件で `done` と `count` を等値演算子で比較する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
