---
id: 6610c2d8d67563174fcf96dc
title: ステップ 52
challengeType: 1
dashedName: step-52
---

# --description--

コード内に明示的に書かれている値、例えば関数内の`"Hello!"`文字列は、<dfn>ハードコーディング</dfn>されたものとみなされます。関数内に値をハードコーディングすると、関数の再利用性が低くなる場合があります。

代わりに、関数に<dfn>パラメータ</dfn>を定義できます。パラメータは、関数を呼び出すときに値が与えられる特別な変数で、関数内で使うことで関数のコードの結果を動的に変えることができます。

関数にパラメータを追加するには、丸括弧の中に変数名を追加します。例えば、この`demo`関数には`name`パラメータがあります。

```js
function demo(name) {

}
```

`name`は便利なパラメータに思えるので、`padRow`関数に追加してください。

# --hints--

あなたの`padRow`関数には`name`パラメータが必要です。

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
