---
id: 6610c538372aa61cc0f5b122
title: ステップ 57
challengeType: 1
dashedName: step-57
---

# --description--

JavaScriptの変数は特定の<dfn>スコープ</dfn>で利用可能です。つまり、変数がどこで宣言されるかによって、その変数がコードのどこで使えるかが決まります。

最初のスコープはグローバルスコープです。関数や`for`ループのような「ブロック」の外で宣言された変数は<dfn>グローバルスコープ</dfn>にあります。あなたの`character`、`count`、および`rows`変数はすべてグローバルスコープにあります。

変数がグローバルスコープにある場合、関数はその定義内でアクセスできます。以下はグローバル`title`変数を使う関数の例です。

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

この例は`"Professor Naomi"`を返します。あなたの`padRow`関数を更新して、`character`パラメータの先頭に`name`変数を連結した値を返すようにしてください。

# --hints--

あなたの`padRow`関数は`character`の先頭に`name`を連結する必要があります。

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

あなたの`padRow`関数は`character + name`の結果を返す必要があります。

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
