---
id: 6610c6541c82551f95e765ab
title: ステップ 58
challengeType: 1
dashedName: step-58
---

# --description--

変数は関数の中でも宣言できます。これらの変数は<dfn>ローカルスコープ</dfn>または<dfn>ブロックスコープ</dfn>にあるとみなされます。関数内で宣言された変数は、その関数の中でしか使えません。関数の外でアクセスしようとすると、参照エラーが発生します。

これを実際に確認するには、`const`を使って`test`関数内に`padRow`変数を宣言してください。値は`"Testing"`で初期化します。

次に、関数の下で`test`をコンソールにログ出力しようとしてください。関数のローカルスコープの外では定義されていないため、エラーが表示されます。その`console.log`を削除すればテストに合格し、先に進めます。

# --hints--

関数内で`test`変数を宣言してください。

```js
assert.match(padRow.toString(), /var\s+test/);
```

`test`を`"Testing"`の値で初期化してください。セミコロンを忘れないでください。

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

`test`変数は`return`キーワードの前に置く必要があります。

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
