---
id: 6610c71600966a2191d3a64a
title: ステップ 59
challengeType: 1
dashedName: step-59
---

# --description--

関数から返された値は、その関数を呼び出すことで使われます。関数呼び出しを返す値として直接使うことも、返された値を変数に代入して使うこともできます。こうすることで、関数内で作成されたローカルスコープの変数に代入された値を、関数の外で使うことができます。

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

`"Testing"` の値を使うには、`padRow` 関数からその値を返すようにして、`return` 文を更新して `test` 変数だけを返すようにしてください。


# --hints--

`padRow` 関数は `test` 変数を返す必要があります。

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
