---
id: 6610c16c4fa0df12c0e30675
title: ステップ 49
challengeType: 1
dashedName: step-49
---

# --description--

`padRow` 関数を呼び出していますが、その関数呼び出しの結果を何もしていません。JavaScript のすべての関数は値を<dfn>返します</dfn>。つまり、呼び出した結果として定義された値を返し、それを他の場所で使うことができます。

`padRow` 関数を呼び出した結果を確認するには、`call` 変数を宣言し、既存の `padRow` 呼び出しをその変数に代入してください。

# --hints--

`call` 変数を宣言する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

`const` を使って `call` 変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

`call` に `padRow` 呼び出しの結果を代入してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
