---
id: 6610c424b7119919b62932f4
title: ステップ 53
challengeType: 1
dashedName: step-53
---

# --description--

関数はハードコードされた値を返す必要はありません。変数に格納された値を返すこともできます。パラメータは関数の特別な変数なので、それらも返すことができます。

`padRow` 関数を変更して、`return` パラメータを直接 `name` してください。

# --hints--

`padRow` 関数は `name` パラメータの値を返すべきです。

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
