---
id: 6610b9f7619764fad5fd516d
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

変数の値を別の変数に代入することもできます。例えば：

```js
let first = "One";
let second = "Two";
second = first;
```

`second` 変数は今、`"One"` の値を持つことになります。

これを実際に確認するには、`secondCharacter` の代入を `"Test"` から `character` 変数に変更してください。

その後、コンソールを開いて何がログに表示されるか見てみましょう。

# --hints--

`"Test"` の値を `secondCharacter` 変数に代入してはいけません。

```js
assert.notEqual(secondCharacter, "Test");
```

`character` 変数の値を `secondCharacter` 変数に代入してください。セミコロンを忘れないでください。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

`secondCharacter` 変数は今、`"World"` の値を持つことになります。

```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
