---
id: 6610bbed59bc2a0194d85533
title: ステップ 22
challengeType: 1
dashedName: step-22
---

# --description--

`rows` 配列の中の値が直接変更されたことに気づきましたか？これは <dfn>ミューテーション</dfn>と呼ばれます。配列についてもっと学ぶと、いつ配列をミューテーションすべきか、いつすべきでないかがわかるようになります。

次に進む前に、配列のよくある使い方を学ぶ絶好の機会です。現在のコードでは `rows[2]` を使って配列の最後の要素にアクセスしています。しかし、最後の要素を取得したいときに配列の要素数がわからない場合があります。

配列の `.length` プロパティを利用できます。これは配列の要素数を返します。どんな配列でも最後の要素を取得するには、次の構文を使います。

```js
array[array.length - 1]
```

`array.length` は配列の要素数を返します。`1` を引くことで、配列の最後の要素のインデックスが得られます。この考え方を `rows` 配列にも適用できます。

`rows[2]` を更新して、`rows` 配列の最後の要素に動的にアクセスするようにしてください。上の例を参考にしてください。

コンソールには何も変化がないはずです。

# --hints--

`.length` 配列の `rows` プロパティを使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

ブラケット記法には `rows.length` を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

ブラケット記法の長さから `1` を引くべきです。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
