---
id: 6610c8cfe4cf4d278e35c156
title: ステップ 63
challengeType: 1
dashedName: step-63
---

# --description--

予想どおり、関数は再び `undefined` を返します。`call` 変数はもう必要ないので、`call` の宣言と `console.log` 変数の `call` を削除してください。

# --hints--

`call` の宣言はしてはいけません。

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

`call` 変数をログに出力してはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
