---
id: 6610c77d50636722e5b6be17
title: ステップ 60
challengeType: 1
dashedName: step-60
---

# --description--

`return` 文の下に、文字列 `"This works!"` をコンソールにログ出力してください。

それを行うと、文字列 `"This works!"` はコンソールに表示されず、`console.log("This works!")` の行はグレーアウトされているのがわかります。

コンソールログをコピーして、`return` 文の上に貼り付けてください。これで文字列 `"This works!"` がコンソールに表示されるはずです。

`return` キーワードについて知っておくべき重要なことは、関数から返す値を定義するだけでなく、関数やブロック文内のコードの実行を停止することです。つまり、`return` 文の後のコードは実行されません。

# --hints--

`padRow` 関数は `test` 変数を返す必要があります。

```js
assert.equal(padRow("Naomi"), "Testing");
```

最初の `console.log` は `return` キーワードの後に来るべきです。

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

2つ目の `console.log` は `return` キーワードの前に来るべきです。

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
