---
id: 6610b9f7619764fad5fd516d
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

ある変数の値を別の変数に代入することも可能です 例えば、次のように記述します。

```js
let first = "One";
let second = "Two";
second = first;
```

これで、変数 `second` に値 `"One"` が代入された状態になります。

では、`secondCharacter` に `"Test"` を代入しているコードを変更して、変数 `character` の値を代入するようにしてください。

その後、コンソールを開いてログを確認してください。

# --hints--

変数 `secondCharacter` に `"Test"` という値を代入しないでください。

```js
assert.notEqual(secondCharacter, "Test");
```

変数 `secondCharacter` に、変数 `character` の値を代入してください。 セミコロンを忘れないようにしましょう。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

これで、変数 `secondCharacter` に値 `"World"` が代入された状態になります。


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
