---
id: 6610c21b3ef82015573ffbbe
title: ステップ 51
challengeType: 1
dashedName: step-51
---

# --description--

あなたの `call` 変数は `undefined` の値を持っていますが、定義したはずです！これは `padRow` 関数が現在値を返していないためです。デフォルトでは、関数は `undefined` を値として返します。

別の値を返すには、`return` キーワードを使う必要があります。以下は文字列 `"Functions are cool!"` を返す関数の例です。

```js
function demo() {
  return "Functions are cool!";
}
```

`return` キーワードを使って、関数が文字列 `"Hello!"` を返すようにしてください。

# --hints--

あなたの `padRow` 関数は `return` キーワードを使うべきです。

```js
assert.match(padRow.toString(), /return/);
```

あなたの `padRow` 関数は文字列 `"Hello!"` を返すべきです。

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
