---
id: 664599653fcd6e97104f9261
title: ステップ 83
challengeType: 1
dashedName: step-83
---

# --description--

すべての `if...else if` 条件が `false` のときに異なるコードを実行したい場合があります。これは `else` ブロックを追加することで実現できます。

`else` ブロックは、`if` および `else if` ブロックの条件が満たされない場合にのみ評価されます。

ここでは `else` ブロックが `else if` ブロックに追加されています。

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

`else` ブロックを `else if` ブロックに追加してください。`else` ブロックの中で、文字列 `"This is the else block"` をコンソールにログ出力してください。

コンソールで結果を確認するには、`<` 文の中の `else if` を手動で `>` に変更できます。そうすると条件が `false` となり、`else` ブロックが実行されます。

# --hints--

`else` ブロックがあるはずです。

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

`else` ブロックは文字列 `"This is the else block"` をコンソールにログ出力する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
