---
id: 66458f0a05df478aa627629e
title: ステップ 82
challengeType: 1
dashedName: step-82
---

# --description--

テキストがまた消えてしまいました！空の文字列は`false`と評価され、<dfn>偽値</dfn>になります。真偽値については今後のプロジェクトでさらに学びます。

`if`文に加えて、JavaScriptには<dfn>else if</dfn>文もあります。`else if`文を使うと、1つのコードブロックで複数の条件をチェックできます。

`else if`文の構文は次の通りです：

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

最初の条件が`false`の場合、JavaScriptはチェーンの次の条件をチェックします。2番目の条件が`false`の場合、JavaScriptは3番目の条件をチェックし、以下同様です。

`if`文の下に、`else if`文を追加して`5`が`10`より小さいかどうかをチェックしてください。次に、`else if`文の本体内で文字列`"5 is less than 10"`をコンソールにログ出力してください。

結果を確認するにはコンソールを見てください。

# --hints--

`else if`文があるはずです。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

`else if`文は`5`が`10`より小さいかどうかをチェックしてください。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

文字列`"5 is less than 10"`をコンソールにログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
