---
id: 664ed9344ba3495e1f67f8ea
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

宣言と初期化をレビューしたので、レビュー用に書いたコードを削除してください。

`character` 変数は削除しないでください。

# --hints--

`profession` 変数は存在してはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

`age` 変数は存在してはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

`console.log` に対する `age` 文は存在してはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

`console.log` に対する `profession` 文は存在してはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

`character` 変数はまだ存在しているはずです。

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
