---
id: 664ed97c55a99e5ffab759dc
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

先に進む前に、これまでに学んだ概念を少しレビューしましょう。

`let` キーワードを使って `profession` 変数と `age` 変数を宣言してください。`profession` には文字列 `"teacher"` を初期化しますが、`age` には値を初期化しないでください。

両方の変数をコンソールにログ出力して結果を確認しましょう。

# --hints--

`profession` 変数を宣言する必要があります。

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

`let` を使って `profession` 変数を宣言してください。

```js
assert.match(code, /let\s+profession/);
```

`"teacher"` という文字列を `profession` 変数に代入してください。

```js
assert.strictEqual(profession, "teacher");
```

`age` 変数を宣言する必要があります。

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

`let` を使って `age` 変数を宣言してください。

```js
assert.match(code, /let\s+age/);
```

`age` 変数には値を代入しないでください。

```js
assert.isUndefined(age);
```

`age` をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

`profession` をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
