---
id: 665732da4815b70bb083915e
title: ステップ 24
challengeType: 1
dashedName: step-24
---


# --description--

ここまでのステップで、配列の操作について学びました。学んだことを振り返ってみましょう。

まず、`cities` 変数を宣言し、文字列 `"London"`、`"New York"`、`"Mumbai"` の配列として初期化してください。その後、その変数をコンソールに出力してください。

出力した後、`cities` の最後の要素を文字列 `"Mexico City"` に変更し、再度 `cities` 変数をコンソールに出力してください。

正しくできていれば、コンソールにこのような出力が表示されるはずです。

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

`let` を使って `cities` 変数を宣言してください。

```js
assert.match(code, /let\s+cities/);
```

文字列 `"London"`、`"New York"`、`"Mumbai"` の配列を `cities` 変数に代入してください。

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

`console.log()` を使って `cities` 配列全体をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

`cities` 配列の最後の要素を文字列 `"Mexico City"` に更新してください。配列の最後の要素には `array[array.length - 1]` を使ってアクセスできることを覚えておいてください。

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

コード内に `console.log(cities)` の文が2つあるはずです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
