---
id: acda2fb1324d9b0fa741e6b5
title: 建置一個 Confirm the Ending Tool
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

在本實驗中，你將實作一個函式，用來檢查字串是否以指定的標的字串結尾。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `confirmEnding` 的函式（程式），該函式接受兩個參數：要檢查的字串和用來比對的字串。
2. 如果第一個字串以第二個字串結尾，該函式（程式）應傳回 `true`，否則傳回 `false`。
3. 你不應該使用 `.endsWith()` 方法；取而代之，請使用其中一個 JavaScript 的子字串方法來達成此目的。

# --hints--

你應該創建一個名為 `confirmEnding` 的函式（程式）。

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` 應該接受 2 個參數。

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` 應該傳回 `false`。

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` 應該傳回 `false`。

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` 應該傳回 `false`。

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` 應該傳回 `false`。

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` 應該傳回 `false`。

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

你的程式碼不應該使用內建方法 `.endsWith()` 來解決此實驗。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
