---
id: 655cd899f8de09431eabb40c
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

El algoritmo de ordenamiento por mezcla es un algoritmo que utiliza el principio de dividir y conquistar para ordenar colecciones de datos. Es decir, "divide" una colección en sub-partes más pequeñas y "conquista" las sub-partes ordenándolas de forma independiente, luego combina las sub-partes ordenadas.

Comienza creando una función llamada `merge_sort` que toma el parámetro `array`. Esta función manejará la tarea de ordenar una lista de números. Usa la palabra clave `pass` en el cuerpo de la función.

# --hints--

Debes declarar una función llamada `merge_sort` con un único parámetro: `array`. No olvides usar la palabra clave `pass` en el cuerpo de la función.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
