---
id: 656627b47bd2d2a4afbc945d
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

En el paso anterior, obtuviste el punto medio. Puedes usarlo para dividir `array` en dos y asignar cada parte a nuevas variables. 

Debes usar la sintaxis de segmento para extraer la mitad izquierda del `array` y asignarla a una variable llamada `left_part`.

# --hints--

Debes tener una variable llamada `left_part` en tu función `merge_sort`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

Debes usar la sintaxis de segmento para extraer la mitad izquierda del array y asignarla a la variable `left_part`.

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
