---
id: 656710d1e0ec62253426db24
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

El `while` que creaste compara un elemento de `left_part` con otro en `right_part`, y luego agrega el elemento menor a la lista principal `array`.

Continuará esta operación hasta que no queden elementos por comparar. Pero `left_part` aún puede tener elementos restantes mientras que `right_part` no tiene ninguno, y viceversa.

Crea otro bucle `while` que se ejecute cuando `left_array_index` sea menor que `len(left_part)`. En los siguientes pasos, lo usarás para copiar los elementos restantes de `left_part` en la lista `array`.

# --hints--

Tu bucle `while` debe tener esta condición: `left_array_index < len(left_part)`. No olvides usar la palabra clave `pass` en el cuerpo del bucle.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

Debe tener la palabra clave `pass` en el cuerpo de tu bucle `while`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
