---
id: 656706afd65547d22bee0b68
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Dentro del bucle `while`, reemplace `pass` con una declaraci칩n `if` que verifique si el elemento en `left_part` es menor que el elemento en `right_part`.

Utilice la palabra clave `pass` en el cuerpo de la declaraci칩n `if`.

# --hints--

Debes reemplazar la palabra clave `pass` con una estructura `if` con esta condici칩n: `left_part[left_array_index] < right_part[right_array_index]`. A침ade la palabra clave `pass` al cuerpo de la estructura `if`.

```js
({
 test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_body().is_equivalent("if left_part[left_array_index] < right_part[right_array_index]:\\n  pass")`))
 })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        pass
--fcc-editable-region--
```
