---
id: 656702f8b4cbd8cbf0a433c6
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Luego, necesitas crear un bucle que continúe mientras queden elementos en ambos `left_part` y `right_part`.

Para eso, crea un bucle `while` con dos condiciones: una que verifica si el `left_array_index` es menor que la longitud de `left_part` y otra condición que verifica si `right_array_index` es menor que la longitud de `right_part`.

# --hints--

Tu bucle `while` debe verificar si `left_array_index` es menor que la longitud de `left_part` y si `right_array_index` es menor que la longitud de `right_part`.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

Usa la palabra clave `pass` en el cuerpo del bucle.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
