---
id: 68da6138d70960c197ec5ef1
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Ora devi identificare quando hai costruito una combinazione completa di parentesi. Una combinazione completa ha una lunghezza uguale al doppio del numero di coppie (dato che ogni coppia contribuisce con una parentesi aperta e una chiusa).

All’interno del tuo loop `while`, aggiungi un’istruzione `if` per verificare se `len(current) == 2 * pairs`. Quando questa condizione è vera, aggiungi `current` alla lista `result`.

# --hints--

Dovresti avere un’istruzione `if` all’interno del tuo loop `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

La tua istruzione `if` dovrebbe verificare se `len(current) == 2 * pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

Dovresti aggiungere `current` alla lista `result` all’interno della tua nuova istruzione `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
--fcc-editable-region--    
    return result

print(gen_parentheses(1))
```
