---
id: 68da6138d70960c197ec5eee
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora imposta la struttura dati per archiviare i tuoi risultati. Crea una variabile chiamata `result` e inizializzala come una lista vuota. Questa lista conterrà tutte le combinazioni valide di parentesi che genererai.

Aggiorna la tua istruzione return per restituire `result` invece di una lista vuota.

# --hints--

Dovresti dichiarare una variabile chiamata `result` nella tua funzione `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

Dovresti inizializzare `result` come una lista vuota.

```js
({ test: () => runPython(`
_result = _Node(_code).find_function("gen_parentheses").find_variable("result")
assert _result.is_equivalent("result = []") or _result.is_equivalent("result = list()")
`) })
```

Dovresti restituire `result` alla fine della tua funzione.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    return []
--fcc-editable-region--
```
