---
id: 68da6138d70960c197ec5ef0
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora implementa il loop principale di BFS. Crea un loop `while` che continua finché `queue` non è vuoto (mentre `queue` valuta a `True`).

All'interno del loop, stampa `queue`.

# --hints--

Dovresti avere un loop `while` nella tua funzione `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

Il tuo loop `while` dovrebbe avere la condizione `queue`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

Dovresti stampare `queue` all'interno del tuo loop `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
