---
id: 68da6138d70960c197ec5ef3
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Ora devi gestire l'aggiunta delle parentesi tonde di chiusura. La regola fondamentale è che puoi aggiungere una parentesi di chiusura solo se mantiene l'equilibrio, cioè deve esserci un numero maggiore di parentesi tonde di apertura usate rispetto a quelle di chiusura.

Sempre all'interno del blocco `else`, aggiungi una seconda istruzione `if` per verificare se `closes_used < opens_used`.

Se questa condizione è vera, aggiungi un'altra nuova tupla alla `queue`: `(current + ')', opens_used, closes_used + 1)`. Questo rappresenta lo stato dopo aver aggiunto una parentesi di chiusura.

# --hints--

Dovresti avere una seconda istruzione `if` all'interno del tuo blocco `else`.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

La tua seconda istruzione annidata `if` dovrebbe verificare se `closes_used < opens_used`.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

Dovresti aggiungere `(current + ')', opens_used, closes_used + 1)` alla coda quando la condizione è vera.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
