---
id: 68da6138d70960c197ec5eed
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Ora devi verificare che il numero di coppie sia almeno uno, dato che non puoi generare combinazioni di parentesi con zero o un numero negativo di coppie.

Aggiungi un'altra istruzione `if` per controllare se `pairs` è minore di `1`. Se questa condizione è vera, restituisci la stringa `The number of pairs should be at least 1`.

# --hints--

Dovresti avere una seconda istruzione `if` nella tua funzione `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

La tua seconda istruzione `if` dovrebbe verificare se `pairs` è minore di `1`.

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

Dovresti restituire la stringa `The number of pairs should be at least 1` quando `pairs` è minore di `1`.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    
    return []
--fcc-editable-region--
```
