---
id: 68da6138d70960c197ec5eef
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Per l'approccio di ricerca in ampiezza, userai una coda per tenere traccia dei diversi stati mentre costruisci le combinazioni di parentesi tonde. Ogni stato sarà rappresentato come una tupla contenente tre elementi:

- La stringa corrente che si sta costruendo
- Il numero di parentesi tonde aperte usate finora
- Il numero di parentesi tonde chiuse usate finora

Crea una variabile chiamata `queue` e inizializzala con una lista contenente una tupla: `('', 0, 0)`. Questo rappresenta lo stato iniziale con una stringa vuota e zero parentesi usate.

# --hints--

Dovresti dichiarare una variabile chiamata `queue` nella tua funzione `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

Dovresti inizializzare `queue` con una lista contenente la tupla `('', 0, 0)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    result = []
    return result
--fcc-editable-region--
```
