---
id: 68da604b8b6012a33e2344fb
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, implementerai una funzione che genera tutte le combinazioni valide di parentesi usando un approccio di ricerca in ampiezza (BFS). Per esempio, le combinazioni valide di due coppie di parentesi sono `(())` e `()()`.

Inizia creando una funzione chiamata `gen_parentheses` con un singolo parametro `pairs`. Per ora, fai in modo che la funzione restituisca una lista vuota.

# --hints--

Dovresti definire una funzione chiamata `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

La tua funzione `gen_parentheses` dovrebbe avere un singolo parametro chiamato `pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

La tua funzione dovrebbe restituire una lista vuota.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
