---
id: 68da6138d70960c197ec5eec
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Prima di implementare l'algoritmo principale, devi validare l'input. Il parametro `pairs` dovrebbe essere un intero, poiché rappresenta il numero di coppie di parentesi tonde da generare.

Aggiungi un'istruzione `if` all'inizio della tua funzione per verificare se `pairs` non è un intero. Usa la funzione `isinstance()` per questo controllo.

Se la condizione è vera, restituisci la stringa `The number of pairs should be an integer`.

# --hints--

Dovresti avere un'istruzione `if` nella tua funzione `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

La tua istruzione `if` dovrebbe verificare se `pairs` non è un intero usando `not isinstance(pairs, int)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

Dovresti restituire la stringa `The number of pairs should be an integer` quando l'input non è un intero.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    
    return []
--fcc-editable-region--
```
