---
id: 68da6138d70960c197ec5ef2
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Se la stringa corrente non è ancora completa, devi esplorare i prossimi stati possibili. Aggiungi una clausola `else` alla tua istruzione `if`.

All'interno del blocco `else` gestirai l'aggiunta delle parentesi tonde di apertura. Aggiungi un'istruzione `if` per verificare se `opens_used < pairs`. Questo assicura che non usi più parentesi tonde di apertura di quelle consentite.

Se questa condizione è vera, aggiungi una nuova tupla alla `queue`: `(current + '(', opens_used + 1, closes_used)`. Questo rappresenta lo stato dopo aver aggiunto una parentesi tonda di apertura.

# --hints--

Dovresti avere una clausola `else` per la tua istruzione `if` all'interno del loop `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

Dovresti avere un'istruzione `if` all'interno del blocco `else` che verifica se `opens_used < pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

Dovresti aggiungere `(current + '(', opens_used + 1, closes_used)` alla coda quando la condizione è vera.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
