---
id: 68e00c5349cbddfef6dbc053
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

All’interno del tuo loop `while`, usa `queue.pop(0)` per rimuovere e ottenere il primo elemento dalla coda. Questo implementa il comportamento first-in-first-out (FIFO) caratteristico del BFS.

Decomprimi questa tupla in tre variabili: `current`, `opens_used` e `closes_used`.

# --hints--

Dovresti usare `queue.pop(0)` per ottenere il primo elemento dalla coda.

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

Dovresti decomprimere il risultato in tre variabili: `current`, `opens_used` e `closes_used`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        
--fcc-editable-region--
    return result
```
