---
id: 689b17e7e63b8e2c13839e98
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Ahora es momento de implementar la funcionalidad real de conteo. 

En lugar de solo registrar el conteo actual, necesitas incrementarlo y actualizar la pantalla.

Elimina el `console.log`.

Verifica si el primer n칰mero actual dentro de `.count` ya es 10 (el texto es `10/10`), y si es as칤, no hagas nada.

Si el n칰mero es menor que 10, incrementa el conteo actual en 1 y actualiza el texto del elemento `.count`. Por ejemplo, si el conteo actual es 2, debe incrementarse a 3 y el texto del elemento `.count` debe ser `3/10`.

# --hints--

Cuando se hace clic en el bot칩n, el conteo debe incrementarse de 0 a 1.

```js
const spy = __helpers.spyOn(console, "log");
const countEl = document.querySelector("#happy-btn .count");
countEl.textContent = "0/10";
document.querySelector("#happy-btn").click();
assert.equal(countEl.textContent, "1/10");
```

Cuando se hace clic en el bot칩n varias veces, el conteo debe incrementarse correctamente.

```js
const countEl = document.querySelector("#happy-btn .count");
countEl.textContent = "3/10";
document.querySelector("#happy-btn").click();
assert.equal(countEl.textContent, "4/10");
```

Cuando el conteo ya est치 en 10, hacer clic no deber칤a incrementarlo m치s.

```js
const countEl = document.querySelector("#happy-btn .count");
countEl.textContent = "10/10";
document.querySelector("#happy-btn").click();
assert.equal(countEl.textContent, "10/10");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>

    <h1 class="title">How are you feeling today?</h1>

    <p class="description">
      Click on the buttons below to rate your emotions.
    </p>
    <div class="btn-container">
      <button id="happy-btn" class="emoji-btn" aria-label="Happy face emoji">
        <span role="img" aria-hidden="true">游땕</span>
        <span class="count">0/10</span>
      </button>
    </div>
  </main>
  <script src="./script.js"></script>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```

```js
--fcc-editable-region--
const happyBtn = document.querySelector("#happy-btn");

happyBtn.addEventListener("click", () => {
  const countEl = happyBtn.querySelector(".count");
  const currCount = +countEl.textContent.split("/")[0];
  console.log("Current count:", currCount);
})


--fcc-editable-region--
```
