---
id: 68a9fa161498c3324ac06fe1
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Es hora de empezar a trabajar en los otros botones. Los elementos `button` restantes ya han sido a침adidos para ti en el archivo HTML.

Para cada nuevo bot칩n (`#confused-btn`, `#sad-btn` y `#loving-btn`), selecci칩nalo usando `querySelector` con su ID, luego agrega un listener de evento click que llame a `updateCount` con ese bot칩n.

Usa los nombres de variable `confusedBtn`, `sadBtn` y `lovingBtn`.

# --hints--

Debes crear una variable llamada `confusedBtn`.

```js
assert.exists(confusedBtn);
```

Debes almacenar el elemento `#confused-btn` en la variable `confusedBtn`.

```js
assert.equal(confusedBtn, document.querySelector('#confused-btn'));
```

Debes crear una variable llamada `sadBtn`.

```js
assert.exists(sadBtn);
```

Debes almacenar el elemento `#sad-btn` en la variable `sadBtn`.

```js
assert.equal(sadBtn, document.querySelector('#sad-btn'));
```

Debes crear una variable llamada `lovingBtn`.

```js
assert.exists(lovingBtn);
```

Debes almacenar el elemento `#loving-btn` en la variable `lovingBtn`.

```js
assert.equal(lovingBtn, document.querySelector('#loving-btn'));
```

El elemento `#confused-btn` debe tener un listener de evento click que llame a `updateCount`.

```js
const spy = __helpers.spyOn(window, "updateCount");
document.querySelector("#confused-btn").click();
assert.equal(spy.calls.length, 1);
assert.equal(spy.calls[0][0], document.querySelector("#confused-btn"));
```

El elemento `#sad-btn` debe tener un listener de evento click que llame a `updateCount`.

```js
const spy = __helpers.spyOn(window, "updateCount");
document.querySelector("#sad-btn").click();
assert.equal(spy.calls.length, 1);
assert.equal(spy.calls[0][0], document.querySelector("#sad-btn"));
```

El elemento `#loving-btn` debe tener un listener de evento click que llame a `updateCount`.

```js
const spy = __helpers.spyOn(window, "updateCount");
document.querySelector("#loving-btn").click();
assert.equal(spy.calls.length, 1);
assert.equal(spy.calls[0][0], document.querySelector("#loving-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>

    <h1 class="title">How are you feeling today?</h1>

    <p class="description">
      Click on the buttons below to rate your emotions.
    </p>
    <div class="btn-container">
      <button id="happy-btn" class="emoji-btn" aria-label="Happy face emoji">
        <span role="img" aria-hidden="true">游땕</span>
        <span class="count">0/10</span>
      </button>
      <button
        id="confused-btn"
        class="emoji-btn"
        aria-label="Confused face emoji"
      >
        <span role="img" aria-hidden="true">游땟</span>
        <span class="count">0/10</span>
      </button>
      <button id="sad-btn" class="emoji-btn" aria-label="Angry face emoji">
        <span role="img" aria-hidden="true">游</span>
        <span class="count">0/10</span>
      </button>
      <button
        id="loving-btn"
        class="emoji-btn"
        aria-label="Loving face emoji"
      >
        <span role="img" aria-hidden="true">游땘</span>
        <span class="count">0/10</span>
      </button>
    </div>
  </main>
  <script src="./script.js"></script>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```

```js
--fcc-editable-region--
function updateCount(btn) {
  const countEl = btn.querySelector(".count");
  let currCount = +countEl.textContent.split("/")[0];
  
  if (currCount === 10) return;
  
  currCount++;
  
  countEl.textContent = `${currCount}/10`;
}

const happyBtn = document.querySelector("#happy-btn");

happyBtn.addEventListener("click", () => updateCount(happyBtn));


--fcc-editable-region--
```
