---
id: 698367b9d1c6914a22095aec
title: Реалізувати алгоритм сортування бульбашкою
challengeType: 26
dashedName: implement-the-bubble-sort-algorithm
---

# --description--

У цій лабораторній роботі ви реалізуєте алгоритм сортування бульбашкою. Він починається з початку несортованого масиву і «випливає» несортовані значення до кінця, проходячи масив ітеративно, доки він повністю не відсортується. Це відбувається шляхом порівняння сусідніх елементів і їхнього обміну місцями, якщо вони розташовані не за порядком. Метод продовжує проходити масив у циклі, доки не відбудеться жодного обміну, після чого масив вважається відсортованим.

Цей метод вимагає кількох ітерацій через масив і для середніх та гірших випадків має квадратичну часову складність. Хоча він простий, зазвичай є непрактичним у більшості ситуацій.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Напишіть функцію `bubbleSort`, яка приймає масив цілих чисел як вхідні дані і повертає масив цих чисел, відсортований за зростанням.

# --hints--

У вас має бути функція з назвою `bubbleSort`.

```js
assert.isFunction(bubbleSort);
```

`bubbleSort` має повертати відсортований масив (від найменшого до найбільшого).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    bubbleSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`bubbleSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` має повертати масив, який не змінений, окрім порядку елементів.

```js
assert.sameMembers(
  bubbleSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`bubbleSort` не повинен використовувати вбудований метод `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    bubbleSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function bubbleSort(array) {
  for (let i = 0; i < array.length; i++) {
    let swapped = false;
    for (let j = 1; j < array.length; j++) {
      if (array[j - 1] > array[j]) {
        let temp = array[j-1];
        array[j-1] =  array[j];
        array[j] = temp;
        swapped = true;
      }
    }
    if (swapped === false) {
      break;
    }
  }
  return array;
}
```
