---
id: 587d774c367417b2b2512a9c
title: Agregar un texto alternativo a las imágenes para accesibilidad de usuarios con dificultades de la vista
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

Probablemente hayas visto un atributo `alt` en una etiqueta `img` en otros desafíos. El texto `alt` describe el contenido de la imagen y proporciona una alternativa en texto para ella. Un atributo `alt` ayuda en casos donde la imagen no se carga o no puede ser vista por un usuario. Los motores de búsqueda también lo usan para entender qué contiene una imagen y así incluirla en los resultados de búsqueda. Aquí tienes un ejemplo:

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

Las personas con dificultades visuales dependen de lectores de pantalla para convertir el contenido web a una interfaz de audio. Por esta razón, no podrán recibir la información si solo se les presenta de manera visual. En el caso de las imágenes, los lectores de pantalla pueden acceder el atributo `alt` y leer su contenido para proporcionar información clave.

Un buen texto `alt` le brinda al lector una breve descripción de la imagen. Siempre deberías incluir el atributo `alt` en tus imágenes. Además, según la especificación de HTML5, esto ahora se considera obligatorio.

# --instructions--

Camper Cat es tanto un ninja de la programación como un ninja real, que está construyendo un sitio web para compartir su conocimiento. La foto de perfil que quiere usar muestra sus habilidades y debería ser apreciada por todos los visitantes del sitio. Agrega un atributo `alt` en la etiqueta `img` que explique que Camper Cat está haciendo karate. (El `src` de la imagen no enlaza a un archivo real, así que deberías ver el texto `alt` en la pantalla.)

# --hints--

Tu etiqueta `img` debe tener un atributo `alt`, y el mismo no debería estar vacío.

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">
```
