---
id: 66ba762af611169359d9d369
title: Costruisci un articolo di giornale
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti impostare la radice `font-size` del tuo documento HTML su `24px`.
1. Dovresti avere un elemento con una classe `newspaper` che contiene tutti gli altri elementi.
1. Il tuo elemento `.newspaper` dovrebbe avere una `font-size` di `16px` e un carattere tipografico `Open Sans` con un carattere di fallback `sans-serif`.
1. All’interno del tuo elemento `.newspaper` dovresti avere almeno sette elementi in più: uno per il nome del giornale con una classe `name`, uno per la data dell’articolo con una classe `date`, uno per il titolo con una classe `headline`, uno per il sottotitolo con una classe `sub-headline`, uno per l’autore con una classe `author` e due paragrafi ciascuno con una classe `text`. Tutti questi elementi devono essere riempiti con le informazioni del tuo articolo.
1. Il tuo elemento `.name` dovrebbe avere una `font-size` che è il doppio della `font-size` dell’elemento radice e dovrebbe usare il carattere tipografico `Times New Roman`, con un carattere di fallback `serif`.
1. I tuoi elementi `.name` e `.author` dovrebbero usare CSS per rendere tutti i loro caratteri maiuscoli.
1. Il tuo elemento `.headline` dovrebbe avere una `font-size` che è il doppio della `font-size` del suo elemento genitore e dovrebbe essere in grassetto.
1. Il tuo elemento `.sub-headline` dovrebbe avere una `font-weight` di `100`, una `font-size` che è `1.5` volte la `font-size` del suo elemento genitore e dovrebbe essere in corsivo.
1. Il tuo `.author` dovrebbe usare CSS per renderlo in grassetto.
1. I tuoi elementi `.text` dovrebbero avere una `text-indent` di `20px`.
1. I tuoi elementi `.text` dovrebbero avere una `line-height` che è il doppio della `font-size` del loro elemento genitore.
1. La prima lettera dei tuoi elementi `.text` dovrebbe essere in grassetto e il doppio della dimensione della `font-size` del loro elemento genitore. Usa il selettore `::first-letter` per questo.

# --hints--

Dovresti avere un selettore `html` che imposta la sua `font-size` su `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

Dovresti avere un elemento con una classe `newspaper`.

```js
assert.exists(document.querySelector('.newspaper'));
```

Dovresti avere un selettore `.newspaper` che imposta la `font-size` dei suoi elementi su `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

Dovresti avere un selettore `.newspaper` che imposta la `font-family` dei suoi elementi su `'Open Sans', sans-serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

Dovresti avere un elemento con una classe `name` all’interno del tuo elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

Il tuo elemento `.name` dovrebbe avere il nome del tuo giornale.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

Dovresti avere un elemento con una classe `date` all’interno del tuo elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

Il tuo elemento `.date` dovrebbe avere la data.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

Dovresti avere un elemento con una classe `headline` all’interno del tuo elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

Il tuo elemento `.headline` dovrebbe avere il titolo del tuo articolo.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

Dovresti avere un elemento con una classe `sub-headline` all’interno del tuo elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

Il tuo elemento `.sub-headline` dovrebbe avere il sottotitolo del tuo articolo.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

Dovresti avere un elemento con una classe `author` all’interno del tuo elemento `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

Il tuo elemento `.author` dovrebbe avere l’autore del tuo articolo.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

Dovresti avere almeno due elementi paragrafo, ciascuno con una classe `text`, all’interno del tuo elemento `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

I tuoi elementi `.text` dovrebbero avere il testo del tuo articolo.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

Dovresti avere un selettore `.name` che imposta la `font-size` dei suoi elementi su `2rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

Dovresti avere un selettore `.name` che imposta la `font-family` dei suoi elementi su `'Times New Roman', serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

L’elemento `.name` dovrebbe avere la proprietà `text-transform` impostata su `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

L’elemento `.author` dovrebbe avere la proprietà `text-transform` impostata su `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Dovresti avere un selettore `.headline` che imposta la `font-size` dei suoi elementi su `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

Dovresti avere un selettore `.headline` che imposta la `font-weight` dei suoi elementi su `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

Dovresti avere un selettore `.sub-headline` che imposta la `font-weight` dei suoi elementi su `100`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

Dovresti avere un selettore `.sub-headline` che imposta la `font-size` dei suoi elementi su `1.5em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

Dovresti avere un selettore `.sub-headline` che imposta la `font-style` dei suoi elementi su `italic`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

Dovresti avere un selettore `.author` che imposta la `font-weight` dei suoi elementi su `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

Dovresti avere un selettore `.text` che imposta la `text-indent` dei suoi elementi su `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

Dovresti avere un selettore `.text` che imposta la `line-height` dei suoi elementi su `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

Dovresti avere un selettore `.text::first-letter` che imposta la `font-weight` dei suoi elementi su `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

Dovresti avere un selettore `.text::first-letter` che imposta la `font-size` dei suoi elementi su `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
