---
id: 66a84111965a0c46df6bbd0a
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

In the previous lecture videos, you learned how to associate a `label` element with an `input` like this:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

The `for` attribute on the `label` element should match the `id` attribute on the `input` element. This is known as an explicit association.

Below your `legend` element, add a `label` element with the text of `Name (required):`. For the `for` attribute, set it to the value of `"full-name"`.

Then below your `label` element, add an `input` element with no attributes. In the next steps, you will add the necessary attributes.

# --hints--

Ви повинні мати початковий тег `label`.

```js
assert.match(code, /<label.*>/);
```

Ви повинні мати кінцевий тег `label`.

```js
assert.match(code, /<\/label>/);
```

Your `label` element should be inside the `fieldset` element.

```js
assert.isNotNull(document.querySelector('fieldset label'));
```

Елемент `label` повинен мати атрибут `for` зі значенням `"full-name"`.

```js
const label = document.querySelector('fieldset label');
assert.strictEqual(label.getAttribute('for'), 'full-name');
```

Your `label` element should have the text of `Name (required):`.

```js
const label = document.querySelector('fieldset label');
assert.strictEqual(label.innerText, 'Name (required):');
```

Ви повинні мати елемент `input`.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
