---
id: 66a83601cd819e37f0dccd14
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Як вступний текст потрібно використати головний заголовок разом із коротким повідомленням про відгук.

Додайте елемент `header` до елемента `body`.

Додайте елемент `h1` до елемента `header`. Текстом елемента `h1` має бути `Hotel Feedback Form`.

Додайте елемент `p` під елементом `h1`. Текстом елемента `p` має бути `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

Ви повинні мати початковий тег `header`.

```js
assert.match(code, /<header>/i);
```

Ви повинні мати кінцевий тег `header`.

```js
assert.match(code, /<\/header>/i);
```

Ви повинні мати початковий тег `h1`.

```js
assert.match(code, /<h1>/i);
```

Ви повинні мати кінцевий тег `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Елемент `h1` повинен містити текст `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

Елемент `h1` має бути в межах елемента `header`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

Ви повинні мати початковий тег `p`.

```js
assert.match(code, /<p>/i);
```

Ви повинні мати кінцевий тег `p`.

```js
assert.match(code, /<\/p>/i);
```

Абзац повинен мати текст `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Абзац має бути в межах елемента `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

Абзац має бути після елемента `h1`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
