---
id: 66a83fec026a7a4631e084d2
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

When working with `fieldset` elements, it is common to use a caption to describe the group of inputs. You can use the <dfn>legend</dfn> element for this.

Here is an example of using a `legend` element:

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

Додайте елемент `legend` з текстом `Personal Information` до елемента `fieldset`.

# --hints--

Ви повинні мати початковий тег `legend`.

```js
assert.match(code, /<legend>/);
```

Ви повинні мати кінцевий тег `legend`.

```js
assert.match(code, /<\/legend>/);
```

Your `legend` element should be inside the `fieldset` element.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

Your `legend` element should contain the text `Personal Information`.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend.innerText, 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
