---
id: 66a83601cd819e37f0dccd14
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

For the introductory text, you will want to display the main title followed by a short note about leaving feedback.

Додайте елемент `header` до елемента `body`.

Додайте елемент `h1` до елемента `header`. Текстом елемента `h1` має бути `Hotel Feedback Form`.

Додайте елемент `p` під елементом `h1`. Текстом елемента `p` має бути `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

Ви повинні мати початковий тег `header`.

```js
assert.match(code, /<header>/i);
```

Ви повинні мати кінцевий тег `header`.

```js
assert.match(code, /<\/header>/i);
```

Ви повинні мати початковий тег `h1`.

```js
assert.match(code, /<h1>/i);
```

Ви повинні мати кінцевий тег `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Елемент `h1` повинен містити текст `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1.innerText, 'Hotel Feedback Form');
```

Your `h1` element should be inside the `header` element.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1.parentElement.tagName, 'HEADER');
```

Ви повинні мати початковий тег `p`.

```js
assert.match(code, /<p>/i);
```

Ви повинні мати кінцевий тег `p`.

```js
assert.match(code, /<\/p>/i);
```

Your paragraph element should have the text `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement.innerText, 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Your paragraph element should be inside your `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement.parentElement.tagName, 'HEADER');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
