---
id: 66a83bdcf425e7446900b7c4
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

У попередніх відеолекціях ви навчилися працювати з елементом `form`:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

Атрибут `action` використовують, щоб вказати, куди мають бути надіслані дані форми, коли її відправлено.

Атрибут `method` використовують, щоб вказати HTTP-метод, який потрібно використати під час надсилання даних форми. Найчастіше використовують методи `GET` та `POST`.

**ПРИМІТКА**: в наступних модулях ви дізнаєтеся, як працюють HTTP-методи.

До елемента `main` додайте елемент `form` з атрибутом `action` зі значенням `"https://hotel-feedback.freecodecamp.org"` та атрибутом `method` зі значенням `"POST"`.

# --hints--

Ви повинні мати початковий тег `form`.

```js
assert.match(code, /<form.*>/i);
```

Ви повинні мати кінцевий тег `form`.

```js
assert.match(code, /<\/form>/i);
```

Елемент `form` повинен мати атрибут `action` зі значенням `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Елемент `form` повинен мати атрибут `method` зі значенням `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```
