---
id: 66a8290a27c2c625e2355042
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn how to work with HTML forms by building a Hotel Feedback Form.

Start by adding the `<!DOCTYPE html>` followed by an `html` element with a `lang` attribute of `en`.

Додайте елемент `head` до елемента `html`.


# --hints--

Ви повинні мати `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

You should have an opening `html` tag with the language set to english.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Ви повинні мати кінцевий тег `html`.

```js
assert.match(code, /<\/html>/i);
```

`DOCTYPE` має бути перед елементом `html`.

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

Ви повинні мати початковий тег `head`.

```js
assert.match(code, /<head>/i);
```

Ви повинні мати кінцевий тег `head`.

```js
assert.match(code, /<\/head>/i);
```

Початковий тег `head` має бути перед кінцевим тегом `head`.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

Елемент `head` має бути в межах елемента `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
