---
id: 66a83bdcf425e7446900b7c4
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

У попередніх відео лекціях ви навчилися працювати з елементом `form` ось так:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

The `action` attribute is used to specify where the form data should be sent when the form is submitted.

The `method` attribute is used to specify the HTTP method to use when sending the form data. The most common methods are `GET` and `POST`.

**NOTE**: You will learn about how HTTP methods work in later modules.

Inside your `main` element, add a `form` element with an `action` attribute set to `"https://hotel-feedback.freecodecamp.org"` and a `method` attribute set to `"POST"`.

# --hints--

Ви повинні мати початковий тег `form`.

```js
assert.match(code, /<form.*>/i);
```

Ви повинні мати кінцевий тег `form`.

```js
assert.match(code, /<\/form>/i);
```

Елемент `form` повинен мати атрибут `action` зі значенням `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Елемент `form` повинен мати атрибут `method` зі значенням `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```
