---
id: 66a83e5e491625454b6f62c3
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Форма складається з елементів `inputs`, де користувачі можуть вводити дані. Пов’язані вхідні дані можна згрупувати за допомогою елемента `fieldset`.

Ось приклад використання елемента `fieldset`:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

Додайте елемент `fieldset` до елемента `form`.

# --hints--

Ви повинні мати початковий тег `fieldset`.

```js
assert.match(code, /<fieldset>/i);
```

Ви повинні мати кінцевий тег `fieldset`.

```js
assert.match(code, /<\/fieldset>/i);
```

Елемент `fieldset` має бути в межах елемента `form`.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">

      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
