---
id: 66a8347f97a33e36ffc81b9f
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Щоб підготуватись до створення вмісту, додайте елемент `body` під елементом `head`.

# --hints--

Ви повинні мати початковий тег `body`.

```js
assert.match(code, /<body>/i);
```

Ви повинні мати кінцевий тег `body`.

```js
assert.match(code, /<\/body>/i);
```

Не змінюйте елемент `head`. Переконайтеся, що не видалили кінцевий тег.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Елемент `body` повинен йти після елемента `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>

  --fcc-editable-region--
</html>
```
