---
id: 66a9419e2d18476c645ce693
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Числовий елемент `input` використовують, щоб створити поле для введення числа.

Ось приклад поля для введення числа:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

Атрибути `min` та `max` використовують, щоб встановити мінімальне та максимальне значення, які можна ввести в поле введення даних.

Під елементом `label` додайте елемент `input` з атрибутом `type` зі значенням `"number"` та `id` зі значенням `"age"`.

Значенням атрибута `name` має бути `"age"`, значенням атрибута `min` має бути `"3"`, а значенням атрибута `max` має бути `"100"`.

# --hints--

Ви повинні мати числовий елемент `input` під елементом `label`.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

Числовий елемент `input` повинен мати `id` зі значенням `age`.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

Числовий елемент `input` повинен мати атрибут `name` зі значенням `age`.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

Числовий елемент `input` повинен мати атрибут `min` зі значенням `"3"`.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

Числовий елемент `input` повинен мати атрибут `max` зі значенням `"100"`.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
      --fcc-editable-region--
          <label for="age">Age (optional):</label>

      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
