---
id: 66a93c95bc58e26a8fe95818
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Тепер додайте елемент `input` з атрибутом `type` зі значенням `"email"` під елементом `label`, який стосується електронної адреси. Цей елемент `input` повинен мати атрибут `id` зі значенням `"email"`. Також встановіть значення атрибута `name` на `"email"`.

Цей елемент `input` також обов’язковий, тому не забудьте додати атрибут `required`.

Вкінці додайте атрибут `placeholder` зі значенням `"example@email.com"`.

# --hints--

Ви повинні мати елемент `input` під елементом `label`, який стосується електронної адреси.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

Елемент `input` повинен мати атрибут `type` зі значенням `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

Елемент `input` повинен мати атрибут `id` зі значенням `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

Елемент `input` повинен мати атрибут `name` зі значенням `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

Елемент `input` повинен мати атрибут `required`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

Елемент `input` повинен мати атрибут `placeholder` зі значенням `"example@email.com"`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>
        --fcc-editable-region--
          <label for="email">Email address (required):</label>

        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
