---
id: 66a954b2bcddba72076c1857
title: Крок 20
challengeType: 0
dashedName: step-20
---

# --description--

Під першою радіокнопкою додайте ще один елемент `label` з атрибутом `for` зі значенням `"no-option"`. Текстом елемента `label` має бути `No`.

Під другим елементом `label` додайте радіокнопку з `id` зі значенням `"no-option"` та атрибутом `name` зі значенням `"hotel-stay"`.

Коли закінчите, спробуйте перевірити радіокнопки, обираючи лише одну опцію за раз.

# --hints--

Ви повинні мати елемент `label` під першою радіокнопкою.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

Елемент `label` повинен мати атрибут `for` зі значенням `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, "no-option");
```

Елемент `label` повинен мати текст `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label[for="no-option"]')?.textContent, 'No');
```

Ви повинні мати другий елемент `input` під другим елементом `label`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(2) label + input').length, 2);
```

Елемент `input` повинен мати `type` зі значенням `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.getAttribute('type'), 'radio');
```

Елемент `input` повинен мати `id` зі значенням `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.id, 'no-option');
```

Елемент `input` повинен мати `name` зі значенням `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.name, 'hotel-stay');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          --fcc-editable-region--
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
