---
id: 66a930b20f589b6664c51cb0
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Якщо користувач надає повне ім’я, то `input` приймає звичайний текст.

У попередніх відеолекціях ви навчилися працювати з атрибутом `type`:

```html
<input type="text">
```

До наявного елемента `input` додайте атрибут `type` зі значенням `"text"`.

Також надайте елементу `input` атрибут `id` зі значенням `"full-name"`.

# --hints--

Не видаляйте наявний елемент `input`.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Не створюйте новий елемент `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Елемент `input` повинен мати атрибут `type` зі значенням `"text"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

Елемент `input` повинен мати атрибут `id` зі значенням `"full-name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
