---
id: 66a9521bc70162712caf118d
title: Крок 19
challengeType: 0
dashedName: step-19
---

# --description--

Якщо ви хочете, щоб користувач вибрав лише одну опцію зі списку, можна використати групу радіокнопок.

Ось приклад двох радіокнопок:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

У цьому прикладі радіокнопки згруповані разом через однакове значення атрибута `name`. Це означає, що одночасно можна вибрати лише одну радіокнопку.

Під елементом `legend` додайте елемент `label` з текстом `Yes` і атрибутом `for` зі значенням `"yes-option"`.

Під елементом `label` додайте радіокнопку з `id` зі значенням `"yes-option"` і атрибутом `name` зі значенням `"hotel-stay"`.

# --hints--

Ви повинні мати елемент `label` під елементом `legend`.

```js
assert(document.querySelector('fieldset:nth-of-type(2) legend + label'));
```

Елемент `label` повинен мати атрибут `for` зі значенням `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + label')?.htmlFor, 'yes-option');
```

Елемент `label` повинен містити текст `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + label')?.innerText, 'Yes');
```

Ви повинні мати радіокнопку під елементом `label`.

```js
assert(document.querySelector('fieldset:nth-of-type(2) label + input[type="radio"]'));
```

Радіокнопка повинна мати атрибут `id` зі значенням `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input[type="radio"]')?.id, 'yes-option');
```

Радіокнопка повинна мати атрибут `name` зі значенням `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input[type="radio"]')?.name, 'hotel-stay');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        --fcc-editable-region--
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
