---
id: 66a962954f4e0d76223b37ed
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Тепер додайте ще один елемент `label` з текстом `Location` і атрибутом `for` зі значенням `"location"`.

Прапорець повинен мати атрибути `id`, `name` та `value` зі значенням `"location"`.

Під цим прапорцем додайте ще один елемент `label` з текстом `Reputation` і атрибутом `for` зі значенням `"reputation"`.

Прапорець повинен мати атрибути `id`, `name` та `value` зі значенням `"reputation"`.

# --hints--

Ви повинні мати елемент `label` з текстом `Location`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[2]?.textContent, 'Location');
```

Елемент `label` повинен мати атрибут `for` зі значенням `"location"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[2]?.getAttribute('for'), 'location');
```

Після елементу `label` має бути третій прапорець.

```js 
assert(document.querySelector("fieldset:nth-of-type(3) label + input:nth-of-type(3)[type='checkbox']"));
```

Третій прапорець повинен мати `id` зі значенням `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[2]?.getAttribute('id'), 'location');
```

Третій прапорець повинен мати `name` зі значенням `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[2]?.getAttribute('name'), 'location');
```

Третій прапорець повинен мати `value` зі значенням `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[2]?.getAttribute('value'), 'location');
```

Під прапорцем, що стосується локації, має бути елемент `label` з текстом `Reputation`.

```js
assert.isNotNull(document.querySelectorAll('fieldset:nth-of-type(3) input[value="location"] + label')?.textContent, 'Reputation');
```

Елемент `label` повинен мати атрибут `for` зі значенням `"reputation"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'reputation');
```

Після елементу `label` має бути четвертий прапорець.

```js 
assert(document.querySelector("fieldset:nth-of-type(3) label + input:nth-of-type(4)[type='checkbox']"));
```

Четвертий прапорець повинен мати `id` зі значенням `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[3]?.getAttribute('id'), 'reputation');
```

Четвертий прапорець повинен мати `name` зі значенням `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[3]?.getAttribute('name'), 'reputation');
```

Четвертий прапорець повинен мати `value` зі значенням `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[3]?.getAttribute('value'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="ads" name="ads" value="ads" />

          <label for="recommendation">Personal Recommendation</label>
          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
