---
id: 66a93bbe65a26169dbf3bc39
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Форма відгуків про готель також має збирати електронну адресу користувачів.

Для початку додайте новий елемент `label` з текстом `Email address (required):` до форми. Елемент `label` повинен мати атрибут `for` зі значенням `"email"`.

# --hints--

Ви повинні мати елемент `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

Елемент `label` повинен містити текст `Email address (required):`.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText, 'Email address (required):');
```

Елемент `label` повинен мати атрибут `for` зі значенням `"email"`.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>
--fcc-editable-region--

--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
