---
id: 66a95d0eff8168747805f1f3
title: Крок 22
challengeType: 0
dashedName: step-22
---

# --description--

Якщо потрібно, щоб користувач вибрав декілька опцій зі списку, можна використати прапорці.

Ось приклад того, як працювати з прапорцями на прикладі їжі:

```html
<fieldset>
  <legend>Food Options</legend>
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="burger">Burger</label>
  <input type="checkbox" id="burger" name="food" value="burger">
</fieldset>
```

Атрибут `value` використовують, щоб вказати значення, яке буде відправлене на сервер при надсиланні форми.

Додайте елемент `label` з текстом `Social Media Ads` під елементом `legend`. Значенням атрибута `for` має бути `"ads"`.

Під елементом `label` додайте прапорець з атрибутами `id`, `name` та `value` зі значенням `"ads"`.

# --hints--

Ви повинні мати елемент `label` під `legend`.

```js
assert(document.querySelector("fieldset:nth-of-type(3) legend + label"));
```

Елемент `label` повинен містити текст `Social Media Ads`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + label")?.textContent, "Social Media Ads");
```

Елемент `label` повинен мати атрибут `for` зі значенням `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + label")?.getAttribute("for"), "ads");
```

Під елементом `label` має бути прапорець.

```js
const inputElement = document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']");
assert.strictEqual(inputElement?.parentElement.tagName, "FIELDSET");
```

Прапорець повинен мати атрибут `id` зі значенням `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.id, "ads");
```

Прапорець повинен мати атрибут `name` зі значенням `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.name, "ads");
```

Прапорець повинен мати атрибут `value` зі значенням `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.value, "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        --fcc-editable-region--
        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
