---
id: 66a969951120be7818d8ee49
title: Крок 27
challengeType: 0
dashedName: step-27
---

# --description--

У наступному розділі форми користувачі зможуть залишати оцінку готелю.

Для початку додайте елемент `fieldset`, який містить елемент `legend`. Елемент `legend` повинен мати текст `Ratings`.

Додайте елемент `label` з текстом `How was the service?` під елементом `legend`. Значенням атрибута `for` має бути `"service"`.

# --hints--

Ви повинні мати елемент `fieldset`.

```js
assert.strictEqual(document.querySelectorAll('fieldset').length, 4);
```

Елемент `fieldset` повинен мати елемент `legend`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend').length, 1);
```

Елемент `legend` повинен мати текст `Ratings`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend')[0]?.textContent, 'Ratings');
```

Елемент `fieldset` повинен мати елемент `label` під `legend`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label').length, 1);
```

Елемент `label` повинен мати текст `How was the service?`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.textContent, 'How was the service?');
```

Елемент `label` повинен мати атрибут `for` зі значенням `"service"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.getAttribute('for'), 'service');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="ads" name="ads" value="ads" />

          <label for="recommendation">Personal Recommendation</label>
          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />

          <label for="location">Location</label>
          <input type="checkbox" id="location" name="location" value="location" />

          <label for="reputation">Reputation</label>
          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />

          <label for="price">Price</label>
          <input type="checkbox" id="price" name="price" value="price" />
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
