---
id: 66a93730719e1f68410cce54
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Атрибут `name` використовують, щоб ідентифікувати дані форми після їх надсилання на сервер.

Ось приклад використання атрибута `name`:

```html
<input type="email" name="email">
```

До елемента `input` додайте атрибут `name` зі значенням `"name"`.

# --hints--

Не видаляйте наявний елемент `input`.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Не створюйте новий елемент `input`.

```js
assert.strictEqual(document.querySelectorAll('fieldset input').length, 1);
```

Елемент `input` повинен мати атрибут `name` зі значенням `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
