---
id: 668f08ea07b99b1f4a91acab
title: Crie uma página de receitas
challengeType: 14
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve ter uma declaração `!DOCTYPE html`.
1. Você deve ter um elemento `html` com `lang` definido como `en`.
1. Você deve ter um elemento `head` contendo um elemento `title` com o nome da receita e um elemento `meta` com um atributo `charset` definido como `UTF-8`.
1. Você deve ter um elemento `body`.
1. Você deve ter um elemento `h1` com o nome da receita.
1. Você deve ter um elemento `p` que introduz a receita abaixo do `h1`.
1. Você deve ter um elemento `h2` com o texto `Ingredients` para a seção de ingredientes.
1. Você deve ter uma lista não ordenada (elemento `ul`) com pelo menos quatro itens de lista (elementos `li`) que listam os ingredientes abaixo do primeiro elemento `h2`.
1. Você deve ter um segundo elemento `h2` com o texto `Instructions` para a seção de instruções.
1. Você deve ter uma lista ordenada (elemento `ol`) com pelo menos quatro itens de lista que listam os passos da receita em ordem, abaixo do segundo elemento `h2`.
1. Você deve ter um elemento `img` com um atributo `src` definido para uma imagem válida, você pode usar `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg` se preferir, além de um atributo `alt` descrevendo a imagem.

# --hints--

A página de receitas deve ter uma declaração `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Você deve ter um elemento `html` com `lang` definido como `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

Você deve ter um elemento `head` dentro do elemento `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Você deve ter um elemento `title` dentro do elemento `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

O elemento `title` deve ter o título da receita.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Você deve ter um elemento `meta` dentro do elemento `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

O elemento `meta` deve ter o atributo `charset` definido como `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Você deve ter um elemento `body` dentro do elemento `html`.

```js
assert(code.match(/<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i));
```

Você deve ter um elemento `h1` com o nome da receita.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Você deve ter apenas um elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Você deve ter um elemento `p` abaixo do elemento `h1` com o texto descrevendo a receita.

```js
assert.strictEqual(document.querySelector('p')?.previousElementSibling.tagName, 'H1');
```

O primeiro elemento `h2` deve ter o texto `Ingredients`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

Você deve ter um elemento de lista não ordenada abaixo do primeiro elemento `h2`.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

Você deve ter pelo menos quatro elementos de item de lista na lista não ordenada com os ingredientes.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

O segundo elemento `h2` deve ter o texto `Instructions`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

Você deve ter um elemento de lista ordenada abaixo do segundo elemento `h2`.

```js
assert.strictEqual(document.querySelector('ol')?.previousElementSibling.tagName, 'H2');
```

Você deve ter pelo menos quatro elementos de item de lista na lista ordenada com as instruções.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Você deve ter pelo menos um elemento `img`.

```js
assert.exists(document.querySelector('img'));
```

Todos os elementos `img` devem ter um atributo e valor de `src` válido.

```js
assert.isAbove(document.querySelector('img')?.src.length, 0);
```

Todos os elementos `img` devem ter um atributo `alt` para descrever a imagem.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
