---
id: 68e6aff5325410118e76e7be
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop, eserciterai a lavorare con ARIA costruendo un controller audio accessibile.

Per iniziare, aggiungi un elemento `h1` con il testo `Audio Controls`.

# --hints--

Dovresti avere un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Il tuo elemento `h1` dovrebbe essere all’interno del tuo elemento `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Il tuo elemento `h1` dovrebbe avere il testo `Audio Controls`.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
