---
id: 68eac57fadcc662ce917e75e
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Successivamente, aggiungi un elemento `button` con l'attributo `type` impostato sul valore `button`.

Il testo dell'elemento `button` dovrebbe essere `Play`.

# --hints--

Dovresti avere un elemento `button`.

```js
assert.exists(document.querySelector('button'));
```

Il tuo elemento `button` dovrebbe avere un attributo `type`.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

Il tuo elemento `button` dovrebbe avere l'attributo `type` impostato su `button`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

Il tuo elemento `button` dovrebbe avere il testo `Play`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
