---
id: 68eb7f2397da251a40e04aa2
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Successivamente, aggiungi un elemento `div` per contenere tutti gli elementi del controllo del volume.

All'interno dell'elemento `div`, annida un elemento `span` con un attributo `id` impostato su `volume-label`.
Inoltre, il testo per l'elemento `span` dovrebbe essere `Volume`.

L'elemento `id` sarà poi referenziato dagli attributi ARIA in modo che le tecnologie assistive possano identificare questo testo come l'etichetta del controllo del volume.

# --hints--

Dovresti avere un elemento `div` sotto il tuo elemento `button`.

```js
assert.exists(document.querySelector('button + div'));
```

Dovresti avere un elemento `span` all'interno del tuo elemento `div`.

```js
assert.exists(document.querySelector('button + div > span'));
```

Il tuo elemento `span` dovrebbe avere un attributo `id`.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

Il tuo elemento `span` dovrebbe avere un attributo `id` con il valore impostato su `volume-label`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

Il tuo elemento `span` dovrebbe avere il testo `Volume`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
