---
id: 68ec0a48c9d6db3594d768c3
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Ora è il momento di aggiungere un altro elemento `span`.

Questo elemento `span` dovrebbe avere un attributo `id` impostato su `volume-description`.
Inoltre, il testo per questo elemento `span` dovrebbe essere `Adjust the sound level`.

Simile al precedente elemento `span`, il riferimento `id` di questo `span` sarà poi utilizzato dagli attributi ARIA in modo che le tecnologie assistive possano identificare questo testo come la descrizione per il controllo del volume.

# --hints--

Dovresti avere due elementi `span` all'interno dell'elemento `div`.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

Il tuo secondo elemento `span` dovrebbe avere un attributo `id`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

Il tuo secondo elemento `span` dovrebbe avere un attributo `id` con il valore `volume-description`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

Il tuo secondo elemento `span` dovrebbe avere il testo `Adjust the sound level`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
