---
id: 68ed4204ab7edc66b6cb2383
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Come ricorderai dalle lezioni precedenti, l'attributo `aria-labelledby` si usa quando c'è del testo già presente nella pagina che può essere usato come etichetta. In questo caso, hai aggiunto l'attributo `id` ai tuoi elementi `span` che ora verranno usati qui.

Aggiungi un attributo `aria-labelledby` all'insieme di elementi `input` impostandolo su `volume-label volume-description`.

# --hints--

Il tuo elemento `input` dovrebbe avere un attributo `aria-labelledby`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

Il tuo elemento `input` dovrebbe avere un attributo `aria-labelledby` con valore `volume-label volume-description`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
