---
id: 68ed488821845990c911a223
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Il tuo controller audio accessibile è quasi completo, ma c’è un ultimo dettaglio da aggiungere.

Aggiungi un elemento `button` con l’attributo `type` impostato su `button` e con il testo `Mute` sotto l’elemento `div`.

Con questa modifica, hai completato questo workshop!

# --hints--

Dovresti avere un elemento `button` sotto l’elemento `div`.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

Il tuo elemento `button` dovrebbe avere un attributo `type`.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

Il tuo elemento `button` dovrebbe avere un attributo `type` con il valore `button`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

Il tuo elemento `button` dovrebbe contenere il testo `Mute`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
