---
id: 68ed3a96e0ebcc36565b39e5
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Successivamente, devi aggiungere un elemento `input` con l'attributo `type` impostato su `range` per creare il cursore del volume. Imposta l'attributo `min` su `0`, l'attributo `max` su `100` e l'attributo `value` su `50` per definire il livello di volume predefinito.

# --hints--

Dovresti avere un elemento `input` sotto il tuo secondo elemento `span`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

Il tuo elemento `input` dovrebbe avere un attributo `type`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

Il tuo elemento `input` dovrebbe avere un attributo `type` con il valore `range`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

Il tuo elemento `input` dovrebbe avere un attributo `min`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

Il tuo elemento `input` dovrebbe avere un attributo `min` con il valore `0`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

Il tuo elemento `input` dovrebbe avere un attributo `max`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

Il tuo elemento `input` dovrebbe avere un attributo `max` con il valore `100`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

Il tuo elemento `input` dovrebbe avere un attributo `value`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

Il tuo elemento `input` dovrebbe avere un attributo `value` con il valore `50`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
