---
id: 587d825a367417b2b2512c89
title: 實作 Quicksort 演算法
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

在這個實驗中你將實作排序演算法：quicksort。

Quicksort 是一種高效的、遞迴的分治法，用來排序陣列。在這個方法中，會在原始陣列中選擇一個樞軸值。接著，陣列會被分割成兩個子陣列，分別包含小於和大於該樞軸值的元素。然後，將對這兩個子陣列遞迴呼叫 quicksort 演算法的結果合併。這個過程會持續，直到達到空陣列或單一項目陣列的基底情況，並傳回該陣列。遞迴呼叫的回溯會傳回排序完成的陣列。

雖然樞軸值的選擇很重要，但對於我們這裡的目的，任何樞軸都可以。為了簡單起見，可以使用第一個或最後一個元素。

Quicksort 是一種非常高效的排序方法，平均提供 *O(nlog(n))* 的效率。它也相對容易實作。這些屬性使它成為一種受歡迎且實用的排序方法。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `quicksort` 函式（程式），它以整數陣列作為輸入。
1. `quicksort` 函式應該傳回一個陣列，該陣列包含與輸入相同的整數，但排序從最低到最高。
1. `quicksort` 函式應該遞迴呼叫自身來排序陣列。

# --hints--

你應該有一個 `quicksort` 函式（程式）。

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` 應該傳回一個除了順序外未改變的陣列。

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` 應該傳回一個已排序的陣列（由小到大）。

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` 不應該使用內建的 `.sort()` 方法。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
