---
id: 56533eb9ac21ba0edf2244b3
title: Costruisci un convertitore da Celsius a Fahrenheit
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

In questo laboratorio, scriverai una funzione che converte la temperatura da Celsius a Fahrenheit. La formula per convertire da Celsius a Fahrenheit è:

```js
fahrenheit = celsius * (9/5) + 32
```

**Obiettivo:** Soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi creare una funzione chiamata `convertCtoF`.
2. `convertCtoF` deve prendere un singolo argomento numerico, che è la temperatura in Celsius.
3. `convertCtoF` deve restituire un numero.

# --hints--

Devi creare una funzione chiamata `convertCtoF`.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` deve prendere un singolo parametro.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` deve restituire un numero.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` deve restituire un valore di `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` deve restituire un valore di `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` deve restituire un valore di `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` deve restituire un valore di `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` deve restituire un valore di `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
