---
id: 6823f9df49cc206af5471a30
title: تصحيح صفحة الملف الشخصي لـ Camperbot
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

يحاول Camperbot بناء صفحة ملف شخصي. طلب من صديق أن يراجع شفرتهم وقال إن بها بعض الأخطاء.

مهمتك هي إصلاح جميع أخطاء Camperbot حتى يتمكن من متابعة بناء صفحة ملفه الشخصي. أكمل العناصر في قصص المستخدم أدناه وانقر على "Check Your Code" لترى إذا أصلحت كل الأخطاء.

**قصص المستخدم:**

1. يحاول Camperbot استخدام عنصر `heading2`، لكن هذا العنصر غير موجود. أصلح تلك الوسوم لتستخدم عنصر عنوان من المستوى الثاني الصحيح.
2. يحاول Camperbot إضافة فقرتين باستخدام `pp`، لكنهما غير موجودتين أيضًا. أصلحهما لتستخدم وسوم الفقرة الصحيحة.
3. يستخدم Camperbot عنصر `h3` للعنوان الفرعي `Background and Interests` لكنه يحتوي على خطأ نحوي. اكتشف المشكلة وحلها.

# --hints--

يجب أن يكون لديك عنصر `h1` واحد فقط على الصفحة.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

لا يجب إزالة العنصر `<h1>Hello from Camperbot!</h1>` من الصفحة.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

يجب أن يكون لديك عنصر `h2` واحد فقط على الصفحة.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

يجب أن يكون لديك عنصر `h2` يحتوي على النص `About`. إليك مثال: `<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

يجب أن يكون لديك عنصر فقرة يحتوي على النص `My name is Camperbot and I love learning new things.` إليك مثال: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

يجب أن يكون لديك عنصر فقرة يحتوي على النص `I enjoy solving puzzles.` إليك مثال: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

يجب أن يكون لديك عنصران فقرتان بالضبط على الصفحة.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

يجب أن يكون لديك عنصر `h3` واحد فقط على الصفحة.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

يجب أن يكون لديك عنصر `h3` يحتوي على النص `Background and Interests`. إليك مثال: `<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

يجب أن يكون لديك خمسة عناصر إجمالًا على الصفحة.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
