---
id: 66faac4139dbbd5dd632c860
title: 求人応募フォームを作成する
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してチャレンジを完了してください。

**ユーザーストーリー:**

1. `div` 要素にクラス `container` があること。
1. `div` 要素の中に `form` 要素があること。
1. フォームには、ユーザーのフルネームを入力するためのタイプ `input`、ID `text` の `name` 要素が含まれていること。
1. ユーザーのメールアドレスを入力するためのタイプ `input`、ID `email` の別の `email` 要素があること。
1. フォームには、職種を選択できる ID `select` の `position` 要素が含まれていること。
1. クラス `fieldset` の `radio-group` 要素があること。
1. `.radio-group` の中に、タイプ `input` の `radio` 要素のセットがあり、利用可能なオプション（例：フルタイム、パートタイム）を選択できること。グループ `name` は `availability` であること。
1. メッセージを入力するための ID `textarea` の `message` 要素があること。
1. すべての `input` 要素に `label` 要素が関連付けられていること。
1. フォームを提出するためのタイプ `button` の `submit` 要素があること。
1. `:focus` と `input` 要素に `textarea` 疑似クラスを追加し、フォーカス時にボーダーの色を変えること。
1. `input`、`select`、`textarea` 要素には、無効な入力が検出されたときにボーダーの色を赤に変える `:invalid` 疑似クラスがあること。
1. `input`、`select`、`textarea` 要素には、有効な入力がされたときにボーダーの色を緑に変える `:valid` 疑似クラスがあること。
1. `button` 要素には、ホバー時に背景色を変える `:hover` 疑似クラスがあること。
1. ラジオボタンが選択されたときにボーダー色、背景色、ボックスシャドウを追加するために `:checked` 疑似クラスを `.radio-group input[type="radio"]` に使うこと。
1. ラジオボタンが選択されたときに関連する `:checked` のテキスト色を変えるために `label` 疑似クラスをラジオボタンに使うこと。
1. 最初の入力フィールドを異なるスタイルにするために `:first-of-type` 疑似クラスを `input` 要素に追加すること。（例：角を丸くする）

**注意:** スタイルシートをHTMLにリンクし、CSSを適用してください。

# --hints--

`div` 要素にクラス `container` があること。

```js
assert.exists(document.querySelector("div.container"));
```

`div` 要素の中に `form` 要素があること。

```js
assert.exists(document.querySelector("div.container > form"));
```

フォームには、ユーザーのフルネームを入力するためのタイプ `input`、ID `text` の `name` 要素が含まれていること。

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

ユーザーのメールアドレスを入力するためのタイプ `input`、ID `email` の別の `email` 要素があること。

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

フォームには、ID `select` の `position` 要素があり、いくつかの `option` 要素が含まれていること。

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

クラス `fieldset` の `section` または `.radio-group` 要素があること。

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

`.radio-group` の中に、利用可能なオプションを選択するためのタイプ `input` の `radio` 要素のグループがあること。グループ `name` は `availability` であること。

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

メッセージを入力するための ID `textarea` の `message` 要素があること。

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

すべての `input` 要素に `label` 要素が関連付けられていること。

```js
const els = document.querySelectorAll('input');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

フォームを提出するためのタイプ `button` の `submit` 要素があること。

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

`:focus` と `input` 要素に `textarea` 疑似クラスを追加し、フォーカス時にボーダーの色を変えること。指定された順序でリストセレクターを使うこと。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'))
```

`input`、`select`、`textarea` 要素には、無効な入力が検出されたときにボーダーの色を赤に変える `:invalid` 疑似クラスがあること。指定された順序でリストセレクターを使うこと。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

`input`、`select`、`textarea` 要素には、有効な入力がされたときにボーダーの色を緑に変える `:valid` 疑似クラスがあること。指定された順序でリストセレクターを使うこと。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

`button` 要素には、ホバー時に背景色を変える `:hover` 疑似クラスがあること。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

ラジオボタンが選択されたときにボーダー色を追加するために `:checked` 疑似クラスを `.radio-group input[type="radio"]` に使うこと。

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

ラジオボタンが選択されたときに背景色を追加するために `:checked` 疑似クラスを `.radio-group input[type="radio"]` に使うこと。

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

ラジオボタンが選択されたときにボックスシャドウを追加するために `:checked` 疑似クラスを `.radio-group input[type="radio"]` に使うこと。

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

ラジオボタンが選択されたときに関連する `:checked` のテキスト色を変えるために `label` 疑似クラスを使うこと。

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

最初の入力フィールドを異なるスタイルにするために `:first-of-type` 疑似クラスを `input` 要素に追加すること。

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
