---
id: 66da326c02141df538f29ba5
title: Створіть платіжну сторінку
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `h1` з текстом `Checkout`.
1. Одразу після елемента `h1` має бути два елементи `section`.
1. У першому розділі має бути елемент `h2` з текстом `Your Cart`.
1. У першому розділі має бути зображення товару з відповідним альтернативним текстом. Можете використати це зображення: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. У другому розділі має бути елемент `h2` з текстом `Payment Information`.
1. У другому розділі має бути елемент `form`.
1. Форма має містити поле для введення даних з `id` та `name` зі значенням `card-name`, а також відповідним елементом `label`.
1. Форма має містити поле для введення даних з `id` та `name` зі значенням `card-number`, а також відповідним елементом `label`.
1. Принаймні два поля для введення даних мають бути обов’язковими та мати атрибут `aria-required` зі значенням `true`.

# --hints--

Ви повинні мати елемент `h1` з текстом `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

На сторінці має бути лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Після елемента `h1` має бути принаймні два елементи `section`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

У першому розділі має бути елемент `h2` з текстом `Your Cart`.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

У першому розділі має бути принаймні одне зображення з альтернативним текстом.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

У другому розділі має бути елемент `h2` з текстом `Payment Information`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

У другому елементі `section` має бути елемент `form`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

Форма має містити елемент `input` з `id` та `name` зі значенням `card-name`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

Форма має містити елемент `input` з `id` та `name` зі значенням `card-number`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

Усі елементи `input`, яким не надано `type` зі значенням `submit`, повинні мати відповідний елемент `label`.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => assert.exists(document.querySelector(`label[for="${input.id}"]`)));
```

Принаймні два елементи `input` мають бути обов’язковими та мати атрибут `aria-required` зі значенням `true`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required && input.getAttribute('aria-required') === 'true');
assert.isAtLeast(requiredInputs?.length, 2)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name</label>
                    <input type="text" id="card-name" name="card-name" required aria-required="true">
                </div>
                <div>
                    <label for="card-number">Card Number</label>
                    <input type="number" id="card-number" name="card-number" required aria-required="true">
                </div>
                <div>
                    <label for="expiry-date">Expiry Date</label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required aria-required="true">
                </div>
                <div>
                    <label for="cvv">CVV</label>
                    <input type="number" id="cvv" name="cvv" required aria-required="true" aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
