---
id: 612e6afc009b450a437940a1
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ابدأ بالهيكل الأساسي لـ HTML. أضف إعلان `DOCTYPE` وعناصر `html` و`head` و`body` و`title`.

حدد لغة هذه الصفحة إلى الإنجليزية. عيّن `title` إلى `Piano`.

# --hints--

يجب أن تحتوي شفرتك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب أن تضع مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب أن تحدد نوع المستند ليكون `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

يجب أن تغلق إعلان `DOCTYPE` بعلامة `>` بعد النوع.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم بداية مع خاصية `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم إغلاق.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يكون إعلان `DOCTYPE` في بداية شفرة HTML الخاصة بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك وسم بداية `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك وسم إغلاق `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك وسم بداية `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك وسم إغلاق `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون عنصرا `head` و`body` أخوين.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

يجب أن تحتوي شفرتك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون لمشروعك عنوان `Piano`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

تذكر، حساسية الحروف والإملاء مهمة للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
