---
id: 612e78af05201622d4bab8aa
title: خطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

حان الوقت للبدء في العمل على البيانو. أنشئ عنصر `div` داخل عنصر `body` مع تعيين `id` إلى `piano`.

# --hints--

يجب أن تنشئ عنصر `div` جديد.

```js
const div = document.querySelector('div');
assert.exists(div);
```

يجب أن يكون عنصر `div` داخل عنصر `body`.

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

يجب أن يحتوي عنصر `div` على `id` معين إلى `piano`.

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
