---
id: 612e77aba7ca691f598feb02
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف وسمين `meta`، واحد لتحسين صفحتك لأجهزة الجوال، وآخر لتحديد `charset` مقبول للصفحة.

# --hints--

يجب أن تضيف عنصرين `meta` إلى صفحتك.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

يجب أن يكون لديك عنصران `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

يجب أن يحتوي أحد عناصر `meta` على خاصية `name` مضبوطة على `viewport`، و`content` مضبوطة على `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

يجب أن يحتوي عنصر `meta` الآخر على خاصية `charset` مضبوطة على `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
