---
id: 612e7d1c29fb872d6384379c
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

قم بتسكين `div` ثاني داخل `div` الموجود لديك، واضبط `class` ليكون `keys`.

# --hints--

يجب أن تنشئ عنصر `div` ثاني.

```js
const divDiv = document.querySelectorAll('div');
assert(divDiv?.length === 2);
```

يجب أن يكون عنصر `div` الجديد داخل عنصر `div` الموجود لديك.

```js
const div = document.querySelector('div');
assert(div?.children?.length === 1);
assert(div?.children?.[0]?.localName === 'div');
```

يجب أن يحتوي عنصر `div` الجديد على `class` مضبوط إلى `keys`.

```js
const div = document.querySelector('div');
assert(div?.children?.[0]?.className === 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
