---
id: 612e8279827a28352ce83a72
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

انسخ الآن مجموعة العناصر السبعة `.key`، والصق مجموعتين إضافيتين داخل عنصر `.keys`.

# --hints--

يجب أن يكون لديك 21 عنصرًا من نوع `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

يجب أن يكون لديك 15 عنصرًا من نوع `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

يجب أن تكون كل عناصر `.key` داخل عنصر `.keys` الخاص بك.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

يجب أن تكون كل عناصر `.black--key` داخل عنصر `.keys` الخاص بك.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
