---
id: 612e89d254fe5d3df7d6693d
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

بعد أن أعدت ضبط نموذج الصندوق `html`، تحتاج إلى تمرير ذلك إلى العناصر الموجودة بداخله أيضًا. للقيام بذلك، يمكنك تعيين خاصية `box-sizing` إلى `inherit`، مما يخبر العناصر المستهدفة باستخدام نفس القيمة التي لدى العنصر الأب.

ستحتاج أيضًا إلى استهداف العناصر المزيفة، وهي كلمات مفتاحية خاصة تتبع المحدد. العنصران المزيفان اللذان ستستخدمهما هما `::before` و`::after`.

محدد `::before` ينشئ عنصرًا مزيفًا يكون الطفل الأول للعنصر المحدد، بينما محدد `::after` ينشئ عنصرًا مزيفًا يكون الطفل الأخير للعنصر المحدد. تُستخدم هذه العناصر المزيفة غالبًا لإنشاء محتوى تجميلي، سترى ذلك لاحقًا في هذا المشروع.

في الوقت الحالي، أنشئ محدد CSS يستهدف كل العناصر باستخدام `*`، وضمّن العناصر المزيفة مع `::before` و`::after`. عيّن خاصية `box-sizing` إلى `inherit`.

# --hints--

يجب أن يكون لديك محدد `*, ::before, ::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

يجب أن يحتوي محدد `*, ::before, ::after` على خاصية `box-sizing` معينة إلى `inherit`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
