---
id: 612ea4c4993aba52ab4aa32e
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

حان الوقت الآن لاستخدام محددات العناصر المزيفة التي أعددتها سابقًا. لإنشاء المفاتيح السوداء، أضف محددًا جديدًا `.key.black--key::after`. سيستهدف هذا العناصر التي تحمل الصنف `key black--key`، ويختار العنصر المزيف الذي يلي هذه العناصر في HTML.

في المحدد الجديد، عيّن خاصية `background-color` إلى `#1d1e22`. كما عيّن خاصية `content` إلى `""`. هذا سيجعل العناصر المزيفة فارغة.

تُستخدم خاصية `content` لتعيين أو تجاوز محتوى العنصر. افتراضيًا، العناصر المزيفة التي تُنشأ بواسطة محددات `::before` و`::after` تكون فارغة، ولن تُعرض على الصفحة. تعيين خاصية `content` إلى سلسلة نصية فارغة `""` يضمن عرض العنصر على الصفحة مع بقائه فارغًا.

إذا أردت التجربة، جرّب إزالة خاصية `background-color` وتعيين قيم مختلفة لخاصية `content`، مثل `"♥"`. تذكر التراجع عن هذه التغييرات عند الانتهاء حتى تجتاز الاختبارات.

# --hints--

يجب أن يكون لديك محدد `.key.black--key::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

يجب أن يحتوي محدد `.key.black--key::after` على خاصية `background-color` مضبوطة إلى `#1d1e22`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

يجب أن يحتوي محدد `.key.black--key::after` على خاصية `content` مضبوطة إلى `""`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
