---
id: 694bf6eb530e19a9c48c59f4
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Concatenare molte stringhe usando `+` e convertire numeri usando `str()` può diventare disordinato e difficile da leggere.

Python 3.6 ha introdotto le **f-string** per risolvere questo problema. Aggiungendo la lettera `f` prima della virgoletta di apertura, puoi inserire variabili ed espressioni all’interno di campi di sostituzione rappresentati da parentesi graffe `{}`. Per esempio:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

Crea una variabile `employee_card` e assegnale un f-string che mostri `Employee:` seguito da uno spazio e dal valore della variabile `full_name`.

# --hints--

Dovresti avere una variabile chiamata `employee_card`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

La tua variabile `employee_card` dovrebbe avere il valore `f'Employee: {full_name}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
