---
id: 694bf6eb530e19a9c48c59f7
title: Passaggio 18
challengeType: 20
dashedName: step-18
---

# --description--

Puoi anche usare numeri negativi per estrarre una porzione da una stringa partendo dalla fine. Per esempio, `-1` indica l’ultimo carattere, `-2` indica il penultimo carattere e così via.

Per ottenere gli ultimi tre caratteri di una stringa, puoi usare `string[-3:]`. Nota come il parametro stop sia omesso dopo i due punti. Questo significa che l’estrazione deve procedere fino al limite della stringa.

Crea una variabile chiamata `last_three`. Usa l’indicizzazione negativa per estrarre gli ultimi tre caratteri da `employee_code` (che rappresentano il numero ID). Infine, stampa `last_three` nel terminale.

Con questo, il workshop della scheda dipendente è completo.

# --hints--

Dovresti avere una variabile chiamata `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

Dovresti estrarre gli ultimi tre caratteri di `employee_code` usando l’indicizzazione negativa e assegnarli alla tua variabile `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

Dovresti stampare `last_three` nel terminale.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
