---
id: 694bf6eb530e19a9c48c59f1
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Ora il tuo indirizzo sembra incompleto. Vuoi anche aggiungere il numero dell’appartamento dove vive il dipendente, quindi dovresti modificare la variabile.

Quando vuoi aggiungere contenuto alla fine di una variabile stringa esistente, puoi usare l’operatore di **assegnazione aumentata** ,`+=`. Questo è più breve che scrivere `var = var + 'new text'`. Per esempio:

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

Ricorda che le stringhe sono immutabili, quindi questa operazione non cambia la stringa originale. Invece crea una nuova stringa e la riassegna alla variabile.

Usa l’operatore `+=` per aggiungere la stringa `, Apartment 4B` alla tua variabile `address`.

# --hints--

Dovresti usare l’operatore `+=` per aggiungere la stringa `, Apartment 4B` alla variabile `address`.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
