---
id: 694bf6eb530e19a9c48c59f6
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

Puoi estrarre una porzione da qualsiasi parte della stringa, non solo dall'inizio. Questo è utile in molti casi.

Crea una variabile `year_code` e assegnale la porzione di `employee_code` dall'indice `4` all'indice `8`. Questo estrarrà `2026`.

Poi crea una variabile `initials` e assegnale la porzione di `employee_code` dall'indice `9` all'indice `11`. Questo estrarrà `JD`.

Infine, stampa entrambe le variabili nel terminale.

# --hints--

Dovresti avere una variabile chiamata `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

Dovresti estrarre una porzione di `employee_code` dall'indice `4` all'indice `8` e assegnare il risultato alla tua variabile `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

Dovresti avere una variabile chiamata `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

Dovresti estrarre una porzione di `employee_code` dall'indice `9` all'indice `11` e assegnare il risultato alla tua variabile `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

Dovresti stampare `year_code` nel terminale.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

Dovresti stampare `initials` nel terminale.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
