---
id: 694bf3af05d10fd89dab05cc
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

In Python puoi unire stringhe usando l'operatore `+`. Questo si chiama concatenazione. Ecco un esempio:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

Crea una variabile `full_name` concatenando `first_name` e `last_name`. Poi stampa `full_name`.

# --hints--

Dovresti avere una variabile chiamata `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

Dovresti assegnare una stringa formata concatenando `first_name` e `last_name` alla tua variabile `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

Dovresti stampare la variabile `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
