---
id: 694bf6eb530e19a9c48c59f3
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Ora userai la funzione `str()` un'altra volta. Proprio come con age, devi convertire qualsiasi variabile numerica in una stringa prima di concatenarla con altro testo.

Crea una variabile chiamata `experience_years` e assegnale l'intero `5`.

Poi, crea una variabile `experience_info`. Assegna a essa una stringa formata concatenando `'Experience: '`, la variabile `experience_years` (convertita in stringa) e `' years'`. Stampa il risultato nel terminale.

# --hints--

Dovresti avere una variabile chiamata `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

Dovresti assegnare l'intero `5` alla tua variabile `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

Dovresti avere una variabile chiamata `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

Dovresti assegnare una stringa formata concatenando `'Experience: '`, `str(experience_years)` e `' years'` alla tua variabile `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

Dovresti stampare `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
