---
id: 694bf6ea530e19a9c48c59f0
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Se concateni due stringhe come `'John' + 'Doe'`, il risultato è `'JohnDoe'` senza spazi. Per risolvere questo, devi concatenare una stringa contenente uno spazio (`' '`) tra di esse.

Aggiorna la tua variabile `full_name` in modo che concateni `first_name`, uno spazio e `last_name`.

# --hints--

Dovresti aggiungere uno spazio (`' '`) tra `first_name` e `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
