---
id: 694bf6eb530e19a9c48c59f2
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Completa la frase aggiornando l'assegnazione `employee_info` per concatenare anche la stringa ` years old` alla fine. Ricorda di includere uno spazio all'inizio della tua stringa.

Infine, stampa `employee_info`.

# --hints--

Dovresti concatenare `' years old'` alla fine della variabile `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

Dovresti stampare `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
