---
id: 694bf6eb530e19a9c48c59f5
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Quando lavori con le stringhe, spesso avrai bisogno di estrarre una porzione specifica di una stringa. Questo si chiama **slicing**.

La sintassi è `string[start:stop]`, dove:

* `start` è l'indice da cui inizia lo slice (**inclusivo**).
* `stop` è l'indice in cui termina lo slice (**esclusivo**).

Per esempio, se `text = 'Python'`, allora `text[0:2]` restituisce `'Py'`.

Definisci `employee_code` come `'DEV-2026-JD-001'`. Dopodiché, crea una variabile `department` e assegnale lo slice di `employee_code` dall'indice `0` all'indice `3`. Poi stampa `department` nel terminale.

# --hints--
Dovresti avere una variabile chiamata `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

Dovresti assegnare la stringa `DEV-2026-JD-001` alla tua variabile `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

Dovresti avere una variabile chiamata `department`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

Dovresti prendere lo slice dei primi tre caratteri di `employee_code` e assegnarlo alla tua variabile `department`.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

Dovresti stampare `department` nel terminale.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
