---
id: 695c9c7c61d446f1451e0ae7
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora crea una variabile `address` per archiviare l'indirizzo del dipendente. Assegna a essa la stringa `123 Main Street` e infine stampa `address`.

# --hints--

Dovresti avere una variabile chiamata `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

Dovresti assegnare la stringa `123 Main Street` alla tua variabile `address`. Ricorda di racchiudere il testo tra virgolette singole o virgolette doppie.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

Dovresti stampare la variabile `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```
