---
id: 695c9f21a8b1d53cf1b97841
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Ora crea una variabile chiamata `employee_age` e assegnale l'intero `28`. 


# --hints--

Dovresti avere una variabile chiamata `employee_age`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

Dovresti assegnare l'intero `28` alla tua variabile `employee_age`. Non racchiudere il testo tra virgolette.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
