---
id: 695c96f82af4fb78e59fa9ed
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Le stringhe sono sequenze di caratteri usate per archiviare dati di testo. Come ricorderai dalle lezioni precedenti, puoi creare una stringa racchiudendo il testo tra virgolette singole (`'`) o virgolette doppie (`"`). Per esempio:

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

Crea una variabile `first_name` che archivia la stringa `John` e una variabile `last_name` che archivia la stringa `Doe`. Poi stampa `first_name` e `last_name`.

# --hints--

Dovresti avere una variabile chiamata `first_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

Dovresti assegnare la stringa `John` alla tua variabile `first_name`. Ricorda di racchiudere il testo tra virgolette singole o doppie.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

Dovresti avere una variabile chiamata `last_name`.

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

Dovresti assegnare la stringa `Doe` alla tua variabile `last_name`. Ricorda di racchiudere il testo tra virgolette singole o doppie.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

Dovresti stampare la tua variabile `first_name` nel terminale.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

Dovresti stampare la tua variabile `last_name` nel terminale.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
