---
id: 695df6480de23758af5a29cd
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Ora vuoi creare una stringa che mostri l'età del dipendente.

Inizia creando una variabile `employee_info` e assegnale il risultato della concatenazione di:

- la variabile `full_name`.
- una stringa composta dai caratteri `is` preceduti e seguiti da uno spazio.

# --hints--

Dovresti avere una variabile chiamata `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

Dovresti assegnare `full_name + ' is '` alla tua variabile `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
