---
id: 695ea1e55c0e619e597fc5c9
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Come puoi vedere, Python ha generato un `TypeError: can only concatenate str (not "int") to str`. Questo succede perché Python non permette di concatenare direttamente testo (stringhe) e numeri (numeri interi).

Per risolvere questo problema, devi prima convertire il numero in una stringa usando la funzione `str()`, che restituisce la versione stringa di un oggetto:

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

Aggiorna la tua assegnazione `employee_info` per convertire `employee_age` in una stringa usando `str(employee_age)`.

# --hints--

Dovresti modificare l’assegnazione della tua variabile `employee_info` convertendo `employee_age` in una stringa con `str(employee_age)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
