---
id: 695fcfbba648b3a96a9495a4
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

L'output nel terminale sta diventando affollato. Prima di continuare, è il momento di fare un po' di pulizia.

Rimuovi tutte le istruzioni `print()` dal tuo codice.

# --hints--

Non dovresti avere nessuna chiamata `print` nel tuo codice.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_calls("print")) == 0`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'  
last_name = 'Doe'  
--fcc-editable-region--  
print(first_name)  
print(last_name)  
full_name = first_name + ' ' + last_name  
print(full_name)  
address = '123 Main Street'  
address += ', Apartment 4B'  
print(address)  
--fcc-editable-region--
```
