---
id: 69727ea5c9606bedfe36d3d2
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

Ora è il momento di aggiungere gli ultimi dettagli alla scheda.

Crea una variabile chiamata `position` con il valore della stringa `Data Analyst` e una variabile chiamata `salary` con il valore dell'intero `75000`.

Poi, aggiorna la tua f-string `employee_card` per includere la posizione e lo stipendio. Deve seguire esattamente questo formato: `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. Sostituisci i segnaposto con le variabili corrispondenti.

Infine, stampa `employee_card` per vedere il risultato.

# --hints--

Dovresti avere una variabile chiamata `position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

Dovresti assegnare la stringa `Data Analyst` alla tua variabile `position`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

Dovresti avere una variabile chiamata `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

Dovresti assegnare l'intero `75000` alla tua variabile `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Dovresti aggiornare `employee_card` per usare la f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

Dovresti stampare `employee_card` nel terminale.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
