---
id: 66d6a7a3e1aa411e94bf2346
title: 建造房屋繪畫
challengeType: 25
dashedName: build-a-house-painting
demoType: onClick
---

# --description--

在本實驗中，你將使用 HTML 來創建房屋的結構。接著，你將使用 CSS 定位來排列房屋中的元素，例如視窗和門。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `div`，其 `id` 為 `house`。
1. 你的 `#house` 應該設定 `position` 為 `relative`，以便其子元素可以相對於它定位。
1. 你的 `#house` 應該有寬度為 `500px` 和高度為 `400px`。不要設定最小值或最大值。
1. 你的 `#house` 應該設定背景顏色和邊框。
1. 你應該在 `#house` 裡有五個 `div` 元素，且它們的 `id` 值分別為 `chimney`、`roof`、`window-1`、`window-2` 和 `door`。
1. `#house` 的所有直接子元素應該有 `position` 為 `absolute`。
1. `#roof`、`#chimney`、`#window-1`、`#window-2` 和 `#door` 應該設定寬度、高度、邊框和背景顏色。
1. 你的 `#roof` 應該有一個頂部值為 `0`。
1. 你的 `#door` 應該放置在你的房屋底部。
1. 你的視窗應該放置在你的 `#roof` 下面，且至少高於你的 `#door` 高度的三分之一。
1. 你的視窗和你的門都應該將 `left` 或 `right` 設定為一個將它們置於房屋邊框內的值。
1. 你的 `#chimney` 應該有一個頂部值，使它位於你的 `#house` 頂端。
1. 你的 `#chimney` 應該有一個 `z-index`，使它位於房子後方。

**注意：** 請務必在你的 HTML 中聯結你的樣式表以套用你的 CSS。

# --hints--

你應該有一個 `div`，其 `id` 為 `house`。

```js
assert.exists(document.querySelector("div#house"));
```

你應該將 `#house` 設為標的，並將它的 `position` 設為 `relative`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("position"), "relative");
```

你應該將標的設定為 `#house` 並設定它的 `border` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("border"));
```

你應該標的 `#house` 並設定它的 `background-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("background-color"));
```

`#house` 不應包含 `min-height` 屬性。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-height"));
```

`#house` 不應包含 `min-width` 屬性。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-width"));
```

`#house` 不應該包含 `max-height` 屬性。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-height"));
```

`#house` 不應包含 `max-width` 屬性。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-width"));
```

你應該將 `#house` 設為標的，並將其寬度設為 `500px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("width"), "500px")
```

你應該將 `#house` 設為標的，並將其高度設為 `400px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("height"), "400px")
```

你應該在 `#house` 中有一個 `id` 為 `chimney` 的 `div`。

```js
assert.exists(document.querySelector("div#house div#chimney"));
```

`#chimney` 應該將 `position` 屬性設為 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("position"), "absolute");
```

你應該將標的設定為 `#chimney` 並設定其 `width` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("width"));
```

你應該將標的設定為 `#chimney` 並設定它的 `height` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("height"));
```

你應該標的 `#chimney` 並設定它的 `border` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("border"));
```

你應該標的 `#chimney` 並設定它的 `background-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("background-color"));
```

你應該在 `#house` 中有一個 `id` 為 `roof` 的 `div`。

```js
assert.exists(document.querySelector("div#house  div#roof"));
```

`#roof` 應該將 `position` 屬性設為 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("position"), "absolute");
```

你應該標的 `#roof` 並設定它的 `width` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("width"));
```

你應該將標的設定為 `#roof` 並設定其 `height` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("height"));
```

你應該將標的設定為 `#roof` 並設定其 `border` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("border"));
```

你應該標的 `#roof` 並設定它的 `background-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("background-color"));
```

你應該在 `#house` 中有一個 `id` 為 `window-1` 的 `div`。

```js
assert.exists(document.querySelector("div#house  div#window-1"));
```

`#window-1` 應該將 `position` 屬性設為 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#window-1"))?.getPropertyValue("position"), "absolute");
```

你應該將標的設定為 `#window-1` 並設定其 `width` 屬性。

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window1 || windows);
```

你應該將標的設定為 `#window-1` 並設定其 `height` 屬性。

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window1 || windows);
```

你應該將標的設定為 `#window-1` 並設定其 `border` 屬性。

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window1 || windows);
```

你應該將標的設定為 `#window-1` 並設定其 `background-color` 屬性。

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window1 || windows);
```

你應該在 `#house` 中有一個 `id` 為 `window-2` 的 `div`。

```js
assert.exists(document.querySelector("div#house  div#window-2"));
```

`#window-2` 應該將 `position` 屬性設為 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#window-2"))?.getPropertyValue("position"), "absolute");
```

你應該將標的設定為 `#window-2` 並設定其 `width` 屬性。

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window2 || windows);
```

你應該將標的設定為 `#window-2` 並設定它的 `height` 屬性。

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window2 || windows);
```

你應該將標的設定為 `#window-2` 並設定其 `border` 屬性。

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window2 || windows);
```

你應該將標的設定為 `#window-2` 並設定其 `background-color` 屬性。

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window2 || windows);
```

你應該在 `#house` 中有一個 `id` 為 `door` 的 `div`。

```js
assert.exists(document.querySelector("div#house  div#door"));
```

`#door` 應該將 `position` 屬性設為 `absolute`。

```js
assert.equal(getComputedStyle(document.querySelector("#door"))?.getPropertyValue("position"), "absolute");
```

你應該將標的設定為 `#door` 並設定其 `width` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("width"));
```

你應該將標的設定為 `#door` 並設定其 `height` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("height"));
```

你應該標的 `#door` 並設定它的 `border` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("border"));
```

你應該標的 `#door` 並設定它的 `background-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("background-color"));
```

你的 `#roof` 應該將 `top` 屬性設為 `0`。

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("top"), "0px");
```

你的 `#door` 應該放置在房子的底部。

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorBottom = Number(door.getPropertyValue("bottom").replace("px", ""));
assert.strictEqual(doorBottom, 0);
```

你應該在你的 `#door` 上設定 `left` 或 `right`，並將它安排在房屋邊框內。

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorLeft = Number(door.getPropertyValue("left").replace("px", ""));
const doorRight = Number(door.getPropertyValue("right").replace("px", ""));
assert.isAbove(doorLeft, 0);
assert.isAbove(doorRight, 0);
```

你的 `#chimney` 應該有一個 `top` 值，使它位於你的 `#house` 頂部。

```js
const chimney = getComputedStyle(document.querySelector("#chimney"));
assert.approximately(parseFloat(chimney?.getPropertyValue("top")), - parseFloat(chimney?.getPropertyValue("height")), 2);
```

你的 `#chimney` 應該有一個 `z-index`，使它位於房子後方。

```js
const houseZIndex = getComputedStyle(document.querySelector("#house"))?.getPropertyValue("z-index");
const chimneyZIndex = getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("z-index");
if (houseZIndex === "auto") {
    assert.isBelow(Number(chimneyZIndex), 0)
} else {
    assert.isBelow(Number(chimneyZIndex), Number(houseZIndex));
}
```

你的視窗應該放置在你的 `#roof` 下面，且至少高於你的 `#door` 高度的三分之一。

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const door = getComputedStyle(document.querySelector("#door"))
const roof = getComputedStyle(document.querySelector("#roof"))
const w1Top = Number(window1.getPropertyValue("top").replace("px", ""));
const w1Bottom = Number(window1.getPropertyValue("bottom").replace("px", ""));
const w2Top = Number(window2.getPropertyValue("top").replace("px", ""));
const w2Bottom = Number(window2.getPropertyValue("bottom").replace("px", ""));
const doorHeight = Number(door.getPropertyValue("height").replace("px", ""));
const roofHeight = Number(roof.getPropertyValue("height").replace("px", ""));
assert.isAbove(w1Top, roofHeight);
assert.isAbove(w1Bottom, doorHeight/3);
assert.isAbove(w2Top, roofHeight);
assert.isAbove(w2Bottom, doorHeight/3);
```

你應該在你的視窗上設定 `left` 或 `right`，並將它們排列在房屋邊框內。

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const w1Left = Number(window1.getPropertyValue("left").replace("px", ""));
const w1Right = Number(window1.getPropertyValue("right").replace("px", ""));
const w2Left = Number(window2.getPropertyValue("left").replace("px", ""));
const w2Right = Number(window2.getPropertyValue("right").replace("px", ""));
assert.isAbove(w1Left, 0);
assert.isAbove(w1Right, 0);
assert.isAbove(w2Left, 0);
assert.isAbove(w2Right, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div id="house">
        <div id="chimney">
          <div id="smoke"></div>
        </div>
        <div id="roof"></div>
        <div id="window-1"></div>
        <div id="window-2"></div>
        <div id="door">
          <div id="door-knob"></div>
        </div>
        <div id="welcome">WELCOME</div>
    </div>
</body>
</html>
```

```css
* {
    box-sizing: border-box;
  }

  body {
    margin: 0;
    padding: 0;
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    background: linear-gradient(
      #b3e6ff,
      #b3e6ff 60%,
      green 60%
    );
  }

  #house {
    flex: 0 0 auto;
    width: 500px;
    height: 400px;
    position: relative;
    background-color: #ff9980;
    border: 6px solid #b35900;
  }

  #chimney {
    position: absolute;
    width: 90px;
    height: 100px;
    top: -100px;
    left: 65%;
    background: repeating-linear-gradient(
      #e6e6e6,
      #e6e6e6 10%,
      black 10%,
      black 12%
    );
    border: 2px solid black;
    z-index: -1;
  }

  #smoke {
    width: 40px;
    height: 40px;
    background-color: rgba(0,0,0,0.3);
    position: absolute;
    top: -20px;
    left: 20px;
    border-radius: 50%;
    box-shadow: 0 -10px 10px 10px rgba(0,0,0,0.3);
    border: none;
  }

  #roof {
    border: 2px;
    width: 100%;
    height: 110px;
    background-color: green;
    position: absolute;
    top: 0;
    background: repeating-linear-gradient(
      45deg,
      #b35900,
      #b35900 2%,
      transparent 2%,
      transparent 5%
    ), repeating-linear-gradient(
      -45deg,
      #b35900,
      #b35900 2%,
      #ff9980 2%,
      #ff9980 5%
    )
  }

  #window-1, #window-2 {
    width: 100px;
    height: 100px;
    background: linear-gradient(
      to right,
      #ffffb3,
      #ffffb3 48%,
      #b35900 48%,
      #b35900 52%,
      #ffffb3 52%,
      #ffffb3 100%
    );
    position: absolute;
    top: 40%;
    border: 6px solid #b35900;
  }

  #window-1 {
    left: 7.5%;
  }

  #window-2 {
    right: 7.5%;
  }

  #door {
    width: 130px;
    height: 180px;
    position: absolute;
    bottom: 0%;
    left: 185px;
    background-color: #e6e6e6;
    border: 6px solid #b35900;
    transform: translateY(6px);
  }

  #door-knob {
    width: 20px;
    height: 20px;
    background-color: #b35900;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    right: 5%;
  }

  #welcome {
    font-family: sans-serif;
    font-weight: bold;
    width: 130px;
    height: 40px;
    position: absolute;
    bottom: -47px;
    left: 178px;
    background-color: #85e085;
    border: 2px solid black;
    transform: skewX(-20deg);
    display: flex;
    align-items: center;
    justify-content: center;
  }
```
