---
id: 5a24c314108439a4d4036161
title: 了解自闭合 JSX 标签
challengeType: 6
forumTopicId: 301396
dashedName: learn-about-self-closing-jsx-tags
---

# --description--

到目前为止，已经看到 JSX 与 HTML 的关键不同在于使用 `className` 还是 `class` 来定义 HTML 的 class。

JSX 不同于 HTML 的另一个重要方面是自闭合标签。

在 HTML 中，几乎所有标签都有开始标签和结束标签：`<div></div>`；结束标签总是在你关闭的标签名称前有一个斜线。然而，HTML 中有一些特殊实例称为<dfn>空元素</dfn>，或者是不需要在另一个元素开始之前同时具有开始标签和结束标签的元素。

例如，换行标签可以写成 `<br>` 或 `<br />`，但绝不应写成 `<br></br>`，因为它不包含任何内容。

在 JSX 中，规则有些不同。任何 JSX 元素都可以用自闭合标签书写，并且每个元素都必须闭合。例如，换行标签必须始终写成 `<br />`，这样才能成为有效且可转译的 JSX。另一方面，`<div>` 可以写成 `<div />` 或 `<div></div>`。区别在于，第一种语法版本中无法在 `<div />` 内包含任何内容。你将在后续挑战中看到这种语法在渲染 React 组件时非常有用。

# --instructions--

修复代码编辑器中的错误，使其成为有效的 JSX 并成功编译。 确保不更改任何内容 -- 只需要在需要的地方关闭标签。

# --hints--

常量`JSX`应该返回一个`div`元素。

```js
assert.strictEqual(JSX.type, 'div');
```

`div` 应该包含一个 `br` 标签。

```js
assert(Enzyme.shallow(JSX).find('br').length === 1);
```

`div` 应该包含一个 `hr` 标签。

```js
assert(Enzyme.shallow(JSX).find('hr').length === 1);
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h2>Welcome to React!</h2> <br >
    <p>Be sure to close all tags!</p>
    <hr >
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div>
  <h2>Welcome to React!</h2> <br />
  <p>Be sure to close all tags!</p>
  <hr />
</div>
);
```
