---
id: 5a24c314108439a4d403617c
title: Use the Lifecycle Method componentWillMount
challengeType: 6
forumTopicId: 301423
dashedName: use-the-lifecycle-method-componentwillmount
---

# --description--

React 组件有几个特殊的方法，它们提供了在组件生命周期的特定时间点执行操作的机会。 这些被称为生命周期方法或生命周期钩子，允许你在特定时间点捕获组件。 这可以是在它们被渲染之前、更新之前、接收 props 之前、卸载之前等等。 以下是一些主要生命周期方法的列表：`componentWillMount()` `componentDidMount()` `shouldComponentUpdate()` `componentDidUpdate()` `componentWillUnmount()` 接下来的几节课将介绍这些生命周期方法的一些基本用例。

**注意：** `componentWillMount` 生命周期方法将在 16.X 的未来版本中被弃用，并在版本 17 中移除。 在这篇 <a href="https://www.freecodecamp.org/news/how-to-safely-use-reacts-life-cycles-with-fiber-s-async-rendering-fd4469ebbd8f/" target="_blank" rel="noopener noreferrer nofollow">文章</a> 中了解更多。

# --instructions--

当组件被挂载到 DOM 时，`componentWillMount()` 方法会在 `render()` 方法之前被调用。 在 `componentWillMount()` 中向控制台记录一些内容——你可能需要打开浏览器控制台来查看输出。

# --hints--

`MyComponent` 应该渲染一个 `div` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').length === 1;
  })()
);
```

`console.log` 应该渲染一个 `componentWillMount` 元素。

```js
assert(
  (function () {
    const lifecycle = React.createElement(MyComponent)
      .type.prototype.componentWillMount.toString()
      .replace(/ /g, '');
    return lifecycle.includes('console.log(');
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line
    console.log('Component is mounting...');
    // Change code above this line
  }
  render() {
    return <div />
  }
};
```
