---
id: 619665c9abd72906f3ad30f9
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ハッピーなFlappy Penguinを作りながら、CSSの変形やアニメーションの練習をさらに進めていきます。

まずはスタイルシートをページにリンクしてください。

# --hints--

コードには`link`要素が必要です。

```js
assert.exists(document.querySelector('link'));
```

`link`要素は`head`要素の中にあるべきです。

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

`link`要素には、値が`rel`の`stylesheet`属性が必要です。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link`要素には、値が`href`の`styles.css`属性が必要です。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
