---
id: 6196928658b6010f28c39484
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

`.ground` 要素に、時計回りに90度傾いた線形グラデーションの `background` を設定してください。グラデーションは `rgb(88, 175, 236)` から始まり、`rgb(182, 255, 255)` で終わります。

# --hints--

`.ground` に `background` の `linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255))` を設定してください。

```js
assert.include(['linear-gradient(90deg,rgb(88,175,236),rgb(182,255,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(88,175,236),rgb(182,255,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.ground')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.ground {
  width: 100vw;
  height: 400px;

}
--fcc-editable-region--
```
