---
id: 5e44412c903586ffb414c94c
title: بناء مشروع منسق العمليات الحسابية
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462359
dashedName: build-an-arithmetic-formatter-project
---

# --description--

غالبًا ما يرتب طلاب المرحلة الابتدائية المسائل الحسابية عموديًا لتسهيل حلها. على سبيل المثال، "235 + 52" تصبح:

```py
  235
+  52
-----
```

أكمل دالة `arithmetic_arranger` التي تستقبل قائمة من السلاسل النصية التي تمثل مسائل حسابية، وتُرجع المسائل مرتبة عموديًا وبجانب بعضها البعض. يجب أن تأخذ الدالة معلمة ثانية اختيارية. عند تعيين المعلمة الثانية إلى `True`، يجب عرض الإجابات.

## مثال

استدعاء الدالة:

```py
arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])
```

الناتج:

```py
   32      3801      45      123
+ 698    -    2    + 43    +  49
-----    ------    ----    -----
```

استدعاء الدالة:

```py
arithmetic_arranger(["32 + 8", "1 - 3801", "9999 + 9999", "523 - 49"], True)
```

الناتج:

```py
  32         1      9999      523
+  8    - 3801    + 9999    -  49
----    ------    ------    -----
  40     -3800     19998      474
```

## القواعد

ستُرجع الدالة التحويل الصحيح إذا كانت المسائل المقدمة منسقة بشكل صحيح، وإلا فستُرجع **سلسلة نصية** تصف خطأً ذا معنى للمستخدم.

- الحالات التي ستُرجع خطأ:
  - إذا تم تزويد الدالة بعدد **كبير جدًا من المسائل**. الحد الأقصى هو **خمسة**، وأي عدد أكبر سيُرجع:
    `'Error: Too many problems.'`
- العمليات الحسابية المقبولة التي ستقبلها الدالة هي **الجمع** و**الطرح** فقط. الضرب والقسمة ستُرجع خطأ. لا حاجة لاختبار عمليات أخرى غير المذكورة في هذه النقطة. الخطأ المُرجع سيكون:
    `"Error: Operator must be '+' or '-'."`
- يجب أن تحتوي كل رقم (المُعامل) على أرقام فقط. وإلا، ستُرجع الدالة:
    `'Error: Numbers must only contain digits.'`
- لكل مُعامل (أي الرقم على كل جانب من العملية) حد أقصى لأربعة أرقام في العرض. وإلا، ستُرجع الدالة رسالة الخطأ:
    `'Error: Numbers cannot be more than four digits.'`
- إذا زود المستخدم المسائل بالتنسيق الصحيح، فإن التحويل الذي تُرجعه سيتبع هذه القواعد:
  - يجب أن يكون هناك فراغ واحد بين العملية وأطول المُعاملين، وستكون العملية في نفس السطر مع المُعامل الثاني، وسيكون ترتيب المُعاملين كما هو مُعطى (الأول في الأعلى والثاني في الأسفل).
  - يجب أن تُحاذى الأرقام إلى اليمين.
  - يجب أن يكون هناك أربعة فراغات بين كل مسألة.
  - يجب أن تكون هناك شرطات في أسفل كل مسألة. يجب أن تمتد الشرطات على طول كل مسألة بشكل منفصل. (المثال أعلاه يوضح الشكل المطلوب.)

ملاحظة: افتح وحدة تحكم المتصفح بالضغط على F12 لرؤية إخراج أكثر تفصيلاً للاختبارات.

# --hints--

`arithmetic_arranger(["3801 - 2", "123 + 49"])` يجب أن تُرجع `  3801      123\n-    2    +  49\n------    -----`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3801 - 2", "123 + 49"]), '  3801      123\\n-    2    +  49\\n------    -----')`);
  }
})
```

`arithmetic_arranger(["1 + 2", "1 - 9380"])` يجب أن تُرجع `  1         1\n+ 2    - 9380\n---    ------`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["1 + 2", "1 - 9380"]), '  1         1\\n+ 2    - 9380\\n---    ------')`);
  }
})
```

`arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"])` يجب أن تُرجع `    3      3801      45      123\n+ 855    -    2    + 43    +  49\n-----    ------    ----    -----`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"]), '    3      3801      45      123\\n+ 855    -    2    + 43    +  49\\n-----    ------    ----    -----')`);
  }
})
```

`arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"])` يجب أن تُرجع `  11      3801      1      123         1\n+  4    - 2999    + 2    +  49    - 9380\n----    ------    ---    -----    ------`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"]), '  11      3801      1      123         1\\n+  4    - 2999    + 2    +  49    - 9380\\n----    ------    ---    -----    ------')`);
  }
})
```

`arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"])` يجب أن تُرجع `'Error: Too many problems.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"]), 'Error: Too many problems.')`);
  }
})
```

`arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"])` يجب أن تُرجع `"Error: Operator must be '+' or '-'."`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Operator must be '+' or '-'.")`);
  }
})
```

`arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"])` يجب أن تُرجع `'Error: Numbers cannot be more than four digits.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers cannot be more than four digits.")`);
  }
})
```

`arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"])` يجب أن تُرجع `'Error: Numbers must only contain digits.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers must only contain digits.")`);
  }
})
```

`arithmetic_arranger(["3 + 855", "988 + 40"], True)` يجب أن تُرجع `    3      988\n+ 855    +  40\n-----    -----\n  858     1028`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "988 + 40"], True), "    3      988\\n+ 855    +  40\\n-----    -----\\n  858     1028")`);
  }
})
```

`arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True)` يجب أن تُرجع `   32         1      45      123      988\n- 698    - 3801    + 43    +  49    +  40\n-----    ------    ----    -----    -----\n -666     -3800      88      172     1028`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True), "   32         1      45      123      988\\n- 698    - 3801    + 43    +  49    +  40\\n-----    ------    ----    -----    -----\\n -666     -3800      88      172     1028")`);
  }
})
```

# --seed--

## --seed-contents--

```py
def arithmetic_arranger(problems, show_answers=False):

    return problems

print(f'\n{arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])}')
```

# --solutions--

```py
def arithmetic_arranger(problems, result=False):

	lin1 = ""
	lin2 = ""
	lin3 = ""
	lin4 = ""

	if len(problems) > 5:
		return 'Error: Too many problems.'
	for problem in problems:
		[num1, sym, num2] = problem.split()
		sign = ['+', '-']
		if sym not in sign:
			return ("Error: Operator must be '+' or '-'.")
		if len(num1) > 4 or len(num2) > 4:
			return ("Error: Numbers cannot be more than four digits.")
		if not num1.isnumeric() or not num2.isnumeric():
			return ("Error: Numbers must only contain digits.")

		lin1 += "  " + num1 + "    " if len(num1) >= len(
		    num2) else " " * (len(num2) + 2 - len(num1)) + num1 + "    "
		lin2 += sym + " " + num2 + "    " if len(num2) >= len(
		    num1) else sym + " " * (len(num1) - len(num2) + 1) + num2 + "    "
		nmax = (len(num1) + 2) if len(num1) >= len(num2) else (len(num2) + 2)
		lin3 += "-" * nmax + "    "
		op = int(num1) + int(num2) if sym == "+" else int(num1) - int(num2)
		lin4 += (" " * (nmax - len(str(op)))) + str(op) + "    "

	arranged_problems = lin1.rstrip() + "\n" + lin2.rstrip(
	) + "\n" + lin3.rstrip()
	if result:
		arranged_problems += "\n" + lin4.rstrip()

	return arranged_problems
```
