---
id: 6894c1b715b8d18a18ef3203
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, practicarás con las clases y objetos de Python construyendo un inventario de instrumentos musicales.

Para definir una clase en Python, usa la palabra clave `class` seguida del nombre de la clase y dos puntos.

```python
class ClassName:
    # class contents go here
```

Recuerda que en Python, los nombres de clase normalmente usan PascalCase, lo que significa que cada palabra comienza con una letra mayúscula sin espacios entre palabras.

Comienza creando una clase llamada `MusicalInstrument`. Añade `pass` como el cuerpo de la clase por ahora.

# --hints--

Debes crear una clase llamada `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
