---
id: 6895e4e023e07603bcef46f9
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Ahora vamos a usar los métodos que has creado. 

Llama al método `play()` en tu instancia `instrument_1`. Esto imprimirá el mensaje sobre tocar el Oboe.

Para llamar a un método en una instancia, usa la notación de punto:

```python
instance_name.method_name()
```

# --hints--

Debes llamar al método `play()` en la instancia `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_1.play()")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
