---
id: 6895e4940e39950288e406c2
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

A continuación, agrega otro método a tu clase llamado `get_fact` que devuelve un dato sobre el instrumento.

Este método debe tener solo el parámetro `self` y devolver un f-string que diga `The [instrument name] is part of the [instrument type] family of instruments.`

A diferencia del método `play`, que imprime un mensaje directamente, este método debe devolver una cadena que pueda usarse en otra parte de tu código.

# --hints--

Debes definir un método llamado `get_fact` con un parámetro `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

Tu método `get_fact` debería devolver un f-string que contenga el nombre y el tipo del instrumento.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
