---
id: 6895e5337b1914050b0af7bd
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Ahora, llama al método `get_fact()` en la instancia `instrument_1` e imprime el resultado.

Dado que `get_fact()` devuelve una cadena en lugar de imprimirla directamente, necesitas usar `print()` para mostrar la cadena devuelta.

# --hints--

Debes llamar a `get_fact()` en `instrument_1` y mostrar el resultado.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()

--fcc-editable-region--
```
