---
id: 6895e3f91c069e0068a42582
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Dentro de tu método `play`, usa `print()` para mostrar un mensaje sobre tocar el instrumento.

Usa una f-string para incluir el nombre del instrumento en el mensaje: `The [instrument name] is fun to play!`. El `[instrument name]` debe ser reemplazado con el nombre real del instrumento.

# --hints--

Tu método `play` debe usar un f-string para mostrar el mensaje `The [instrument name] is fun to play!`. Usa `self.name` para acceder al nombre del instrumento.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
