---
id: 6895e197180c9af6988e1e19
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Crea una segunda instancia llamada `instrument_2` para un `Trumpet` que es un instrumento de `brass`.

Observa cómo puedes crear múltiples instancias a partir de la misma plantilla de clase. Cada instancia almacena sus propios valores de atributos de forma independiente de otras instancias.

# --hints--

Debes crear una instancia llamada `instrument_2` de la clase `MusicalInstrument`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

Debes pasar `Trumpet` y `brass` como argumentos al crear `instrument_2`.

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')

--fcc-editable-region--
```
