---
id: 6895e235935106f88281b648
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Puedes acceder a los atributos de tus instancias usando la notación de punto:

```py
instance_name.attribute_name
```

Esto te permite recuperar los valores que se almacenaron durante la inicialización.

Intenta imprimir los atributos `name` y `instrument_type` de ambas instancias. Podrás ver los valores que asignaste cuando las creaste.

# --hints--

Debes imprimir el atributo `name` de `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

Debes imprimir el atributo `instrument_type` de `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

Debes imprimir el atributo `name` de `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

Debes imprimir el atributo `instrument_type` de `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
