---
id: 6895e2becb7890fad818ad4d
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Ahora que entiendes cómo acceder a los atributos del objeto, elimina las llamadas existentes a `print`.

# --hints--

Debes eliminar la instrucción `print(instrument_1.name)`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

Debes eliminar la instrucción `print(instrument_1.instrument_type)`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

Debes eliminar la sentencia `print(instrument_2.name)`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

Debes eliminar la sentencia `print(instrument_2.instrument_type)`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
print(instrument_1.name)
print(instrument_1.instrument_type)

print(instrument_2.name)
print(instrument_2.instrument_type)

--fcc-editable-region--
```
