---
id: 6895e0452b3370f26f365a46
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

En el método `__init__`, necesitas asignar los parámetros a los atributos de la instancia. Los atributos de la instancia son variables que pertenecen a una instancia específica de una clase.

Dentro de tu método `__init__`, asigna el parámetro `name` a `self.name` y el parámetro `instrument_type` a `self.instrument_type`.

# --hints--

Debes asignar el parámetro `name` a `self.name` dentro del método `__init__`.

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

Debes asignar el parámetro `instrument_type` a `self.instrument_type` dentro del método `__init__`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        pass
--fcc-editable-region--
```
