---
id: 6895e356fd877efd6e3657ee
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--


Dentro de la clase `MusicalInstrument`, define otro método llamado `play` con solo el parámetro `self`. El parámetro `self` permite que los métodos accedan a los atributos de la instancia específica.

# --hints--

Debes definir un método llamado `play` con un parámetro `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
