---
id: 6895e0bdb0348bf46952ab67
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora es momento de crear instancias de tu clase `MusicalInstrument`.

Para crear una instancia de una clase, llamas al nombre de la clase como una función, pasando los argumentos requeridos para el método `__init__` (excepto `self`, que Python maneja automáticamente).

```python
my_instance = ClassName(arg1, arg2)
```

Cuando creas esta instancia, Python llama a tu método `__init__` y almacena los valores como atributos de instancia que pertenecen solo a esta instancia específica.

Fuera de la definición de la clase, crea una instancia llamada `instrument_1` para un `Oboe` que es un instrumento de `woodwind`.

# --hints--

Debes crear una instancia de la clase `MusicalInstrument` llamada `instrument_1`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

Debes pasar `Oboe` y `woodwind` como argumentos al crear `instrument_1`.

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
