---
id: 6895dffc64e97df0e62e9ef9
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Los métodos son funciones definidas dentro de una clase que pueden realizar acciones usando los datos de la clase. Aquí te mostramos cómo crear métodos:

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

En Python, el método `__init__` es un método especial conocido como el método inicializador, que se llama cuando creas una nueva instancia de la clase.

Dentro de tu clase `MusicalInstrument`, agrega un método `__init__` con tres parámetros: `self`, `name` y `instrument_type`.

El parámetro `self` es una forma estándar de referirse a la instancia de la clase y es obligatorio como primer parámetro para todos los métodos de instancia en una clase.

# --hints--

Debes definir un método `__init__` dentro de tu clase `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

Tu método `__init__` debe tener tres parámetros: `self`, `name` y `instrument_type`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
