---
id: 69aaafeaa4037197c75587cd
title: Implementa una Cola
challengeType: 26
dashedName: implement-a-queue
---

# --description--

En este laboratorio, implementarás una estructura de datos de cola usando funciones. Una cola es una estructura de datos First-In-First-Out (FIFO) donde los elementos se agregan al final de la cola y se eliminan desde el frente.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una `enqueue` function que agrega un elemento al final de la cola.
1. Debes tener una `dequeue` function que elimine y devuelva el elemento frontal de la cola.
1. Debes tener una función `front` que devuelva el elemento frontal de la cola sin eliminarlo.
1. Debes tener una función `size` que devuelva el número de elementos en la cola.
1. Debes tener una función `isEmpty` que devuelva `true` si la cola está vacía y `false` en caso contrario.

# --hints--

Debes tener una función `enqueue`.

```js
assert.isFunction(enqueue);
```

La función `enqueue` debe agregar un elemento al final de la cola.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

Debes tener una función `dequeue`.

```js
assert.isFunction(dequeue);
```

La función `dequeue` debe eliminar y devolver el elemento frontal de la cola.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

`dequeue` debe devolver `undefined` cuando se llama en una cola vacía.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` debe continuar agregando elementos al final de la cola después de los `dequeues`.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

Diferentes instancias de cola no deben compartir estado.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

Debes tener una función `front`.

```js
assert.isFunction(front);
```

La función `front` debe devolver el primer elemento de la cola sin eliminarlo.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`front` debe devolver `undefined` cuando se llama en una cola vacía.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

Debes tener una función `size`.

```js
assert.isFunction(size);
```

La función `size` debe devolver el número correcto de elementos en la cola.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

Debes tener una función `isEmpty`.

```js
assert.isFunction(isEmpty);
```

La función `isEmpty` debe devolver `true` para una cola vacía.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

La función `isEmpty` debe devolver `false` para una cola no vacía.

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

La función `isEmpty` debe devolver `true` después de que todos los elementos sean eliminados de la cola.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
