---
id: a5229172f011153519423690
title: Erstellen Sie einen Rechner für die Summe ungerader Fibonacci-Zahlen
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

In diesem Labor erstellen Sie einen Rechner für die Summe ungerader Fibonacci-Zahlen, der die Summe aller ungeraden Fibonacci-Zahlen berechnet, die kleiner oder gleich einer gegebenen positiven ganzen Zahl sind.

**Ziel**: Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `sumFibs`-Funktion haben, die eine Zahl als Argument akzeptiert.
1. Die `sumFibs`-Funktion sollte die Summe aller ungeraden Fibonacci-Zahlen zurückgeben, die kleiner oder gleich der gegebenen Zahl sind.
1. Die Fibonacci-Folge beginnt mit `0` und `1`, und jede nachfolgende Zahl ist die Summe der beiden vorherigen.
1. Nur die ungeraden Fibonacci-Zahlen sollten zur Summe hinzugefügt werden.

# --hints--

Sie sollten eine `sumFibs`-Funktion haben.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` sollte `1785` zurückgeben.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` sollte `4613732` zurückgeben.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` sollte `5` zurückgeben.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` sollte `60696` zurückgeben.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` sollte `135721` zurückgeben.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
