---
id: 68e45a2b4ab9b2f48b1cc6df
title: تصميم صفحة تخطيط خطط التسعير
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تحتوي صفحتك على عنصر `h1` بالنص `Pricing Plans`.
2. يجب أن تحتوي صفحتك على عنصر `div` بالصنف `pricing-container` أسفل عنصر `h1`.
    - يجب أن يحتوي محدِّد `.pricing-container` على خاصية `display` بقيمة `flex` وخاصية `flex-wrap` بقيمة `wrap`.
3. داخل عنصر `.pricing-container`، يجب أن تحتوي على ثلاثة عناصر `div` بالصنف `pricing-card` لتمثيل خطط التسعير.
    - يجب أن يحتوي أحد عناصر `.pricing-card` على الصنف `basic-plan` بالإضافة إلى صنف `pricing-card`.
    - يجب أن يحتوي أحد عناصر `.pricing-card` على الصنف `pro-plan` بالإضافة إلى صنف `pricing-card`.
    - يجب أن يحتوي أحد عناصر `.pricing-card` على الصنف `premium-plan` بالإضافة إلى صنف `pricing-card`.
4. يجب أن يحتوي عنصر `.basic-plan` على عنصر `h2` بالنص `Basic`.
5. يجب أن يحتوي عنصر `.basic-plan` على عنصر `p` بالنص `$9/month`.
6. يجب أن يحتوي عنصر `.pro-plan` على عنصر `h2` بالنص `Pro`.
7. يجب أن يحتوي عنصر `.pro-plan` على عنصر `p` بالنص `$19/month`.
8. يجب أن يحتوي عنصر `.premium-plan` على عنصر `h2` بالنص `Premium`.
9. يجب أن يحتوي عنصر `.premium-plan` على عنصر `p` بالنص `$29/month`.
10. يجب أن:
    - تستخدم كل عناصر `.pricing-card` خاصية Flexbox لترتيب محتواها عمودياً وتوزيع المساحة بين العناصر الفرعية باستخدام `space-between`.
    - تضبط خاصية `flex` إلى `0 0 200px` لمنحها عرضًا ثابتًا ومنعها من التمدد أو الانكماش في التخطيط.
    - تضبط خاصية `border` إلى `2px solid black` لرؤية كيف تشغل البطاقات المختلفة المساحة.
11. يجب أن يظهر عنصر `.basic-plan` أولاً في التخطيط. استخدم خاصية `order` لذلك.
12. يجب أن يظهر عنصر `.pro-plan` ثانيًا في التخطيط. استخدم خاصية `order` واضبط خاصية `flex-grow` إلى `2` ليشغل مساحة أكبر من الخطط الأخرى.
13. يجب أن يظهر عنصر `.premium-plan` أخيرًا في التخطيط. استخدم خاصية `order` لذلك.

**ملاحظة:** تأكد من ربط ورقة الأنماط في ملف HTML وتطبيق CSS الخاص بك.

# --hints--

يجب أن يحتوي على عنصر `h1`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن تحتوي صفحتك على عنصر `h1` بالنص `Pricing Plans`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

يجب أن يحتوي على عنصر `div` أسفل عنصر `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

يجب أن يحتوي عنصر `div` على الصنف `pricing-container`.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

يجب أن يكون لديك محدِّد `.pricing-container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

يجب أن يحتوي محدِّد `pricing-container` على الخاصية `display` بقيمة `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

يجب أن يحتوي محدِّد `.pricing-container` على الخاصية `flex-wrap` بقيمة `wrap`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

يجب أن يحتوي عنصر `.pricing-container` على ثلاثة عناصر `div` بالصنف `pricing-card`.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

يجب أن يحتوي أحد عناصر `.pricing-card` على الصنف `basic-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

يجب أن يحتوي أحد عناصر `.pricing-card` على الصنف `premium-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

يجب أن يحتوي أحد عناصر `.pricing-card` على الصنف `pro-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

يجب أن يكون لديك محدِّد `.pricing-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

يجب أن يحتوي محدِّد `.pricing-card` على الخاصية `display` بقيمة `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

يجب أن يحتوي محدِّد `.pricing-card` على الخاصية `flex-direction` بقيمة `column`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

يجب أن يحتوي محدِّد `.pricing-card` على الخاصية `justify-content` بقيمة `space-between`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

يجب أن يحتوي محدِّد `.pricing-card` على الخاصية `flex` بقيمة `0 0 200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

يجب أن يحتوي محدِّد `.pricing-card` على الخاصية `border` بقيمة `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

يجب أن يحتوي محدِّد `.basic-plan` على الخاصية `order` بقيمة `0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

يجب أن يحتوي عنصر `.basic-plan` على عنصر `h2` بالنص `Basic`.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

يجب أن يحتوي عنصر `.basic-plan` على عنصر `p` بالنص `$9/month`.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

يجب أن يحتوي عنصر `.pro-plan` على عنصر `h2` بالنص `Pro`.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

يجب أن يحتوي عنصر `.pro-plan` على عنصر `p` بالنص `$19/month`.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

يجب أن يحتوي عنصر `.pro-plan` على الخاصية `order` بقيمة `1`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

يجب أن يحتوي عنصر `.pro-plan` على الخاصية `flex-grow` بقيمة `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

يجب أن يحتوي عنصر `.premium-plan` على عنصر `h2` بالنص `Premium`.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

يجب أن يحتوي عنصر `.premium-plan` على عنصر `p` بالنص `$29/month`.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

يجب أن يكون عنصر `premium-plan` هو العنصر الأخير في التخطيط ويحتوي على الخاصية `order` بالقيمة `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```

