---
id: 66ed41f912d0bb1dc62da5dd
title: Costruisci un bot trivia in JavaScript
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti stampare `"Hello! I'm your coding fun fact guide!"` nella console come messaggio di saluto per l'utente.
1. Dovresti creare tre variabili: `botName`, `botLocation` e `favoriteLanguage`, che archiviano rispettivamente il nome del bot, da dove viene e il suo linguaggio di programmazione preferito.
1. Dovresti stampare `"My name is (botName) and I live on (botLocation)."` nella console.
1. Dovresti stampare `"My favorite programming language is (favoriteLanguage)."` nella console.
1. Dovresti usare `let` per creare una variabile `codingFact` e assegnarle una stringa che sia un fatto divertente sul linguaggio di programmazione preferito del tuo bot, includendo l'uso della variabile `favoriteLanguage`.
1. Dovresti stampare `codingFact` nella console.
1. Dovresti riassegnare la variabile `codingFact` a un nuovo fatto sul linguaggio preferito del bot usando di nuovo la variabile `favoriteLanguage`.
1. Dovresti stampare di nuovo `codingFact` nella console.
1. Dovresti riassegnare di nuovo la variabile `codingFact` a un altro nuovo fatto sul linguaggio preferito del bot usando la variabile `favoriteLanguage`.
1. Dovresti stampare per la terza volta `codingFact` nella console.
1. Dovresti stampare `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` nella console come messaggio di addio del bot.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti stampare `"Hello! I'm your coding fun fact guide!"` nella console.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

Dovresti dichiarare una variabile `botName`. Controlla bene che non ci siano errori di ortografia o maiuscole.

```js
assert.exists(botName);
```

La tua variabile `botName` dovrebbe essere una stringa.

```js
assert.isString(botName);
```

Dovresti dichiarare una variabile `botLocation`. Controlla bene che non ci siano errori di ortografia o maiuscole.

```js
assert.exists(botLocation);
```

La tua variabile `botLocation` dovrebbe essere una stringa.

```js
assert.isString(botLocation);
```

Dovresti dichiarare una variabile `favoriteLanguage`. Controlla bene che non ci siano errori di ortografia o maiuscole.

```js
assert.exists(favoriteLanguage);
```

La tua variabile `favoriteLanguage` dovrebbe essere una stringa.

```js
assert.isString(favoriteLanguage);
```

Dovresti stampare nella console `"My name is (botName) and I live on (botLocation)."` e aggiungere le variabili alla stringa.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

Dovresti stampare nella console `"My favorite programming language is (favoriteLanguage)."` e aggiungere la variabile alla stringa.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

Dovresti usare `let` per dichiarare una nuova variabile `codingFact`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

Dovresti assegnare a `codingFact` un valore che includa `favoriteLanguage`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

Dovresti stampare `codingFact` nella console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

Dovresti assegnare un nuovo valore a `codingFact` che contenga anche `favoriteLanguage`, e stamparlo nella console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

Dovresti assegnare un valore a `codingFact` per la terza volta che contenga anche `favoriteLanguage`, e stamparlo nella console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

Dovresti stampare nella console `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` e aggiungere i valori delle variabili.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
