---
id: 66fbcf750a62784cf11f562e
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

创建一个具有以下属性和值的 `recipe2` 对象：

| Key | Value   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

你应该创建一个名为 `recipe2` 的对象。

```js
assert.isObject(recipe2);
```

你的 `recipe2` 对象应具有一个 `name` 属性。

```js
assert.property(recipe2, "name");
```

你的 `name` 属性应设置为 `Chicken Curry`。

```js
assert.equal(recipe2.name, "Chicken Curry");
```

你的 `recipe2` 对象应具有一个 `ingredients` 属性。

```js
assert.property(recipe2, "ingredients");
```

你的 `ingredients` 属性应具有数组值。

```js
assert.isArray(recipe2.ingredients);
```

你的 `ingredients` 属性的数组值应包含 `chicken breast`、`coconut milk`、`curry powder`、`onion` 和 `garlic`。

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

你的 `recipe2` 对象应具有 `cookingTime` 属性。

```js
assert.property(recipe2, "cookingTime");
```

你的 `cookingTime` 属性值应为数字。

```js
assert.isNumber(recipe2.cookingTime);
```

你的 `cookingTime` 属性应设置为 `42`。

```js
assert.equal(recipe2.cookingTime, 42);
```

你的 `recipe2` 对象应有一个 `totalIngredients` 属性。

```js
assert.property(recipe2, "totalIngredients");
```

你的 `totalIngredients` 属性应设置为 `null`。

```js
assert.isNull(recipe2.totalIngredients);
```

你的 `recipe2` 对象应具有 `difficultyLevel` 属性。

```js
assert.property(recipe2, "difficultyLevel");
```

你的 `difficultyLevel` 属性应设置为空字串。

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
