---
id: 66fbcf750a62784cf11f5634
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

是时候测试每个函数了。你可以使用任何配方来完成此操作，但本教程将从 `recipe1` 开始。

创建两个新变量：`recipe1TotalIngredients` 和 `recipe1DifficultyLevel`。通过调用对应的函数并传入适当的 `recipe1` 属性，为它们赋值。

最后，将每个变量记录到控制台以查看结果。

# --hints--

你应该创建一个 `recipe1TotalIngredients` 变量。

```js
assert.isNotNull(recipe1TotalIngredients);
```

你的 `recipe1TotalIngredients` 变量应设置为 `getTotalIngredients(recipe1.ingredients)`。

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

你应该将 `recipe1TotalIngredients` 记录到控制台日志。

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

你应该创建一个 `recipe1DifficultyLevel` 变量。

```js
assert.isNotNull(recipe1DifficultyLevel);
```

你的 `recipe1DifficultyLevel` 变量应设置为 `getDifficultyLevel(recipe1.cookingTime)`。

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

你应该将 `recipe1DifficultyLevel` 记录到控制台日志。

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
