---
id: 66fbcf750a62784cf11f5630
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

你现在应该将这三个对象推入 `recipes` 数组。为此，你可以使用 `push()` 方法。

使用 `push()` 方法将所有配方对象推入 `recipes` 数组。确保按顺序推入 `recipe1`、`recipe2` 和 `recipe3`。

还要删除 `recipe1Name`、`recipe2CookingTime` 和 `recipe3Ingredients` 变量，以及记录这些变量的 `console.log` 语句。

# --hints--

你应该将所有配方对象推入 `recipes` 数组。确保按照它们声明的顺序推入。

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

你应该移除 `recipe1Name` 变量。

```js
assert.throws(() => recipe1Name, ReferenceError);
```

你应该移除 `console.log(recipe1Name);` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

你应该移除 `recipe2CookingTime` 变量。

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

你应该删除 `console.log(recipe2CookingTime);` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

你应该删除 `recipe3Ingredients` 变量。

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

你应该删除 `console.log(recipe3Ingredients);` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
