---
id: 66fbcf750a62784cf11f562d
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

创建一个名为 `recipe1` 的对象。在 `recipe1` 对象内，创建一个值为 `"Spaghetti Carbonara"` 的 `name` 属性。

同样在 `recipe1` 对象内，创建一个 `ingredients` 属性，其值为一个数组。该数组应包含 `spaghetti`、`Parmesan cheese`、`pancetta` 和 `black pepper`。

# --hints--

你应该创建一个名为 `recipe1` 的对象。

```js
assert.isObject(recipe1);
```

你的 `recipe1` 对象应具有一个 `name` 属性。

```js
assert.property(recipe1, 'name');
```

你的 `name` 属性应设置为 `Spaghetti Carbonara`。

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

你的 `recipe1` 对象应具有一个 `ingredients` 属性。

```js
assert.property(recipe1, 'ingredients');
```

你的 `ingredients` 属性应具有数组值。

```js
assert.isArray(recipe1.ingredients);
```

你的 `ingredients` 属性的数组值应包含 `spaghetti`、`Parmesan cheese`、`pancetta` 和 `black pepper`。

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
