---
id: 66fbcf750a62784cf11f5635
title: 步骤11
challengeType: 1
dashedName: step-11
---

# --description--

你现在可以为 `recipes` 数组的每一项填写 `totalIngredients` 和 `difficultyLevel` 属性的值。

现在，访问 `recipe1` 的 `totalIngredients` 和 `difficultyLevel`，并将它们设置为函数调用和参数的相应结果。

# --hints--

你应该访问 `recipe1` 的 `totalIngredients` 属性。

```js
assert.isNotNull(recipe1.totalIngredients);
```

你应该将调用 `getTotalIngredients` 并传入 `recipe1.ingredients` 的结果赋值给 `recipe1` 的 `totalIngredients` 属性。

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

你应该访问 `recipe1` 的 `difficultyLevel` 属性。

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

你应该将调用 `getDifficultyLevel` 并传入 `recipe1.cookingTime` 的结果赋值给 `recipe1` 的 `cookingTime` 属性。

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
