---
id: 66fbcf750a62784cf11f5633
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

创建一个 `getDifficultyLevel` 函数，该函数以表示烹饪时间的数字作为参数。

如果烹饪时间小于或相等于 `30`，函数应返回 `"easy"`。如果小于或相等于 `60`，函数应返回 `"medium"`。否则，函数应返回 `"hard"`。

# --hints--

你应该创建一个 `getDifficultyLevel` 函数。

```js
assert.isFunction(getDifficultyLevel);
```

你的 `getDifficultyLevel` 函数应只有一个参数。

```js
assert.lengthOf(getDifficultyLevel, 1);
```

当烹饪时间小于或相等于 `30` 时，你的 `getDifficultyLevel` 函数应返回 `"easy"`。

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

当烹饪时间大于 `31` 且小于或相等于 `60` 时，你的 `getDifficultyLevel` 函数应返回 `"medium"`。

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

当烹饪时间大于 `60` 时，你的 `getDifficultyLevel` 函数应返回 `"hard"`。

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
