---
id: 66ff02792b3042a2b811f475
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

为你的 `recipe1` 对象添加以下属性：

| Key | Value   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

属性值为 `null` 和空字串的属性将在你计算后更新。

# --hints--

你的 `recipe1` 对象应有一个 `cookingTime` 属性。

```js
assert.property(recipe1, "cookingTime");
```

你的 `cookingTime` 属性应设置为 `22`。

```js
assert.equal(recipe1.cookingTime, 22);
```

你的 `recipe1` 对象应具有 `totalIngredients` 属性。

```js
assert.property(recipe1, "totalIngredients");
```

你的 `totalIngredients` 属性应设置为 `null`。

```js
assert.isNull(recipe1.totalIngredients);
```

你的 `recipe1` 对象应具有 `difficultyLevel` 属性。

```js
assert.property(recipe1, "difficultyLevel");
```

你的 `difficultyLevel` 属性应设置为空字串。

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
