---
id: 670e4f45f7116c0f216a5177
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

在继续之前，你应该练习如何从一个对象中访问属性。

你可以使用点号（`.`）或方括弧（`[]`）表示法来实现这一点。以下是一个示例：

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

访问 `recipe1` 的 `name` 属性，并将其赋值给变量 `recipe1Name`。

接下来，访问 `recipe2` 的 `cookingTime` 属性并将其赋值给变量 `recipe2CookingTime`。

最后，访问 `recipe3` 的 `ingredients` 属性并将其赋值给变量 `recipe3Ingredients`。

确保所有你创建的变量都已记录到控制台日志中。

# --hints--

你应该创建一个 `recipe1Name` 变量。

```js
assert.isDefined(recipe1Name);
```

你应该将 `recipe1` 的 `name` 属性的值赋值给你的 `recipe1name` 变量。

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

你应该将 `recipe1Name` 记录到控制台日志。

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

你应该创建一个 `recipe2CookingTime` 变量。

```js
assert.isDefined(recipe2CookingTime);
```

你应该将 `recipe2` 的 `cookingTime` 属性的值赋值给你的 `recipe2CookingTime` 变量。

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

你应该将 `recipe2CookingTime` 记录到控制台日志。

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

你应该创建一个 `recipe3Ingredients` 变量。

```js
assert.isDefined(recipe3Ingredients);
```

你应该将 `recipe3` 的 `ingredients` 属性的值赋值给你的 `recipe3Ingredients` 变量。

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

你应该将 `recipe3Ingredients` 记录到控制台日志。

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
