---
id: 69161f4c0c9c8c48b3c843c5
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

创建一个具有以下属性和值的 `recipe3` 对象：

|Key|Value|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

你应该创建一个名为 `recipe3` 的对象。

```js
assert.isObject(recipe3);
```

你的 `recipe3` 对象应具有一个 `name` 属性。

```js
assert.property(recipe3, "name");
```

你的 `name` 属性应设置为 `Vegetable Stir Fry`。

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

你的 `recipe3` 对象应有一个 `ingredients` 属性。

```js
assert.property(recipe3, "ingredients");
```

你的 `ingredients` 属性应具有数组值。

```js
assert.isArray(recipe3.ingredients);
```

你的 `ingredients` 属性的数组值应包含 `broccoli`、`carrot` 和 `bell pepper`。

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

你的 `recipe3` 对象应有一个 `cookingTime` 属性。

```js
assert.property(recipe3, "cookingTime");
```

你的 `cookingTime` 属性值应为数字。

```js
assert.isNumber(recipe3.cookingTime);
```

你的 `cookingTime` 属性应设置为 `15`。

```js
assert.equal(recipe3.cookingTime, 15);
```

你的 `recipe3` 对象应具有 `totalIngredients` 属性。

```js
assert.property(recipe3, "totalIngredients");
```

你的 `totalIngredients` 属性应设置为 `null`。

```js
assert.isNull(recipe3.totalIngredients);
```

你的 `recipe3` 对象应具有 `difficultyLevel` 属性。

```js
assert.property(recipe3, "difficultyLevel");
```

你的 `difficultyLevel` 属性应设置为空字串。

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
