---
id: 673b567e3ba535dda140d278
title: Costruire un sistema di votazione
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

In questo laboratorio costruirai un sistema di votazione che usa `Map` per creare un sondaggio e `Set` per evitare votazioni duplicate.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti inizializzare una variabile `poll` a un nuovo oggetto `Map`.

2. Dovresti avere una funzione `addOption` che accetta un parametro `option`.

3. Nella funzione `addOption`:

   - Se `option` non esiste già nel sondaggio, dovrebbe essere aggiunto al sondaggio con un `Set` vuoto come suo valore per tracciare gli elettori. Dovresti anche restituire il messaggio `Option "<option>" added to the poll.`

   - Se `option` esiste già, dovrebbe restituire il messaggio `Option "<option>" already exists.`

   - Se provi ad aggiungere un'opzione vuota, la funzione dovrebbe restituire il messaggio `Option cannot be empty.`

4. Dovresti avere una funzione `vote` che accetta due parametri, `option` (l'opzione per cui votare) e `voterId` (un ID unico per l'elettore).

5. Nella funzione `vote`:

   - Se `option` non esiste nel sondaggio, la funzione dovrebbe restituire il messaggio `Option "<option>" does not exist.`

   - Se `option` esiste, la funzione dovrebbe controllare se `voterId` ha già votato per questo `option`.

   - Se l'elettore ha già votato, la funzione dovrebbe restituire il messaggio `Voter <voterId> has already voted for "<option>".`

   - Se l'elettore non ha votato, `voterId` dovrebbe essere aggiunto al `Set` di `voters` per questa opzione. La funzione dovrebbe restituire il messaggio `Voter <voterId> voted for "<option>".`

6. Dovresti avere almeno tre opzioni nel tuo `poll`.

7. Il tuo `poll` dovrebbe avere almeno tre voti.

8. Dovresti avere una funzione `displayResults` che restituisce i risultati del sondaggio nel formato seguente:

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

Dovresti avere una variabile `poll` inizializzata a un nuovo oggetto `Map`.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

Dovresti definire una funzione `addOption` che accetta un parametro `option`.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

Dovresti definire una funzione `vote` che accetta due parametri, `option` e `voterId`.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

Dovresti definire una funzione `displayResults` per mostrare i risultati del sondaggio.

```js
assert.isFunction(displayResults);
```

Dovresti avere almeno tre opzioni nel tuo `poll`.

```js
assert.isAtLeast(poll.size, 3);
```

Il tuo `poll` dovrebbe avere almeno tre voti.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

Dovresti assicurarti che ogni opzione di voto corrisponda a un oggetto `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` dovrebbe restituire `Option "Egypt" added to the poll.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Aggiungere un'opzione vuota dovrebbe restituire `"Option cannot be empty."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

Quando `Turkey` è già aggiunto, `addOption("Turkey")` dovrebbe restituire `Option "Turkey" already exists.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quando `Malaysia` esiste nelle opzioni di voto, `vote("Malaysia", "traveler1")` dovrebbe restituire `Voter traveler1 voted for "Malaysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` dovrebbe aggiornare il `Set` degli elettori per un'opzione.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quando `traveler1` prova a votare di nuovo per `Algeria`, `vote("Algeria", "traveler1")` dovrebbe restituire `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

I voti duplicati non dovrebbero aumentare la dimensione del `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quando `Nigeria` non è nelle opzioni di voto, `vote("Nigeria", "traveler2")` dovrebbe restituire `Option "Nigeria" does not exist.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Un'opzione unica dovrebbe poter ricevere voti multipli.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Quando l'opzione `Turkey` riceve due voti, l'opzione `Morocco` uno, e l'opzione `Spain` nessun voto, `displayResults()` dovrebbe restituire `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` dovrebbe restituire i risultati nel formato corretto.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
