---
id: 587d78af367417b2b2512b04
title: Створіть посадкову сторінку продукту
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301144
dashedName: build-a-product-landing-page
---

# --description--

**Мета:** Створіть застосунок, який функціонально схожий на <a href="https://product-landing-page.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://product-landing-page.freecodecamp.rocks</a>. **Не копіюйте цей демонстраційний проєкт**.

**Історії користувача:**

1. Ваша посадкова сторінка продукту повинна мати `header` елемент із відповідним `id="header"`
1. Ви повинні бачити зображення всередині `header` елемента з відповідним `id="header-img"` (логотип буде гарним зображенням тут)
1. Усередині `#header` елемента ви повинні бачити `nav` елемент із відповідним `id="nav-bar"`
1. Ви повинні бачити принаймні три активні елементи всередині `nav` елемента, кожен із класом `nav-link`
1. Коли ви натискаєте на `.nav-link` кнопку в `nav` елементі, вас направляють до відповідного розділу посадкової сторінки
1. Ви можете переглядати вставлене відео продукту з `id="video"`
1. Ваша посадкова сторінка має `form` елемент із відповідним `id="form"`
1. Усередині форми є поле введення `input` з `id="email"`, де можна ввести адресу е-пошти
1. Поле введення `#email` має містити текст-заповнювач, щоб користувачі знали, для чого це поле
1. Поле введення `#email` використовує HTML5 валідацію, щоб підтвердити, що введений текст є адресою е-пошти
1. Усередині форми є кнопка відправлення `input` із відповідним `id="submit"`
1. Коли ви натискаєте на `#submit` елемент, е-пошта відправляється на статичну сторінку (використовуйте цей тестовий URL: `https://www.freecodecamp.com/email-submit`)
1. Панель навігації завжди має бути зверху вікна перегляду
1. Ваша посадкова сторінка продукту повинна мати принаймні один медіазапит
1. Ваша посадкова сторінка продукту повинна використовувати CSS гнучкий контейнер принаймні один раз

Виконайте історії користувача та пройдіть усі тести нижче, щоб завершити цей проєкт. Додайте свій особистий стиль. Успішного кодування!

**Примітка:** Обов’язково додайте `<link rel="stylesheet" href="styles.css">` у ваш HTML, щоб направити на ваш файл стилів і застосувати CSS

# --hints--

Ви повинні мати `header` елемент із `id` значенням `header`.

```js
const el = document.getElementById('header')
assert(!!el && el.tagName === 'HEADER')
```

Ви повинні мати `img` елемент із `id` значенням `header-img`.

```js
const el = document.getElementById('header-img')
assert(!!el && el.tagName === 'IMG')
```

Ваш `#header-img` має бути нащадком `#header`.

```js
const els = document.querySelectorAll('#header #header-img')
assert(els.length > 0)
```

Ваш `#header-img` має мати атрибут `src`.

```js
const el = document.getElementById('header-img')
assert(!!el && !!el.src)
```

Значення `#header-img` у `src` має бути дійсним URL (починається з `http`).

```js
const el = document.getElementById('header-img')
assert(!!el && /^http/.test(el.src))
```

Ви повинні мати `nav` елемент із `id` значенням `nav-bar`.

```js
const el = document.getElementById('nav-bar')
assert(!!el && el.tagName === 'NAV')
```

Ваш `#nav-bar` має бути нащадком `#header`.

```js
const els = document.querySelectorAll('#header #nav-bar')
assert(els.length > 0)
```

Ви повинні мати принаймні 3 `.nav-link` елементи всередині `#nav-bar`.

```js
const els = document.querySelectorAll('#nav-bar .nav-link')
assert(els.length >= 3)
```

Кожен `.nav-link` елемент має мати атрибут `href`.

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  if (!el.href) assert(false)
})
assert(els.length > 0)
```

Кожен `.nav-link` елемент має направляти на відповідний елемент на посадковій сторінці (має `href` з значенням ідентифікатора іншого елемента, наприклад `#footer`).

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  const linkDestination = el.getAttribute('href').slice(1)
  if (!document.getElementById(linkDestination)) assert(false)
})
assert(els.length > 0)
```

Ви повинні мати `video` або `iframe` елемент із `id` значенням `video`.

```js
const el = document.getElementById('video')
assert(!!el && (el.tagName === 'VIDEO' || el.tagName === 'IFRAME'))
```

Ваш `#video` має мати атрибут `src`.

```js
let el = document.getElementById('video')
const sourceNode = el.children;
let sourceElement = null;
if (sourceNode.length) {
  sourceElement = [...video.children].filter(el => el.localName === 'source')[0];
}
if (sourceElement) {
  el = sourceElement;
}
assert(el.hasAttribute('src'));
```

Ви повинні мати `form` елемент із `id` значенням `form`.

```js
const el = document.getElementById('form')
assert(!!el && el.tagName === 'FORM')
```

Ви повинні мати `input` елемент із `id` значенням `email`.

```js
const el = document.getElementById('email')
assert(!!el && el.tagName === 'INPUT')
```

Ваш `#email` має бути нащадком `#form`.

```js
const els = document.querySelectorAll('#form #email')
assert(els.length > 0)
```

Ваш `#email` має атрибут `placeholder` із текстом-заповнювачем.

```js
const el = document.getElementById('email')
assert(!!el && !!el.placeholder && el.placeholder.length > 0)
```

Ваш `#email` має використовувати HTML5 валідацію, встановивши `type` у `email`.

```js
const el = document.getElementById('email')
assert(!!el && el.type === 'email')
```

Ви повинні мати `input` елемент із `id` значенням `submit`.

```js
const el = document.getElementById('submit')
assert(!!el && el.tagName === 'INPUT')
```

Ваш `#submit` має бути нащадком `#form`.

```js
const els = document.querySelectorAll('#form #submit')
assert(els.length > 0)
```

Ваш `#submit` має `type` значення `submit`.

```js
const el = document.getElementById('submit')
assert(!!el && el.type === 'submit')
```

Ваш `#form` має атрибут `action` зі значенням `https://www.freecodecamp.com/email-submit`.

```js
const el = document.getElementById('form')
assert(!!el && el.action === 'https://www.freecodecamp.com/email-submit')
```

Ваш `#email` має атрибут `name` зі значенням `email`.

```js
const el = document.getElementById('email')
assert(!!el && el.name === 'email')
```

Ваш `#nav-bar` завжди має бути зверху вікна перегляду.

```js
  const timeout = (milliseconds) => new Promise((resolve) => setTimeout(resolve, milliseconds));

  const header = document.getElementById('header');
  const headerChildren = header.children;
  const navbarCandidates = [header, ...headerChildren];
  
  // Return smallest top position of all navbar candidates
  const getNavbarPosition = (candidates = []) => {
    return candidates.reduce(
      (min, candidate) =>
        Math.min(min, Math.abs(candidate?.getBoundingClientRect().top)),
      Infinity
    );
  };
  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the viewport '
  );

  window.scroll(0, 500);
  await timeout(1);

  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the ' +
      'viewport even after scrolling '
  );
    
  window.scroll(0, 0);
```

Ваша посадкова сторінка продукту має використовувати принаймні один медіазапит.

```js
const htmlSourceAttr = Array.from(document.querySelectorAll('source')).map(el => el.getAttribute('media'))
const cssCheck = new __helpers.CSSHelp(document).getCSSRules('media')
assert(cssCheck.length > 0 || htmlSourceAttr.length > 0);
```

Ваша посадкова сторінка продукту має використовувати CSS гнучкий контейнер принаймні один раз.

```js
const hasFlex = (rule) => ["flex", "inline-flex"].includes(rule.style?.display)
const stylesheet = new __helpers.CSSHelp(document).getStyleSheet()
const cssRules = new __helpers.CSSHelp(document).styleSheetToCssRulesArray(stylesheet)
const mediaRules = new __helpers.CSSHelp(document).getCSSRules('media')
const usesFlex = cssRules.find(rule => hasFlex(rule))
const usesFlexMedia = mediaRules.find(mediaRule => {
  return [...mediaRule.cssRules].find(rule => hasFlex(rule))
})
assert(usesFlex || usesFlexMedia)
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Product Landing Page</title>
  </head>
  <body>
    <header id="header">
      <nav id="nav-bar">
        <img
          id="header-img"
          src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
          max-height="50px"
        />
        <a href="#Features" class="nav-link">Features</a> |
        <a href="#Video" class="nav-link">See our facility!</a> |
        <a href="#Pricing" class="nav-link">Purchase</a>
        <hr />
      </nav>
    </header>
    <main>
      <h1>
        Pokemon Daycare Service
      </h1>
      <section id="Features">
        <h2>What we offer</h2>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">Guaranteed friendly and loving staff!</div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Comfortable environment for Pokemon to explore and play!
          </div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Multiple membership plans to fit your lifestyle!
          </div>
        </div>
      </section>
      <section id="Video">
        <h2>Check us out!</h2>
        A sneak peek into our facility:
        <br />
        <iframe
          id="video"
          width="520"
          height="347"
          src="https://www.youtube.com/embed/Nw-ksH2r6AQ"
          frameborder="0"
          allowfullscreen
          alt="A video tour of our facility"
        >
        </iframe>
      </section>
      <section id="Pricing">
        <h2>Membership Plans</h2>
        <div class="flex-mem">
          <div class="flex-mem-box">
            <font size="+2">Basic Membership</font><br />
            <ul>
              <li>One Pokemon</li>
              <li>Food and berries provided</li>
            </ul>
            <em>$9.99/month</em>
          </div>
          <div class="flex-mem-box">
            <font size="+2">Silver Membership</font><br />
            <ul>
              <li>Up to Three Pokemon</li>
              <li>Food and berries provided</li>
              <li>Grooming and accessories included</li>
            </ul>
            <em>$19.99/month</em>
          </div>
          <div class="flex-mem-box">
            <font size="+2">Gold Membership</font><br />
            <ul>
              <li>Up to six Pokemon!</li>
              <li>Food and berries provided</li>
              <li>Grooming and accessories included</li>
              <li>Personal training for each Pokemon</li>
              <li>Breeding and egg hatching</li>
            </ul>
            <em>$29.99/month</em>
          </div>
        </div>
      </section>
      <form id="form" action="https://www.freecodecamp.com/email-submit">
        <p>Sign up for our newsletter!</p>
        <label for="email"><p>Email:</p><input name="email" id="email" type="email" placeholder="johnsmith@email.com" required></label>
        <input type="submit" id="submit">
      </form>
      <footer>
        <a href="../">Return to Project List</a> |
        <a href="https://www.nhcarrigan.com">Return to HomePage</a>
      </footer>
    </main>
  </body>
</html>
```

```css
body {
  background-color: #3a3240;
  color: white;
}
main {
  max-width: 750px;
  margin: 50px auto;
}
input {
  background-color: #92869c;
}
a:not(.nav-link) {
  color: white;
}
#header-img {
  max-height: 25px;
}
#nav-bar {
  position: fixed;
  width: 100%;
  text-align: center;
  top: 0%;
  background-color: #92869c;
}
h1 {
  text-align: center;
}
body {
  text-align: center;
}
footer {
  text-align: center;
}
#bullet {
  max-height: 25px;
}
.flex-here {
  display: flex;
  justify-content: center;
}
.flex-left {
  height: 25px;
}
.flex-mem {
  display: flex;
  justify-content: center;
}
.flex-mem-box {
  background-color: #92869c;
  border-color: black;
  border-width: 5px;
  border-style: solid;
  margin: 10px;
  padding: 10px;
  color: black;
}
@media (max-width: 350px) {
  #video {
    width: 300;
    height: 200;
  }
}
```
