---
id: 5d5a813321b9e3db6c106a46
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript es un poderoso lenguaje que le permite construir sitios web que sean interactivos. 

*Nota*: Para todos los proyectos restantes en este plan de estudios, necesitarás un nivel básico de conocimiento en HTML y CSS. Si eres nuevo en HTML y CSS, por favor revisa la [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

Para comenzar, crea una plantilla standard de HTML con `DOCTYPE`, `html`, `head` y `body`, añadiendo a continuación un tag `meta` para `charset`. Añade un elemento `title` con el texto `RPG - Dragon Repeller`. Incluye un tag `link` para la hoja de estilos, enlazando el fichero `styles.css`.

Por último, anidado en el `body`, crea un elemento `div` con el atributo `id` establecido a `game`.

# --hints--

Su código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio despues de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debe definir el tipo de documento para que sea `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tu `html` debe tener una etiqueta de apertura. No olvides el atributo `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Su declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Debes tener solo una etiqueta de apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Debes tener una etiqueta de cierre `head` tag.

```js
assert(code.match(/<\/head\s*>/i));
```

Debes tener una etiqueta de apertura `body`.

```js
assert(code.match(/<body\s*>/i));
```

Debes tener una etiqueta de cierre `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Las etiquetas `head` y `body` deben ser hermanas.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

La etiqueta `head` debe estar dentro de la etiqueta `html`.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

La etiqueta `body` debe estar dentro de la etiqueta `html`.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

Tu código debe tener un elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

El elemento `meta` debe tener un atributo `charset` con el valor `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Debes tener una etiqueta `title` de cierre.

```js
assert(code.match(/<\/title\s*>/i));
```

Su código debe tener un elemento <code>link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Debe tener un elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Tu elemento `div` debe tener un atributo `id` con el valor `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Tu elemento `div` debe estar dentro del elemento `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
