---
id: 62a1166ed9a56d439c0770e7
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Crea tres elementos `span` dentro de tu elemento `#stats`. Dale a cada uno de ellos la clase `stat`. Luego dale al primer uno el texto `XP: 0`, al segundo uno el texto `Health: 100`, y al tercer uno el texto `Gold: 50`.

# --hints--

Debes tener tres elementos `span` dentro del elemento `#stats`.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Debes dar a los tres nuevos elementos `span` una clase `stat`.

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

Tu primer elemento `.stat` debe tener el texto proporcionado `XP: 0`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

Tu segundo elemento `.stat` debe tener el texto proporcionado `Health: 100`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

Tu tercer elemento `.stat` debe tener el texto proporcionado `Gold: 50`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
