---
id: 62a240c67f3dbb1a1e6d95ee
title: Pasos 23
challengeType: 0
dashedName: step-23
---

# --description--

Dale a tu `#game` un ancho máximo de `500px` y un alto máximo de `400px`. Establecer el `background-color` a `#ffffff` y el `color` a `#ffffff`. 

Usa márgenes para centrarlo estableciendo el margen superior en `30px`, el margen inferior en `0px`, y el margen izquierdo y derecho en `auto`. 

Finalmente, dele `10px` de relleno en los cuatro lados.

# --hints--

Debes tener un selector `#game`.

```js
const game = new __helpers.CSSHelp(document).getStyle('#game');
assert.exists(game);
```

Tu selector `#game` debe tener un `max-width` de `500px`.

```js
const maxWidth = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-width');
assert.equal(maxWidth, '500px');
```

Tu selector `#game` debe tener un `max-height` de `400px`.

```js
const maxHeight = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-height');
assert.equal(maxHeight, '400px');
```

Tu `#game` selector debe tener un `Color de fondo` de `#ffffff`.

```js
const background = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('background-color');
assert.equal(background, 'rgb(255, 255, 255)');
```

Tu selector `#game` debe tener un `color` de `#ffffff`.

```js
const color = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('color');
assert.equal(color, 'rgb(255, 255, 255)');
```

Tu selector `#game` debe tener un `margin-top` establecido en `30px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-top');
assert.equal(margin, '30px');
```

Tu selector `#game` debe tener un `margin-left` establecido en `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-left');
assert.equal(margin, 'auto');
```

Tu selector `#game` debe tener un `margin-right` establecido en `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-right');
assert.equal(margin, 'auto');
```


Tu selector `#game` debe tener un `margin-bottom` establecido en `0px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-bottom');
assert.equal(margin, '0px');
```

Tu selector `#game` debe tener `10px` de relleno en todos los lados.

```js
const padding = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('padding');
assert.equal(padding, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
