---
id: 62a23cb9bc467a147516b500
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Para tu `#controls` elemento, crear tres `button` elementos. El primer debe tener el `id` establecido a `button1`, y el texto `Go to store`. El segundo debe tener el `id` establecido a `button2`, y el texto `Go to cave`. El tercer debe tener el `id` establecido a `button3`, y el texto `Fight dragon`.

# --hints--

Debes agregar tres `button` elementos a tu `#controls` elemento.

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

Tu primer botón debe tener el `id` establecido a `button1`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

Tu primer botón debe tener el texto `Go to store`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

Tu segundo botón debe tener el `id` establecido a `button2`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

Tu segundo botón debe tener el texto `Go to cave`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

Tu tercer botón debe tener el `id` establecido a `button3`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

Tu tercer botón debe tener el texto `Fight dragon`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
--fcc-editable-region--
      <div id="controls">

      </div>
--fcc-editable-region--
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
