---
id: 62a255dae245b52317da824a
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Una de las herramientas más poderosas es la consola de desarrollador. Dependiendo de tu navegador, esta puede abrirse pulsando `F12` o `Ctrl+Shift+I`. En Mac, puedes presionar `Opción + ⌘ + C` y seleccionar "Console". También puede hacer clic en el botón "Consola" sobre la ventana de vista previa para ver nuestra consola incorporada. 

La consola del desarrollador incluirá errores producidos por tu código, pero también puede usarlo para ver los valores de las variables en su código, lo que es útil para la depuración.

Añade una línea `console.log("Hello World");` dentro de las etiquetas `script`. Luego, haz clic en el botón "Consola" para abrir la consola JavaScript. Debes ver el texto `"Hello World"`.


# --hints--

Debes añadir una línea `console.log("Hello World");` a tu código. 

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

La línea `console.log("Hola Mundo");` debe estar entre las etiquetas `script`.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
