---
id: 62a3a7e4f1060e2fc5ffb34b
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Crea otra variable llamada `currentWeaponIndex` con el valor `0`.

# --hints--

Debes usar `let` para declarar una variable llamada `currentWeaponIndex`.

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

Debes escribir el nombre de la variable usando el criterio camelCase.

```js
assert.match(code, /currentWeaponIndex/);
```

Tu variable `currentWeaponIndex` debe tener el valor `0`.

```js
assert.equal(currentWeaponIndex, 0);
```

Debes inicializar la variable a `0`.

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
