---
id: 62a3b3eab50e193608c19fc6
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Declara una variable llamada `fighting`, pero no la inicialice con un valor. 

# --hints--

Debes usar `let` para declarar una variable `fighting`.

```js
assert.match(code, /let\s+fighting/);
```

Tu variable `fighting` no debe tener un valor.

```js
assert.isUndefined(fighting);
```

No debes asignarle un valor a tu variable `fighting`. 

```js
assert.match(code, /let\s+fighting\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
--fcc-editable-region--
```
