---
id: 62a39f5a5790eb27c1e5d4bf
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Tu vista ha sido cambiada a su nuevo archivo `script.js`. Recuerde que puede utilizar las pestañas de arriba para cambiar entre archivos.

Agrega su línea `console.log("Hello World");` a este archivo y véala aparecer en su consola.

# --hints--

Debe tener una línea `console.log("Hello World");` en su código.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
