---
id: 62a3a0a3c0a4b32915d26a6e
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Elimina la línea `console.log("Hello World");` para comenzar a escribir el código del proyecto.

En el proyecto anterior, aprendiste a trabajar con variables y la palabra reservada `let`, de esta manera:

```js
let camperbot;
```

También aprendiste cómo inicializar una variable con un valor, de esta manera:

```js
let age = 32;
```

Utiliza la palabra reservada `let` para declarar una variable llamada `xp` y asígnale el valor `0`, un número.

# --hints--

No debe tener una línea `console.log("Hello World");` en el código.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Debes usar la palabra reservada `let` para declarar tu variable.

```js
assert.match(code, /let/);
```

Debes declarar una variable llamada `xp`.

```js
assert.match(code, /let\s+xp/);
```

`xp` debe tener un valor de `0`.

```js
assert.equal(xp, 0);
```

Debes inicializar la variable `xp` a `0`. 

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
