---
id: 62a8c2bbbd8aa82052f47c53
title: Paso 71
challengeType: 0
dashedName: step-71
---

# --description--

Agrega un tercer objeto al arreglo `locations`. Dale las mismas propiedades que los otros dos objetos.

Establece `name` a `cave`. Establece `button text` a un arreglo con las cadenas `"Fight slime"`, `"Fight fanged beast"`, y `"Go to town square"`. Establece los `"button functions"` a un arreglo con las variables `fightSlime`, `fightBeast`, y `goTown`. Establece la propiedad <code>text` a `"You enter the cave. You see some monsters."</code>.

# --hints--

Debes tener tres valores en tu arreglo `locations`.

```js
assert.lengthOf(locations, 3);
```

Tu tercer valor de `locations` debe ser un objeto.

```js
assert.isObject(locations[2]);
```

Tu tercer objeto de `locations` debe tener una propiedad `name` con el valor de `"cave"`.

```js
assert.equal(locations[2].name, "cave");
```

Tu tercer objeto de `locations` debe tener una propiedad `"button text"` que sea un arreglo.

```js
assert.isArray(locations[2]["button text"]);
```

Tu propiedad `"button text"` debe tener los valores de cadena `"Fight slime"`, `"Fight fanged beast"`, y `"Go to town square"`.

```js
assert.equal(locations[2]["button text"][0], "Fight slime");
assert.equal(locations[2]["button text"][1], "Fight fanged beast");
assert.equal(locations[2]["button text"][2], "Go to town square");
```

Tu tercer objeto de `locations` debe tener una propiedad `"button functions"` que sea un arreglo.

```js
assert.isArray(locations[2]["button functions"]);
```

Tu propiedad `"button functions"` debe tener los valores de función `fightSlime`, `fightBeast`, y `goTown`.

```js
assert.equal(locations[2]["button functions"][0], fightSlime);
assert.equal(locations[2]["button functions"][1], fightBeast);
assert.equal(locations[2]["button functions"][2], goTown);
```

Tu tercer objeto de `locations` debe tener una propiedad `text` que sea una cadena.

```js
assert.isString(locations[2].text);
```

Tu tercer objeto de `locations` debe tener una propiedad `text` con el valor de `"You enter the cave. You see some monsters."`.

```js
assert.equal(locations[2].text, "You enter the cave. You see some monsters.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
--fcc-editable-region--

--fcc-editable-region--
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}

function buyHealth() {

}

function buyWeapon() {

}

function fightSlime() {

}

function fightBeast() {

}
```
