---
id: 66d6a7a3e1aa411e94bf2346
title: بناء لوحة منزل
challengeType: 25
dashedName: build-a-house-painting
demoType: onClick
---

# --description--

في هذا المختبر، ستستخدم HTML لإنشاء هيكل منزل. ثم ستستخدم تحديد المواقع في CSS لترتيب عناصر منزلك مثل النوافذ والأبواب.

**الهدف:** حقق قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يحتوي الكود على عنصر `div` مع `id` قيمته `house`.
1. يجب أن يكون لـ `#house` خاصية `position` مضبوطة على `relative` حتى يمكن وضع العناصر الفرعية بالنسبة له.
1. يجب أن يكون عرض `#house` يساوي `500px` وارتفاعه يساوي `400px`. لا تضبط أي قيم للحد الأدنى أو الحد الأقصى.
1. يجب تعيين لون خلفية وحدود لـ `#house`.
1. يجب أن يحتوي `#house` على خمسة عناصر `div` بداخلها مع قيم `id` التالية: `chimney`، `roof`، `window-1`، `window-2`، و`door`.
1. يجب أن يكون لكل العناصر الفرعية المباشرة لـ `#house` خاصية `position` بقيمة `absolute`.
1. يجب أن تحتوي العناصر `#roof`، `#chimney`، `#window-1`، `#window-2`، و`#door` على عرض وارتفاع وحدود ولون خلفية محددة.
1. يجب أن يكون لـ `#roof` خاصية top بقيمة `0`.
1. يجب وضع `#door` في أسفل المنزل.
1. يجب وضع النوافذ أسفل `#roof` وعلى الأقل أعلى من ثلث ارتفاع `#door`.
1. يجب أن تحتوي كل من النوافذ والباب على خاصية `left` أو `right` بقيمة تضعها داخل حدود المنزل.
1. يجب أن يكون لـ `#chimney` قيمة top تضعه في أعلى `#house`.
1. يجب أن يكون لـ `#chimney` خاصية `z-index` تضعه خلف المنزل.

**ملاحظة:** تأكد من ربط ورقة الأنماط الخاصة بك في ملف HTML لتطبيق CSS.

# --hints--

يجب أن يكون لديك عنصر `div` مع `id` بقيمة `house`.

```js
assert.exists(document.querySelector("div#house"));
```

يجب استهداف `#house` وضبط خاصية `position` على `relative`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("position"), "relative");
```

يجب استهداف `#house` وضبط خاصية `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("border"));
```

يجب استهداف `#house` وضبط خاصية `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("background-color"));
```

يجب ألا يحتوي `#house` على خاصية `min-height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-height"));
```

يجب ألا يحتوي `#house` على خاصية `min-width`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-width"));
```

يجب ألا يحتوي `#house` على خاصية `max-height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-height"));
```

يجب ألا يحتوي `#house` على خاصية `max-width`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-width"));
```

يجب استهداف `#house` وضبط عرضه على `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("width"), "500px")
```

يجب استهداف `#house` وضبط ارتفاعه على `400px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("height"), "400px")
```

يجب أن يكون لديك عنصر `div` يحمل `id` باسم `chimney` داخل `#house`.

```js
assert.exists(document.querySelector("div#house div#chimney"));
```

يجب أن تكون خاصية `position` لـ `#chimney` مضبوطة على `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("position"), "absolute");
```

يجب أن تستهدف `#chimney` وتحدد خاصية `width` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("width"));
```

يجب أن تستهدف `#chimney` وتحدد خاصية `height` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("height"));
```

يجب أن تستهدف `#chimney` وتضبط خاصية `border` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("border"));
```

يجب استهداف `#chimney` وتعيين خاصية `background-color` له.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("background-color"));
```

يجب أن يكون لديك `div` يحمل `id` باسم `roof` داخل `#house`.

```js
assert.exists(document.querySelector("div#house  div#roof"));
```

يجب أن تكون خاصية `position` لـ `#roof` مضبوطة على `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("position"), "absolute");
```

يجب أن تستهدف `#roof` وتحدد خاصية `width` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("width"));
```

يجب أن تستهدف `#roof` وتحدد خاصية `height` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("height"));
```

يجب أن تستهدف `#roof` وتضبط خاصية `border` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("border"));
```

يجب أن تستهدف `#roof` وتضبط خاصية `background-color` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("background-color"));
```

يجب أن يكون لديك `div` يحمل `id` باسم `window-1` داخل `#house`.

```js
assert.exists(document.querySelector("div#house  div#window-1"));
```

يجب أن تكون خاصية `position` لـ `#window-1` مضبوطة على `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#window-1"))?.getPropertyValue("position"), "absolute");
```

يجب أن تستهدف `#window-1` وتحدد خاصية `width` الخاصة به.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window1 || windows);
```

يجب أن تستهدف `#window-1` وتضبط خاصية `height`.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window1 || windows);
```

يجب أن تستهدف `#window-1` وتضبط خاصية `border` الخاصة به.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window1 || windows);
```

يجب أن تستهدف `#window-1` وتضبط خاصية `background-color`.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window1 || windows);
```

يجب أن يكون لديك `div` يحمل `id` باسم `window-2` داخل `#house`.

```js
assert.exists(document.querySelector("div#house  div#window-2"));
```

يجب أن تكون خاصية `position` لـ `#window-2` مضبوطة على `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#window-2"))?.getPropertyValue("position"), "absolute");
```

يجب أن تستهدف `#window-2` وتحدد خاصية `width` الخاصة به.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window2 || windows);
```

يجب أن تستهدف `#window-2` وتحدد خاصية `height` الخاصة به.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window2 || windows);
```

يجب أن تستهدف `#window-2` وتضبط خاصية `border` الخاصة به.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window2 || windows);
```

يجب أن تستهدف `#window-2` وتعيّن خاصية `background-color` الخاصة به.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window2 || windows);
```

يجب أن يكون لديك `div` يحمل `id` باسم `door` داخل `#house`.

```js
assert.exists(document.querySelector("div#house  div#door"));
```

يجب أن تكون خاصية `position` لـ `#door` مضبوطة على `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#door"))?.getPropertyValue("position"), "absolute");
```

يجب أن تستهدف `#door` وتحدد خاصية `width` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("width"));
```

يجب أن تستهدف `#door` وتضبط خاصية `height` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("height"));
```

يجب أن تستهدف `#door` وتضبط خاصية `border` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("border"));
```

يجب أن تستهدف `#door` وتضبط خاصية `background-color` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("background-color"));
```

يجب أن تكون خاصية `top` لـ `#roof` مضبوطة على `0`.

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("top"), "0px");
```

يجب وضع `#door` في أسفل المنزل.

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorBottom = Number(door.getPropertyValue("bottom").replace("px", ""));
assert.strictEqual(doorBottom, 0);
```

يجب ضبط خاصية `left` أو `right` على `#door` وترتيبها داخل حدود المنزل.

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorLeft = Number(door.getPropertyValue("left").replace("px", ""));
const doorRight = Number(door.getPropertyValue("right").replace("px", ""));
assert.isTrue(doorLeft > 0 || doorRight > 0);
```

يجب أن تكون خاصية `top` لـ `#chimney` مضبوطة لتضعه في أعلى `#house`.

```js
const chimney = getComputedStyle(document.querySelector("#chimney"));
assert.approximately(parseFloat(chimney?.getPropertyValue("top")), - parseFloat(chimney?.getPropertyValue("height")), 2);
```

يجب أن تكون خاصية `z-index` لـ `#chimney` مضبوطة لتضعه خلف المنزل.

```js
const houseZIndex = getComputedStyle(document.querySelector("#house"))?.getPropertyValue("z-index");
const chimneyZIndex = getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("z-index");
if (houseZIndex === "auto") {
    assert.isBelow(Number(chimneyZIndex), 0)
} else {
    assert.isBelow(Number(chimneyZIndex), Number(houseZIndex));
}
```

يجب وضع النوافذ أسفل `#roof` وعلى ارتفاع لا يقل عن ثلث ارتفاع `#door`.

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const door = getComputedStyle(document.querySelector("#door"))
const roof = getComputedStyle(document.querySelector("#roof"))
const w1Top = Number(window1.getPropertyValue("top").replace("px", ""));
const w1Bottom = Number(window1.getPropertyValue("bottom").replace("px", ""));
const w2Top = Number(window2.getPropertyValue("top").replace("px", ""));
const w2Bottom = Number(window2.getPropertyValue("bottom").replace("px", ""));
const doorHeight = Number(door.getPropertyValue("height").replace("px", ""));
const roofHeight = Number(roof.getPropertyValue("height").replace("px", ""));
assert.isAbove(w1Top, roofHeight);
assert.isAbove(w1Bottom, doorHeight/3);
assert.isAbove(w2Top, roofHeight);
assert.isAbove(w2Bottom, doorHeight/3);
```

يجب ضبط خاصية `left` أو `right` على النوافذ وترتيبها داخل حدود المنزل.

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const w1Left = Number(window1.getPropertyValue("left").replace("px", ""));
const w1Right = Number(window1.getPropertyValue("right").replace("px", ""));
const w2Left = Number(window2.getPropertyValue("left").replace("px", ""));
const w2Right = Number(window2.getPropertyValue("right").replace("px", ""));
assert.isTrue(w1Left > 0 || w1Right > 0);
assert.isTrue(w2Left > 0 || w2Right > 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div id="house">
        <div id="chimney">
          <div id="smoke"></div>
        </div>
        <div id="roof"></div>
        <div id="window-1"></div>
        <div id="window-2"></div>
        <div id="door">
          <div id="door-knob"></div>
        </div>
        <div id="welcome">WELCOME</div>
    </div>
</body>
</html>
```

```css
* {
    box-sizing: border-box;
  }

  body {
    margin: 0;
    padding: 0;
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    background: linear-gradient(
      #b3e6ff,
      #b3e6ff 60%,
      green 60%
    );
  }

  #house {
    flex: 0 0 auto;
    width: 500px;
    height: 400px;
    position: relative;
    background-color: #ff9980;
    border: 6px solid #b35900;
  }

  #chimney {
    position: absolute;
    width: 90px;
    height: 100px;
    top: -100px;
    left: 65%;
    background: repeating-linear-gradient(
      #e6e6e6,
      #e6e6e6 10%,
      black 10%,
      black 12%
    );
    border: 2px solid black;
    z-index: -1;
  }

  #smoke {
    width: 40px;
    height: 40px;
    background-color: rgba(0,0,0,0.3);
    position: absolute;
    top: -20px;
    left: 20px;
    border-radius: 50%;
    box-shadow: 0 -10px 10px 10px rgba(0,0,0,0.3);
    border: none;
  }

  #roof {
    border: 2px;
    width: 100%;
    height: 110px;
    background-color: green;
    position: absolute;
    top: 0;
    background: repeating-linear-gradient(
      45deg,
      #b35900,
      #b35900 2%,
      transparent 2%,
      transparent 5%
    ), repeating-linear-gradient(
      -45deg,
      #b35900,
      #b35900 2%,
      #ff9980 2%,
      #ff9980 5%
    )
  }

  #window-1, #window-2 {
    width: 100px;
    height: 100px;
    background: linear-gradient(
      to right,
      #ffffb3,
      #ffffb3 48%,
      #b35900 48%,
      #b35900 52%,
      #ffffb3 52%,
      #ffffb3 100%
    );
    position: absolute;
    top: 40%;
    border: 6px solid #b35900;
  }

  #window-1 {
    left: 7.5%;
  }

  #window-2 {
    right: 7.5%;
  }

  #door {
    width: 130px;
    height: 180px;
    position: absolute;
    bottom: 0%;
    left: 185px;
    background-color: #e6e6e6;
    border: 6px solid #b35900;
    transform: translateY(6px);
  }

  #door-knob {
    width: 20px;
    height: 20px;
    background-color: #b35900;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    right: 5%;
  }

  #welcome {
    font-family: sans-serif;
    font-weight: bold;
    width: 130px;
    height: 40px;
    position: absolute;
    bottom: -47px;
    left: 178px;
    background-color: #85e085;
    border: 2px solid black;
    transform: skewX(-20deg);
    display: flex;
    align-items: center;
    justify-content: center;
  }
```
