---
id: 56533eb9ac21ba0edf2244b3
title: 將攝氏度轉換為華氏度
challengeType: 1
forumTopicId: 16806
dashedName: convert-celsius-to-fahrenheit
---

# --description--

攝氏轉華氏的公式是攝氏溫度乘以 `9/5`，再加上 `32`。

你會得到一個代表攝氏溫度的變數 `celsius`。使用已定義的變數 `fahrenheit`，並將其指定為與給定攝氏溫度等價的華氏溫度。使用上述提到的公式來幫助將攝氏溫度轉換為華氏溫度。

# --hints--

`convertCtoF(0)` 應該傳回一個數字

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` 應該傳回一個值為 `-22`

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` 應該傳回一個值為 `14`

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` 應該傳回一個值為 `32`

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` 應該傳回一個值為 `68`

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` 應該傳回一個值為 `86`

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js
function convertCtoF(celsius) {
  let fahrenheit;
  return fahrenheit;
}

convertCtoF(30);
```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
