---
id: a302f7aae1aa3152a5b413bc
title: 計算一個數字的階乘
challengeType: 1
forumTopicId: 16013
dashedName: factorialize-a-number
---

# --description--

傳回所提供整數的階乘。

如果整數以字母 `n` 表示，階乘是所有小於或相等於 `n` 的正整數的乘積。

階乘通常以簡寫符號 `n!` 表述。

例如：`5! = 1 * 2 * 3 * 4 * 5 = 120`

只有大於或相等於零的整數會被傳遞給該函式（程式）。

# --hints--

`factorialize(5)` 應該傳回一個數字。

```js
assert.isNumber(factorialize(5));
```

`factorialize(5)` 應該傳回 `120`。

```js
assert.strictEqual(factorialize(5), 120);
```

`factorialize(10)` 應該傳回 `3628800`。

```js
assert.strictEqual(factorialize(10), 3628800);
```

`factorialize(20)` 應該傳回 `2432902008176640000`。

```js
assert.strictEqual(factorialize(20), 2432902008176640000);
```

`factorialize(0)` 應該傳回 `1`。

```js
assert.strictEqual(factorialize(0), 1);
```

# --seed--

## --seed-contents--

```js
function factorialize(num) {
  return num;
}

factorialize(5);
```

# --solutions--

```js
function factorialize(num) {
  return num < 1 ? 1 : num * factorialize(num - 1);
}

factorialize(5);
```
