---
id: a6e40f1041b06c996f7b2406
title: 先到先得
challengeType: 1
forumTopicId: 16016
dashedName: finders-keepers
---

# --description--

創建一個函式（程式），用來檢查陣列 `arr` 並傳回第一個通過「真值測試」的元素。這表示給定一個元素 `x`，如果 `func(x)` 是 `true`，則該元素通過「真值測試」。如果沒有元素通過測試，則傳回 `undefined`。

# --hints--

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` 應該傳回 `8`。

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` 應該傳回 `undefined`。

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js
function findElement(arr, func) {
  let num = 0;
  return num;
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
