---
id: a77dbc43c33f39daa4429b4f
title: 嗚嗚
challengeType: 1
forumTopicId: 16000
dashedName: boo-who
---

# --description--

檢查一個值是否被分類為布林值原始類型。傳回 `true` 或 `false`。

布林值原始值是 `true` 和 `false`。

# --hints--

`booWho(true)` 應該傳回 `true`。

```js
assert.isTrue(booWho(true));
```

`booWho(false)` 應該傳回 `true`。

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` 應該傳回 `false`。

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` 應該傳回 `false`。

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` 應該傳回 `false`。

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` 應該傳回 `false`。

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` 應該傳回 `false`。

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` 應該傳回 `false`。

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` 應該傳回 `false`。

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` 應該傳回 `false`。

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js
function booWho(bool) {
  return bool;
}

booWho(null);
```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

booWho(null);
```
