---
id: a9bd25c716030ec90084d8a1
title: 厚猴子
challengeType: 1
forumTopicId: 16005
dashedName: chunky-monkey
---

# --description--

撰寫一個函式（程式），將陣列（第一個引數）依 `size`（第二個引數）長度拆分成多個群組，並將它們傳回為二維陣列。

# --hints--

`chunkArrayInGroups(["a", "b", "c", "d"], 2)` 應該傳回 `[["a", "b"], ["c", "d"]]`。

```js
assert.deepEqual(chunkArrayInGroups(['a', 'b', 'c', 'd'], 2), [
  ['a', 'b'],
  ['c', 'd']
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3)` 應該傳回 `[[0, 1, 2], [3, 4, 5]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3), [
  [0, 1, 2],
  [3, 4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2)` 應該傳回 `[[0, 1], [2, 3], [4, 5]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2), [
  [0, 1],
  [2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4)` 應該傳回 `[[0, 1, 2, 3], [4, 5]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4), [
  [0, 1, 2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3)` 應該傳回 `[[0, 1, 2], [3, 4, 5], [6]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3), [
  [0, 1, 2],
  [3, 4, 5],
  [6]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4)` 應該傳回 `[[0, 1, 2, 3], [4, 5, 6, 7], [8]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4), [
  [0, 1, 2, 3],
  [4, 5, 6, 7],
  [8]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2)` 應該傳回 `[[0, 1], [2, 3], [4, 5], [6, 7], [8]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2), [
  [0, 1],
  [2, 3],
  [4, 5],
  [6, 7],
  [8]
]);
```

# --seed--

## --seed-contents--

```js
function chunkArrayInGroups(arr, size) {
  return arr;
}

chunkArrayInGroups(['a', 'b', 'c', 'd'], 2);
```

# --solutions--

```js
function chunkArrayInGroups(arr, size) {
  let out = [];

  for (let i = 0; i < arr.length; i += size) {
    out.push(arr.slice(i, i + size));
  }

  return out;
}

chunkArrayInGroups(['a', 'b', 'c', 'd'], 2);
```
