---
id: ab6137d4e35944e21037b769
title: 將句子轉為標題大小寫
challengeType: 1
forumTopicId: 16088
dashedName: title-case-a-sentence
---

# --description--

傳回提供的字串，並將每個字的第一個字母大寫。確保該字的其餘部分為小寫。

為了這個練習的目的，你也應該將連接詞如 `the` 和 `of` 大寫。

# --hints--

`titleCase("I'm a little tea pot")` 應該傳回一個字串。

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` 應該傳回字串 `I'm A Little Tea Pot`。

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` 應該傳回字串 `Short And Stout`。

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` 應該傳回字串 `Here Is My Handle Here Is My Spout`。

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js
function titleCase(str) {
  return str;
}

titleCase("I'm a little tea pot");
```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}

titleCase("I'm a little tea pot");
```
