---
id: acda2fb1324d9b0fa741e6b5
title: 確認結尾
challengeType: 1
forumTopicId: 16006
dashedName: confirm-the-ending
---

# --description--

檢查字串（第一個引數，`str`）是否以指定的標的字串（第二個引數，`target`）結尾。

這個挑戰*可以*用 `.endsWith()` 方法解決，該方法是在 ES2015 中介紹的。但為了這個挑戰的目的，我們希望你改用其中一個 JavaScript 的子字串方法。

# --hints--

`confirmEnding("Bastian", "n")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` 應該傳回 `false`。

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` 應該傳回 `false`。

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` 應該傳回 `false`。

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` 應該傳回 `false`。

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing", "mountain")` 應該傳回 `false`。

```js
assert.isFalse(
  confirmEnding(
    'If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing',
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` 應該傳回 `true`。

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

你的程式碼不應該使用內建方法 `.endsWith()` 來解決此挑戰。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js
function confirmEnding(str, target) {
  return str;
}

confirmEnding('Bastian', 'n');
```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
