---
id: afcc8d540bea9ea2669306b6
title: 重複字串 重複字串
challengeType: 1
forumTopicId: 16041
dashedName: repeat-a-string-repeat-a-string
---

# --description--

重複給定的字串 `str`（第一個引數）共 `num` 次（第二個引數）。如果 `num` 不是正數，則傳回空字串。為了這個挑戰，*不要* 使用內建的 `.repeat()` 方法。

# --hints--

`repeatStringNumTimes("*", 3)` 應該傳回字串 `***`。

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` 應該傳回字串 `abcabcabc`。

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` 應該傳回字串 `abcabcabcabc`。

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` 應該傳回字串 `abc`。

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` 應該傳回字串 `********`。

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` 應該傳回一個空字串（`""`）。

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

內建的 `repeat()` 方法不應該被使用。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

`repeatStringNumTimes("abc", 0)` 應該傳回 `""`。

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

# --seed--

## --seed-contents--

```js
function repeatStringNumTimes(str, num) {
  return str;
}

repeatStringNumTimes('abc', 3);
```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
