---
id: 587d78af367417b2b2512b03
title: 建置調查表單
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該在一個具有 `id` 為 `title` 的 `h1` 元素中放置頁面標題。
1. 你應該在具有 `id` 為 `description` 的 `p` 元素中放置簡短的說明。
1. 你應該有一個 `form` 元素，其 `id` 為 `survey-form`。
1. 在表單元素內，你應該有一個必填的 `input` 欄位來輸入你的名字，該欄位的 `id` 為 `name`，且 `type` 為 `text`。
1. 在表單元素內，你應該有一個必填的 `input` 欄位來輸入你的電子郵件，該欄位的 `id` 為 `email`。
1. 如果你輸入的電子郵件格式不正確，你應該會看到 HTML5 驗證錯誤。
1. 在表單元素內，你應該有一個 `input` 欄位用來輸入一個具有 `id` 為 `number` 的數字。
1. 數字輸入不應接受非數字，無論是透過阻止你輸入，或是顯示 HTML5 驗證錯誤（視你的瀏覽器而定）。
1. 如果你輸入的數字超出由 `min` 和 `max` 屬性定義的數字輸入範圍，應該會看到 HTML5 驗證錯誤。
1. 對於名稱、電子郵件和數字輸入欄，表單中應該有對應的 `label` 元素，使用以下 id 來描述每個欄的用途：`name-label`、`email-label` 和 `number-label`。
1. 對於名稱、電子郵件和號碼輸入欄，應該有一個佔位文字，提供每個欄的描述或說明。
1. 在表單元素內，你應該有一個 `select` 下拉選單元素，其 `id` 為 `dropdown`，並且至少有兩個可供選擇的選項。
1. 在 `form` 元素內，你可以從使用 `name` 屬性群組的至少兩個圓鈕群組中選擇一個選項。
1. 在表單元素內，你可以從一系列核取方塊中選擇多個欄，每個核取方塊都必須有一個 `value` 屬性。
1. 在表單元素內，你應該有一個 `textarea` 用於額外評論。
1. 在表單元素內，你應該有一個 `id` 為 `submit` 的按鈕來提交所有輸入。

# --hints--

你應該有一個 `h1` 元素，其 `id` 為 `title`。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

你的 `#title` 不應該是空的。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你應該有一個 `p` 元素，其 `id` 為 `description`。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

你的 `#description` 不應該是空的。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你應該有一個 `form` 元素，其 `id` 為 `survey-form`。

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

你應該有一個 `input` 元素，其 `id` 為 `name`。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

你的 `#name` 應該有 `text` 的 `type`。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

你的 `#name` 應該需要輸入。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

你的 `#name` 應該是 `#survey-form` 的後代。

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

你應該有一個 `input` 元素，其 `id` 為 `email`。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

你的 `#email` 應該有 `email` 的 `type`。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

你的 `#email` 應該要求輸入。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

你的 `#email` 應該是 `#survey-form` 的子孫元素。

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

你應該有一個 `input` 元素，其 `id` 為 `number`。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

你的 `#number` 應該是 `#survey-form` 的後代。

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

你的 `#number` 應該有 `number` 的 `type`。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

你的 `#number` 應該有一個帶有數值的 `min` 屬性。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

你的 `#number` 應該有一個帶有數值的 `max` 屬性。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

你應該有一個 `label` 元素，其 `id` 為 `name-label`。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

你應該有一個 `label` 元素，其 `id` 為 `email-label`。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

你應該有一個 `label` 元素，其 `id` 為 `number-label`。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

你的 `#name-label` 應該包含描述輸入的文字。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你的 `#email-label` 應包含描述輸入的文字。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你的 `#number-label` 應包含描述輸入的文字。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你的 `#name-label` 應該是 `#survey-form` 的子孫。

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

你的 `#email-label` 應該是 `#survey-form` 的子孫。

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

你的 `#number-label` 應該是 `#survey-form` 的子孫。

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

你的 `#name` 應該有一個 `placeholder` 屬性和值。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

你的 `#email` 應該有一個 `placeholder` 屬性和值。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

你的 `#number` 應該有一個 `placeholder` 屬性和值。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

你應該有一個 `select` 欄位，其 `id` 為 `dropdown`。

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

你的 `#dropdown` 應該至少有兩個可選擇（未停用）的 `option` 元素。

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

你的 `#dropdown` 應該是 `#survey-form` 的子孫元素。

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

你應該至少有兩個 `input` 元素，其 `type` 為 `radio`（圓鈕、單選按鈕）。

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

你應該至少有兩個作為 `#survey-form` 子孫的圓鈕。

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

你所有的圓鈕都應該有一個 `value` 屬性和值。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

你所有的圓鈕都應該有一個 `name` 屬性和值。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

每個圓鈕群組應該至少有 2 個圓鈕。

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

你應該至少有兩個 `input` 元素，其 `type` 為 `checkbox`（核取方塊），且是 `#survey-form` 的子孫元素。

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

`#survey-form` 內所有的核取方塊都應該有 `value` 屬性和值。

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

你應該至少有一個 `textarea` 元素是 `#survey-form` 的子孫。

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

你應該有一個 `id` 為 `submit` 的 `input` 或 `button` 元素。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

你的 `#submit` 應該有 `submit` 的 `type`。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

你的 `#submit` 應該是 `#survey-form` 的子孫元素。

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
