---
id: 5895f70ef9fc0f352b528e6b
title: Cómo elaborar un perfil
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

Una vez asegurado que el usuario que accede a `/profile` está autenticado, se puede aprovechar la información contenida en `req.user` y hacer uso de ella en la página que estamos construyendo.

Pasa un objeto con la propiedad `username` y valor `req.user.username` como segundo argumento al método `render` de la vista profile.

Luego, ve a tu perfil `profile.pug` y agrega la siguiente línea debajo del elemento existente `h1` y al mismo nivel de indentación:

```pug
h2.center#welcome Welcome, #{username}!
```

Esto crea un elemento `h2` con la clase `center` y el id `welcome` que contiene el texto `Welcome, ` seguido del nombre de usuario.

Añade, también en `profile.pug`, un enlace a la ruta `/logout`, la cual albergará la lógica para desautenticar un usuario:

```pug
a(href='/logout') Logout
```

Envía tu página cuando creas que está correcta. Si tienes errores, puedes <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">consultar el proyecto completado hasta este punto</a>.

# --hints--

Debes añadir correctamente la variable de renderizado Pug a `/profile`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );
```

