---
id: 594810f028c0303b75339ad8
title: Zigzag-matrix
challengeType: 1
forumTopicId: 302348
dashedName: zig-zag-matrix
---

# --description--

Eine „Zigzag“-Matrix ist eine quadratische Anordnung der ersten $N^2$ ganzen Zahlen, bei der die Zahlen sequentiell an den Antidiagonalen der Matrix entlang in Zickzack-Reihenfolge ansteigen.

Zum Beispiel sollte für die Eingabe `5` folgendes Ergebnis erzeugt werden:

<pre>
 0  1  5  6 14
 2  4  7 13 15
 3  8 12 16 21
 9 11 17 20 22
10 18 19 23 24
</pre>

# --instructions--

Schreiben Sie eine Funktion, die die Größe der Zigzag-Matrix entgegennimmt und die entsprechende Matrix als zweidimensionales Array zurückgibt.

# --before-each--

```js
const zm1 = [[0]];
const zm2 = [[0, 1], [2, 3]];
const zm5 = [
  [0, 1, 5, 6, 14],
  [2, 4, 7, 13, 15],
  [3, 8, 12, 16, 21],
  [9, 11, 17, 20, 22],
  [10, 18, 19, 23, 24]
];
```

# --hints--

ZigZagMatrix sollte eine Funktion sein.

```js
assert.equal(typeof ZigZagMatrix, 'function');
```

ZigZagMatrix sollte ein Array zurückgeben.

```js
assert.equal(typeof ZigZagMatrix(1), 'object');
```

ZigZagMatrix sollte ein Array aus verschachtelten Arrays zurückgeben.

```js
assert.equal(typeof ZigZagMatrix(1)[0], 'object');
```

ZigZagMatrix(1) sollte \[[0]] zurückgeben.

```js
assert.deepEqual(ZigZagMatrix(1), zm1);
```

ZigZagMatrix(2) sollte \[[0, 1], [2, 3]] zurückgeben.

```js
assert.deepEqual(ZigZagMatrix(2), zm2);
```

ZigZagMatrix(5) sollte die angegebene Matrix zurückgeben.

```js
assert.deepEqual(ZigZagMatrix(5), zm5);
```

# --seed--

## --seed-contents--

```js
function ZigZagMatrix(n) {

  return [[], []];
}
```

# --solutions--

```js
function ZigZagMatrix(n) {
  const mtx = [];
  for (let i = 0; i < n; i++) {
    mtx[i] = [];
  }

  let i = 1;
  let j = 1;
  for (let e = 0; e < n * n; e++) {
    mtx[i - 1][j - 1] = e;
    if ((i + j) % 2 === 0) {
      // Even stripes
      if (j < n) j++;
      else i += 2;
      if (i > 1) i--;
    } else {
      // Odd stripes
      if (i < n) i++;
      else j += 2;
      if (j > 1) j--;
    }
  }
  return mtx;
}
```
