---
id: 594810f028c0303b75339ace
title: Akkumulator-Fabrik
challengeType: 1
forumTopicId: 302222
dashedName: accumulator-factory
---

# --description--

Ein von Paul Graham gestelltes Problem besteht darin, eine Funktion zu erstellen, die ein einzelnes (numerisches) Argument entgegennimmt und eine weitere Funktion zurückgibt, die ein Akkumulator ist. Die zurückgegebene Akkumulatorfunktion nimmt wiederum ein einzelnes numerisches Argument entgegen und gibt die Summe aller bisher an diesen Akkumulator übergebenen numerischen Werte zurück (einschließlich des Anfangswerts, der bei der Erstellung des Akkumulators übergeben wurde).

# --instructions--

Erstellen Sie eine Funktion, die eine Zahl $n$ entgegennimmt und Akkumulatorfunktionen erzeugt, die die Summe aller jemals an sie übergebenen Zahlen zurückgeben.

**Regeln:**

Verwenden Sie keine globalen Variablen.

**Hinweis:**

Closures speichern äußeren Zustand.

# --hints--

`accumulator` sollte eine Funktion sein.

```js
assert(typeof accumulator === 'function');
```

`accumulator(0)` sollte eine Funktion zurückgeben.

```js
assert(typeof accumulator(0) === 'function');
```

`accumulator(0)(2)` sollte eine Zahl zurückgeben.

```js
assert(typeof accumulator(0)(2) === 'number');
```

Die Übergabe der Werte 3, -4, 1.5 und 5 sollte 5,5 zurückgeben.

```js
assert(testFn && testFn(5) === 5.5);
```

# --before-each--

```js
const testFn = typeof accumulator === 'function' && accumulator(3);
if (testFn) {
  testFn(-4);
  testFn(1.5);
}
```

# --seed--

## --seed-contents--

```js
function accumulator(sum) {

}
```

# --solutions--

```js
function accumulator(sum) {
  return function(n) {
    return sum += n;
  };
}
```
