---
id: 594810f028c0303b75339ad4
title: Zeilenumbruch
challengeType: 1
forumTopicId: 302344
dashedName: word-wrap
---

# --description--

Auch heute, mit proportionalen Schriftarten und komplexen Layouts, gibt es noch Fälle, in denen Sie Text an einer bestimmten Spalte umbrechen müssen. Die grundlegende Aufgabe besteht darin, einen Absatz auf einfache Weise umzubrechen.

# --instructions--

Schreiben Sie eine Funktion, die diesen Text auf eine beliebige Anzahl von Zeichen umbrechen kann. Beachten Sie, dass der Eingabetext bereits Zeilenumbrüche enthält, die Ihre Funktion entsprechend verarbeiten sollte. Als Beispiel sollte der auf 80 Zeichen umgebrochene Text wie folgt aussehen:

<pre>
Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX
algorithm. If your language provides this, you get easy extra credit, but you
must reference documentation indicating that the algorithm is something better
than a simple minimum length algorithm.
</pre>

# --before-each--

```js
const text =
`Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX algorithm.
If your language provides this, you get easy extra credit,
but you ''must reference documentation'' indicating that the algorithm
is something better than a simple minimum length algorithm.`;

const wrapped80 = wrap(text, 80);
const wrapped42 = wrap(text, 42);

const firstRow80 =
    'Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX';

const firstRow42 = 'Wrap text using a more sophisticated';
```

# --hints--

`wrap` sollte eine Funktion sein.

```js
assert.equal(typeof wrap, 'function');
```

`wrap` sollte einen String zurückgeben.

```js
assert.equal(typeof wrap('abc', 10), 'string');
```

`wrap(text,80)` sollte 4 Zeilen zurückgeben.

```js
assert(wrapped80.split('\n').length === 4);
```

Ihre `wrap`-Funktion sollte den erwarteten Text zurückgeben.

```js
assert.equal(wrapped80.split('\n')[0], firstRow80);
```

`wrap(text,42)` sollte 7 Zeilen zurückgeben.

```js
assert(wrapped42.split('\n').length === 7);
```

Ihre `wrap`-Funktion sollte den erwarteten Text zurückgeben.

```js
assert.equal(wrapped42.split('\n')[0], firstRow42);
```

# --seed--

## --seed-contents--

```js
function wrap(text, limit) {
  return text;
}
```

# --solutions--

```js
function wrap(text, limit) {
  const noNewlines = text.replace('\n', '');
  if (noNewlines.length > limit) {
    // find the last space within limit
    const edge = noNewlines.slice(0, limit).lastIndexOf(' ');
    if (edge > 0) {
      const line = noNewlines.slice(0, edge);
      const remainder = noNewlines.slice(edge + 1);
      return line + '\n' + wrap(remainder, limit);
    }
  }
  return text;
}
```
