---
id: 594810f028c0303b75339ad3
title: Skalarprodukt von Vektoren
challengeType: 1
forumTopicId: 302343
dashedName: vector-dot-product
---

# --description--

Ein Vektor kann einen oder mehrere Werte haben, die durch eine geordnete Sammlung dargestellt werden. Beispiele könnten (x), (x, y) oder (x, y, z) sein.

# --instructions--

Schreiben Sie eine Funktion, die zwei Vektoren (dargestellt als eindimensionale Arrays) als Eingabe erhält und ihr Skalarprodukt berechnet. Ihre Funktion soll `null` zurückgeben bei ungültigen Eingaben, wie Vektoren unterschiedlicher Länge oder wenn etwas anderes als zwei Vektoren übergeben wird.

# --hints--

`dotProduct` sollte eine Funktion sein.

```js
assert.equal(typeof dotProduct, 'function');
```

`dotProduct()` sollte `null` zurückgeben.

```js
assert.equal(dotProduct(), null);
```

`dotProduct([1], [1])` sollte `1` zurückgeben.

```js
assert.equal(dotProduct([1], [1]), 1);
```

`dotProduct([1], [1, 2])` sollte `null` zurückgeben.

```js
assert.equal(dotProduct([1], [1, 2]), null);
```

`dotProduct([1, 3, -5], [4, -2, -1])` sollte `3` zurückgeben.

```js
assert.equal(dotProduct([1, 3, -5], [4, -2, -1]), 3);
```

`dotProduct([3, 2, 1], [2, 4, 2], [5, 3, 1])` sollte `null` zurückgeben.

```js
assert.equal(dotProduct([3, 2, 1], [2, 4, 2], [5, 3, 1]), null);
```

`dotProduct([ 0, 3, 6, 9, 12 ], [ 0, 4, 8, 12, 16 ])` sollte `360` zurückgeben.

```js
assert.equal(dotProduct([ 0, 3, 6, 9, 12 ], [ 0, 4, 8, 12, 16 ]), 360);
```

# --seed--

## --seed-contents--

```js
function dotProduct(...vectors) {

}
```

# --solutions--

```js
function dotProduct(...vectors) {
  if (!vectors || !vectors.length || vectors.length > 2 || vectors[0].length !== vectors[1].length) {
    return null;
  }
  const vectorLen = vectors[0].length;

  let prod = 0;
  let sum = 0;
  let j = vectorLen;
  let i = 2;
  // Sum terms
  while (j--) {
    i = 2;
    prod = 1;

    while (i--) {
      prod *= vectors[i][j];
    }
    sum += prod;
  }
  return sum;
}
```
