---
id: 594810f028c0303b75339ad1
title: Glückliche Zahlen
challengeType: 1
forumTopicId: 302280
dashedName: happy-numbers
---

# --description--

Eine glückliche Zahl wird durch folgenden Prozess definiert:

Ausgehend von einer beliebigen positiven ganzen Zahl ersetzen Sie die Zahl durch die Summe der Quadrate ihrer Ziffern und wiederholen den Vorgang, bis die Zahl `1` erreicht (wo sie bleibt) oder sich endlos in einem Zyklus wiederholt, der `1` nicht enthält. Die Zahlen, bei denen dieser Prozess in `1` endet, sind glückliche Zahlen, während diejenigen, die nicht in `1` enden, unglückliche Zahlen sind.

# --instructions--

Implementieren Sie eine Funktion, die true zurückgibt, wenn die Zahl glücklich ist, oder false, wenn nicht.

# --hints--

`happy` sollte eine Funktion sein.

```js
assert(typeof happy === 'function');
```

`happy(1)` sollte einen Boolean zurückgeben.

```js
assert(typeof happy(1) === 'boolean');
```

`happy(1)` sollte `true` zurückgeben.

```js
assert(happy(1));
```

`happy(2)` sollte `false` zurückgeben.

```js
assert(!happy(2));
```

`happy(7)` sollte `true` zurückgeben.

```js
assert(happy(7));
```

`happy(10)` sollte `true` zurückgeben.

```js
assert(happy(10));
```

`happy(13)` sollte `true` zurückgeben.

```js
assert(happy(13));
```

`happy(19)` sollte `true` zurückgeben.

```js
assert(happy(19));
```

`happy(23)` sollte `true` zurückgeben.

```js
assert(happy(23));
```

`happy(28)` sollte `true` zurückgeben.

```js
assert(happy(28));
```

`happy(31)` sollte `true` zurückgeben.

```js
assert(happy(31));
```

`happy(32)` sollte `true` zurückgeben.

```js
assert(happy(32));
```

`happy(33)` sollte `false` zurückgeben.

```js
assert(!happy(33));
```

# --seed--

## --seed-contents--

```js
function happy(number) {

}
```

# --solutions--

```js
function happy (number) {
  let m;
  let digit;
  const cycle = [];

  while (number !== 1 && cycle[number] !== true) {
    cycle[number] = true;
    m = 0;
    while (number > 0) {
      digit = number % 10;
      m += Math.pow(digit, 2);
      number = (number - digit) / 10;
    }
    number = m;
  }
  return (number === 1);
}
```
