---
id: 594810f028c0303b75339acf
title: Ackermann-Funktion
challengeType: 1
forumTopicId: 302223
dashedName: ackermann-function
---

# --description--

Die Ackermann-Funktion ist ein klassisches Beispiel für eine rekursive Funktion, die besonders bemerkenswert ist, weil sie keine primitive rekursive Funktion ist. Ihr Wert wächst sehr schnell, ebenso wie die Größe ihres Aufrufbaums.

Die Ackermann-Funktion wird üblicherweise wie folgt definiert:

$A(m, n) = \\begin{cases} n+1 & \\mbox{if } m = 0 \\\\ A(m-1, 1) & \\mbox{if } m > 0 \\mbox{ and } n = 0 \\\\ A(m-1, A(m, n-1)) & \\mbox{if } m > 0 \\mbox{ and } n > 0. \\end{cases}$

Ihre Argumente sind niemals negativ, und sie terminiert immer.

# --instructions--

Schreiben Sie eine Funktion, die den Wert von $A(m, n)$ zurückgibt. Beliebige Genauigkeit ist bevorzugt (da die Funktion so schnell wächst), aber nicht erforderlich.

# --hints--

`ack` sollte eine Funktion sein.

```js
assert(typeof ack === 'function');
```

`ack(0, 0)` sollte 1 zurückgeben.

```js
assert(ack(0, 0) === 1);
```

`ack(1, 1)` sollte 3 zurückgeben.

```js
assert(ack(1, 1) === 3);
```

`ack(2, 5)` sollte 13 zurückgeben.

```js
assert(ack(2, 5) === 13);
```

`ack(3, 3)` sollte 61 zurückgeben.

```js
assert(ack(3, 3) === 61);
```

# --seed--

## --seed-contents--

```js
function ack(m, n) {

}
```

# --solutions--

```js
function ack(m, n) {
  return m === 0 ? n + 1 : ack(m - 1, n === 0 ? 1 : ack(m, n - 1));
}
```
