---
id: 594810f028c0303b75339ad0
title: Spalten ausrichten
challengeType: 1
forumTopicId: 302224
dashedName: align-columns
---

# --description--

Gegeben ist ein Array mit vielen Zeilen, wobei die Felder innerhalb einer Zeile durch ein einzelnes `$`-Zeichen getrennt sind. Schreiben Sie ein Programm, das jede Spalte von Feldern so ausrichtet, dass die Wörter in jeder Spalte durch mindestens ein Leerzeichen getrennt sind. Erlauben Sie außerdem, dass jedes Wort in einer Spalte entweder linksbündig, rechtsbündig oder zentriert innerhalb seiner Spalte ausgerichtet wird.

# --instructions--

Verwenden Sie den folgenden Text, um Ihre Programme zu testen:

```js
const testText = [
  'Given$a$text$file$of$many$lines',
  'where$fields$within$a$line$',
  'are$delineated$by$a$single$"dollar"$character',
  'write$a$program',
  'that$aligns$each$column$of$fields',
  'by$ensuring$that$words$in$each$',
  'column$are$separated$by$at$least$one$space.',
  'Further,$allow$for$each$word$in$a$column$to$be$either$left$',
  'justified,$right$justified',
  'or$center$justified$within$its$column.'
];
```

**Beachten Sie:**

- Die Beispiel-Eingabetextzeilen können am Ende Dollarzeichen enthalten oder auch nicht.
- Alle Spalten sollten dieselbe Ausrichtung haben.
- Aufeinanderfolgende Leerzeichen, die am Ende von Zeilen entstehen, sind für die Aufgabe unerheblich.
- Der Ausgabetext wird in einer Monospace-Schriftart in einem einfachen Texteditor oder Terminal angezeigt. Die Zeilen sollten mit dem Zeilenumbruchzeichen (`\n`) verbunden werden.
- Der minimale Abstand zwischen den Spalten soll aus dem Text berechnet werden und darf nicht fest kodiert sein.
- Es ist nicht erforderlich, Trennzeichen zwischen oder um die Spalten hinzuzufügen.

Zum Beispiel eine der Zeilen aus `testText`, jeweils rechtsbündig, linksbündig und zentriert ausgerichtet:

```js
'    column        are separated     by     at    least       one space.\n'
'column     are        separated by     at     least    one       space.\n'
'  column      are     separated   by     at    least      one    space.\n'
```

# --before-each--

```js
const _testText = [
  'Given$a$text$file$of$many$lines',
  'where$fields$within$a$line$',
  'are$delineated$by$a$single$"dollar"$character',
  'write$a$program',
  'that$aligns$each$column$of$fields$',
  'by$ensuring$that$words$in$each$',
  'column$are$separated$by$at$least$one$space.',
  'Further,$allow$for$each$word$in$a$column$to$be$either$left$',
  'justified,$right$justified',
  'or$center$justified$within$its$column.'
];

const rightAligned = '     Given          a      text   file     of     many     lines\n' +
'     where     fields    within      a   line \n' +
'       are delineated        by      a single "dollar" character\n' +
'     write          a   program\n' +
'      that     aligns      each column     of   fields \n' +
'        by   ensuring      that  words     in     each \n' +
'    column        are separated     by     at    least       one space.\n' +
'  Further,      allow       for   each   word       in         a column to be either left \n' +
'justified,      right justified\n' +
'        or     center justified within    its  column.';

const leftAligned = 'Given      a          text      file   of     many     lines    \n' +
'where      fields     within    a      line   \n' +
'are        delineated by        a      single "dollar" character\n' +
'write      a          program  \n' +
'that       aligns     each      column of     fields   \n' +
'by         ensuring   that      words  in     each     \n' +
'column     are        separated by     at     least    one       space.\n' +
'Further,   allow      for       each   word   in       a         column to be either left \n' +
'justified, right      justified\n' +
'or         center     justified within its    column. ';

const centerAligned = '  Given        a        text     file    of     many     lines  \n' +
'  where      fields    within     a     line  \n' +
'   are     delineated    by       a    single \"dollar\" character\n' +
'  write        a       program \n' +
'   that      aligns     each    column   of    fields  \n' +
'    by      ensuring    that    words    in     each   \n' +
'  column      are     separated   by     at    least      one    space.\n' +
' Further,    allow       for     each   word     in        a     column to be either left \n' +
'justified,   right    justified\n' +
'    or       center   justified within  its   column. ';
```

# --hints--

`formatText` sollte eine Funktion sein.

```js
assert(typeof formatText === 'function');
```

`formatText(testText, 'right')` sollte Text mit rechtsbündig ausgerichteten Spalten erzeugen.

```js
assert.strictEqual(formatText(_testText, 'right'), rightAligned);
```

`formatText(testText, 'left')` sollte Text mit linksbündig ausgerichteten Spalten erzeugen.

```js
assert.strictEqual(formatText(_testText, 'left'), leftAligned);
```

`formatText(testText, 'center')` sollte Text mit zentriert ausgerichteten Spalten erzeugen.

```js
assert.strictEqual(formatText(_testText, 'center'), centerAligned);
```

# --seed--

## --seed-contents--

```js
function formatText(input, justification) {

}

const testText = [
  'Given$a$text$file$of$many$lines',
  'where$fields$within$a$line$',
  'are$delineated$by$a$single$"dollar"$character',
  'write$a$program',
  'that$aligns$each$column$of$fields$',
  'by$ensuring$that$words$in$each$',
  'column$are$separated$by$at$least$one$space.',
  'Further,$allow$for$each$word$in$a$column$to$be$either$left$',
  'justified,$right$justified',
  'or$center$justified$within$its$column.'
];
```

# --solutions--

```js
String.prototype.repeat = function (n) { return new Array(1 + parseInt(n)).join(this); };

function formatText(input, justification) {
  let x, y, max, cols = 0, diff, left, right;
  for (x = 0; x < input.length; x++) {
    input[x] = input[x].split('$');
    if (input[x].length > cols) {
      cols = input[x].length;
    }
  }
  for (x = 0; x < cols; x++) {
    max = 0;
    for (y = 0; y < input.length; y++) {
      if (input[y][x] && max < input[y][x].length) {
        max = input[y][x].length;
      }
    }
    for (y = 0; y < input.length; y++) {
      if (input[y][x]) {
        diff = (max - input[y][x].length) / 2;
        left = ' '.repeat(Math.floor(diff));
        right = ' '.repeat(Math.ceil(diff));
        if (justification === 'left') {
          right += left; left = '';
        }
        if (justification === 'right') {
          left += right; right = '';
        }
        input[y][x] = left + input[y][x] + right;
      }
    }
  }
  for (x = 0; x < input.length; x++) {
    input[x] = input[x].join(' ');
  }
  input = input.join('\n');
  return input;
}
```
