---
id: 594810f028c0303b75339ad7
title: Zhang-Suen-Dünnungsalgorithmus
challengeType: 1
forumTopicId: 302347
dashedName: zhang-suen-thinning-algorithm
---

# --description--

Dies ist ein Algorithmus, der verwendet wird, um schwarz-weiße, also ein Bit pro Pixel Bilder zu dünnen. Zum Beispiel erzeugt er bei einem Eingabebild wie:

```js
const testImage1 = [
 '                               ',
 '#########       ########       ',
 '###   ####     ####  ####      ',
 '###    ###     ###    ###      ',
 '###   ####     ###             ',
 '#########      ###             ',
 '### ####       ###    ###      ',
 '###  ####  ### ####  #### ###  ',
 '###   #### ###  ########  ###  ',
 '                               '
];
```

das gedünnte Ausgabebild:

```js
[ '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               ' ];
```

## Algorithmus

Gehen Sie davon aus, dass schwarze Pixel Eins und weiße Pixel Null sind und dass das Eingabebild ein rechteckiges N mal M Array aus Einsen und Nullen ist. Der Algorithmus arbeitet auf allen schwarzen Pixeln P1, die acht Nachbarn haben können. Die Nachbarn sind in folgender Reihenfolge angeordnet:

$$\begin{array}{|c|c|c|}
  \\hline
  P9 & P2              & P3\\\\ \\hline
  P8 & \boldsymbol{P1} & P4\\\\ \\hline
  P7 & P6              & P5\\\\ \\hline
\end{array}$$

Offensichtlich können die Randpixel des Bildes nicht alle acht Nachbarn haben.

- Definieren Sie $A(P1)$ als die Anzahl der Übergänge von Weiß zu Schwarz ($0 \to 1$) in der Folge P2, P3, P4, P5, P6, P7, P8, P9, P2. (Beachten Sie das zusätzliche P2 am Ende – die Folge ist zirkulär).
- Definieren Sie $B(P1)$ als die Anzahl der schwarzen Nachbarpixel von P1. ($= \\sum(P2 \ldots P9)$)

**Schritt 1:**

Alle Pixel werden getestet, und Pixel, die alle folgenden Bedingungen (gleichzeitig) erfüllen, werden in diesem Schritt nur notiert.

1. Das Pixel ist schwarz und hat acht Nachbarn
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. Mindestens eines von $P2$, $P4$ und $P6$ ist weiß
5. Mindestens eines von $P4$, $P6$ und $P8$ ist weiß

Nachdem über das Bild iteriert und alle Pixel gesammelt wurden, die alle Bedingungen von Schritt 1 erfüllen, werden diese Pixel auf weiß gesetzt.

**Schritt 2:**

Alle Pixel werden erneut getestet, und Pixel, die alle folgenden Bedingungen erfüllen, werden in diesem Schritt nur notiert.

1. Das Pixel ist schwarz und hat acht Nachbarn
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. Mindestens eines von $P2$, $P4$ und $P8$ ist weiß
5. Mindestens eines von $P2$, $P6$ und $P8$ ist weiß

Nachdem über das Bild iteriert und alle Pixel gesammelt wurden, die alle Bedingungen von Schritt 2 erfüllen, werden diese Pixel erneut auf weiß gesetzt.

**Iteration:**

Wenn in dieser Runde von Schritt 1 oder Schritt 2 irgendwelche Pixel geändert wurden, werden alle Schritte wiederholt, bis keine Pixel mehr geändert werden.

# --instructions--

Schreiben Sie eine Routine, die Zhang-Suen-Dünnung auf das bereitgestellte `image`, ein Array von Strings, anwendet, wobei jeder String eine einzelne Zeile des Bildes darstellt. Im String steht `#` für ein schwarzes Pixel, und Leerzeichen für ein weißes Pixel. Die Funktion soll das gedünnte Bild mit derselben Darstellung zurückgeben.

# --before-each--

```js
const _testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
const expected1 = [
  '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               '
];
const _testImage2 = [
  '                                                          ',
  ' #################                   #############        ',
  ' ##################               ################        ',
  ' ###################            ##################        ',
  ' ########     #######          ###################        ',
  '   ######     #######         #######       ######        ',
  '   ######     #######        #######                      ',
  '   #################         #######                      ',
  '   ################          #######                      ',
  '   #################         #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######         #######       ######        ',
  ' ########     #######          ###################        ',
  ' ########     ####### ######    ################## ###### ',
  ' ########     ####### ######      ################ ###### ',
  ' ########     ####### ######         ############# ###### ',
  '                                                          '
];
const expected2 = [
  '                                                          ',
  '                                                          ',
  '    # ##########                       #######            ',
  '     ##        #                   ####       #           ',
  '     #          #                 ##                      ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     ############               #                         ',
  '     #          #               #                         ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #                            ##                      ',
  '     #                             ############           ',
  '                       ###                          ###   ',
  '                                                          ',
  '                                                          '
];
```

# --hints--

`thinImage` sollte eine Funktion sein.

```js
assert.equal(typeof thinImage, 'function');
```

`thinImage` sollte ein Array zurückgeben.

```js
assert(Array.isArray(thinImage(_testImage1)));
```

`thinImage` sollte ein Array von Strings zurückgeben.

```js
assert.equal(typeof thinImage(_testImage1)[0], 'string');
```

`thinImage(testImage1)` sollte ein gedünntes Bild wie im Beispiel zurückgeben.

```js
assert.deepEqual(thinImage(_testImage1), expected1);
```

`thinImage(testImage2)` sollte ein gedünntes Bild zurückgeben.

```js
assert.deepEqual(thinImage(_testImage2), expected2);
```

# --seed--

## --seed-contents--

```js
function thinImage(image) {

}

const testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
```

# --solutions--

```js
function Point(x, y) {
  this.x = x;
  this.y = y;
}

const ZhangSuen = (function () {
  function ZhangSuen() {
  }

  ZhangSuen.nbrs = [[0, -1], [1, -1], [1, 0], [1, 1], [0, 1], [-1, 1], [-1, 0], [-1, -1], [0, -1]];

  ZhangSuen.nbrGroups = [[[0, 2, 4], [2, 4, 6]], [[0, 2, 6], [0, 4, 6]]];

  ZhangSuen.toWhite = [];

  ZhangSuen.main = function (image) {
    ZhangSuen.grid = new Array(image);
    for (let r = 0; r < image.length; r++) {
      ZhangSuen.grid[r] = image[r].split('');
    }
    ZhangSuen.thinImage();
    return ZhangSuen.getResult();
  };

  ZhangSuen.thinImage = function () {
    let firstStep = false;
    let hasChanged;
    do {
      hasChanged = false;
      firstStep = !firstStep;
      for (let r = 1; r < ZhangSuen.grid.length - 1; r++) {
        for (let c = 1; c < ZhangSuen.grid[0].length - 1; c++) {
          if (ZhangSuen.grid[r][c] !== '#') {
            continue;
          }
          const nn = ZhangSuen.numNeighbors(r, c);
          if (nn < 2 || nn > 6) {
            continue;
          }
          if (ZhangSuen.numTransitions(r, c) !== 1) {
            continue;
          }
          if (!ZhangSuen.atLeastOneIsWhite(r, c, firstStep ? 0 : 1)) {
            continue;
          }
          ZhangSuen.toWhite.push(new Point(c, r));
          hasChanged = true;
        }
      }
      for (let i = 0; i < ZhangSuen.toWhite.length; i++) {
        const p = ZhangSuen.toWhite[i];
        ZhangSuen.grid[p.y][p.x] = ' ';
      }
      ZhangSuen.toWhite = [];
    } while ((firstStep || hasChanged));
  };

  ZhangSuen.numNeighbors = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === '#') {
        count++;
      }
    }
    return count;
  };

  ZhangSuen.numTransitions = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === ' ') {
        if (ZhangSuen.grid[r + ZhangSuen.nbrs[i + 1][1]][c + ZhangSuen.nbrs[i + 1][0]] === '#') {
          count++;
        }
      }
    }
    return count;
  };

  ZhangSuen.atLeastOneIsWhite = function (r, c, step) {
    let count = 0;
    const group = ZhangSuen.nbrGroups[step];
    for (let i = 0; i < 2; i++) {
      for (let j = 0; j < group[i].length; j++) {
        const nbr = ZhangSuen.nbrs[group[i][j]];
        if (ZhangSuen.grid[r + nbr[1]][c + nbr[0]] === ' ') {
          count++;
          break;
        }
      }
    }
    return count > 1;
  };

  ZhangSuen.getResult = function () {
    const result = [];
    for (let i = 0; i < ZhangSuen.grid.length; i++) {
      const row = ZhangSuen.grid[i].join('');
      result.push(row);
    }
    return result;
  };
  return ZhangSuen;
}());

function thinImage(image) {
  return ZhangSuen.main(image);
}
```
