---
id: 594810f028c0303b75339acb
title: 100 Türen
challengeType: 1
forumTopicId: 302217
dashedName: 100-doors
---

# --description--

Es gibt 100 Türen in einer Reihe, die alle anfangs geschlossen sind. Sie gehen 100 Mal an den Türen vorbei. Beim ersten Durchgang besuchen Sie jede Tür und „schalten“ sie um (wenn die Tür geschlossen ist, öffnen Sie sie; wenn sie offen ist, schließen Sie sie). Beim zweiten Durchgang besuchen Sie nur jede 2. Tür (also Tür Nr. 2, Nr. 4, Nr. 6, ...) und schalten sie um. Beim dritten Durchgang besuchen Sie jede 3. Tür (also Tür Nr. 3, Nr. 6, Nr. 9, ...), und so weiter, bis Sie nur noch die 100. Tür besuchen.

# --instructions--

Implementieren Sie eine Funktion, um den Zustand der Türen nach dem letzten Durchgang zu bestimmen. Geben Sie das Endergebnis in einem Array zurück, wobei nur die Türnummer im Array enthalten sein soll, wenn die Tür offen ist.

# --hints--

`getFinalOpenedDoors` sollte eine Funktion sein.

```js
assert(typeof getFinalOpenedDoors === 'function');
```

`getFinalOpenedDoors` sollte ein Array zurückgeben.

```js
assert(Array.isArray(getFinalOpenedDoors(100)));
```

`getFinalOpenedDoors` sollte das korrekte Ergebnis liefern.

```js
const solution = [1, 4, 9, 16, 25, 36, 49, 64, 81, 100];
assert.deepEqual(getFinalOpenedDoors(100), solution);
```

# --seed--

## --seed-contents--

```js
function getFinalOpenedDoors(numDoors) {

}
```

# --solutions--

```js
function getFinalOpenedDoors(numDoors) {
  // this is the final pattern (always squares).
  // thus, the most efficient solution simply returns an array of squares up to numDoors).
  const finalState = [];
  let i = 1;
  while (Math.pow(i, 2) <= numDoors) {
    finalState.push(Math.pow(i, 2));
    i++;
  }
  return finalState;
}
```
