---
id: 594810f028c0303b75339acd
title: 'Klassifizierungen von überreichen, mangelhaften und vollkommenen Zahlen'
challengeType: 1
forumTopicId: 302221
dashedName: abundant-deficient-and-perfect-number-classifications
---

# --description--

Diese definieren drei Klassifizierungen von positiven ganzen Zahlen basierend auf ihren echten Teilern.

Sei $P(n)$ die Summe der echten Teiler von `n`, wobei echte Teiler alle positiven ganzen Zahlen `n` außer `n` selbst sind.

Wenn `P(n) < n`, dann wird `n` als `deficient` klassifiziert.

Wenn `P(n) === n`, dann wird `n` als `perfect` klassifiziert.

Wenn `P(n) > n`, dann wird `n` als `abundant` klassifiziert.

**Beispiel**: `6` hat echte Teiler `1`, `2` und `3`. `1 + 2 + 3 = 6`, daher wird `6` als vollkommene Zahl klassifiziert.

# --instructions--

Implementieren Sie eine Funktion, die berechnet, wie viele der ganzen Zahlen von `1` bis `num` (einschließlich) in jede der drei Klassen fallen. Geben Sie das Ergebnis als Array im folgenden Format `[deficient, perfect, abundant]` aus.

# --hints--

`getDPA` sollte eine Funktion sein.

```js
assert(typeof getDPA === 'function');
```

`getDPA(5000)` sollte ein Array zurückgeben.

```js
assert(Array.isArray(getDPA(5000)));
```

`getDPA(5000)` zurückgegebenes Array sollte eine Länge von `3` haben.

```js
assert(getDPA(5000).length === 3);
```

`getDPA(5000)` sollte `[3758, 3, 1239]` zurückgeben.

```js
assert.deepEqual(getDPA(5000), [3758, 3, 1239]);
```

`getDPA(10000)` sollte `[7508, 4, 2488]` zurückgeben.

```js
assert.deepEqual(getDPA(10000), [7508, 4, 2488]);
```

`getDPA(20000)` sollte `[15043, 4, 4953]` zurückgeben.

```js
assert.deepEqual(getDPA(20000), [15043, 4, 4953]);
```

# --seed--

## --seed-contents--

```js
function getDPA(num) {

}
```

# --solutions--

```js
function getDPA(num) {
  const dpa = [1, 0, 0];
  for (let n = 2; n <= num; n += 1) {
    let ds = 1;
    const e = Math.sqrt(n);
    for (let d = 2; d < e; d += 1) {
      if (n % d === 0) {
        ds += d + (n / d);
      }
    }
    if (n % e === 0) {
      ds += e;
    }
    dpa[ds < n ? 0 : ds === n ? 1 : 2] += 1;
  }
  return dpa;
}
```
