---
id: 5949b579404977fbaefcd736
title: "9 Milliarden Namen Gottes, die ganze Zahl"
challengeType: 1
forumTopicId: 302219
dashedName: 9-billion-names-of-god-the-integer
---

# --description--

Diese Aufgabe ist eine Variation der Kurzgeschichte von Arthur C. Clarke.

(Lösende sollten sich der Konsequenzen bewusst sein, die das Abschließen dieser Aufgabe mit sich bringt.)

Im Detail, um zu spezifizieren, was mit einem „Namen“ gemeint ist:

<ul>
  <li>Die ganze Zahl 1 hat 1 Namen „1“.</li>
  <li>Die ganze Zahl 2 hat 2 Namen „1+1“ und „2“.</li>
  <li>Die ganze Zahl 3 hat 3 Namen „1+1+1“, „2+1“ und „3“.</li>
  <li>Die ganze Zahl 4 hat 5 Namen „1+1+1+1“, „2+1+1“, „2+2“, „3+1“, „4“.</li>
  <li>Die ganze Zahl 5 hat 7 Namen „1+1+1+1+1“, „2+1+1+1“, „2+2+1“, „3+1+1“, „3+2“, „4+1“, „5“.</li>
</ul>

Dies kann in folgender Form visualisiert werden:

<pre>          1
        1   1
      1   1   1
    1   2   1   1
  1   2   2   1   1
1   3   3   2   1   1
</pre>

Dabei entspricht die Zeile $n$ der ganzen Zahl $n$, und jede Spalte $C$ in der Zeile $m$ von links nach rechts entspricht der Anzahl der Namen, die mit $C$ beginnen.

Optional: Beachten Sie, dass die Summe der $n$-ten Zeile $P(n)$ die Funktion der ganzzahligen Partition ist.

# --instructions--

Implementieren Sie eine Funktion, die die Summe der $n$-ten Zeile zurückgibt.

# --hints--

`numberOfNames` sollte function sein.

```js
assert(typeof numberOfNames === 'function');
```

`numberOfNames(5)` sollte 7 ergeben.

```js
assert.equal(numberOfNames(5), 7);
```

`numberOfNames(12)` sollte 77 ergeben.

```js
assert.equal(numberOfNames(12), 77);
```

`numberOfNames(18)` sollte 385 ergeben.

```js
assert.equal(numberOfNames(18), 385);
```

`numberOfNames(23)` sollte 1255 ergeben.

```js
assert.equal(numberOfNames(23), 1255);
```

`numberOfNames(42)` sollte 53174 ergeben.

```js
assert.equal(numberOfNames(42), 53174);
```

`numberOfNames(123)` sollte 2552338241 ergeben.

```js
assert.equal(numberOfNames(123), 2552338241);
```

# --seed--

## --seed-contents--

```js
function numberOfNames(num) {

  return true;
}
```

# --solutions--

```js
function numberOfNames(num) {
  const cache = [
    [1]
  ];
  for (let l = cache.length; l < num + 1; l++) {
    let Aa;
    let Mi;
    const r = [0];
    for (let x = 1; x < l + 1; x++) {
      r.push(r[r.length - 1] + (Aa = cache[l - x < 0 ? cache.length - (l - x) : l - x])[(Mi = Math.min(x, l - x)) < 0 ? Aa.length - Mi : Mi]);
    }
    cache.push(r);
  }
  return cache[num][cache[num].length - 1];
}
```
