---
id: 594da033de4190850b893874
title: Mittelwerte/Quadratisches Mittel
challengeType: 1
forumTopicId: 302228
dashedName: averagesroot-mean-square
---

# --description--

Berechnen Sie das quadratische Mittel (Root Mean Square, RMS) der Zahlen von 1 bis einschließlich 10.

Das <abbr title="Root mean square">RMS</abbr> wird berechnet, indem die Quadratwurzel des Mittels der Quadrate der Zahlen genommen wird, dargestellt durch die Gleichung:

$$x\_{\\mathrm{rms}} = \\sqrt {{{x_1}^2 + {x_2}^2 + \\cdots + {x_n}^2} \\over n}. $$

# --hints--

`rms` sollte eine Funktion sein.

```js
assert(typeof rms === 'function');
```

`rms([1, 2, 3, 4, 5, 6, 7, 8, 9, 10])` sollte gleich `6.2048368229954285` sein.

```js
const arr1 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const answer1 = 6.2048368229954285;
assert.equal(rms(arr1), answer1);
```

# --seed--

## --seed-contents--

```js
function rms(arr) {

}
```

# --solutions--

```js
function rms(arr) {
  const sumOfSquares = arr.reduce((s, x) => s + x * x, 0);
  return Math.sqrt(sumOfSquares / arr.length);
}
```
