---
id: 594d8d0ab97724821379b1e6
title: Mittelwerte/Modus
challengeType: 1
forumTopicId: 302226
dashedName: averagesmode
---

# --description--

Schreiben Sie eine Funktion `mode`, die den Wert findet, der in einem Array am häufigsten vorkommt.

Der Fall, dass die Sammlung leer ist, kann ignoriert werden. Achten Sie darauf, den Fall zu behandeln, in dem der Modus nicht eindeutig ist.

Wenn es nicht angemessen oder möglich ist, eine allgemeine Sammlung zu unterstützen, verwenden Sie wenn möglich einen Vektor (Array). Wenn es nicht angemessen oder möglich ist, einen unbestimmten Werttyp zu unterstützen, verwenden Sie ganze Zahlen.

# --before-each--

```js
const arr1 = [1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17];
const arr2 = [1, 2, 4, 4, 1];
```

# --hints--

`mode` sollte eine Funktion sein.

```js
assert(typeof mode === 'function');
```

`mode([1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17])` sollte gleich `[6]` sein.

```js
assert.deepEqual(mode(arr1), [6]);
```

`mode([1, 2, 4, 4, 1])` sollte gleich `[1, 4]` sein.

```js
assert.deepEqual(mode(arr2).sort(), [1, 4]);
```

# --seed--

## --seed-contents--

```js
function mode(arr) {

  return true;
}
```

# --solutions--

```js
function mode(arr) {
  const counter = {};
  let result = [];
  let max = 0;
  // for (const i in arr) {
  arr.forEach(el => {
    if (!(el in counter)) {
      counter[el] = 0;
    }
    counter[el]++;

    if (counter[el] === max) {
      result.push(el);
    }
    else if (counter[el] > max) {
      max = counter[el];
      result = [el];
    }
  });
  return result;
}
```
