---
id: 594faaab4e2a8626833e9c3d
title: Einen string mit Escape-Zeichen tokenisieren
challengeType: 1
forumTopicId: 302338
dashedName: tokenize-a-string-with-escaping
---

# --description--

Schreiben Sie eine Funktion oder ein Programm, das einen String an jeder nicht-escaped auftretenden Stelle eines Trennzeichens aufteilen kann.

Es soll drei Eingabeparameter akzeptieren:

<ul>
  <li>Den <strong>String</strong></li>
  <li>Das <strong>Trennzeichen</strong></li>
  <li>Das <strong>Escape-Zeichen</strong></li>
</ul>

Es soll eine Liste von Strings ausgeben.

Regeln für das Aufteilen:

<ul>
  <li>Die Felder, die durch die Trennzeichen getrennt wurden, werden zu den Elementen der Ausgabeliste.</li>
  <li>Leere Felder sollen erhalten bleiben, auch am Anfang und Ende.</li>
</ul>

Regeln für das Escapen:

<ul>
  <li>"Escaped" bedeutet, dass es von einem Escape-Zeichen unmittelbar davor begleitet wird, das selbst nicht escaped ist.</li>
  <li>Wenn das Escape-Zeichen einem Zeichen ohne besondere Bedeutung vorausgeht, zählt es trotzdem als Escape (hat aber keine spezielle Wirkung).</li>
  <li>Jede Vorkommnis des Escape-Zeichens, das zum Escapen verwendet wurde, soll nicht Teil der Ausgabe werden.</li>
</ul>

Zeigen Sie, dass Ihre Funktion den folgenden Testfall erfüllt:

Gegeben ist der String

<pre>one^|uno||three^^^^|four^^^|^cuatro|</pre>

und unter Verwendung von `|` als Trennzeichen und `^` als Escape-Zeichen soll Ihre Funktion das folgende Array ausgeben:

<pre>  ['one|uno', '', 'three^^', 'four^|cuatro', '']
</pre>

# --before-each--

```js
const testStr1 = 'one^|uno||three^^^^|four^^^|^cuatro|';
const res1 = ['one|uno', '', 'three^^', 'four^|cuatro', ''];

const testStr2 = 'a@&bcd&ef&&@@hi';
const res2 = ['a&bcd', 'ef', '', '@hi'];

const testStr3 = 'hello^|world|how^are^you^|';
const res3 = ['hello|world', 'howareyou|'];
```

# --hints--

`tokenize` sollte eine Funktion sein.

```js
assert(typeof tokenize === 'function');
```

`tokenize` sollte ein Array zurückgeben.

```js
assert(typeof tokenize('a', 'b', 'c') === 'object');
```

`tokenize('one^|uno||three^^^^|four^^^|^cuatro|', '|', '^')` sollte `['one|uno', '', 'three^^', 'four^|cuatro', '']` zurückgeben.

```js
assert.deepEqual(tokenize(testStr1, '|', '^'), res1);
```

`tokenize('a@&bcd&ef&&@@hi', '&', '@')` sollte `['a&bcd', 'ef', '', '@hi']` zurückgeben.

```js
assert.deepEqual(tokenize(testStr2, '&', '@'), res2);
```

`tokenize('hello^|world|how^are^you^|', '|', '^')` sollte `['hello|world', 'howareyou|']` zurückgeben.

```js
assert.deepEqual(tokenize(testStr3, '|', '^'), res3);
```

# --seed--

## --seed-contents--

```js
function tokenize(str, sep, esc) {
  return true;
}
```

# --solutions--

```js
// tokenize :: String -> Character -> Character -> [String]
function tokenize(str, charDelim, charEsc) {
  const dctParse = str.split('')
    .reduce((a, x) => {
      const blnEsc = a.esc;
      const blnBreak = !blnEsc && x === charDelim;
      const blnEscChar = !blnEsc && x === charEsc;

      return {
        esc: blnEscChar,
        token: blnBreak ? '' : (
          a.token + (blnEscChar ? '' : x)
        ),
        list: a.list.concat(blnBreak ? a.token : [])
      };
    }, {
      esc: false,
      token: '',
      list: []
    });

  return dctParse.list.concat(
    dctParse.token
  );
}
```
