---
id: 594fa2746886f41f7d8bf225
title: Topologische Sortierung
challengeType: 1
forumTopicId: 302340
dashedName: topological-sort
---

# --description--

Gegeben ist eine Abbildung zwischen Elementen und den Elementen, von denen sie abhĂ¤ngen. Eine topologische Sortierung ordnet die Elemente so, dass kein Element vor einem Element steht, von dem es abhĂ¤ngt. Es gibt zwei bekannte Algorithmen fĂĽr die topologische Sortierung: Kahns (1962) topologische Sortierung und Tiefensuche.

# --instructions--

Schreiben Sie eine Funktion, die eine Liste mit einer gĂĽltigen Kompilierreihenfolge von Bibliotheken basierend auf deren AbhĂ¤ngigkeiten zurĂĽckgibt.

- Gehen Sie davon aus, dass Bibliotheksnamen aus einem einzigen Wort bestehen.
- Elemente, die nur als AbhĂ¤ngige genannt werden, haben keine eigenen AbhĂ¤ngigkeiten, aber ihre Kompilierreihenfolge muss angegeben werden.
- SelbstabhĂ¤ngigkeiten sollen ignoriert werden.
- Nicht sortierbare AbhĂ¤ngigkeiten sollen ignoriert werden.

Verwenden Sie die folgenden Daten als Beispiel:

<pre>
BIBLIOTHEK      BIBLIOTHEKSABHĂ„NGIGKEITEN
=======         ===========================
des_system_lib  std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
dw01            ieee dw01 dware gtech
dw02            ieee dw02 dware
dw03            std synopsys dware dw03 dw02 dw01 ieee gtech
dw04            dw04 ieee dw01 dware gtech
dw05            dw05 ieee dware
dw06            dw06 ieee dware
dw07            ieee dware
dware           ieee dware
gtech           ieee gtech
ramlib          std ieee
std_cell_lib    ieee std_cell_lib
synopsys
</pre>

Das Kompilieren einer Bibliothek in der VHDL-Sprache hat die EinschrĂ¤nkung, dass eine Bibliothek erst nach allen Bibliotheken kompiliert werden darf, von denen sie abhĂ¤ngt. Die obigen Daten wĂ¤ren nicht sortierbar, wenn zum Beispiel `dw04` zur Liste der AbhĂ¤ngigkeiten von `dw01` hinzugefĂĽgt wĂĽrde.

Die Eingabe der Funktion ist ein mehrzeiliger String, wobei jede Zeile aus dem Namen der Bibliothek und anschlieĂźend deren AbhĂ¤ngigkeiten (falls vorhanden) besteht.

Zum Beispiel:

```js
const libsSimple =
  `aaa bbb
  bbb`;
```

# --before-each--

```js
const libsSimple =
  `aaa bbb
  bbb`;

const libsVHDL =
  `des_system_lib   std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
  dw01             ieee dw01 dware gtech
  dw02             ieee dw02 dware
  dw03             std synopsys dware dw03 dw02 dw01 ieee gtech
  dw04             dw04 ieee dw01 dware gtech
  dw05             dw05 ieee dware
  dw06             dw06 ieee dware
  dw07             ieee dware
  dware            ieee dware
  gtech            ieee gtech
  ramlib           std ieee
  std_cell_lib     ieee std_cell_lib
  synopsys`;

const libsCustom =
  `a b c d
  b c d
  d c
  c base
  base`;

const libsUnorderable =
  `TestLib Base MainLib
  MainLib TestLib
  Base`;
```

# --hints--

`topologicalSort` sollte eine Funktion sein.

```js
assert(typeof topologicalSort === 'function');
```

`topologicalSort(libsSimple)` sollte ein Array zurĂĽckgeben.

```js
assert(Array.isArray(topologicalSort(libsSimple)));
```

`topologicalSort(libsSimple)` sollte `['bbb', 'aaa']` zurĂĽckgeben.

```js
assert.deepEqual(topologicalSort(libsSimple), ['bbb', 'aaa']);
```

`topologicalSort(libsVHDL)` sollte `['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']` zurĂĽckgeben.

```js
assert.deepEqual(topologicalSort(libsVHDL), ['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']);
```

`topologicalSort(libsCustom)` sollte `['base', 'c', 'd', 'b', 'a']` zurĂĽckgeben.

```js
assert.deepEqual(topologicalSort(libsCustom), ['base', 'c', 'd', 'b', 'a']);
```

`topologicalSort` sollte nicht sortierbare AbhĂ¤ngigkeiten ignorieren.

```js
assert.deepEqual(topologicalSort(libsUnorderable), ['Base']);
```

# --seed--

## --seed-contents--

```js
function topologicalSort(libs) {

  return true;
}
```

# --solutions--

```js
function topologicalSort(libs) {
  // A map of the input data, with the keys as the packages, and the values as
  // and array of packages on which it depends.
  const D = libs
    .split('\n')
    .map(e => e.split(' ').filter(ep => ep !== ''))
    .reduce((p, c) =>
      p.set(c[0], c.filter((e, i) => (i > 0 && e !== c[0] ? e : null))), new Map());
  [].concat(...D.values()).forEach(e => {
    D.set(e, D.get(e) || []);
  });

  // The above map rotated so that it represents a DAG of the form
  // Map {
  //    A => [ A, B, C],
  //    B => [C],
  //    C => []
  // }
  // where each key represents a node, and the array contains the edges.
  const G = [...D.keys()].reduce((p, c) =>
    p.set(
      c,
      [...D.keys()].filter(e => D.get(e).includes(c))),
    new Map()
  );

  // An array of leaf nodes; nodes with 0 in degrees.
  const Q = [...D.keys()].filter(e => D.get(e).length === 0);

  // The result array.
  const S = [];
  while (Q.length) {
    const u = Q.pop();
    S.push(u);
    G.get(u).forEach(v => {
      D.set(v, D.get(v).filter(e => e !== u));
      if (D.get(v).length === 0) {
        Q.push(v);
      }
    });
  }

  return S;
}
```
