---
id: 595011cba5a81735713873bd
title: Top-Rang pro Gruppe
challengeType: 1
forumTopicId: 302339
dashedName: top-rank-per-group
---

# --description--

Finden Sie die Daten mit dem höchsten `n` Rang in jeder Gruppe, wobei `n` als Parameter übergeben wird. Der Name des Rangs und der Gruppe werden ebenfalls als Parameter angegeben.

Gegeben sind die folgenden Daten:

```js
testData1 = [
  { name: 'Tyler Bennett', id: 'E10297', salary: 32000, dept: 'D101' },
  { name: 'John Rappl', id: 'E21437', salary: 47000, dept: 'D050' },
  { name: 'George Woltman', id: 'E00127', salary: 53500, dept: 'D101' },
  { name: 'Adam Smith', id: 'E63535', salary: 18000, dept: 'D202' },
  { name: 'Claire Buckman', id: 'E39876', salary: 27800, dept: 'D202' },
  { name: 'David McClellan', id: 'E04242', salary: 41500, dept: 'D101' },
  { name: 'Rich Holcomb', id: 'E01234', salary: 49500, dept: 'D202' },
  { name: 'Nathan Adams', id: 'E41298', salary: 21900, dept: 'D050' },
  { name: 'Richard Potter', id: 'E43128', salary: 15900, dept: 'D101' },
  { name: 'David Motsinger', id: 'E27002', salary: 19250, dept: 'D202' },
  { name: 'Tim Sampair', id: 'E03033', salary: 27000, dept: 'D101' },
  { name: 'Kim Arlich', id: 'E10001', salary: 57000, dept: 'D190' },
  { name: 'Timothy Grove', id: 'E16398', salary: 29900, dept: 'D190' }
];
```

Man könnte die Top 10 Mitarbeiter in jeder Abteilung folgendermaßen ranken:

```js
topRankPerGroup(10, testData1, 'dept', 'salary')
```

Gegeben sind die folgenden Daten:

```js
testData2 = [
  { name: 'Friday 13th', genre: 'horror', rating: 9.9 },
  { name: "Nightmare on Elm's Street", genre: 'horror', rating: 5.7 },
  { name: 'Titanic', genre: 'drama', rating: 7.3 },
  { name: 'Maze Runner', genre: 'scifi', rating: 7.1 },
  { name: 'Blade runner', genre: 'scifi', rating: 8.9 }
];
```

Man könnte den bestbewerteten Film in jedem Genre folgendermaßen ranken:

```js
topRankPerGroup(1, testData2, 'genre', 'rating')
```

Die Funktion soll ein Array zurückgeben, das für jede Gruppe ein Array mit den Top `n` Objekten enthält.

Zum Beispiel, bei den gegebenen Daten:

```js
[
  { name: 'Claire Buckman', id: 'E39876', salary: 27800, dept: 'D101' },
  { name: 'Rich Holcomb', id: 'E01234', salary: 49500, dept: 'D050' },
  { name: 'David Motsinger', id: 'E27002', salary: 19250, dept: 'D050' },
  { name: 'Tim Sampair', id: 'E03033', salary: 27000, dept: 'D101' },
  { name: 'Kim Arlich', id: 'E10001', salary: 57000, dept: 'D050' },
  { name: 'Timothy Grove', id: 'E16398', salary: 29900, dept: 'D101' }
];
```

Die zwei bestplatzierten Mitarbeiter in jeder Abteilung nach Gehalt wären:

```js
[ [ { name: 'Kim Arlich', id: 'E10001', salary: 57000, dept: 'D050' },
    { name: 'Rich Holcomb', id: 'E01234', salary: 49500, dept: 'D050' } ],
  [ { name: 'Timothy Grove', id: 'E16398', salary: 29900, dept: 'D101' },
    { name: 'Claire Buckman', id: 'E39876', salary: 27800, dept: 'D101' } ] ]
```

# --before-each--

```js
const testData1 = [
  { name: 'Tyler Bennett', id: 'E10297', salary: 32000, dept: 'D101' },
  { name: 'John Rappl', id: 'E21437', salary: 47000, dept: 'D050' },
  { name: 'George Woltman', id: 'E00127', salary: 53500, dept: 'D101' },
  { name: 'Adam Smith', id: 'E63535', salary: 18000, dept: 'D202' },
  { name: 'Claire Buckman', id: 'E39876', salary: 27800, dept: 'D202' },
  { name: 'David McClellan', id: 'E04242', salary: 41500, dept: 'D101' },
  { name: 'Rich Holcomb', id: 'E01234', salary: 49500, dept: 'D202' },
  { name: 'Nathan Adams', id: 'E41298', salary: 21900, dept: 'D050' },
  { name: 'Richard Potter', id: 'E43128', salary: 15900, dept: 'D101' },
  { name: 'David Motsinger', id: 'E27002', salary: 19250, dept: 'D202' },
  { name: 'Tim Sampair', id: 'E03033', salary: 27000, dept: 'D101' },
  { name: 'Kim Arlich', id: 'E10001', salary: 57000, dept: 'D190' },
  { name: 'Timothy Grove', id: 'E16398', salary: 29900, dept: 'D190' }
];

const testData2 = [
  { name: 'Friday 13th', genre: 'horror', rating: 9.9 },
  { name: "Nightmare on Elm's Street", genre: 'horror', rating: 5.7 },
  { name: 'Titanic', genre: 'drama', rating: 7.3 },
  { name: 'Maze Runner', genre: 'scifi', rating: 7.1 },
  { name: 'Blade runner', genre: 'scifi', rating: 8.9 }
];

function getResults_() {
  return {
    res1: topRankPerGroup(10, testData1, 'dept', 'salary'),
    res2: topRankPerGroup(1, testData2, 'genre', 'rating'),
    res3: topRankPerGroup(2, testData2, 'genre', 'rating')
  };
}
```

# --hints--

`topRankPerGroup` sollte eine Funktion sein.

```js
assert(typeof topRankPerGroup === 'function');
```

`topRankPerGroup` sollte bei negativen n-Werten undefined zurückgeben.

```js
assert(typeof topRankPerGroup(-1, []) === 'undefined');
```

Für `topRankPerGroup(10, testData1, 'dept', 'salary')` sollte das erste Ergebnis in der ersten Gruppe `{ name: 'John Rappl', id: 'E21437', salary: 47000, dept: 'D050'}` sein.

```js
assert.deepEqual(getResults_().res1[0][0], { name: 'John Rappl', id: 'E21437', salary: 47000, dept: 'D050'});
```

Für `topRankPerGroup(10, testData1, 'dept', 'salary')` sollte das letzte Ergebnis in der letzten Gruppe `{ name: 'Adam Smith', id: 'E63535', salary: 18000, dept: 'D202' }` sein.

```js
assert.deepEqual(getResults_().res1[3][3], { name: 'Adam Smith', id: 'E63535', salary: 18000, dept: 'D202' });
```

`topRankPerGroup(1, ...)` sollte nur das bestplatzierte Ergebnis pro Gruppe zurückgeben.

```js
const { res2 } = getResults_();
assert.equal(res2[2].length, 1);
```

`topRankPerGroup(2, ...)` sollte zwei bestplatzierte Ergebnisse pro Gruppe zurückgeben.

```js
const { res3 } = getResults_();
assert.equal(res3[2][1].name, 'Maze Runner');
```

# --seed--

## --seed-contents--

```js
function topRankPerGroup(n, data, groupName, rankName) {

  return true;
}
```

# --solutions--

```js
const collectDept = function (arrOfObj, groupName) {
  const collect = arrOfObj.reduce((rtnObj, obj) => {
    if (rtnObj[obj[groupName]] === undefined) {
      rtnObj[obj[groupName]] = [];
    }
    rtnObj[obj[groupName]].push(obj);
    return rtnObj;
  }, {} // initial value to reduce
  );

  return Object.keys(collect).sort().map(key => collect[key]);
};

const sortRank = function (arrOfRankArrs, rankName) {
  return arrOfRankArrs.map(item => item.sort((a, b) => {
    if (a[rankName] > b[rankName]) { return -1; }
    if (a[rankName] < b[rankName]) { return 1; }
    return 0;
  }));
};

function topRankPerGroup(n, data, groupName, rankName) {
  if (n < 0) { return; }
  return sortRank(collectDept(data, groupName),
    rankName).map(list => list.slice(0, n));
}
```
