---
id: 5951815dd895584b06884620
title: Kreise mit gegebenem Radius durch zwei Punkte
challengeType: 1
forumTopicId: 302231
dashedName: circles-of-given-radius-through-two-points
---

# --description--

Gegeben sind zwei Punkte in einer Ebene und ein Radius. Üblicherweise können zwei Kreise mit dem gegebenen Radius durch diese Punkte gezeichnet werden.

**Ausnahmen:**

<ul>
  <li>Ein Radius von null sollte so behandelt werden, dass er niemals Kreise beschreibt (außer im Fall, dass die Punkte übereinstimmen).</li>
  <li>Wenn die Punkte übereinstimmen, können unendlich viele Kreise mit dem Punkt auf ihrem Umfang gezeichnet werden, es sei denn, der Radius ist ebenfalls null, was die Kreise dann zu einem Punkt zusammenfallen lässt.</li>
  <li>Wenn die Punkte einen Durchmesser bilden, dann geben Sie einen einzelnen Kreis zurück.</li>
  <li>Wenn die Punkte zu weit auseinanderliegen, können keine Kreise gezeichnet werden.</li>
</ul>

# --instructions--

Implementieren Sie eine Funktion, die zwei Punkte und einen Radius entgegennimmt und die zwei Kreise durch diese Punkte zurückgibt. Für jeden resultierenden Kreis geben Sie die Koordinaten des Mittelpunkts an, gerundet auf vier Dezimalziffern. Geben Sie jede Koordinate als Array zurück und die Koordinaten als Array von Arrays.

**Für Randfälle geben Sie Folgendes zurück:**

<ul>
  <li>Wenn die Punkte auf dem Durchmesser liegen, geben Sie einen Punkt zurück. Wenn der Radius jedoch ebenfalls null ist, geben Sie <code>"Radius Zero"</code> zurück.</li>
  <li>Wenn die Punkte übereinstimmen, geben Sie <code>"Coincident point. Infinite solutions"</code> zurück.</li>
  <li>Wenn die Punkte weiter auseinanderliegen als der Durchmesser, geben Sie <code>"No intersection. Points further apart than circle diameter"</code> zurück.</li>
</ul>

**Beispiel-Eingaben:**

<pre>      p1                p2           r
0.1234, 0.9876    0.8765, 0.2345    2.0
0.0000, 2.0000    0.0000, 0.0000    1.0
0.1234, 0.9876    0.1234, 0.9876    2.0
0.1234, 0.9876    0.8765, 0.2345    0.5
0.1234, 0.9876    0.1234, 0.9876    0.0
</pre>

# --before-each--

```js
const testCases = [
  [[0.1234, 0.9876], [0.8765, 0.2345], 2.0],
  [[0.0000, 2.0000], [0.0000, 0.0000], 1.0],
  [[0.1234, 0.9876], [0.1234, 0.9876], 2.0],
  [[0.1234, 0.9876], [0.8765, 0.2345], 0.5],
  [[0.1234, 0.9876], [0.1234, 0.9876], 0.0]
];
const answers = [
  [[1.8631, 1.9742], [-0.8632, -0.7521]],
  [0, 1],
  'Coincident point. Infinite solutions',
  'No intersection. Points further apart than circle diameter',
  'Radius Zero'
];
```

# --hints--

`getCircles` sollte eine Funktion sein.

```js
assert(typeof getCircles === 'function');
```

`getCircles([0.1234, 0.9876], [0.8765, 0.2345], 2.0)` sollte `[[1.8631, 1.9742], [-0.8632, -0.7521]]` zurückgeben.

```js
assert.deepEqual(getCircles(...testCases[0]), answers[0]);
```

`getCircles([0.0000, 2.0000], [0.0000, 0.0000], 1.0)` sollte `[0, 1]` zurückgeben.

```js
assert.deepEqual(getCircles(...testCases[1]), answers[1]);
```

`getCircles([0.1234, 0.9876], [0.1234, 0.9876], 2.0)` sollte `Coincident point. Infinite solutions` zurückgeben.

```js
assert.deepEqual(getCircles(...testCases[2]), answers[2]);
```

`getCircles([0.1234, 0.9876], [0.8765, 0.2345], 0.5)` sollte `No intersection. Points further apart than circle diameter` zurückgeben.

```js
assert.deepEqual(getCircles(...testCases[3]), answers[3]);
```

`getCircles([0.1234, 0.9876], [0.1234, 0.9876], 0.0)` sollte `Radius Zero` zurückgeben.

```js
assert.deepEqual(getCircles(...testCases[4]), answers[4]);
```

# --seed--

## --seed-contents--

```js
function getCircles(...args) {

  return true;
}
```

# --solutions--

```js
const hDist = (p1, p2) => Math.hypot(...p1.map((e, i) => e - p2[i])) / 2;
const pAng = (p1, p2) => Math.atan(p1.map((e, i) => e - p2[i]).reduce((p, c) => c / p, 1));
const solveF = (p, r) => t => [parseFloat((r * Math.cos(t) + p[0]).toFixed(4)), parseFloat((r * Math.sin(t) + p[1]).toFixed(4))];
const diamPoints = (p1, p2) => p1.map((e, i) => parseFloat((e + (p2[i] - e) / 2).toFixed(4)));

function getCircles(...args) {
  const [p1, p2, s] = args;
  const solve = solveF(p1, s);
  const halfDist = hDist(p1, p2);

  let msg = [];
  switch (Math.sign(s - halfDist)) {
    case 0:
      msg = s ? diamPoints(p1, p2) :
        'Radius Zero';
      break;
    case 1:
      if (!halfDist) {
        msg = 'Coincident point. Infinite solutions';
      }
      else {
        const theta = pAng(p1, p2);
        const theta2 = Math.acos(halfDist / s);
        [1, -1].map(e => solve(theta + e * theta2)).forEach(
          e => msg.push(e));
      }
      break;
    case -1:
      msg = 'No intersection. Points further apart than circle diameter';
      break;
    default:
      msg = 'Reached the default';
  }
  return msg;
}
```
