---
id: 5951ed8945deab770972ae56
title: Türme von Hanoi
challengeType: 1
forumTopicId: 302341
dashedName: towers-of-hanoi
---

# --description--

Lösen Sie das Problem der Türme von Hanoi. Die Anzahl der Objekte wird als erster Parameter angegeben, gefolgt von den Strings, die verwendet werden, um jeden Stapel von Objekten zu identifizieren. Erstellen Sie ein verschachteltes Array, das die Liste der Züge enthält, `["source", "destination"]`.

Zum Beispiel führen die Parameter `(4, 'A', 'B', 'C')` zu einem verschachtelten Array von Zügen `[['A', 'C'], ['B', 'A']]`, das anzeigt, dass der erste Zug darin bestand, ein Objekt vom Stapel `A` zu `C` zu bewegen, und der zweite Zug darin bestand, ein Objekt vom Stapel `B` zu `A` zu bewegen.

# --instructions--

Schreiben Sie eine Funktion, die die Züge zurückgibt, um die Objekte in einem verschachtelten Array zu stapeln.

# --before-each--

```js
const res3Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B']];
const res7First10Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B'], ['A', 'C'], ['B', 'C'], ['B', 'A']];

function getTowerResults_() {
  return {
    res3: towerOfHanoi(3, 'A', 'B', 'C'),
    res5: towerOfHanoi(5, 'X', 'Y', 'Z'),
    res7: towerOfHanoi(7, 'A', 'B', 'C')
  };
}
```

# --hints--

`towerOfHanoi` sollte eine Funktion sein.

```js
assert(typeof towerOfHanoi === 'function');
```

`towerOfHanoi(3, ...)` sollte 7 Züge zurückgeben.

```js
assert(getTowerResults_().res3.length === 7);
```

`towerOfHanoi(3, 'A', 'B', 'C')` sollte `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B']]` zurückgeben.

```js
assert.deepEqual(getTowerResults_().res3, res3Moves);
```

`towerOfHanoi(5, "X", "Y", "Z")` 10. Zug sollte Y -> X sein.

```js
const { res5 } = getTowerResults_();
assert.deepEqual(res5[9], ['Y', 'X']);
```

`towerOfHanoi(7, 'A', 'B', 'C')` die ersten zehn Züge sollten `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B'], ['A','C'], ['B','C'], ['B','A']]` sein.

```js
const { res7 } = getTowerResults_();
assert.deepEqual(res7.slice(0, 10), res7First10Moves);
```

# --seed--

## --seed-contents--

```js
function towerOfHanoi(n, a, b, c) {

  return [[]];
}
```

# --solutions--

```js
function towerOfHanoi(n, a, b, c) {
  const res = [];
  towerOfHanoiHelper(n, a, c, b, res);
  return res;
}

function towerOfHanoiHelper(n, a, b, c, res) {
  if (n > 0) {
    towerOfHanoiHelper(n - 1, a, c, b, res);
    res.push([a, c]);
    towerOfHanoiHelper(n - 1, b, a, c, res);
  }
}
```
