---
id: 595608ff8bcd7a50bd490181
title: Hailstone-Folge
challengeType: 1
forumTopicId: 302279
dashedName: hailstone-sequence
---

# --description--

Die Hailstone-Folge von Zahlen kann aus einer positiven Startzahl, `n`, erzeugt werden durch:

- Wenn `n` gleich `1` ist, endet die Folge
- Wenn `n` gleich `even` ist, dann ist das nächste `n` der Folge `= n/2`
- Wenn `n` gleich `odd` ist, dann ist das nächste `n` der Folge `= (3 * n) + 1`

Die (unbewiesene) Collatz-Vermutung besagt, dass die Hailstone-Folge für jede Startzahl immer endet.

Die Hailstone-Folge ist auch als Hailstone-Zahlen bekannt (weil die Werte meist mehreren Ab- und Aufstiegen wie Hagelkörner in einer Wolke unterliegen) oder als Collatz-Folge.

# --instructions--

1. Erstellen Sie eine Routine, um die Hailstone-Folge für eine Zahl zu erzeugen
2. Ihre Funktion soll ein Array zurückgeben mit der Zahl kleiner als `limit`, die die längste Hailstone-Folge hat, und der Länge dieser Folge. (Aber zeigen Sie nicht die tatsächliche Folge!)

# --hints--

`hailstoneSequence` sollte eine Funktion sein.

```js
assert(typeof hailstoneSequence === 'function');
```

`hailstoneSequence(30)` sollte ein Array zurückgeben.

```js
assert(Array.isArray(hailstoneSequence(30)));
```

`hailstoneSequence(30)` sollte `[27, 112]` zurückgeben.

```js
assert.deepEqual(hailstoneSequence(30), [27, 112]);
```

`hailstoneSequence(50000)` sollte `[35655, 324]` zurückgeben.

```js
assert.deepEqual(hailstoneSequence(50000), [35655, 324]);
```

`hailstoneSequence(100000)` sollte `[77031, 351]` zurückgeben.

```js
assert.deepEqual(hailstoneSequence(100000), [77031, 351]);
```

# --seed--

## --seed-contents--

```js
function hailstoneSequence(limit) {
  const res = [];


  return res;
}
```

# --solutions--

```js
function hailstoneSequence (limit) {
  function hailstone(n) {
    const seq = [n];
    while (n > 1) {
      n = n % 2 ? 3 * n + 1 : n / 2;
      seq.push(n);
    }
    return seq;
  }

  let n = 0;
  let max = 0;
  for (let i = limit; --i;) {
    const seq = hailstone(i);
    const sLen = seq.length;

    if (sLen > max) {
      n = i;
      max = sLen;
    }
  }

  return [n, max];
}
```
