---
id: 5956795bc9e2c415eb244de1
title: Hash-Join
challengeType: 1
forumTopicId: 302284
dashedName: hash-join
---

# --description--

Ein Inner Join ist eine Operation, die zwei Datentabellen basierend auf übereinstimmenden Spaltenwerten zu einer Tabelle kombiniert. Die einfachste Methode, diese Operation zu implementieren, ist der Algorithmus des verschachtelten Schleifen-Joins, aber eine skalierbarere Alternative ist der Hash-Join-Algorithmus.

Der „Hash-Join“-Algorithmus besteht aus zwei Schritten:

<ol>
  <li><strong>Hash-Phase:</strong> Erstellen Sie eine Multimap aus einer der beiden Tabellen, die jeden Join-Spaltenwert auf alle Zeilen abbildet, die diesen enthalten.</li>
  <ul>
    <li>Die Multimap muss eine hash-basierte Suche unterstützen, die besser skaliert als eine einfache lineare Suche, denn genau das ist der Zweck dieses Algorithmus.</li>
    <li>Idealerweise sollten wir die Multimap für die kleinere Tabelle erstellen, um die Erstellungszeit und den Speicherbedarf zu minimieren.</li>
  </ul>
  <li><strong>Join-Phase:</strong> Scannen Sie die andere Tabelle und finden Sie passende Zeilen, indem Sie in der zuvor erstellten Multimap nachschlagen.</li>
</ol>

In Pseudocode könnte der Algorithmus wie folgt ausgedrückt werden:

<pre><strong>let</strong> <i>A</i> = die erste Eingabetabelle (oder idealerweise die größere)
<strong>let</strong> <i>B</i> = die zweite Eingabetabelle (oder idealerweise die kleinere)
<strong>let</strong> <i>j<sub>A</sub></i> = die Join-Spalten-ID der Tabelle <i>A</i>
<strong>let</strong> <i>j<sub>B</sub></i> = die Join-Spalten-ID der Tabelle <i>B</i>
<strong>let</strong> <i>M<sub>B</sub></i> = eine Multimap, die einzelne Werte auf mehrere Zeilen der Tabelle <i>B</i> abbildet (anfangs leer)
<strong>let</strong> <i>C</i> = die Ausgabetabelle (anfangs leer)
<strong>für jede</strong> Zeile <i>b</i> in Tabelle <i>B</i>:
  <strong>füge</strong> <i>b</i> in die Multimap <i>M<sub>B</sub></i> unter dem Schlüssel <i>b(j<sub>B</sub>)</i> ein
<strong>für jede</strong> Zeile <i>a</i> in Tabelle <i>A</i>:
  <strong>für jede</strong> Zeile <i>b</i> in der Multimap <i>M<sub>B</sub></i> unter dem Schlüssel <i>a(j<sub>A</sub>)</i>:
    <strong>let</strong> <i>c</i> = die Verkettung der Zeilen <i>a</i> und <i>b</i>
    <strong>füge</strong> Zeile <i>c</i> in Tabelle <i>C</i> ein
</pre>

# --instructions--

Implementieren Sie den „Hash-Join“-Algorithmus als Funktion und zeigen Sie, dass er den unten aufgeführten Testfall besteht. Die Funktion soll zwei Arrays von Objekten akzeptieren und ein Array kombinierter Objekte zurückgeben.

**Eingabe**

<table>
  <tr>
    <td style="padding: 4px; margin: 5px;">
      <table style="border:none; border-collapse:collapse;">
        <tr>
          <td style="border:none"><i>A =</i></td>
          <td style="border:none">
            <table>
              <tr>
                <th style="padding: 4px; margin: 5px;">Alter</th>
                <th style="padding: 4px; margin: 5px;">Name</th>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">27</td>
                <td style="padding: 4px; margin: 5px;">Jonah</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">18</td>
                <td style="padding: 4px; margin: 5px;">Alan</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">28</td>
                <td style="padding: 4px; margin: 5px;">Glory</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">18</td>
                <td style="padding: 4px; margin: 5px;">Popeye</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">28</td>
                <td style="padding: 4px; margin: 5px;">Alan</td>
              </tr>
            </table>
          </td>
          <td style="border:none; padding-left:1.5em;" rowspan="2"></td>
          <td style="border:none"><i>B =</i></td>
          <td style="border:none">
            <table>
              <tr>
                <th style="padding: 4px; margin: 5px;">Charakter</th>
                <th style="padding: 4px; margin: 5px;">Nemesis</th>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">Jonah</td>
                <td style="padding: 4px; margin: 5px;">Whales</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">Jonah</td>
                <td style="padding: 4px; margin: 5px;">Spiders</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">Alan</td>
                <td style="padding: 4px; margin: 5px;">Ghosts</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">Alan</td>
                <td style="padding: 4px; margin: 5px;">Zombies</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">Glory</td>
                <td style="padding: 4px; margin: 5px;">Buffy</td>
              </tr>
            </table>
          </td>
        </tr>
        <tr>
          <td style="border:none">
            <i>j<sub>A</sub> =</i>
          </td>
          <td style="border:none">
            <i><code>Name</code> (also Spalte 1)</i>
          </td>
          <td style="border:none">
            <i>j<sub>B</sub> =</i>
          </td>
          <td style="border:none">
            <i><code>Character</code> (also Spalte 0)</i>
          </td>
        </tr>
      </table>
    </td>
  </tr>
</table>

**Ausgabe**

| A_alter | A_name | B_character | B_nemesis |
| ------- | ------ | ----------- | --------- |
| 27      | Jonah  | Jonah       | Whales    |
| 27      | Jonah  | Jonah       | Spiders   |
| 18      | Alan   | Alan        | Ghosts    |
| 18      | Alan   | Alan        | Zombies   |
| 28      | Glory  | Glory       | Buffy     |
| 28      | Alan   | Alan        | Ghosts    |
| 28      | Alan   | Alan        | Zombies   |

Die Reihenfolge der Zeilen in der Ausgabetabelle ist nicht relevant.

# --hints--

`hashJoin` sollte eine Funktion sein.

```js
assert(typeof hashJoin === 'function');
```

`hashJoin([{ age: 27, name: "Jonah" }, { age: 18, name: "Alan" }, { age: 28, name: "Glory" }, { age: 18, name: "Popeye" }, { age: 28, name: "Alan" }], [{ character: "Jonah", nemesis: "Whales" }, { character: "Jonah", nemesis: "Spiders" }, { character: "Alan", nemesis: "Ghosts" }, { character:"Alan", nemesis: "Zombies" }, { character: "Glory", nemesis: "Buffy" }, { character: "Bob", nemesis: "foo" }])` sollte `[{"A_age": 27,"A_name": "Jonah", "B_character": "Jonah", "B_nemesis": "Whales"}, {"A_age": 27,"A_name": "Jonah", "B_character": "Jonah", "B_nemesis": "Spiders"}, {"A_age": 18,"A_name": "Alan", "B_character": "Alan", "B_nemesis": "Ghosts"}, {"A_age": 18,"A_name": "Alan", "B_character": "Alan", "B_nemesis": "Zombies"}, {"A_age": 28,"A_name": "Glory", "B_character": "Glory", "B_nemesis": "Buffy"}, {"A_age": 28,"A_name": "Alan", "B_character": "Alan", "B_nemesis": "Ghosts"}, {"A_age": 28,"A_name": "Alan", "B_character": "Alan", "B_nemesis": "Zombies"}]` zurückgeben.

```js
const hash1 = [
    { age: 27, name: 'Jonah' },
    { age: 18, name: 'Alan' },
    { age: 28, name: 'Glory' },
    { age: 18, name: 'Popeye' },
    { age: 28, name: 'Alan' }
];

const hash2 = [
    { character: 'Jonah', nemesis: 'Whales' },
    { character: 'Jonah', nemesis: 'Spiders' },
    { character: 'Alan', nemesis: 'Ghosts' },
    { character: 'Alan', nemesis: 'Zombies' },
    { character: 'Glory', nemesis: 'Buffy' },
    { character: 'Bob', nemesis: 'foo' }
];

const res = [
    { A_age: 27, A_name: 'Jonah', B_character: 'Jonah', B_nemesis: 'Whales' },
    { A_age: 27, A_name: 'Jonah', B_character: 'Jonah', B_nemesis: 'Spiders' },
    { A_age: 18, A_name: 'Alan', B_character: 'Alan', B_nemesis: 'Ghosts' },
    { A_age: 18, A_name: 'Alan', B_character: 'Alan', B_nemesis: 'Zombies' },
    { A_age: 28, A_name: 'Glory', B_character: 'Glory', B_nemesis: 'Buffy' },
    { A_age: 28, A_name: 'Alan', B_character: 'Alan', B_nemesis: 'Ghosts' },
    { A_age: 28, A_name: 'Alan', B_character: 'Alan', B_nemesis: 'Zombies' }
];

const bench1 = [{ name: 'u2v7v', num: 1 }, { name: 'n53c8', num: 10 }, { name: 'oysce', num: 9 }, { name: '0mto2s', num: 1 }, { name: 'vkh5id', num: 4 }, { name: '5od0cf', num: 8 }, { name: 'uuulue', num: 10 }, { name: '3rgsbi', num: 9 }, { name: 'kccv35r', num: 4 }, { name: '80un74', num: 9 }, { name: 'h4pp3', num: 6 }, { name: '51bit', num: 7 }, { name: 'j9ndf', num: 8 }, { name: 'vf3u1', num: 10 }, { name: 'g0bw0om', num: 10 }, { name: 'j031x', num: 7 }, { name: 'ij3asc', num: 9 }, { name: 'byv83y', num: 8 }, { name: 'bjzp4k', num: 4 }, { name: 'f3kbnm', num: 10 }];
const bench2 = [{ friend: 'o8b', num: 8 }, { friend: 'ye', num: 2 }, { friend: '32i', num: 5 }, { friend: 'uz', num: 3 }, { friend: 'a5k', num: 4 }, { friend: 'uad', num: 7 }, { friend: '3w5', num: 10 }, { friend: 'vw', num: 10 }, { friend: 'ah', num: 4 }, { friend: 'qv', num: 7 }, { friend: 'ozv', num: 2 }, { friend: '9ri', num: 10 }, { friend: '7nu', num: 4 }, { friend: 'w3', num: 9 }, { friend: 'tgp', num: 8 }, { friend: 'ibs', num: 1 }, { friend: 'ss7', num: 6 }, { friend: 'g44', num: 9 }, { friend: 'tab', num: 9 }, { friend: 'zem', num: 10 }];
assert.deepEqual(hashJoin(hash1, hash2), res);
```

# --seed--

## --seed-contents--

```js
function hashJoin(hash1, hash2) {

  return [];
}
```

# --solutions--

```js
function hashJoin(hash1, hash2) {
  const hJoin = (tblA, tblB, strJoin) => {
    const [jA, jB] = strJoin.split('=');
    const M = tblB.reduce((a, x) => {
      const id = x[jB];
      return (
        a[id] ? a[id].push(x) : (a[id] = [x]),
        a
      );
    }, {});

    return tblA.reduce((a, x) => {
      const match = M[x[jA]];
      return match ? (
                a.concat(match.map(row => dictConcat(x, row)))
            ) : a;
    }, []);
  };

  const dictConcat = (dctA, dctB) => {
    const ok = Object.keys;
    return ok(dctB).reduce(
            (a, k) => (a[`B_${k}`] = dctB[k]) && a,
            ok(dctA).reduce(
                (a, k) => (a[`A_${k}`] = dctA[k]) && a, {}
            )
        );
  };

  return hJoin(hash1, hash2, 'name=character');
}
```
