---
id: 59667989bf71cf555dd5d2ff
title: S-Ausdrücke
challengeType: 1
forumTopicId: 302303
dashedName: s-expressions
---

# --description--

<a href="https://rosettacode.org/wiki/S-expressions" target="_blank" rel="noopener noreferrer nofollow">S-Ausdrücke</a> sind eine praktische Möglichkeit, Daten zu parsen und zu speichern.

# --instructions--

Schreiben Sie einen einfachen Reader/Parser für S-Ausdrücke, der zitierte und unzitierte Strings, ganze Zahlen und Fließkommazahlen verarbeitet.

Die Funktion soll einen einzelnen, aber verschachtelten S-Ausdruck aus einem String lesen und als (verschachteltes) Array zurückgeben.

Zeilenumbrüche und andere Leerzeichen können ignoriert werden, sofern sie nicht innerhalb eines zitierten Strings enthalten sind.

"`()`" innerhalb zitierter Strings werden nicht interpretiert, sondern als Teil des Strings behandelt.

Das Verarbeiten von escaped quotes innerhalb eines Strings ist optional; daher kann "`(foo"bar)`" entweder als String "`foo"bar`" behandelt werden oder als Fehler.

Dafür muss der Reader `\` für Escaping nicht erkennen, sollte aber zusätzlich Zahlen erkennen, falls die Sprache entsprechende Datentypen hat.

Beachten Sie, dass mit Ausnahme von `()"` (bzw. `\` wenn Escaping unterstützt wird) und Leerzeichen keine Sonderzeichen existieren. Alles andere ist ohne Anführungszeichen erlaubt.

Der Reader sollte in der Lage sein, folgende Eingabe zu lesen:

<pre>((data "quoted data" 123 4.5)
(data (!@# (4.5) "(more" "data)")))
</pre>

und daraus eine native Datenstruktur erzeugen.

# --before-each--

```js
const simpleSExpr = '(data1 data2 data3)';
const simpleSolution = ['data1', 'data2', 'data3'];

const basicSExpr = '((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))';
const basicSolution = [["data","\"quoted data\"",123,4.5],["data",["!@#",[4.5],"\"(more\"","\"data)\""]]];
```

# --hints--

`parseSexpr` sollte eine Funktion sein.

```js
assert(typeof parseSexpr === 'function');
```

`parseSexpr('(data1 data2 data3)')` sollte `['data1', 'data2', 'data3']` zurückgeben.

```js
assert.deepEqual(parseSexpr(simpleSExpr), simpleSolution);
```

`parseSexpr('((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))')` sollte `[['data', '"quoted data"', 123, 4.5], ['data', ['!@#', [4.5], '"(more"', '"data)"']]]` zurückgeben.

```js
assert.deepEqual(parseSexpr(basicSExpr), basicSolution);
```

# --seed--

## --seed-contents--

```js
function parseSexpr(str) {

  return true;
}
```

# --solutions--

```js
function parseSexpr(str) {
  const t = str.match(/\s*("[^"]*"|\(|\)|"|[^\s()"]+)/g);
  for (var o, c = 0, i = t.length - 1; i >= 0; i--) {
    var n,
      ti = t[i].trim();
    if (ti == '"') return;
    else if (ti == '(') t[i] = '[', c += 1;
    else if (ti == ')') t[i] = ']', c -= 1;
    else if ((n = +ti) == ti) t[i] = n;
    else t[i] = `'${ti.replace('\'', '\\\'')}'`;
    if (i > 0 && ti != ']' && t[i - 1].trim() != '(') t.splice(i, 0, ',');
    if (!c) if (!o) o = true; else return;
  }
  return c ? undefined : eval(t.join(''));
}
```
