---
id: 59669d08d75b60482359409f
title: Datumsformat
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

Geben Sie ein Array mit zwei Datumsstrings des aktuellen Datums zurück, die folgende Vorgaben erfüllen:

- Die Reihenfolge des Datums im ersten String soll Jahr, Monat und Tag sein, getrennt durch Bindestriche (`-`).
- Das Jahr im ersten String soll vierstellig sein.
- Monat und Tag im ersten String dürfen keine führenden Nullen enthalten.
- Die Wochentags- und Monatsnamen im zweiten String dürfen nicht abgekürzt sein.
- Der Tag im zweiten String darf keine führenden Nullen enthalten.

Beispielausgaben:

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

`getDateFormats` sollte eine Funktion sein.

```js
assert(typeof getDateFormats === 'function');
```

`getDateFormats` sollte ein Objekt zurückgeben.

```js
assert(typeof getDateFormats() === 'object');
```

`getDateFormats` sollte ein Array mit 2 Elementen zurückgeben.

```js
assert(getDateFormats().length === 2);
```

`getDateFormats` sollte das korrekte Datum im richtigen Format zurückgeben.

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
