---
id: 596e414344c3b2872167f0fe
title: Komma-Streitigkeiten
challengeType: 1
forumTopicId: 302234
dashedName: comma-quibbling
---

# --description--

<a href="https://rosettacode.org/wiki/Comma_quibbling" target="_blank" rel="noopener noreferrer nofollow">Comma quibbling</a> ist eine Aufgabe, die ursprünglich von Eric Lippert in seinem Blog gestellt wurde.

# --instructions--

Schreiben Sie eine Funktion, die eine String-Ausgabe erzeugt, welche die Verkettung von Eingabewörtern aus einer Liste/Folge ist, wobei gilt:

<ol>
  <li>Eine Eingabe ohne Wörter erzeugt die Ausgabekette mit nur den beiden Klammerzeichen (<code>"{}"</code>)</li>
  <li>Eine Eingabe mit nur einem Wort, z. B. <code>["ABC"]</code>, erzeugt die Ausgabekette mit dem Wort innerhalb der beiden Klammern, z. B. <code>"{ABC}"</code></li>
  <li>Eine Eingabe mit zwei Wörtern, z. B. <code>["ABC", "DEF"]</code>, erzeugt die Ausgabekette mit den beiden Wörtern innerhalb der beiden Klammern, getrennt durch den String <code>" and "</code>, z. B. <code>"{ABC and DEF}"</code></li>
  <li>Eine Eingabe mit drei oder mehr Wörtern, z. B. <code>["ABC", "DEF", "G", "H"]</code>, erzeugt die Ausgabekette mit allen bis auf das letzte Wort, getrennt durch <code>", "</code>, wobei das letzte Wort durch <code>" and "</code> getrennt ist und alles innerhalb der Klammern steht; z. B. <code>"{ABC, DEF, G and H}"</code></li>
</ol>

Testen Sie Ihre Funktion mit der folgenden Reihe von Eingaben und zeigen Sie Ihre Ausgabe hier auf dieser Seite:

<ul>
  <li>[] # (Keine Eingabewörter).</li>
  <li>["ABC"]</li>
  <li>["ABC", "DEF"]</li>
  <li>["ABC", "DEF", "G", "H"]</li>
</ul>

**Hinweis:** Gehen Sie bei dieser Aufgabe davon aus, dass die Wörter nicht-leere Strings aus Großbuchstaben sind.

# --before-each--

```js
const testCases = [[], ["ABC"], ["ABC", "DEF"], ["ABC", "DEF", "G", "H"]];
const results = ["{}", "{ABC}", "{ABC and DEF}", "{ABC, DEF, G and H}"];
```

# --hints--

`quibble` sollte eine Funktion sein.

```js
assert(typeof quibble === 'function');
```

`quibble(["ABC"])` sollte einen String zurückgeben.

```js
assert(typeof quibble(['ABC']) === 'string');
```

`quibble([])` sollte "{}" zurückgeben.

```js
assert.equal(quibble(testCases[0]), results[0]);
```

`quibble(["ABC"])` sollte `"{ABC}"` zurückgeben.

```js
assert.equal(quibble(testCases[1]), results[1]);
```

`quibble(["ABC", "DEF"])` sollte `"{ABC and DEF}"` zurückgeben.

```js
assert.equal(quibble(testCases[2]), results[2]);
```

`quibble(["ABC", "DEF", "G", "H"])` sollte `"{ABC, DEF, G and H}"` zurückgeben.

```js
assert.equal(quibble(testCases[3]), results[3]);
```

# --seed--

## --seed-contents--

```js
function quibble(words) {

  return true;
}
```

# --solutions--

```js
function quibble(words) {
  return "{" +
    words.slice(0, words.length - 1).join(", ") +
   (words.length > 1 ? " and " : "") +
   (words[words.length - 1] || '') +
  "}";
}
```
