---
id: 596fda99c69f779975a1b67d
title: Vorkommen eines Teilstrings zählen
challengeType: 1
forumTopicId: 302237
dashedName: count-occurrences-of-a-substring
---

# --description--

Erstellen Sie eine Funktion oder zeigen Sie eine eingebaute Funktion, um die Anzahl der nicht überlappenden Vorkommen eines Teilstrings innerhalb eines Strings zu zählen.

Die Funktion soll zwei Argumente entgegennehmen:

<ul>
  <li>das erste Argument ist der zu durchsuchende String, und</li>
  <li>das zweite ein zu suchender Teilstring.</li>
</ul>

Sie soll eine ganze Zahl als Ergebnis zurückgeben.

Die Übereinstimmung soll die größtmögliche Anzahl nicht überlappender Treffer liefern.

Im Allgemeinen bedeutet dies im Wesentlichen, von links nach rechts oder von rechts nach links zu suchen.

# --before-each--

```js
const testCases = ['the three truths', 'ababababab', 'abaabba*bbaba*bbab'];
const searchString = ['th', 'abab', 'a*b'];
const results = [3, 2, 2];
```

# --hints--

`countSubstring` sollte eine Funktion sein.

```js
assert(typeof countSubstring === 'function');
```

`countSubstring("the three truths", "th")` sollte `3` zurückgeben.

```js
assert.equal(countSubstring(testCases[0], searchString[0]), results[0]);
```

`countSubstring("ababababab", "abab")` sollte `2` zurückgeben.

```js
assert.equal(countSubstring(testCases[1], searchString[1]), results[1]);
```

`countSubstring("abaabba*bbaba*bbab", "a*b")` sollte `2` zurückgeben.

```js
assert.equal(countSubstring(testCases[2], searchString[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function countSubstring(str, subStr) {

  return true;
}
```

# --solutions--

```js
function countSubstring(str, subStr) {
  const escapedSubStr = subStr.replace(/[.+*?^$[\]{}()|/]/g, '\\$&');
  const matches = str.match(new RegExp(escapedSubStr, 'g'));
  return matches ? matches.length : 0;
}
```
