---
id: 596fd036dc1ab896c5db98b1
title: Sekunden in zusammengesetzte Dauer umwandeln
challengeType: 1
forumTopicId: 302236
dashedName: convert-seconds-to-compound-duration
---

# --description--

Implementieren Sie eine Funktion, die:

<ul>
  <li>eine positive ganze Zahl als Eingabe erhält, die eine Dauer in Sekunden darstellt (z. B. <code>100</code>), und</li>
  <li>einen String zurückgibt, der dieselbe Dauer in Wochen, Tagen, Stunden, Minuten und Sekunden zerlegt darstellt, wie unten beschrieben (z. B. <code>1 min, 40 sec</code>).</li>
</ul>

Zeigen Sie, dass sie die folgenden drei Testfälle besteht:

<div style='font-size:115%; font-weight: bold;'>Testfälle</div>

| Eingabezahl | Ausgabezahl                         |
| ------------ | ------------------------------------- |
| 7259         | <code>2 hr, 59 sec</code>             |
| 86400        | <code>1 d</code>                      |
| 6000000      | <code>9 wk, 6 d, 10 hr, 40 min</code> |

<div style="font-size:115%; font-weight: bold;">Details</div>
<ul>
  <li>
    Die folgenden fünf Einheiten sollen verwendet werden:

| Einheit | Im Output verwendete Abkürzung | Umrechnung            |
| ------ | ------------------------------ | --------------------- |
| Woche  | <code>wk</code>                | 1 Woche = 7 Tage      |
| Tag    | <code>d</code>                 | 1 Tag = 24 Stunden    |
| Stunde | <code>hr</code>                | 1 Stunde = 60 Minuten |
| Minute | <code>min</code>               | 1 Minute = 60 Sekunden|
| Sekunde| <code>sec</code>               | ---                   |

  </li>
  <li>
    Allerdings sollen <strong>nur</strong> Mengen mit einem Wert ungleich Null im Output erscheinen (z. B. <code>1 d</code> und nicht <code>0 wk, 1 d, 0 hr, 0 min, 0 sec</code>).
  </li>
  <li>
    Größere Einheiten sollen so weit wie möglich Vorrang vor kleineren haben (z. B. <code>2 min, 10 sec</code> und nicht <code>1 min, 70 sec</code> oder <code>130 sec</code>).
  </li>
  <li>
    Orientieren Sie sich an der Formatierung der Testfälle (Mengen sortiert von der größten zur kleinsten Einheit und durch Komma plus Leerzeichen getrennt; Wert und Einheit jeder Menge durch Leerzeichen getrennt).
  </li>
</ul>

# --before-each--

```js
const testCases = [7259, 86400, 6000000];
const results = ['2 hr, 59 sec', '1 d', '9 wk, 6 d, 10 hr, 40 min'];
```

# --hints--

`convertSeconds` sollte eine Funktion sein.

```js
assert(typeof convertSeconds === 'function');
```

`convertSeconds(7259)` sollte `2 hr, 59 sec` zurückgeben.

```js
assert.equal(convertSeconds(testCases[0]), results[0]);
```

`convertSeconds(86400)` sollte `1 d` zurückgeben.

```js
assert.equal(convertSeconds(testCases[1]), results[1]);
```

`convertSeconds(6000000)` sollte `9 wk, 6 d, 10 hr, 40 min` zurückgeben.

```js
assert.equal(convertSeconds(testCases[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function convertSeconds(sec) {

  return true;
}
```

# --solutions--

```js
function convertSeconds(sec) {
  const localNames = ['wk', 'd', 'hr', 'min', 'sec'];
  // compoundDuration :: [String] -> Int -> String
  const compoundDuration = (labels, intSeconds) =>
    weekParts(intSeconds)
    .map((v, i) => [v, labels[i]])
    .reduce((a, x) =>
      a.concat(x[0] ? [`${x[0]} ${x[1] || '?'}`] : []), []
    )
    .join(', ');

    // weekParts :: Int -> [Int]
  const weekParts = intSeconds => [0, 7, 24, 60, 60]
    .reduceRight((a, x) => {
      const r = a.rem;
      const mod = x !== 0 ? r % x : r;

      return {
        rem: (r - mod) / (x || 1),
        parts: [mod].concat(a.parts)
      };
    }, {
      rem: intSeconds,
      parts: []
    })
    .parts;

  return compoundDuration(localNames, sec);
}
```
