---
id: 59713bd26bdeb8a594fb9413
title: Münzen zählen
challengeType: 1
forumTopicId: 302238
dashedName: count-the-coins
---

# --description--

Es gibt vier Arten von gängigen Münzen in der US-Währung:

<ul>
  <li>Quarters (25 Cent)</li>
  <li>Dimes (10 Cent)</li>
  <li>Nickels (5 Cent) und</li>
  <li>Pennies (1 Cent)</li>
</ul>

<p>Es gibt sechs Möglichkeiten, 15 Cent mit diesen Münzen zu wechseln:</p>

<ul>
  <li>Ein Dime und ein Nickel</li>
  <li>Ein Dime und 5 Pennies</li>
  <li>3 Nickels</li>
  <li>2 Nickels und 5 Pennies</li>
  <li>Ein Nickel und 10 Pennies</li>
  <li>15 Pennies</li>
</ul>

# --instructions--

Implementieren Sie eine Funktion, um zu bestimmen, wie viele Möglichkeiten es gibt, für eine gegebene Eingabe, `cents`, die einen Betrag in US-Pennies darstellt, mit diesen gängigen Münzen zu wechseln.

# --hints--

`countCoins` sollte eine Funktion sein.

```js
assert(typeof countCoins === 'function');
```

`countCoins(15)` sollte `6` zurückgeben.

```js
assert.equal(countCoins(15), 6);
```

`countCoins(85)` sollte `163` zurückgeben.

```js
assert.equal(countCoins(85), 163);
```

`countCoins(100)` sollte `242` zurückgeben.

```js
assert.equal(countCoins(100), 242);
```

# --seed--

## --seed-contents--

```js
function countCoins(cents) {

  return true;
}
```

# --solutions--

```js
function countCoins(cents) {
  const operands = [1, 5, 10, 25];
  const targetsLength = cents + 1;
  const operandsLength = operands.length;
  const t = [1];

  for (let a = 0; a < operandsLength; a++) {
    for (let b = 1; b < targetsLength; b++) {
      // initialise undefined target
      t[b] = t[b] ? t[b] : 0;

      // accumulate target + operand ways
      t[b] += (b < operands[a]) ? 0 : t[b - operands[a]];
    }
  }

  return t[targetsLength - 1];
}
```
