---
id: 597f1e7fbc206f0e9ba95dc4
title: Faktoren einer ganzen Zahl
challengeType: 1
forumTopicId: 302265
dashedName: factors-of-an-integer
---

# --description--

Schreiben Sie eine Funktion, die die Faktoren einer positiven ganzen Zahl als Array zurückgibt.

Diese Faktoren sind die positiven ganzen Zahlen, durch die die zu faktorisierende Zahl geteilt werden kann, sodass ein positives ganzzahliges Ergebnis entsteht.

# --before-each--

```js
const ans=[[1,3,5,9,15,45],[1,53],[1,2,4,8,16,32,64]];
```

# --hints--

`factors` sollte eine Funktion sein.

```js
assert(typeof factors === 'function');
```

`factors(45)` sollte `[1,3,5,9,15,45]` zurückgeben.

```js
assert.deepEqual(factors(45), ans[0]);
```

`factors(53)` sollte `[1,53]` zurückgeben.

```js
assert.deepEqual(factors(53), ans[1]);
```

`factors(64)` sollte `[1,2,4,8,16,32,64]` zurückgeben.

```js
assert.deepEqual(factors(64), ans[2]);
```

# --seed--

## --seed-contents--

```js
function factors(num) {

}
```

# --solutions--

```js
function factors(num)
{
 let n_factors = [], i, sqr=Math.floor(Math.sqrt(num));

 for (i = 1; i <=sqr ; i += 1)
  if (num % i === 0)
  {
   n_factors.push(i);
   if (num / i !== i)
    n_factors.push(num / i);
  }
 n_factors.sort(function(a, b){return a - b;});
 return n_factors;
}
```
