---
id: 599d1566a02b571412643b84
title: Äthiopische Multiplikation
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

Die äthiopische Multiplikation ist eine Methode, ganze Zahlen nur mit Addition, Verdopplung und Halbierung zu multiplizieren.

**Methode:**

<ol>
  <li>Nehmen Sie zwei Zahlen, die multipliziert werden sollen, und schreiben Sie sie oben in zwei Spalten</li>
  <li>Halbieren Sie in der linken Spalte wiederholt die letzte Zahl, verwerfen Sie dabei alle Reste, und schreiben Sie das Ergebnis unter die letzte Zahl in derselben Spalte, bis Sie den Wert <code>1</code> schreiben</li>
  <li>Verdoppeln Sie in der rechten Spalte wiederholt die letzte Zahl und schreiben Sie das Ergebnis darunter. Stoppen Sie, wenn Sie eine Zahl in derselben Zeile hinzufügen, in der die linke Spalte <code>1</code> zeigt</li>
  <li>Untersuchen Sie die erzeugte Tabelle und streichen Sie jede Zeile, in der der Wert in der linken Spalte gerade ist</li>
  <li>Summieren Sie die verbleibenden Werte in der rechten Spalte, um das Ergebnis der Multiplikation der ursprünglichen beiden Zahlen zu erhalten</li>
</ol>

**Zum Beispiel:** `17 × 34`

<pre>17   34
</pre>

Halbierung der ersten Spalte:

<pre>17   34
8
4
2
1
</pre>

Verdopplung der zweiten Spalte:

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

Streichen Sie die Zeilen, deren erste Zelle gerade ist:

<pre>17   34
8    <strike>68</strike>
4   <strike>136</strike>
2   <strike>272</strike>
1   544
</pre>

Summieren Sie die verbleibenden Zahlen in der rechten Spalte:

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

Also ist `17` multipliziert mit `34` nach der äthiopischen Methode `578`.

# --instructions--

Die Aufgabe ist, drei benannte Funktionen/Methoden/Prozeduren/Subroutinen zu definieren:

<ol>
  <li>eine, um eine ganze Zahl zu halbieren,</li>
  <li>eine, um eine ganze Zahl zu verdoppeln, und</li>
  <li>eine, die angibt, ob eine ganze Zahl gerade ist</li>
</ol>

Verwenden Sie diese Funktionen, um eine Funktion zu erstellen, die die äthiopische Multiplikation durchführt.

<!-- markdownlint-disable MD046-->

# --hints--

`eth_mult` sollte eine Funktion sein.

```js
assert(typeof eth_mult === 'function');
```

`eth_mult(17,34)` sollte `578` zurückgeben.

```js
assert.equal(eth_mult(17, 34), 578);
```

`eth_mult(23,46)` sollte `1058` zurückgeben.

```js
assert.equal(eth_mult(23, 46), 1058);
```

`eth_mult(12,27)` sollte `324` zurückgeben.

```js
assert.equal(eth_mult(12, 27), 324);
```

`eth_mult(56,98)` sollte `5488` zurückgeben.

```js
assert.equal(eth_mult(56, 98), 5488);
```

`eth_mult(63,74)` sollte `4662` zurückgeben.

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
