---
id: 59e0a8df964e4540d5abe599
title: Brain**** ausführen
challengeType: 1
forumTopicId: 302261
dashedName: execute-brain
---

# --description--

Schreiben Sie eine Funktion, um einen Brain****-Interpreter zu implementieren. Die Funktion erhält einen String als Parameter und soll einen String als Ausgabe zurückgeben. Weitere Details finden Sie unten:

RCBF ist eine Sammlung von <a href="https://rosettacode.org/wiki/Brainf***" target="_blank" rel="noopener noreferrer nofollow">Brainf***</a>-Compilern und -Interpretern, die für Rosetta Code in verschiedenen Sprachen geschrieben wurden.

Unten finden Sie Links zu jeder der Versionen von RCBF.

Eine Implementierung muss nur die folgenden Befehle korrekt implementieren:

| Befehl           | Beschreibung                                                                        |
| ---------------- | ---------------------------------------------------------------------------------- |
| <code>></code>    | Verschiebt den Zeiger nach rechts                                                  |
| <code>&lt;</code> | Verschiebt den Zeiger nach links                                                   |
| <code>+</code>    | Erhöht die Speicherzelle unter dem Zeiger                                         |
| <code>-</code>    | Verringert die Speicherzelle unter dem Zeiger                                     |
| <code>.</code>    | Gibt das durch die Speicherzelle am Zeiger dargestellte Zeichen aus                |
| <code>,</code>    | Nimmt ein Zeichen als Eingabe und speichert es in der Speicherzelle am Zeiger      |
| <code>\[</code>   | Springt hinter das passende <code>]</code>, wenn die Speicherzelle unter dem Zeiger 0 ist |
| <code>]</code>    | Springt zurück zum passenden <code>\[</code>, wenn die Speicherzelle unter dem Zeiger ungleich 0 ist |

Jede Zellgröße ist erlaubt, EOF (*E*nd-*O*-*F*ile) Unterstützung ist optional, ebenso wie begrenzter oder unbegrenzter Speicher.

# --hints--

`brain(globalThis.hello)` sollte einen String zurückgeben

```js
assert(typeof brain(globalThis.hello) === 'string');
```

`brain("++++++[>++++++++++<-]>+++++.")` sollte "A" zurückgeben

```js
assert.equal(brain('++++++[>++++++++++<-]>+++++.'), 'A');
```

`brain(globalThis.hello)` sollte `Hello World!\n` zurückgeben

```js
assert.equal(brain(globalThis.hello), 'Hello World!\n');
```

`brain(globalThis.fib)` sollte `1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89` zurückgeben

```js
assert.equal(brain(globalThis.fib), '1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89');
```

# --before-all--

```js
globalThis.fib=`+

++

+++

++++

+>+>>

>>++++

+++++++

++++++++

+++++++++

++++++++++

++++++>++++

++++++++++++

+++++++++++++

+++<<<<<<[>[>>

>>>>+>+<<<<<<<-

]>>>>>>>[<<<<<<<

+>>>>>>>-]<[>++++

++++++[-<-[>>+>+<<

<-]>>>[<<<+>>>-]+<[

>[-]<[-]]>[<<[>>>+<<

<-]>>[-]]<<]>>>[>>+>+

<<<-]>>>[<<<+>>>-]+<[>

[-]<[-]]>[<<+>>[-]]<<<<

<<<]>>>>>[++++++++++++++

+++++++++++++++++++++++++

+++++++++.[-]]++++++++++<[

->-<]>+++++++++++++++++++++

+++++++++++++++++++++++++++.

[-]<<<<<<<<<<<<[>>>+>+<<<<-]>

>>>[<<<<+>>>>-]<-[>>.>.<<<[-]]

<<[>>+>+<<<-]>>>[<<<+>>>-]<<[<+

>-]>[<+>-]<<<-]`;
globalThis.hello='++++++++[>++++[>++>+++>+++>+<<<<-]>+>+>->>+[<]<-]>>.>---.+++++++..+++.>>.<-.<.+++.------.--------.>>+.>++.'
```

# --seed--

## --seed-contents--

```js
function brain(prog) {

}
```

# --solutions--

```js
function brain(prog){
  var output="";
    var code; // formatted code
  var ip = 0; // current instruction within code
  var nest = 0; // current bracket nesting (for Out button)
  var ahead = []; // locations of matching brackets

  var data = [0]; // data array (mod by +, -)
  var dp = 0; // index into data (mod by <, >)

  var inp = 0; // current input character (fetch with ,)
  var quit = 0;
    var commands = {
    '>':function() { if (++dp >= data.length) data[dp]=0 },
    '<':function() { if (--dp < 0) quit++ },
    '+':function() { ++data[dp] },
    '-':function() { --data[dp] },
    '[':function() { if (!data[dp]) ip = ahead[ip]; else ++nest },
    ']':function() { if ( data[dp]) ip = ahead[ip]; else --nest },
    ',':function() {
        var c = document.getElementById("input").value.charCodeAt(inp++);
        data[dp] = isNaN(c) ? 0 : c; // EOF: other options are -1 or no change
    },
    '.':function() {
            output+=String.fromCharCode(data[dp]);
            /*var s = document.getElementById("output").innerHTML)
             + String.fromCharCode(data[dp]);
            s = s.replace(/\n/g,"<br>").replace(/ /g,"&amp;nbsp;");
            document.getElementById("output").innerHTML = s;*/
        },
    };

    let ar=prog.split('');
    var st = [], back, error = -1;
    for (ip=0; ip<ar.length; ip++) {
        switch(ar[ip]) {
        case '[':
            st.push(ip);
            break;
        case ']':
            if (st.length == 0) error = ip;
            back = st.pop();
            ahead[ip] = back;
            ahead[back] = ip;
            break;
        }
    }

    for(ip=0;ip<ar.length;ip++){
    if(commands.hasOwnProperty(ar[ip]))
          commands[ar[ip]]();
    }

    return output;
}
```
