---
id: 5a23c84252665b21eecc8017
title: Soundex
challengeType: 1
forumTopicId: 302320
dashedName: soundex
---

# --description--

**Soundex-Algorithmus** beschäftigt sich mit den *Absichten* der Wörter. Er erstellt eine Darstellung für ähnlich klingende Wörter.

Er wird für die Suche nach <em>Namen</em> und <em>Adressen</em> verwendet. Das bedeutet, dass sich die Person, die den <em>Namen</em> eingibt, auf die Aussprache konzentrieren kann, anstatt die Schreibweise der <em>Namen</em> zu korrigieren.

Zum Beispiel:

Wenn Sie den Namen `Quenci` zum ersten Mal hören und falsch schreiben, erhalten Sie den **Soundex**-Code `Q520`.

Wenn Sie den Namen `Quincy` beim nächsten Mal richtig schreiben, erhalten Sie trotzdem denselben Code `Q520`, was bedeutet, dass Sie mehrere Namensaussprachen derselben <em>Person</em> zuordnen können, ohne jede Schreibweise hinzufügen zu müssen.

Hier sind die Regeln:

<ul>
  <li>Wenn ein Vokal (A, E, I, O, U) zwei Konsonanten mit demselben Soundex-Code trennt, wird der Konsonant rechts vom Vokal codiert. Tymczak wird als T-522 codiert (T, 5 für das M, 2 für das C, Z ignoriert (siehe die Regel „Nebeneinander“ oben), 2 für das K). Da der Vokal „A“ das Z und K trennt, wird das K codiert.</li>
  <li>Wenn „H“ oder „W“ zwei Konsonanten mit demselben Soundex-Code trennen, wird der Konsonant rechts vom Vokal nicht codiert. Beispiel: Ashcraft wird als A-261 codiert (A, 2 für das S, C ignoriert, 6 für das R, 1 für das F). Es wird nicht als A-226 codiert.</li>
</ul>

# --instructions--

Schreiben Sie eine Funktion, die einen String als Parameter erhält und den codierten String zurückgibt.

# --hints--

`soundex` sollte eine Funktion sein.

```js
assert(typeof soundex == 'function');
```

`soundex("Soundex")` sollte einen String zurückgeben.

```js
assert(typeof soundex('Soundex') == 'string');
```

`soundex("Soundex")` sollte `"S532"` zurückgeben.

```js
assert.equal(soundex('Soundex'), 'S532');
```

`soundex("Example")` sollte `"E251"` zurückgeben.

```js
assert.equal(soundex('Example'), 'E251');
```

`soundex("Sownteks")` sollte `"S532"` zurückgeben.

```js
assert.equal(soundex('Sownteks'), 'S532');
```

`soundex("Ekzampul")` sollte `"E251"` zurückgeben.

```js
assert.equal(soundex('Ekzampul'), 'E251');
```

`soundex("Euler")` sollte `"E460"` zurückgeben.

```js
assert.equal(soundex('Euler'), 'E460');
```

`soundex("Gauss")` sollte `"G200"` zurückgeben.

```js
assert.equal(soundex('Gauss'), 'G200');
```

`soundex("Hilbert")` sollte `"H416"` zurückgeben.

```js
assert.equal(soundex('Hilbert'), 'H416');
```

`soundex("Knuth")` sollte `"K530"` zurückgeben.

```js
assert.equal(soundex('Knuth'), 'K530');
```

`soundex("Lloyd")` sollte `"L300"` zurückgeben.

```js
assert.equal(soundex('Lloyd'), 'L300');
```

`soundex("Lukasiewicz")` sollte `"L222"` zurückgeben.

```js
assert.equal(soundex('Lukasiewicz'), 'L222');
```

# --seed--

## --seed-contents--

```js
function soundex(s) {

}
```

# --solutions--

```js
function soundex(s) {
  var a = s.toLowerCase().split('');
  var f = a.shift(),
    r = '',
    codes = {
      a: '',
      e: '',
      i: '',
      o: '',
      u: '',
      b: 1,
      f: 1,
      p: 1,
      v: 1,
      c: 2,
      g: 2,
      j: 2,
      k: 2,
      q: 2,
      s: 2,
      x: 2,
      z: 2,
      d: 3,
      t: 3,
      l: 4,
      m: 5,
      n: 5,
      r: 6
    };
  r =
    f +
    a
      .map(function(v, i, a) {
        return codes[v];
      })
      .filter(function(v, i, a) {
        return i === 0 ? v !== codes[f] : v !== a[i - 1];
      })
      .join('');

  return (r + '000').slice(0, 4).toUpperCase();
}
```
