---
id: 5a23c84252665b21eecc8028
title: Stern-Brocot-Folge
challengeType: 1
forumTopicId: 302324
dashedName: stern-brocot-sequence
---

# --description--

Für diese Aufgabe soll die Stern-Brocot-Folge mit einem Algorithmus erzeugt werden, der ähnlich dem bei der <a href="https://rosettacode.org/wiki/Fibonacci_sequence" target="_blank" rel="noopener noreferrer nofollow">Fibonacci-Folge</a> verwendeten ist.

<ol>
  <li>Die ersten beiden Glieder der Folge sind beide 1:</li>
    <ul><li>1, 1</li></ul>
  <li>Beginnen Sie mit dem zweiten Glied der Folge</li>
  <li>Summieren Sie das betrachtete Glied der Folge und sein Vorgänger, (1 + 1) = 2, und fügen Sie es ans Ende der
    Folge an:</li>
    <ul><li>1, 1, 2</li></ul>
  <li>Fügen Sie das betrachtete Glied der Folge ans Ende der Folge an:</li>
    <ul><li>1, 1, 2, 1</li></ul>
  <li>Betrachten Sie das nächste Glied der Folge, (das dritte Glied, also 2)</li>
  <li>GEHEN SIE ZU 3 </li>
    <ul>
      <li></li>
      <li> ─── Wenn man eine weitere Schleife ausführt, erhält man: ───</li>
      <li></li>
    </ul>
  <li>Summieren Sie das betrachtete Glied der Folge und sein Vorgänger, (2 + 1) = 3, und fügen Sie es ans Ende der
    Folge an:</li>
    <ul><li>1, 1, 2, 1, 3</li></ul>
  <li>Fügen Sie das betrachtete Glied der Folge ans Ende der Folge an:</li>
    <ul><li>1, 1, 2, 1, 3, 2</li></ul>
  <li>Betrachten Sie das nächste Glied der Folge, (das vierte Glied, also 1)</li>
</ol>

# --instructions--

Erstellen Sie eine Funktion, die die Position in der Stern-Brocot-Folge zurückgibt, an der $ n $ zum ersten Mal vorkommt, wobei die Folge mit der oben beschriebenen Methode erzeugt wird. Beachten Sie, dass diese Folge eine 1-basierte Indexierung verwendet.

# --hints--

`sternBrocot` sollte eine Funktion sein.

```js
assert(typeof sternBrocot == 'function');
```

`sternBrocot(2)` sollte eine Zahl zurückgeben.

```js
assert(typeof sternBrocot(2) == 'number');
```

`sternBrocot(2)` sollte `3` zurückgeben.

```js
assert.equal(sternBrocot(2), 3);
```

`sternBrocot(3)` sollte `5` zurückgeben.

```js
assert.equal(sternBrocot(3), 5);
```

`sternBrocot(5)` sollte `11` zurückgeben.

```js
assert.equal(sternBrocot(5), 11);
```

`sternBrocot(7)` sollte `19` zurückgeben.

```js
assert.equal(sternBrocot(7), 19);
```

`sternBrocot(10)` sollte `39` zurückgeben.

```js
assert.equal(sternBrocot(10), 39);
```

# --seed--

## --seed-contents--

```js
function sternBrocot(num) {

}
```

# --solutions--

```js
function sternBrocot(num) {
  function f(n) {
    return n < 2
      ? n
      : n & 1
      ? f(Math.floor(n / 2)) + f(Math.floor(n / 2 + 1))
      : f(Math.floor(n / 2));
  }

  function gcd(a, b) {
    return a ? (a < b ? gcd(b % a, a) : gcd(a % b, b)) : b;
  }
  var n;
  for (n = 1; f(n) != num; n++);
  return n;
}
```
