---
id: 5a23c84252665b21eecc7eca
title: Kaprekar-Zahlen
challengeType: 1
forumTopicId: 302296
dashedName: kaprekar-numbers
---

# --description--

Eine positive ganze Zahl ist eine Kaprekar-Zahl, wenn:

<ul>
  <li>sie 1 ist, oder</li>
  <li>die dezimale Darstellung ihres Quadrats einmal in zwei Teile aufgeteilt werden kann, die aus positiven ganzen Zahlen bestehen und deren Summe die ursprüngliche Zahl ergibt.</li>
</ul>

Beachten Sie, dass eine Aufteilung, bei der ein Teil nur aus 0en besteht, nicht gültig ist, da 0 nicht als positiv gilt.

Beispiel für Kaprekar-Zahlen:

<ul>
  <li><code>2223</code> ist eine Kaprekar-Zahl, da <code>2223 * 2223 = 4941729</code>, <code>4941729</code> in <code>494</code> und <code>1729</code> aufgeteilt werden kann, und <code>494 + 1729 = 2223</code> ergibt</li>
  <li>Die Folge der Kaprekar-Zahlen ist als A006886 bekannt und beginnt mit <code>1, 9, 45, 55, ...</code></li>
</ul>

# --instructions--

Schreiben Sie eine Funktion, die eine Zahl $n$, eine Basis $bs$ entgegennimmt und true zurückgibt, wenn die Zahl eine Kaprekar-Zahl für die gegebene Basis ist. Andernfalls gibt die Funktion false zurück.

# --hints--

`isKaprekar` sollte eine Funktion sein.

```js
assert(typeof isKaprekar == 'function');
```

`isKaprekar(1, 10)` sollte einen Boolean zurückgeben.

```js
assert(typeof isKaprekar(1, 10) == 'boolean');
```

`isKaprekar(1, 10)` sollte `true` zurückgeben.

```js
assert.equal(isKaprekar(1, 10), true);
```

`isKaprekar(9, 10)` sollte `true` zurückgeben.

```js
assert.equal(isKaprekar(9, 10), true);
```

`isKaprekar(2223, 10)` sollte `true` zurückgeben.

```js
assert.equal(isKaprekar(2223, 10), true);
```

`isKaprekar(22823, 10)` sollte `false` zurückgeben.

```js
assert.equal(isKaprekar(22823, 10), false);
```

`isKaprekar(9, 17)` sollte `false` zurückgeben.

```js
assert.equal(isKaprekar(9, 17), false);
```

`isKaprekar(225, 17)` sollte `true` zurückgeben.

```js
assert.equal(isKaprekar(225, 17), true);
```

`isKaprekar(999, 17)` sollte `false` zurückgeben.

```js
assert.equal(isKaprekar(999, 17), false);
```

# --seed--

## --seed-contents--

```js
function isKaprekar(n, bs) {

}
```

# --solutions--

```js
function isKaprekar(n, bs) {
  if (n < 1) return false;
  if (n == 1) return true;
  for (var a = n * n, b = 0, s = 1; a; s *= bs) {
    b += (a % bs) * s;
    a = Math.floor(a / bs);
    if (b && a + b == n) return true;
  }
  return false;
}
```
