---
id: 5a23c84252665b21eecc7edf
title: Kleinstes gemeinsames Vielfaches
challengeType: 1
forumTopicId: 302301
dashedName: least-common-multiple
---

# --description--

Das kleinste gemeinsame Vielfache von 12 und 18 ist 36, denn 12 ist ein Faktor (12 × 3 = 36) und 18 ist ein Faktor (18 × 2 = 36), und es gibt keine positive ganze Zahl kleiner als 36, die beide Faktoren hat. Ein Sonderfall ist, wenn entweder $m$ oder $n$ Null ist, dann ist das kleinste gemeinsame Vielfache Null. Eine Möglichkeit, das kleinste gemeinsame Vielfache zu berechnen, besteht darin, alle Vielfachen von $m$ durchzugehen, bis man eines findet, das auch ein Vielfaches von $n$ ist. If you already have $gcd$ for <a href="https://rosettacode.org/wiki/Greatest_common_divisor" target="_blank" rel="noopener noreferrer nofollow">greatest common divisor</a>, then this formula calculates $lcm$.

$$ \\operatorname{lcm}(m, n) = \\frac{|m \\times n|}{\\operatorname{gcd}(m, n)} $$

# --instructions--

Berechne das kleinste gemeinsame Vielfache eines Arrays von Ganzzahlen. Bei *m* und *n* ist das kleinste gemeinsame Vielfache die kleinste positive Ganzzahl, die sowohl *m* als auch *n* als Faktoren hat.

# --hints--

`LCM` sollte eine Funktion sein.

```js
assert(typeof LCM == 'function');
```

`LCM([2, 4, 8])` sollte eine Zahl zurückgeben.

```js
assert(typeof LCM([2, 4, 8]) == 'number');
```

`LCM([2, 4, 8])` sollte `8` zurückgeben.

```js
assert.equal(LCM([2, 4, 8]), 8);
```

`LCM([4, 8, 12])` sollte `24` zurückgeben.

```js
assert.equal(LCM([4, 8, 12]), 24);
```

`LCM([3, 4, 5, 12, 40])` sollte `120` zurückgeben.

```js
assert.equal(LCM([3, 4, 5, 12, 40]), 120);
```

`LCM([11, 33, 90])` sollte `990` zurückgeben.

```js
assert.equal(LCM([11, 33, 90]), 990);
```

`LCM([-50, 25, -45, -18, 90, 447])` sollte `67050` zurückgeben.

```js
assert.equal(LCM([-50, 25, -45, -18, 90, 447]), 67050);
```

# --seed--

## --seed-contents--

```js
function LCM(A) {

}
```

# --solutions--

```js
function LCM(A) {
  var n = A.length,
    a = Math.abs(A[0]);
  for (var i = 1; i < n; i++) {
    var b = Math.abs(A[i]),
      c = a;
    while (a && b) {
      a > b ? (a %= b) : (b %= a);
    }
    a = Math.abs(c * A[i]) / (a + b);
  }
  return a;
}
```
