---
id: 5a23c84252665b21eecc8036
title: Steuerzeichen und erweiterte Zeichen aus einem String entfernen
challengeType: 1
forumTopicId: 302327
dashedName: strip-control-codes-and-extended-characters-from-a-string
---

# --description--

Die Aufgabe besteht darin, Steuerzeichen und erweiterte Zeichen aus einem String zu entfernen. Die Lösung soll zeigen, wie man jedes der folgenden Ergebnisse erzielt: Ein String, bei dem Steuerzeichen und erweiterte Zeichen entfernt wurden. Im ASCII-Code haben die Steuerzeichen die Dezimalwerte von 0 bis 31 und 127. Auf einem ASCII-basierten System würde ein String, bei dem die Steuerzeichen entfernt wurden, nur noch Zeichen im Bereich von 32 bis 126 Dezimal in der ASCII-Tabelle enthalten. Auf einem nicht ASCII-basierten System betrachten wir Zeichen, die kein entsprechendes Glyphensymbol in der ASCII-Tabelle (im Bereich von 32 bis 126 Dezimal) haben, als erweiterte Zeichen für diese Aufgabe.

# --hints--

`strip` sollte eine Funktion sein.

```js
assert(typeof strip == 'function');
```

`strip("abc")` sollte einen String zurückgeben.

```js
assert(typeof strip('abc') == 'string');
```

`strip("\ba\x00b\n\rc\fd\xc3")` sollte `"abcd"` zurückgeben.

```js
assert.equal(strip('\ba\x00b\n\rc\fd\xc3'), 'abcd');
```

`strip("\u0000\n abc\u00E9def\u007F")` sollte `" abcdef"` zurückgeben.

```js
assert.equal(strip('\u0000\n abc\u00E9def\u007F'), ' abcdef');
```

`strip("a\n\tb\u2102d\u2147f")` sollte `"abdf"` zurückgeben.

```js
assert.equal(strip('a\n\tb\u2102d\u2147f'), 'abdf');
```

`strip("Français.")` sollte `"Franais."` zurückgeben.

```js
assert.equal(strip('Français.'), 'Franais.');
```

`strip("123\tabc\u0007DEF\u007F+-*/€æŧðłþ")` sollte `"123abcDEF+-*/"` zurückgeben.

```js
assert.equal(strip('123\tabc\u0007DEF\u007F+-*/€æŧðłþ'), '123abcDEF+-*/');
```

# --seed--

## --seed-contents--

```js
function strip(s) {

}
```

# --solutions--

```js
function strip(s) {
  return s
    .split('')
    .filter(function(x) {
      var n = x.charCodeAt(0);

      return 31 < n && 127 > n;
    })
    .join('');
}
```
