---
id: 5a23c84252665b21eecc8005
title: Sortieralgorithmen/Comb sort
challengeType: 1
forumTopicId: 302313
dashedName: sorting-algorithmscomb-sort
---

# --description--

Implementieren Sie einen *Comb sort*.

Der **Comb Sort** ist eine Variante des Bubble Sort.

Wie beim Shell sort erhöht der Comb Sort die Lücke, die bei Vergleichen und Vertauschungen verwendet wird.

Das Teilen der Lücke durch $(1-e^{-\\varphi})^{-1} \\approx 1.247330950103979$ funktioniert am besten, aber 1,3 kann praktischer sein.

Einige Implementierungen verwenden den Insertion Sort, sobald die Lücke kleiner als ein bestimmter Wert ist.

Varianten:

<ul>
  <li>Combsort11 stellt sicher, dass die Lücke in (11, 8, 6, 4, 3, 2, 1) endet, was deutlich schneller ist als die beiden anderen möglichen Endungen.</li>
  <li>Combsort mit verschiedenen Endungen wechselt zu einem effizienteren Sortierverfahren, wenn die Daten fast sortiert sind (wenn die Lücke klein ist). Comb sort mit einer kleinen Lücke ist nicht viel besser als der Bubble Sort.</li>
</ul>

Pseudocode:

<pre><b>function</b> combsort(<b>array</b> input)
  gap := input<b>.size</b> <i>//initialisiere die Lückengröße</i>
  <b>loop until</b> gap = 1 <b>and</b> swaps = 0
    <i>//aktualisiere den Lückenwert für den nächsten Comb. Unten ein Beispiel</i>
    gap := int(gap / 1.25)
    <b>if</b> gap &#x3C; 1 
      <i>//Minimum der Lücke ist 1</i>
      gap := 1
    <b>end if</b>
    i := 0
    swaps := 0 <i>//siehe <a href='https://rosettacode.org/wiki/Sorting_algorithms/Bubble_sort' target='_blank'>Bubble Sort</a> für eine Erklärung</i>
    <i>//ein einzelner "Comb"-Durchlauf über die Eingabeliste</i>
    <b>loop until</b> i + gap >= input<b>.size</b> <i>//siehe <a href='https://rosettacode.org/wiki/Sorting_algorithms/Shell_sort' target='_blank'>Shell sort</a> für eine ähnliche Idee</i>
      <b>if</b> input[i] > input[i+gap]
        <b>swap</b>(input[i], input[i+gap])
        swaps := 1 <i>// Markiere, dass ein Tausch stattgefunden hat, also</i>
            <i>// ist die Liste nicht garantiert sortiert</i>
      <b>end if</b>
      i := i + 1
    <b>end loop</b>
  <b>end loop</b>
<b>end function</b>
</pre>

# --instructions--

Schreiben Sie eine Funktion, die ein gegebenes Array mit Comb sort sortiert.

# --hints--

`combSort` sollte eine Funktion sein.

```js
assert(typeof combSort == 'function');
```

`combSort([25, 32, 12, 7, 20])` sollte ein Array zurückgeben.

```js
assert(Array.isArray(combSort([25, 32, 12, 7, 20])));
```

`combSort([25, 32, 12, 7, 20])` sollte `[7, 12, 20, 25, 32]` zurückgeben.

```js
assert.deepEqual(combSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`combSort([38, 45, 35, 8, 13])` sollte `[8, 13, 35, 38, 45]` zurückgeben.

```js
assert.deepEqual(combSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`combSort([43, 36, 20, 34, 24])` sollte `[20, 24, 34, 36, 43]` zurückgeben.

```js
assert.deepEqual(combSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`combSort([12, 33, 26, 18, 1, 16, 38])` sollte `[1, 12, 16, 18, 26, 33, 38]` zurückgeben.

```js
assert.deepEqual(combSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`combSort([3, 39, 48, 16, 1, 4, 29])` sollte `[1, 3, 4, 16, 29, 39, 48]` zurückgeben.

```js
assert.deepEqual(combSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function combSort(arr) {

}
```

# --solutions--

```js
function combSort(arr) {
  function is_array_sorted(arr) {
    var sorted = true;
    for (var i = 0; i < arr.length - 1; i++) {
      if (arr[i] > arr[i + 1]) {
        sorted = false;
        break;
      }
    }
    return sorted;
  }
  var iteration_count = 0;
  var gap = arr.length - 2;
  var decrease_factor = 1.25;

  // Until array is not sorted, repeat iterations
  while (!is_array_sorted(arr)) {
    // If not first gap
    if (iteration_count > 0)
      // Calculate gap
      gap = gap == 1 ? gap : Math.floor(gap / decrease_factor);

    // Set front and back elements and increment to a gap
    var front = 0;
    var back = gap;
    while (back <= arr.length - 1) {
      // If elements are not ordered swap them
      if (arr[front] > arr[back]) {
        var temp = arr[front];
        arr[front] = arr[back];
        arr[back] = temp;
      }

      // Increment and re-run swapping
      front += 1;
      back += 1;
    }
    iteration_count += 1;
  }

  return arr;
}
```
