---
id: 5a23c84252665b21eecc7eb0
title: I vor E außer nach C
challengeType: 1
forumTopicId: 302288
dashedName: i-before-e-except-after-c
---

# --description--

**<a href="http://www.rosettacode.org/wiki/I_before_E_except_after_C" target="_blank" rel="noopener noreferrer nofollow">„I vor E, außer nach C“</a>** ist eine allgemeine Regel für die englische Rechtschreibung. Wenn man unsicher ist, ob ein Wort mit dem Digraph `ei` oder `ie` geschrieben wird, legt der Reim nahe, dass die korrekte Reihenfolge `ie` ist, es sei denn, der vorhergehende Buchstabe ist `c`, in welchem Fall es `ei` sein kann.

Verwenden Sie die angegebenen Wörter, um zu prüfen, ob die beiden Teilklauseln des Satzes einzeln plausibel sind:

<ol>
  <li>
    <i>„I vor E, wenn nicht von C vorausgegangen“. </i>
  </li>
  <li>
    <i>„E vor I, wenn von C vorausgegangen“. </i>
  </li>
</ol>

Wenn beide Teilphrasen plausibel sind, kann man sagen, dass der ursprüngliche Satz plausibel ist.

# --instructions--

Schreiben Sie eine Funktion, die ein Wort entgegennimmt und überprüft, ob das Wort dieser Regel folgt. Die Funktion soll true zurückgeben, wenn das Wort der Regel entspricht, und false, wenn nicht.

# --hints--

`IBeforeExceptC` sollte eine Funktion sein.

```js
assert(typeof IBeforeExceptC == 'function');
```

`IBeforeExceptC("receive")` sollte einen Boolean zurückgeben.

```js
assert(typeof IBeforeExceptC('receive') == 'boolean');
```

`IBeforeExceptC("receive")` sollte `true` zurückgeben.

```js
assert.equal(IBeforeExceptC('receive'), true);
```

`IBeforeExceptC("science")` sollte `false` zurückgeben.

```js
assert.equal(IBeforeExceptC('science'), false);
```

`IBeforeExceptC("imperceivable")` sollte `true` zurückgeben.

```js
assert.equal(IBeforeExceptC('imperceivable'), true);
```

`IBeforeExceptC("inconceivable")` sollte `true` zurückgeben.

```js
assert.equal(IBeforeExceptC('inconceivable'), true);
```

`IBeforeExceptC("insufficient")` sollte `false` zurückgeben.

```js
assert.equal(IBeforeExceptC('insufficient'), false);
```

`IBeforeExceptC("omniscient")` sollte `false` zurückgeben.

```js
assert.equal(IBeforeExceptC('omniscient'), false);
```

# --seed--

## --seed-contents--

```js
function IBeforeExceptC(word) {

}
```

# --solutions--

```js
function IBeforeExceptC(word)
{
    if(word.indexOf("c")==-1 && word.indexOf("ie")!=-1)
        return true;
    else if(word.indexOf("cei")!=-1)
        return true;
    return false;
}
```
