---
id: 5a23c84252665b21eecc7edf
title: kleinster gemeinsamer Vielfacher
challengeType: 1
forumTopicId: 302301
dashedName: least-common-multiple
---

# --description--

Das kleinste gemeinsame Vielfache von 12 und 18 ist 36, weil 12 ein Faktor ist (12 × 3 = 36) und 18 ein Faktor ist (18 × 2 = 36) und es keine positive ganze Zahl kleiner als 36 gibt, die beide Faktoren hat. Als Sonderfall gilt: Wenn entweder $m$ oder $n$ null ist, dann ist das kleinste gemeinsame Vielfache null. Eine Möglichkeit, das kleinste gemeinsame Vielfache zu berechnen, besteht darin, alle Vielfachen von $m$ zu iterieren, bis Sie eines finden, das auch ein Vielfaches von $n$ ist. Wenn Sie bereits den $gcd$ für den <a href="https://rosettacode.org/wiki/Greatest_common_divisor" target="_blank" rel="noopener noreferrer nofollow">größten gemeinsamen Teiler</a> haben, dann berechnet diese Formel das $lcm$. 

$$
\\operatorname{lcm}(m, n) = \\frac{|m \\times n|}{\\operatorname{gcd}(m, n)}
$$

# --instructions--

Berechnen Sie das kleinste gemeinsame Vielfache eines Arrays von ganzen Zahlen. Gegeben sind *m* und *n*, das kleinste gemeinsame Vielfache ist die kleinste positive ganze Zahl, die sowohl *m* als auch *n* als Faktoren hat.

# --hints--

`LCM` sollte eine Funktion sein.

```js
assert(typeof LCM == 'function');
```

`LCM([2, 4, 8])` sollte eine Zahl zurückgeben.

```js
assert(typeof LCM([2, 4, 8]) == 'number');
```

`LCM([2, 4, 8])` sollte `8` zurückgeben.

```js
assert.equal(LCM([2, 4, 8]), 8);
```

`LCM([4, 8, 12])` sollte `24` zurückgeben.

```js
assert.equal(LCM([4, 8, 12]), 24);
```

`LCM([3, 4, 5, 12, 40])` sollte `120` zurückgeben.

```js
assert.equal(LCM([3, 4, 5, 12, 40]), 120);
```

`LCM([11, 33, 90])` sollte `990` zurückgeben.

```js
assert.equal(LCM([11, 33, 90]), 990);
```

`LCM([-50, 25, -45, -18, 90, 447])` sollte `67050` zurückgeben.

```js
assert.equal(LCM([-50, 25, -45, -18, 90, 447]), 67050);
```

# --seed--

## --seed-contents--

```js
function LCM(A) {

}
```

# --solutions--

```js
function LCM(A) {
  var n = A.length,
    a = Math.abs(A[0]);
  for (var i = 1; i < n; i++) {
    var b = Math.abs(A[i]),
      c = a;
    while (a && b) {
      a > b ? (a %= b) : (b %= a);
    }
    a = Math.abs(c * A[i]) / (a + b);
  }
  return a;
}
```
