---
id: 5a23c84252665b21eecc7e7b
title: Generator/Exponential
challengeType: 1
forumTopicId: 302275
dashedName: generatorexponential
---

# --description--

Ein Generator ist eine ausführbare Einheit (wie eine Funktion oder Prozedur), die Code enthält, der eine Folge von Werten jeweils einzeln erzeugt, sodass bei jedem Aufruf des Generators der nächste Wert der Folge bereitgestellt wird.

Generatoren basieren oft auf Koroutinen oder Objekten, sodass der interne Zustand des Objekts „natürlich“ verwaltet wird.

Generatoren werden häufig in Situationen verwendet, in denen eine Folge potenziell unendlich ist und der nächste Wert der Folge mit nur minimalem Zustand konstruiert werden kann.

# --instructions--

Schreiben Sie eine Funktion, die Generatoren verwendet, um Quadratzahlen und Kubikzahlen zu erzeugen. Erstellen Sie einen neuen Generator, der alle Kubikzahlen aus dem Generator der Quadratzahlen herausfiltert.

Die Funktion soll den \\( n^{th} \\) Wert des gefilterten Generators zurückgeben.

Zum Beispiel soll die Funktion für \\(n=7\\) den Wert 81 zurückgeben, da die Folge 4, 9, 16, 25, 36, 49, 81 lautet. Hier wird 64 herausgefiltert, da es eine Kubikzahl ist.

# --hints--

`exponentialGenerator` sollte eine Funktion sein.

```js
assert(typeof exponentialGenerator == 'function');
```

`exponentialGenerator()` sollte eine Zahl zurückgeben.

```js
assert(typeof exponentialGenerator(10) == 'number');
```

`exponentialGenerator(10)` sollte `144` zurückgeben.

```js
assert.equal(exponentialGenerator(10), 144);
```

`exponentialGenerator(12)` sollte `196` zurückgeben.

```js
assert.equal(exponentialGenerator(12), 196);
```

`exponentialGenerator(14)` sollte `256` zurückgeben.

```js
assert.equal(exponentialGenerator(14), 256);
```

`exponentialGenerator(20)` sollte `484` zurückgeben.

```js
assert.equal(exponentialGenerator(20), 484);
```

`exponentialGenerator(25)` sollte `784` zurückgeben.

```js
assert.equal(exponentialGenerator(25), 784);
```

# --seed--

## --seed-contents--

```js
function exponentialGenerator(n) {

}
```

# --solutions--

```js
function exponentialGenerator(n){
  function* PowersGenerator(m) {
    var n=0;
    while(1) {
        yield Math.pow(n, m);
        n += 1;
    }
  }

  function* FilteredGenerator(g, f){
    var value = g.next().value;
    var filter = f.next().value;
    while(1) {
        if( value < filter ) {
            yield value;
            value = g.next().value;
        } else if ( value > filter ) {
            filter = f.next().value;
        } else {
            value = g.next().value;
            filter = f.next().value;
        }
    }
  }

  var squares = PowersGenerator(2);
  var cubes = PowersGenerator(3);

  var filtered = FilteredGenerator(squares, cubes);

  var curr=0;
  for(var i=0;i<n;i++) curr=filtered.next();

  return curr.value;
}
```
