---
id: 5e4ce2f5ac708cc68c1df261
title: Linearer kongruenter Generator
challengeType: 1
forumTopicId: 385266
dashedName: linear-congruential-generator
---

# --description--

Ein linearer kongruenter Generator (LCG) ist ein <em>Algorithmus</em>, der eine Folge von pseudozufälligen Zahlen erzeugt, die mit einer stückweise linearen, diskontinuierlichen Gleichung berechnet werden. Alle linearen kongruenten Generatoren verwenden diese Formel:

$$r_{n + 1} = (a \times r_n + c) \bmod m$$

Dabei gilt:

<ul>
<li>$ r_0 $ ist ein Startwert (Seed).</li>
<li>$r_1$, $r_2$, $r_3$, ..., sind die Zufallszahlen.</li>
<li>$a$, $c$, $m$ sind Konstanten.</li>
</ul>

Wenn man die Werte von $a$, $c$ und $m$ sorgfältig auswählt, erzeugt der Generator eine gleichmäßige Verteilung ganzer Zahlen von $0$ bis $m - 1$.

<abbr title="linear congruential generator">LCG</abbr>-Zahlen haben eine schlechte Qualität. $r_n$ und $r\_{n + 1}$ sind nicht unabhängig, wie es bei echten Zufallszahlen der Fall wäre. Wer $r_n$ kennt, kann $r\_{n + 1}$ vorhersagen, daher ist der <abbr title="linear congruential generator">LCG</abbr> nicht kryptografisch sicher. Der <abbr title="linear congruential generator">LCG</abbr> ist jedoch ausreichend für einfache Aufgaben wie den Miller-Rabin-Primzahltest oder FreeCell-Ausgaben. Zu den Vorteilen des <abbr title="linear congruential generator">LCG</abbr> gehört, dass man eine Zahlenfolge leicht reproduzieren kann, wenn man denselben $r_0$ verwendet. Man kann eine solche Folge auch in einer anderen Programmiersprache reproduzieren, da die Formel so einfach ist.

# --instructions--

Schreiben Sie eine Funktion, die $r_0,a,c,m,n$ als Parameter erhält und $r_n$ zurückgibt.

# --hints--

`linearCongGenerator` sollte eine Funktion sein.

```js
assert(typeof linearCongGenerator == 'function');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` sollte eine Zahl zurückgeben.

```js
assert(typeof linearCongGenerator(324, 1145, 177, 2148, 3) == 'number');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` sollte `855` zurückgeben.

```js
assert.equal(linearCongGenerator(324, 1145, 177, 2148, 3), 855);
```

`linearCongGenerator(234, 11245, 145, 83648, 4)` sollte `1110` zurückgeben.

```js
assert.equal(linearCongGenerator(234, 11245, 145, 83648, 4), 1110);
```

`linearCongGenerator(85, 11, 1234, 214748, 5)` sollte `62217` zurückgeben.

```js
assert.equal(linearCongGenerator(85, 11, 1234, 214748, 5), 62217);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 1)` sollte `12345` zurückgeben.

```js
assert.equal(linearCongGenerator(0, 1103515245, 12345, 2147483648, 1), 12345);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 2)` sollte `1406932606` zurückgeben.

```js
assert.equal(
  linearCongGenerator(0, 1103515245, 12345, 2147483648, 2),
  1406932606
);
```

# --seed--

## --seed-contents--

```js
function linearCongGenerator(r0, a, c, m, n) {

}
```

# --solutions--

```js
function linearCongGenerator(r0, a, c, m, n) {
    for (let i = 0; i < n; i++) {
        r0 = (a * r0 + c) % m;
    }
    return r0;
}
```
