---
id: 5e6dd14797f5ce267c2f19d0
title: Look-and-say-Folge
challengeType: 1
forumTopicId: 385277
dashedName: look-and-say-sequence
---

# --description--

Die Look-and-say-Folge ist eine rekursiv definierte Zahlenfolge.

Definition der Folge

<ul><li>Nehmen Sie eine Dezimalzahl</li>
<li><span>Betrachten</span> Sie die Zahl und gruppieren Sie visuell aufeinanderfolgende gleiche Ziffern.</li>
<li><span>Sagen</span> Sie die Zahl von links nach rechts, Gruppe für Gruppe; wie viele dieser Ziffer vorhanden sind – gefolgt von der gruppierten Ziffer.</li></ul><span> Dies wird zur nächsten Zahl der Folge.</span>

Ein Beispiel:

<ul><li>Ausgehend von der Zahl 1 haben Sie <span>eine</span> 1, was 11 ergibt</li>
<li>Ausgehend von 11 haben Sie <span>zwei</span> 1en. Also: 21</li>
<li>Ausgehend von 21 haben Sie <span>eine</span> 2, dann <span>eine</span> 1. Also: (12)(11), was 1211 ergibt</li>
<li>Ausgehend von 1211 haben Sie <span>eine</span> 1, <span>eine</span> 2, dann <span>zwei</span> 1en. Also: (11)(12)(21), was 111221 ergibt</li></ul>

# --instructions--

Schreiben Sie eine Funktion, die einen String als Parameter annimmt, ihn verarbeitet und den resultierenden String zurückgibt.

# --hints--

`lookAndSay` sollte eine Funktion sein.

```js
assert(typeof lookAndSay == 'function');
```

`lookAndSay("1")` sollte einen String zurückgeben.

```js
assert(typeof lookAndSay('1') == 'string');
```

`lookAndSay("1")` sollte `"11"` zurückgeben.

```js
assert.equal(lookAndSay('1'), '11');
```

`lookAndSay("11")` sollte `"21"` zurückgeben.

```js
assert.equal(lookAndSay('11'), '21');
```

`lookAndSay("21")` sollte `"1211"` zurückgeben.

```js
assert.equal(lookAndSay('21'), '1211');
```

`lookAndSay("1211")` sollte `"111221"` zurückgeben.

```js
assert.equal(lookAndSay('1211'), '111221');
```

`lookAndSay("3542")` sollte `"13151412"` zurückgeben.

```js
assert.equal(lookAndSay('3542'), '13151412');
```

# --seed--

## --seed-contents--

```js
function lookAndSay(str) {

}
```

# --solutions--

```js
function lookAndSay(str) {
    return str.replace(/(.)\1*/g, function(seq, p1) {
      return seq.length.toString() + p1;
    });
}
```
