---
id: 5e9ddb06ec35240f39657419
title: FizzBuzz
challengeType: 1
forumTopicId: 385370
dashedName: fizzbuzz
---

# --description--

Schreiben Sie ein Programm, das ein Array von ganzen Zahlen von 1 bis 100 (einschließlich) erzeugt. Aber:

<ul>
    <li>für Vielfache von 3 fügen Sie <code>"Fizz"</code> anstelle der Zahl zum Array hinzu</li>
    <li>für Vielfache von 5 fügen Sie <code>"Buzz"</code> anstelle der Zahl zum Array hinzu</li>
    <li>für Vielfache von 3 und 5 fügen Sie <code>"FizzBuzz"</code> anstelle der Zahl zum Array hinzu</li>
</ul>

# --instructions--

Ihr Programm soll ein Array zurückgeben, das die Ergebnisse basierend auf den obigen Regeln enthält.

# --hints--

`fizzBuzz` sollte eine Funktion sein.

```js
assert(typeof fizzBuzz == 'function');
```

`fizzBuzz()` sollte ein Array zurückgeben.

```js
assert(Array.isArray(fizzBuzz()) == true);
```

Zahlen, die nur durch 3 teilbar sind, sollten `"Fizz"` zurückgeben.

```js
assert.equal(fizzBuzz()[2], 'Fizz');
```

Zahlen, die nur durch 5 teilbar sind, sollten `"Buzz"` zurückgeben.

```js
assert.equal(fizzBuzz()[99], 'Buzz');
```

Zahlen, die durch sowohl 3 als auch 5 teilbar sind, sollten `"FizzBuzz"` zurückgeben.

```js
assert.equal(fizzBuzz()[89], 'FizzBuzz');
```

Zahlen, die weder durch 3 noch durch 5 teilbar sind, sollten die Zahl selbst zurückgeben.

```js
assert.equal(fizzBuzz()[12], 13);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz() {

}
```

# --solutions--

```js
function fizzBuzz() {
    let res=[];
    for (let i =1; i < 101; i++) {
        if (i % 3 === 0  && i % 5 === 0) {
            res.push("FizzBuzz");
        }
        else if (i % 3 === 0) {
            res.push("Fizz");
        }
        else if (i % 5 === 0) {
            res.push("Buzz");
        } 
        else {
            res.push(i);
        }
    }
    return res;
}
```
