---
id: 5ea28156e79528a9ab248f27
title: Luhn-Test für Kreditkartennummern
challengeType: 1
forumTopicId: 385284
dashedName: luhn-test-of-credit-card-numbers
---

# --description--

Der Luhn-Test wird von einigen Kreditkartenunternehmen verwendet, um gültige Kreditkartennummern von einer zufälligen Auswahl von Ziffern zu unterscheiden.

Diese Unternehmen verwenden Kreditkartennummern, die durch den Luhn-Test validiert werden können und die folgenden Test bestehen:

<ol>
  <li> Drehen Sie die Reihenfolge der Ziffern in der Nummer um.</li>
  <li> Nehmen Sie die erste, dritte ... und jede weitere ungerade Ziffer in den umgedrehten Ziffern und summieren Sie sie, um die Teilsumme s1 zu bilden.</li>
  <li> Nehmen Sie die zweite, vierte ... und jede weitere gerade Ziffer in den umgedrehten Ziffern:</li>
    <ol>
      <li>Multiplizieren Sie jede Ziffer mit zwei und summieren Sie die Ziffern, falls das Ergebnis größer als neun ist, um Teilsummen für die geraden Ziffern zu bilden.</li>
      <li>Summieren Sie die Teilsummen der geraden Ziffern, um s2 zu bilden.</li>
    </ol>
  <li>Wenn s1 + s2 auf null endet, dann ist die ursprüngliche Nummer in der Form einer gültigen Kreditkartennummer, wie durch den Luhn-Test überprüft.</li>
</ol>

Zum Beispiel, wenn die Testnummer 49927398716 ist:

```md
Reverse the digits:
  61789372994
Sum the odd digits:
  6 + 7 + 9 + 7 + 9 + 4 = 42 = s1
The even digits:
    1, 8, 3, 2, 9
  Two times each even digit:
    2, 16, 6, 4, 18
  Sum the digits of each multiplication:
    2, 7, 6, 4, 9
  Sum the last:
    2 + 7 + 6 + 4 + 9 = 28 = s2

s1 + s2 = 70 which ends in zero which means that 49927398716 passes the Luhn test.
```

# --instructions--

Schreiben Sie eine Funktion, die eine Nummer mit dem Luhn-Test validiert. Geben Sie true zurück, wenn es eine gültige Nummer ist. Andernfalls geben Sie false zurück.

# --hints--

`luhnTest` sollte eine Funktion sein.

```js
assert(typeof luhnTest === 'function');
```

`luhnTest("4111111111111111")` sollte einen Boolean zurückgeben.

```js
assert(typeof luhnTest('4111111111111111') === 'boolean');
```

`luhnTest("4111111111111111")` sollte `true` zurückgeben.

```js
assert.equal(luhnTest('4111111111111111'), true);
```

`luhnTest("4111111111111112")` sollte `false` zurückgeben.

```js
assert.equal(luhnTest('4111111111111112'), false);
```

`luhnTest("49927398716")` sollte `true` zurückgeben.

```js
assert.equal(luhnTest('49927398716'), true);
```

`luhnTest("49927398717")` sollte `false` zurückgeben.

```js
assert.equal(luhnTest('49927398717'), false);
```

`luhnTest("1234567812345678")` sollte `false` zurückgeben.

```js
assert.equal(luhnTest('1234567812345678'), false);
```

`luhnTest("1234567812345670")` sollte `true` zurückgeben.

```js
assert.equal(luhnTest('1234567812345670'), true);
```

# --seed--

## --seed-contents--

```js
function luhnTest(str) {

}
```

# --solutions--

```js
function luhnTest(str) {
  var luhnArr = [0, 2, 4, 6, 8, 1, 3, 5, 7, 9];
  var counter = 0;
  var incNum;
  var odd = false;
  var temp = String(str).replace(/[^\d]/g, '');
  if (temp.length == 0) return false;
  for (var i = temp.length - 1; i >= 0; --i) {
    incNum = parseInt(temp.charAt(i), 10);
    counter += (odd = !odd) ? incNum : luhnArr[incNum];
  }
  return counter % 10 == 0;
}
```
