---
id: 5eaf48389ee512d4d103684b
title: Selbstbeschreibende Zahlen
challengeType: 1
forumTopicId: 385289
dashedName: self-describing-numbers
---

# --description--

Es gibt mehrere sogenannte „selbstbeschreibende“ ganze Zahlen.

Eine ganze Zahl gilt als „selbstbeschreibend“, wenn sie die Eigenschaft hat, dass bei der Bezeichnung der Ziffernpositionen von 0 bis N-1 die Ziffer an jeder Position der Anzahl der Vorkommen dieser Ziffer in der Zahl entspricht.

Zum Beispiel ist **2020** eine vierstellige selbstbeschreibende Zahl:

<ul>
    <li> Position 0 hat den Wert 2 und es gibt zwei 0en in der Zahl; </li>
    <li> Position 1 hat den Wert 0 und es gibt keine 1en in der Zahl; </li>
    <li> Position 2 hat den Wert 2 und es gibt zwei 2en; </li>
    <li> Position 3 hat den Wert 0 und es gibt null 3en; </li>
</ul>

Selbstbeschreibende Zahlen &lt; 100.000.000 sind: 1210, 2020, 21200, 3211000, 42101000.

# --instructions--

Schreiben Sie eine Funktion, die eine positive ganze Zahl als Parameter erhält. Wenn sie selbstbeschreibend ist, geben Sie true zurück. Andernfalls geben Sie false zurück.

# --hints--

`isSelfDescribing` sollte eine Funktion sein.

```js
assert(typeof isSelfDescribing == 'function');
```

`isSelfDescribing()` sollte einen Boolean zurückgeben.

```js
assert(typeof isSelfDescribing(2020) == 'boolean');
```

`isSelfDescribing(2020)` sollte `true` zurückgeben.

```js
assert.equal(isSelfDescribing(2020), true);
```

`isSelfDescribing(3021)` sollte `false` zurückgeben.

```js
assert.equal(isSelfDescribing(3021), false);
```

`isSelfDescribing(3211000)` sollte `true` zurückgeben.

```js
assert.equal(isSelfDescribing(3211000), true);
```

# --seed--

## --seed-contents--

```js
function isSelfDescribing(n) {

}
```

# --solutions--

```js
function isSelfDescribing(n) {
    let digits = String(n).split("");
    digits = digits.map(function(e) {return parseInt(e)});
    let count = digits.map((x) => {return 0})
    digits.forEach((d) =>{
        if (d >= count.length) {
            return false
        }
        count[d] += 1;
    });

     if (digits === count) {
        return true;
    }
    if (digits.length != count.length) {
        return false;
    }
    
    for (let i=0; i< digits.length; i++){
      if (digits[i] !== count[i]) {
        return false;
      }
    }
    return true;
}
```
