---
id: 5eb3e497b8d6d7f63c5517ea
title: Eine Liste von Datensätzen durchsuchen
challengeType: 1
forumTopicId: 385315
dashedName: search-a-list-of-records
---

# --description--

Ein Datensatz besteht aus Attributen, die eine Entität beschreiben. Jedes Attribut hat einen Namen und einen Wert. Zum Beispiel kann eine Person ein Attribut `age` mit dem Wert 25 haben. Eine wichtige Operation bei einer Liste von Datensätzen ist es, einen Datensatz mit einem bestimmten Attributwert zu finden.

# --instructions--

Schreiben Sie eine Funktion, die einen String als Parameter erhält. Die Funktion soll den Index des Elements in `list` zurückgeben, für das der Wert des `name`-Attributs mit dem angegebenen String übereinstimmt.

# --hints--

`searchCity` sollte eine Funktion sein.

```js
assert(typeof searchCity === 'function');
```

`searchCity("Dar Es Salaam")` sollte eine Zahl zurückgeben.

```js
assert(typeof searchCity('Dar Es Salaam') === 'number');
```

`searchCity("Dar Es Salaam")` sollte `6` zurückgeben.

```js
assert.equal(searchCity('Dar Es Salaam'), 6);
```

`searchCity("Casablanca")` sollte `9` zurückgeben.

```js
assert.equal(searchCity('Casablanca'), 9);
```

`searchCity("Cairo")` sollte `1` zurückgeben.

```js
assert.equal(searchCity('Cairo'), 1);
```

`searchCity("Mogadishu")` sollte `4` zurückgeben.

```js
assert.equal(searchCity('Mogadishu'), 4);
```

`searchCity("Lagos")` sollte `0` zurückgeben.

```js
assert.equal(searchCity('Lagos'), 0);
```

# --seed--

## --seed-contents--

```js
function searchCity(name) {

}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```

# --solutions--

```js
function searchCity(name) {
  return list.findIndex(item => item.name === name);
}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```
