---
id: 66c06d618d075c7f7f1b890a
title: 占い機能を作る
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

このラボでは、5つの占い結果を定義し、その中からランダムに1つを選んでユーザーに表示します。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. 5つの変数 `fortune1`、`fortune2`、`fortune3`、`fortune4`、`fortune5` を任意の文字列で初期化してください。以下の選択肢を使ってもかまいません:

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. 1から5までのランダムな数字を選び、変数 `randomNumber` に代入してください。

3. `selectedFortune` という変数を作成し、以下のルールに従って適切な占い結果を代入してください:

   - `randomNumber` が1なら、`fortune1` の値を `selectedFortune` に代入する。
   - `randomNumber` が2なら、`fortune2` の値を `selectedFortune` に代入する。
   - `randomNumber` が3なら、`fortune3` の値を `selectedFortune` に代入する。
   - `randomNumber` が4なら、`fortune4` の値を `selectedFortune` に代入する。
   - `randomNumber` が5なら、`fortune5` の値を `selectedFortune` に代入する。

4. `selectedFortune` をコンソールに出力してください。

# --hints--

`fortune1` を文字列で初期化してください。

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

`fortune2` を文字列で初期化してください。

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

`fortune3` を文字列で初期化してください。

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

`fortune4` を文字列で初期化してください。

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

`fortune5` を文字列で初期化してください。

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

ランダムな数字を生成するには `Math.random()` メソッドを使ってください。

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

1から5までのランダムな数字を生成し、変数 `randomNumber` に代入してください。

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

`selectedFortune` という変数を作成し、`randomNumber` の値に基づいて値を代入してください。

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` は対応する占い結果と一致する必要があります。例えば、`randomNumber` が1なら `selectedFortune` は `fortune1` と等しくなるべきです。

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

`selectedFortune` をコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
