---
id: 615f34ecc1091b4fd5a8a484
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

All'interno del tuo elemento `head`, aggiungi un elemento `link` con l'attributo `rel` impostato su `stylesheet` e l'attributo `href` impostato su `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Questo importerà la famiglia di caratteri `Open Sans`, con i valori di peso del carattere `400`, `700` e `800`.

Aggiungi anche un elemento `link` per collegare il tuo file `styles.css`.

# --hints--

Il tuo codice dovrebbe avere due elementi `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

I tuoi elementi `link` dovrebbero essere elementi void, non devono avere un tag di chiusura `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

I tuoi due elementi `link` dovrebbero essere all'interno dell'elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Entrambi i tuoi elementi `link` dovrebbero avere l'attributo `rel` impostato su `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Uno dei tuoi elementi `link` dovrebbe avere un attributo `href` impostato su `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

Uno dei tuoi elementi `link` dovrebbe avere un attributo `href` impostato su `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
