---
id: 612e8279827a28352ce83a72
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

Now, copy the set of seven `.key` elements and paste two additional sets into the `.keys` div.

# --hints--

`.key` 要素が 21 個必要です。

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

`.black--key` 要素が 15 個必要です。

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

`.key` 要素はすべて `.keys` 要素の中にある必要があります。

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

`.black--key` 要素はすべて `.keys` 要素の中にある必要があります。

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
