---
id: 587d7fb4367417b2b2512bfe
title: додайте ліцензію до свого package.json
challengeType: 2
forumTopicId: 301523
dashedName: add-a-license-to-your-package-json
---

# --description--

Поле `license` — це місце, де ви повідомляєте користувачам, що їм дозволено робити з вашим проєктом.

Деякі загальні ліцензії для проєктів з відкритим кодом включають MIT і BSD. Інформація про ліцензію не є обов’язковою, і закони про авторське право в більшості країн за замовчуванням надають вам власність на те, що ви створюєте. Проте завжди корисно чітко вказувати, що користувачі можуть і не можуть робити. Ось приклад поля license:

```json
"license": "MIT",
```

# --instructions--

Заповніть поле `license` у файлі package.json вашого проєкту так, як вважаєте за потрібне.

# --hints--

package.json має містити дійсний ключ "license"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.license, '"license" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
