---
id: 587d7fb3367417b2b2512bfc
title: додайте опис до свого package.json
challengeType: 2
forumTopicId: 301522
dashedName: add-a-description-to-your-package-json
---

# --description--

Наступною частиною хорошого package.json є поле `description`, де має бути короткий, але інформативний опис вашого проєкту.

Якщо колись ви плануєте опублікувати пакет у npm, це рядок, який має переконати користувача встановити ваш пакет чи ні. Проте це не єдиний випадок використання опису — це чудовий спосіб підсумувати, що робить проєкт. Він так само важливий у будь-якому проєкті Node.js, щоб допомогти іншим розробникам, майбутнім підтримувачам або навіть вашому майбутньому «я» швидко зрозуміти проєкт.

Незалежно від того, що ви плануєте для свого проєкту, опис однозначно рекомендований. Ось приклад:

```json
"description": "A project that does something awesome",
```

# --instructions--

Додайте `description` до package.json вашого проєкту.

**Примітка:** Пам’ятайте використовувати подвійні лапки для назв полів (") і коми (,) для розділення полів.

# --hints--

package.json має містити дійсний ключ "description"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.description, '"description" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
