---
id: 587d7fb5367417b2b2512c03
title: Використання карета для використання останньої мінорної версії залежності
challengeType: 2
forumTopicId: 301531
dashedName: use-the-caret-character-to-use-the-latest-minor-version-of-a-dependency
---

# --description--

Подібно до того, як тильда, яку ми вивчали в останньому завданні, дозволяє npm встановлювати останній PATCH для залежності, карет (`^`) дозволяє npm також встановлювати майбутні оновлення. Різниця в тому, що карет дозволяє як MINOR оновлення, так і PATCH.

Поточна версія `@freecodecamp/example` має бути `~1.2.13`, що дозволяє npm встановлювати останню `1.2.x` версію. Якби ви використали карет (^) як префікс версії, npm дозволив би оновлювати до будь-якої `1.x.x` версії.

```json
"package": "^1.3.8"
```

Це дозволило б оновлення до будь-якої `1.x.x` версії пакета.

# --instructions--

Використайте карет (`^`) як префікс версії `@freecodecamp/example` у своїх залежностях, щоб npm міг оновлювати її до будь-якого нового MINOR релізу.

**Примітка:** Самі номери версій змінювати не потрібно.

# --hints--

`"dependencies"` має містити `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Версія `"@freecodecamp/example"` має відповідати `"^1.x.x"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\^1\./,
        'Wrong version of "@freecodecamp/example". It should be ^1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
