---
id: 587d7fb4367417b2b2512bfd
title: додайте ключові слова до свого package.json
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

Поле `keywords` — це місце, де ви можете описати свій проєкт за допомогою відповідних ключових слів. Ось приклад:

```json
"keywords": [ "descriptive", "related", "words" ],
```

Як бачите, це поле має структуру масиву рядків у подвійних лапках.

# --instructions--

Додайте масив відповідних рядків до поля `keywords` у файлі package.json вашого проєкту.

Одне з ключових слів має бути "freecodecamp".

# --hints--

package.json має містити дійсний ключ "keywords"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Поле "keywords" має бути масивом

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

"keywords" має містити "freecodecamp"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
