---
id: 587d7fb5367417b2b2512c04
title: видалити пакет із залежностей
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

Ви вже випробували кілька способів керування залежностями свого проєкту, використовуючи розділ dependencies у package.json. Ви також додавали зовнішні пакети, додаючи їх у файл, і навіть вказували npm, які типи версій ви хочете, використовуючи спеціальні символи, такі як тильда або карет.

Але що, якщо ви хочете видалити зовнішній пакет, який вам більше не потрібен? Ви, мабуть, вже здогадалися: просто видаліть відповідну пару ключ-значення для цього пакета з ваших залежностей.

Цей самий метод застосовується і для видалення інших полів у вашому package.json.

# --instructions--

Видаліть пакет `@freecodecamp/example` зі своїх залежностей.

**Примітка:** Переконайтеся, що після видалення залишилась правильна кількість ком.

# --hints--

`"dependencies"` не повинен містити `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
