---
id: 587d7fb4367417b2b2512c00
title: Розширте свій проєкт зовнішніми пакетами з npm
challengeType: 2
forumTopicId: 301527
dashedName: expand-your-project-with-external-packages-from-npm
---

# --description--

Одна з головних причин використовувати менеджер пакетів — це його потужне управління залежностями. Замість того, щоб вручну перевіряти всі залежності щоразу, коли ви налаштовуєте проєкт на новому комп’ютері, npm автоматично встановлює все за вас. Але як npm може точно знати, що потрібно вашому проєкту? Знайомтеся з розділом `dependencies` у вашому файлі package.json.

У цьому розділі зберігаються пакети, які потрібні вашому проєкту, у такому форматі:

```json
"dependencies": {
  "package-name": "version",
  "express": "4.14.0"
}

```

# --instructions--

Додайте версію `1.1.0` пакета `@freecodecamp/example` у поле `dependencies` вашого файлу `package.json`.

**Примітка:** `@freecodecamp/example` — це фіктивний пакет, який використовується як навчальний інструмент.

# --hints--

`"dependencies"` має містити `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Версія `"@freecodecamp/example"` має бути `"1.1.0"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^[\^\~]?1\.1\.0/,
        'Wrong version of "@freecodecamp/example" installed. It should be 1.1.0'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
