---
id: 587d7fb5367417b2b2512c02
title: "Використовуйте символ тильда, щоб завжди використовувати останню патч-версію залежності"
challengeType: 2
forumTopicId: 301532
dashedName: use-the-tilde-character-to-always-use-the-latest-patch-version-of-a-dependency
---

# --description--

У попередньому завданні ви вказали npm використовувати лише конкретну версію пакета. Це корисний спосіб зафіксувати свої залежності, якщо потрібно переконатися, що різні частини проєкту залишаються сумісними між собою. Але в більшості випадків ви не хочете пропустити виправлення помилок, оскільки вони часто містять важливі патчі безпеки і (сподіваємося) не ламають функціональність.

Щоб дозволити npm-залежності оновлюватися до останньої патч-версії, можна додати префікс тильда (`~`) перед версією залежності. Ось приклад, як дозволити оновлення до будь-якої `1.3.x` версії.

```json
"package": "~1.3.8"
```

# --instructions--

У файлі package.json ваше поточне правило для оновлення `@freecodecamp/example` npm — використовувати конкретну версію (`1.2.13`). Але тепер ви хочете дозволити останню `1.2.x` версію.

Використайте символ тильда (`~`) як префікс для версії `@freecodecamp/example` у ваших залежностях, щоб npm міг оновлювати її до будь-якого нового _патч_ релізу.

**Примітка:** Самі номери версій змінювати не потрібно.

# --hints--

`"dependencies"` має містити `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Версія `"@freecodecamp/example"` має відповідати `"~1.2.13"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\~1\.2\.13/,
        'Wrong version of "@freecodecamp/example". It should be ~1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
