---
id: 6551f1e8737dbe6e0cb08ea4
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Adesso stampa la tua nuova variabile.

# --hints--

Non dovresti rimuovere la tua chiamata `print(type(text))` esistente.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Dovresti chiamare la funzione `print()` passando la variabile `shift`. Fai attenzione a mettere la chiamata della funzione all'inizio della riga.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
