---
id: 6551f5756c20146fc94f8675
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Puoi consultare il numero di caratteri in una stringa con la funzione integrata `len()`.

Modifica la tua chiamata `print()` esistente passando `len(text)` invece di `text[-1]`.

# --hints--


Dovresti chiamare la funzione `len()`.

```js
assert.match(code, /len\s*\(.*\)/)
```

Dovresti passare la variabile `text` alla funzione `len()` includendola tra le parentesi tonde.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

Dovresti passare `len(text)` alla funzione `print()` includendola tra le parentesi tonde. Fai attenzione a posizionare la chiamata della funzione all'inizio della riga.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

Non dovresti avere `print(text[-1])` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
