---
id: 6551f3bfed58796ebb1268dc
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Ogni carattere di una stringa può essere consultato tramite un indice numerico. Il conteggio degli indici parte da zero. Quindi il primo carattere di una stringa ha indice `0`. Per esempio, nella stringa `'Hello World'`, `'H'` si trova all'indice `0`, `'e'` si trova all'indice `1` e così via.

Ogni carattere di una stringa può essere consultato usando la notazione a parentesi. Devi scrivere il nome della variabile seguito dalle parentesi quadre e inserire l'indice del carattere tra le parentesi:

```py
text = 'Hello World'
r = text[8]
```

Ora, invece di stampare `text`, stampa solo il carattere all'indice `6`.

# --hints--

Dovresti ancora chiamare la funzione `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Dovresti passare `text[6]` alla funzione `print()` includendolo tra le parentesi tonde. Fai attenzione a posizionare la chiamata della funzione all'inizio della riga.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
