---
id: 6551f22653e6ce6e4f1d1370
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Modifica la tua chiamata `print(shift)` per stampare il tipo di dato della tua variabile `shift`.

# --hints--

Dovresti mantenere la tua chiamata `print(type(text))` esistente.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Dovresti passare `shift` alla funzione `type()`.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

Dovresti passare `type(shift)` alla tua funzione `print()`. Fai attenzione a posizionare la chiamata della funzione all'inizio della riga.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

Non dovresti avere `print(shift)` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
