---
id: 6551ef504d91936d2d4e54f8
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Un *argomento* è un oggetto o un'espressione passata a una funzione — aggiunta tra le parentesi tonde di apertura e chiusura — quando viene chiamata:

```py
greet = 'Hello!'
print(greet)
```

Il codice nell'esempio sopra stamperebbe la stringa `'Hello!'`, che è il valore della variabile `greet` passata a `print()` come argomento.

Stampa la tua variabile `text` sullo schermo passando la variabile `text` come argomento alla funzione `print()`.

# --hints--

Dovresti passare `text` alla funzione `print()` aggiungendo il nome di questa variabile tra le parentesi tonde di apertura e chiusura. Fai attenzione a mettere la chiamata della funzione all'inizio della riga.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
