---
id: 6551f4259d66026ee28cb205
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Puoi anche consultare i caratteri di una stringa partendo dalla fine della stringa. L’ultimo carattere ha un indice pari a `-1`, il penultimo `-2` e così via.

Ora modifica la tua chiamata `print()` esistente per stampare l’ultimo carattere della tua stringa.

# --hints--

Dovresti ancora chiamare la funzione `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Dovresti passare `text[-1]` alla funzione `print()` includendolo tra le parentesi tonde. Fai attenzione a posizionare la chiamata della funzione all’inizio della riga.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
