---
id: 6551ecf597038c36637f3759
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Le variabili sono una parte essenziale di Python e di qualsiasi linguaggio di programmazione. Una variabile è un nome che fa riferimento o punta a un oggetto. Puoi dichiarare una variabile scrivendo il nome della variabile sul lato sinistro dell'operatore di assegnazione `=` e specificando il valore da assegnare a quella variabile sul lato destro dell'operatore di assegnazione:

```py
variable_name = value
```

Crea una variabile chiamata `number` e assegna il valore `5` alla tua nuova variabile.

# --hints--

Dovresti dichiarare una variabile chiamata `number`. Fai attenzione a mettere il nome della variabile all'inizio della riga.

```js
assert.match(code, /^number\s*=/m)
```

Dovresti assegnare il valore `5` alla tua variabile `number`.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
