---
id: 6551fb531ce7f773071a0b84
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Come puoi vedere, l'output della stampa di `type(text)` è `<class 'str'>`, il che significa che la tua variabile è una stringa, indicata come `str`.

Ora vai a una nuova riga e crea un'altra variabile chiamata `shift` e assegna a questa variabile il valore `3`.

# --hints--

Dovresti dichiarare una variabile chiamata `shift`. Fai attenzione a mettere il nome della variabile all'inizio della riga.

```js
assert.match(code, /^shift\s*=/m)
```

Dovresti assegnare il valore `3` alla tua variabile `shift`.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
