---
id: 6551f6f9ec6db270eb83f3db
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Puoi vedere `11` stampato sul terminale perché `'Hello World'` contiene 11 caratteri.

Un'altra funzione integrata utile è `type()`, che restituisce il tipo di dato di una variabile. Modifica la tua chiamata a `print()` per stampare il tipo di dato di `text`.

# --hints--

Dovresti chiamare la funzione `type()`.

```js
assert.match(code, /type\s*\(.*\)/)
```

Dovresti passare `text` alla funzione `type()` includendola tra le parentesi tonde.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

Dovresti passare `type(text)` alla funzione `print()` includendola tra le parentesi tonde. Fai attenzione a posizionare la chiamata alla funzione all'inizio della riga.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Non dovresti avere `print(len(text))` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
