---
id: 655261b2e1f2c197093f3993
title: Passaggio 41
challengeType: 20
dashedName: step-41
---

# --description--

Attualmente, gli spazi vengono criptati come `'c'`. Per mantenere la spaziatura originale nel messaggio in chiaro, ti servirà un'istruzione condizionale `if`. Questa è composta dalla parola chiave `if`, una condizione e due punti `:`.

```py
if x != 0:
    print(x)
```

Nell'esempio sopra, la condizione dell'istruzione `if` è `x != 0`. Il codice `print(x)`, all'interno del corpo dell'istruzione `if`, viene eseguito solo quando la condizione risulta `True` (in questo esempio, significa che `x` è diverso da zero).

All'inizio del tuo for loop, sostituisci `print(char == ' ')` con un'istruzione `if`. La condizione di questa istruzione `if` dovrebbe valutarsi come `True` se `char` è uno spazio vuoto e `False` altrimenti. All'interno del corpo `if`, stampa la stringa `'space!'`. Ricordati di indentare questa riga.

# --hints--

Non dovresti avere `print(char == ' ')` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

Dovresti sostituire `print(char == ' ')` con un'istruzione `if` che si attiva quando `char == ' '`. Non usare parentesi tonde per racchiudere la condizione `if` e ricorda di includere i due punti finali.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

Dovresti stampare la stringa `'space!'` all'interno della tua nuova istruzione `if`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
