---
id: 6552487e689f6e895f658717
title: Passaggio 33
challengeType: 20
dashedName: step-33
---

# --description--

Ora devi creare una variabile `new_char` alla fine del corpo del tuo loop. Imposta il suo valore su `alphabet[new_index]`.

# --hints--

Dovresti creare una variabile `new_char` all'interno del tuo loop `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_char\s*=/));
```

Dovresti impostare la tua variabile `new_char` su `alphabet[new_index]` alla fine del corpo del tuo loop.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
