---
id: 65521203d7165c7b84b22ad4
title: Passaggio 23
challengeType: 20
dashedName: step-23
---

# --description--

Un loop ti permette di scorrere sistematicamente una sequenza di elementi ed eseguire azioni su ciascuno di essi.

In questo caso, userai un loop `for`. Ecco come puoi iterare su `text`:

```py
for i in text:
```

`for` è la parola chiave che indica il tipo di loop. `i` è una variabile che assume sequenzialmente il valore degli elementi in `text`. L’istruzione termina con due punti, `:`.

Sotto la riga in cui hai dichiarato `alphabet`, scrivi un loop `for` per iterare su `text`. Usa `i` come variabile del loop.

Facendo così, c’è un errore nel terminale. Lo scoprirai nel passaggio successivo.

# --hints--

Dovresti usare la parola chiave `for` per creare un loop. Assicurati di mettere la parola chiave `for` all’inizio della riga e lascia uno spazio bianco dopo la parola chiave.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

Dovresti scrivere la variabile `i` dopo la parola chiave `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

Dovresti scrivere la parola chiave `in` dopo `for i `. Assicurati di lasciare uno spazio intorno alla parola chiave `in`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

Dovresti scrivere `text` dopo `for i in `. Non dimenticare di aggiungere il `:` finale.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

Il tuo loop `for` dovrebbe essere posizionato sotto la riga di codice `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
