---
id: 65520e6f2b9678799977f24d
title: Passaggio 18
challengeType: 20
dashedName: step-18
---

# --description--

Rimuovi l'ultima chiamata `print()`. Poi, invece di `text[0]`, passa `text[0].lower()` come argomento alla tua chiamata `.find()` e osserva l'output.

# --hints--

Non dovresti avere `print(text.lower())` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

Dovresti ancora stampare la tua variabile `index`.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

Dovresti aggiornare la tua chiamata `alphabet.find(text[0])` per usare `text[0].lower()` come argomento. Fai attenzione a posizionare la chiamata al metodo all'inizio della riga.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
