---
id: 65524d2a1a253b8bb5197ae2
title: Passaggio 28
challengeType: 20
dashedName: step-28
---

# --description--

`find` restituisce ancora `-1` per le lettere maiuscole e anche per il carattere spazio. Ti occuperai dello spazio più avanti.

Per ora, invece di iterare su `text`, modifica il loop `for` per iterare su `text.lower()`.

# --hints--

Dovresti modificare il tuo loop `for` per iterare su `text.lower()` invece che su `text`.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
