---
id: 655251308f31958d06cdf267
title: Passaggio 36
challengeType: 20
dashedName: step-36
---

# --description--

Al momento, il carattere criptato viene aggiornato a ogni iterazione. Sarebbe meglio archiviare la stringa criptata in una nuova variabile. Prima del tuo loop `for`, dichiara una variabile chiamata `encrypted_text` e assegna a questa variabile una stringa vuota (`''`).

# --hints--

Dovresti dichiarare una variabile chiamata `encrypted_text` prima del tuo loop `for`.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

Dovresti assegnare una stringa vuota alla tua variabile `encrypted_text`.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
