---
id: 655209a4a27dd37873c4cac3
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` restituisce l’indice del carattere corrispondente all’interno della stringa. Se il carattere non viene trovato, restituisce `-1`. Come puoi vedere, il primo carattere in `text`, `'H'` maiuscolo, non viene trovato, dato che `alphabet` contiene solo lettere minuscole.

Puoi trasformare una stringa nel suo equivalente minuscolo con il metodo `.lower()`. Aggiungi un’altra chiamata `print()` per stampare `text.lower()` e vedere l’output.

# --hints--

Dovresti avere ancora `print(index)` nel tuo codice. Fai attenzione a mettere la chiamata della funzione all’inizio della riga.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Dovresti stampare `text.lower()`. Fai attenzione a posizionare la chiamata della funzione all’inizio della riga.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
