---
id: 65525e359ca28d938baa82c5
title: Passaggio 39
challengeType: 20
dashedName: step-39
---

# --description--

Puoi ottenere lo stesso effetto di `a = a + b` usando l'operatore di assegnazione con addizione:

```py
a += b
```

L'operatore di assegnazione con addizione ti permette di aggiungere un valore a una variabile e poi assegnare il risultato a quella variabile.

Usa l'operatore `+=` per aggiungere un valore e assegnarlo contemporaneamente a `encrypted_text`.

# --hints--

Dovresti usare l'operatore di assegnazione con addizione per aggiungere `alphabet[new_index]` al valore corrente di `encrypted_text` all'interno del corpo del tuo loop.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
