---
id: 65522255d5b9cd80f335c6f2
title: Passaggio 29
challengeType: 20
dashedName: step-29
---

# --description--

Alla fine del corpo del tuo loop, dichiara una variabile chiamata `new_index` e assegna a questa variabile il valore di `index + shift`.

# --hints--

Dovresti dichiarare una variabile chiamata `new_index` all'interno del tuo loop `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

Dovresti assegnare `index + shift` alla tua nuova variabile alla fine del corpo del tuo loop `for`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
