---
id: 655208d59b131e7816f18c96
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

La funzione `print()` ti dà solo un output nella console, ma funzioni e metodi possono avere un valore di ritorno che puoi usare nel tuo codice.

Ora assegna `alphabet.find(text[0])` a una variabile chiamata `index`. In questo modo, `index` archivierà il valore restituito da `alphabet.find(text[0])`.

# --hints--

Dovresti dichiarare una variabile chiamata `index`. Fai attenzione a mettere il nome della variabile all’inizio della riga.

```js
assert.match(code, /^index\s*=/m)
```

Dovresti assegnare il valore restituito da `alphabet.find(text[0])` alla tua variabile `index`.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
