---
id: 65521fc818947e800bffe48a
title: Passaggio 26
challengeType: 20
dashedName: step-26
---

# --description--

All'interno del loop `for`, prima di stampare il carattere corrente, dichiara una variabile chiamata `index` e assegna a questa variabile il valore restituito da `alphabet.find(char)`.

# --hints--

Dovresti dichiarare una nuova variabile chiamata `index` all'inizio del tuo loop `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

Dovresti assegnare `alphabet.find(char)` alla tua nuova variabile `index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
