---
id: 6552008c0d9d9075cbec9772
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Userai il metodo `.find()` per trovare la posizione nell'alfabeto di ogni lettera nel tuo messaggio. Un metodo è simile a una funzione, ma appartiene a un oggetto.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

Qui sopra, il metodo `.find()` viene *chiamato su* `sentence` (la stringa in cui cercare), e `'r'` (il carattere da localizzare) viene passato come argomento. La chiamata `sentence.find('r')` restituirà `4`, che è l'indice della prima occorrenza di `'r'` in `sentence`.

Alla fine del tuo codice, chiama `.find()` su `alphabet` e passa `'z'` come argomento al metodo.

# --hints--

Dovresti chiamare il metodo `.find()`.

```js
assert.match(code, /find\s*\(.*\)/)
```

Dovresti chiamare il metodo `.find()` sulla variabile `alphabet`.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


Dovresti chiamare il metodo `.find()` sulla variabile `alphabet` e passare `'z'` al metodo. Fai attenzione a mettere la chiamata al metodo all'inizio della riga.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
