---
id: 655220953ba90d80514d7ee2
title: Passaggio 27
challengeType: 20
dashedName: step-27
---

# --description--

Attualmente, la funzione `print()` prende un solo argomento `char`, ma può prenderne multipli, separati da una virgola.

Aggiungi un secondo argomento a `print(char)` in modo che stampi il carattere e il suo indice nell'alfabeto.

# --hints--

Dovresti aggiungere `index` come secondo argomento alla tua chiamata `print(char)` esistente. Non dimenticare di separare gli argomenti con una virgola.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
