---
id: 65520f8282faf57a0db4f7fe
title: Passaggio 21
challengeType: 20
dashedName: step-21
---

# --description--

Come puoi vedere dall'output, `'h'` si trova all'indice `7` nella stringa `alphabet`. Ora devi trovare la lettera all'indice `7` più il valore di `shift`. Per questo, puoi usare l'operatore di addizione, `+`, nello stesso modo in cui lo useresti per un'addizione matematica.

Modifica la tua variabile `shifted` in modo che contenga il valore di `alphabet` all'indice `index + shift`.

# --hints--

Dovresti assegnare `alphabet[index + shift]` alla tua variabile `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
