---
id: 6553f8c570f9982e013a8886
title: Passaggio 57
challengeType: 20
dashedName: step-57
---

# --description--

Attualmente, ogni singola lettera viene sempre criptata con la stessa lettera, a seconda dell'offset specificato. E se l'offset fosse diverso per ogni lettera? Sarebbe molto più difficile da decifrare. Questo algoritmo è chiamato cifrario di Vigenère, dove l'offset per ogni lettera è determinato da un altro testo, chiamato *chiave*.

Inizia a trasformare il tuo cifrario di Cesare in un cifrario di Vigenère rimuovendo le due chiamate di funzione.

# --hints--

Dovresti rimuovere le tue due chiamate di funzione `caesar()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
