---
id: 6553efd6ada3f42aa2d75448
title: Passaggio 53
challengeType: 20
dashedName: step-53
---

# --description--

Ora dovresti vedere di nuovo l'output. Anche se questo approccio funziona, non migliora significativamente la riutilizzabilità del codice. Chiamare ripetutamente la tua funzione produrrà lo stesso risultato. Tuttavia, le funzioni possono essere dichiarate con *parametri* per introdurre versatilità e personalizzazione:

```py
def function_name(param_1, param_2):
    <code>
```

I parametri sono variabili che puoi usare all'interno della tua funzione. Una funzione può essere dichiarata con un numero diverso di parametri. Nell'esempio sopra, `param_1` e `param_2` sono parametri.

Modifica la dichiarazione della tua funzione in modo che accetti due parametri chiamati `message` e `offset`.

Dopo vedrai comparire un errore nel terminale. Scoprirai come risolverlo nei passaggi successivi.

# --hints--

La tua funzione `caesar` dovrebbe accettare `message` e `offset` come parametri. Ricorda di separare i parametri con una virgola.

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
