---
id: 6553980e0527fa115c705646
title: Passaggio 42
challengeType: 20
dashedName: step-42
---

# --description--

Ora, invece di stampare `'space!'`, usa l'operatore di assegnazione con addizione per aggiungere lo spazio (attualmente archiviato in `char`) al valore corrente di `encrypted_text`.

# --hints--

Non dovresti avere `print('space!')` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

Dovresti usare l'operatore `+=` per aggiungere `char` al valore corrente di `encrypted_text` all'interno della tua nuova istruzione `if`.

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
