---
id: 6553a572f7a65718f1e42e18
title: Passaggio 49
challengeType: 20
dashedName: step-49
---

# --description--

Una funzione è essenzialmente un blocco di codice riutilizzabile. Hai già incontrato alcune funzioni integrate, come `print()`, `find()` e `len()`. Ma puoi anche definire funzioni personalizzate in questo modo:

```py
def function_name():
    <code>
```

Una dichiarazione di funzione inizia con la parola chiave `def` seguita dal nome della funzione — un nome di variabile valido — e da una coppia di parentesi tonde. La dichiarazione termina con due punti.

Subito dopo la tua variabile `shift`, dichiara una funzione chiamata `caesar` e rientra tutte le righe successive per dare un corpo alla tua nuova funzione.

# --hints--

Dovresti usare la parola chiave `def` per dichiarare una nuova funzione.

```js
assert.match(code, /^def\s+/m)
```

Dovresti scrivere `caesar` come nome della funzione dopo la parola chiave `def`. Ricorda di aggiungere uno spazio dopo `def`.

```js
assert.match(code, /^def\s+caesar/m)
```

Dovresti aggiungere una coppia di parentesi tonde dopo il nome della funzione. Non dimenticare i due punti finali.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

Dovresti rientrare tutte le righe dopo `shift = 3` in modo che diventino il corpo della tua nuova funzione.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
