---
id: 6553ed69ece88d29594748aa
title: Passaggio 52
challengeType: 20
dashedName: step-52
---

# --description--

Per eseguire, una funzione deve essere chiamata (o invocata) aggiungendo una coppia di parentesi tonde dopo il suo nome, così:

```py
function_name()
```

Alla fine del tuo codice, chiama la tua funzione `caesar`. Fai attenzione all'indentazione.

# --hints--

Dovresti chiamare la tua funzione `caesar`. Assicurati di scrivere la chiamata della funzione all'inizio della riga.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
