---
id: 6553a44b1801991847d8cc69
title: Passaggio 47
challengeType: 20
dashedName: step-47
---

# --description--

Adesso, modifica la tua chiamata `print()` per stampare `'encrypted text:', encrypted_text` e mettila all’esterno del loop `for`, in modo che la stringa criptata venga stampata una sola volta.

# --hints--

Dovresti stampare `'encrypted text:', encrypted_text` dopo il tuo loop for.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/m) })
```

Non dovresti avere `print('char:', char, 'encrypted text:', encrypted_text)` nel tuo codice.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.notMatch(commentless_code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
