---
id: 6553f6086add4b2cbb99fd78
title: Passaggio 56
challengeType: 20
dashedName: step-56
---

# --description--

Alla fine del tuo codice, dopo la chiamata `caesar(text, shift)` esistente, chiama di nuovo la tua funzione.

Questa volta, passa la variabile `text` e l'intero `13` come argomenti.

# --hints--

Dovresti chiamare di nuovo la tua funzione, questa volta passando `text` e `13` come argomenti.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

Dovresti mantenere la chiamata alla funzione esistente.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
