---
id: 6553f4f66099802c6ae94613
title: Passaggio 55
challengeType: 20
dashedName: step-55
---

# --description--

Attualmente, il tuo codice genera un `TypeError`, perché la funzione `caesar` è definita con due parametri (`message` e `offset`), quindi si aspetta di essere chiamata con due *argomenti*.

Chiamare `caesar()` senza gli argomenti richiesti interrompe l'esecuzione del codice.

Assegna valori a `message` e `offset`, passando `text` e `shift` come argomenti alla chiamata della funzione `caesar`.

# --hints--

Dovresti passare `text` e `shift` come argomenti alla chiamata della tua funzione includendoli all'interno delle parentesi tonde. Non dimenticare di separare gli argomenti con una virgola.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
