---
id: 6553a755879b131a445e664c
title: Passaggio 44
challengeType: 20
dashedName: step-44
---

# --description--

Prova ad assegnare la stringa `'Hello Zaira'` alla tua variabile `text` e guarda cosa succede nel terminale.

Vedrai un'eccezione `string index out of range`. Non preoccuparti, presto capirai come risolverla!

# --hints--

Dovresti avere una variabile `text`.

```js
assert.match(code, /^text\s*=/m)
```

Dovresti assegnare la stringa `'Hello Zaira'` alla tua variabile `text`.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
