---
id: 6553995f412dd8122ed38e4a
title: Passaggio 43
challengeType: 20
dashedName: step-43
---

# --description--

Un'istruzione condizionale può anche avere una clausola `else`. Questa clausola può essere aggiunta alla fine di un'istruzione `if` per eseguire codice alternativo se la condizione dell'istruzione `if` è falsa:

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

Come puoi vedere nel tuo output, quando le iterazioni del loop raggiungono lo spazio, uno spazio viene aggiunto alla stringa criptata. Poi il codice esterno al blocco `if` viene eseguito e un `c` viene aggiunto alla stringa criptata.

Per risolvere, aggiungi una clausola `else` dopo `encrypted_text += char` e rientra tutte le righe di codice successive tranne la chiamata `print()`.

# --hints--

Dovresti creare una clausola `else`. Ricorda di includere i due punti finali.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

Dovresti rientrare le righe di codice dopo la clausola `else` tranne la chiamata `print()`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

Il tuo codice contiene sintassi e/o rientri non validi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
