---
id: 6553a7d8d05cbb1ae335a665
title: Passaggio 45
challengeType: 20
dashedName: step-45
---

# --description--

Quando il loop raggiunge la lettera `Z`, la somma `index + shift` supera l’ultimo indice della stringa `alphabet`. Di conseguenza, `alphabet[new_index]` sta cercando di usare un indice non valido, il che causa il lancio di un `IndexError`.

Puoi notare che l’output nel terminale si ferma allo spazio immediatamente prima del `Z`, l’ultimo `print` prima che venga lanciato l’errore.

In questo caso, l’operatore modulo (`%`) può essere usato per restituire il resto della divisione tra due numeri. Per esempio: `5 % 2` è uguale a `1`, perché 5 diviso 2 ha un quoziente di 2 e un resto di 1.

Racchiudi `index + shift` tra parentesi tonde, e calcola il modulo dell’espressione con `26`, che è la lunghezza di `alphabet`.

# --hints--

Dovresti avere `new_index = (index + shift) % 26` nella tua istruzione `else`.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

Il tuo codice contiene sintassi non valida e/o indentazione non valida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
