---
id: 65549f90cf78131c96ebcf28
title: Passaggio 61
challengeType: 20
dashedName: step-61
---

# --description--

Quando programmi, la leggibilità è fondamentale. I commenti funzionano come note efficaci, spiegando la logica dietro il tuo codice. Sono preziosi quando torni a un progetto dopo un po' di tempo e aiutano anche i tuoi colleghi a capire il codice.

In Python, puoi scrivere un commento usando un `#`. Tutto ciò che segue il `#` non verrà eseguito.

Prima della tua istruzione `if`, aggiungi un commento che dica `Append space to the message`.

# --hints--

Dovresti aggiungere un commento che dica `Append space to the message`. Non dimenticare il carattere `#` all'inizio.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
