---
id: 6554a88d5af937226f4a9121
title: Passaggio 65
challengeType: 20
dashedName: step-65
---

# --description--

Il metodo `.index()` è identico al metodo `.find()` ma lancia un'eccezione `ValueError` se non riesce a trovare la sottostringa.

Un `ValueError` è un'eccezione integrata che viene sollevata quando a una funzione viene passato un argomento con il tipo corretto ma un valore inappropriato.

Dopo aver incrementato `key_index`, dichiara una variabile chiamata `offset`. Trova l'indice che `key_char` ha in `alphabet` e assegnalo a `offset`. Usa il `.index()` per trovare l'indice.

# --hints--

Dovresti dichiarare una variabile chiamata `offset`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

La tua variabile `offset` dovrebbe archiviare il valore di `alphabet.index(key_char)`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
