---
id: 65549561463f0016876e852c
title: Passaggio 60
challengeType: 20
dashedName: step-60
---

# --description--

Poiché la tua chiave è più corta del testo che dovrai cifrare, dovrai ripeterla per generare l'intero testo cifrato.
All'inizio del corpo della tua funzione, dichiara una variabile `key_index` e impostala a zero.

# --hints--

Dovresti dichiarare una variabile chiamata `key_index` all'inizio del corpo della tua funzione.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

Dovresti assegnare `0` alla tua variabile `key_index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
