---
id: 6554930320d70414e7b6acc6
title: Passaggio 59
challengeType: 20
dashedName: step-59
---

# --description--

Elimina la tua variabile `shift`. Poi, dichiara una nuova variabile chiamata `custom_key` e assegna a questa variabile il valore della stringa `'python'`.

# --hints--

Dovresti eliminare la variabile `shift` e il suo valore.

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

Dovresti dichiarare una variabile chiamata `custom_key`.

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

Dovresti assegnare la stringa `'python'` alla tua variabile `custom_key`.

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
