---
id: 6554d25dc5ceaa354307a77e
title: Passaggio 70
challengeType: 20
dashedName: step-70
---

# --description--

Crittografia e decrittografia sono processi opposti e la tua funzione può fare entrambi con qualche piccola modifica.

Aggiungi un terzo parametro chiamato `direction` alla definizione della tua funzione. Inoltre, commenta le ultime due righe di codice per evitare errori nella console.

# --hints--

Dovresti trasformare le ultime due righe del tuo codice in commenti. Metti un `#` all'inizio di ogni riga.

```js
({ test: () => assert.match(code, /#\s*encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)\s*#\s*print\s*\(\s*encryption\s*\)/) })
```

Dovresti aggiungere `direction` come terzo parametro della tua funzione.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
encryption = vigenere(text, custom_key)
print(encryption)
--fcc-editable-region--
```
