---
id: 655491bd5b98b813fa5bedca
title: Passaggio 58
challengeType: 20
dashedName: step-58
---

# --description--

Ora modifica la dichiarazione della tua funzione: cambia il nome della funzione in `vigenere` e il secondo parametro in `key`.

# --hints--

Dovresti modificare il nome della tua funzione in `vigenere`.

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

La tua funzione `vigenere` dovrebbe prendere `message` e `key` come parametri.

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
