---
id: 6555d6e12c95701172b55709
title: Passaggio 77
challengeType: 20
dashedName: step-77
---

# --description--

Le funzioni possono essere chiamate con argomenti predefiniti. Un argomento predefinito indica il valore che la funzione dovrebbe assumere se l'argomento non viene passato. Per esempio, la funzione qui sotto accetta tre argomenti ma puoi chiamarla passando due argomenti. Il terzo assumerà il valore specificato per impostazione predefinita:

```py
def foo(a, b, c=value):
    <code>
```

Modifica la funzione `vigenere` in modo che il suo parametro `direction` abbia un valore predefinito di `1`.

# --hints--

Il parametro `direction` della tua funzione `vigenere` dovrebbe avere un valore predefinito di `1`.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

--fcc-editable-region--
def vigenere(message, key, direction):
--fcc-editable-region--
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
