---
id: 6555231eac4d9f19bd3d44b5
title: Passaggio 83
challengeType: 20
dashedName: step-83
---

# --description--

La parola chiave `pass` può essere usata come segnaposto per codice futuro. Non ha alcun effetto nel tuo codice ma può evitarti errori in caso di codice incompleto:

```py
def foo():
    pass
```

Chiamare `vigenere` con `1` per criptare e `-1` per decriptare va bene ma potrebbe risultare un po’ criptico. Crea una nuova funzione chiamata `encrypt` che prenda i parametri `message` e `key`, e usa `pass` per riempire il corpo della funzione.


# --hints--

Dovresti dichiarare una nuova funzione chiamata `encrypt`. Ricorda di usare la parola chiave `pass` all’interno del corpo.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

La tua funzione `encrypt` dovrebbe prendere `message` e `key` come parametri. L’ordine è importante.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
