---
id: 6555e1bfcd374e18c6be8e58
title: Passaggio 91
challengeType: 20
dashedName: step-91
---

# --description--

In Python, esiste un modo semplice per formattare le stringhe. Le *f-string* ti permettono di interpolare valori nelle tue stringhe.

Interpolare significa scrivere segnaposto che verranno sostituiti dai valori specificati quando il programma viene eseguito. Per esempio, puoi ottenere lo stesso risultato di `'Encrypted text: ' + text` con `f'Encrypted text: {text}'`. Devi mettere un `f` prima delle virgolette per creare la f-string e scrivere le variabili o le espressioni che vuoi interpolare tra parentesi graffe.

Modifica la prima chiamata `print()` per usare una f-string.

# --hints--

Dovresti modificare la chiamata `print('Encrypted text: ' + text)` in modo che stampi lo stesso risultato usando una f-string.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
