---
id: 6555e3f387381a19d5e00333
title: Passaggio 92
challengeType: 20
dashedName: step-92
---

# --description--

Adesso, modifica `print('Key: ' + custom_key)` per usare una f-string.

# --hints--

Dovresti modificare la chiamata `print('Key: ' + custom_key)` in modo che stampi lo stesso risultato usando una f-string.

```js
({ test: () => assert.match(code, /^print\s*\(\s*f("|')Key:\s\{\s*custom_key\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print(f'Encrypted text: {text}')
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
