---
id: 6555df0c81300b175308557d
title: Passaggio 90
challengeType: 20
dashedName: step-90
---

# --description--

Sotto la chiamata `print()` che hai appena aggiunto, aggiungi un’altra chiamata `print()` per stampare `Key: python` concatenando la stringa `'Key: '` e il valore della variabile `custom_key`.

# --hints--

Dovresti avere una chiamata `print()` che stampa `Key: python` concatenando la stringa `'Key: '` e il valore della variabile `custom_key`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^print\s*\(\s*("|')Key:\s\1\s*\+\s*custom_key\s*\)/m);
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)

--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
