---
id: 655a2a7210094920069b117c
title: Passaggio 40
challengeType: 20
dashedName: step-40
---

# --description--

Gli operatori di confronto ti permettono di confrontare due oggetti in base ai loro valori. Puoi usare un operatore di confronto posizionandolo tra gli oggetti che vuoi confrontare.
Restituiscono un valore *booleano* — cioè `True` o `False` — a seconda della veridicità dell'espressione.

Python ha i seguenti operatori di confronto:

<table>
  <thead>
    <tr>
      <th>Operatore</th>
      <th>Significato</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Uguale</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Non uguale</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Maggiore di</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Minore di</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>Maggiore o uguale a</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>Minore o uguale a</td>
    </tr>
  </tbody>
</table>

All'inizio del corpo del tuo loop, stampa il risultato del confronto tra `char` e uno spazio (`' '`). Usa l'operatore di uguaglianza `==` per farlo.

# --hints--

Dovresti confrontare `char` con uno spazio usando l'operatore di uguaglianza all'interno del tuo loop `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

Dovresti stampare il risultato del confronto tra `char` e uno spazio all'interno del tuo loop `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

Dovresti stampare il risultato del confronto tra `char` e uno spazio all'inizio del tuo loop.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
