---
id: 65688b5a1655a7a6caede847
title: Passaggio 31
challengeType: 20
dashedName: step-31
---

# --description--

Quando provi a modificare i singoli caratteri di una stringa come hai fatto nel passaggio precedente, ottieni un `TypeError`, che si verifica quando nel codice viene usato un oggetto di tipo non appropriato.

Come puoi vedere dal messaggio di errore, le stringhe non supportano l’assegnazione di elementi, perché sono immutabili. Tuttavia, una variabile può essere riassegnata a un’altra stringa:

```py
message = 'Hello World'
message = 'Hello there!'
```

Elimina la riga `text[0]` e riassegna a `text` la stringa `'Albatross'`.

# --hints--

Non dovresti avere `text[0] = 'C'` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

Dovresti riassegnare a `text` la stringa `'Albatross'`. Non modificare `text = 'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
