---
id: 65688a50e6c998a21d8e41d3
title: Passaggio 30
challengeType: 20
dashedName: step-30
---

# --description--

Le stringhe sono immutabili, il che significa che non possono essere modificate una volta create. Per esempio, potresti pensare che il seguente codice cambi il valore di `my_string` nella stringa `'train'`, ma questo non è valido:

```py
my_string = 'brain'
my_string[0] = 't'
```

Conferma consultando la prima lettera in `text` usando la notazione a parentesi e prova a cambiarla in un carattere a tua scelta. Vedrai che l'output scomparirà e apparirà un errore.

# --hints--

Dovresti ancora avere `text = 'Hello World'` nel tuo codice.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

Dovresti consultare la prima lettera in `text` con `text[0]`.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

Dovresti usare l'operatore `=` per assegnare un carattere a tua scelta a `text[0]`. Non dimenticare di racchiudere il carattere tra virgolette singole o virgolette doppie.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
