---
id: 6568b7a40c7ba0ccbc4b4425
title: Passaggio 32
challengeType: 20
dashedName: step-32
---

# --description--

Come puoi vedere, ogni carattere della stringa riassegnata viene stampato all’interno del loop.

Torna alla stringa originale eliminando la riassegnazione `text`.

# --hints--

Non dovresti avere `text = 'Albatross'` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

Dovresti ancora avere una variabile `text` con il valore `'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
