---
id: 656886ec4e4d9b96ded3c565
title: Passaggio 25
challengeType: 20
dashedName: step-25
---

# --description--

La variabile di iterazione può avere qualsiasi nome valido, ma è comodo darle un nome significativo.

Rinomina la tua variabile `i` in `char`.

# --hints--

Dovresti sostituire `i` con `char`. 

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
