---
id: 65687d2f8c7ee27b0446cef3
title: Passaggio 24
challengeType: 20
dashedName: step-24
---

# --description--

Il codice da eseguire a ogni iterazione — posizionato dopo il `:` — costituisce il corpo del loop. Questo codice deve essere indentato. In Python, si consiglia di usare 4 spazi per ogni livello di indentazione. Questo livello indentato è un blocco di codice.

```py
for i in text:
    print(i)
```

Python si basa sull'indentazione per indicare i blocchi di codice. Due punti alla fine di una riga segnalano che seguirà un nuovo blocco di codice indentato.

Quindi, quando non si trova un blocco indentato dopo i due punti finali, l'esecuzione del codice si ferma e viene lanciato un `IndentationError`. Questo codice non mostrerà l'output e invece lancerà un `IndentationError`:

```py
for i in text:
print(i)
```

Dai un corpo al tuo loop `for` aggiungendo una chiamata a `print(i)`. Ricorda di indentare il corpo del loop.

# --hints--

Dovresti aggiungere `print(i)` al corpo del tuo loop `for`. Fai attenzione all'indentazione.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
