---
id: 659526f3e5424b3eac4b6aa7
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Successivamente, stampa la variabile `index` nella console.

# --hints--

Dovresti stampare la tua variabile `index`. Fai attenzione a mettere la chiamata della funzione all'inizio della riga.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
