---
id: 65952b70924edb449faf1950
title: Passaggio 19
challengeType: 20
dashedName: step-19
---

# --description--

Dichiara una nuova variabile chiamata `shifted`. Usa la notazione a parentesi per consultare il valore di `alphabet` all'indice `index` e assegnalo alla tua nuova variabile.

# --hints--

Dovresti dichiarare una variabile chiamata `shifted`. Fai attenzione a mettere il nome della variabile all'inizio della riga.

```js
assert.match(code, /^shifted\s*=/m)
```

Dovresti assegnare `alphabet[index]` alla tua variabile `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
