---
id: 65951b637ddec63611c482c0
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Puoi usare la funzione integrata `print()` per stampare l’output del tuo codice sul terminale.

Le funzioni sono blocchi di codice riutilizzabili che puoi chiamare, o invocare, per eseguire il loro codice quando ne hai bisogno. Per chiamare una funzione, devi solo scrivere una coppia di parentesi tonde accanto al suo nome. Imparerai di più sulle funzioni molto presto.

Per ora, vai a una nuova riga e aggiungi una chiamata vuota alla funzione `print()`. Non dovresti ancora vedere alcun output.

# --hints--

Dovresti avere `print()` nel tuo codice. Fai attenzione a posizionare la chiamata alla funzione all’inizio della riga.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
