---
id: 659530ca823d6c47f61daa8a
title: Passaggio 20
challengeType: 20
dashedName: step-20
---

# --description--

Stampa la tua variabile `shifted`.

# --hints--

Dovresti stampare la tua variabile `shifted`. Fai attenzione a mettere la chiamata della funzione all'inizio della riga.

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
