---
id: 659db3e3670d3e712be82593
title: Passaggio 22
challengeType: 20
dashedName: step-22
---

# --description--

Ripetere il processo di individuare la lettera nell'alfabeto e determinare la lettera spostata per ogni carattere in `text` può essere noioso. Per fortuna, puoi semplificarlo usando un loop.

Per ora, elimina tutte le righe di codice sotto la dichiarazione della variabile `alphabet`.

# --hints--

Dovresti ancora avere `text = 'Hello World'` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

Dovresti ancora avere `shift = 3` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

Dovresti ancora avere `alphabet = 'abcdefghijklmnopqrstuvwxyz'` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

Dovresti eliminare la variabile `index` e il suo valore.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

Non dovresti avere `print(index)` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

Dovresti eliminare la variabile `shifted` e il suo valore.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

Non dovresti avere `print(shifted)` nel tuo codice.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

Il tuo codice contiene sintassi e/o indentazione non valide.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
