---
id: 6624e0ec0bf47017eed84b1f
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Il primo tipo di cifrario che costruirai si chiama cifrario di *Cesare*. In particolare, prenderai ogni lettera del tuo messaggio, troverai la sua posizione nell'alfabeto, prenderai la lettera che si trova 3 posizioni dopo nell'alfabeto e sostituirai la lettera originale con quella nuova.

Per implementare questo, userai il metodo `.find()` discusso nel passaggio precedente. Modifica la tua chiamata `.find()` esistente passando `text[0]` come argomento invece di `'z'`.

# --hints--

Dovresti modificare la tua chiamata `alphabet.find('z')` esistente passando `text[0]` al metodo.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
