---
id: 66605a64a9822f01ec5ae8ce
title: Passaggio 50
challengeType: 20
dashedName: step-50
---

# --description--

In Python, l'*ambito* di una variabile determina dove quella variabile può essere consultata:

- Le variabili definite all'esterno di una funzione hanno un ambito *globale* e possono essere consultate da qualsiasi parte del tuo codice.

- Le variabili definite all'interno di una funzione sono *locali* a quella funzione e non possono essere consultate all'esterno di essa.

Per vedere questo in azione, prova a stampare la variabile `alphabet` alla fine del tuo codice. Questo genererà un'eccezione `NameError`.

Dovresti vedere un messaggio di errore che indica che `alphabet` non è definito. Questo perché `alphabet` è definito all'interno della funzione `caesar` e non è consultabile all'esterno di essa.

# --hints--

Dovresti provare a stampare la variabile `alphabet` all'esterno della funzione caesar.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
