---
id: 587d78a7367417b2b2512ae1
title: CSSアニメーションを使って動きを作る
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7amZfW'
forumTopicId: 301051
dashedName: create-movement-using-css-animation
---

# --description--

要素に`position`（`fixed`や`relative`など）が指定されている場合、CSSオフセットプロパティの`right`、`left`、`top`、`bottom`をアニメーションルールで使って動きを作ることができます。

以下の例のように、`top`の`50%`プロパティを50pxに設定し、最初（`0%`）と最後（`100%`）のキーフレームでは0pxに設定することで、アイテムを下に押し下げてから上に戻す動きを作れます。

```css
@keyframes rainbow {
  0% {
    background-color: blue;
    top: 0px;
  }
  50% {
    background-color: green;
    top: 50px;
  }
  100% {
    background-color: yellow;
    top: 0px;
  }
}
```

# --instructions--

`div`アニメーションに横方向の動きを追加してください。`left`オフセットプロパティを使い、`@keyframes`ルールに次のように追加します。`0%`で0ピクセル、`50%`で25ピクセル、`100%`で-25ピクセルに動くようにします。エディター内の`top`プロパティは置き換えないでください。アニメーションには縦方向と横方向の両方の動きが必要です。

# --hints--

`@keyframes`の`0%`ルールは、0pxの`left`オフセットを使うべきです。

```js
assert.match(code,/[^50]0%\s*?{[\s\S]*?left:\s*?0px(;[\s\S]*?|\s*?)}/gi);
```

`@keyframes`の`50%`ルールは、25pxの`left`オフセットを使うべきです。

```js
assert.match(code,/50%\s*?{[\s\S]*?left:\s*?25px(;[\s\S]*?|\s*?)}/gi);
```

`@keyframes`の`100%`ルールは、-25pxの`left`オフセットを使うべきです。

```js
assert.match(code,/100%\s*?{[\s\S]*?left:\s*?-25px(;[\s\S]*?|\s*?)}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;

    }
    50% {
      background-color: green;
      top: 50px;

    }
    100% {
      background-color: yellow;
      top: 0px;

    }
  }
</style>

<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;
      left: 0px;
    }
    50% {
      background-color: green;
      top: 50px;
      left: 25px;
    }
    100% {
      background-color: yellow;
      top: 0px;
      left: -25px;
    }
  }
</style>
<div id="rect"></div>
```
