---
id: 587d78a5367417b2b2512ada
title: ホバー時に要素を拡大するためにCSSのtransform scaleプロパティを使う
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLPJuM'
forumTopicId: 301077
dashedName: use-the-css-transform-scale-property-to-scale-an-element-on-hover
---

# --description--

`transform` プロパティには、要素を拡大・移動・回転・傾斜させるなど、さまざまな機能があります。`:hover` のような要素の特定の状態を指定する疑似クラスと組み合わせて使うと、`transform` プロパティで簡単に要素にインタラクティブな動きを加えられます。

以下は、ユーザーが段落要素にホバーしたときに、その要素を元のサイズの2.1倍に拡大する例です。

```css
p:hover {
  transform: scale(2.1);
}
```

**注意:** `div` 要素にtransformを適用すると、そのdiv内に含まれる子要素にも影響します。

# --instructions--

`hover` の`div` 状態に対するCSSルールを追加し、ユーザーがホバーしたときに`transform` 要素を元のサイズの1.1倍に拡大するために`div` プロパティを使ってください。

# --hints--

ユーザーがホバーしたときに`div` 要素のサイズが1.1倍に拡大するようにしてください。

```js
assert.match(code, /div:hover\s*?{\s*?transform:\s*?scale\(1\.1\);/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }



</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }
  div:hover {
    transform: scale(1.1);
  }
</style>
<div></div>
```
