---
id: 587d78a4367417b2b2512ad2
title: 三次色について学ぼう
challengeType: 0
forumTopicId: 301057
dashedName: learn-about-tertiary-colors
---

# --description--

コンピュータのモニターやデバイスの画面は、赤、緑、青の光の量を組み合わせてさまざまな色を作り出します。これは現代の色彩理論でRGB加法混色モデルとして知られています。赤（R）、緑（G）、青（B）は原色と呼ばれます。2つの原色を混ぜると、シアン（G + B）、マゼンタ（R + B）、イエロー（R + G）という二次色ができます。これらの色は補色のチャレンジで見ました。これらの二次色は、作成に使われなかった原色の補色であり、カラーホイール上でその原色の反対側に位置します。例えば、マゼンタは赤と青で作られ、緑の補色です。

三次色は、原色とその隣接する二次色を組み合わせた結果です。例えば、RGBカラーモデルでは、赤（原色）とイエロー（二次色）を混ぜるとオレンジ（三次色）になります。これにより、シンプルなカラーホイールに6色が追加され、合計12色になります。

デザインで調和のとれた組み合わせを作るために、さまざまな色の選び方があります。三次色を使う例の一つにスプリット補色配色があります。この配色は基準となる色から始まり、その補色の隣接する2色と組み合わせます。この3色はデザインに強い視覚的コントラストを与えますが、2つの補色を使うよりも控えめです。

以下はスプリット補色配色で作られた3色です：

<table><thead><tr><th>色</th><th>16進コード</th></tr></thead><thead></thead><tbody><tr><td>オレンジ</td><td>#FF7F00</td></tr><tr><td>シアン</td><td>#00FFFF</td></tr><tr><td>ラズベリー</td><td>#FF007F</td></tr></tbody></table>

# --instructions--

`background-color`、`orange`、`cyan` クラスの `raspberry` プロパティをそれぞれの色に変更してください。色名ではなく16進コードを使うことを確認してください。

# --hints--

クラス `div` の `orange` 要素はオレンジの `background-color` を持つべきです。

```js
const orangeElement = document.querySelector('.orange');
const orangeStyle = window.getComputedStyle(orangeElement);
assert.equal(orangeStyle?.backgroundColor, 'rgb(255, 127, 0)');
```

クラス `div` の `cyan` 要素はシアンの `background-color` を持つべきです。

```js
const cyanElement = document.querySelector('.cyan');
const cyanStyle = window.getComputedStyle(cyanElement);
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

クラス `div` の `raspberry` 要素はラズベリーの `background-color` を持つべきです。

```js
const raspberryElement = document.querySelector('.raspberry');
const raspberryStyle = window.getComputedStyle(raspberryElement);
assert.equal(raspberryStyle?.backgroundColor, 'rgb(255, 0, 127)');
```

色クラスのすべての `background-color` 値は色名ではなく16進コードであるべきです。

```js
assert.notMatch(code, /background-color:\s(orange|cyan|raspberry)/);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .raspberry {
    background-color: #000000;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>

<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #FF7F00;
  }

  .cyan {
    background-color: #00FFFF;
  }

  .raspberry {
    background-color: #FF007F;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>
<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```
