---
id: 587d78a9367417b2b2512ae8
title: ベジェ曲線の仕組みを学ぼう
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bDrs8'
forumTopicId: 301058
dashedName: learn-how-bezier-curves-work
---

# --description--

前回のチャレンジでは`animation-timing-function`プロパティと、アニメーションの再生時間にわたる速度を変えるいくつかのキーワードを紹介しました。CSSにはキーワード以外に、ベジェ曲線を使ってアニメーションの再生方法をさらに細かく制御できるオプションがあります。

CSSアニメーションでは、ベジェ曲線は`cubic-bezier`関数とともに使われます。曲線の形状がアニメーションの再生の仕方を表します。この曲線は1×1の座標系上に存在します。この座標系のX軸はアニメーションの時間（時間軸のように考えてください）、Y軸はアニメーションの変化を表します。

`cubic-bezier`関数は、この1×1のグリッド上にある4つの主要な点で構成されます：`p0`、`p1`、`p2`、`p3`です。`p0`と`p3`はあらかじめ設定されており、それぞれ原点(0, 0)と(1, 1)に常に位置しています。残りの2つの点のxとyの値を設定し、それらをグリッド上のどこに置くかでアニメーションがたどる曲線の形が決まります。これはCSSで、`p1`と`p2`の「アンカー」ポイントのxとyの値を`(x1, y1, x2, y2)`の形で宣言することで行います。まとめると、以下はCSSコードでのベジェ曲線の例です：

```css
animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
```

上の例では、各点のxとyの値が等しく（x1 = 0.25 = y1、x2 = 0.75 = y2）、幾何学の授業で習ったように、原点から点(1, 1)まで伸びる直線になります。このアニメーションはアニメーションの長さにわたる要素の線形変化であり、`linear`キーワードを使うのと同じです。つまり、一定の速度で変化します。

# --instructions--

idが`ball1`の要素について、`animation-timing-function`プロパティの値を`linear`から同等の`cubic-bezier`関数の値に変更してください。上の例で示したポイントの値を使ってください。

# --hints--

idが`animation-timing-function`の要素の`ball1`プロパティの値は、線形に相当する`cubic-bezier`関数であるべきです。

```js
const ballOne = document.querySelector('#ball1');
const ballOneStyle = window.getComputedStyle(ballOne); 
assert.equal(
  ballOneStyle?.animationTimingFunction,
    'cubic-bezier(0.25, 0.25, 0.75, 0.75)'
);
```

idが`animation-timing-function`の要素の`ball2`プロパティの値は変更しないでください。

```js
const ballTwoElement = document.querySelector('#ball2');
const ballTwoStyle = window.getComputedStyle(ballTwoElement); 
const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(
  ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
