---
id: 587d781e367417b2b2512aca
title: CSSオフセットで相対位置指定要素を移動する
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bQEA4'
forumTopicId: 301065
dashedName: move-a-relatively-positioned-element-with-css-offsets
---

# --description--

`top` または `bottom` の CSSオフセット、そして `left` または `right` は、ブラウザーに対して要素を文書の通常のフロー上の位置からどれだけずらすかを指示します。要素を特定の位置からずらすことで、参照側から離れた方向（つまり反対方向）に要素が移動します。前のチャレンジで見たように、`top` のオフセットを使うと `h2` が下方向に移動しました。同様に、`left` のオフセットを使うと要素は右方向に移動します。

# --instructions--

CSSオフセットを使って `h2` を右に15ピクセル、上に10ピクセル移動させてください。

# --hints--

コードでは CSSオフセットを使って `h2` を上に10px相対位置指定してください。つまり、通常の位置の `bottom` から10px離すように移動させます。

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.bottom, '10px');
```

コードでは CSSオフセットを使って `h2` を右に15px相対位置指定してください。つまり、通常の位置の `left` から15px離すように移動させます。

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.left , '15px');
```

# --seed--

## --seed-contents--

```html
<head>
<style>
  h2 {
    position: relative;


  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<head>
<style>
  h2 {
    position: relative;
    left: 15px;
    bottom: 10px;
  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
