---
id: 587d78a3367417b2b2512ad1
title: 補色について学ぼう
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MD3Tr'
forumTopicId: 301056
dashedName: learn-about-complementary-colors
---

# --description--

色彩理論とデザインへの影響は奥深いテーマで、以下のチャレンジでは基本のみを扱います。ウェブサイトでは、色がコンテンツに注目を集めたり、感情を呼び起こしたり、視覚的な調和を生み出したりします。色の組み合わせを変えることでウェブサイトの見た目が大きく変わり、コンテンツに合ったカラーパレットを選ぶには多くの工夫が必要です。

カラーホイールは色同士の関係を視覚化する便利なツールで、似た色相が隣り合い、異なる色相は離れて配置されています。ホイール上で向かい合う2色は補色と呼ばれます。これらの色は組み合わせると互いに「打ち消し合い」、グレーの色になります。しかし、並べて配置すると色がより鮮やかに見え、強い視覚的コントラストを生み出します。

補色の例とその16進コードは以下の通りです：

<blockquote>赤 (#FF0000) とシアン (#00FFFF)<br>緑 (#00FF00) とマゼンタ (#FF00FF)<br>青 (#0000FF) と黄色 (#FFFF00)</blockquote>

これは、学校で教わったことがある多くの人に馴染みのある古いRYB色モデルとは異なり、RYBは異なる原色と補色を持ちます。現代の色彩理論では、加法混色のRGBモデル（コンピューター画面のような）と減法混色のCMY(K)モデル（印刷のような）を使います。

オンラインには補色を見つけるオプションがある色選択ツールが多数あります。

**注意：** 色はページに視覚的な興味を加える強力な方法ですが、色だけで重要な情報を伝えるべきではありません。視覚障害のあるユーザーがその内容を理解できない場合があるためです。この問題はApplied Accessibilityチャレンジで詳しく扱います。

# --instructions--

`background-color`と`blue`クラスの`yellow`プロパティをそれぞれの色に変更してください。色が白い背景と比べて隣り合うときにどのように見えるかに注目しましょう。

# --hints--

`div`クラスの`blue`要素は青の`background-color`を持つべきです。

```js
const blueElement = document.querySelector('.blue');
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

`div`クラスの`yellow`要素は黄色の`background-color`を持つべきです。

```js
const yellowElement = document.querySelector('.yellow');
const yellowStyle = window.getComputedStyle(yellowElement);
assert.equal(yellowStyle?.backgroundColor, 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }
  .blue {
    background-color: #000000;
  }
  .yellow {
    background-color: #000000;
  }
  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="blue"></div>
<div class="yellow"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }
  .blue {
    background-color: blue;
  }
  .yellow {
    background-color: yellow;
  }
  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="blue"></div>
<div class="yellow"></div>
```
