---
id: 587d78a8367417b2b2512ae7
title: キーワードでアニメーションのタイミングを変更する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKvwCM'
forumTopicId: 301045
dashedName: change-animation-timing-with-keywords
---

# --description--

CSSアニメーションでは、`animation-timing-function` プロパティがアニメーションの期間中にアニメーション要素がどれだけ速く変化するかを制御します。もしアニメーションが、ある時間内に地点Aから地点Bへ移動する車だとすると（あなたの `animation-duration`）、`animation-timing-function` はその移動中に車がどのように加速・減速するかを示します。

よく使われるオプションには、いくつかの定義済みキーワードがあります。例えば、デフォルト値は `ease` で、これはゆっくり始まり、途中で速くなり、最後にまたゆっくりになる動きです。他には、最初に速くて後で遅くなる `ease-out`、最初に遅くて最後に速くなる `ease-in`、またはアニメーション全体を通して一定の速度で動く `linear` などがあります。

# --instructions--

idが `ball1` と `ball2` の要素に、それぞれ `animation-timing-function` プロパティを追加し、`#ball1` を `linear` に、`#ball2` を `ease-out` に設定してください。アニメーション中の動きの違いに注目してください。ただし、両方の要素は `animation-duration` が2秒で共通なので、最後は一緒に終わります。

# --hints--

idが `animation-timing-function` の要素の `ball1` プロパティの値は `linear` にしてください。

```js
const ballOne =document.querySelector('#ball1'); 
const ballOneStyle = window.getComputedStyle(ballOne); 

const ball1Animation = __helpers.removeWhiteSpace(ballOneStyle?.animationTimingFunction);
assert.isTrue(ball1Animation == 'linear' || ball1Animation == 'cubic-bezier(0,0,1,1)');
```

idが `animation-timing-function` の要素の `ball2` プロパティの値は `ease-out` にしてください。

```js
const ballTwo = document.querySelector('#ball2'); 
const ballTwoStyle = window.getComputedStyle(ballTwo); 

const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)');
```

# --seed--

## --seed-contents--

```html
<style>

  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;

  }
  #ball2 {
    left:56%;

  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left:56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
