---
id: 587d78a6367417b2b2512adb
title: CSSのtransformプロパティのskewXを使って要素をX軸方向に傾ける
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLP8Sr'
forumTopicId: 301074
dashedName: use-the-css-transform-property-skewx-to-skew-an-element-along-the-x-axis
---

# --description--

`transform`プロパティの次の機能は`skewX()`で、指定した角度だけ選択した要素をX（水平）軸方向に傾けます。

次のコードは段落要素をX軸方向に-32度傾けています。

```css
p {
  transform: skewX(-32deg);
}
```

# --instructions--

`bottom`プロパティを使って、idが`transform`の要素をX軸方向に24度傾けてください。

# --hints--

idが`bottom`の要素はX軸方向に24度傾ける必要があります。

```js
assert.match(code, /#bottom\s*?{\s*?.*?\s*?transform:\s*?skewX\(24deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;

  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
