---
id: 587d7791367417b2b2512ab3
title: text-align プロパティを使って視覚的なバランスを作る
challengeType: 0
videoUrl: 'https://scrimba.com/c/c3b4EAp'
forumTopicId: 301053
dashedName: create-visual-balance-using-the-text-align-property
---

# --description--

このカリキュラムのセクションは応用視覚デザインに焦点を当てています。最初のグループのチャレンジでは、与えられたカードレイアウトを基にいくつかの基本的な原則を示します。

テキストはウェブコンテンツの大部分を占めることが多いです。CSSには`text-align`プロパティでテキストの配置を調整するいくつかの方法があります。

`text-align: justify;`はテキストの各行を均等な幅にスペース調整します。

`text-align: center;`はテキストを中央揃えにします。

`text-align: right;`はテキストを右揃えにします。

そして`text-align: left;`（デフォルト）はテキストを左揃えにします。

# --instructions--

「Google」と書かれた`h4`タグのテキストを中央揃えにしてください。次に、Googleの創立に関する情報が書かれた段落タグを両端揃えにしてください。

# --hints--

`h4`タグに対してtext-alignプロパティを使い、値を`center`に設定してください。

```js
const h4Element =document.querySelector('h4')
const h4Style = window.getComputedStyle(h4Element);
assert.equal(h4Style?.textAlign, 'center');
```

`p`タグに対してtext-alignプロパティを使い、値を`justify`に設定してください。

```js
const pElement =document.querySelector('p')
const pStyle = window.getComputedStyle(pElement);
assert.equal(pStyle?.textAlign, 'justify');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {

  }
  p {

  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
  }
  p {
    text-align: justify;
  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
