---
id: 587d78a9367417b2b2512aea
title: ベジェ曲線を使って動きをより自然にする
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7akWUv'
forumTopicId: 301063
dashedName: make-motion-more-natural-using-a-bezier-curve
---

# --description--

このチャレンジでは、ジャグリングされているボールの動きを再現するために要素をアニメーションさせます。これまでのチャレンジでは`linear`と`ease-out`の三次ベジェ曲線を扱いましたが、どちらもジャグリングの動きを正確に表現していません。ここではベジェ曲線をカスタマイズする必要があります。

`animation-timing-function`は`animation-iteration-count`が無限に設定されているとき、各キーフレームで自動的にループします。アニメーションの途中（`50%`）にキーフレームルールが設定されているため、ボールの上昇と下降の動きで同じアニメーション進行が2回繰り返されます。

次の三次ベジェ曲線はジャグリングの動きをシミュレートします：

```css
cubic-bezier(0.3, 0.4, 0.5, 1.6);
```

y2の値が1より大きいことに注目してください。三次ベジェ曲線は1×1の座標系にマッピングされ、x値は0から1の範囲でしか受け付けませんが、y値は1より大きい数値を設定できます。これにより、ジャグリングボールの跳ねる動きを理想的に表現できます。

# --instructions--

IDが`animation-timing-function`の要素の`green`の値を、x1、y1、x2、y2の値をそれぞれ0.311、0.441、0.444、1.649に設定した`cubic-bezier`関数に変更してください。

# --hints--

IDが`animation-timing-function`の要素の`green`プロパティの値は、指定されたx1、y1、x2、y2の値を持つ`cubic-bezier`関数であるべきです。

```js
const greenElement = document.querySelector('#green');
 const greenStyle = window.getComputedStyle(greenElement);
assert.equal(greenStyle?.animationTimingFunction, 'cubic-bezier(0.311, 0.441, 0.444, 1.649)');
```

# --seed--

## --seed-contents--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.69, 0.1, 1, 0.1);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.311, 0.441, 0.444, 1.649);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```
