---
id: 587d781e367417b2b2512acc
title: 固定位置で要素をブラウザーウィンドウに固定する
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDNUR'
forumTopicId: 301061
dashedName: lock-an-element-to-the-browser-window-with-fixed-positioning
---

# --description--

CSSが提供する次のレイアウト方式は`fixed`位置で、これはブラウザーウィンドウに対して要素を固定する絶対位置指定の一種です。絶対位置指定と同様に、CSSオフセットプロパティと一緒に使われ、要素は文書の通常のフローから外れます。他の要素はその位置を「認識」しなくなるため、他の部分のレイアウト調整が必要になる場合があります。

`fixed`位置と`absolute`位置の大きな違いの一つは、固定位置の要素はユーザーがスクロールしても動かないことです。

# --instructions--

コード内のナビゲーションバーには`navbar`というidが付いています。その`position`を`fixed`に変更し、`top`から0ピクセル、`left`から0ピクセルのオフセットを設定してください。コードを追加したら、プレビューウィンドウをスクロールしてナビゲーションが固定されていることを確認してください。

# --hints--

`#navbar`要素には`position`が`fixed`に設定されている必要があります。

```js
const navbarElement = document.querySelector('#navbar');
const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.position, 'fixed');
```

コードでは`top`要素の`#navbar` CSSオフセットを0ピクセルに設定してください。

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
assert.equal(navbarStyle?.top, '0px');
```

コードでは`left`要素の`#navbar` CSSオフセットを0ピクセルに設定してください。

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.left, '0px');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {



    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```

# --solutions--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {
    position: fixed;
    top: 0;
    left: 0;
    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```
