---
id: 587d78a8367417b2b2512ae3
title: 無限のアニメーション回数を使って要素を継続的にアニメーションさせる
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDVfq'
forumTopicId: 301041
dashedName: animate-elements-continually-using-an-infinite-animation-count
---

# --description--

前のチャレンジでは、いくつかのアニメーションプロパティと`@keyframes`ルールの使い方を学びました。もう一つのアニメーションプロパティが`animation-iteration-count`で、これはアニメーションを何回ループさせるかを制御できます。例を見てみましょう。

```css
animation-iteration-count: 3;
```

この場合、アニメーションは3回実行した後に停止しますが、その値を`infinite`に設定するとアニメーションを継続的に実行できます。

# --instructions--

右側のボールが継続的に跳ね続けるように、`animation-iteration-count`プロパティを`infinite`に変更してください。

# --hints--

`animation-iteration-count`プロパティの値は`infinite`にしてください。

```js
const ballElement = document.querySelector('#ball');
const ballStyle = window.getComputedStyle(ballElement); 
assert.equal(ballStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: 3;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```
