---
id: 587d781e367417b2b2512acb
title: 絶対位置指定で要素を親に固定する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJ7c3'
forumTopicId: 301060
dashedName: lock-an-element-to-its-parent-with-absolute-positioning
---

# --description--

CSSの`position`プロパティの次の選択肢は`absolute`で、要素を親コンテナに対して固定します。`relative`の位置指定とは異なり、これは要素を文書の通常のフローから外すため、周囲の要素はそれを無視します。位置調整にはCSSオフセットプロパティ（topまたはbottom、leftまたはright）が使われます。

絶対位置指定の注意点は、最も近い*位置指定された*祖先に対して固定されることです。親要素にpositionルールを追加し忘れると（通常は`position: relative;`で設定します）、ブラウザは上位の要素を探し続け、最終的に`body`タグを基準にします。

# --instructions--

`#searchbar`要素を`section`の親の右上に固定するために、`position`を`absolute`に設定してください。`top`と`right`のオフセットをそれぞれ50ピクセルにしてください。

# --hints--

`#searchbar`要素は`position`を`absolute`に設定する必要があります。

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.position, 'absolute');
```

コードでは`top`のCSSオフセットを`#searchbar`要素に50ピクセル使うべきです。

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.top, '50px');
```

コードでは`right`のCSSオフセットを`#searchbar`要素に50ピクセル使うべきです。

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.right, '50px');
```

# --seed--

## --seed-contents--

```html
<style>
  #searchbar {



  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```

# --solutions--

```html
<style>
  #searchbar {
    position: absolute;
    top: 50px;
    right: 50px;
  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```
