---
id: 587d78a3367417b2b2512acf
title: z-index プロパティで重なった要素の位置を変更する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM94aHk'
forumTopicId: 301046
dashedName: change-the-position-of-overlapping-elements-with-the-z-index-property
---

# --description--

要素が重なるように配置されている場合（つまり `position: absolute | relative | fixed | sticky` を使っている場合）、HTMLマークアップで後に書かれた要素がデフォルトでは他の要素の上に表示されます。しかし、`z-index` プロパティを使うと、要素が重なって表示される順序を指定できます。この値は整数（小数ではなく整数）でなければならず、`z-index` プロパティの値が大きい要素ほど、値が小さい要素よりも上に表示されます。

# --instructions--

クラス名が `z-index` の要素（赤い長方形）に `first` プロパティを追加し、値を 2 に設定して、もう一方の要素（青い長方形）より上に表示されるようにしてください。

# --hints--

クラス `first` の要素は `z-index` の値を 2 にする必要があります。

```js
const firstElement = document.querySelector('.first');
const firstStyle = window.getComputedStyle(firstElement); 
assert.equal(firstStyle?.zIndex, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;

  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>
```

# --solutions--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;
    z-index: 2;
  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>
<div class="first"></div>
<div class="second"></div>
```
