---
id: 587d781b367417b2b2512abc
title: テキストの background-color プロパティを調整する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDqwA6'
forumTopicId: 301032
dashedName: adjust-the-background-color-property-of-text
---

# --description--

前景のテキストを読みやすくするために全体の背景やテキストの色を調整する代わりに、強調したいテキストを含む要素に `background-color` を追加できます。このチャレンジでは `rgba()` を `hex` コードや通常の `rgb()` の代わりに使います。

<blockquote>rgba は以下の意味です:<br>  r = 赤<br>  g = 緑<br>  b = 青<br>  a = アルファ／不透明度のレベル</blockquote>

RGB の値は 0 から 255 の範囲で指定できます。アルファ値は 1（完全に不透明、つまりソリッドカラー）から 0（完全に透明、つまりクリア）までの範囲です。`rgba()` はこの場合に最適で、不透明度を調整できます。つまり背景を完全に隠す必要がありません。

このチャレンジでは `background-color: rgba(45, 45, 45, 0.1)` を使います。不透明度が 0.1 と低いため、ほぼ透明なダークグレーの色が生成されます。

# --instructions--

テキストをより目立たせるために、`background-color` 要素の `h4` を指定された `rgba()` の値に調整してください。

また `h4` については、`height` プロパティを削除し、`padding` を 10px に設定してください。

# --hints--

コードは `background-color` 要素に `h4` プロパティを追加し、`rgba(45, 45, 45, 0.1)` に設定する必要があります。

```js
assert(
  /(background-color|background):rgba\(45,45,45,0?\.1\)(;?}|;)/gi.test(
    code.replace(/\s/g, '')
  )
);
```

コードは `padding` 要素に `h4` プロパティを追加し、10ピクセルに設定する必要があります。

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.equal(h4style?.paddingTop, "10px");
assert.equal(h4style?.paddingRight, "10px");
assert.equal(h4style?.paddingBottom, "10px");
assert.equal(h4style?.paddingLeft, "10px");
```

`height` 要素の `h4` プロパティは削除してください。

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.notEqual(h4style?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;


  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    padding: 10px;
    background-color: rgba(45, 45, 45, 0.1);
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
