---
id: 587d78a9367417b2b2512ae9
title: ベジェ曲線を使ってグラフィックを動かす
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bnRCK'
forumTopicId: 301071
dashedName: use-a-bezier-curve-to-move-a-graphic
---

# --description--

前のチャレンジでは、アニメーションの最初は速く、最後にゆっくりになる変化を表す`ease-out`キーワードについて説明しました。右側では、青い要素に対する`ease-out`キーワードと赤い要素に対する`linear`キーワードの違いが示されています。`ease-out`キーワードに似たアニメーションの進行は、カスタムの3次ベジェ曲線関数を使って実現できます。

一般的に、`p1`と`p2`のアンカー点を変えることで異なるベジェ曲線が作られ、アニメーションの時間経過に沿った進行を制御します。以下は、`ease-out`スタイルを模倣する値を使ったベジェ曲線の例です。

```css
animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
```

すべての`cubic-bezier`関数は`p0`で(0, 0)から始まり、`p3`で(1, 1)に終わることを覚えておいてください。この例では、曲線はY軸方向に速く動きます（0から始まり、`p1`のy値0を経て、`p2`のy値1へ）一方、X軸方向は（0から始まり、`p1`で0、`p2`で0.58まで）動きます。その結果、アニメーション要素の変化は、その区間のアニメーション時間よりも速く進みます。曲線の終わりに近づくと、xとyの値の変化の関係が逆転し、y値は1から1へ（変化なし）、x値は0.58から1へ動くため、アニメーションの変化はアニメーションの時間に比べて遅く進みます。

# --instructions--

このベジェ曲線の効果を実際に見るには、idが`animation-timing-function`の要素の`red`を、x1、y1、x2、y2をそれぞれ`cubic-bezier`に設定した`0, 0, 0.58, 1`関数に変更してください。これにより両方の要素が同じようにアニメーションを進行します。

# --hints--

idが`animation-timing-function`の要素の`red`プロパティの値は、x1、y1、x2、y2をそれぞれ0、0、0.58、1に設定した`cubic-bezier`関数である必要があります。

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.equal(
  redStyle?.animationTimingFunction, 'cubic-bezier(0, 0, 0.58, 1)'
);
```

idが`red`の要素は、もはや`animation-timing-function`の`linear`プロパティを持っていないはずです。

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.notEqual(redStyle?.animationTimingFunction, 'linear');
```

idが`animation-timing-function`の要素の`blue`プロパティの値は変わらないはずです。

```js
const blueElement = document.querySelector('#blue');
const blueStyle = window.getComputedStyle( blueElement);
const blueBallAnimation = __helpers.removeWhiteSpace(
  blueStyle?.animationTimingFunction
);
assert.isTrue(
  blueBallAnimation == 'ease-out' ||
    blueBallAnimation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```

# --solutions--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```
