---
id: 587d78a6367417b2b2512ade
title: CSSとHTMLを使ってより複雑な形を作成する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpz4fr'
forumTopicId: 301050
dashedName: create-a-more-complex-shape-using-css-and-html
---

# --description--

世界で最も人気のある形の一つがハート形で、このチャレンジでは純粋なCSSを使ってそれを作成します。しかしまず、`::before`と`::after`の疑似要素を理解する必要があります。`::before`は選択された要素の最初の子要素となる疑似要素を作成し、`::after`は選択された要素の最後の子要素となる疑似要素を作成します。以下の例では、`::before`疑似要素を使って`heart`クラスの要素に長方形を追加しています。

```css
.heart::before {
  content: "";
  background-color: yellow;
  border-radius: 25%;
  position: absolute;
  height: 50px;
  width: 70px;
  top: -50px;
  left: 5px;
}
```

`::before`と`::after`の疑似要素が正しく機能するには、`content`プロパティが定義されている必要があります。このプロパティは通常、写真や文字列のようなものを選択された要素に追加するために使われます。`::before`と`::after`の疑似要素を使って形を作る場合でも、`content`プロパティは必要ですが、空の文字列に設定されます。上の例では、`heart`クラスの要素に`::before`疑似要素があり、高さと幅がそれぞれ`50px`と`70px`の黄色い長方形を生成しています。この長方形は25%の`border-radius`で角が丸くなっており、要素の左から`5px`、上から`50px`の位置に絶対配置されています。

# --instructions--

画面上の要素をハート形に変形してください。`.heart::after`セレクターでは、`background-color`を`pink`に、`border-radius`を50%に変更してください。

次に、`heart`クラス（`heart`だけ）をターゲットにして、`transform`プロパティを埋めてください。-45度の`rotate()`関数を使います。

最後に、`.heart::before`セレクターで、その`content`プロパティを空の文字列に設定してください。

# --hints--

`background-color`セレクターの`.heart::after`プロパティは`pink`であるべきです。

```js
const heartAfter = code.match(/\.heart::after\s*{[\s\S]+?[^\}]}/g)?.[0];
assert.match(heartAfter, /({|;)\s*background-color\s*:\s*pink\s*(;|})/g);
```

`border-radius`セレクターの`.heart::after`は50%であるべきです。

```js
assert.lengthOf(code.match(/border-radius\s*?:\s*?50%/gi),2);
```

`transform`クラスの`heart`プロパティは、-45度に設定された`rotate()`関数を使うべきです。

```js
assert.match(code,/transform\s*?:\s*?rotate\(\s*?-45deg\s*?\)/gi);
```

`content`セレクターの`.heart::before`は空の文字列であるべきです。

```js
assert.match(code,/\.heart::before\s*?{\s*?content\s*?:\s*?("|')\1\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: ;
  }
  .heart::after {
    background-color: blue;
    content: "";
    border-radius: 25%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: ;
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
  }
  .heart::after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: "";
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```
