---
id: 587d781c367417b2b2512ac3
title: 複数の見出し要素のフォントウェイトを設定する
challengeType: 0
videoUrl: 'https://scrimba.com/c/crVWRHq'
forumTopicId: 301069
dashedName: set-the-font-weight-for-multiple-heading-elements
---

# --description--

前のチャレンジで各見出しタグの`font-size`を設定しましたが、ここでは`font-weight`を調整します。

`font-weight`プロパティは、テキストのセクション内の文字の太さを設定します。

# --instructions--

<ul><li><code>h1</code>タグの<code>font-weight</code>を800に設定してください。</li><li><code>h2</code>タグの<code>font-weight</code>を600に設定してください。</li><li><code>h3</code>タグの<code>font-weight</code>を500に設定してください。</li><li><code>h4</code>タグの<code>font-weight</code>を400に設定してください。</li><li><code>h5</code>タグの<code>font-weight</code>を300に設定してください。</li><li><code>h6</code>タグの<code>font-weight</code>を200に設定してください。</li></ul>

# --hints--

コードは`font-weight`タグの`h1`プロパティを800に設定する必要があります。

```js
const h1Element =document.querySelector('h1')
const h1Style = window.getComputedStyle(h1Element);
assert.equal(h1Style?.fontWeight, '800');
```

コードは`font-weight`タグの`h2`プロパティを600に設定する必要があります。

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.fontWeight, '600');
```

コードは`font-weight`タグの`h3`プロパティを500に設定する必要があります。

```js
const h3Element =document.querySelector('h3')
const h3Style = window.getComputedStyle(h3Element);
assert.equal(h3Style?.fontWeight, '500');
```

コードは`font-weight`タグの`h4`プロパティを400に設定する必要があります。

```js
const h4Element =document.querySelector('h4')
const h4Style = window.getComputedStyle(h4Element);
assert.equal(h4Style?.fontWeight, '400');
```

コードは`font-weight`タグの`h5`プロパティを300に設定する必要があります。

```js
const h5Element =document.querySelector('h5')
const h5Style = window.getComputedStyle(h5Element);
assert.equal(h5Style?.fontWeight, '300');
```

コードは`font-weight`タグの`h6`プロパティを200に設定する必要があります。

```js
const h6Element =document.querySelector('h6')
const h6Style = window.getComputedStyle(h6Element);
assert.equal(h6Style?.fontWeight, '200');
```

# --seed--

## --seed-contents--

```html
<style>
  h1 {
    font-size: 68px;

  }
  h2 {
    font-size: 52px;

  }
  h3 {
    font-size: 40px;

  }
  h4 {
    font-size: 32px;

  }
  h5 {
    font-size: 21px;

  }
  h6 {
    font-size: 14px;

  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
    font-weight: 800;
  }
  h2 {
    font-size: 52px;
    font-weight: 600;
  }
  h3 {
    font-size: 40px;
    font-weight: 500;
  }
  h4 {
    font-size: 32px;
    font-weight: 400;
  }
  h5 {
    font-size: 21px;
    font-weight: 300;
  }
  h6 {
    font-size: 14px;
    font-weight: 200;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
