---
id: 587d78a3367417b2b2512ace
title: floatプロパティで要素を左または右に寄せる
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDqu2'
forumTopicId: 301066
dashedName: push-elements-left-or-right-with-the-float-property
---

# --description--

次の配置ツールは実際には`position`を使わず、要素の`float`プロパティを設定します。浮動要素は文書の通常のフローから外され、親要素の`left`または`right`に寄せられます。これは、浮動要素が必要とする水平スペースを指定するために`width`プロパティと一緒に使われることが多いです。

# --instructions--

与えられたマークアップは、`section`と`aside`の要素が隣り合う2カラムレイアウトとしてうまく機能します。`#left`の項目に`float`を`left`に設定し、`#right`の項目に`float`を`right`に設定してください。

# --hints--

idが`left`の要素には`float`を`left`に設定してください。

```js
const leftElement = document.querySelector('#left');
const leftStyle = window.getComputedStyle(leftElement);
assert.equal(leftStyle?.float, 'left');
```

idが`right`の要素には`float`を`right`に設定してください。

```js
const rightElement = document.querySelector('#right');
const rightStyle = window.getComputedStyle(rightElement);
assert.equal(rightStyle?.float, 'right');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
    #left {

      width: 50%;
    }
    #right {

      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```

# --solutions--

```html
<head>
  <style>
    #left {
      float: left;
      width: 50%;
    }
    #right {
      float: right;
      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```
