---
id: 587d78a3367417b2b2512ad0
title: マージンプロパティを使って要素を水平方向に中央揃えする
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJqU4'
forumTopicId: 301043
dashedName: center-an-element-horizontally-using-the-margin-property
---

# --description--

もう一つの配置の手法は、ブロック要素を水平方向に中央揃えすることです。その方法の一つは、`margin`をautoに設定することです。

この手法は画像にも使えます。画像はデフォルトではインライン要素ですが、`display`プロパティを`block`に設定するとブロック要素に変更できます。

# --instructions--

`div`をページの中央に配置するために、`margin`プロパティを`auto`の値で追加してください。

# --hints--

`div`には`margin`を`auto`に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('div')?.margin ,'auto');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;

  }
</style>
<div></div>
```

# --solutions--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;
    margin: auto;
  }
</style>
<div></div>
```
