---
id: 587d781d367417b2b2512ac8
title: アンカーのホバー状態を調整する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakRGcm'
forumTopicId: 301035
dashedName: adjust-the-hover-state-of-an-anchor-tag
---

# --description--

このチャレンジでは疑似クラスの使い方に触れます。疑似クラスとは、セレクターに追加して要素の特定の状態を選択するためのキーワードです。

例えば、アンカーのホバー状態のスタイルを`:hover`疑似クラスセレクターを使って変更できます。以下はアンカーの`color`をホバー時に赤に変えるCSSです。

```css
a:hover {
  color: red;
}
```

# --instructions--

コードエディターにはすべての`a`タグを黒くするCSSルールがあります。ユーザーが`a`タグにホバーしたときに`color`が青くなるようにルールを追加してください。

# --hints--

アンカーの`color`は黒のままにして、`:hover`状態のCSSルールだけを追加してください。

```js
const anchorElement = document.querySelector("a"); 
const anchorStyle = window.getComputedStyle(anchorElement);
assert.equal(anchorStyle?.color, 'rgb(0, 0, 0)');
```

アンカーはホバー時に`color`が青くなるはずです。

```js
assert.match(code,
    /a:hover\s*?{\s*?color:\s*?(blue|rgba\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?,\s*?1\s*?\)|#00F|rgb\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?\))\s*?;\s*?}/gi
  );
```

# --seed--

## --seed-contents--

```html
<style>
  a {
    color: #000;
  }



</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```

# --solutions--

```html
<style>
  a {
    color: #000;
  }
  a:hover {
    color: rgba(0,0,255,1);
  }
</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```
