---
id: 587d78a4367417b2b2512ad5
title: 色のトーンを調整する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDJvT7'
forumTopicId: 301038
dashedName: adjust-the-tone-of-a-color
---

# --description--

CSSの`hsl()`オプションを使うと、色のトーンを簡単に調整できます。純粋な色相に白を混ぜるとその色のティントができ、黒を加えるとシェードになります。あるいは、グレーを加えるか、ティントとシェードの両方を使うことでトーンが作られます。`hsl()`の「s」と「l」はそれぞれ彩度と明度を表すことを思い出してください。彩度のパーセントはグレーの量を変え、明度のパーセントは色に含まれる白または黒の量を決めます。これは、好きな基本の色相があるけれど、そのバリエーションをいくつか作りたい場合に便利です。

# --instructions--

すべての要素にはデフォルトで`background-color`が`transparent`に設定されています。現在、`nav`要素は`cyan`の背景に見えています。なぜなら、その背後の要素に`background-color`が`cyan`に設定されているからです。`background-color`要素に`nav`を追加して、同じ`cyan`の色相を使いながら、`80%`の彩度と`25%`の明度の値を変えてトーンとシェードを調整してください。

# --hints--

`nav`要素には、`background-color`プロパティを使って調整したシアンのトーンの`hsl()`を設定してください。

```js
// Computed style of hsl(180, 80%, 25%) results in rgb(13,115,115)
assert.equal(
  new __helpers.CSSHelp(document).getStyle('nav').getPropVal('background-color', true), 
  'rgb(13,115,115)'
)
```

# --seed--

## --seed-contents--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {

  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>

<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```

# --solutions--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {
    background-color: hsl(180, 80%, 25%);
  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```
