---
id: 587d78a6367417b2b2512adc
title: CSSのtransformプロパティのskewYを使って要素をY軸に沿って傾ける
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZ2uB'
forumTopicId: 301075
dashedName: use-the-css-transform-property-skewy-to-skew-an-element-along-the-y-axis
---

# --description--

`skewX()`関数が選択した要素を指定した角度でX軸に沿って傾けるのと同様に、`skewY()`プロパティは要素をY（垂直）軸に沿って傾けます。

# --instructions--

`top`プロパティを使って、idが`transform`の要素をY軸に沿って-10度傾けてください。

# --hints--

idが`top`の要素はY軸に沿って-10度傾ける必要があります。

```js
assert.match(code, /#top\s*?{\s*?.*?\s*?transform:\s*?skewY\(-10deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;

  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;
    transform: skewY(-10deg);
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
