---
id: 587d78a5367417b2b2512ad8
title: 背景画像としてさりげないパターンを追加してテクスチャを作成する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQdwJC8'
forumTopicId: 301052
dashedName: create-texture-by-adding-a-subtle-pattern-as-a-background-image
---

# --description--

背景にテクスチャやアクセントを加えてより目立たせる方法の一つは、さりげないパターンを追加することです。重要なのはバランスで、背景があまりにも目立ちすぎて前景の邪魔にならないようにすることです。`background` プロパティは、選んだテクスチャやパターンの画像にリンクするために `url()` 関数をサポートしています。リンクのアドレスは丸括弧の中に引用符で囲んで記述します。

# --instructions--

`https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png` の URL を使って、`background` セレクターでページ全体の `body` を設定してください。

# --hints--

`body` 要素には、指定されたリンクを使った `background` が設定された `url()` プロパティが必要です。

```js
assert.match(code,/background(-image)?:\s*?url\(\s*("|'|)https:\/\/cdn-media-1\.freecodecamp\.org\/imgr\/MJAkxbh\.png\2\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {

  }
</style>
```

# --solutions--

```html
<style>
  body {
    background: url("https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png");
  }
</style>
```
