---
id: 587d78a5367417b2b2512ad6
title: CSSで徐々に変化する線形グラデーションを作成する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4dpt9'
forumTopicId: 301047
dashedName: create-a-gradual-css-linear-gradient
---

# --description--

HTML要素に色を適用するのは単一の平坦な色だけに限りません。CSSでは、グラデーションとも呼ばれる色の移り変わりを要素に使うことができます。これは`background`プロパティの`linear-gradient()`関数でアクセスできます。一般的な構文は以下の通りです。

```css
background: linear-gradient(gradient_direction, color 1, color 2, color 3, ...);
```

最初の引数は色の移り変わりが始まる方向を指定します。これは度数で表現でき、`90deg`は水平方向のグラデーション（左から右へ）、`45deg`は斜めのグラデーション（左下から右上へ）を作ります。続く引数はグラデーションに使う色の順序を指定します。

例：

```css
background: linear-gradient(90deg, red, yellow, rgb(204, 204, 255));
```

# --instructions--

`linear-gradient()`要素の`div`に`background`を使い、方向を35度に設定して`#CCFFFF`から`#FFCCCC`へ色を変化させてください。

# --hints--

`div`要素は指定された方向と色で`linear-gradient` `background`を持つ必要があります。

```js
const divElement = document.querySelector('div');
const divStyle = window.getComputedStyle(divElement); 
assert.match(divStyle?.background, /linear-gradient\(35deg, rgb\(204, 255, 255\), rgb\(255, 204, 204\)\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;

  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;
    background: linear-gradient(35deg, #CCFFFF, #FFCCCC);
  }
</style>
<div></div>
```
