---
id: 587d78a7367417b2b2512ae0
title: ボタンのホバー状態をCSSアニメーションで変更する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4vZAa'
forumTopicId: 301073
dashedName: use-css-animation-to-change-the-hover-state-of-a-button
---

# --description--

CSS `@keyframes`を使って、ボタンのホバー状態の色を変更できます。

ホバー時に画像の幅を変える例を示します。

```html
<style>
  img {
    width: 30px;
  }
  img:hover {
    animation-name: width;
    animation-duration: 500ms;
  }

  @keyframes width {
    100% {
      width: 40px;
    }
  }
</style>

<img src="https://cdn.freecodecamp.org/curriculum/applied-visual-design/google-logo.png" alt="Google's Logo" />
```

# --instructions--

`ms`はミリ秒を表し、1000msは1秒に相当することに注意してください。

CSS `@keyframes`を使って、ユーザーがホバーしたときに`background-color`要素の`button`を`#4791d0`に変えるようにしてください。`@keyframes`ルールには`100%`のエントリーだけが必要です。

# --hints--

@keyframesルールは`animation-name`のbackground-colorを使うべきです。

```js
assert.match(code, /@keyframes\s+?background-color\s*?{/g);
```

`@keyframes`の下に、100%の時点で`background-color`を`#4791d0`に変えるルールが1つあるべきです。

```js
assert.match(code, /100%\s*?{\s*?background-color:\s*?#4791d0;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }


</style>

<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }

  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
