---
id: 587d781c367417b2b2512ac2
title: 複数の見出し要素のフォントサイズを設定する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpQNT3'
forumTopicId: 301067
dashedName: set-the-font-size-for-multiple-heading-elements
---

# --description--

`font-size` プロパティは、特定の要素内のテキストの大きさを指定するために使います。このルールは複数の要素に適用して、ページ上のテキストの視覚的一貫性を作ることができます。このチャレンジでは、`h1` から `h6` までのすべてのタグの値を設定して、見出しのサイズのバランスを取ります。

# --instructions--

`style` タグ内で、次のタグの `font-size` を設定してください：

- `h1` タグを 68px に。
- `h2` タグを 52px に。
- `h3` タグを 40px に。
- `h4` タグを 32px に。
- `h5` タグを 21px に。
- `h6` タグを 14px に。

# --hints--

`font-size` タグの `h1` プロパティを 68 ピクセルに設定してください。

```js
 const fontSizeOfh1 = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh1 ,'68px');
```

`font-size` タグの `h2` プロパティを 52 ピクセルに設定してください。

```js
 const fontSizeOfh2 = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh2 ,'52px');
```

`font-size` タグの `h3` プロパティを 40 ピクセルに設定してください。

```js
 const fontSizeOfh3 = new __helpers.CSSHelp(document).getStyle('h3')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh3 ,'40px');
```

`font-size` タグの `h4` プロパティを 32 ピクセルに設定してください。

```js
 const fontSizeOfh4 = new __helpers.CSSHelp(document).getStyle('h4')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh4 , '32px');
```

`font-size` タグの `h5` プロパティを 21 ピクセルに設定してください。

```js
 const fontSizeOfh5 = new __helpers.CSSHelp(document).getStyle('h5')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh5 ,'21px');
```

`font-size` タグの `h6` プロパティを 14 ピクセルに設定してください。

```js
 const fontSizeOfh6 = new __helpers.CSSHelp(document).getStyle('h6')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh6 , '14px');
```

# --seed--

## --seed-contents--

```html
<style>






</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
  }
  h2 {
    font-size: 52px;
  }
  h3 {
    font-size: 40px;
  }
  h4 {
    font-size: 32px;
  }
  h5 {
    font-size: 21px;
  }
  h6 {
    font-size: 14px;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
