---
id: 587d78a4367417b2b2512ad4
title: 色の色相を調整する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp38TZ'
forumTopicId: 301036
dashedName: adjust-the-hue-of-a-color
---

# --description--

色には色相、彩度、明度などの特徴があります。CSS3では、これらの特徴を直接指定して色を選ぶ方法として`hsl()`関数が導入されました。

**色相**は一般的に「色」として考えられるものです。左に赤、中央に緑、右に青がある色のスペクトルを思い浮かべると、色相はこの線上のどこに色が位置するかを示します。`hsl()`では、色相はスペクトルの代わりにカラーホイールの概念を使い、円の上の色の角度を0から360の値で表します。

**彩度**は色に含まれる灰色の量です。彩度が100%の色は灰色が全くなく、彩度がほぼ0%の色はほとんど灰色です。これはパーセンテージで表され、100%が完全に彩度の高い色です。

**明度**は色に含まれる白または黒の量です。0%（黒）から100%（白）までのパーセンテージで表され、50%が通常の色の明るさです。

以下は、彩度が完全で明度が通常の色を`hsl()`関数で指定した例です。

<table><thead><tr><th>色</th><th>HSL</th></tr></thead><tbody><tr><td>赤</td><td>hsl(0, 100%, 50%)</td></tr><tr><td>黄</td><td>hsl(60, 100%, 50%)</td></tr><tr><td>緑</td><td>hsl(120, 100%, 50%)</td></tr><tr><td>シアン</td><td>hsl(180, 100%, 50%)</td></tr><tr><td>青</td><td>hsl(240, 100%, 50%)</td></tr><tr><td>マゼンタ</td><td>hsl(300, 100%, 50%)</td></tr></tbody></table>

# --instructions--

`background-color`を`div`を使って、クラス名（`green`、`cyan`、または`blue`）に基づいてそれぞれの`hsl()`要素の色相を変更してください。3つとも彩度は完全で明度は通常にしてください。

# --hints--

緑色を宣言するには`hsl()`関数を使ってください。

```js
assert.match(code,/\.green\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

シアン色を宣言するには`hsl()`関数を使ってください。

```js
assert.match(code,/\.cyan\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

青色を宣言するには`hsl()`関数を使ってください。

```js
assert.match(code,/\.blue\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

クラス`div`の`green`要素は緑色の`background-color`を持つべきです。

```js
const greenElement = document.querySelector(".green");
const greenStyle = window.getComputedStyle(greenElement); 
assert.equal(greenStyle?.backgroundColor, 'rgb(0, 255, 0)');
```

クラス`div`の`cyan`要素はシアン色の`background-color`を持つべきです。

```js
const cyanElement = document.querySelector(".cyan");
const cyanStyle = window.getComputedStyle(cyanElement); 
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

クラス`div`の`blue`要素は青色の`background-color`を持つべきです。

```js
const blueElement = document.querySelector(".blue");
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .blue {
    background-color: #000000;
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>

<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: hsl(120, 100%, 50%);
  }

  .cyan {
    background-color:   hsl(180, 100%, 50%);
  }

  .blue {
    background-color: hsl(240, 100%, 50%);
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```
