---
id: 58a7a6ebf9a6318348e2d5aa
title: アニメーションの Fill Mode を変更する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDmcE'
forumTopicId: 301064
dashedName: modify-fill-mode-of-an-animation
---

# --description--

いいですね、でもまだ正しく動作していません。`500ms` が過ぎるとアニメーションがリセットされて、ボタンが元の色に戻ってしまうのに気づきましたか？ボタンはハイライトされたままでいてほしいです。

これは `animation-fill-mode` プロパティを `forwards` に設定することで実現できます。`animation-fill-mode` はアニメーションが終了したときに要素に適用されるスタイルを指定します。次のように設定できます。

```css
animation-fill-mode: forwards;
```

# --instructions--

ユーザーがボタンにホバーしたときにボタンがハイライトされたままになるように、`animation-fill-mode` の `button:hover` プロパティを `forwards` に設定してください。

# --hints--

`button:hover` は `animation-fill-mode` の値を持つ `forwards` プロパティを持つべきです。

```js
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-fill-mode\s*?:\s*?forwards\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-name\s*?:\s*?background-color\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-duration\s*?:\s*?500ms\s*?;[\s\S]*}/gi );
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    /* Only change code below this line */

    /* Only change code above this line */
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    animation-fill-mode: forwards;
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
