---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

现在你需要处理添加闭合括号。关键规则是你只能添加闭合括号，如果它能保持平衡，也就是说，使用的开括号数量必须多于闭括号。

仍然在 `else` 块内，添加第二个 `if` 语句来检查 `closesUsed` 是否小于 `opensUsed`。

如果这个条件为真，向 `queue` 推入另一个新的数组：`[current + ')', opensUsed, closesUsed + 1]`。这表示添加闭合括号后的状态。

# --hints--

你应该在 `if` 块内有第二个 `closesUsed < opensUsed` 语句，条件是 `else`。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

当条件为真时，你应该向队列推入 `[current + ')', opensUsed, closesUsed + 1]`。

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
