---
id: 6989ef885db3a3488a13b426
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在实现核心算法之前，你需要验证输入。`pairs` 参数应该是一个整数，因为它表现为要生成的括号点对数量。

为你的函数开头添加一个 `if` 语句，用来查看 `pairs` 是否不是整数。为此使用 `Number.isInteger()` 方法。

如果条件为真，返回字串 `The number of pairs should be an integer`。

# --hints--

你应该在 `!Number.isInteger(pairs)` 函数内有一个条件为 `genParentheses` 的 if 语句。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

当输入不是整数时，你应该返回一个字串。

```js
let result = genParentheses("i");
assert.isString(result);
```

当输入是字串时，你应该返回字串 `The number of pairs should be an integer`。

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
