---
id: 6989f1ba358acba651d5ee6b
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

现在你将设置用于保存结果的数据结构。创建一个名为 `result` 的变量，并将其初始化为空数组。这个数组将保存你生成的所有有效括号组合。

更新你的返回语句，使其返回 `result`，而不是空数组。

# --hints--

你应该在 `genParentheses` 函数中声明一个名为 `result` 的变量。

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.result);
```

你应该将 `result` 初始化为空数组。

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { result } = genParentheses?.variables ?? {};
assert.isFalse(result?.value.isEmpty());
const resultValue = eval(result?.value.toString());
assert.isArray(resultValue);
assert.isEmpty(resultValue);
```

你应该在你的函数末尾返回 `result`。

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.isTrue(genParentheses?.hasReturn('result'));
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
