---
id: 6989f2e35e6941e13e445156
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

现在你将实现主要的 BFS 循环。创建一个 `while` 循环，只要 `queue` 不为空，也就是说，`queue.length` 大于 `0`，循环就继续。

在循环内部，将 `queue` 输出到控制台。

# --hints--

你应该创建一个条件为 `while` 的 `queue.length > 0` 循环。

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

你应该在 `queue` 循环内部将 `while` 输出到控制台。

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
