---
id: 6989eb65324d97775c4965fe
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会中，你将实现一个函数，使用广度优先查找（BFS）方法生成所有有效的括号组合。例如，两个点对括号的有效组合是 `(())` 和 `()()`。

首先创建一个名为 `genParentheses` 的函数，它有一个参数 `pairs`。目前，函数应返回一个空数组。

# --hints--

你应该定义一个名为 `genParentheses` 的函数。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
```

你的 `genParentheses` 函数应有一个名为 `pairs` 的参数。

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.equal(genParentheses?.parameters.length, 1);
assert.equal(genParentheses?.parameters[0].toString(), 'pairs');
```

你的函数应返回一个空数组。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
