---
id: 6989eb65324d97775c4965fe
title: 第一步
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将实现一个函数，使用广度优先搜索（BFS）方法生成所有有效的括号组合。例如，两个点对括号的有效组合是 `(())` 和 `()()`。

首先创建一个名为 `genParentheses` 的函数，它有一个参数 `pairs`。目前，函数应返回一个空数组。

# --hints--

你应该定义一个名为 `genParentheses` 的函数。

```js
assert.isFunction(genParentheses);
```

你的 `genParentheses` 函数应有一个名为 `pairs` 的参数。

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

你的函数应返回一个空数组。

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
