---
id: 6989f10a72ad424d4ad6d514
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

接下来，你需要验证点对的数量至少为 1，因为你无法生成零个或负数点对的括号组合。

添加另一个 `if` 语句来查看 `pairs` 是否小于 `1`。如果此条件为真，返回字串 `The number of pairs should be at least 1`。

# --hints--

你应该有第二个 `if` 语句，条件是 `pairs < 1`。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

当 `The number of pairs should be at least 1` 小于 `pairs` 时，你应该返回字串 `1`。

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
