---
id: 6989f207306332c06b0add39
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

对于广度优先查找方法，你将使用一个队列来跟踪不同的状态，同时构建括号组合。每个状态将表现为一个包含三个元素的数组：

- 当前正在构建的字串
- 到目前为止使用的左括号数量
- 到目前为止使用的右括号数量

创建一个名为 `queue` 的变量，并将其初始化为包含一个数组 `['', 0, 0]` 的数组。这表现了起始状态，字串为空且括号数量为零。

# --hints--

你应该在你的 `queue` 函数中声明一个名为 `genParentheses` 的变量。

```js
assert.match(genParentheses.toString(), /queue/);
```

你应该用包含数组 `queue` 的数组初始化 `['', 0, 0]`。

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  return result;
}
```
