---
id: 6989f48c3f0292c84f8d8938
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

现在你需要识别何时构建了一个完整的括号组合。一个完整的组合长度等于点对数量的两倍（因为每个点对包含一个左括号和一个右括号）。

在你的 `while` 循环内，添加一个 `if` 语句来查看 `current.length === 2 * pairs`。当此条件为真时，将 `current` 添加到 `result` 数组中。

# --hints--

你应该在 `if` 循环内有一个条件为 `current.length === 2 * pairs` 的 `while` 语句。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

你应该在新的 `current` 语句内将 `result` 添加到 `if` 数组中。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
