---
id: 6989f3632f8b8bdf168aa321
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

在你的 `while` 循环内，使用 `queue.shift()` 来移除并获取队列中的第一个元素。这实现了 BFS 特有的先进先出（FIFO）行为。

将这个数组解构为三个变量：`current`、`opensUsed` 和 `closesUsed`。

# --hints--

你应该使用 `queue.shift()` 来获取队列中的第一个元素。

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

你应该将结果解构为三个变量：`current`、`opensUsed` 和 `closesUsed`。

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return result;
}
```
