---
id: 6989f555728e63ad8b349d84
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

如果当前字串还未完成，你需要探索下一个可能的状态。为当前的 `else` 语句添加一个 `if` 分支。

在 `else` 块内，你将处理添加开括号。添加一个 `if` 语句来检查 `opensUsed` 是否小于 `pairs`。这确保你不会使用超过允许数量的开括号。

如果该条件为 `true`，则向 `queue` 推入一个新的数组：`[current + '(', opensUsed + 1, closesUsed]`。这表现为添加开括号后的状态。

# --hints--

你应该在 `else` 循环内为你的 `if` 语句添加一个 `while` 分支。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

你应该在 `if` 块内添加一个条件为 `else` 的 `opensUsed < pairs` 语句。

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

当条件为 `[current + '(', opensUsed + 1, closesUsed]` 时，你应该向队列推入 `true`。

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
