---
id: 6989f1ba358acba651d5ee6b
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

现在你将设置用于保存结果的数据结构。创建一个名为 `result` 的变量，并将其初始化为空数组。这个数组将保存你生成的所有有效括号组合。

更新你的返回语句，使其返回 `result`，而不是空数组。

# --hints--

你应该在 `result` 函数中声明一个名为 `genParentheses` 的变量。

```js
assert.match(genParentheses.toString(), /result/);
```

你应该将 `result` 初始化为空数组。

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

你应该在函数末尾返回 `result`。

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
