---
id: 587d7db8367417b2b2512ba0
title: 英数字以外のすべてにマッチさせる
challengeType: 1
forumTopicId: 301353
dashedName: match-everything-but-letters-and-numbers
---

# --description--

英数字にマッチさせるショートカットを `[A-Za-z0-9_]` を使って `\w` で呼び出せることを学びました。検索したい自然なパターンの1つは、英数字の反対のものです。

`\w` を使って `\W` の反対を検索できます。反対のパターンは大文字を使うことに注意してください。このショートカットは `[^A-Za-z0-9_]` と同じです。

```js
let shortHand = /\W/;
let numbers = "42%";
let sentence = "Coding!";
numbers.match(shortHand);
sentence.match(shortHand);
```

最初の `match` の呼び出しは `["%"]` を返し、2回目は `["!"]` を返します。

# --instructions--

ショートハンドの文字クラス `\W` を使って、さまざまな引用文や文字列の中の英数字以外の文字の数を数えてください。

# --hints--

正規表現にはグローバルフラグを使ってください。

```js
assert(nonAlphabetRegex.global);
```

正規表現は文字列 `The five boxing wizards jump quickly.` の中で6つの英数字以外の文字を見つける必要があります。

```js
assert(
  'The five boxing wizards jump quickly.'.match(nonAlphabetRegex).length == 6
);
```

正規表現は英数字以外の文字にマッチするショートハンド文字を使う必要があります。

```js
assert(/\\W/.test(nonAlphabetRegex.source));
```

正規表現は文字列 `Pack my box with five dozen liquor jugs.` の中で8つの英数字以外の文字を見つける必要があります。

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(nonAlphabetRegex).length == 8
);
```

正規表現は文字列 `How vexingly quick daft zebras jump!` の中で6つの英数字以外の文字を見つける必要があります。

```js
assert(
  'How vexingly quick daft zebras jump!'.match(nonAlphabetRegex).length == 6
);
```

正規表現は文字列 `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.` の中で12の英数字以外の文字を見つける必要があります。

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(nonAlphabetRegex)
    .length == 12
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let nonAlphabetRegex = /change/; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards_jump quickly.";
let nonAlphabetRegex = /\W/g; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```
