---
id: 587d7db5367417b2b2512b96
title: アルファベットの文字にマッチさせる
challengeType: 1
forumTopicId: 301354
dashedName: match-letters-of-the-alphabet
---

# --description--

<dfn>文字セット</dfn>を使ってマッチさせたい文字のグループを指定できることは見ましたが、たとえばアルファベットのすべての文字のように広い範囲の文字にマッチさせたい場合は入力が多くなってしまいます。幸い、これを短く簡単にする組み込みの機能があります。

文字セットの中で、ハイフン文字を使ってマッチさせたい文字の範囲を指定できます：`-`。

たとえば、小文字の`a`から`e`までにマッチさせたい場合は`[a-e]`を使います。

```js
let catStr = "cat";
let batStr = "bat";
let matStr = "mat";
let bgRegex = /[a-e]at/;
catStr.match(bgRegex);
batStr.match(bgRegex);
matStr.match(bgRegex);
```

順番に、3つの`match`の呼び出しは`["cat"]`、`["bat"]`、`null`の値を返します。

# --instructions--

文字列`quoteSample`の中のすべての文字にマッチさせてください。

**注意**：大文字と小文字の両方にマッチさせるようにしてください。

# --hints--

あなたの正規表現`alphabetRegex`は35個の項目にマッチするはずです。

```js
assert(result.length == 35);
```

あなたの正規表現`alphabetRegex`はグローバルフラグを使うべきです。

```js
assert(alphabetRegex.flags.match(/g/).length == 1);
```

あなたの正規表現`alphabetRegex`は大文字小文字を区別しないフラグを使うべきです。

```js
assert(alphabetRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /change/; // Change this line
let result = alphabetRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /[a-z]/gi; // Change this line
let result = quoteSample.match(alphabetRegex); // Change this line
```
