---
id: 587d7db7367417b2b2512b9f
title: すべての文字と数字にマッチさせる
challengeType: 1
forumTopicId: 301346
dashedName: match-all-letters-and-numbers
---

# --description--

文字クラスを使って、`[a-z]`でアルファベットのすべての文字を検索できました。この種の文字クラスはよく使われるため、ショートカットも用意されていますが、いくつか余分な文字も含まれています。

JavaScriptでアルファベットに最も近い文字クラスは`\w`です。このショートカットは`[A-Za-z0-9_]`と同じ意味です。この文字クラスは大文字と小文字の文字に加えて数字にもマッチします。なお、この文字クラスにはアンダースコア文字（`_`）も含まれています。

```js
let longHand = /[A-Za-z0-9_]+/;
let shortHand = /\w+/;
let numbers = "42";
let varNames = "important_var";
longHand.test(numbers);
shortHand.test(numbers);
longHand.test(varNames);
shortHand.test(varNames);
```

これら4つの`test`の呼び出しはすべて`true`を返します。

これらのショートカット文字クラスは<dfn>省略文字クラス</dfn>とも呼ばれます。

# --instructions--

省略文字クラス`\w`を使って、さまざまな引用文や文字列の英数字の数を数えてください。

# --hints--

正規表現にはグローバルフラグを使ってください。

```js
assert(alphabetRegexV2.global);
```

すべての英数字にマッチするために省略文字クラス`\w`を使ってください。

```js
assert(/\\w/.test(alphabetRegexV2.source));
```

文字列`The five boxing wizards jump quickly.`には31個の英数字が含まれています。

```js
assert(
  'The five boxing wizards jump quickly.'.match(alphabetRegexV2).length === 31
);
```

文字列`Pack my box with five dozen liquor jugs.`には32個の英数字が含まれています。

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(alphabetRegexV2).length ===
    32
);
```

文字列`How vexingly quick daft zebras jump!`には30個の英数字が含まれています。

```js
assert(
  'How vexingly quick daft zebras jump!'.match(alphabetRegexV2).length === 30
);
```

文字列`123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`には36個の英数字が含まれています。

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(alphabetRegexV2)
    .length === 36
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /change/; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /\w/g; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```
