---
id: 587d7db7367417b2b2512b9e
title: 文字列の末尾パターンにマッチさせる
challengeType: 1
forumTopicId: 301352
dashedName: match-ending-string-patterns
---

# --description--

前回のチャレンジでは、キャレット文字を使って文字列の先頭のパターンを検索する方法を学びました。文字列の末尾のパターンを検索する方法もあります。

正規表現の末尾にドル記号`$`を付けることで、文字列の末尾を検索できます。

```js
let theEnding = "This is a never ending story";
let storyRegex = /story$/;
storyRegex.test(theEnding);
let noEnding = "Sometimes a story will have to end";
storyRegex.test(noEnding);
```

最初の`test`の呼び出しは`true`を返し、2番目は`false`を返します。

# --instructions--

アンカー文字（`$`）を使って、文字列`caboose`の末尾にある文字列`caboose`にマッチさせてください。

# --hints--

正規表現でドル記号`caboose`のアンカーを使って`$`を検索してください。

```js
assert(lastRegex.source == 'caboose$');
```

正規表現にフラグは使わないでください。

```js
assert(lastRegex.flags == '');
```

文字列`caboose`の末尾に`The last car on a train is the caboose`をマッチさせてください。

```js
lastRegex.lastIndex = 0;
assert(lastRegex.test('The last car on a train is the caboose'));
```

# --seed--

## --seed-contents--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /change/; // Change this line
let result = lastRegex.test(caboose);
```

# --solutions--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /caboose$/; // Change this line
let result = lastRegex.test(caboose);
```
