---
id: 587d7dbb367417b2b2512bac
title: 文字列の先頭と末尾の空白を削除する
challengeType: 1
forumTopicId: 301362
dashedName: remove-whitespace-from-start-and-end
---

# --description--

文字列の前後に不要な空白文字が含まれている場合があります。文字列を処理するときは、通常、先頭と末尾の空白を削除します。

# --instructions--

正規表現を作成し、適切な文字列メソッドを使って文字列の先頭と末尾の空白を削除してください。

**注意:** `String.prototype.trim()` メソッドでも対応できますが、このチャレンジでは正規表現を使って解決してください。

# --hints--

`result` は文字列 `Hello, World!` と等しくなるはずです。

```js
assert(result === 'Hello, World!');
```

解答では `String.prototype.trim()` メソッドを使わないでください。

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?trim/));
```

`result` 変数には直接文字列を代入しないでください。

```js
assert(!__helpers.removeJSComments(code).match(/result\s*=\s*["'`].*?["'`]/));
```

`hello` 変数の値は変更しないでください。

```js
assert(hello === '   Hello, World!  ');
```

# --seed--

## --seed-contents--

```js
let hello = "   Hello, World!  ";
let wsRegex = /change/; // Change this line
let result = hello; // Change this line
```

# --solutions--

```js
let hello = "   Hello, World!  ";
let wsRegex = /^(\s+)(.+[^\s])(\s+)$/;
let result = hello.replace(wsRegex, '$2');
```
