---
id: 587d7dba367417b2b2512ba8
title: すべてまたはなしのチェック
challengeType: 1
forumTopicId: 301338
dashedName: check-for-all-or-none
---

# --description--

検索したいパターンの一部が存在する場合もあれば、存在しない場合もあります。しかし、それらを確認することが重要な場合があります。

要素の存在があるかもしれないことを、疑問符 `?` で指定できます。これは直前の要素が0個か1個であることをチェックします。この記号は、直前の要素が任意であることを示していると考えられます。

たとえば、アメリカ英語とイギリス英語にはわずかな違いがあり、疑問符を使って両方の綴りにマッチさせることができます。

```js
let american = "color";
let british = "colour";
let rainbowRegex= /colou?r/;
rainbowRegex.test(american);
rainbowRegex.test(british);
```

`test` メソッドの両方の使い方は `true` を返します。

# --instructions--

正規表現 `favRegex` を変更して、アメリカ英語（`favorite`）とイギリス英語（`favourite`）の両方の単語にマッチさせてください。

# --hints--

正規表現にはオプション記号 `?` を使う必要があります。

```js
favRegex.lastIndex = 0;
assert(favRegex.source.match(/\?/).length > 0);
```

正規表現は文字列 `favorite` にマッチする必要があります。

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favorite'));
```

正規表現は文字列 `favourite` にマッチする必要があります。

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favourite'));
```

正規表現は文字列 `fav` にはマッチしないようにしてください。

```js
favRegex.lastIndex = 0;
assert(!favRegex.test('fav'));
```

# --seed--

## --seed-contents--

```js
let favWord = "favorite";
let favRegex = /change/; // Change this line
let result = favRegex.test(favWord);
```

# --solutions--

```js
let favWord = "favorite";
let favRegex = /favou?r/;
let result = favRegex.test(favWord);
```
