---
id: 587d7db5367417b2b2512b97
title: アルファベットの数字と文字の範囲をマッチさせる
challengeType: 1
forumTopicId: 301356
dashedName: match-numbers-and-letters-of-the-alphabet
---

# --description--

ハイフン (`-`) を使って文字の範囲をマッチさせる方法は、文字だけに限りません。数字の範囲をマッチさせることにも使えます。

例えば、`/[0-5]/` は `0` から `5` までの数字のいずれかにマッチし、`0` と `5` も含まれます。

また、文字の範囲と数字の範囲を1つの文字セットに組み合わせることも可能です。

```js
let jennyStr = "Jenny8675309";
let myRegex = /[a-z0-9]/ig;
jennyStr.match(myRegex);
```

# --instructions--

`h` から `s` までの文字の範囲と、`2` から `6` までの数字の範囲にマッチする単一の正規表現を作成してください。正規表現には適切なフラグを含めることを忘れないでください。

# --hints--

あなたの正規表現 `myRegex` は17個の項目にマッチする必要があります。

```js
assert(result.length == 17);
```

あなたの正規表現 `myRegex` はグローバルフラグを使う必要があります。

```js
assert(myRegex.flags.match(/g/).length == 1);
```

あなたの正規表現 `myRegex` は大文字小文字を区別しないフラグを使う必要があります。

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /[h-s2-6]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
