---
id: 587d7db7367417b2b2512b9d
title: 文字列の先頭パターンにマッチさせる
challengeType: 1
forumTopicId: 301349
dashedName: match-beginning-string-patterns
---

# --description--

これまでのチャレンジでは、正規表現を使って複数のマッチを探す方法を学びました。正規表現は文字列の特定の位置にあるパターンを検索するためにも使われます。

以前のチャレンジでは、キャレット文字（`^`）を文字セットの中で使い、`[^thingsThatWillNotBeMatched]`のような否定文字セットを作りました。文字セットの外では、キャレットは文字列の先頭にあるパターンを検索するために使われます。

```js
let firstString = "Ricky is first and can be found.";
let firstRegex = /^Ricky/;
firstRegex.test(firstString);
let notFirst = "You can't find Ricky now.";
firstRegex.test(notFirst);
```

最初の`test`の呼び出しは`true`を返し、2つ目は`false`を返します。

# --instructions--

正規表現でキャレット文字を使い、文字列`Cal`の先頭にだけ`rickyAndCal`を見つけてください。

# --hints--

正規表現は大文字の`Cal`という文字列を検索する必要があります。

```js
assert(calRegex.source == '^Cal');
```

正規表現にフラグは使わないでください。

```js
assert(calRegex.flags == '');
```

正規表現は文字列の先頭にある`Cal`にマッチする必要があります。

```js
calRegex.lastIndex = 0;
assert(calRegex.test('Cal and Ricky both like racing.'));
```

正規表現は文字列の途中にある`Cal`にはマッチしないようにしてください。

```js
calRegex.lastIndex = 0;
assert(!calRegex.test('Ricky and Cal both like racing.'));
```

# --seed--

## --seed-contents--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /change/; // Change this line
let result = calRegex.test(rickyAndCal);
```

# --solutions--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /^Cal/; // Change this line
let result = calRegex.test(rickyAndCal);
```
