---
id: 587d7db3367417b2b2512b8f
title: リテラル文字列の一致
challengeType: 1
forumTopicId: 301355
dashedName: match-literal-strings
---

# --description--

前のチャレンジでは、正規表現 `Hello` を使って単語 `/Hello/` を検索しました。その正規表現は文字列 `Hello` のリテラル一致を検索していました。ここに文字列 `Kevin` のリテラル一致を検索する別の例があります：

```js
let testStr = "Hello, my name is Kevin.";
let testRegex = /Kevin/;
testRegex.test(testStr);
```

この `test` 呼び出しは `true` を返します。

他の形式の `Kevin` は一致しません。例えば、正規表現 `/Kevin/` は `kevin` や `KEVIN` には一致しません。

```js
let wrongRegex = /kevin/;
wrongRegex.test(testStr);
```

この `test` 呼び出しは `false` を返します。

将来のチャレンジで他の形式の一致方法も示します。

# --instructions--

文字列 `waldoRegex` の中で `"Waldo"` をリテラル一致で見つけるように、正規表現 `waldoIsHiding` を完成させてください。

# --hints--

正規表現 `waldoRegex` は文字列 `Waldo` を見つけるべきです。

```js
waldoRegex.lastIndex = 0;
assert(waldoRegex.test(waldoIsHiding));
```

正規表現 `waldoRegex` は他のものを検索してはいけません。

```js
waldoRegex.lastIndex = 0;
assert(!waldoRegex.test('Somewhere is hiding in this text.'));
```

正規表現でリテラル文字列の一致を行うべきです。

```js
assert(!/\/.*\/i/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /search/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```

# --solutions--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /Waldo/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```
