---
id: 587d7db9367417b2b2512ba7
title: 正確なマッチ数を指定する
challengeType: 1
forumTopicId: 301365
dashedName: specify-exact-number-of-matches
---

# --description--

量指定子を使って中括弧でパターンの最小数と最大数を指定できます。時には特定の数だけマッチさせたい場合があります。

特定の数のパターンを指定するには、中括弧の間にその数だけを書けばいいです。

例えば、単語`hah`が文字`a`を`3`回だけ含む場合にマッチさせたいなら、正規表現は`/ha{3}h/`となります。

```js
let A4 = "haaaah";
let A3 = "haaah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleHA = /ha{3}h/;
multipleHA.test(A4);
multipleHA.test(A3);
multipleHA.test(A100);
```

順に、3つの`test`の呼び出しは`false`、`true`、`false`を返します。

# --instructions--

正規表現`timRegex`を変更して、単語`Timber`が文字`m`を4つだけ持つ場合にのみマッチするようにしてください。

# --hints--

正規表現には中括弧を使う必要があります。

```js
assert(timRegex.source.match(/{.*?}/).length > 0);
```

正規表現は文字列`Timber`にはマッチしてはいけません。

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timber'));
```

正規表現は文字列`Timmber`にはマッチしてはいけません。

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmber'));
```

正規表現は文字列`Timmmber`にはマッチしてはいけません。

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmmber'));
```

正規表現は文字列`Timmmmber`にはマッチする必要があります。

```js
timRegex.lastIndex = 0;
assert(timRegex.test('Timmmmber'));
```

正規表現は文字列`Timber`（30個の`m`を含む）にはマッチしてはいけません。

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Ti' + 'm'.repeat(30) + 'ber'));
```

# --seed--

## --seed-contents--

```js
let timStr = "Timmmmber";
let timRegex = /change/; // Change this line
let result = timRegex.test(timStr);
```

# --solutions--

```js
let timStr = "Timmmmber";
let timRegex = /Tim{4}ber/; // Change this line
let result = timRegex.test(timStr);
```
