---
id: 587d7db4367417b2b2512b93
title: 最初の一致以上を見つける
challengeType: 1
forumTopicId: 301342
dashedName: find-more-than-the-first-match
---

# --description--

これまで、パターンを一度だけ抽出または検索できました。

```js
let testStr = "Repeat, Repeat, Repeat";
let ourRegex = /Repeat/;
testStr.match(ourRegex);
```

ここで `match` は `["Repeat"]` を返します。

パターンを複数回検索または抽出するには、グローバル検索フラグ `g` を使えます。

```js
let repeatRegex = /Repeat/g;
testStr.match(repeatRegex);
```

ここで `match` は値 `["Repeat", "Repeat", "Repeat"]` を返します。

# --instructions--

正規表現 `starRegex` を使って、文字列 `Twinkle` から `twinkleStar` の両方の単語を見つけて抽出してください。

**注意**  
正規表現には `/search/gi` のように複数のフラグを付けることができます。

# --hints--

正規表現 `starRegex` はグローバルフラグ `g` を使うべきです。

```js
assert(starRegex.flags.match(/g/).length == 1);
```

正規表現 `starRegex` は大文字小文字を区別しないフラグ `i` を使うべきです。

```js
assert(starRegex.flags.match(/i/).length == 1);
```

マッチは単語 `Twinkle` の両方の出現にマッチするべきです。

```js
assert(
  result.sort().join() ==
    twinkleStar
      .match(/twinkle/gi)
      .sort()
      .join()
);
```

マッチ `result` は2つの要素を含むべきです。

```js
assert(result.length == 2);
```

# --seed--

## --seed-contents--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /change/; // Change this line
let result = twinkleStar; // Change this line
```

# --solutions--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /twinkle/gi;
let result = twinkleStar.match(starRegex);
```
