---
id: 587d7db6367417b2b2512b98
title: 指定しない単一文字にマッチさせる
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

これまでに、マッチさせたい文字のセットを作成しましたが、マッチさせたくない文字のセットを作成することもできます。このような文字セットは<dfn>否定文字セット</dfn>と呼ばれます。

否定文字セットを作成するには、開きブラケットの後、マッチさせたくない文字の前にキャレット文字（`^`）を置きます。

たとえば、`/[^aeiou]/gi`は母音ではないすべての文字にマッチします。`.`、`!`、`[`、`@`、`/`や空白のような文字もマッチすることに注意してください。否定された母音の文字セットは母音文字だけを除外します。

# --instructions--

数字または母音ではないすべての文字にマッチする単一の正規表現を作成してください。正規表現には適切なフラグを含めることを忘れないでください。

# --hints--

あなたの正規表現`myRegex`は9つの項目にマッチするはずです。

```js
assert(result.length == 9);
```

あなたの正規表現`myRegex`はグローバルフラグを使うべきです。

```js
assert(myRegex.flags.match(/g/).length == 1);
```

あなたの正規表現`myRegex`は大文字小文字を区別しないフラグを使うべきです。

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
