---
id: 587d7db4367417b2b2512b92
title: マッチの抽出
challengeType: 1
forumTopicId: 301340
dashedName: extract-matches
---

# --description--

これまで、文字列の中にパターンが存在するかどうかを確認するだけでした。`.match()` メソッドを使うと、実際に見つかったマッチを抽出することもできます。

`.match()` メソッドを使うには、文字列に対してこのメソッドを適用し、丸括弧の中に正規表現を渡します。

例を示します：

```js
"Hello, World!".match(/Hello/);
let ourStr = "Regular expressions";
let ourRegex = /expressions/;
ourStr.match(ourRegex);
```

ここで最初の `match` は `["Hello"]` を返し、2番目は `["expressions"]` を返します。

`.match` の構文は、これまで使ってきた `.test` メソッドの「逆」だということに注意してください：

```js
'string'.match(/regex/);
/regex/.test('string');
```

# --instructions--

`.match()` メソッドを使って文字列 `coding` を抽出してください。

# --hints--

`result` は文字列 `coding` を含んでいるはずです。

```js
assert(result.join() === 'coding');
```

正規表現 `codingRegex` は文字列 `coding` を検索する必要があります。

```js
assert(codingRegex.source === 'coding');
```

`.match()` メソッドを使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/\.match\(.*\)/));
```

# --seed--

## --seed-contents--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /change/; // Change this line
let result = extractStr; // Change this line
```

# --solutions--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /coding/; // Change this line
let result = extractStr.match(codingRegex); // Change this line
```
