---
id: 587d7db6367417b2b2512b9b
title: "遅延マッチングで文字を見つける"
challengeType: 1
forumTopicId: 301341
dashedName: find-characters-with-lazy-matching
---

# --description--

正規表現では、<dfn>貪欲</dfn>マッチは正規表現パターンに合う文字列の中で最も長い部分を見つけてマッチとして返します。これに対して、<dfn>遅延</dfn>マッチは正規表現パターンを満たす最も短い部分を見つけます。

正規表現 `/t[a-z]*i/` を文字列 `"titanic"` に適用できます。この正規表現は基本的に `t` で始まり、`i` で終わり、その間にいくつかの文字があるパターンです。

正規表現はデフォルトでは貪欲なので、マッチは `["titani"]` を返します。これはパターンに合う最大の部分文字列を見つけたことになります。

しかし、`?` 文字を使うと遅延マッチングに変更できます。`"titanic"` を調整した正規表現 `/t[a-z]*?i/` に対してマッチさせると `["ti"]` が返されます。

**注意:** 正規表現でHTMLを解析するのは避けるべきですが、HTML文字列のパターンマッチングに正規表現を使うのは問題ありません。

# --instructions--

正規表現 `/<.*>/` を修正して、テキスト `<h1>` ではなくHTMLタグ `"<h1>Winter is coming</h1>"` を返すようにしてください。正規表現のワイルドカード `.` は任意の文字にマッチすることを覚えておいてください。

# --hints--

`result` 変数は `<h1>` を含む配列であるべきです

```js
assert(result[0] == '<h1>');
```

`myRegex` は遅延マッチングを使うべきです

```js
assert(/[^\\][\*\+\?]\?/.test(myRegex));
```

`myRegex` は文字列 `h1` を含まないようにしてください

```js
assert(!myRegex.source.match('h1'));
```

# --seed--

## --seed-contents--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*>/; // Change this line
let result = text.match(myRegex);
```

# --solutions--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*?>/; // Change this line
let result = text.match(myRegex);
```
