---
id: 587d7db6367417b2b2512b99
title: 1回以上連続して現れる文字の一致
challengeType: 1
forumTopicId: 301350
dashedName: match-characters-that-occur-one-or-more-times
---

# --description--

時々、1回以上連続して現れる文字（または文字のグループ）に一致させる必要があります。これは少なくとも1回は現れ、繰り返される場合もあることを意味します。

それが当てはまるかどうかを確認するには、`+` 文字を使うことができます。文字やパターンは連続して存在しなければなりません。つまり、文字が1つずつ連続して繰り返される必要があります。

例えば、`/a+/g` は `abc` の中で1つの一致を見つけ、`["a"]` を返します。`+` のため、`aabc` の中でも1つの一致を見つけて `["aa"]` を返します。

もし `abab` という文字列をチェックしていたら、2つの一致を見つけて `["a", "a"]` を返します。なぜなら `a` の文字は連続しておらず、間に `b` があるからです。最後に、`a` の文字列に `bcd` がないため、一致は見つかりません。

# --instructions--

文字 `s` が `Mississippi` の中で1回以上現れる場合に一致を見つけたいです。`+` 記号を使った正規表現を書いてください。

# --hints--

あなたの正規表現 `myRegex` は、1回以上の `+` 文字に一致させるために `s` 記号を使うべきです。

```js
assert(/\+/.test(myRegex.source));
```

あなたの正規表現 `myRegex` は2つの項目に一致するはずです。

```js
assert(result.length == 2);
```

`result` 変数は `ss` の2つの一致を含む配列であるべきです。

```js
assert(result[0] == 'ss' && result[1] == 'ss');
```

# --seed--

## --seed-contents--

```js
let difficultSpelling = "Mississippi";
let myRegex = /change/; // Change this line
let result = difficultSpelling.match(myRegex);
```

# --solutions--

```js
let difficultSpelling = "Mississippi";
let myRegex = /s+/g; // Change this line
let result = difficultSpelling.match(myRegex);
```
