---
id: 587d7db9367417b2b2512ba6
title: マッチ数の下限のみを指定する
challengeType: 1
forumTopicId: 301366
dashedName: specify-only-the-lower-number-of-matches
---

# --description--

量指定子を使って、括弧で囲んだパターンの下限と上限の数を指定できます。時には、上限を指定せずに下限だけを指定したい場合があります。

パターンの下限のみを指定するには、最初の数字の後にカンマを付けてください。

例えば、文字 `hah` に対して、文字 `a` が少なくとも `3` 回出現する文字列にマッチさせたい場合、正規表現は `/ha{3,}h/` となります。

```js
let A4 = "haaaah";
let A2 = "haah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleA = /ha{3,}h/;
multipleA.test(A4);
multipleA.test(A2);
multipleA.test(A100);
```

順番に、3つの `test` の呼び出しは `true`、`false`、`true` を返します。

# --instructions--

正規表現 `haRegex` を変更して、文字 `Hazzah` が4回以上続く場合にのみ単語 `z` にマッチするようにしてください。

# --hints--

正規表現には括弧を使ってください。

```js
assert(haRegex.source.match(/{.*?}/).length > 0);
```

文字列 `Hazzah` にはマッチしないようにしてください。

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzah'));
```

文字列 `Hazzzah` にはマッチしないようにしてください。

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzzah'));
```

文字列 `Hazzzzah` にはマッチするようにしてください。

```js
assert('Hazzzzah'.match(haRegex)[0].length === 8);
```

文字列 `Hazzzzzah` にはマッチするようにしてください。

```js
assert('Hazzzzzah'.match(haRegex)[0].length === 9);
```

文字列 `Hazzzzzzah` にはマッチするようにしてください。

```js
assert('Hazzzzzzah'.match(haRegex)[0].length === 10);
```

文字列 `Hazzah`（`z` が30回続く）にもマッチするようにしてください。

```js
assert('Hazzzzzzzzzzzzzzzzzzzzzzzzzzzzzzah'.match(haRegex)[0].length === 34);
```

# --seed--

## --seed-contents--

```js
let haStr = "Hazzzzah";
let haRegex = /change/; // Change this line
let result = haRegex.test(haStr);
```

# --solutions--

```js
let haStr = "Hazzzzah";
let haRegex = /Haz{4,}ah/; // Change this line
let result = haRegex.test(haStr);
```
