---
id: 587d7db3367417b2b2512b8e
title: "test メソッドの使い方"
challengeType: 1
forumTopicId: 301369
dashedName: using-the-test-method
---

# --description--

正規表現はプログラミング言語で文字列の一部をマッチさせるために使われます。マッチングを助けるパターンを作成します。

文字列 `the` の中から単語 `The dog chased the cat` を見つけたい場合、次の正規表現を使うことができます：`/the/`。正規表現の中に引用符は必要ないことに注意してください。

JavaScript には正規表現を使う方法がいくつかあります。正規表現をテストする方法の一つが `.test()` メソッドです。`.test()` メソッドは正規表現を受け取り、それを丸括弧内に入れた文字列に適用し、パターンが見つかれば `true` を、見つからなければ `false` を返します。

```js
let testStr = "freeCodeCamp";
let testRegex = /Code/;
testRegex.test(testStr);
```

ここでの `test` メソッドは `true` を返します。

# --instructions--

正規表現 `myRegex` を文字列 `myString` に `.test()` メソッドで適用してください。

# --hints--

正規表現をテストするには `.test()` を使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

結果は `true` を返すはずです。

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex; // Change this line
```

# --solutions--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex.test(myString); // Change this line
```
