---
id: 587d7db9367417b2b2512ba4
title: 空白以外の文字にマッチさせる
challengeType: 1
forumTopicId: 18210
dashedName: match-non-whitespace-characters
---

# --description--

`\s`（小文字の`s`）を使って空白の検索について学びました。空白以外のすべての文字を検索することもできます。

大文字の`\S`である`s`を使って空白以外を検索してください。このパターンは空白、復帰、タブ、改ページ、改行文字にはマッチしません。`[^ \r\t\f\n\v]`という文字クラスに似ていると考えてください。

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let nonSpaceRegex = /\S/g;
whiteSpace.match(nonSpaceRegex).length;
```

`.length`メソッドが返す値は`32`になります。

# --instructions--

正規表現`countNonWhiteSpace`を変更して、文字列内の複数の空白以外の文字を探してください。

# --hints--

正規表現にはグローバルフラグを使う必要があります。

```js
assert(countNonWhiteSpace.global);
```

すべての空白以外の文字にマッチさせるために、短縮文字`\S`を使う必要があります。

```js
assert(/\\S/.test(countNonWhiteSpace.source));
```

正規表現は文字列`Men are from Mars and women are from Venus.`の中で35個の空白以外の文字を見つける必要があります。

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countNonWhiteSpace)
    .length == 35
);
```

正規表現は文字列`Space: the final frontier.`の中で23個の空白以外の文字を見つける必要があります。

```js
assert('Space: the final frontier.'.match(countNonWhiteSpace).length == 23);
```

正規表現は文字列`MindYourPersonalSpace`の中で21個の空白以外の文字を見つける必要があります。

```js
assert('MindYourPersonalSpace'.match(countNonWhiteSpace).length == 21);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /change/; // Change this line
let result = sample.match(countNonWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /\S/g; // Change this line
let result = sample.match(countNonWhiteSpace);
```
