---
id: 6140c7e645d8e905819f1dd4
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Kuanza, ongeza kipengele cha `link` kwa ajili ya faili la `./styles.css`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.match(code, /<link/);
```

Unapaswa kuwa na kipengele kimoja cha `link`.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `rel` yenye thamani `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `href` yenye thamani `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>

  </body>
</html>

```

```css

```
