---
id: 6140c9d35015ae0ba0c250e8
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Ongeza `div` ndani ya kipengele chako cha `body` na kipa `class` ya `wheel`.

Ndani ya `div` mpya, ongeza vipengele sita vya `span` vyenye `class` iliyowekwa kuwa `line`, na vipengele sita vya `div` vyenye `class` iliyowekwa kuwa `cabin`.

# --hints--

Unapaswa kuunda `div` mpya ndani ya kipengele chako cha `body`.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert.lengthOf(divs, 1);
```

`div` yako ndani ya kipengele cha `body` inapaswa kuwa na `class` ya `wheel`.

```js
assert.isTrue(document.querySelector('body div')?.classList?.contains('wheel'));
```

Kipengele chako cha `.wheel` kinapaswa kuwa na vipengele sita vya `span` ndani yake.

```js
assert.lengthOf(document.querySelectorAll('.wheel span'), 6);
```

Vipengele vyako sita vya `span` ndani ya kipengele cha `.wheel` vinapaswa kuwa na `class` ya `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert.isTrue(spans?.every(span => span?.classList?.contains('line')));
assert.lengthOf(document.querySelectorAll('.line'), 6);
```

Kipengele chako cha `.wheel` kinapaswa kuwa na vipengele sita vya `div` ndani yake.

```js
assert.lengthOf(document.querySelectorAll('.wheel div'), 6);
```

Vipengele vyako sita vya `div` ndani ya kipengele cha `.wheel` vinapaswa kuwa na `class` ya `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert.isTrue(divs?.every(div => div?.classList?.contains('cabin')));
assert.lengthOf(document.querySelectorAll('.cabin'), 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
