---
id: 6140df547f09402144e40b92
title: Hatua 20
challengeType: 0
dashedName: step-20
---

# --description--

Sifa ya `animation-name` hutumika kuunganisha sheria ya `@keyframes` na kichaguzi cha CSS. Thamani ya sifa hii inapaswa kufanana na jina la sheria ya `@keyframes`. Toa kichaguzi chako cha `.wheel` sifa ya `animation-name` iliyowekwa kuwa `wheel`.

Sifa ya `animation-duration` hutumika kuweka muda ambao mfululizo wa uhuishaji unapaswa kumalizika. Muda unapaswa kuelezwa kwa sekunde (`s`) au millisekunde (`ms`). Weka kichaguzi chako cha `.wheel` kuwa na sifa ya `animation-duration` yenye thamani ya `10s`.

# --hints--

Kichaguzi chako cha `.wheel` kinapaswa kuwa na sifa ya `animation-name` iliyowekwa kuwa `wheel`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationName, 'wheel');
```

Kichaguzi chako cha `.wheel` kinapaswa kuwa na sifa ya `animation-duration` iliyowekwa kuwa `10s`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationDuration, '10s');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

@keyframes wheel {
  0% {
    transform: rotate(0deg);
  }
  100% {
    transform: rotate(360deg);
  }
}
```
