---
id: 6140d0069049f5139d78da40
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Sifa ya `transform-origin` hutumika kuweka sehemu ambayo mabadiliko ya CSS yanatumika kuzunguka. Kwa mfano, unapofanya mabadiliko ya `rotate` (kama utakavyofanya baadaye katika mradi huu), `transform-origin` huamua sehemu ambayo kipengele kinazungushwa kuzunguka.

Toa kichaguzi cha `.line` sifa ya `transform-origin` yenye thamani ya `0% 0%`. Hii itasogeza sehemu ya asili ya `0%` kutoka kushoto na `0%` kutoka juu, ikiifanya iwe kona ya juu kushoto ya kipengele.

# --hints--

Kichaguzi chako cha `.line` kinapaswa kuwa na sifa ya `transform-origin` iliyowekwa kuwa `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert.oneOf(transformOrigin, ['0% 0%', '0% 0% 0px']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
