---
id: 6140cdebd39d6a101e747529
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Tengeneza kichaguzi kwa vipengele vyako vya `.line`. Anza kwa kuweka `background-color` kuwa `black`, `width` kuwa `50%`, na `height` kuwa `2px`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.line`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line'));
```

Kichaguzi chako cha `.line` kinapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor, "black");
```

Kichaguzi chako cha `.line` kinapaswa kuwa na sifa ya `width` iliyowekwa kuwa `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.width, "50%");
```

Kichaguzi chako cha `.line` kinapaswa kuwa na sifa ya `height` iliyowekwa kuwa `2px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.height, "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
