---
id: 68b94f23f7b3ee14078c89c5
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、`iframe` 要素を使って動画を表示します。基本的なHTMLのボイラープレートは用意されています。

まず、テキストが `h1` の `iframe Video Display` 要素を作成してください。

# --hints--

`h1` 要素があるはずです。

```js
assert.exists(document.querySelector('h1'))
```

`h1` 要素にはテキスト `iframe Video Display` があるべきです。

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
