---
id: 68ba9c4f1688914d72876093
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

今、`iframe` 要素を作成してください。まだ何も入れないでください。

# --hints--

`iframe` 要素があるはずです。

```js
assert.exists(document.querySelector('iframe'))
```

あなたの `iframe` 要素は空であるべきです。

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
