---
id: 68ba9c4f1688914d72876094
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`iframe` 要素の最初のレッスンで、`img` と同様に置き換えられた要素であることを学びました。つまり、`width` と `height` の属性を使って、高さと幅を指定できるということです。

`iframe` 要素に `width` の `560` と `height` の `315` を設定してください。

# --hints--

`iframe` 要素には `width` に設定された `560` 属性が必要です。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

`iframe` 要素には `height` に設定された `315` 属性が必要です。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
