---
id: 68ba9c4f1688914d72876098
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

次に追加する属性は `referrerpolicy` です。これは、ページが別のページに接続するときにどの程度の詳細を共有するかを決めるルールです。

`referrerpolicy` 属性を追加し、`strict-origin-when-cross-origin` に設定してください。これにより、同じサイト内では完全なアドレスが共有され、他のサイトではサイト名のみが共有され、安全でないサイトには何も共有されません。

# --hints--

`iframe` 要素には、`referrerpolicy` 属性を `strict-origin-when-cross-origin` に設定してください。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
