---
id: 68ba9c4f1688914d72876096
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

属性の1つに `allow` があります。これは、ブラウザーに `iframe` が使用できる機能を伝える許可リストのようなものです。

`iframe` 属性を持つ `allow` 要素の例を示します。

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

`allow` 属性を値 `accelerometer`、`autoplay`、`clipboard-write` で追加してください。

`accelerometer` は `iframe` にモーションセンサーの使用を許可し、デバイスの傾きや回転などを検出できるようにします。`autoplay` は動画を自動再生できるようにし、`clipboard-write` は iframe がユーザーのクリップボードにデータを書き込めるようにします。

# --hints--

あなたの `iframe` 要素には、`allow` 属性が `accelerometer` に設定されている必要があります。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

あなたの `allow` 要素の `iframe` 属性には、値の1つとして `autoplay` が含まれている必要があります。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

あなたの `allow` 要素の `iframe` 属性には、値の1つとして `clipboard-write` が含まれている必要があります。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
