---
id: 68ba9c4f1688914d72876095
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`iframe` 要素は、表示するリソースの URL またはパスを示す値を持つ `src` 属性も取ります。

`src` 要素に `https://www.youtube.com/embed/I0_951_MPE0` の `iframe` 属性を追加してください。

この時点で、ページに動画が表示されているはずですが、追加する必要がある属性がまだいくつかあります。

# --hints--

`iframe` 要素には、`src` に設定された `https://www.youtube.com/embed/I0_951_MPE0` 属性が必要です。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
