---
id: 68ba9c4f1688914d72876097
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

既存の `encrypted-media` 属性の値に `gyroscope`、`web-share`、および `allow` を追加してください。

これら3つは、動画を保護するための暗号化メディア拡張機能の使用、デバイスの動きや向きのセンサーへのアクセスの許可、そしてデバイスのネイティブ共有ダイアログを通じて iframe コンテンツを共有することを可能にします。

# --hints--

`allow` 要素の `iframe` 属性には、値の一つとして `encrypted-media` を含める必要があります。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

`allow` 要素の `iframe` 属性には、値の一つとして `gyroscope` を含める必要があります。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

`allow` 要素の `iframe` 属性には、値の一つとして `web-share` を含める必要があります。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
