---
id: 5dc1798ff86c76b9248c6eb3
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1` から `h6` までの見出し要素は、その下のコンテンツの重要度を示すために使われます。数字が小さいほど重要度が高く、`h2` 要素は `h1` 要素よりも重要度が低いです。

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

ページごとに `h1` 要素は1つだけ使い、重要度の低い見出しは重要度の高い見出しの下に配置してください。

`h1` 要素の下に、次のテキストを含む `h2` 要素を追加してください。

`Cat Photos`

# --hints--

`h1` 要素には開始タグが必要です。開始タグの構文は `<elementName>` です。

```js
assert(document.querySelector('h1'));
```

`h1` 要素には終了タグが必要です。終了タグは `/` 文字の直後に `<` があります。

```js
assert(code.match(/<\/h1\>/));
```

`h1` 要素は1つだけにしてください。余分なものは削除してください。

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

`h1` 要素のテキストは `CatPhotoApp` であるべきです。テキストが抜けているか、タイプミスがあります。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

`h2` タグは小文字で書いてください。慣例として、すべてのHTMLタグは小文字で記述します。

```js
assert.notMatch(code, /<\/?H2>/);
```

`h2` 要素には開始タグが必要です。開始タグの構文は `<elementName>` です。

```js
assert(document.querySelector('h2'));
```

`h2` 要素には終了タグが必要です。終了タグは `/` 文字の直後に `<` があります。

```js
assert(code.match(/<\/h2\>/));
```

`h2` 要素のテキストは `Cat Photos` であるべきです。開閉 `Cat Photos` タグの間にはテキスト `h2` のみを置いてください。

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

`h2` 要素は `h1` 要素の下に配置してください。`h1` 要素はより重要で、`h2` 要素の上にある必要があります。

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```
