---
id: 5dc17d3bf86c76b9248c6eb4
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`p` 要素はウェブサイト上で段落のテキストを作成するために使います。`p` 要素の下に `h2` 要素を作成し、次のテキストを入れてください。

`Everyone loves cute cats online!`

# --hints--

`p` 要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert(document.querySelector('p'));
```

`p` タグは小文字で書いてください。慣例として、すべてのHTMLタグは小文字で記述されます。

```js
assert.notMatch(code, /<\/?P>/);
```

`p` 要素には終了タグが必要です。終了タグは `/` の直後に `<` があります。

```js
assert(code.match(/<\/p\>/));
```

`p` 要素のテキストは `Everyone loves cute cats online!` であるべきです。テキストが抜けているか、タイプミスがあります。

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

`p` 要素は `h2` 要素の下にあるべきです。順番が間違っています。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```
