---
id: 5dc17dc8f86c76b9248c6eb5
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

コメントはブラウザーの表示に影響を与えずにメッセージを残すことができます。また、コードを無効にすることもできます。HTMLのコメントは`<!--`で始まり、任意の行数のテキストを含み、`-->`で終わります。

以下は`TODO: Remove h1`を使ったコメントの例です。

```html
<!-- TODO: Remove h1 -->
```

`p`要素の上に、次のテキストを含むコメントを追加してください。

`TODO: Add link to cat photos`

# --hints--

コメントは`<!--`で始まる必要があります。コメントの開始を示す文字が1つ以上欠けています。

```js
assert(code.match(/<!--/));
```

コメントは`-->`で終わる必要があります。コメントの終了を示す文字が1つ以上欠けています。

```js
assert(code.match(/-->/));
```

コードに余分なコメント開始または終了文字があってはいけません。ブラウザーに余分な`<!--`または`-->`が表示されています。

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

コメントは`p`要素の上にある必要があります。順序が間違っています。

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

コメントにはテキスト`TODO: Add link to cat photos`を含めてください。

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
