---
id: 5dc174fcf86c76b9248c6eb2
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML要素には`<h1>`のような開始タグと`</h1>`のような終了タグがあります。要素が表示するテキストは、開始タグと終了タグの間に入ります。

下の`h1`要素のテキストを`Hello World`から`CatPhotoApp`に変更し、ブラウザのプレビューで変化を確認してください。

終わったら、「Check Your Code」ボタンを押して正しいかどうか確認してください。

# --hints--

コード内に`CatPhotoApp`というテキストがあるはずです。スペルを確認してください。

```js
assert(code.match(/catphotoapp/i));
```

`h1`要素には開始タグが必要です。開始タグの構文は`<elementName>`のようになっています。

```js
assert(document.querySelector('h1'));
```

`h1`タグは小文字で書く必要があります。慣例として、すべてのHTMLタグは小文字で記述します。

```js
assert.notMatch(code, /<\/?H1>/);
```

`h1`要素には終了タグが必要です。終了タグは`/`文字の直後に`<`があります。

```js
assert(code.match(/<\/h1\>/));
```

`h1`要素が複数あります。余分な`h1`要素を削除してください。

```js
assert(document.querySelectorAll('h1').length === 1);
```

`h1`要素のテキストは`CatPhotoApp`であるべきです。テキストが抜けているか、タイプミスがあるか、`h1`要素の開始タグと終了タグの間にテキストがありません。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```
