---
id: 5dc24073f86c76b9248c6ebb
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

HTMLの<dfn>属性</dfn>は、要素の開始タグ内で使われる特別な単語で、要素の動作を制御します。`src`属性は`img`要素内で画像のURL（画像の場所）を指定します。

以下は、freeCodeCampのロゴを指す`img`属性を持つ`src`要素の例です。

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

既存の`img`要素内に、次のURLを持つ`src`属性を追加してください。

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

コードには`img`要素が必要です。`img`要素を削除してしまったか、`src`属性の値を引用符で囲んでいない可能性があります。

```js
assert(document.querySelector('img'));
```

`img`要素には`src`属性が必要です。属性を省略しているか、タイプミスがあります。要素名と属性名の間にスペースがあることを確認してください。

```js
assert(document.querySelector('img').src);
```

`img`要素の`src`属性は'`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'に設定してください。URLを省略しているか、タイプミスがあります。URLの大文字・小文字は重要です。

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

`img`要素の`src`を正しいURLに設定していても、属性の値は常に引用符で囲むことをお勧めします。

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

`img`要素は`>`または`/>`で閉じる必要があります。

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
