---
id: 5dc24165f86c76b9248c6ebc
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

すべての`img`要素には`alt`属性が必要です。`alt`属性のテキストはスクリーンリーダーでアクセシビリティを向上させるために使われ、画像が読み込めなかった場合に表示されます。

以下は`img`要素に`alt`属性を付けた例です：

```html
<img src="cat.jpg" alt="A cat">
```

`img`要素の中に、次のテキストを持つ`alt`属性を追加してください：

`A cute orange cat lying on its back`

# --hints--

コードには`img`要素が必要です。前のステップで`img`要素を削除しました。

```js
assert(document.querySelector('img'));
```

`img`要素に`alt`属性がありません。開始タグの名前の後にスペースがあるか、すべての属性名の前にスペースがあるか確認してください。

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

`img`要素の`alt`属性の値が`A cute orange cat lying on its back`以外に設定されています。`alt`属性の値が引用符で囲まれていることを確認してください。

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
