---
id: 5dc23991f86c76b9248c6eb8
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

前のステップでは、`h1`、`h2`、コメント、そして`p`要素を`main`要素の中に入れました。これは*ネストする*と呼ばれます。ネストされた要素は、ネストされている要素よりも右に2文字分スペースを空けて配置します。このスペースの空け方をインデントと呼び、HTMLを読みやすくするために使われます。

以下はネストとインデントの例です：

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

下のコードでは、`h1`要素、`h2`要素、そしてコメントが`main`要素よりも2文字分多くインデントされています。キーボードのスペースキーを使って、`p`要素の前に2文字分のスペースを追加し、正しくインデントしてください。

# --hints--

コードには`h2`要素があり、そのテキストは`Cat Photos`である必要があります。誤って削除したか、開始タグと終了タグの両方が欠けているか、テキストが変更されている可能性があります。

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

例にある`ul`や`li`要素は追加しないでください。

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

`h2`要素の行のインデントは変更しないでください。開始タグは行の先頭から6文字分の位置にある必要があります。元のインデントに戻すにはステップを再開してください。

```js
assert(code.toLowerCase().match(/<\/h1\>\s*\n\s{6}<h2>/));
```

コードにはコメントが必要です。前のステップでコメントを削除しています。

```js
assert(code.match(/<!--.*-->/));
```

コメントのテキストは`TODO: Add link to cat photos`である必要があります。コメントのテキストやスペースは変更しないでください。

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

コメント要素の行のインデントは変更しないでください。開始タグは行の先頭から6文字分の位置にある必要があります。元のインデントに戻すにはステップを再開してください。

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

コードには`p`要素が必要です。前のステップで`p`要素を削除しています。

```js
assert(document.querySelector('p'));
```

`p`要素のテキストは`Everyone loves cute cats online!`である必要があります。`p`要素のテキスト、スペース、句読点は変更しないでください。

```js
assert.match(
  document
    .querySelector('p')
    .innerText.toLowerCase(),
    /everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

開始`p`タグのインデントは、`h2`要素やコメント要素と一致している必要があります。開始タグは行の先頭から6文字分の位置にある必要があります。

```js
assert(code.toLowerCase().match(/-->\s*\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
