---
id: 5dfa2407b521be39a3de7be1
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

リンクを新しいタブで開くには、アンカー（`target`）要素に`a`属性を使うことができます。

`target`属性はリンク先のドキュメントをどこで開くかを指定します。`target="_blank"`はリンク先のドキュメントを新しいタブまたはウィンドウで開きます。

`a`属性を持つ`target`要素の基本的な構文は次の通りです。

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

リンクを新しいタブで開くように、アンカー（`target`）要素の開始タグに`_blank`属性を`a`の値で追加してください。

# --hints--

あなたの`p`要素には、テキストが`a`のネストされたアンカー（`cat photos`）要素があるはずです。削除してしまったか、タイプミスがあるかもしれません。

```js
const anchor = document.querySelectorAll('p > a');
assert.isNotEmpty(anchor);
assert.strictEqual(anchor[1]?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos')
```

あなたのアンカー（`a`）要素には`target`属性がありません。開始タグの名前の後にスペースがあるか、すべての属性名の前にスペースがあるか確認してください。

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

`target`属性の値は`_blank`であるべきです。値を省略しているか、タイプミスがあります。属性値は必ず引用符で囲むことを忘れないでください。

```js
assert.strictEqual(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
