---
id: 5dfa3589eacea3f48c6300ae
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

2番目の`section`要素の中に、テキスト`h2`を持つ新しい`Cat Lists`要素を追加してください。

# --hints--

あなたの`section`要素は開始タグを持つ必要があります。開始タグの構文は次の通りです：`<elementName>`。

```js
assert(
  document.querySelectorAll('section').length === 2 &&
    code.match(/<\/section>/g).length === 2
);
```

あなたの`h2`要素は開始タグを持つ必要があります。開始タグの構文は次の通りです：`<elementName>`。

```js
assert(document.querySelectorAll('h2').length === 2);
```

あなたの`h2`要素は終了タグを持つ必要があります。終了タグは`/`文字の直後に`<`があります。

```js
assert(code.match(/<\/h2\>/g).length === 2);
```

2番目の`h2`要素は2番目の`section`要素の終了タグのすぐ上にある必要があります。正しい位置にありません。

```js
const secondSection = document.querySelectorAll('section')[1];
assert(secondSection.lastElementChild.nodeName === 'H2');
```

2番目の`h2`要素はテキスト`Cat Lists`を持つ必要があります。テキストを省略しているか、タイプミスがあります。

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase() === 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
