---
id: 5dfb6a35eacea3f48c6300b4
title: ステップ 26
challengeType: 0
dashedName: step-26
---

# --description--

`figcaption` 要素は、`figure` 要素内に含まれる画像を説明するキャプションを追加するために使います。

以下は、キャプションが `A cute cat` の `figcaption` 要素の例です。

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

`figure` 要素にネストされた画像の後に、テキストを次のように設定した `figcaption` 要素を追加してください。

`Cats love lasagna.`

# --hints--

`figcaption` 要素は開始タグを持つ必要があります。開始タグの構文は次の通りです：`<elementName>`。

```js
assert(document.querySelector('figcaption'));
```

`figcaption` 要素は終了タグを持つ必要があります。終了タグは `/` の直後にあります。

```js
assert(code.match(/<\/figcaption\>/));
```

`figcaption` 要素は `figure` 要素にネストされている必要があります。

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

ラザニアの `img` 要素は `figure` 要素にネストされている必要があります。

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

`figure` 要素にネストされた `figcaption` 要素は、`img` 要素の下に配置してください。順序が逆になっています。

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

`figcaption` 要素のテキストは `Cats love lasagna.` である必要があります。テキストが抜けているか、タイプミスがあります。

```js
assert(
  document.querySelector('figcaption').innerText.trim().match(/^Cats love lasagna\.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          
        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
