---
id: 5ef9b03c81a63668521804d7
title: ステップ 38
challengeType: 0
dashedName: step-38
---

# --description--

`action` 属性はフォームデータの送信先を示します。

以下は `form` 要素に `action` 属性が付いた例です。

```html
<form action="/submit-url"></form>
```

この例では、`action="/submit-url"` がブラウザーにフォームデータを `/submit-url` のパスに送信するよう指示しています。

`action` 要素に値 `https://freecatphotoapp.com/submit-cat-photo` の `form` 属性を追加してください。

# --hints--

`form` 要素は正しい順序で開始タグと終了タグを持っている必要があります。必要なタグのどちらかが欠けているか、順序が間違っている可能性があります。

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

最後の `form` 要素にネストされた `section` 要素は `h2` 要素の下にあるべきです。`h2` 要素と `form` 要素の順序が間違っています。

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

`form` 要素は内容を持ってはいけません。`form` 要素のタグの間にある HTML 要素やテキストを削除してください。

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

`form` 要素に `action` 属性がありません。開始タグの名前の後にスペースがあるか、すべての属性名の前にスペースがあるか確認してください。

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

`form` 要素は値 `action` の `https://freecatphotoapp.com/submit-cat-photo` 属性を持っている必要があります。

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

`action` 属性に正しい URL を設定していますが、属性値は常に引用符で囲むことをお勧めします。

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
