---
id: 5ef9b03c81a63668521804e3
title: ステップ 60
challengeType: 0
dashedName: step-60
---

# --description--

先ほど追加したチェックボックスの後に、もう一つチェックボックスを追加してください。`id` 属性の値は `lazy` にし、`name` 属性の値は最後のチェックボックスと同じにしてください。

また、新しいチェックボックスの右側に `label` 要素を追加し、テキストは `Lazy` にしてください。`label` 要素を `for` 属性で新しいチェックボックスに関連付けることを忘れないでください。

# --hints--

新しいチェックボックスを追加する必要があります。

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'), 2);
```

新しいチェックボックスには、`id` 属性が `lazy` の値を持ち、`name` 属性が `personality` の値を持つべきです。開始タグの名前の後やすべての属性名の前にスペースがあるか確認してください。

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'lazy' && checkbox.getAttribute('name') === 'personality'
  )
);
```

新しいチェックボックスは最初のチェックボックスの後に配置してください。順番が間違っています。

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('loving') < checkboxes.indexOf('lazy'));
```

新しいチェックボックスの右側には、テキストが `label` の `Lazy` 要素があるべきです。

```js
const nextElementSibling = document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling;
assert.isOk(nextElementSibling?.nodeName === 'LABEL' && nextElementSibling?.innerText?.replace(/\s+/g, '')?.match(/^Lazy$/i)
);
```

新しい `label` には、新しいチェックボックスの `for` 属性と同じ値の `id` 属性が必要です。値を省略しているか、タイプミスがあります。

```js
assert.strictEqual(
  document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling?.getAttribute('for'),
    'lazy'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
