---
id: 5ef9b03c81a63668521804e9
title: ステップ 66
challengeType: 0
dashedName: step-66
---

# --description--

既存の`freeCodeCamp.org`テキストをアンカー（`a`）要素で囲んでリンクにしてください。`href`属性は`https://www.freecodecamp.org`に設定する必要があります。

# --hints--

アンカー（`a`）要素は`p`内の`footer`要素の中にネストしてください。アンカー（`a`）要素の開始タグと終了タグを必ず追加してください。

```js
assert(document.querySelectorAll('footer > p > a').length);
```

アンカー（`a`）要素には終了タグが必要です。終了タグは`/`文字の直後に`<`があります。

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.lengthOf(aElemClosingTags, 4);
```

アンカー（`a`）要素には、値が`href`の`https://www.freecodecamp.org`属性が必要です。属性や値を省略しているか、タイプミスがあるかもしれません。

```js
const nestedAnchor = document.querySelector('footer > p > a');
assert.strictEqual(nestedAnchor?.getAttribute('href')?.toLowerCase(), 'https://www.freecodecamp.org');
```

リンクのテキストは`freeCodeCamp.org`である必要があります。テキストを省略しているか、タイプミスがある可能性があります。

```js
const nestedAnchor = document.querySelector('footer > p > a');
assert.strictEqual(
  nestedAnchor?.innerText?.toLowerCase()?.replace(/\s+/g, ' '),
  'freecodecamp.org'
);
```

アンカー（`a`）要素をネストした後、ブラウザに表示される唯一の`p`要素の内容は`No Copyright - freeCodeCamp.org`であるべきです。`p`とネストされたアンカー要素のテキスト、スペース、句読点をもう一度確認してください。

```js
const pText = document.querySelector('footer > p')?.innerText?.toLowerCase()?.replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving" checked> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic">Energetic</label>
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
