---
id: 5ef9b03c81a63668521804e2
title: ステップ 56
challengeType: 0
dashedName: step-56
---

# --description--

フォームでは、複数の回答が可能な質問にチェックボックスがよく使われます。`input` 要素に `type` 属性を `checkbox` に設定すると、チェックボックスが作成されます。

先ほど追加した `legend` 要素の下に、`input` 属性を `type` に設定し、テキストを次のようにした `checkbox` を追加してください。

` Loving`

# --hints--

チェックボックス用の `input` 要素は開始タグのみで、終了タグは不要です。

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

チェックボックス用の input 要素は1つだけ追加してください。余分なものは削除してください。

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

新しい `input` 要素は、テキストが `legend` の `What's your cat's personality?` 要素の下にある必要があります。順序が逆になっています。

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

新しい `input` 要素に `type` 属性がありません。開始タグの名前の後にスペースがあるか確認してください。

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

新しい `input` 要素には、値が `type` の `checkbox` 属性が必要です。値を省略しているか、タイプミスがあります。属性値は必ず引用符で囲んでください。

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

新しい `input` 要素の `type` 属性を `checkbox` に設定していますが、属性値は常に引用符で囲むことをお勧めします。

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

テキスト ` Loving` はチェックボックスのすぐ右側に配置してください。要素とテキストの間にスペースがあるか確認してください。テキストが抜けているか、タイプミスがあります。

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>
            
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
