---
id: 5ef9b03c81a63668521804d6
title: ステップ 37
challengeType: 0
dashedName: step-37
---

# --description--

これからユーザーから情報を集めるためのウェブフォームを追加します。

`form` 要素は、名前やメールアドレスなどユーザーの情報を取得するために使われます。

`Cat Form` 見出しの後に、`form` 要素を追加してください。

# --hints--

`form` 要素は開始タグと終了タグの両方が必要です。どちらか一方または両方が欠けているか、順序が間違っている可能性があります。

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

`form` 要素のタグの順序が正しくありません。

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

最後の `form` 要素にネストされた `section` 要素は、`h2` 要素の下にあるべきです。`h2` 要素と `form` 要素の順序が間違っています。

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

`form` 要素は内容を持ってはいけません。`form` 要素のタグの間にあるHTML要素やテキストを削除してください。

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
