---
id: 5ef9b03c81a63668521804dc
title: ステップ 46
challengeType: 0
dashedName: step-46
---

# --description--

複数の選択肢の中から1つだけ答えを選んでほしい質問には、ラジオボタンを使えます。

以下は、テキストが `cat` に設定されたラジオボタンの例です。

```html
<input type="radio"> cat
```

`input` 要素は空要素であることを覚えておいてください。

テキスト入力の前に、テキストが次のように設定されたラジオボタンを追加してください：

`Indoor`

# --hints--

ラジオボタンには `input` 要素を作成してください。構文を確認しましょう。

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

`input` 要素は開始タグのみで、終了タグは不要です。

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

ラジオボタンには `input` 要素を1つだけ追加してください。余分なものは削除しましょう。

```js
assert(document.querySelectorAll('form > input'), 2);
```

新しい `input` 要素は、既存の `input` の上にあり、`type` 属性が `text` に設定されている必要があります。順番が間違っています。

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

新しい `input` 要素に `type` 属性がありません。開始タグの名前の後にスペースがあるか確認してください。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

新しい `input` 要素は `type` 属性を1つだけ持つべきです。余分なものは削除してください。

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

新しい `input` 要素は、値が `type` の `radio` 属性を持つべきです。値を省略しているか、タイプミスがあります。属性値は必ず引用符で囲むことを忘れないでください。

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

新しい `input` 要素の `type` 属性を `radio` に設定していますが、属性値は常に引用符で囲むことをお勧めします。

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

`radio` ボタンの ` Indoor` テキストは、ボタンの前ではなく後ろに配置してください。

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

テキスト ` Indoor` は `radio` ボタンのすぐ右側に配置してください。テキストが抜けているか、タイプミスがあります。

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
