---
id: 5ef9b03c81a63668521804e5
title: ステップ 63
challengeType: 0
dashedName: step-63
---

# --description--

チェックボックスをデフォルトでチェック済みにしたり、ラジオボタンを選択済みにするには、`checked` 属性を追加する必要があります。

以下は `checked` 属性が付いたラジオボタンの例です：

```html
<input checked type="radio" name="meal" value="breakfast"> Breakfast
```

`checked` 属性に値を設定する必要はありません。代わりに、`checked` 要素に `input` という単語を追加し、他の属性との間にスペースを入れてください。

最初のラジオボタンと最初のチェックボックスをデフォルトで選択済みにしてください。

# --hints--

それぞれの `fieldset` 要素に、ラジオボタンが2つ、チェックボックスが3つネストされたままであることを確認してください。

```js
assert.isTrue(
  document.querySelectorAll('input[type="radio"]').length === 2 &&
     document.querySelectorAll('fieldset > input[type="checkbox"]').length === 3
);
```

`id` が `indoor` に設定された最初のラジオボタンには、`checked` 属性が必要です。

```js
assert.isTrue(document.querySelector('input[type="radio"]')?.hasAttribute('checked'));
```

2番目のラジオボタンには `checked` 属性を付けないでください。

```js
assert.isFalse(document.querySelectorAll('input[type="radio"]')?.[1]?.hasAttribute('checked'));
```

最初のチェックボックスには `checked` 属性がありません。

```js
assert.isTrue(document.querySelector('input[type="checkbox"]')?.hasAttribute('checked'));
```

2番目のチェックボックスには `checked` 属性を付けないでください。

```js
assert.isFalse(document.querySelectorAll('input[type="checkbox"]')?.[1]?.hasAttribute('checked'));
```

3番目のチェックボックスには `checked` 属性を付けないでください。

```js
assert.isFalse(document.querySelectorAll('input[type="checkbox"]')?.[2]?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic"> Energetic</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
