---
id: 5ef9b03c81a63668521804d1
title: ステップ 28
challengeType: 0
dashedName: step-28
---

# --description--

`figure` 要素の後に、テキストが次のようになっている別の `h3` 要素を追加してください：

`Top 3 things cats hate:`

# --hints--

2番目の `h3` 要素の終了タグのすぐ上に `section` 要素があるはずです。開始タグと終了タグがあることを確認してください。

```js
assert(
  document.querySelectorAll('main > section')[1].lastElementChild.nodeName ===
    'H3' && code.match(/<\/h3\>/g).length === 2
);
```

新しい `h3` 要素にはテキスト `Top 3 things cats hate:` が含まれている必要があります。テキストの最後にコロンがあることを必ず確認してください。

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'top 3 things cats hate:'
);
```

新しい `figure` 要素の上に `h3` があるはずです。`figure` 要素を誤って削除してしまった可能性があります。

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  .lastElementChild;
assert(
  secondSectionLastElemNode.nodeName === 'H3' &&
    secondSectionLastElemNode.previousElementSibling.nodeName === 'FIGURE'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
