---
id: 5ef9b03c81a63668521804d2
title: ステップ 29
challengeType: 0
dashedName: step-29
---

# --description--

順序付きリスト (ordered list、`ol`) のコードは順序なしリストと似ていますが、順序付きリストのリスト項目は番号付きで表示されます。

`h3` 要素の下に、次の 3 つの項目を持つ順序付きリストを追加してください:

`flea treatment` `thunder` `other cats`

# --hints--

`ol` 要素には開始タグが必要です。 Opening tags have this syntax: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

`ol` 要素には終了タグが必要です。 Closing tags have a `/` just after the `<` character.

```js
assert(code.match(/<\/ol>/));
```

`ol` 要素は、2 つ目の `section` 要素の終了タグの上にある必要があります。 順番に誤りがあるようです。

```js
assert.equal(document.querySelectorAll('main > section')?.[1]?.lastElementChild?.nodeName, 'OL');
```

3 つの `li` 要素は、`ol` 要素の中に入れ子にする必要があります。

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

3 つの `li` 要素に、テキスト `flea treatment`、`thunder`、`other cats` が任意の順番で設定されている必要があります。

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

`ol` 要素が 1 つだけ必要です。

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

