---
id: 5ef9b03c81a63668521804d9
title: ステップ 42
challengeType: 0
dashedName: step-42
---

# --description--

プレースホルダーテキストは、どのような情報を入力すればよいかのヒントをユーザーに伝えるために使われます。

以下は、プレースホルダーが`input`に設定された`e.g. Jane Doe`要素の例です。

```html
<input type="text" placeholder="e.g. Jane Doe">
```

`cat photo URL`要素にプレースホルダーテキスト`input`を追加してください。

# --hints--

`input`要素を削除したか、構文が無効です。すべての属性値は引用符で囲む必要があります。

```js
assert(document.querySelectorAll('input').length);
```

`form`には`input`要素のみを含めてください。`form`要素内の追加のHTML要素やテキストは削除してください。

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

`input`要素に`placeholder`属性がありません。開始タグ名の後にスペースがあるか、すべての属性名の前にスペースがあるか確認してください。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('placeholder'));
```

`input`要素には値が`placeholder`の`cat photo URL`属性が必要です。値を省略しているか、タイプミスがあります。属性値は引用符で囲む必要があることを忘れないでください。

```js
assert(
  document.querySelector('input')
    ?.getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

`input`要素の`placeholder`属性を`cat photo URL`に設定していますが、属性値は常に引用符で囲むことをお勧めします。

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
