---
id: 5efb2c990dc218d6c85f89b2
title: ステップ 45
challengeType: 0
dashedName: step-45
---

# --description--

テキスト入力の下にボタンを追加しても、ページ上では隣り合って表示されます。これは、`input` と `button` の両方が<dfn>インライン要素</dfn>であり、新しい行に表示されないためです。

追加したボタンはデフォルトでフォームを送信します。しかし、デフォルトの動作に頼ると混乱を招く場合があります。`type` に `submit` 属性を値 `button` で追加して、送信ボタンであることを明確にしてください。

# --hints--

あなたの `button` 要素には開始タグが必要です。開始タグの構文は `<elementName>` です。

```js
assert(document.querySelector('button'));
```

あなたの `button` 要素には終了タグが必要です。終了タグは `/` の直後に `<` があります。

```js
assert(code.match(/<\/button\>/));
```

あなたの `button` 要素には `type` 属性がありません。開始タグの名前の後にスペースがあるか確認してください。

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('type'));
```

あなたの `button` 要素には値 `type` の `submit` 属性が必要です。値を省略しているか、タイプミスがあります。属性値は必ず引用符で囲むことを忘れないでください。

```js
assert.match(
  document.querySelector('button')?.getAttribute('type'),
  /^submit$/i
);
```

`button` 要素の `type` 属性を `submit` に設定していますが、属性値は常に引用符で囲むことをお勧めします。

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
