---
id: 5efb23e70dc218d6c85f89b1
title: ステップ 40
challengeType: 0
dashedName: step-40
---

# --description--

`type` 属性を使って作成できる入力にはさまざまな種類があります。パスワードフィールドやリセットボタン、ユーザーが自分のコンピューターからファイルを選択できるコントロールなどを簡単に作成できます。

`type` 要素に `text` 属性を値 `input` で追加して、ユーザーからテキスト入力を受け取るテキストフィールドを作成してください。

# --hints--

`input` 要素を削除したか、構文が無効です。属性を追加した場合は、その値が引用符で囲まれていることを確認してください。

```js
assert(document.querySelectorAll('input').length);
```

`form` 要素には `input` 要素のみを含めるべきです。`form` 要素のタグの間に余分なHTML要素やテキストがあれば削除してください。

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

`input` 要素に `type` 属性が値 `text` で設定されていません。開始タグの名前の後にスペースがあるか、すべての属性名の前にスペースがあるか確認してください。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

`input` 要素には `type` 属性が値 `text` で設定されている必要があります。値を省略しているか、タイプミスがあります。属性値は引用符で囲む必要があることを覚えておいてください。

```js
assert.match(
  document.querySelector('input')?.getAttribute('type').replace(/\s+/g, ' '),
  /^text$/i
);
```

`input` 要素の `type` 属性を `text` に設定していますが、属性値は常に引用符で囲むことをお勧めします。

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
