---
id: 5f05a1d8e233dff4a68508d8
title: ステップ 49
challengeType: 0
dashedName: step-49
---

# --description--

最初のラジオボタンの下にもう一つラジオボタンを作成してください。`label` 要素の中にネストし、`Outdoor` を `label` テキストとして使います。ラジオボタンには `id` 属性を付け、値は `outdoor` にしてください。

# --hints--

元の `Indoor` ラジオボタンだけに `id` が `indoor` に設定されている必要があります。必要ならステップを再起動して元のHTMLを戻せます。

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

`Indoor` ラジオボタンには変更を加えないでください。必要ならステップを再起動して元のHTMLを戻せます。

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

新しい `input` 要素を1つだけ、新しい `label` 要素の中にネストして追加してください。新しい `label` は開始タグと終了タグの両方を持つことを確認してください。

```js
assert(document.querySelectorAll('label input').length === 2);
```

どちらの開始 `label` タグにも属性を追加しないでください。

```js
assert(code.match(/<label\s*>/g).length === 2);
```

新しい `input` には `id` 属性を付け、値は `outdoor` にしてください。

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

新しい `input` には `type` 属性を付け、値は `radio` にしてください。

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

`input` を `label` の中にネストして追加する以外に新しい要素を追加しないでください。

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

新しい `input` の左側にテキストがあってはいけません。

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

新しいラジオボタンのラベルテキストは `Outdoor` と完全に一致している必要があります。大文字小文字も含みます。

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

新しいラジオボタンとラベルは `Indoor` ラジオボタンとラベルのすぐ下／後に配置してください。その間に他のタグがあってはいけません。

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
