---
id: 5f07fb1579dc934717801375
title: ステップ 35
challengeType: 0
dashedName: step-35
---

# --description--

新しいセクションを追加する時間です。2番目の`section`要素の下に3番目の`section`要素を追加してください。

# --hints--

`section`要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert(document.querySelectorAll('section').length >= 3);
```

開始`section`タグは1つだけ追加してください。余分なものは削除してください。

```js
assert(document.querySelectorAll('section').length === 3);
```

`section`要素には終了タグが必要です。終了タグは`/`で、`<`文字のすぐ後にあります。

```js
assert(code.match(/<\/section>/g).length >= 3);
```

終了`section`タグは1つだけ追加してください。余分なものは削除してください。

```js
assert(code.match(/<\/section>/g).length === 3);
```

すべての`section`要素は`main`要素の開始タグと終了タグの間にある必要があります。

```js
const childrenOfMain = [...document.querySelector('main').children];
const sectionElemsFound = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(sectionElemsFound.length === 3);
```

最後の`section`要素には内容があってはいけません。`section`要素内のHTML要素やテキストは削除してください。

```js
assert.lengthOf(document.querySelectorAll('main > section')?.[2]?.children, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
