---
id: 5f0d4ab1b435f13ab6550052
title: ステップ 54
challengeType: 0
dashedName: step-54
---

# --description--

次に、新しいフォーム `input` 要素を追加します。現在の `fieldset` 要素のすぐ下に、もう1つの `fieldset` 要素を追加してください。

# --hints--

新しい `fieldset` 要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert(document.querySelectorAll('fieldset').length >= 2);
```

開始 `fieldset` タグは1つだけ追加してください。余分なものは削除してください。

```js
assert(document.querySelectorAll('fieldset').length === 2);
```

新しい `fieldset` 要素には終了タグが必要です。終了タグは `/` のすぐ後ろに `<` 文字があります。

```js
assert(code.match(/<\/fieldset>/g).length >= 2);
```

終了 `fieldset` タグは1つだけ追加してください。余分なものは削除してください。

```js
assert(code.match(/<\/fieldset>/g).length === 2);
```

2番目の `fieldset` 要素は、最初の `fieldset` 要素の中にネストされてはいけません。

```js
const childrenOf1stFieldset = [
  ...document.querySelector('form > fieldset').children
];
const foundElems = childrenOf1stFieldset.filter((child) => {
  return child.nodeName === 'FIELDSET';
});
assert(foundElems.length === 0);
```

両方の `fieldset` 要素はテキストフィールドの上にあるべきです。順序が間違っています。

```js
const formChildren = document.querySelector('form')?.children;
assert.isTrue(
  formChildren?.[0]?.nodeName === 'FIELDSET' &&
  formChildren?.[1]?.nodeName === 'FIELDSET' &&
  formChildren?.[2]?.nodeName === 'INPUT' &&
  formChildren?.[2]?.getAttribute('type') === 'text'
);
```

新しい `fieldset` 要素は既存の `fieldset` 要素の下にあるべきです。順序が逆になっています。

```js
const fieldsetChildren = [...document.querySelectorAll('fieldset')].map(
  (elem) => elem.children
);
assert(fieldsetChildren[0].length > fieldsetChildren[1].length);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
