---
id: 5f1a80975fc4bcae0edb3497
title: ステップ 51
challengeType: 0
dashedName: step-51
---

# --description--

`Indoor` ラジオボタンを選択してフォームを提出すると、そのボタンのフォームデータは `name` と `value` 属性に基づきます。あなたのラジオボタンには `value` 属性がないため、フォームデータには `indoor-outdoor=on` が含まれますが、複数のボタンがある場合には役に立ちません。

両方のラジオボタンに `value` 属性を追加してください。利便性のために、ボタンの `value` 属性を `id` 属性と同じ値に設定してください。

# --hints--

両方のラジオボタンは、開始タグと終了タグの `label` 要素の間に配置されている必要があります。

```js
const labelChildNodes = [...document.querySelectorAll('form > label')].map(
  (node) => node.childNodes
);
assert(
  labelChildNodes.filter((childNode) => childNode[0].nodeName === 'INPUT')
    .length === 2
);
```

両方のラジオボタンには `value` 属性が必要です。開始タグの名前の後ろや、すべての属性名の前にスペースがあるか確認してください。

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(radioButtons.every((btn) => btn.hasAttribute('value')));
```

両方のラジオボタンには `type` 属性が必要です。開始タグの名前の後ろや、すべての属性名の前にスペースがあるか確認してください。

```js
const radioButtons = [...document.querySelectorAll('input')];
assert(radioButtons.every((btn) => btn.hasAttribute('type')));
```

`Indoor` ラジオボタンの `value` 属性は `indoor` に設定されている必要があります。値を省略しているか、タイプミスがあります。属性値は必ず引用符で囲んでください。

```js
const indoorRadioButton = document.querySelector('#indoor');
assert(indoorRadioButton.getAttribute('value').match(/^indoor$/));
```

`Outdoor` ラジオボタンの `value` 属性は `outdoor` に設定されている必要があります。値を省略しているか、タイプミスがあります。属性値は必ず引用符で囲んでください。

```js
const outdoorRadioButton = document.querySelector('#outdoor');
assert(outdoorRadioButton.getAttribute('value').match(/^outdoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
          <label><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
