---
id: a6b0bb188d873cb2c8729495
title: Implementa un Convertidor de Entidades HTML
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

Este laboratorio trata sobre convertir caracteres especiales en una cadena con sus correspondientes entidades HTML.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Deberías tener una función `convertHTML` que acepte una cadena como argumento.
1. La función `convertHTML` debe devolver una nueva cadena convirtiendo los caracteres especiales en la cadena argumento a sus entidades HTML correspondientes.

   - `&` debe convertirse en `&amp;`.
   - `<` debe convertirse en `&lt;`.
   - `>` debe convertirse en `&gt;`.
   - `&quot;` debería convertirse en `&quot;`.
   - `'` debe convertirse en `&apos;`.

# --hints--

Debes tener una función `convertHTML`.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` debería devolver la cadena `Dolce &amp; Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` debería devolver la cadena `Hamburgers &lt; Pizza &lt; Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` debería devolver la cadena `Sixty &gt; twelve`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` debería devolver la cadena `Stuff in &quot;quotation marks&quot;`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` debería devolver la cadena `Schindler&apos;s List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` debería devolver la cadena `&lt;&gt;`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` debe devolver la cadena `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
