---
id: 6814d8e1516e86b171929de4
title: "Sfida 1: equilibrio delle vocali"
challengeType: 28
dashedName: challenge-1
---

# --description--

Data una stringa, determina se il numero di vocali nella prima metà della stringa è uguale al numero di vocali nella seconda metà.

- La stringa può contenere qualsiasi carattere.
- Le lettere `a`, `e`, `i`, `o` e `u`, sia in maiuscolo sia in minuscolo, sono considerate vocali.
- Se la stringa ha un numero dispari di caratteri, ignora il carattere centrale.

# --hints--

`isBalanced("racecar")` dovrebbe restituire `true`.

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` dovrebbe restituire `true`.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` dovrebbe restituire `false`.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` dovrebbe restituire `false`.

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` dovrebbe restituire `true`.

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` dovrebbe restituire `false`.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` dovrebbe restituire `true`.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const half = Math.floor(s.length / 2);

  let firstHalf = s.slice(0, half);
  let secondHalf = s.length % 2 === 0 ? s.slice(half) : s.slice(half + 1);

  const countVowels = str =>
    str
      .toLowerCase()
      .split('')
      .filter(c => vowels.includes(c))
      .length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
