---
id: 681cb1b0dab50c87ddb2e518
title: "Sfida 7: somma mirata"
challengeType: 28
dashedName: challenge-7
---

# --description--

Dato un array di numeri e un intero obiettivo, trova due numeri unici nell'array che sommati diano il valore obiettivo. Restituisci un array con gli indici di quei due numeri, oppure `"Target not found"` se nessuna coppia di numeri somma al valore obiettivo.

- L'array restituito deve avere gli indici in ordine crescente.

# --hints--

`findTarget([2, 7, 11, 15], 9)` dovrebbe restituire `[0, 1]`.

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)` dovrebbe restituire `[1, 2]`.

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)` dovrebbe restituire `[4, 5]`.

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)` dovrebbe restituire `"Target not found"`.

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
