---
id: 681cb1a2dab50c87ddb2e514
title: "Sfida 3: sequenza di Fibonacci"
challengeType: 28
dashedName: challenge-3
---

# --description--

La sequenza di Fibonacci è una serie di numeri in cui ogni numero è la somma dei due precedenti. Partendo da `0` e `1`, i primi 10 numeri della sequenza sono `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

Dato un array contenente i primi due numeri di una sequenza di Fibonacci e un intero che rappresenta la lunghezza della sequenza, restituisci un array contenente la sequenza della lunghezza indicata.

- La tua funzione dovrebbe gestire sequenze di qualsiasi lunghezza maggiore o uguale a zero.
- Se la lunghezza è zero, restituisci un array vuoto.
- Nota che i numeri iniziali fanno parte della sequenza.

# --hints--

`fibonacciSequence([0, 1], 20)` dovrebbe restituire `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)` dovrebbe restituire `[21]`.

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)` dovrebbe restituire `[]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)` dovrebbe restituire `[10, 20]`.

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)` dovrebbe restituire `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
