---
id: 681cb1b0dab50c87ddb2e51b
title: "Sfida 10: 3 tentativi"
challengeType: 28
dashedName: challenge-10
---

# --description--

Dato un numero intero tra 1 e 10.000, restituisci il conteggio di quanti numeri da 1 fino a quel numero hanno il quadrato che contiene almeno una cifra 3.

# --hints--

`squaresWithThree(1)` dovrebbe restituire `0`.

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` dovrebbe restituire `1`.

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` dovrebbe restituire `19`.

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` dovrebbe restituire `326`.

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` dovrebbe restituire `4531`.

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
