---
id: 681cb05adab50c87ddb2e513
title: "Sfida 2: verifica della base"
challengeType: 28
dashedName: challenge-2
---

# --description--

Data una stringa che rappresenta un numero e una base intera da 2 a 36, determina se il numero è valido in quella base.

- La stringa può contenere numeri interi e caratteri maiuscoli o minuscoli.
- Il controllo deve essere case-insensitive.
- La base può essere qualsiasi numero da 2 a 36.
- Un numero è valido se ogni carattere è una cifra valida nella base data.
- Esempi di cifre valide per le basi:
  - Base 2: 0-1
  - Base 8: 0-7
  - Base 10: 0-9
  - Base 16: 0-9 e A-F
  - Base 36: 0-9 e A-Z

# --hints--

`isValidNumber("10101", 2)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` dovrebbe restituire `false`.

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` dovrebbe restituire `false`.

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` dovrebbe restituire `false`.

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` dovrebbe restituire `false`.

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` dovrebbe restituire `false`.

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` dovrebbe restituire `true`.

```js
assert.isTrue(isValidNumber("z", 36))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
