---
id: 681cb1b0dab50c87ddb2e519
title: "Sfida 8: calcolo del fattoriale"
challengeType: 28
dashedName: challenge-8
---

# --description--

Dato un numero intero da zero a 20, restituisci il fattoriale di quel numero. Il fattoriale di un numero è il prodotto di tutti i numeri compresi tra 1 e il numero dato.

- Il fattoriale di zero è 1.

# --hints--

`factorial(0)` dovrebbe restituire `1`.

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` dovrebbe restituire `120`.

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` dovrebbe restituire `2432902008176640000`.

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
