---
id: 681cb1b0dab50c87ddb2e51a
title: "Sfida 9: somma dei quadrati"
challengeType: 28
dashedName: challenge-9
---

# --description--

Dato un numero intero positivo fino a 1.000, restituisci la somma di tutti i numeri interi elevati al quadrato da 1 fino a quel numero.

# --hints--

`sumOfSquares(5)` dovrebbe restituire `55`.

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` dovrebbe restituire `385`.

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` dovrebbe restituire `5525`.

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` dovrebbe restituire `41791750`.

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` dovrebbe restituire `333833500`.

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
