---
id: 681cb1afdab50c87ddb2e515
title: "Sfida 4: s p a c e j a m"
challengeType: 28
dashedName: challenge-4
---

# --description--

Data una stringa, rimuovi tutti gli spazi dalla stringa, inserisci due spazi tra ogni carattere, trasforma tutte le lettere alfabetiche in maiuscolo e restituisci il risultato.

- I caratteri non alfabetici devono rimanere invariati (eccetto gli spazi).

# --hints--

`spaceJam("freeCodeCamp")` dovrebbe restituire `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` dovrebbe restituire `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")` dovrebbe restituire `"H  E  L  L  O  W  O  R  L  D  ?  !"`.

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` dovrebbe restituire `"C  @  T  $  &  D  0  G  $"`.

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` dovrebbe restituire `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
