---
id: 681cb1afdab50c87ddb2e517
title: "Sfida 6: verifica anagrammi"
challengeType: 28
dashedName: challenge-6
---

# --description--

Date due stringhe, determina se sono anagrammi l’una dell’altra (contengono gli stessi caratteri in qualsiasi ordine).

- Ignora maiuscole, minuscole e spazi bianchi.

# --hints--

`areAnagrams("listen", "silent")` dovrebbe restituire `true`.

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` dovrebbe restituire `true`.

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` dovrebbe restituire `true`.

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` dovrebbe restituire `false`.

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` dovrebbe restituire `false`.

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` dovrebbe restituire `false`.

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
