---
id: 681cb1afdab50c87ddb2e516
title: "Sfida 5: testo jbelmud"
challengeType: 28
dashedName: challenge-5
---

# --description--

Data una stringa, restituisci una versione mescolata di quella stringa in cui ogni parola viene trasformata seguendo queste regole:

- La prima e l'ultima lettera delle parole rimangono al loro posto
- Tutte le lettere tra la prima e l'ultima lettera sono ordinate alfabeticamente.
- Le stringhe di input non conterranno punteggiatura e saranno interamente in minuscolo.

# --hints--

`jbelmu("hello world")` dovrebbe restituire `"hello wlord"`.

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` dovrebbe restituire `"i love jbelmud text"`.

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` dovrebbe restituire `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` dovrebbe restituire `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
