---
id: 6821ebdf237de8297eaee793
title: "Sfida 17: ordine non convenzionale delle operazioni"
challengeType: 28
dashedName: challenge-17
---

# --description--

Dato un array di numeri interi e un array di operatori stringa, applica le operazioni ai numeri in sequenza da sinistra a destra. Ripeti le operazioni secondo necessità finché tutti i numeri non sono stati usati. Restituisci il risultato finale.

Per esempio, dati `[1, 2, 3, 4, 5]` e `['+', '*']`, restituisci il risultato della valutazione di `1 + 2 * 3 + 4 * 5` da sinistra a destra ignorando il normale ordine delle operazioni.

- Gli operatori validi sono `+`, `-`, `*`, `/` e `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` dovrebbe restituire `3`

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` dovrebbe restituire `38`

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` dovrebbe restituire `60`

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` dovrebbe restituire `30`

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` dovrebbe restituire `-2`

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
