---
id: 6821ebce237de8297eaee790
title: "Sfida 14: battaglia tra caratteri"
challengeType: 28
dashedName: challenge-14
---

# --description--

Date due stringhe che rappresentano il tuo esercito e un esercito avversario, ogni carattere del tuo esercito combatte contro il carattere nella stessa posizione dell'esercito avversario seguendo queste regole:

- I caratteri `a-z` hanno una forza di `1-26`, rispettivamente.
- I caratteri `A-Z` hanno una forza di `27-52`, rispettivamente.
- Le cifre `0-9` hanno una forza pari al loro valore facciale.
- Tutti gli altri caratteri hanno un valore pari a zero.
- Ogni carattere può combattere solo una battaglia.

Per ogni battaglia, vince il carattere più forte. L'esercito con più vittorie vince la guerra. Restituisci i seguenti valori:

- `"Opponent retreated"` se il tuo esercito ha più caratteri dell'esercito avversario.
- `"We retreated"` se l'esercito avversario ha più caratteri del tuo.
- `"We won"` se il tuo esercito ha vinto più battaglie.
- `"We lost"` se l'esercito avversario ha vinto più battaglie.
- `"It was a tie"` se entrambi gli eserciti hanno vinto lo stesso numero di battaglie.

# --hints--

`battle("Hello", "World")` dovrebbe restituire `"We lost"`.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` dovrebbe restituire `"We won"`.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` dovrebbe restituire `"We won"`.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` dovrebbe restituire `"Opponent retreated"`.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` dovrebbe restituire `"We retreated"`.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` dovrebbe restituire `"It was a tie"`.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` dovrebbe restituire `"It was a tie"`.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
