---
id: 6821ebfd237de8297eaee799
title: "Sfida 23: da RGB a esadecimale"
challengeType: 28
dashedName: challenge-23
---

# --description--

Dato un colore CSS `rgb(r, g, b)` come stringa, restituisci il suo equivalente esadecimale.

Ecco alcuni esempi di output per un dato input:

| Input   | Output   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- Trasforma tutte le lettere in minuscolo.
- Restituisci un `#` seguito da sei caratteri. Non usare valori abbreviati.

# --hints--

`rgbToHex("rgb(255, 255, 255)")` dovrebbe restituire `"#ffffff"`.

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` dovrebbe restituire `"#010b6f"`.

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` dovrebbe restituire `"#add8e6"`.

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` dovrebbe restituire `"#4f7bc9"`.

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
