---
id: 6821ebd4237de8297eaee791
title: "Sfida 15: camelCase"
challengeType: 28
dashedName: challenge-15
---

# --description--

Data una stringa, restituisci la sua versione camelCase seguendo queste regole:

- Le parole nell'argomento stringa sono separate da uno o più caratteri appartenenti al seguente insieme: spazio (` `), trattino (`-`) o underscore (`_`). Considera qualsiasi sequenza di questi come una separazione tra parole.
- La prima parola deve essere tutta in minuscolo.
- Ogni parola successiva deve iniziare con una lettera maiuscola, con il resto in minuscolo.
- Tutti gli spazi e i separatori devono essere rimossi.

# --hints--

`toCamelCase("hello world")` dovrebbe restituire `"helloWorld"`.

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` dovrebbe restituire `"helloWorld"`.

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")` dovrebbe restituire `"secretAgentX"`.

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` dovrebbe restituire `"freeCodeCamp"`.

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` dovrebbe restituire `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
