---
id: 6821ebf8237de8297eaee798
title: "Sfida 22: sequenza tribonacci"
challengeType: 28
dashedName: challenge-22
---

# --description--

La sequenza tribonacci è una serie di numeri in cui ogni numero è la somma dei tre precedenti. Partendo da `0`, `0` e `1`, i primi 10 numeri della sequenza sono `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24`, `44`.

Dato un array contenente i primi tre numeri di una sequenza tribonacci e un intero che rappresenta la lunghezza della sequenza, restituisci un array contenente la sequenza della lunghezza indicata.

- La tua funzione dovrebbe gestire sequenze di qualsiasi lunghezza maggiore o uguale a zero.
- Se la lunghezza è zero, restituisci un array vuoto.
- Nota che i numeri iniziali fanno parte della sequenza.

# --hints--

`tribonacciSequence([0, 0, 1], 20)` dovrebbe restituire `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)` dovrebbe restituire `[21]`.


```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)` dovrebbe restituire `[]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)` dovrebbe restituire `[10, 20]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)` dovrebbe restituire `[10, 20, 30]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)` dovrebbe restituire `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
