---
id: 6821ebea237de8297eaee795
title: "Sfida 19: luce di candela"
challengeType: 28
dashedName: challenge-19
---

# --description--

Dato un numero intero che rappresenta il numero di candele con cui inizi, e un numero intero che indica quante candele bruciate servono per crearne una nuova, restituisci il numero di candele che avrai usato dopo aver creato e bruciato quante più puoi.

Per esempio, se hai 7 candele e servono 2 candele bruciate per farne una nuova:

1. Brucia 7 candele per ottenere 7 avanzi,
2. Ricicla 6 avanzi per ottenere 3 nuove candele (rimane 1 avanzo),
3. Brucia 3 candele per ottenere altri 3 avanzi (4 in totale),
4. Ricicla 4 avanzi per ottenere 2 nuove candele,
5. Brucia 2 candele per ottenere 2 avanzi,
6. Ricicla 2 avanzi per ottenere 1 nuova candela,
7. Brucia 1 candela.

Avrai bruciato in totale 13 candele nell’esempio.

# --hints--

`burnCandles(7, 2)` dovrebbe restituire `13`

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` dovrebbe restituire `12`

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` dovrebbe restituire `29`

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` dovrebbe restituire `22`

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` dovrebbe restituire `3517`

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```
