---
id: 6821ebee237de8297eaee796
title: "Sfida 20: duplicati in un array"
challengeType: 28
dashedName: challenge-20
---

# --description--

Dato un array di numeri interi, restituisci un array di numeri interi che compaiono più di una volta nell'array iniziale, ordinati in ordine crescente. Se nessun valore compare più di una volta, restituisci un array vuoto.

- Includi solo una istanza di ogni valore nell'array restituito.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` dovrebbe restituire `[]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` dovrebbe restituire `[1, 2]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` dovrebbe restituire `[-6, 0, 2, 4, 5, 23]`.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
