---
id: 6821ebc9237de8297eaee78f
title: "Sfida 13: numero primo non naturale"
challengeType: 28
dashedName: challenge-13
---

# --description--

Dato un numero intero, determina se quel numero è un numero primo o un numero primo negativo.

- Un numero primo è un intero positivo maggiore di 1 che è divisibile solo per 1 e per se stesso.
- Un numero primo negativo è la versione negativa di un numero primo positivo.
- `1` e `0` non sono considerati numeri primi.

# --hints--

`isUnnaturalPrime(1)` dovrebbe restituire `false`.

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` dovrebbe restituire `false`.

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` dovrebbe restituire `true`.

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` dovrebbe restituire `true`.

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` dovrebbe restituire `false`.

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` dovrebbe restituire `true`.

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` dovrebbe restituire `true`.

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` dovrebbe restituire `false`.

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` dovrebbe restituire `false`.

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
