---
id: 68216ef80f957572e7c340c5
title: "Sfida 12: decodificatore di messaggi"
challengeType: 28
dashedName: challenge-12
---

# --description--

Data una stringa di messaggio segreto e un numero intero che rappresenta il numero di lettere con cui il messaggio è stato spostato per codificarlo, restituisci la stringa decodificata.

- Un numero positivo significa che il messaggio è stato spostato in avanti nell'alfabeto.
- Un numero negativo significa che il messaggio è stato spostato indietro nell'alfabeto.
- La distinzione tra maiuscole e minuscole è importante, i caratteri decodificati devono mantenere la stessa maiuscola o minuscola dei corrispondenti codificati.
- I caratteri non alfabetici non devono essere decodificati.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` dovrebbe restituire `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` dovrebbe restituire `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` dovrebbe restituire `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` dovrebbe restituire `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
