---
id: 6821ebe4237de8297eaee794
title: "Sfida 18: il secondo migliore"
challengeType: 28
dashedName: challenge-18
---

# --description--

Dato un array di numeri interi che rappresentano il prezzo di diversi laptop e un numero intero che rappresenta il tuo budget, restituisci:

1. Il secondo laptop più costoso se rientra nel tuo budget, oppure
2. Il laptop più costoso che rientra nel tuo budget, oppure
3. `0` se nessun laptop rientra nel tuo budget.

- I prezzi duplicati devono essere ignorati.

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)` dovrebbe restituire `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)` dovrebbe restituire `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)` dovrebbe restituire `1899`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)` dovrebbe restituire `0`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` dovrebbe restituire `1400`

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
