---
id: 68adce01c0e1144d0a902956
title: "Sfida 25: ripetitore di vocali"
challengeType: 28
dashedName: challenge-25
---

# --description--

Data una stringa, restituisci una nuova versione della stringa in cui ogni vocale è duplicata un numero di volte maggiore di uno rispetto alla vocale precedente incontrata. Per esempio, la prima vocale nella frase rimane invariata. La seconda vocale deve apparire due volte di seguito. La terza vocale deve apparire tre volte di seguito, e così via.

- Le lettere `a`, `e`, `i`, `o` e `u`, sia in maiuscolo che in minuscolo, sono considerate vocali.
- La vocale originale deve mantenere il suo caso.
- Le vocali ripetute devono essere in minuscolo.
- Tutti i caratteri non vocali devono mantenere il loro caso originale.

# --hints--

`repeatVowels("hello world")` dovrebbe restituire `"helloo wooorld"`.

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` dovrebbe restituire `"freeeCooodeeeeCaaaaamp"`.

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` dovrebbe restituire `"AEeIiiOoooUuuuu"`.

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` dovrebbe restituire `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
