---
id: 68adce01c0e1144d0a90295e
title: "Sfida 29: generatore di acronimi"
challengeType: 28
dashedName: challenge-29
---

# --description--

Dato una stringa contenente una o più parole, restituisci un acronimo delle parole usando le seguenti regole:

- L'acronimo deve consistere nella prima lettera di ogni parola in maiuscolo, a meno che non sia indicato diversamente.
- L'acronimo deve ignorare la prima lettera di queste parole a meno che non siano la prima parola della stringa data: `a`, `for`, `an`, `and`, `by` e `of`.
- Le lettere dell'acronimo devono essere restituite nell'ordine in cui appaiono.
- L'acronimo non deve contenere spazi.

# --hints--

`buildAcronym("Search Engine Optimization")` dovrebbe restituire `"SEO"`.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` dovrebbe restituire `"FAQ"`.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` dovrebbe restituire `"NASA"`.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` dovrebbe restituire `"FBI"`.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` dovrebbe restituire `"FYI"`.

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")` dovrebbe restituire `"BTW"`.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` dovrebbe restituire `"AUHWPOTIMSH"`.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
