---
id: 68adce01c0e1144d0a902958
title: "Sfida 26: validatore IPv4"
challengeType: 28
dashedName: challenge-26
---

# --description--

Data una stringa, determina se è un indirizzo IPv4 valido. Un indirizzo IPv4 valido è composto da quattro numeri interi separati da punti (`.`). Ogni numero deve soddisfare le seguenti condizioni:

- Deve essere compreso tra 0 e 255 inclusi.
- Non deve avere zeri iniziali (ad esempio 0 è consentito, 01 no).
- Sono consentiti solo caratteri numerici.

# --hints--

`isValidIPv4("192.168.1.1")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
