---
id: 68adce01c0e1144d0a90295c
title: "Sfida 28: parser di numeri romani"
challengeType: 28
dashedName: challenge-28
---

# --description--

Dato una stringa che rappresenta un numero romano, restituisci il suo valore intero.

I numeri romani sono composti dai seguenti simboli e valori:

| Simbolo | Valore |
|---------|--------|
| I       | 1      |
| V       | 5      |
| X       | 10     |
| L       | 50     |
| C       | 100    |
| D       | 500    |
| M       | 1000   |

- I numeri si leggono da sinistra a destra. Se un numero più piccolo precede uno più grande, il valore viene sottratto. Altrimenti, i valori si sommano.

# --hints--

`parseRomanNumeral("III")` dovrebbe restituire `3`.

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` dovrebbe restituire `4`.

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` dovrebbe restituire `26`.

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` dovrebbe restituire `99`.

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` dovrebbe restituire `460`.

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` dovrebbe restituire `504`.

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` dovrebbe restituire `2025`.

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
