---
id: 68af0687ef34c76c28ffa54d
title: "Sfida 33: tempo davanti allo schermo"
challengeType: 28
dashedName: challenge-33
---

# --description--

Dato un array di sette numeri interi, che rappresentano il tempo di una settimana, dove ogni numero intero è la quantità di ore trascorse al telefono quel giorno, determina se è troppo tempo davanti allo schermo basandoti su questi vincoli:

- Se un singolo giorno ha 10 ore o più, è troppo.
- Se la media di tre giorni consecutivi è maggiore o uguale a 8 ore, è troppo.
- Se la media dei sette giorni è maggiore o uguale a 6 ore, è troppo.

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` dovrebbe restituire `false`.

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` dovrebbe restituire `false`.

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` dovrebbe restituire `false`.

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` dovrebbe restituire `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` dovrebbe restituire `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` dovrebbe restituire `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` dovrebbe restituire `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```
