---
id: 68af0687ef34c76c28ffa547
title: "Sfida 30: caratteri unici"
challengeType: 28
dashedName: challenge-30
---

# --description--

Data una stringa, determina se tutti i caratteri nella stringa sono unici.

- Le lettere maiuscole e minuscole devono essere considerate caratteri diversi.

# --hints--

`allUnique("abc")` dovrebbe restituire `true`.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` dovrebbe restituire `true`.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` dovrebbe restituire `true`.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` dovrebbe restituire `true`.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` dovrebbe restituire `false`.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` dovrebbe restituire `false`.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` dovrebbe restituire `false`.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
