---
id: 68af0687ef34c76c28ffa54b
title: "Sfida 32: inverti la frase"
challengeType: 28
dashedName: challenge-32
---

# --description--

Data una stringa di parole, restituisci una nuova stringa con le parole in ordine inverso. Per esempio, la prima parola dovrebbe trovarsi alla fine della stringa restituita e l'ultima parola all'inizio della stringa restituita.

- Nella stringa data, le parole possono essere separate da uno o più spazi.
- La stringa restituita dovrebbe avere solo uno spazio tra le parole.

# --hints--

`reverseSentence("world hello")` dovrebbe restituire `"hello world"`.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` dovrebbe restituire `"git commit push"`.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` dovrebbe restituire `"sudo install npm"`.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` dovrebbe restituire `"export function default import"`.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
