---
id: 68b06e589bf2273243814771
title: "Sfida 36: regolatore del termostato"
challengeType: 28
dashedName: challenge-36
---

# --description--

Data la temperatura attuale di una stanza e una temperatura obiettivo, restituisci una stringa che indichi come regolare la temperatura della stanza in base a questi vincoli:

- Restituisci `"heat"` se la temperatura attuale è inferiore all'obiettivo.
- Restituisci `"cool"` se la temperatura attuale è superiore all'obiettivo.
- Restituisci `"hold"` se la temperatura attuale è uguale all'obiettivo.

# --hints--

`adjustThermostat(68, 72)` dovrebbe restituire `"heat"`.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` dovrebbe restituire `"cool"`.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` dovrebbe restituire `"hold"`.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` dovrebbe restituire `"heat"`.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` dovrebbe restituire `"cool"`.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` dovrebbe restituire `"hold"`.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
