---
id: 68b06e589bf227324381476f
title: "Sfida 35: frequenza delle parole"
challengeType: 28
dashedName: challenge-35
---

# --description--

Dato un paragrafo, restituisci un array con le tre parole che compaiono più frequentemente.

- Le parole nel paragrafo saranno separate da spazi.
- Ignora le maiuscole e minuscole nel paragrafo dato. Per esempio, considera `Hello` e `hello` come la stessa parola.
- Ignora la punteggiatura nel paragrafo dato. La punteggiatura consiste in virgole (`,`), punti (`.`) e punti esclamativi (`!`).
- L'array restituito deve contenere tutte parole in minuscolo.
- L'array restituito deve essere in ordine decrescente, con la parola che compare più frequentemente per prima.

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` dovrebbe restituire `["coding", "python", "in"]`.

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` dovrebbe restituire `["i", "like", "coding"]`.

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` dovrebbe restituire `["debug", "test", "deploy"]`.

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
