---
id: 68b06e589bf2273243814777
title: "Sfida 39: riempi il serbatoio"
challengeType: 28
dashedName: challenge-39
---

# --description--

Dato il volume di un serbatoio di carburante, il livello attuale di carburante e il prezzo per gallone, restituisci il costo per riempire completamente il serbatoio.

- `tankSize` è la capacità totale del serbatoio in galloni.
- `fuelLevel` è la quantità attuale di carburante nel serbatoio in galloni.
- `pricePerGallon` è il costo di un gallone di carburante.
- Il valore restituito dovrebbe essere arrotondato a due cifre decimali nel formato: `"$d.dd"`.

# --hints--

`costToFill(20, 0, 4.00)` dovrebbe restituire `"$80.00"`.

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` dovrebbe restituire `"$17.50"`.

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` dovrebbe restituire `"$29.25"`.

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` dovrebbe restituire `"$0.00"`.

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` dovrebbe restituire `"$21.89"`.

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
