---
id: 68b06e589bf2273243814775
title: "Sfida 38: generatore di slug"
challengeType: 28
dashedName: challenge-38
---

# --description--

Data una stringa, restituisci una versione della stringa adatta a un URL usando le seguenti regole:

- Tutte le lettere devono essere minuscole.
- Tutti i caratteri che non sono lettere, numeri o spazi devono essere rimossi.
- Tutti gli spazi devono essere sostituiti con il codice spazio codificato per URL `%20`.
- Spazi consecutivi devono essere sostituiti con un singolo `%20`.
- La stringa restituita non deve avere `%20` all'inizio o alla fine.

# --hints--

`generateSlug("helloWorld")` dovrebbe restituire `"helloworld"`.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` dovrebbe restituire `"hello%20world"`.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` dovrebbe restituire `"helloworld"`.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` dovrebbe restituire `"hello%20world"`.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` dovrebbe restituire `"h3110%20w0r1d"`.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
