---
id: 68b1f72371a5ac895ac70a0a
title: "Sfida 44: specchio di stringa"
challengeType: 28
dashedName: challenge-44
---

# --description--

Date due stringhe, determina se la seconda stringa è lo specchio della prima.

- Una stringa è considerata uno specchio se contiene le stesse lettere in ordine inverso.
- Considera le lettere maiuscole e minuscole come distinte.
- Ignora tutti i caratteri non alfabetici.

# --hints--

`isMirror("helloworld", "helloworld")` dovrebbe restituire `false`.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` dovrebbe restituire `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` dovrebbe restituire `true`.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` dovrebbe restituire `false`.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` dovrebbe restituire `false`.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` dovrebbe restituire `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` dovrebbe restituire `true`.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
