---
id: 68b1f72371a5ac895ac70a04
title: "Sfida 41: archivio file"
challengeType: 28
dashedName: challenge-41
---

# --description--

Dato un file size, un'unità per il file size e la capacità del disco rigido in gigabyte (GB), restituisci il numero di file che il disco rigido può archiviare usando le seguenti condizioni:

- L'unità per il file size può essere byte (`"B"`), kilobyte (`"KB"`) o megabyte (`"MB"`).
- Restituisci il numero di file interi che il disco può contenere.
- Usa le seguenti conversioni:

| Unità | Equivalente |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

Per esempio, dati `500`, `"KB"` e `1` come argomenti, determina quanti file da 500 KB possono entrare in un disco rigido da 1 GB.

# --hints--

`numberOfFiles(500, "KB", 1)` dovrebbe restituire `2000`.

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` dovrebbe restituire `20000`.

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` dovrebbe restituire `200`.

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` dovrebbe restituire `366210`.

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` dovrebbe restituire `453514`.

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` dovrebbe restituire `166666`.

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
