---
id: 68b1f72371a5ac895ac70a08
title: "Sfida 43: cifre contro lettere"
challengeType: 28
dashedName: challenge-43
---

# --description--

Data una stringa, restituisci `"digits"` se la stringa contiene più cifre che lettere, `"letters"` se contiene più lettere che cifre, e `"tie"` se ha lo stesso numero di cifre e lettere.

- Le cifre sono costituite da `0-9`.
- Le lettere sono costituite da `a-z` in maiuscolo o minuscolo.
- Ignora tutti gli altri caratteri.

# --hints--

`digitsOrLetters("abc123")` dovrebbe restituire `"tie"`.

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` dovrebbe restituire `"letters"`.

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` dovrebbe restituire `"digits"`.

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` dovrebbe restituire `"letters"`.

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` dovrebbe restituire `"digits"`.

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` dovrebbe restituire `"tie"`.

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
