---
id: 68b1f72371a5ac895ac70a06
title: "Sfida 42: archiviazione video"
challengeType: 28
dashedName: challenge-42
---

# --description--

Dato un formato video, un'unità per la dimensione del video, una capacità del disco rigido e un'unità per il disco rigido, restituisci il numero di video che il disco rigido può archiviare usando le seguenti condizioni:

- L'unità per la dimensione del video può essere byte (`"B"`), kilobyte (`"KB"`), megabyte (`"MB"`) o gigabyte (`"GB"`).
- Se non viene fornita una delle unità video sopra, restituisci `"Invalid video unit"`.
- L'unità della capacità del disco rigido può essere gigabyte (`"GB"`) o terabyte (`"TB"`).
- Se non viene fornita una delle unità del disco rigido sopra, restituisci `"Invalid drive unit"`.
- Restituisci il numero intero di video che il disco può contenere.
- Usa le seguenti conversioni:

| Unità | Equivalente |
|:-----:|:-----------:|
| 1 B   | 1 B         |
| 1 KB  | 1000 B      |
| 1 MB  | 1000 KB     |
| 1 GB  | 1000 MB     |
| 1 TB  | 1000 GB     |

Per esempio, dati `500`, `"MB"`, `100` e `"GB"` come argomenti, determina quanti video da 500 MB possono entrare in un disco rigido da 100 GB.

# --hints--

`numberOfVideos(500, "MB", 100, "GB")` dovrebbe restituire `200`.

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")` dovrebbe restituire `"Invalid video unit"`.

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")` dovrebbe restituire `"Invalid drive unit"`.

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")` dovrebbe restituire `4000`.

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")` dovrebbe restituire `1466`.

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
