---
id: 68b1f72371a5ac895ac70a02
title: "Sfida 40: archivio di foto"
challengeType: 28
dashedName: challenge-40
---

# --description--

Dato un formato di foto in megabyte (MB) e la capacità di un disco rigido in gigabyte (GB), restituisci il numero di foto che il disco può archiviare usando le seguenti condizioni:

- 1 gigabyte equivale a 1000 megabyte.
- Restituisci il numero intero di foto che il disco può archiviare.

# --hints--

`numberOfPhotos(1, 1)` dovrebbe restituire `1000`.

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)` dovrebbe restituire `500`.

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)` dovrebbe restituire `64000`.

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)` dovrebbe restituire `214285`.

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)` dovrebbe restituire `1571`.

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
