---
id: 68b7cadffed0e75a517da673
title: "Sfida 52: da binario a decimale"
challengeType: 28
dashedName: challenge-52
---

# --description--

Dato una stringa che rappresenta un numero binario, restituisci il suo equivalente decimale come numero.

Un numero binario usa solo le cifre `0` e `1` per rappresentare qualsiasi numero. Per convertire un numero binario in decimale, moltiplica ogni cifra per una potenza di `2` e somma i risultati. Inizia moltiplicando la cifra più a destra per `2^0`, la cifra successiva a sinistra per `2^1` e così via. Una volta moltiplicate tutte le cifre per una potenza di `2`, somma il risultato.

Per esempio, il numero binario `101` equivale a `5` in decimale perché:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")` dovrebbe restituire `5`.

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")` dovrebbe restituire `10`.

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")` dovrebbe restituire `18`.

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")` dovrebbe restituire `85`.

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
