---
id: 68b7cadffed0e75a517da677
title: "Sfida 54: forza della p@ssw0rd!"
challengeType: 28
dashedName: challenge-54
---

# --description--

Data una stringa password, restituisci `"weak"`, `"medium"` o `"strong"` in base alla forza della password.

Una password viene valutata secondo le seguenti regole:

- Deve essere lunga almeno 8 caratteri.
- Deve contenere sia lettere maiuscole che minuscole.
- Deve contenere almeno un numero.
- Deve contenere almeno un carattere speciale da questo insieme: `!`, `@`, `#`, `$`, `%`, `^`, `&` o `*`.

Restituisci `"weak"` se la password soddisfa meno di due regole.
Restituisci `"medium"` se la password soddisfa 2 o 3 regole.
Restituisci `"strong"` se la password soddisfa tutte e 4 le regole.

# --hints--

`checkStrength("123456")` dovrebbe restituire `"weak"`.

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")` dovrebbe restituire `"weak"`.

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")` dovrebbe restituire `"weak"`.

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("PASSWORD")` dovrebbe restituire `"weak"`.

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")` dovrebbe restituire `"medium"`.

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")` dovrebbe restituire `"medium"`.

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")` dovrebbe restituire `"medium"`.

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")` dovrebbe restituire `"strong"`.

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")` dovrebbe restituire `"strong"`.

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
