---
id: 68b7cadffed0e75a517da675
title: "Sfida 53: da decimale a binario"
challengeType: 28
dashedName: challenge-53
---

# --description--

Dato un numero intero non negativo, restituisci la sua rappresentazione binaria come stringa.

Un numero binario usa solo le cifre `0` e `1` per rappresentare qualsiasi numero. Per convertire un numero decimale in binario, dividi ripetutamente il numero per `2` e registra il resto. Ripeti finché il numero non è zero. Leggi i resti dall'ultimo registrato al primo. Per esempio, per convertire `12` in binario:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` in binario è `1100`.

# --hints--

`toBinary(5)` dovrebbe restituire `"101"`.

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)` dovrebbe restituire `"1100"`.

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)` dovrebbe restituire `"110010"`.

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)` dovrebbe restituire `"1100011"`.

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
