---
id: 68b7687dded630607aceccb1
title: "Sfida 48: rilevatore di spam"
challengeType: 28
dashedName: challenge-48
---

# --description--

Dato un numero di telefono nel formato `"+A (BBB) CCC-DDDD"`, dove ogni lettera rappresenta una cifra come segue:

- `A` rappresenta il prefisso internazionale e può essere composto da un numero qualsiasi di cifre.
- `BBB` rappresenta il prefisso dell'area ed è sempre composto da tre cifre.
- `CCC` e `DDDD` rappresentano il numero locale e sono sempre lunghi rispettivamente tre e quattro cifre.

Determina se si tratta di un numero spam in base ai seguenti criteri:

- Il prefisso internazionale è più lungo di 2 cifre oppure non inizia con uno zero (`0`).
- Il prefisso dell'area è maggiore di 900 o minore di 200.
- La somma delle prime tre cifre del numero locale appare all'interno delle ultime quattro cifre del numero locale.
- Il numero contiene la stessa cifra ripetuta quattro o più volte di seguito (ignorando i caratteri di formattazione).

# --hints--

`isSpam("+0 (200) 234-0182")` dovrebbe restituire `false`.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` dovrebbe restituire `true`.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` dovrebbe restituire `true`.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` dovrebbe restituire `true`.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` dovrebbe restituire `true`.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` dovrebbe restituire `true`.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` dovrebbe restituire `true`.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` dovrebbe restituire `false`.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
