---
id: 68b7687dded630607aceccb3
title: "Sfida 49: parser di intestazioni CSV"
challengeType: 28
dashedName: challenge-49
---

# --description--

Data la prima riga di un file CSV (valori separati da virgola), restituisci un array contenente le intestazioni.

- La prima riga di un file CSV contiene intestazioni separate da virgole.
- Rimuovi eventuali spazi bianchi all'inizio o alla fine di ogni intestazione.

# --hints--

`getHeadings("name,age,city")` dovrebbe restituire `["name", "age", "city"]`.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` dovrebbe restituire `["first name", "last name", "phone"]`.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` dovrebbe restituire `["username", "email", "signup date"]`.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
