---
id: 68b7687dded630607aceccab
title: "Sfida 45: quadrato perfetto"
challengeType: 28
dashedName: challenge-45
---

# --description--

Dato un numero intero, determina se è un quadrato perfetto.

- Un numero è un quadrato perfetto se puoi moltiplicare un intero per se stesso per ottenerlo. Per esempio, 9 è un quadrato perfetto perché puoi moltiplicare 3 per se stesso per ottenerlo.

# --hints--

`isPerfectSquare(9)` dovrebbe restituire `true`.

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` dovrebbe restituire `true`.

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` dovrebbe restituire `true`.

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` dovrebbe restituire `false`.

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` dovrebbe restituire `false`.

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` dovrebbe restituire `false`.

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` dovrebbe restituire `true`.

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` dovrebbe restituire `true`.

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
