---
id: 68c1a929005bf54d342aa8d2
title: "Sfida 55: settimana dello spazio giorno 1: classificazione stellare"
challengeType: 28
dashedName: challenge-55
---

# --description--

Il 4 ottobre segna l'inizio della settimana mondiale dello spazio. I prossimi sette giorni ti proporranno sfide di programmazione a tema astronomia.

Per la sfida di oggi, ti viene data la temperatura superficiale di una stella in kelvin (K) e devi determinare la sua classificazione stellare in base ai seguenti intervalli:

- `"O"`: 30.000 K o superiore
- `"B"`: da 10.000 K a 29.999 K
- `"A"`: da 7.500 K a 9.999 K
- `"F"`: da 6.000 K a 7.499 K
- `"G"`: da 5.200 K a 5.999 K
- `"K"`: da 3.700 K a 5.199 K
- `"M"`: da 0 K a 3.699 K

- Restituisci la classificazione della stella data.

# --hints--

`classification(5778)` dovrebbe restituire `"G"`.

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` dovrebbe restituire `"M"`.

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` dovrebbe restituire `"A"`.

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` dovrebbe restituire `"K"`.

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` dovrebbe restituire `"M"`.

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` dovrebbe restituire `"O"`.

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` dovrebbe restituire `"F"`.

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` dovrebbe restituire `"B"`.

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
