---
id: 68c1a929005bf54d342aa8d6
title: "Sfida 59: settimana dello spazio giorno 5: zona abitabile"
challengeType: 28
dashedName: challenge-59
---

# --description--

Per il quinto giorno della settimana dello spazio, calcolerai la "zona abitabile" di una stella — la regione intorno a una stella dove le condizioni sono "giuste" per la presenza di acqua liquida.

Data la massa di una stella, restituisci un array con le distanze di inizio e fine della sua zona abitabile in unità astronomiche.

Per calcolare la zona abitabile:

1. Trova la luminosità della stella elevando la sua massa alla potenza di 3,5.
2. L'inizio della zona è 0,95 volte la radice quadrata della sua luminosità.
3. La fine della zona è 1,37 volte la radice quadrata della sua luminosità.

- Restituisci le distanze arrotondate a due cifre decimali.

Per esempio, data `1` come massa, restituisci `[0.95, 1.37]`.

# --hints--

`goldilocksZone(1)` dovrebbe restituire `[0.95, 1.37]`.

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` dovrebbe restituire `[0.28, 0.41]`.

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` dovrebbe restituire `[21.85, 31.51]`.

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` dovrebbe restituire `[9.38, 13.52]`.

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` dovrebbe restituire `[179.69, 259.13]`.

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
