---
id: 68c1a929005bf54d342aa8d3
title: "Sfida 56: settimana dello spazio giorno 2: ricerca di esopianeti"
challengeType: 28
dashedName: challenge-56
---

# --description--

Per il secondo giorno della settimana dello spazio, ti viene fornita una stringa in cui ogni carattere rappresenta la lettura di luminosità di una stella. Determina se le letture hanno rilevato un esopianeta usando il metodo del transito. Il metodo del transito si verifica quando un pianeta passa davanti a una stella, riducendo la sua luminosità osservata.

- Le letture di luminosità comprendono solo i caratteri `0-9` e `A-Z` dove ogni lettura corrisponde ai seguenti valori numerici:
- I caratteri `0-9` corrispondono ai livelli di luminosità `0-9`.
- I caratteri `A-Z` corrispondono ai livelli di luminosità `10-35`.

Una stella è considerata avere un esopianeta se una qualsiasi singola lettura è minore o uguale all'80% della media di tutte le letture. Per esempio, se la luminosità media di una stella è 10, sarà considerata avere un esopianeta se una singola lettura è 8 o meno.

# --hints--

`hasExoplanet("665544554")` dovrebbe restituire `false`.

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")` dovrebbe restituire `true`.

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")` dovrebbe restituire `false`.

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")` dovrebbe restituire `true`.

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")` dovrebbe restituire `false`.

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` dovrebbe restituire `true`.

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
