---
id: 68c1a929005bf54d342aa8d4
title: "Sfida 57: settimana dello spazio giorno 3: telefono a casa"
challengeType: 28
dashedName: challenge-57
---

# --description--

Per il terzo giorno della settimana dello spazio, ti viene fornito un array di numeri che rappresentano le distanze (in chilometri) tra te, i satelliti e il tuo pianeta natale in una rotta di comunicazione. Determina quanto tempo impiegherà un messaggio inviato lungo la rotta per raggiungere il pianeta di destinazione usando le seguenti condizioni:

- Il primo valore nell'array è la distanza dalla tua posizione al primo satellite.
- Ogni valore successivo, tranne l'ultimo, è la distanza fino al satellite successivo.
- L'ultimo valore nell'array è la distanza dall'ultimo satellite al tuo pianeta natale.
- Il messaggio viaggia a 300.000 km/s.
- Ogni satellite che il messaggio **attraversa** aggiunge un ritardo di trasmissione di 0,5 secondi.
- Restituisci un numero arrotondato a 4 cifre decimali, eliminando gli zeri finali.

# --hints--

`sendMessage([300000, 300000])` dovrebbe restituire `2.5`.

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` dovrebbe restituire `3.0627`.

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` dovrebbe restituire `364.5`.

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` dovrebbe restituire `1674.3333`.

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` dovrebbe restituire `2.4086`.

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` dovrebbe restituire `21.1597`.

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
