---
id: 68c497f3aaefc9fd9f1b0e26
title: "Sfida 62: da esadecimale a decimale"
challengeType: 28
dashedName: challenge-62
---

# --description--

Data una stringa che rappresenta un numero esadecimale (base 16), restituisci il suo valore decimale (base 10) come intero.

L'esadecimale è un sistema numerico che usa 16 cifre:

- `0-9` rappresentano i valori da `0` a `9`.
- `A-F` rappresentano i valori da `10` a `15`.

Ecco una tabella di conversione parziale:

| Esadecimale | Decimale |
|:-----------:|:--------:|
| 0           | 0        |
| 1           | 1        |
| ...         | ...      |
| 9           | 9        |
| A           | 10       |
| ...         | ...      |
| F           | 15       |
| 10          | 16       |
| ...         | ...      |
| 9F          | 159      |
| A0          | 160      |
| ...         | ...      |
| FF          | 255      |
| 100         | 256      |

- La stringa conterrà solo caratteri `0–9` e `A–F`.

# --hints--

`hexToDecimal("A")` dovrebbe restituire `10`.

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")` dovrebbe restituire `21`.

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")` dovrebbe restituire `46`.

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")` dovrebbe restituire `255`.

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")` dovrebbe restituire `2623`.

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```
