---
id: 68c497f3aaefc9fd9f1b0e24
title: "Sfida 60: settimana dello spazio giorno 6: fase della luna"
challengeType: 28
dashedName: challenge-60
---

# --description--

Per il sesto giorno della settimana dello spazio, ti verrà fornita una data nel formato `"YYYY-MM-DD"` e dovrai determinare la fase della luna per quel giorno usando le seguenti regole:

Usa un ciclo lunare semplificato di 28 giorni, diviso in quattro fasi uguali:

- `"New"`: giorni da 1 a 7
- `"Waxing"`: giorni da 8 a 14
- `"Full"`: giorni da 15 a 21
- `"Waning"`: giorni da 22 a 28

Dopo il giorno 28, il ciclo ricomincia con il giorno 1, luna nuova.

- Usa `"2000-01-06"` come luna nuova di riferimento (giorno 1 del ciclo) per determinare la fase del giorno dato.
- Non ti verranno date date precedenti alla data di riferimento.
- Restituisci la fase corretta come stringa.

**Nota:** Il giorno 1 rappresenta il giorno della luna nuova, cioè 0 giorni sono passati dall’ultima luna nuova.

# --hints--

`moonPhase("2000-01-12")` dovrebbe restituire `"New"`.

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` dovrebbe restituire `"Waxing"`.

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` dovrebbe restituire `"Full"`.

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` dovrebbe restituire `"Waning"`.

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` dovrebbe restituire `"New"`.

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
