---
id: 68cae5b538ff798bbd4da007
title: "Sfida 69: calzini mancanti"
challengeType: 28
dashedName: challenge-69
---

# --description--

Dato un numero intero che rappresenta il numero di paia di calzini con cui hai iniziato e un altro numero intero che rappresenta quanti cicli di lavaggio hai completato, restituisci il numero di paia complete di calzini che hai attualmente usando le seguenti regole:

- Ogni 2 cicli di lavaggio, perdi un singolo calzino.
- Ogni 3 cicli di lavaggio, trovi un singolo calzino mancante.
- Ogni 5 cicli di lavaggio, un singolo calzino si consuma e deve essere buttato.
- Ogni 10 cicli di lavaggio, compri una paio di calzini.
- Non puoi mai avere meno di zero calzini in totale.
- Le regole possono sovrapporsi. Per esempio, al ciclo di lavaggio 10, perderai un calzino, ne butterai uno e comprerai una nuova paio.
- Restituisci il numero di paia complete di calzini.

# --hints--

`sockPairs(2, 5)` dovrebbe restituire `1`.

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` dovrebbe restituire `0`.

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` dovrebbe restituire `4`.

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` dovrebbe restituire `3`.

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` dovrebbe restituire `0`.

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```
