---
id: 68cae5b538ff798bbd4da006
title: "Sfida 68: mascherare numeri di carta di credito"
challengeType: 28
dashedName: challenge-68
---

# --description--

Data una stringa di numeri di carta di credito, restituisci una versione mascherata usando le seguenti regole:

- La stringa conterrà quattro insiemi di quattro cifre (`0-9`), con tutti gli insiemi separati da uno spazio singolo o da un singolo trattino (`-`).
- Sostituisci tutti i numeri, tranne gli ultimi quattro, con un asterisco (`*`).
- Lascia invariati gli altri caratteri.

Per esempio, data `"4012-8888-8888-1881"` restituisci `"****-****-****-1881"`.

# --hints--

`mask("4012-8888-8888-1881")` dovrebbe restituire `"****-****-****-1881"`.

```js
assert.equal(mask("4012-8888-8888-1881"), "****-****-****-1881");
```

`mask("5105 1051 0510 5100")` dovrebbe restituire `"**** **** **** 5100"`.

```js
assert.equal(mask("5105 1051 0510 5100"), "**** **** **** 5100");
```

`mask("6011 1111 1111 1117")` dovrebbe restituire `"**** **** **** 1117"`.

```js
assert.equal(mask("6011 1111 1111 1117"), "**** **** **** 1117");
```

`mask("2223-0000-4845-0010")` dovrebbe restituire `"****-****-****-0010"`.

```js
assert.equal(mask("2223-0000-4845-0010"), "****-****-****-0010");
```

# --seed--

## --seed-contents--

```js
function mask(card) {

  return card;
}
```

# --solutions--

```js
function mask(card) {
  const split = card.split(card.includes('-') ? '-' : ' ');
  
  return card.includes('-') ?
    `****-****-****-${split[3]}` :
    `**** **** **** ${split[3]}`;
}
```
