---
id: 68cae5b538ff798bbd4da00a
title: "Sfida 72: regolatore del termostato 2"
challengeType: 28
dashedName: challenge-72
---

# --description--

Data la temperatura attuale di una stanza in gradi Fahrenheit e una temperatura obiettivo in gradi Celsius, restituisci una stringa che indichi come regolare la temperatura della stanza in base a queste condizioni:

- Restituisci `"Heat: X degrees Fahrenheit"` se la temperatura attuale è inferiore all'obiettivo. Con `X` che indica il numero di gradi in Fahrenheit da riscaldare per raggiungere l'obiettivo, arrotondato a 1 cifra decimale.
- Restituisci `"Cool: X degrees Fahrenheit"` se la temperatura attuale è superiore all'obiettivo. Con `X` che indica il numero di gradi in Fahrenheit da raffreddare per raggiungere l'obiettivo, arrotondato a 1 cifra decimale.
- Restituisci `"Hold"` se la temperatura attuale è uguale all'obiettivo.

Per convertire Celsius in Fahrenheit, moltiplica la temperatura in Celsius per 1,8 e aggiungi 32 al risultato (`F = (C * 1.8) + 32`).

# --hints--

`adjustThermostat(32, 0)` dovrebbe restituire `"Hold"`.

```js
assert.equal(adjustThermostat(32, 0), "Hold");
```

`adjustThermostat(70, 25)` dovrebbe restituire `"Heat: 7.0 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(70, 25), "Heat: 7.0 degrees Fahrenheit");
```

`adjustThermostat(72, 18)` dovrebbe restituire `"Cool: 7.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(72, 18), "Cool: 7.6 degrees Fahrenheit");
```

`adjustThermostat(212, 100)` dovrebbe restituire `"Hold"`.

```js
assert.equal(adjustThermostat(212, 100), "Hold");
```

`adjustThermostat(59, 22)` dovrebbe restituire `"Heat: 12.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(59, 22), "Heat: 12.6 degrees Fahrenheit");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(currentF, targetC) {

  return currentF;
}
```

# --solutions--

```js
function adjustThermostat(currentF, targetC) {
  const targetF = targetC * 1.8 + 32;
  const diff = Math.abs(targetF - currentF).toFixed(1);

  if (currentF < targetF) {
    return `Heat: ${diff} degrees Fahrenheit`;
  } else if (currentF > targetF) {
    return `Cool: ${diff} degrees Fahrenheit`;
  } else {
    return "Hold"
  }
}
```
