---
id: 68cae5b538ff798bbd4da008
title: "Sfida 70: estrattore di attributi HTML"
challengeType: 28
dashedName: challenge-70
---

# --description--

Dato un stringa di un elemento HTML valido, restituisci gli attributi dell'elemento seguendo questi criteri:

- Ti verrà fornito solo uno elemento.
- Gli attributi saranno nel formato: `attribute="value"`.
- Restituisci un array di stringhe con ogni proprietà e valore dell'attributo, separati da una virgola, in questo formato: `["attribute1, value1", "attribute2, value2"]`.
- Restituisci gli attributi nell'ordine in cui sono dati.
- Se non vengono trovati attributi, restituisci un array vuoto.

# --hints--

`extractAttributes('<span class="red"></span>')` dovrebbe restituire `["class, red"]`.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` dovrebbe restituire `["charset, UTF-8"]`.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` dovrebbe restituire `[]`.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` dovrebbe restituire `["name, email", "type, email", "required, true"]`.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` dovrebbe restituire `["id, submit", "class, btn btn-primary"]`.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
