---
id: 68cae5b538ff798bbd4da002
title: "Sfida 64: da 24 a 12"
challengeType: 28
dashedName: challenge-64
---

# --description--

Data una stringa che rappresenta un orario nel formato 24 ore di `"HHMM"`, restituisci l'orario nel suo formato equivalente a 12 ore di `"H:MM AM"` o `"H:MM PM"`.

- L'input fornito sarà sempre una stringa di quattro cifre nel formato orario a 24 ore, da `"0000"` a `"2359"`.

# --hints--

`to12("1124")` dovrebbe restituire `"11:24 AM"`.

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")` dovrebbe restituire `"9:00 AM"`.

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")` dovrebbe restituire `"2:55 PM"`.

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")` dovrebbe restituire `"11:46 PM"`.

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")` dovrebbe restituire `"12:30 AM"`.

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
