---
id: 68cae5b538ff798bbd4da005
title: "Sfida 67: validatore di email"
challengeType: 28
dashedName: challenge-67
---

# --description--

Data una stringa, determina se è un indirizzo email valido usando i seguenti vincoli:

- Deve contenere esattamente un simbolo `@`.
- La parte locale (prima del `@`):
  - Può contenere solo lettere (`a-z`, `A-Z`), cifre (`0-9`), punti (`.`), underscore (`_`) o trattini (`-`).
  - Non può iniziare o finire con un punto.
- La parte del dominio (dopo il `@`):
  - Deve contenere almeno un punto.
  - Deve terminare con un punto seguito da almeno due lettere.
- Né la parte locale né quella del dominio possono avere due punti consecutivi.

# --hints--

`validate("a@b.cd")` dovrebbe restituire `true`.

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` dovrebbe restituire `true`.

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` dovrebbe restituire `false`.

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` dovrebbe restituire `false`.

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` dovrebbe restituire `false`.

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` dovrebbe restituire `true`.

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` dovrebbe restituire `false`.

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` dovrebbe restituire `false`.

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")` dovrebbe restituire `false`.

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")` dovrebbe restituire `false`.

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
