---
id: 68cae5b538ff798bbd4da004
title: "Sfida 66: rimuovere i tag HTML"
challengeType: 28
dashedName: challenge-66
---

# --description--

Dato un stringa di codice HTML, rimuovi i tag e restituisci il contenuto di testo semplice.

- La stringa di input conterrà solo HTML valido.
- I tag HTML possono essere annidati.
- Rimuovi i tag e tutti gli attributi.

Ad esempio, `'<a href="#">Click here</a>'` dovrebbe restituire `"Click here"`.

# --hints--

`stripTags('<a href="#">Click here</a>')` dovrebbe restituire `"Click here"`.

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` dovrebbe restituire `"Hello World!"`.

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` dovrebbe restituire una stringa vuota (`""`).

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` dovrebbe restituire `sectionsection`.

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
