---
id: 68cae5b538ff798bbd4da001
title: "Sfida 63: battaglia di parole"
challengeType: 28
dashedName: challenge-63
---

# --description--

Date due frasi che rappresentano la tua squadra e una squadra avversaria, dove ogni parola della tua squadra combatte contro la parola corrispondente della squadra avversaria, determina quale squadra vince seguendo queste regole:

- Le frasi date conterranno sempre lo stesso numero di parole.
- Le parole sono separate da uno spazio singolo e conterranno solo lettere.
- Il valore di ogni parola è la somma delle sue lettere.
- Le lettere `a` a `z` corrispondono ai valori `1` fino a `26`. Per esempio, `a` vale `1`, e `z` vale `26`.
- Una lettera maiuscola raddoppia il valore della lettera. Per esempio, `A` vale `2`, e `Z` vale `52`.
- Le parole combattono in ordine: la prima parola della tua squadra combatte contro la prima parola della squadra avversaria, e così via.
- Una parola vince se il suo valore è maggiore del valore della parola avversaria.
- La squadra con più parole vincenti è la vincitrice.

Restituisci `"We win"` se la tua squadra vince, `"We lose"` se la tua squadra perde, e `"Draw"` se entrambe le squadre hanno lo stesso numero di vittorie.

# --hints--

`battle("hello world", "hello word")` dovrebbe restituire `"We win"`.

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")` dovrebbe restituire `"We win"`.

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")` dovrebbe restituire `"We lose"`.

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")` dovrebbe restituire `"Draw"`.

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")` dovrebbe restituire `"We win"`.

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` dovrebbe restituire `"We lose"`.

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")` dovrebbe restituire `"Draw"`.

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
