---
id: 68cae5b538ff798bbd4da003
title: "Sfida 65: conteggio delle stringhe"
challengeType: 28
dashedName: challenge-65
---

# --description--

Date due stringhe, determina quante volte la seconda stringa appare nella prima.

- La stringa pattern può sovrapporsi nella prima stringa. Per esempio, `"aaa"` contiene `"aa"` due volte. Le prime due `a` e le seconde due.

# --hints--

`count('abcdefg', 'def')` dovrebbe restituire `1`.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` dovrebbe restituire `0`.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` dovrebbe restituire `2`.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` dovrebbe restituire `3`.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` dovrebbe restituire `10`.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
