---
id: 68cae5b538ff798bbd4da009
title: "Sfida 71: calcolatore di mance"
challengeType: 28
dashedName: challenge-71
---

# --description--

Dato il prezzo del tuo pasto e una percentuale di mancia personalizzata, restituisci un array con tre valori di mancia: 15%, 20% e l'importo personalizzato.

- I prezzi saranno forniti nel formato: `"$N.NN"`.
- Le percentuali di mancia personalizzate saranno fornite in questo formato: `"25%"`.
- Restituisci gli importi nello stesso formato `"$N.NN"`, arrotondati a due cifre decimali.

Per esempio, dato un prezzo del pasto `"$10.00"` e un valore di mancia personalizzato `"25%"`, restituisci `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculateTips("$10.00", "25%")` dovrebbe restituire `["$1.50", "$2.00", "$2.50"]`.

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` dovrebbe restituire `["$13.45", "$17.93", "$23.31"]`.

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` dovrebbe restituire `["$2.98", "$3.97", "$1.79"]`.

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
