---
id: 68d2ba1468508398389487d0
title: "Sfida 75: tesoro nascosto"
challengeType: 28
dashedName: challenge-75
---

# --description--

Dato un array 2D che rappresenta una mappa del fondale oceanico contenente un tesoro nascosto, e un array con le coordinate (`[row, column]`) per la prossima immersione della tua ricerca del tesoro, restituisci `"Empty"`, `"Found"` o `"Recovered"` seguendo queste regole:

- L'array 2D dato conterrà esattamente un tesoro non ancora recuperato, che occuperà più celle.
- Ogni cella nell'array 2D conterrà uno dei seguenti valori:
  - `"-"`: nessun tesoro.
  - `"O"`: una parte del tesoro non ancora trovata.
  - `"X"`: una parte del tesoro già trovata.

- Se la posizione dell'immersione non contiene tesoro, restituisci `"Empty"`.
- Se la posizione dell'immersione trova tesoro, ma almeno un'altra parte del tesoro rimane non trovata, restituisci `"Found"`.
- Se la posizione dell'immersione trova l'ultima parte non trovata del tesoro, restituisci `"Recovered"`.

Per esempio, dato:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

E `[2, 1]` per le coordinate della posizione dell'immersione, restituisci `"Recovered"` perché l'immersione ha trovato l'ultima parte non trovata del tesoro.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` dovrebbe restituire `"Recovered"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` dovrebbe restituire `"Empty"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` dovrebbe restituire `"Found"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` dovrebbe restituire `"Found"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` dovrebbe restituire `"Recovered"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` dovrebbe restituire `"Empty"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
