---
id: 68d2ba1468508398389487ce
title: "Sfida 73: parla con saggezza, devi"
challengeType: 28
dashedName: challenge-73
---

# --description--

Data una frase, restituisci una versione che suoni come un consiglio di un insegnante saggio usando le seguenti regole:

- Le parole sono separate da uno spazio singolo.
- Trova la prima occorrenza di una delle seguenti parole nella frase: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Sposta tutte le parole prima e inclusa quella parola alla fine della frase e:
  - conserva l’ordine delle parole quando le sposti.
  - rendile tutte minuscole.
  - aggiungi una virgola e uno spazio prima di esse.
- Metti in maiuscolo la prima lettera della nuova prima parola della frase.
- Tutte le frasi date finiranno con un segno di punteggiatura singolo. Mantieni la punteggiatura originale della frase e spostala alla fine della nuova frase.
- Restituisci la nuova frase, assicurandoti che ci sia uno spazio singolo tra ogni parola e nessuno spazio all’inizio o alla fine della frase.

Per esempio, data `"You must speak wisely."` restituisci `"Speak wisely, you must."`

# --hints--

`wiseSpeak("You must speak wisely.")` dovrebbe restituire `"Speak wisely, you must."`

```js
assert.equal(wiseSpeak("You must speak wisely."), "Speak wisely, you must.");
```

`wiseSpeak("You can do it!")` dovrebbe restituire `"Do it, you can!"`

```js
assert.equal(wiseSpeak("You can do it!"), "Do it, you can!");
```

`wiseSpeak("Do you think you will complete this?")` dovrebbe restituire `"Complete this, do you think you will?"`

```js
assert.equal(wiseSpeak("Do you think you will complete this?"), "Complete this, do you think you will?");
```

`wiseSpeak("All your base are belong to us.")` dovrebbe restituire `"Belong to us, all your base are."`

```js
assert.equal(wiseSpeak("All your base are belong to us."), "Belong to us, all your base are.");
```

`wiseSpeak("You have much to learn.")` dovrebbe restituire `"Much to learn, you have."`

```js
assert.equal(wiseSpeak("You have much to learn."), "Much to learn, you have.");
```

# --seed--

## --seed-contents--

```js
function wiseSpeak(sentence) {

  return sentence;
}
```

# --solutions--

```js
function wiseSpeak(sentence) {
  const triggers = ["have", "must", "are", "will", "can"];
  const punctuation = sentence[sentence.length - 1];
  const words = sentence.split(" ");

  const triggerIndex = words.findIndex(w => triggers.includes(w));

  const toMove = words.slice(0, triggerIndex + 1).map(w => w.toLowerCase());
  const remaining = words.slice(triggerIndex + 1);

  const newStart = remaining.map((w, i) => {
    return i === 0 ? w[0].toUpperCase() + w.slice(1) :
      i === remaining.length - 1 ? w.slice(0, w.length - 1) : w;
  })

  return newStart.join(" ") + ", " + toMove.join(" ") + punctuation;
}
```
