---
id: 68d2ba1468508398389487cf
title: "Sfida 74: canzoni preferite"
challengeType: 28
dashedName: challenge-74
---

# --description--

Ti ricordi gli iPod? Il primo modello è uscito 24 anni fa, il 23 ottobre 2001.

Dato un array di oggetti canzone che rappresentano la tua playlist iPod, restituisci un array con i titoli delle due canzoni più ascoltate, mettendo per prime quelle più ascoltate.

- Ogni oggetto avrà una proprietà `"title"` (stringa) e una proprietà `"plays"` (intero).

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` dovrebbe restituire `["Sync or Swim", "Earbud Blues"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` dovrebbe restituire `["Clickwheel Love", "99 Downloads"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` dovrebbe restituire `["Song B", "Song C"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```
