---
id: 68d30845cc08266018fc46be
title: "Sfida 78: sequenza di numeri interi"
challengeType: 28
dashedName: challenge-78
---

# --description--

Dato un numero intero positivo, restituisci una stringa con tutti i numeri interi da `1` fino a, e compreso, il numero dato, in ordine numerico.

Per esempio, dato `5`, restituisci `"12345"`.

# --hints--

`sequence(5)` dovrebbe restituire `"12345"`.

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)` dovrebbe restituire `"12345678910"`.

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)` dovrebbe restituire `"1"`.

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)` dovrebbe restituire `"123456789101112131415161718192021222324252627"`.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
