---
id: 68d30845cc08266018fc46bc
title: "Sfida 76: DNA complementare"
challengeType: 28
dashedName: challenge-76
---

# --description--

Data una stringa che rappresenta una sequenza di DNA, restituisci il suo filamento complementare usando le seguenti regole:

- Il DNA è composto dalle lettere `"A"`, `"C"`, `"G"` e `"T"`.
- Le lettere `"A"` e `"T"` si completano a vicenda.
- Le lettere `"C"` e `"G"` si completano a vicenda.

Per esempio, data `"ACGT"`, restituisci `"TGCA"`.

# --hints--

`complementaryDNA("ACGT")` dovrebbe restituire `"TGCA"`.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` dovrebbe restituire `"TACGCATGCAATCG"`.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` dovrebbe restituire `"CCGAATGCTAGCTTC"`.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` dovrebbe restituire `"CTAGATCGATCCGATCGATC"`.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
