---
id: 68d30fc57588d97fd3027b30
title: "Sfida 79: Navigator"
challengeType: 28
dashedName: challenge-79
---

# --description--

Il 28 ottobre 1994 è stato rilasciato Netscape Navigator, che ha aiutato milioni di persone a esplorare il web agli albori.

Dato un array di comandi del browser che hai eseguito su Netscape Navigator, restituisci la pagina corrente su cui ti trovi dopo aver eseguito tutti i comandi seguendo queste regole:

- Parti sempre dalla pagina `"Home"`, che non sarà inclusa nell'array dei comandi.
- I comandi validi sono:
  - `"Visit Page"`: dove `"Page"` è il nome della pagina che stai visitando. Per esempio, `"Visit About"` ti porta alla pagina `"About"`. Quando visiti una nuova pagina, assicurati di scartare qualsiasi cronologia avanti che hai.
  - `"Back"`: ti porta alla pagina precedente nella cronologia o rimane sulla pagina corrente se non ce n'è una.
  - `"Forward"`: ti porta avanti nella cronologia alla pagina da cui sei venuto o rimane sulla pagina corrente se non ce n'è una.

Per esempio, dato `["Visit About Us", "Back", "Forward"]`, restituisci `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` dovrebbe restituire `"About Us"`.

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` dovrebbe restituire `"Home"`.

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` dovrebbe restituire `"Home"`.

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` dovrebbe restituire `"Gallery"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` dovrebbe restituire `"Home"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` dovrebbe restituire `"Contact"`.

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` dovrebbe restituire `"Visit Us"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```
