---
id: 68d30845cc08266018fc46bd
title: "Sfida 77: formattatore di durata"
challengeType: 28
dashedName: challenge-77
---

# --description--

Dato un numero intero di secondi, restituisci una stringa che rappresenta la stessa durata nel formato `"H:MM:SS"`, dove `"H"` è il numero di ore, `"MM"` è il numero di minuti e `"SS"` è il numero di secondi. Restituisci il tempo usando le seguenti regole:

- Secondi: devono essere sempre due cifre.
- Minuti: devono omettere gli zeri iniziali quando non sono necessari. Usa `"0"` se la durata è inferiore a un minuto.
- Ore: devono essere incluse solo se sono maggiori di zero.

# --hints--

`format(500)` dovrebbe restituire `"8:20"`.

```js
assert.equal(format(500), "8:20");
```

`format(4000)` dovrebbe restituire `"1:06:40"`.

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` dovrebbe restituire `"0:01"`.

```js
assert.equal(format(1), "0:01");
```

`format(5555)` dovrebbe restituire `"1:32:35"`.

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` dovrebbe restituire `"27:46:39"`.

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
