---
id: 68e39ed6106dac2f0a98fd65
title: "Sfida 83: validazione della firma"
challengeType: 28
dashedName: challenge-83
---

# --description--

Data una stringa di messaggio, una stringa di chiave segreta e un numero di firma, determina se la firma è valida usando questo metodo di codifica:

- Le lettere nel messaggio e nella chiave segreta hanno questi valori:
  - `a` a `z` hanno rispettivamente i valori `1` a `26`.
  - `A` a `Z` hanno rispettivamente i valori `27` a `52`.
- Tutti gli altri caratteri non hanno valore.
- Calcola la firma sommando la somma del messaggio più la somma della chiave segreta.

Per esempio, dato il messaggio `"foo"` e la chiave segreta `"bar"`, la firma sarebbe `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

Verifica se la firma calcolata corrisponde alla firma fornita.

# --hints--

`verify("foo", "bar", 57)` dovrebbe restituire `true`.

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` dovrebbe restituire `false`.

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` dovrebbe restituire `true`.

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` dovrebbe restituire `false`.

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` dovrebbe restituire `true`.

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` dovrebbe restituire `true`.

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
