---
id: 68e39ed6106dac2f0a98fd62
title: "Sfida 80: ordinatore di email"
challengeType: 28
dashedName: challenge-80
---

# --description--

Il 29 ottobre 1971 è stata inviata la prima email di sempre, introducendo il formato `username@domain` che usiamo ancora oggi. Ora ci sono miliardi di indirizzi email.

In questa sfida, ti viene fornita una lista di indirizzi email e devi ordinarli alfabeticamente prima per nome di dominio (la parte dopo `@`) e poi per nome utente (la parte prima di `@`).

- L'ordinamento deve essere case-insensitive.
- Se più email hanno lo stesso dominio, ordinale per nome utente.
- Restituisci un array con gli indirizzi ordinati.
- Gli indirizzi restituiti devono mantenere il loro caso originale.

Per esempio, dato `["jill@mail.com", "john@example.com", "jane@example.com"]`, restituisci `["jane@example.com", "john@example.com", "jill@mail.com"]`.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` dovrebbe restituire `["jane@example.com", "john@example.com", "jill@mail.com"]`.

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` dovrebbe restituire `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`.

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` dovrebbe restituire `["user@x.com", "user@y.com", "user@z.com"]`.

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` dovrebbe restituire `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`.

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` dovrebbe restituire `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`.

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```
