---
id: 68e39ed6106dac2f0a98fd63
title: "Sfida 81: n-esimo numero primo"
challengeType: 28
dashedName: challenge-81
---

# --description--

Un numero primo è un intero positivo maggiore di 1 che è divisibile solo per 1 e per se stesso. I primi cinque numeri primi sono `2`, `3`, `5`, `7` e `11`.

Dato un intero positivo `n`, restituisci il `n`-esimo numero primo. Per esempio, dato `5` restituisci il 5° numero primo: `11`.

# --hints--

`nthPrime(5)` dovrebbe restituire `11`.

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` dovrebbe restituire `29`.

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` dovrebbe restituire `53`.

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` dovrebbe restituire `523`.

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` dovrebbe restituire `7919`.

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
