---
id: 68e39ed6106dac2f0a98fd64
title: "Sfida 82: SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

Data una stringa che rappresenta il nome di una variabile, convertila in "spooky case" seguendo queste regole:

- Sostituisci tutti gli underscore (`_`) e i trattini (`-`) con una tilde (`~`).
- Metti in maiuscolo la prima lettera della stringa e ogni altra lettera alternata dopo quella. Ignora il carattere tilde nel conteggio. Tutte le altre lettere devono essere in minuscolo.

Per esempio, data `hello_world`, restituisci `HeLlO~wOrLd`.

# --hints--

`spookify("hello_world")` dovrebbe restituire `"HeLlO~wOrLd"`.

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")` dovrebbe restituire `"SpOoKy~CaSe"`.

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")` dovrebbe restituire `"TrIcK~oR~tReAt"`.

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")` dovrebbe restituire `"C~a~N~d~Y~~b~O~w~L"`.

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` dovrebbe restituire `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
