---
id: 68ee9e3066cfd4eb2328e8a4
title: "Sfida 85: contatore di parole"
challengeType: 28
dashedName: challenge-85
---

# --description--

Data una stringa che rappresenta una frase, restituisci il numero di parole presenti nella frase.

- Le parole sono qualsiasi sequenza di caratteri non spazi e sono separate da uno spazio singolo.

# --hints--

`countWords("Hello world")` dovrebbe restituire `2`.

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` dovrebbe restituire `9`.

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` dovrebbe restituire `4`.

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` dovrebbe restituire `7`.

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` dovrebbe restituire `7`.

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
