---
id: 68ee9e3066cfd4eb2328e8a7
title: "Sfida 88: trovare il giorno della settimana"
challengeType: 28
dashedName: challenge-88
---

# --description--

Data una stringa che rappresenta una data nel formato `YYYY-MM-DD`, restituisci il giorno della settimana.

I giorni validi da restituire sono:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Assicurati di ignorare i fusi orari.

# --hints--

`getWeekday("2025-11-06")` dovrebbe restituire `Thursday`.

```js
assert.equal(getWeekday("2025-11-06"), "Thursday");
```

`getWeekday("1999-12-31")` dovrebbe restituire `Friday`.

```js
assert.equal(getWeekday("1999-12-31"), "Friday");
```

`getWeekday("1111-11-11")` dovrebbe restituire `Saturday`.

```js
assert.equal(getWeekday("1111-11-11"), "Saturday");
```

`getWeekday("2112-12-21")` dovrebbe restituire `Wednesday`.

```js
assert.equal(getWeekday("2112-12-21"), "Wednesday");
```

`getWeekday("2345-10-01")` dovrebbe restituire `Monday`.

```js
assert.equal(getWeekday("2345-10-01"), "Monday");
```

# --seed--

## --seed-contents--

```js
function getWeekday(dateString) {

  return dateString;
}
```

# --solutions--

```js
function getWeekday(dateString) {
  const days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"];
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  return days[date.getUTCDay()];
}
```
