---
id: 68ee9e3066cfd4eb2328e8a6
title: "Sfida 87: costruttore di matrici"
challengeType: 28
dashedName: challenge-87
---

# --description--

Dati due numeri interi (un numero di righe e un numero di colonne), restituisci una matrice (un array di array) riempita di zeri (`0`) della dimensione indicata.

Per esempio, dati `2` e `3`, restituisci:

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`buildMatrix(2, 3)` dovrebbe restituire `[[0, 0, 0], [0, 0, 0]]`.

```js
assert.deepEqual(buildMatrix(2, 3), [[0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(3, 2)` dovrebbe restituire `[[0, 0], [0, 0], [0, 0]]`.

```js
assert.deepEqual(buildMatrix(3, 2), [[0, 0], [0, 0], [0, 0]]);
```

`buildMatrix(4, 3)` dovrebbe restituire `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
assert.deepEqual(buildMatrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(9, 1)` dovrebbe restituire `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`.

```js
assert.deepEqual(buildMatrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]]);
```

# --seed--

## --seed-contents--

```js
function buildMatrix(rows, cols) {

  return rows;
}
```

# --solutions--

```js
function buildMatrix(rows, cols) {
  const matrix = [];
  for (let i = 0; i < rows; i++) {
    const row = new Array(cols).fill(0);
    matrix.push(row);
  }
  return matrix;
}
```
