---
id: 68ee9e3066cfd4eb2328e8a8
title: "Sfida 89: contare le carte"
challengeType: 28
dashedName: challenge-89
---

# --description--

Un mazzo standard di carte da gioco ha 13 carte uniche per ogni seme. Dato un numero intero che rappresenta il numero di carte da estrarre dal mazzo, restituisci il numero di combinazioni uniche di carte che puoi ottenere.

- L'ordine non conta. Estrarre la carta A e poi la carta B è lo stesso che estrarre la carta B e poi la carta A.

Per esempio, dato `52`, restituisci `1`. C'è solo una combinazione di 52 carte da estrarre da un mazzo di 52 carte. E dato `2`, restituisci `1326`. Ci sono 1326 combinazioni di carte possibili quando si estraggono 2 carte dal mazzo.

# --hints--

`combinations(52)` dovrebbe restituire `1`.

```js
assert.equal(combinations(52), 1);
```

`combinations(1)` dovrebbe restituire `52`.

```js
assert.equal(combinations(1), 52);
```

`combinations(2)` dovrebbe restituire `1326`.

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)` dovrebbe restituire `2598960`.

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)` dovrebbe restituire `15820024220`.

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)` dovrebbe restituire `1326`.

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
