---
id: 68f6587287ad1f4ad39b0c7d
title: "Sfida 91: ricerca di parole"
challengeType: 28
dashedName: challenge-91
---

# --description--

Data una matrice (un array di array) di singole lettere e una parola da trovare, restituisci gli indici di inizio e fine della parola nella matrice.

- La matrice data sarà riempita con tutte lettere minuscole (`a-z`).
- La parola da trovare sarà sempre presente nella matrice esattamente una volta.
- La parola da trovare sarà sempre in linea retta in una di queste direzioni:
  - da sinistra a destra
  - da destra a sinistra
  - dall’alto verso il basso
  - dal basso verso l’alto

Per esempio, data la matrice:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

E la parola `"cat"`, restituisci:

```md
[[0, 1], [2, 1]]
```

Dove `[0, 1]` sono gli indici per l’inizio della parola `"c"`, e `[2, 1]` sono gli indici per la fine della parola `"t"`.

# --hints--

`findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` dovrebbe restituire `[[0, 1], [2, 1]]`.

```js
assert.deepEqual(findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]]);
```

`findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` dovrebbe restituire `[[0, 0], [0, 2]]`.

```js
assert.deepEqual(findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]]);
```

`findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` dovrebbe restituire `[[3, 3], [0, 3]]`.

```js
assert.deepEqual(findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]]);
```

`findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` dovrebbe restituire `[[1, 3], [1, 1]]`.

```js
assert.deepEqual(findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]]);
```

# --seed--

## --seed-contents--

```js
function findWord(matrix, word) {

  return matrix;
}
```

# --solutions--

```js
function findWord(matrix, word) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const len = word.length;

  const directions = [
    [0, 1],
    [0, -1],
    [1, 0],
    [-1, 0]
  ];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      for (let [dr, dc] of directions) {
        let match = true;
        for (let i = 0; i < len; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols || matrix[nr][nc] !== word[i]) {
            match = false;
            break;
          }
        }
        if (match) {
          return [
            [r, c],
            [r + dr * (len - 1), c + dc * (len - 1)]
          ];
        }
      }
    }
  }
}
```
