---
id: 68f6587287ad1f4ad39b0c80
title: "Sfida 94: generatore di firma email"
challengeType: 28
dashedName: challenge-94
---

# --description--

Dato il nome, il titolo e l'azienda di una persona come stringhe, restituisci una firma email come singola stringa usando le seguenti regole:

- Il nome deve apparire per primo, preceduto da un prefisso che dipende dalla prima lettera del nome. Per nomi che iniziano con (ignorando maiuscole/minuscole):
  - `A-I`: usa `>>` come prefisso.
  - `J-R`: usa `--` come prefisso.
  - `S-Z`: usa `::` come prefisso.
- Dopo il nome deve seguire una virgola e uno spazio (`, `).
- Il titolo e l'azienda devono seguire la virgola e lo spazio, separati da `" at "` (con spazi intorno).

Per esempio, dati `"Quinn Waverly"`, `"Founder and CEO"` e `"TechCo"` restituisci `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` dovrebbe restituire `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` dovrebbe restituire `">>Alice Reed, Engineer at TechCo"`.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` dovrebbe restituire `"::Tina Vaughn, Developer at example.com"`.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` dovrebbe restituire `">>B. B., Product Tester at AcmeCorp"`.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` dovrebbe restituire `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
