---
id: 68f6587287ad1f4ad39b0c84
title: "Sfida 98: conteggio dei rettangoli"
challengeType: 28
dashedName: challenge-98
---

# --description--

Dati due numeri interi positivi che rappresentano la larghezza e l'altezza di un rettangolo, determina quanti rettangoli possono essere contenuti in quello dato.

- Conta solo i rettangoli con larghezza e altezza intere.

Per esempio, dati `1` e `3`, restituisci `6`. Tre rettangoli 1x1, due rettangoli 1x2 e un rettangolo 1x3.

# --hints--

`countRectangles(1, 3)` dovrebbe restituire `6`.

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` dovrebbe restituire `18`.

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` dovrebbe restituire `3`.

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` dovrebbe restituire `150`.

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` dovrebbe restituire `12540`.

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
