---
id: 68f6587287ad1f4ad39b0c7c
title: "Sfida 90: limite di caratteri"
challengeType: 28
dashedName: challenge-90
---

# --description--

In questa sfida, ti viene data una stringa e devi determinare se può essere inserita in un post sui social media. Restituisci le seguenti stringhe in base alle regole indicate:

- `"short post"` se rientra nel limite di 40 caratteri.
- `"long post"` se supera i 40 caratteri ma rientra nel limite di 80 caratteri.
- `"invalid post"` se è troppo lunga per rientrare in entrambi i limiti.

# --hints--

`canPost("Hello world")` dovrebbe restituire `"short post"`.

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")` dovrebbe restituire `"long post"`.

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")` dovrebbe restituire `"invalid post"`.

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```
