---
id: 68f6587287ad1f4ad39b0c81
title: "Sfida 95: shift di un array"
challengeType: 28
dashedName: challenge-95
---

# --description--

Dato un array e un intero che indica di quante posizioni spostare l'array, restituisci l'array spostato.

- Un intero positivo sposta l'array verso sinistra.
- Un intero negativo sposta l'array verso destra.
- Lo spostamento avvolge l'array.

Per esempio, dato `[1, 2, 3]` e `1`, sposta l'array di 1 verso sinistra, restituendo `[2, 3, 1]`.

# --hints--

`shiftArray([1, 2, 3], 1)` dovrebbe restituire `[2, 3, 1]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` dovrebbe restituire `[3, 1, 2]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` dovrebbe restituire `["charlie", "alpha", "bravo"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` dovrebbe restituire `["bravo", "charlie", "alpha"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` dovrebbe restituire `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
