---
id: 68f6587287ad1f4ad39b0c85
title: "Sfida 99: test delle impronte digitali"
challengeType: 28
dashedName: challenge-99
---

# --description--

Date due stringhe che rappresentano impronte digitali, determina se corrispondono usando le seguenti regole:

- Ogni impronta digitale sarà composta solo da lettere minuscole (`a-z`).
- Due impronte digitali sono considerate corrispondenti se:
  - Hanno la stessa lunghezza.
  - Il numero di caratteri diversi non supera il 10% della lunghezza dell'impronta.

# --hints--

`isMatch("helloworld", "helloworld")` dovrebbe restituire `true`.

```js
assert.isTrue(isMatch("helloworld", "helloworld"));
```

`isMatch("helloworld", "helloworlds")` dovrebbe restituire `false`.

```js
assert.isFalse(isMatch("helloworld", "helloworlds"));
```

`isMatch("helloworld", "jelloworld")` dovrebbe restituire `true`.

```js
assert.isTrue(isMatch("helloworld", "jelloworld"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")` dovrebbe restituire `true`.

```js
assert.isTrue(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"));
```

`isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")` dovrebbe restituire `true`.

```js
assert.isTrue(isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")` dovrebbe restituire `false`.

```js
assert.isFalse(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"));
```

# --seed--

## --seed-contents--

```js
function isMatch(fingerprintA, fingerprintB) {

  return fingerprintA;
}
```

# --solutions--

```js
function isMatch(fingerprintA, fingerprintB) {
  if (fingerprintA.length !== fingerprintB.length) return false;

  const length = fingerprintA.length;
  let mismatches = 0;

  for (let i = 0; i < length; i++) {
    if (fingerprintA[i] !== fingerprintB[i]) {
      mismatches++;
      if (mismatches > length * 0.1) return false;
    }
  }

  return true;
}
```
