---
id: 68f6587287ad1f4ad39b0c82
title: "Sfida 96: è il weekend?"
challengeType: 28
dashedName: challenge-96
---

# --description--

Dato un giorno nel formato `"YYYY-MM-DD"`, restituisci il numero di giorni mancanti al weekend.

- Il weekend inizia il sabato.
- Se il giorno dato è sabato o domenica, restituisci `"It's the weekend!"`.
- Altrimenti, restituisci `"X days until the weekend."`, dove `X` è il numero di giorni fino al sabato.
- Se `X` è `1`, usa `"day"` (singolare) invece di `"days"` (plurale).
- Assicurati che il calcolo ignori il fuso orario locale.

# --hints--

`daysUntilWeekend("2025-11-14")` dovrebbe restituire `"1 day until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` dovrebbe restituire `"3 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` dovrebbe restituire `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` dovrebbe restituire `"4 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` dovrebbe restituire `"5 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` dovrebbe restituire `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
