---
id: 68f6587287ad1f4ad39b0c7f
title: "Sfida 93: vocali e consonanti"
challengeType: 28
dashedName: challenge-93
---

# --description--

Data una stringa, restituisci un array con il numero di vocali e il numero di consonanti presenti nella stringa.

- Le vocali sono `a`, `e`, `i`, `o`, `u` in qualsiasi caso.
- Le consonanti sono tutte le altre lettere in qualsiasi caso.
- Ignora tutti i caratteri che non sono lettere.

Per esempio, data `"Hello World"`, restituisci `[3, 7]`.

# --hints--

`count("Hello World")` dovrebbe restituire `[3, 7]`.

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` dovrebbe restituire `[3, 7]`.

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` dovrebbe restituire `[1, 5]`.

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` dovrebbe restituire `[5, 7]`.

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` dovrebbe restituire `[3, 7]`.

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` dovrebbe restituire `[11, 24]`.

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
