---
id: 68f6587287ad1f4ad39b0c7e
title: "Sfida 92: Estrazione dell'estensione"
challengeType: 28
dashedName: challenge-92
---

# --description--

Data una stringa che rappresenta un nome di file, restituisci l'estensione del file.

- L'estensione è la parte del nome del file che segue l'ultimo punto (`.`).
- Se il nome del file non contiene un punto o termina con un punto, restituisci `"none"`.
- L'estensione deve essere restituita così com'è, mantenendo il caso.

# --hints--

`getExtension("document.txt")` dovrebbe restituire `"txt"`.

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` dovrebbe restituire `"none"`.

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` dovrebbe restituire `"PNG"`.

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` dovrebbe restituire `"gitignore"`.

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` dovrebbe restituire `"gz"`.

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` dovrebbe restituire `"none"`.

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
