---
id: 68f6587287ad1f4ad39b0c83
title: "Sfida 97: MCD"
challengeType: 28
dashedName: challenge-97
---

# --description--

Dati due numeri interi positivi, restituisci il loro massimo comune divisore (MCD).

- Il MCD di due numeri interi è il numero più grande che divide entrambi senza lasciare resto.

Ad esempio, i divisori di `4` sono `1`, `2` e `4`. I divisori di `6` sono `1`, `2`, `3` e `6`. Quindi, dati `4` e `6`, restituisci `2`, il numero più grande che appare in entrambi gli insiemi di divisori.

# --hints--

`gcd(4, 6)` dovrebbe restituire `2`.

```js
assert.equal(gcd(4, 6), 2);
```

`gcd(20, 15)` dovrebbe restituire `5`.

```js
assert.equal(gcd(20, 15), 5);
```

`gcd(13, 17)` dovrebbe restituire `1`.

```js
assert.equal(gcd(13, 17), 1);
```

`gcd(654, 456)` dovrebbe restituire `6`.

```js
assert.equal(gcd(654, 456), 6);
```

`gcd(3456, 4320)` dovrebbe restituire `864`.

```js
assert.equal(gcd(3456, 4320), 864);
```

# --seed--

## --seed-contents--

```js
function gcd(x, y) {

  return x;
}
```

# --solutions--

```js
function gcd(x, y) {
  while (y !== 0) {
    const temp = y;
    y = x % y;
    x = temp;
  }
  return x;
}
```
