---
id: 68ffb91507a5b645769328cc
title: "Sfida 109: Quanti anni ho di nuovo?"
challengeType: 28
dashedName: challenge-109
---

# --description--

Dato il compleanno di una persona nel formato `YYYY-MM-DD`, restituisci l'età della persona al 27 novembre 2025.

- Considera che tutti i compleanni sono date valide anteriori al 27 novembre 2025.
- Restituisci l'età come numero intero.
- Assicurati di considerare se la persona ha già compiuto gli anni nel 2025.

# --hints--

`calculateAge("2000-11-20")` dovrebbe restituire `25`.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` dovrebbe restituire `24`.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` dovrebbe restituire `11`.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` dovrebbe restituire `31`.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` dovrebbe restituire `30`.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
