---
id: 68ffb91507a5b645769328c4
title: "Sfida 101: convertitore di intestazioni markdown"
challengeType: 28
dashedName: challenge-101
---

# --description--

Data una stringa che rappresenta un'intestazione Markdown, restituisci l'intestazione HTML equivalente.

Un'intestazione Markdown valida deve:

- Iniziare con zero o più spazi, seguiti da
- Da 1 a 6 caratteri hash (`#`) consecutivi, poi
- Almeno uno spazio. Infine,
- Il testo dell'intestazione.

Il numero di simboli hash determina il livello dell'intestazione. Per esempio, un simbolo hash corrisponde a un tag `h1`, e sei simboli hash corrispondono a un tag `h6`.

Se la stringa fornita non ha esattamente questo formato, restituisci `"Invalid format"`.

Per esempio, data `"# My level 1 heading"`, restituisci `"<h1>My level 1 heading</h1>"`.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`convert("# My level 1 heading")` dovrebbe restituire `"<h1>My level 1 heading</h1>"`.

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")` dovrebbe restituire `"Invalid format"`.

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")` dovrebbe restituire `"<h5>My level 5 heading</h5>"`.

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")` dovrebbe restituire `"Invalid format"`.

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")` dovrebbe restituire `"<h3>My level 3 heading</h3>"`.

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")` dovrebbe restituire `"Invalid format"`.

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")` dovrebbe restituire `"<h2>My #2 heading</h2>"`.

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```
