---
id: 68ffb91507a5b645769328c7
title: "Sfida 104: scala ricette"
challengeType: 28
dashedName: challenge-104
---

# --description--

Dato un array di ingredienti di una ricetta e un numero per scalare la ricetta, restituisci un array con le quantità scalate di conseguenza.

- Ogni elemento nell'array fornito sarà una stringa nel formato: `"quantity unit ingredient"`. Per esempio `"2 C Flour"`.
- Scala la quantità in base al numero dato. Non includere zeri finali e non convertire le unità.
- Restituisci gli elementi scalati nello stesso ordine in cui sono forniti.

# --hints--

`scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2)` dovrebbe restituire `["4 C Flour", "3 T Sugar"]`.

```js
assert.deepEqual(scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2), ["4 C Flour", "3 T Sugar"]);
```

`scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5)` dovrebbe restituire `["6 T Flour", "1.5 C Milk", "3 T Oil"]`.

```js
assert.deepEqual(scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5), ["6 T Flour", "1.5 C Milk", "3 T Oil"]);
```

`scaleRecipe(["3 C Milk", "2 C Oats"], 0.5)` dovrebbe restituire `["1.5 C Milk", "1 C Oats"]`.

```js
assert.deepEqual(scaleRecipe(["3 C Milk", "2 C Oats"], 0.5), ["1.5 C Milk", "1 C Oats"]);
```

`scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5)` dovrebbe restituire `["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]`.

```js
assert.deepEqual(scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5), ["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]);
```

# --seed--

## --seed-contents--

```js
function scaleRecipe(ingredients, scale) {

  return ingredients;
}
```

# --solutions--

```js
function scaleRecipe(ingredients, scale) {
  return ingredients.map(item => {
    const [quantityStr, unit, ...rest] = item.split(" ");
    const ingredientName = rest.join(" ");
    
    const quantity = parseFloat(quantityStr) * scale;
    
    return `${quantity} ${unit} ${ingredientName}`;
  });
}
```
