---
id: 68ffb91507a5b645769328c9
title: "Sfida 106: validatore di messaggi"
challengeType: 28
dashedName: challenge-106
---

# --description--

Dato un messaggio come stringa e una stringa di validazione, determina se il messaggio è valido.

- Un messaggio è valido se ogni parola nel messaggio inizia con la lettera corrispondente nella stringa di validazione, nell'ordine.
- Le lettere non fanno distinzione tra maiuscole e minuscole.
- Le parole nel messaggio sono separate da spazi singoli.

# --hints--

`isValidMessage("hello world", "hw")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
