---
id: 68ffb91507a5b645769328c3
title: "Sfida 100: 100 caratteri"
challengeType: 28
dashedName: challenge-100
---

# --description--

Benvenuto alla 100ª sfida di coding quotidiana!

Data una stringa, ripeti i suoi caratteri finché il risultato non è esattamente lungo 100 caratteri. Se le ripetizioni superano i 100 caratteri, taglia l'eccesso in modo che sia esattamente 100.

# --hints--

`oneHundred("One hundred ")` dovrebbe restituire `"One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One "`.

```js
assert.equal(oneHundred("One hundred "), "One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One ");
```

`oneHundred("freeCodeCamp ")` dovrebbe restituire `"freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC"`.

```js
assert.equal(oneHundred("freeCodeCamp "), "freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC");
```

`oneHundred("daily challenges ")` dovrebbe restituire `"daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge"`.

```js
assert.equal(oneHundred("daily challenges "), "daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge");
```

`oneHundred("!")` dovrebbe restituire `"!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!"`.

```js
assert.equal(oneHundred("!"), "!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!");
```

# --seed--

## --seed-contents--

```js
function oneHundred(chars) {

  return chars;
}
```

# --solutions--

```js
function oneHundred(chars) {
  let result = "";
  let i = 0;

  while (result.length < 100) {
    result += chars[i % chars.length];
    i++;
  }

  return result.slice(0, 100);
}
```
