---
id: 68ffb91507a5b645769328c5
title: "Sfida 102: parola più lunga"
challengeType: 28
dashedName: challenge-102
---

# --description--

Data una stringa che rappresenta una frase, restituisci la parola più lunga nella frase.

- Le parole sono separate da uno spazio singolo.
- Solo le lettere (`a-z`, senza distinzione tra maiuscole e minuscole) contano per la lunghezza della parola.
- Se ci sono più parole con la stessa lunghezza, restituisci la prima che appare.
- Restituisci la parola così come appare nella stringa data, senza la punteggiatura.

# --hints--

`longestWord("The quick red fox")` dovrebbe restituire `"quick"`.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` dovrebbe restituire `"challenge"`.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` dovrebbe restituire `"This"`.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` dovrebbe restituire `"exclamation"`.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` dovrebbe restituire `"tie"`.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` dovrebbe restituire `"Wouldnt"`.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
