---
id: 68ffb91507a5b645769328c6
title: "Sfida 103: minimo comune multiplo"
challengeType: 28
dashedName: challenge-103
---

# --description--

Dati due numeri interi, restituisci il minimo comune multiplo (LCM) dei due numeri.

Il minimo comune multiplo di due numeri è il più piccolo numero intero positivo che è multiplo di entrambi i numeri. Per esempio, dati `4` e `6`, restituisci `12` perché:

- I multipli di `4` sono `4`, `8`, `12` e così via.
- I multipli di `6` sono `6`, `12`, `18` e così via.
- `12` è il numero più piccolo che è multiplo di entrambi.

# --hints--

`lcm(4, 6)` dovrebbe restituire `12`.

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)` dovrebbe restituire `18`.

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)` dovrebbe restituire `100`.

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)` dovrebbe restituire `221`.

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)` dovrebbe restituire `630`.

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```
