---
id: 69162d64f96574d9bb629f00
title: "Sfida 115: convertitore di elementi di lista ordinata in markdown"
challengeType: 28
dashedName: challenge-115
---

# --description--

Data una stringa che rappresenta un elemento di lista ordinata in markdown, restituisci la stringa HTML equivalente.

Un elemento valido di lista ordinata in markdown deve:

- Iniziare con zero o più spazi, seguito da
- Un numero (1 o superiore) e un punto (`.`), seguito da
- Almeno uno spazio, e poi
- Il testo dell'elemento della lista.

Se la stringa non ha esattamente questo formato, restituisci `"Invalid format"`. Altrimenti, racchiudi il testo dell'elemento della lista in tag `li` e restituisci la stringa.

Per esempio, data `"1. My item"`, restituisci `"<li>My item</li>"`.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`convertListItem("1. My item")` dovrebbe restituire `"<li>My item</li>"`.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` dovrebbe restituire `"<li>Another item</li>"`.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` dovrebbe restituire `"Invalid format"`.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` dovrebbe restituire `"<li>list item text</li>"`.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` dovrebbe restituire `"Invalid format"`.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` dovrebbe restituire `"Invalid format"`.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
