---
id: 69162d64f96574d9bb629f01
title: "Sfida 116: conteggio delle permutazioni"
challengeType: 28
dashedName: challenge-116
---

# --description--

Data una stringa, restituisci il numero di permutazioni distinte che si possono formare con i suoi caratteri.

- Una permutazione è qualsiasi riordinamento dei caratteri nella stringa.
- Non contare permutazioni duplicate.
- Se la stringa contiene caratteri ripetuti, le disposizioni ripetute devono essere contate solo una volta.
- La stringa conterrà solo lettere (`A-Z`, `a-z`).

Per esempio, data `"abb"`, restituisci `3` perché ci sono tre modi unici per disporre le lettere: `"abb"`, `"bab"` e `"bba"`.

# --hints--

`countPermutations("abb")` dovrebbe restituire `3`.

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` dovrebbe restituire `6`.

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` dovrebbe restituire `630`.

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` dovrebbe restituire `39916800`.

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
