---
id: 69162d64f96574d9bb629f04
title: "Sfida 119: compressione di stringhe"
challengeType: 28
dashedName: challenge-119
---

# --description--

Data una stringa frase, restituisci una versione compressa della frase in cui le parole duplicate consecutive sono sostituite dalla parola seguita dal numero di volte in cui si ripete tra parentesi tonde.

- Vengono compresse solo le parole duplicate consecutive.
- Le parole sono separate da spazi singoli.

Per esempio, data `"yes yes yes please"`, restituisci `"yes(3) please"`.

# --hints--

`compressString("yes yes yes please")` dovrebbe restituire `"yes(3) please"`.

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")` dovrebbe restituire `"I have(3) apples"`.

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")` dovrebbe restituire `"one(2) three and to the(4)"`.

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")` dovrebbe restituire `"route(6) tee(6)"`.

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
