---
id: 69162d64f96574d9bb629eff
title: "Sfida 114: da camel case a snake case"
challengeType: 28
dashedName: challenge-114
---

# --description--

Data una stringa in camel case, restituisci la versione in snake case della stringa seguendo queste regole:

- La stringa di input conterrà solo lettere (`A-Z` e `a-z`) e inizierà sempre con una lettera minuscola.
- Ogni lettera maiuscola nella stringa camel case indica l'inizio di una nuova parola.
- Converti tutte le lettere in minuscolo.
- Separa le parole con un underscore (`_`).

# --hints--

`toSnake("helloWorld")` dovrebbe restituire `"hello_world"`.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` dovrebbe restituire `"my_variable_name"`.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` dovrebbe restituire `"freecodecamp_daily_challenges"`.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
