---
id: 69162d64f96574d9bb629efd
title: "Sfida 112: rilevatore di intelligenza artificiale"
challengeType: 28
dashedName: challenge-112
---

# --description--

La sfida di oggi è ispirata al rilascio di ChatGPT il 30 novembre 2022.

Data una stringa composta da una o più frasi, determina se è probabile che sia stata generata da un'IA usando le seguenti regole:

- Contiene due o più trattini (`-`).
- Contiene due o più insiemi di parentesi (`()`). Il testo può essere all'interno delle parentesi.
- Contiene tre o più parole con 7 o più lettere.

- Le parole sono separate da uno spazio singolo e consistono solo di lettere (`A-Z`). Non includere la punteggiatura o altri caratteri non alfabetici come parte di una parola.

Se la frase data soddisfa una qualsiasi delle regole sopra, restituisci `"AI"`, altrimenti restituisci `"Human"`.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` dovrebbe restituire `"Human"`.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` dovrebbe restituire `"Human"`.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` dovrebbe restituire `"AI"`.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` dovrebbe restituire `"AI"`.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` dovrebbe restituire `"AI"`.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
