---
id: 69162d64f96574d9bb629efb
title: "Sfida 110: indovina la parola"
challengeType: 28
dashedName: challenge-110
---

# --description--

Date due stringhe della stessa lunghezza, una parola segreta e un tentativo, confronta il tentativo con la parola segreta usando le seguenti regole:

- La parola segreta e il tentativo saranno composti solo da lettere maiuscole (`"A"` a `"Z"`);
- Per ogni lettera nel tentativo, sostituiscila con un numero in base a come corrisponde alla parola segreta:
  - `"2"` se la lettera è nella parola segreta e nella posizione corretta.
  - `"1"` se la lettera è nella parola segreta ma nella posizione sbagliata.
  - `"0"` se la lettera non è nella parola segreta.
- Ogni lettera nella parola segreta può essere usata al massimo una volta.
- Le corrispondenze esatte (`"2"`) vengono assegnate per prime, poi le corrispondenze parziali (`"1"`) vengono assegnate da sinistra a destra per le lettere rimanenti.
- Se una lettera compare più volte nel tentativo, può corrispondere solo tante volte quante ne compaiono nella parola segreta.

Per esempio, data una parola segreta `"APPLE"` e un tentativo `"POPPA"`, restituisci `"10201"`:

La prima `"P"` non è nella posizione corretta (`"1"`), la `"O"` non è nella parola segreta (`"0"`), la seconda `"P"` è nella posizione corretta (`"2"`), la terza `"P"` è uno zero (`"0"`) perché le due `"P"` nella parola segreta sono già state usate, e la `"A"` non è nella posizione corretta (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` dovrebbe restituire `"10201"`.

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")` dovrebbe restituire `"11221"`.

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")` dovrebbe restituire `"000000"`.

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")` dovrebbe restituire `"0000222222"`.

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")` dovrebbe restituire `"110200"`.

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")` dovrebbe restituire `"10021000"`.

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```
