---
id: 69162d64f96574d9bb629efe
title: "Sfida 113: miglia in chilometri"
challengeType: 28
dashedName: challenge-113
---

# --description--

Dato un numero che rappresenta una distanza in miglia, restituisci la distanza equivalente in chilometri.

- L'input sarà sempre un numero non negativo.
- 1 miglio equivale a 1,60934 chilometri.
- Arrotonda il risultato a due cifre decimali.
- Rimuovi gli zeri finali non necessari dal risultato arrotondato.

# --hints--

`convertToKm(1)` dovrebbe restituire `1.61`.

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` dovrebbe restituire `33.8`.

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` dovrebbe restituire `5.63`.

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` dovrebbe restituire `0`.

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` dovrebbe restituire `1`.

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
