---
id: 691b559495c5cb5a37b9b484
title: "Sfida 124: aggiornamento dell'inventario"
challengeType: 28
dashedName: challenge-124
---

# --description--

Dato un array 2D che rappresenta l'inventario del tuo archivio e un altro array 2D che rappresenta una spedizione ricevuta, restituisci il tuo inventario aggiornato.

- Ogni elemento negli array avrà il formato: `[quantity, "item"]`, dove `quantity` è un intero e `"item"` è una stringa.
- Aggiorna gli elementi nell'inventario aggiungendo la quantità di eventuali elementi corrispondenti dalla spedizione.
- Se un elemento ricevuto non esiste nell'inventario attuale, aggiungilo come nuova voce alla fine dell'inventario.
- Restituisci l'inventario nell'ordine in cui è stato fornito, con i nuovi elementi alla fine nell'ordine in cui appaiono nella spedizione.

Per esempio, dato un inventario di `[[2, "apples"], [5, "bananas"]]` e una spedizione di `[[1, "apples"], [3, "bananas"]]`, restituisci `[[3, "apples"], [8, "bananas"]]`.

# --hints--

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` dovrebbe restituire `[[3, "apples"], [8, "bananas"]]`.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]]);
```

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` dovrebbe restituire `[[3, "apples"], [8, "bananas"], [4, "oranges"]]`.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]]);
```

`updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` dovrebbe restituire `[[10, "apples"], [30, "bananas"], [20, "oranges"]]`.

```js
assert.deepEqual(updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]]);
```

`updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` dovrebbe restituire `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]`.

```js
assert.deepEqual(updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]);
```

# --seed--

## --seed-contents--

```js
function updateInventory(inventory, shipment) {
  return inventory;
}
```

# --solutions--

```js
function updateInventory(inventory, shipment) {
  const inventoryMap = new Map();
  inventory.forEach(([qty, item], index) => {
    inventoryMap.set(item, index);
  });

  shipment.forEach(([qty, item]) => {
    if (inventoryMap.has(item)) {
      const index = inventoryMap.get(item);
      inventory[index][0] += qty;
    } else {
      inventory.push([qty, item]);
      inventoryMap.set(item, inventory.length - 1);
    }
  });

  return inventory;
}
```
