---
id: 691b559495c5cb5a37b9b489
title: "Sfida 129: parser di citazioni in blocco Markdown"
challengeType: 28
dashedName: challenge-129
---

# --description--

Data una stringa che include una citazione in blocco in Markdown, restituisci la stringa HTML equivalente.

Una citazione in blocco in Markdown è qualsiasi riga che:

- Inizia con zero o più spazi
- Seguita da un segno maggiore di (`>`)
- Poi, uno o più spazi
- E infine, il testo della citazione in blocco.

Restituisci il testo della citazione in blocco racchiuso dai tag HTML `blockquote` di apertura e chiusura.

Per esempio, data `"> This is a quote"`, restituisci `<blockquote>This is a quote</blockquote>`.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parseBlockquote("> This is a quote")` dovrebbe restituire `"<blockquote>This is a quote</blockquote>"`.

```js
assert.equal(parseBlockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>");
```

`parseBlockquote(" > This is also a quote")` dovrebbe restituire `"<blockquote>This is also a quote</blockquote>"`.

```js
assert.equal(parseBlockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>");
```

`parseBlockquote("  >    So  Is  This")` dovrebbe restituire `"<blockquote>So  Is  This</blockquote>"`.

```js
assert.equal(parseBlockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>");
```

# --seed--

## --seed-contents--

```js
function parseBlockquote(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBlockquote(markdown) {
  const blockquoteRegex = /^\s*>\s+(.+)$/;
  const match = markdown.match(blockquoteRegex);

  const text = match[1];
  return `<blockquote>${text}</blockquote>`;
}
```
