---
id: 691b559495c5cb5a37b9b482
title: "Sfida 122: parser per grassetto in markdown"
challengeType: 28
dashedName: challenge-122
---

# --description--

Data una stringa che può includere del testo in grassetto in Markdown, restituisci la stringa HTML equivalente.

- Il testo in grassetto in Markdown è qualsiasi testo che inizia e finisce con due asterischi (`**`) o due underscore (`__`).
- Non ci possono essere spazi tra il testo e gli asterischi o gli underscore, ma possono esserci spazi all’interno del testo stesso.
- Converti tutte le occorrenze di grassetto in tag HTML `b` e restituisci la stringa.

Per esempio, data `"**This is bold**"`, restituisci `"<b>This is bold</b>"`.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parseBold("**This is bold**")` dovrebbe restituire `"<b>This is bold</b>"`.

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` dovrebbe restituire `"<b>This is also bold</b>"`.

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` dovrebbe restituire `"**This is not bold **"`.

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` dovrebbe restituire `"__ This is also not bold__"`.

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` dovrebbe restituire `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
