---
id: 691b559495c5cb5a37b9b480
title: "Sfida 120: da libbre a chilogrammi"
challengeType: 28
dashedName: challenge-120
---

# --description--

Dato un peso in libbre come numero, restituisci la stringa `"(lbs) pounds equals (kgs) kilograms."`.

- Sostituisci `"(lbs)"` con il numero di input.
- Sostituisci `"(kgs)"` con l'input convertito in chilogrammi, arrotondato a due decimali e includendo sempre due cifre decimali nel valore.
- 1 libbra equivale a 0,453592 chilogrammi.
- Se l'input è `1`, usa `"pound"` invece di `"pounds"`.
- Se il valore convertito è `1`, usa `"kilogram"` invece di `"kilograms"`.

# --hints--

`convertToKgs(1)` dovrebbe restituire `"1 pound equals 0.45 kilograms."`.

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` dovrebbe restituire `"0 pounds equals 0.00 kilograms."`.

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` dovrebbe restituire `"100 pounds equals 45.36 kilograms."`.

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` dovrebbe restituire `"2.5 pounds equals 1.13 kilograms."`.

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` dovrebbe restituire `"2.20462 pounds equals 1.00 kilogram."`.

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
