---
id: 691b559495c5cb5a37b9b487
title: "Sfida 127: controllo della velocità"
challengeType: 28
dashedName: challenge-127
---

# --description--

Dato la velocità a cui stai viaggiando in miglia orarie (MPH) e un limite di velocità in chilometri orari (KPH), determina se stai superando il limite e se riceverai un avviso o una multa.

- 1 miglio equivale a 1,60934 chilometri.
- Se stai viaggiando a una velocità inferiore o uguale al limite, restituisci `"Not Speeding"`.
- Se stai viaggiando fino a 5 KPH oltre il limite, restituisci `"Warning"`.
- Se stai viaggiando a più di 5 KPH oltre il limite, restituisci `"Ticket"`.

# --hints--

`speedCheck(30, 70)` dovrebbe restituire `"Not Speeding"`.

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` dovrebbe restituire `"Warning"`.

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` dovrebbe restituire `"Not Speeding"`.

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` dovrebbe restituire `"Ticket"`.

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` dovrebbe restituire `"Warning"`.

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` dovrebbe restituire `"Ticket"`.

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
