---
id: 691b559495c5cb5a37b9b486
title: "Sfida 126: metti in maiuscolo la prima lettera"
challengeType: 28
dashedName: challenge-126
---

# --description--

Data una stringa title, restituisci una nuova stringa formattata in stile titolo seguendo queste regole:

- Metti in maiuscolo la prima lettera di ogni parola.
- Rendi tutte le altre lettere di ogni parola minuscole.
- Le parole sono sempre separate da uno spazio singolo.

# --hints--

`titleCase("hello world")` dovrebbe restituire `"Hello World"`.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` dovrebbe restituire `"The Quick Brown Fox"`.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` dovrebbe restituire `"Javascript And Python"`.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` dovrebbe restituire `"Avocado Toast For Breakfast"`.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
