---
id: 691b559495c5cb5a37b9b483
title: "Sfida 123: costruisci numeri romani"
challengeType: 28
dashedName: challenge-123
---

# --description--

Dato un numero intero, restituisci il suo valore equivalente in numeri romani.

I numeri romani usano questi simboli:

| Simbolo | Valore |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

I numeri romani si scrivono dal più grande al più piccolo, da sinistra a destra, seguendo queste regole:

- Si usa l'addizione quando un simbolo è seguito da uno di valore uguale o inferiore. Per esempio, `18` si scrive come `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- Si usa la sottrazione quando un simbolo più piccolo precede uno più grande, per rappresentare 4 o 9 in qualsiasi posizione. Per esempio, 19 si scrive come `XIX` (10 + (10 - 1)).
- Nessun simbolo può essere ripetuto più di tre volte di seguito. Si usa la sottrazione quando altrimenti dovresti scrivere un simbolo più di tre volte di fila. Quindi il numero più grande che puoi scrivere è 3999.

Ecco un altro esempio: dato `1464`, restituisci `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`toRoman(18)` dovrebbe restituire `"XVIII"`.

```js
assert.equal(toRoman(18), "XVIII");
```

`toRoman(19)` dovrebbe restituire `"XIX"`.

```js
assert.equal(toRoman(19), "XIX");
```

`toRoman(1464)` dovrebbe restituire `"MCDLXIV"`.

```js
assert.equal(toRoman(1464), "MCDLXIV");
```

`toRoman(2025)` dovrebbe restituire `"MMXXV"`.

```js
assert.equal(toRoman(2025), "MMXXV");
```

`toRoman(3999)` dovrebbe restituire `"MMMCMXCIX"`.

```js
assert.equal(toRoman(3999), "MMMCMXCIX");
```

# --seed--

## --seed-contents--

```js
function toRoman(num) {
  return num;
}
```

# --solutions--

```js
function toRoman(num) {
  const symbols = [
    ["M", 1000],
    ["CM", 900],
    ["D", 500],
    ["CD", 400],
    ["C", 100],
    ["XC", 90],
    ["L", 50],
    ["XL", 40],
    ["X", 10],
    ["IX", 9],
    ["V", 5],
    ["IV", 4],
    ["I", 1]
  ];

  let result = "";
  for (const [sym, val] of symbols) {
    while (num >= val) {
      result += sym;
      num -= val;
    }
  }
  return result;
}
```
