---
id: 691b5597f0f3e85a588a5c94
title: "Sfida 130: scacchiera"
challengeType: 28
dashedName: challenge-130
---

# --description--

Dato un array con due numeri, il primo indica il numero di righe e il secondo il numero di colonne, restituisci una matrice (un array di array) riempita con i caratteri `"X"` e `"O"` della dimensione indicata.

- I caratteri devono alternarsi come in una scacchiera.
- La cella in alto a sinistra deve essere sempre `"X"`.

Per esempio, dato `[3, 3]`, restituisci:

```json
[
  ["X", "O", "X"],
  ["O", "X", "O"],
  ["X", "O", "X"]
]
```

# --hints--

`createBoard([3, 3])` dovrebbe restituire `[["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]`.

```js
assert.deepEqual(createBoard([3, 3]), [["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]);
```

`createBoard([6, 1])` dovrebbe restituire `[["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]`.

```js
assert.deepEqual(createBoard([6, 1]), [["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]);
```

`createBoard([2, 10])` dovrebbe restituire `[["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]`.

```js
assert.deepEqual(createBoard([2, 10]), [["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]);
```

`createBoard([5, 4])` dovrebbe restituire `[["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]`.

```js
assert.deepEqual(createBoard([5, 4]), [["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]);
```

# --seed--

## --seed-contents--

```js
function createBoard(dimensions) {
  return dimensions;
}
```

# --solutions--

```js
function createBoard(dimensions) {
  const [rows, cols] = dimensions;
  const board = [];

  for (let r = 0; r < rows; r++) {
    const row = [];
    for (let c = 0; c < cols; c++) {
      row.push((r + c) % 2 === 0 ? "X" : "O");
    }
    board.push(row);
  }

  return board;
}
```
