---
id: 691b559495c5cb5a37b9b485
title: "Sfida 125: gioco della vita"
challengeType: 28
dashedName: challenge-125
---

# --description--

Data una matrice (array di array) che rappresenta lo stato attuale nel gioco della vita di Conway, restituisci il prossimo stato della matrice usando queste regole:

- Ogni cella è o `1` (viva) o `0` (morta).
- I vicini di una cella sono fino a otto celle circostanti (verticalmente, orizzontalmente e diagonalmente).
- Le celle ai bordi hanno meno di otto vicini.

Regole per aggiornare ogni cella:

- Qualsiasi cella viva con meno di due vicini vivi muore (sottopopolazione).
- Qualsiasi cella viva con due o tre vicini vivi sopravvive.
- Qualsiasi cella viva con più di tre vicini vivi muore (sovrappopolazione).
- Qualsiasi cella morta con esattamente tre vicini vivi diventa viva (riproduzione).

Per esempio, dato:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

restituisci:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Ogni cella si aggiorna in base al numero di vicini vivi. Per esempio, `[0][0]` rimane morta (2 vicini vivi), `[0][1]` rimane viva (2 vicini vivi), `[0][2]` muore (3 vicini vivi), e così via.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` dovrebbe restituire `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` dovrebbe restituire `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` dovrebbe restituire `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` dovrebbe restituire `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
