---
id: 691b559495c5cb5a37b9b488
title: "Sfida 128: conta le consonanti"
challengeType: 28
dashedName: challenge-128
---

# --description--

Data una stringa e un numero obiettivo, determina se la stringa contiene esattamente il numero obiettivo di consonanti.

- Le consonanti sono tutti i caratteri alfabetici tranne `"a"`, `"e"`, `"i"`, `"o"` e `"u"` in qualsiasi caso.
- Ignora cifre, punteggiatura, spazi e altri caratteri non alfabetici durante il conteggio.

# --hints--

`hasConsonantCount("helloworld", 7)` dovrebbe restituire `true`.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` dovrebbe restituire `false`.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` dovrebbe restituire `true`.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` dovrebbe restituire `false`.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` dovrebbe restituire `true`.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
