---
id: 6925e2068081f40f549ced1d
title: "Sfida 139: carta, forbice, sasso"
challengeType: 28
dashedName: challenge-139
---

# --description--

Date due stringhe, la prima che rappresenta il giocatore 1 e la seconda che rappresenta il giocatore 2, determina il vincitore di una partita di carta, forbice, sasso.

- Le stringhe di input saranno sempre `"Rock"`, `"Paper"` o `"Scissors"`.
- `"Rock"` batte `"Scissors"`.
- `"Paper"` batte `"Rock"`.
- `"Scissors"` batte `"Paper"`.

Restituisci:

- `"Player 1 wins"` se vince il giocatore 1.
- `"Player 2 wins"` se vince il giocatore 2.
- `"Tie"` se entrambi i giocatori scelgono la stessa opzione.

# --hints--

`rockPaperScissors("Rock", "Rock")` dovrebbe restituire `"Tie"`.

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` dovrebbe restituire `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` dovrebbe restituire `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` dovrebbe restituire `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` dovrebbe restituire `"Tie"`.

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` dovrebbe restituire `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
