---
id: 6925e2068081f40f549ced1c
title: "Sfida 138: somma dei divisori"
challengeType: 28
dashedName: challenge-138
---

# --description--

Dato un numero intero positivo, restituisci la somma di tutti i suoi divisori.

- Un divisore è un numero intero che divide il numero senza resto (il resto è `0`).
- Conta ogni divisore una sola volta.

Per esempio, dato `6`, restituisci `12` perché i divisori di `6` sono `1`, `2`, `3` e `6`, e la somma di questi è `12`.

# --hints--

`sumDivisors(6)` dovrebbe restituire `12`.

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)` dovrebbe restituire `14`.

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)` dovrebbe restituire `56`.

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)` dovrebbe restituire `224`.

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)` dovrebbe restituire `806`.

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)` dovrebbe restituire `23520`.

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
