---
id: 6925e2068081f40f549ced1b
title: "Sfida 137: generatore di fiocchi di neve"
challengeType: 28
dashedName: challenge-137
---

# --description--

Dato un stringa multilinea che usa caratteri di nuova riga (`\n`) per rappresentare un'interruzione di riga, restituisci una nuova stringa in cui ogni riga è specchiata orizzontalmente e allegata alla fine della riga originale.

- Specchia una riga invertendo tutti i suoi caratteri, inclusi gli spazi.

Per esempio, dato `"* \n *\n* "`, che stampa sulla console:

```sh
* 
 *
* 
```

Restituisci `"*  *\n ** \n*  *"`, che stampa sulla console:

```sh
*  *
 ** 
*  *
```

Fai molta attenzione agli spazi nelle stringhe date e restituite. Assicurati di non eliminarne nessuno.

# --hints--

`generateSnowflake("* \n *\n* ")` dovrebbe restituire `"*  *\n ** \n*  *"`.

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` dovrebbe restituire `"X=~~=X"`.

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` dovrebbe restituire `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` dovrebbe restituire `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` dovrebbe restituire `"*  --  *\n * -- * \n*  --  *"`.

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
