---
id: 6925e2068081f40f549ced1a
title: "Sfida 136: parser di immagini in markdown"
challengeType: 28
dashedName: challenge-136
---

# --description--

Dato un stringa che rappresenta un'immagine in Markdown, restituisci la stringa HTML equivalente.

Un'immagine in Markdown ha il seguente formato: `"![alt text](image_url)"`. Dove:

- `alt text` è la descrizione dell'immagine (il valore dell'attributo `alt`).
- `image_url` è l'URL sorgente dell'immagine (il valore dell'attributo `src`).

Restituisci una stringa del tag HTML `img` con l'insieme `src` impostato sull'URL dell'immagine e l'insieme `alt` impostato sul testo alternativo.

Per esempio, dato `"![Cute cat](cat.png)"` restituisci `'<img src="cat.png" alt="Cute cat">'`;

- Assicurati che il tag, l'ordine degli attributi, gli spazi e l'utilizzo delle virgolette siano uguali a quelli sopra.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi — controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parseImage("![Cute cat](cat.png)")` dovrebbe restituire `'<img src="cat.png" alt="Cute cat">'`.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` dovrebbe restituire `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` dovrebbe restituire `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
