---
id: 69272dcf1c24b44fd79137c6
title: "Sfida 143: parser per corsivo in markdown"
challengeType: 28
dashedName: challenge-143
---

# --description--

Dato un testo che può includere del testo in corsivo in Markdown, restituisci la stringa HTML equivalente.

- Il testo in corsivo in Markdown è qualsiasi testo che inizia e finisce con un singolo asterisco (`*`) o un singolo underscore (`_`).
- Non ci possono essere spazi tra il testo e l’asterisco o l’underscore, ma possono esserci spazi all’interno del testo stesso.
- Converti tutte le occorrenze in corsivo in tag HTML `i` e restituisci la stringa.

Per esempio, dato `"*This is italic*"`, restituisci `"<i>This is italic</i>"`.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parseItalics("*This is italic*")` dovrebbe restituire `"<i>This is italic</i>"`.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` dovrebbe restituire `"<i>This is also italic</i>"`.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` dovrebbe restituire `"*This is not italic *"`.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` dovrebbe restituire `"_ This is also not italic_"`.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` dovrebbe restituire `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
