---
id: 69272dcf1c24b44fd79137c4
title: "Sfida 141: carburante per il decollo"
challengeType: 28
dashedName: challenge-141
---

# --description--

Dato il numero di galloni di carburante attualmente nel tuo aereo e il numero richiesto di litri di carburante per raggiungere la destinazione, determina quanti galloni di carburante aggiuntivi dovresti aggiungere.

- 1 gallone equivale a 3,78541 litri.
- Se l'aereo ha già abbastanza carburante, restituisci `0`.
- Puoi aggiungere solo galloni interi.
- Non includere decimali nel numero restituito.

# --hints--

`fuelToAdd(0, 1)` dovrebbe restituire `1`.

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)` dovrebbe restituire `6`.

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)` dovrebbe restituire `0`.

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)` dovrebbe restituire `4520`.

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)` dovrebbe restituire `12209`.

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
