---
id: 69272dcf1c24b44fd79137c3
title: "Sfida 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

Data una stringa che rappresenta il nome di una variabile, restituisci il nome della variabile convertito in SCREAMING_SNAKE_CASE.

I nomi delle variabili forniti saranno scritti in uno dei seguenti formati:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

Nei formati sopra, le parole sono separate da un underscore (`_`), un trattino (`-`) oppure una nuova parola inizia con una lettera maiuscola.

Per convertire in SCREAMING_SNAKE_CASE:

- Trasforma tutte le lettere in maiuscolo
- Separa le parole con un underscore (`_`)

# --hints--

`toScreamingSnakeCase("userEmail")` dovrebbe restituire `"USER_EMAIL"`.

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` dovrebbe restituire `"USER_PASSWORD"`.

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` dovrebbe restituire `"USER_ID"`.

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` dovrebbe restituire `"USER_ADDRESS"`.

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` dovrebbe restituire `"USERNAME"`.

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

`toScreamingSnakeCase("my_variable_name")` dovrebbe restituire `"MY_VARIABLE_NAME"`.

```js
assert.equal(toScreamingSnakeCase("my_variable_name"), "MY_VARIABLE_NAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();
}
```
