---
id: 69272dcf1c24b44fd79137c5
title: "Sfida 142: somma la stringa"
challengeType: 28
dashedName: challenge-142
---

# --description--

Data una stringa contenente cifre e altri caratteri, restituisci la somma di tutti i numeri presenti nella stringa.

- Considera le cifre consecutive come un unico numero. Per esempio, `"13"` conta come 13, non come 1 + 3.
- Ignora tutti i caratteri non numerici.

# --hints--

`stringSum("3apples2bananas")` dovrebbe restituire `5`.

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` dovrebbe restituire `17`.

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` dovrebbe restituire `125344`.

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` dovrebbe restituire `321`.

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` dovrebbe restituire `1653`.

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
