---
id: 69306364df283fcaff2e1ad6
title: "Sfida 145: n-esimo numero di Fibonacci"
challengeType: 28
dashedName: challenge-145
---

# --description--

Dato un intero `n`, restituisci il `n`-esimo numero nella sequenza di Fibonacci.

La sequenza di Fibonacci è una serie di numeri in cui ogni numero è la somma dei due precedenti. I primi 10 numeri della sequenza sono `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

# --hints--

`nthFibonacci(4)` dovrebbe restituire `2`.

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` dovrebbe restituire `34`.

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` dovrebbe restituire `377`.

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` dovrebbe restituire `63245986`.

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` dovrebbe restituire `1304969544928657`.

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
