---
id: 69306364df283fcaff2e1ada
title: "Sfida 149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

Data una stringa, restituisci una nuova stringa in cui tutte le vocali sono convertite in maiuscolo e tutti gli altri caratteri alfabetici sono convertiti in minuscolo.

- Le vocali sono `"a"`, `"e"`, `"i"`, `"o"` e `"u"` in qualsiasi caso.
- I caratteri non alfabetici devono rimanere invariati.

# --hints--

`vowelCase("vowelcase")` dovrebbe restituire `"vOwElcAsE"`.

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` dovrebbe restituire `"cOdIng Is fUn"`.

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` dovrebbe restituire `"hEllO, wOrld!"`.

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` dovrebbe restituire `"gIt chErry-pIck"`.

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` dovrebbe restituire `"hEAd~1"`.

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
