---
id: 69306364df283fcaff2e1ad9
title: "Sfida 148: pressione degli pneumatici"
challengeType: 28
dashedName: challenge-148
---

# --description--

Dato un array con quattro numeri che rappresentano la pressione degli pneumatici in psi dei quattro pneumatici del tuo veicolo, e un altro array con due numeri che rappresentano la pressione minima e massima per i tuoi pneumatici in bar, restituisci un array di quattro stringhe che descrivano lo stato di ciascun pneumatico.

- 1 bar equivale a 14,5038 psi.

Restituisci un array con i seguenti valori per ogni pneumatico:

- `"Low"` se la pressione dello pneumatico è inferiore al minimo consentito.
- `"Good"` se è compresa tra il minimo e il massimo consentito.
- `"High"` se è superiore al massimo consentito.

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])` dovrebbe restituire `["Good", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])` dovrebbe restituire `["Good", "Low", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])` dovrebbe restituire `["Low", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])` dovrebbe restituire `["High", "High", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])` dovrebbe restituire `["Good", "Good", "Good", "Good"]`.

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
