---
id: 69306364df283fcaff2e1ad5
title: "Sfida 144: serie di risoluzioni"
challengeType: 28
dashedName: challenge-144
---

# --description--

Dato un array di array, dove ogni sotto-array rappresenta un giorno delle tue attività di risoluzione e contiene tre numeri interi: il numero di passi fatti quel giorno, i minuti di tempo davanti allo schermo e il numero di pagine lette; determina se stai mantenendo le tue risoluzioni.

- Il primo sotto-array è il giorno 1, il secondo il giorno 2 e così via.

Un giorno è considerato riuscito se tutti e tre i seguenti obiettivi sono raggiunti:

- Hai camminato almeno 10.000 passi.
- Hai avuto non più di 120 minuti di tempo davanti allo schermo.
- Hai letto almeno cinque pagine.

Se tutti i giorni dati sono riusciti, restituisci `"Resolution on track: N day streak."` dove `N` è il numero di giorni riusciti.

Se uno o più giorni non sono riusciti, restituisci `"Resolution failed on day X: N day streak."`. Dove `X` è il numero del primo giorno non riuscito e `N` è il numero di giorni riusciti prima del primo giorno non riuscito.

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` dovrebbe restituire `"Resolution on track: 3 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])` dovrebbe restituire `"Resolution failed on day 2: 1 day streak."`

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` dovrebbe restituire `"Resolution failed on day 3: 2 day streak."`

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` dovrebbe restituire `"Resolution failed on day 4: 3 day streak."`

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` dovrebbe restituire `"Resolution on track: 10 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```
