---
id: 69306364df283fcaff2e1ad8
title: "Sfida 147: calcolatore di anni bisestili"
challengeType: 28
dashedName: challenge-147
---

# --description--

Dato un anno intero, determina se è un anno bisestile.

Un anno è bisestile se soddisfa le seguenti regole:

- L'anno è divisibile per 4 senza resto, e
- L'anno non è divisibile per 100 senza resto, a meno che
- L'anno sia divisibile per 400 senza resto.

# --hints--

`isLeapYear(2024)` dovrebbe restituire `true`.

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` dovrebbe restituire `false`.

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` dovrebbe restituire `false`.

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` dovrebbe restituire `true`.

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` dovrebbe restituire `false`.

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` dovrebbe restituire `true`.

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` dovrebbe restituire `false`.

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
