---
id: 69306364df283fcaff2e1ad7
title: "Sfida 146: sedile per mancini a tavola"
challengeType: 28
dashedName: challenge-146
---

# --description--

Data una matrice 4x2 (array di array) che rappresenta la disposizione dei posti a tavola per un pasto, determina quanti sedili sono disponibili per una persona mancina.

- Una persona mancina non può sedersi dove a sinistra immediata ci sarebbe una persona destrimane.

Nella matrice data:

- Un `"R"` è un sedile occupato da una persona destrimane.
- Un `"L"` è un sedile occupato da una persona mancina.
- Un `"U"` è un sedile libero.
- Solo i sedili liberi sono disponibili per sedersi.
- I sedili nella fila superiore sono rivolti verso il basso, e quelli nella fila inferiore sono rivolti verso l’alto (come a una tavola), quindi sinistra e destra sono relativi all’orientamento del sedile.
- I sedili agli angoli hanno solo un sedile accanto.

Per esempio, nella matrice data:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Il sedile in alto a sinistra non può essere occupato perché c’è una persona destrimane a sinistra. Gli altri due sedili liberi possono essere occupati perché non c’è una persona destrimane a sinistra.

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` dovrebbe restituire `2`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` dovrebbe restituire `8`.

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` dovrebbe restituire `0`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` dovrebbe restituire `1`.

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` dovrebbe restituire `5`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```
