---
id: 69373793f5a867f769cde137
title: "Sfida 152: numero primo circolare"
challengeType: 28
dashedName: challenge-152
---

# --description--

Dato un numero intero, determina se è un numero primo circolare.

Un numero primo circolare è un intero per cui tutte le rotazioni delle sue cifre sono a loro volta numeri primi.

Per esempio, `197` è un numero primo circolare perché tutte le rotazioni delle sue cifre: `197`, `971` e `719`, sono numeri primi.

# --hints--

`isCircularPrime(197)` dovrebbe restituire `true`.

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)` dovrebbe restituire `false`.

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)` dovrebbe restituire `true`.

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)` dovrebbe restituire `false`.

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)` dovrebbe restituire `true`.

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
