---
id: 69373793f5a867f769cde138
title: "Sfida 153: tris"
challengeType: 28
dashedName: challenge-153
---

# --description--

Data una matrice 3×3 (un array di array) che rappresenta una partita di tris completata, determina il vincitore.

- Ogni elemento nella matrice data è o `"X"` o `"O"`.

Un giocatore vince se ha tre dei suoi caratteri in fila — orizzontalmente, verticalmente o diagonalmente.

Restituisci:

- `"X wins"` se il giocatore X ha tre in fila.
- `"O wins"` se il giocatore O ha tre in fila.
- `"Draw"` se nessun giocatore ha tre in fila.

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` dovrebbe restituire `"X wins"`.

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])` dovrebbe restituire `"X wins"`.

```js
assert.equal(ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` dovrebbe restituire `"Draw"`.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]])` dovrebbe restituire `"O wins"`.

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` dovrebbe restituire `"X wins"`.

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` dovrebbe restituire `"Draw"`.

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```
