---
id: 69373793f5a867f769cde135
title: "Sfida 150: parser di lista non ordinata in markdown"
challengeType: 28
dashedName: challenge-150
---

# --description--

Data una stringa che rappresenta una lista non ordinata valida in Markdown, restituisci la stringa HTML equivalente.

Una lista non ordinata è composta da uno o più elementi di lista. Un elemento di lista valido appare su una riga a sé stante e:

- Inizia con un trattino (`"-"`), seguito da
- Almeno uno spazio, e poi
- Il testo dell'elemento di lista.

La lista è fornita come una singola stringa con le righe separate dal carattere di nuova riga (`"\n"`). Non includere i caratteri di nuova riga nel testo dell'elemento.

Racchiudi ogni elemento di lista nei tag HTML `li` e l'intera lista di elementi nei tag `ul`.

Per esempio, data `"- Item A\n- Item B"`, restituisci `"<ul><li>Item A</li><li>Item B</li></ul>"`.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parseUnorderedList("- Item A\n- Item B")` dovrebbe restituire `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")` dovrebbe restituire `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` dovrebbe restituire `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")` dovrebbe restituire `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
