---
id: 69373793f5a867f769cde13a
title: "Sfida 155: pianta il raccolto"
challengeType: 28
dashedName: challenge-155
---

# --description--

Dato un numero intero che rappresenta la dimensione del tuo campo agricolo, e `"acres"` o `"hectares"` che rappresentano l'unità di misura della dimensione del tuo campo, e un tipo di coltura, determina quante piante di quel tipo puoi piantare nel tuo campo.

- 1 acro corrisponde a 4046,86 metri quadrati.
- 1 ettaro corrisponde a 10.000 metri quadrati.

Ecco una lista di colture che saranno fornite come input e lo spazio che occupa una singola pianta:

|Coltura|Spazio per pianta|
|-|-|
|`"corn"`|1 metro quadrato|
|`"wheat"`|0,1 metri quadrati|
|`"soybeans"`|0,5 metri quadrati|
|`"tomatoes"`|0,25 metri quadrati|
|`"lettuce"`|0,2 metri quadrati|

Restituisci il numero di piante che possono essere piantate nel campo, arrotondato per difetto all'intero più vicino.

# --hints--

`getNumberOfPlants(1, "acres", "corn")` dovrebbe restituire `4046`.

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` dovrebbe restituire `100000`.

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` dovrebbe restituire `161874`.

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` dovrebbe restituire `150000`.

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` dovrebbe restituire `271139`.

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
