---
id: 69373793f5a867f769cde136
title: "Sfida 151: array ordinato?"
challengeType: 28
dashedName: challenge-151
---

# --description--

Dato un array di numeri, determina se i numeri sono ordinati in ordine crescente, decrescente o nessuno dei due.

Se l'array dato è:

- In ordine crescente (dal più basso al più alto), restituisci `"Ascending"`.
- In ordine decrescente (dal più alto al più basso), restituisci `"Descending"`.
- Non ordinato né in ordine crescente né decrescente, restituisci `"Not sorted"`.

# --hints--

`isSorted([1, 2, 3, 4, 5])` dovrebbe restituire `"Ascending"`.

```js
assert.equal(isSorted([1, 2, 3, 4, 5]), "Ascending");
```

`isSorted([10, 8, 6, 4, 2])` dovrebbe restituire `"Descending"`.

```js
assert.equal(isSorted([10, 8, 6, 4, 2]), "Descending");
```

`isSorted([1, 3, 2, 4, 5])` dovrebbe restituire `"Not sorted"`.

```js
assert.equal(isSorted([1, 3, 2, 4, 5]), "Not sorted");
```

`isSorted([3.14, 2.71, 1.61, 0.57])` dovrebbe restituire `"Descending"`.

```js
assert.equal(isSorted([3.14, 2.71, 1.61, 0.57]), "Descending");
```

`isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])` dovrebbe restituire `"Ascending"`.

```js
assert.equal(isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending");
```

`isSorted([0.4, 0.5, 0.3])` dovrebbe restituire `"Not sorted"`.

```js
assert.equal(isSorted([0.4, 0.5, 0.3]), "Not sorted");
```

# --seed--

## --seed-contents--

```js
function isSorted(arr) {

  return arr;
}
```

# --solutions--

```js
function isSorted(arr) {
  let ascending = true;
  let descending = true;

  for (let i = 1; i < arr.length; i++) {
    if (arr[i] < arr[i - 1]) ascending = false;
    if (arr[i] > arr[i - 1]) descending = false;
  }

  if (ascending) return "Ascending";
  if (descending) return "Descending";
  return "Not sorted";
}
```
