---
id: 69373793f5a867f769cde139
title: "Sfida 154: par per la buca"
challengeType: 28
dashedName: challenge-154
---

# --description--

Dati due numeri interi, il par per una buca da golf e il numero di tratti che un golfista ha impiegato su quella buca, restituisci il punteggio del golfista usando i termini del golf.

Restituisci:

- `"Hole in one!"` se ha impiegato un tratto.
- `"Eagle"` se ha impiegato due tratti in meno rispetto al par.
- `"Birdie"` se ha impiegato un tratto in meno rispetto al par.
- `"Par"` se ha impiegato lo stesso numero di tratti del par.
- `"Bogey"` se ha impiegato un tratto in più rispetto al par.
- `"Double bogey"` se ha impiegato due tratti in più rispetto al par.

# --hints--

`golfScore(3, 3)` dovrebbe restituire `"Par"`.

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` dovrebbe restituire `"Birdie"`.

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` dovrebbe restituire `"Hole in one!"`.

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` dovrebbe restituire `"Double bogey"`.

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` dovrebbe restituire `"Bogey"`.

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` dovrebbe restituire `"Eagle"`.

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
