---
id: 6939b873185d8e00d453563e
title: "Sfida 159: ipotenusa intera"
challengeType: 28
dashedName: challenge-159
---

# --description--

Dati due numeri interi positivi che rappresentano le lunghezze dei due cateti (i due lati corti) di un triangolo rettangolo, determina se l'ipotenusa è un numero intero.

La lunghezza dell'ipotenusa si calcola sommando i quadrati delle lunghezze dei due cateti e poi prendendo la radice quadrata di quel totale (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

`isIntegerHypotenuse(3, 4)` dovrebbe restituire `true`.

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` dovrebbe restituire `false`.

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` dovrebbe restituire `true`.

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` dovrebbe restituire `false`.

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` dovrebbe restituire `true`.

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` dovrebbe restituire `false`.

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
