---
id: 6939b873185d8e00d453563f
title: "Sfida 160: mosse del cavallo"
challengeType: 28
dashedName: challenge-160
---

# --description--

Data la posizione di un cavallo su una scacchiera, restituisci il numero di caselle valide in cui il cavallo può muoversi.

Una scacchiera standard è 8x8, con le colonne etichettate `A` fino a `H` (da sinistra a destra) e le righe etichettate `1` fino a `8` (dal basso verso l'alto). Si presenta così:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Un cavallo si muove a "L": due caselle in una direzione (orizzontale o verticale) e una casella nella direzione perpendicolare.

Questo significa che un cavallo può muoversi fino a otto posizioni possibili, ma meno se si trova vicino ai bordi della scacchiera. Per esempio, se un cavallo fosse in `A1`, potrebbe muoversi solo in `B3` o `C2`.

# --hints--

`knightMoves("A1")` dovrebbe restituire `2`.

```js
assert.equal(knightMoves("A1"), 2);
```

`knightMoves("D4")` dovrebbe restituire `8`.

```js
assert.equal(knightMoves("D4"), 8);
```

`knightMoves("G6")` dovrebbe restituire `6`.

```js
assert.equal(knightMoves("G6"), 6);
```

`knightMoves("B8")` dovrebbe restituire `3`.

```js
assert.equal(knightMoves("B8"), 3);
```

`knightMoves("H3")` dovrebbe restituire `4`.

```js
assert.equal(knightMoves("H3"), 4);
```

# --seed--

## --seed-contents--

```js
function knightMoves(position) {

  return position;
}
```

# --solutions--

```js
function knightMoves(position) {
  const col = position[0].charCodeAt(0) - 65;
  const row = parseInt(position[1]) - 1;

  const moves = [
    [2, 1], [2, -1],
    [-2, 1], [-2, -1],
    [1, 2], [1, -2],
    [-1, 2], [-1, -2]
  ];

  let valid = 0;

  for (const [dx, dy] of moves) {
    const newCol = col + dx;
    const newRow = row + dy;

    if (newCol >= 0 && newCol < 8 && newRow >= 0 && newRow < 8) {
      valid++;
    }
  }

  return valid;
}
```
