---
id: 6939b873185d8e00d453563d
title: "Sfida 158: scambio di array"
challengeType: 28
dashedName: challenge-158
---

# --description--

Dato un array con due valori, restituisci un array con i valori scambiati.

Per esempio, dato `["A", "B"]` restituisci `["B", "A"]`.

# --hints--

`arraySwap(["A", "B"])` dovrebbe restituire `["B", "A"]`.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` dovrebbe restituire `[20, 25]`.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` dovrebbe restituire `[false, true]`.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` dovrebbe restituire `[1, "1"]`.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
