---
id: 6939b873185d8e00d453563c
title: "Sfida 157: parser di collegamenti markdown"
challengeType: 28
dashedName: challenge-157
---

# --description--

Dato il contenuto di un collegamento in Markdown, restituisci la stringa HTML equivalente.

Un'immagine Markdown ha il seguente formato: `"[link_text](link_url)"`. Restituisci la stringa del tag HTML `a` con l'insieme `href` impostato su `link_url` e `link_text` come contenuto del tag.

Per esempio, dato `"[freeCodeCamp](https://freecodecamp.org/)"` restituisci `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi — controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")` dovrebbe restituire `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` dovrebbe restituire `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` dovrebbe restituire `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
