---
id: 6939b873185d8e00d453563b
title: "Sfida 156: dispari o pari?"
challengeType: 28
dashedName: challenge-156
---

# --description--

Dato un numero intero positivo, restituisci `"Odd"` se è un numero dispari e `"Even"` se è pari.

# --hints--

`oddOrEven(1)` dovrebbe restituire `"Odd"`.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` dovrebbe restituire `"Even"`.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` dovrebbe restituire `"Odd"`.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` dovrebbe restituire `"Even"`.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` dovrebbe restituire `"Odd"`.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
