---
id: 694596b0585c11170ac7c7fa
title: "Sfida 162: consumo di energia"
challengeType: 28
dashedName: challenge-162
---

# --description--

Dato il numero di calorie bruciate durante un allenamento e il numero di wattora consumati dai tuoi dispositivi elettronici durante quell’allenamento, determina quale dei due ha usato più energia.

Per confrontarli, converti entrambi i valori in joule usando le seguenti conversioni:

- 1 caloria equivale a 4184 joule.
- 1 wattora equivale a 3600 joule.

Restituisci:

- `"Workout"` se l’allenamento ha usato più energia.
- `"Devices"` se il dispositivo ha usato più energia.
- `"Equal"` se entrambi hanno usato la stessa quantità di energia.

# --hints--

`compareEnergy(250, 50)` dovrebbe restituire `"Workout"`.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` dovrebbe restituire `"Devices"`.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` dovrebbe restituire `"Equal"`.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` dovrebbe restituire `"Workout"`.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` dovrebbe restituire `"Devices"`.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` dovrebbe restituire `"Equal"`.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
