---
id: 694596b0585c11170ac7c7fd
title: "Sfida 165: media della classe"
challengeType: 28
dashedName: challenge-165
---

# --description--

Dato un array di voti d'esame (numeri), restituisci la media dei voti sotto forma di voto lettera secondo il seguente grafico:

| Media dei voti | Voto lettera |
| - | - |
| 97-100 | `"A+"` |
| 93-96 | `"A"` |
| 90-92 | `"A-"` |
| 87-89 | `"B+"` |
| 83-86 | `"B"` |
| 80-82 | `"B-"` |
| 77-79 | `"C+"` |
| 73–76 | `"C"` |
| 70-72 | `"C-"` |
| 67-69 | `"D+"` |
| 63-66 | `"D"` |
| 60–62 | `"D-"` |
| sotto 60 | `"F"` |

Calcola la media sommando tutti i voti nell'array e dividendo per il numero totale di voti.

# --hints--

`getAverageGrade([92, 91, 90, 94, 89, 93])` dovrebbe restituire `"A-"`.

```js
assert.equal(getAverageGrade([92, 91, 90, 94, 89, 93]), "A-");
```

`getAverageGrade([84, 89, 85, 100, 91, 88, 79])` dovrebbe restituire `"B+"`.

```js
assert.equal(getAverageGrade([84, 89, 85, 100, 91, 88, 79]), "B+");
```

`getAverageGrade([63, 69, 65, 66, 71, 64, 65])` dovrebbe restituire `"D"`.

```js
assert.equal(getAverageGrade([63, 69, 65, 66, 71, 64, 65]), "D");
```

`getAverageGrade([97, 98, 99, 100, 96, 97, 98, 99, 100])` dovrebbe restituire `"A+"`.

```js
assert.equal(getAverageGrade([97, 98, 99, 100, 96, 97, 98, 99, 100]), "A+");
```

`getAverageGrade([75, 100, 88, 79, 80, 78, 64, 60])` dovrebbe restituire `"C+"`.

```js
assert.equal(getAverageGrade([75, 100, 88, 79, 80, 78, 64, 60]), "C+");
```

`getAverageGrade([45, 48, 50, 52, 100, 54, 56, 58, 59])` dovrebbe restituire `"F"`.

```js
assert.equal(getAverageGrade([45, 48, 50, 52, 100, 54, 56, 58, 59]), "F");
```

# --seed--

## --seed-contents--

```js
function getAverageGrade(scores) {

  return scores;
}
```

# --solutions--

```js
function getAverageGrade(scores) {
  const avg =
    scores.reduce((sum, score) => sum + score, 0) / scores.length;

  if (avg >= 97) return "A+";
  if (avg >= 93) return "A";
  if (avg >= 90) return "A-";

  if (avg >= 87) return "B+";
  if (avg >= 83) return "B";
  if (avg >= 80) return "B-";

  if (avg >= 77) return "C+";
  if (avg >= 73) return "C";
  if (avg >= 70) return "C-";

  if (avg >= 67) return "D+";
  if (avg >= 63) return "D";
  if (avg >= 60) return "D-";

  return "F";
}
```
