---
id: 694596b0585c11170ac7c7fc
title: "Sfida 164: parser di codice inline in markdown"
challengeType: 28
dashedName: challenge-164
---

# --description--

Dato un stringa di Markdown che include uno o più blocchi di codice inline, restituisci la stringa HTML equivalente.

I blocchi di codice inline in Markdown usano un singolo backtick (``` ` ```) all'inizio e alla fine del testo del blocco di codice.

Restituisci la stringa data con tutti i blocchi di codice convertiti in tag HTML `code`.

Per esempio, data la stringa ```"Use `let` to declare the variable."```, return `"Usa <code>let</code> per dichiarare la variabile."``.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

```parseInlineCode("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
assert.equal(parseInlineCode("Usa `let` per dichiarare la variabile."), "Usa <code>let</code> per dichiarare la variabile.");
```

```parseInlineCode("Usa `let` o `const` per dichiarare una variabile.")``` should return `"Usa <code>let</code> o <code>const</code> per dichiarare una variabile."`.

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

```parseInlineCode("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
assert.equal(parseInlineCode("Esegui `npm install` poi `npm start`."), "Esegui <code>npm install</code> poi <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
