---
id: 694596b0585c11170ac7c7fb
title: "Sfida 163: caso consonante"
challengeType: 28
dashedName: challenge-163
---

# --description--

Data una stringa che rappresenta un nome di variabile, convertila in caso consonante seguendo queste regole:

- Tutte le consonanti devono essere convertite in maiuscolo.
- Tutte le vocali (`a`, `e`, `i`, `o`, `u` in qualsiasi caso) devono essere convertite in minuscolo.
- Tutti i trattini (`-`) devono essere convertiti in underscore (`_`).

# --hints--

`toConsonantCase("helloworld")` dovrebbe restituire `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` dovrebbe restituire `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")` dovrebbe restituire `"_HeLLo_WoRLD_"`.

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` dovrebbe restituire `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
