---
id: 694596b0585c11170ac7c7f9
title: "Sfida 161: spedizione gratuita"
challengeType: 28
dashedName: challenge-161
---

# --description--

Dato un array di stringhe che rappresentano gli elementi nel tuo carrello e un numero per l'importo minimo dell'ordine per qualificarsi alla spedizione gratuita, determina se gli elementi nel tuo carrello soddisfano i requisiti per la spedizione gratuita.

L'array fornito conterrà elementi dalla lista seguente:

| Elemento | Prezzo |
| - | - |
| `"shirt"`	| 34.25 |
| `"jeans"`	| 48.50 |
| `"shoes"` |	75.00 |
| `"hat"`	| 19.95 |
| `"socks"` | 15.00 |
| `"jacket"` | 109.95 |

# --hints--

`getsFreeShipping(["shoes"], 50)` dovrebbe restituire `true`.

```js
assert.isTrue(getsFreeShipping(["shoes"], 50));
```

`getsFreeShipping(["hat", "socks"], 50)` dovrebbe restituire `false`.

```js
assert.isFalse(getsFreeShipping(["hat", "socks"], 50));
```

`getsFreeShipping(["jeans", "shirt", "jacket"], 75)` dovrebbe restituire `true`.

```js
assert.isTrue(getsFreeShipping(["jeans", "shirt", "jacket"], 75));
```

`getsFreeShipping(["socks", "socks", "hat"], 75)` dovrebbe restituire `false`.

```js
assert.isFalse(getsFreeShipping(["socks", "socks", "hat"], 75));
```

`getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100)` dovrebbe restituire `true`.

```js
assert.isTrue(getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100));
```

`getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200)` dovrebbe restituire `false`.

```js
assert.isFalse(getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200));
```

# --seed--

## --seed-contents--

```js
function getsFreeShipping(cart, minimum) {

  return cart;
}
```

# --solutions--

```js
function getsFreeShipping(cart, minimum) {
  const prices = {
    shirt: 34.25,
    jeans: 48.50,
    shoes: 75.00,
    hat: 19.95,
    socks: 15.00,
    jacket: 109.95
  };

  let total = 0;

  for (const item of cart) {
    total += prices[item];
  }

  return total >= minimum;
}
```
