---
id: 696655d24b614176d4c9b78c
title: "Sfida 169: mini fizzbuzz"
challengeType: 28
dashedName: challenge-169
---

# --description--

Dato un numero intero, restituisci una stringa basata sulle seguenti regole:

- Se il numero è divisibile per 3, restituisci `"Fizz"`.
- Se il numero è divisibile per 5, restituisci `"Buzz"`.
- Se il numero è divisibile sia per 3 che per 5, restituisci `"FizzBuzz"`.
- Altrimenti, restituisci il numero dato come stringa.

# --hints--

`fizzBuzzMini(3)` dovrebbe restituire `"Fizz"`.

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` dovrebbe restituire `"4"`.

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` dovrebbe restituire `"Buzz"`.

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` dovrebbe restituire `"FizzBuzz"`.

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` dovrebbe restituire `"98"`.

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
