---
id: 696655d24b614176d4c9b789
title: "Sfida 166: validatore esadecimale"
challengeType: 28
dashedName: challenge-166
---

# --description--

Data una stringa, determina se è un colore esadecimale CSS valido. Un colore esadecimale CSS valido deve:

- Iniziare con un `#`, e
- essere seguito da 3 o 6 caratteri esadecimali.

I caratteri esadecimali sono i numeri `0` fino a `9` e le lettere `a` fino a `f` (non fa differenza tra maiuscole e minuscole).

# --hints--

`isValidHex("#123")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
