---
id: 696655d24b614176d4c9b78d
title: "Sfida 170: giorno dispari o pari"
challengeType: 28
dashedName: challenge-170
---

# --description--

Dato un timestamp (numero di millisecondi dall'epoca Unix), restituisci:

- `"odd"` se il giorno del mese per quel timestamp è dispari.
- `"even"` se il giorno del mese per quel timestamp è pari.

Per esempio, dato `1769472000000`, un timestamp per il 27 gennaio 2026, restituisci `"odd"` perché il numero del giorno (`27`) è un numero dispari.

Nota: il timestamp è in millisecondi e devi usare la data nel fuso orario UTC, non nel tuo orario locale.

# --hints--

`oddOrEvenDay(1769472000000)` dovrebbe restituire `"odd"`.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` dovrebbe restituire `"even"`.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` dovrebbe restituire `"odd"`.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` dovrebbe restituire `"odd"`.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` dovrebbe restituire `"even"`.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
