---
id: 696655d24b614176d4c9b78b
title: "Sfida 168: immagine scalata"
challengeType: 28
dashedName: challenge-168
---

# --description--

Data una stringa che rappresenta la larghezza e l'altezza di un'immagine, e un numero per scalare l'immagine, restituisci la larghezza e l'altezza scalate.

- La stringa di input è nel formato `"WxH"`. Per esempio, `"800x600"`.
- La scala è un numero per moltiplicare larghezza e altezza.

Restituisci le dimensioni scalate nello stesso formato `"WxH"`.

# --hints--

`scaleImage("800x600", 2)` dovrebbe restituire `"1600x1200"`.

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)` dovrebbe restituire `"1000x1000"`.

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)` dovrebbe restituire `"512x384"`.

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)` dovrebbe restituire `"450x300"`.

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
