---
id: 69738771fb5a7b8b24cca2a2
title: "Sfida 176: il giorno della marmotta"
challengeType: 28
dashedName: challenge-176
---

# --description--

Oggi è il giorno della marmotta, in cui una marmotta predice il tempo in base al fatto che veda o meno la sua ombra.

Dato un valore che rappresenta l'apparizione della marmotta, restituisci la previsione corretta:

- Se il valore dato è il booleano `true` (la marmotta ha visto la sua ombra), restituisci `"Looks like we'll have six more weeks of winter."`.
- Se il valore è il booleano `false` (la marmotta non ha visto la sua ombra), restituisci `"It's going to be an early spring."`.
- Se il valore è qualsiasi altro (la marmotta non si è presentata), restituisci `"No prediction this year."`.

# --hints--

`groundhogDayPrediction(true)` dovrebbe restituire `"Looks like we'll have six more weeks of winter."`.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)` dovrebbe restituire `"It's going to be an early spring."`.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)` dovrebbe restituire `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` dovrebbe restituire `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` dovrebbe restituire `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
