---
id: 69738771fb5a7b8b24cca2a3
title: "Sfida 177: specchio di stringa"
challengeType: 28
dashedName: challenge-177
---

# --description--

Data una stringa, restituisci una nuova stringa che consiste nella stringa data con una copia invertita di essa aggiunta in coda.

# --hints--

`mirror("freeCodeCamp")` dovrebbe restituire `"freeCodeCamppmaCedoCeerf"`.

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` dovrebbe restituire `"RaceCarraCecaR"`.

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` dovrebbe restituire `"helloworlddlrowolleh"`.

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` dovrebbe restituire `"The quick brown fox......xof nworb kciuq ehT"`.

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
