---
id: 69738771fb5a7b8b24cca29f
title: "Sfida 173: mosse valide del pedone"
challengeType: 28
dashedName: challenge-173
---

# --description--

Data la posizione di uno dei tuoi pedoni su una scacchiera, restituisci un array con tutte le caselle valide in cui può muoversi in ordine crescente.

Una scacchiera standard è 8x8, con le colonne etichettate `A` fino a `H` (da sinistra a destra) e le righe etichettate `1` fino a `8` (dal basso verso l'alto). Ha questo aspetto:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Per questa sfida:

- Sei il giocatore in basso sulla scacchiera.
- I pedoni possono muoversi solo di una casella "verso l'alto".
- A meno che il pedone non sia nella riga di partenza (riga 2), in quel caso può muoversi di una o due caselle verso l'alto.

Ad esempio, dato `"D4"`, restituisci `["D5"]`, l'unica casella in cui il pedone può muoversi.
Dato `"B2"`, restituisci `["B3", "B4"]`, perché si trova nella riga di partenza e deve essere in ordine crescente.

# --hints--

`findPawnMoves("D4")` dovrebbe restituire `["D5"]`.

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")` dovrebbe restituire `["B3", "B4"]`.

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")` dovrebbe restituire `["A8"]`.

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")` dovrebbe restituire `["G3", "G4"]`.

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")` dovrebbe restituire `["E4"]`.

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
