---
id: 69738771fb5a7b8b24cca29e
title: "Sfida 172: lettere e numeri"
challengeType: 28
dashedName: challenge-172
---

# --description--

Data una stringa contenente solo lettere e numeri, restituisci una nuova stringa in cui viene inserito un trattino (`-`) ogni volta che la stringa passa da una lettera a un numero o da un numero a una lettera.

# --hints--

`separateLettersAndNumbers("ABC123")` dovrebbe restituire `"ABC-123"`.

```js
assert.equal(separateLettersAndNumbers("ABC123"), "ABC-123");
```

`separateLettersAndNumbers("Route66")` dovrebbe restituire `"Route-66"`.

```js
assert.equal(separateLettersAndNumbers("Route66"), "Route-66");
```

`separateLettersAndNumbers("H3LL0W0RLD")` dovrebbe restituire `"H-3-LL-0-W-0-RLD"`.

```js
assert.equal(separateLettersAndNumbers("H3LL0W0RLD"), "H-3-LL-0-W-0-RLD");
```

`separateLettersAndNumbers("a1b2c3d4")` dovrebbe restituire `"a-1-b-2-c-3-d-4"`.

```js
assert.equal(separateLettersAndNumbers("a1b2c3d4"), "a-1-b-2-c-3-d-4");
```

# --seed--

## --seed-contents--

```js
function separateLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function separateLettersAndNumbers(str) {
  if (str.length === 0) return str;

  let result = str[0];

  for (let i = 1; i < str.length; i++) {
    const prev = str[i - 1];
    const curr = str[i];

    const prevIsLetter = /[a-zA-Z]/.test(prev);
    const currIsLetter = /[a-zA-Z]/.test(curr);

    if (prevIsLetter !== currIsLetter) {
      result += "-";
    }

    result += curr;
  }

  return result;
}
```
