---
id: 69738771fb5a7b8b24cca2a5
title: "Sfida 179: spiccioli in tasca"
challengeType: 28
dashedName: challenge-179
---

# --description--

Dato un array di numeri interi che rappresentano le monete nella tua tasca, con ogni numero intero che indica il valore di una moneta in centesimi, restituisci l'importo totale nel formato `"$D.CC"`.

- 100 centesimi equivalgono a 1 dollaro.
- Nel valore di ritorno, includi uno zero iniziale per importi inferiori a un dollaro e sempre esattamente due cifre per i centesimi.

# --hints--

`countChange([25, 10, 5, 1])` dovrebbe restituire `"$0.41"`.

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` dovrebbe restituire `"$1.43"`.

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` dovrebbe restituire `"$37.55"`.

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` dovrebbe restituire `"$0.70"`.

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` dovrebbe restituire `"$0.01"`.

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` dovrebbe restituire `"$1.00"`.

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
