---
id: 69738771fb5a7b8b24cca2a4
title: "Sfida 178: troncare il testo"
challengeType: 28
dashedName: challenge-178
---

# --description--

Data una stringa, restituiscila così com'è se è lunga 20 caratteri o meno. Se è più lunga di 20 caratteri, tronca la stringa ai primi 17 caratteri e aggiungi `"..."` alla fine (così la lunghezza totale sarà di 20 caratteri) e restituisci il risultato.

# --hints--

`truncateText("Hello, world!")` dovrebbe restituire `"Hello, world!"`.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` dovrebbe restituire `"This string shoul..."`.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` dovrebbe restituire `"Exactly twenty chars"`.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` dovrebbe restituire `"...................."`.

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
