---
id: 69738771fb5a7b8b24cca2a0
title: "Sfida 174: trovare il segno zodiacale"
challengeType: 28
dashedName: challenge-174
---

# --description--

Dato un valore di data nel formato `"YYYY-MM-DD"`, restituisci il segno zodiacale corrispondente a quella data usando il seguente grafico:

| Intervallo di date | Segno zodiacale |
| - | - |
| 21 marzo - 19 aprile | `"Aries"` |
| 20 aprile - 20 maggio | `"Taurus"` |
| 21 maggio - 20 giugno | `"Gemini"` |
| 21 giugno - 22 luglio | `"Cancer"` |
| 23 luglio - 22 agosto | `"Leo"` |
| 23 agosto - 22 settembre | `"Virgo"` |
| 23 settembre - 22 ottobre | `"Libra"` |
| 23 ottobre - 21 novembre | `"Scorpio"` |
| 22 novembre - 21 dicembre | `"Sagittarius"` |
| 22 dicembre - 19 gennaio | `"Capricorn"` |
| 20 gennaio - 18 febbraio | `"Aquarius"` |
| 19 febbraio - 20 marzo | `"Pisces"` |

- I segni zodiacali si basano solo sul mese e sul giorno, puoi ignorare l'anno.

# --hints--

`getSign("2026-01-31")` dovrebbe restituire `"Aquarius"`.

```js
assert.equal(getSign("2026-01-31"), "Aquarius");
```

`getSign("2001-06-10")` dovrebbe restituire `"Gemini"`.

```js
assert.equal(getSign("2001-06-10"), "Gemini");
```

`getSign("1985-09-07")` dovrebbe restituire `"Virgo"`.

```js
assert.equal(getSign("1985-09-07"), "Virgo");
```

`getSign("2023-03-19")` dovrebbe restituire `"Pisces"`.

```js
assert.equal(getSign("2023-03-19"), "Pisces");
```

`getSign("2045-11-05")` dovrebbe restituire `"Scorpio"`.

```js
assert.equal(getSign("2045-11-05"), "Scorpio");
```

`getSign("1985-12-06")` dovrebbe restituire `"Sagittarius"`.

```js
assert.equal(getSign("1985-12-06"), "Sagittarius");
```

`getSign("2025-12-30")` dovrebbe restituire `"Capricorn"`.

```js
assert.equal(getSign("2025-12-30"), "Capricorn");
```

`getSign("2018-10-08")` dovrebbe restituire `"Libra"`.

```js
assert.equal(getSign("2018-10-08"), "Libra");
```

`getSign("1958-05-04")` dovrebbe restituire `"Taurus"`.

```js
assert.equal(getSign("1958-05-04"), "Taurus");
```

# --seed--

## --seed-contents--

```js
function getSign(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getSign(dateStr) {
  const [, monthStr, dayStr] = dateStr.split("-");
  const month = Number(monthStr);
  const day = Number(dayStr);

  if ((month === 3 && day >= 21) || (month === 4 && day <= 19)) return "Aries";
  if ((month === 4 && day >= 20) || (month === 5 && day <= 20)) return "Taurus";
  if ((month === 5 && day >= 21) || (month === 6 && day <= 20)) return "Gemini";
  if ((month === 6 && day >= 21) || (month === 7 && day <= 22)) return "Cancer";
  if ((month === 7 && day >= 23) || (month === 8 && day <= 22)) return "Leo";
  if ((month === 8 && day >= 23) || (month === 9 && day <= 22)) return "Virgo";
  if ((month === 9 && day >= 23) || (month === 10 && day <= 22)) return "Libra";
  if ((month === 10 && day >= 23) || (month === 11 && day <= 21)) return "Scorpio";
  if ((month === 11 && day >= 22) || (month === 12 && day <= 21)) return "Sagittarius";
  if ((month === 12 && day >= 22) || (month === 1 && day <= 19)) return "Capricorn";
  if ((month === 1 && day >= 20) || (month === 2 && day <= 18)) return "Aquarius";
  if ((month === 2 && day >= 19) || (month === 3 && day <= 20)) return "Pisces";
}
```
