---
id: 697a49e6ff50d756c9b69365
title: "Sfida 188: Giochi invernali 2026 giorno 9: skeleton"
challengeType: 28
dashedName: challenge-188
---

# --description--

Data una stringa che rappresenta le curve di una pista di skeleton, determina la difficoltà della pista.

- La stringa fornita conterrà solo le lettere:
  - `"L"` per una curva a sinistra
  - `"R"` per una curva a destra
  - `"S"` per un tratto dritto

- Ogni cambio di direzione aggiunge 15 punti (un `"L"` seguito da un `"R"` o viceversa).
- Tutte le altre curve aggiungono 5 punti ciascuna (tutti gli altri caratteri `"L"` o `"R"`).
- I tratti dritti aggiungono 0 punti.

La difficoltà della pista si basa sul punteggio totale. Restituisci:

- `"Easy"` se il totale è da 0 a 100
- `"Medium"` se il totale è da 101 a 200
- `"Hard"` se il totale è superiore a 200

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` dovrebbe restituire `"Easy"`.

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` dovrebbe restituire `"Hard"`.

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` dovrebbe restituire `"Medium"`.

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` dovrebbe restituire `"Hard"`.

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` dovrebbe restituire `"Medium"`.

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` dovrebbe restituire `"Easy"`.

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
