---
id: 697a49e6ff50d756c9b69366
title: "Sfida 189: giochi invernali 2026 giorno 10: sci alpino"
challengeType: 28
dashedName: challenge-189
---

# --description--

Dato il dislivello verticale, la distanza orizzontale e il tipo di una pista da sci, determina la valutazione della difficoltà della pista.

Per determinare la valutazione:

- Calcola la pendenza della pista dividendo il dislivello per la distanza.
- Poi, calcola la pendenza aggiustata in base al tipo di pista:
  - `"Downhill"`: moltiplica la pendenza per 1.2
  - `"Slalom"`: moltiplica la pendenza per 0.9
  - `"Giant Slalom"`: moltiplica la pendenza per 1.0

Restituisci:

- `"Green"` se la pendenza aggiustata è minore o uguale a 0.1
- `"Blue"` se la pendenza aggiustata è maggiore di 0.1 e minore o uguale a 0.25
- `"Black"` se la pendenza aggiustata è maggiore di 0.25

# --hints--

`getHillRating(95, 900, "Slalom")` dovrebbe restituire `"Green"`.

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")` dovrebbe restituire `"Black"`.

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")` dovrebbe restituire `"Blue"`.

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")` dovrebbe restituire `"Green"`.

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")` dovrebbe restituire `"Blue"`.

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")` dovrebbe restituire `"Black"`.

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```
