---
id: 697a49e6ff50d756c9b6935e
title: "Sfida 181: giochi invernali 2026 giorno 2: snowboard"
challengeType: 28
dashedName: challenge-181
---

# --description--

Dato lo stance iniziale di uno snowboarder e una rotazione in gradi, determina lo stance di atterraggio.

- Lo stance di uno snowboarder è o `"Regular"` o `"Goofy"`.
- Le rotazioni dei trick sono multipli di 90 gradi. Un valore positivo indica una rotazione in senso orario, un valore negativo in senso antiorario.
- Lo stance di atterraggio si inverte ogni 180 gradi di rotazione.

Per esempio, dato `"Regular"` e `90`, restituisci `"Regular"`. Dato `"Regular"` e `180` gradi, restituisci `"Goofy"`.

# --hints--

`getLandingStance("Regular", 90)` dovrebbe restituire `"Regular"`.

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` dovrebbe restituire `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` dovrebbe restituire `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` dovrebbe restituire `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` dovrebbe restituire `"Goofy"`.

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` dovrebbe restituire `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
