---
id: 697a49e9860d24853adef67c
title: "Sfida 191: giochi invernali 2026 giorno 12: bob"
challengeType: 28
dashedName: challenge-191
---

# --description--

Dato un array che rappresenta i pesi degli atleti di una squadra di bob e un numero che rappresenta il peso del bob, determina se la squadra è idonea a gareggiare.

- La lunghezza dell'array determina la dimensione della squadra: squadre da 1, 2 o 4 persone.
- Tutti i valori di peso forniti sono in chilogrammi (kg).

Il bob (slitta da solo) deve avere un peso minimo di:

- 162 kg per una squadra da 1 persona
- 170 kg per una squadra da 2 persone
- 210 kg per una squadra da 4 persone

Il peso totale del bob (atleti più slitta) non deve superare:

- 247 kg per una squadra da 1 persona
- 390 kg per una squadra da 2 persone
- 630 kg per una squadra da 4 persone

Restituisci `"Eligible"` se la squadra soddisfa tutti i requisiti, oppure `"Not Eligible"` se la squadra non soddisfa uno o più requisiti.

# --hints--

`checkEligibility([78], 165)` dovrebbe restituire `"Eligible"`.

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` dovrebbe restituire `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` dovrebbe restituire `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` dovrebbe restituire `"Eligible"`.

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` dovrebbe restituire `"Not Eligible"`.

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` dovrebbe restituire `"Not Eligible"`.

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` dovrebbe restituire `"Eligible"`.

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` dovrebbe restituire `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` dovrebbe restituire `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
