---
id: 697a49e6ff50d756c9b69364
title: "Sfida 187: giochi invernali 2026 giorno 8: slittino"
challengeType: 28
dashedName: challenge-187
---

# --description--

Dato un array di cinque numeri, ognuno rappresentante il tempo (in secondi) impiegato da uno slittinista per completare un segmento di una pista, determina quale segmento ha avuto la velocità più alta e qual è stata quella velocità.

La pista è divisa nei seguenti segmenti:

- Segmento 1: 320 metri
- Segmento 2: 280 metri
- Segmento 3: 350 metri
- Segmento 4: 300 metri
- Segmento 5: 250 metri

Il primo valore nell'array dato corrisponde al tempo per il segmento 1, il secondo valore al segmento 2, e così via.

Per calcolare la velocità (in metri al secondo) di un segmento, dividi la distanza per il tempo.

Restituisci `"The luger's fastest speed was X m/s on segment Y."`, dove `X` è la velocità più alta, arrotondata a due decimali, e `Y` è il numero del segmento in cui si è verificata la velocità più alta.

# --hints--

`getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128])` dovrebbe restituire `"The luger's fastest speed was 35.07 m/s on segment 5."`

```js
assert.equal(getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.");
```

`getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284])` dovrebbe restituire `"The luger's fastest speed was 37.75 m/s on segment 2."`

```js
assert.equal(getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.");
```

`getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912])` dovrebbe restituire `"The luger's fastest speed was 38.49 m/s on segment 3."`

```js
assert.equal(getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.");
```

`getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840])` dovrebbe restituire `"The luger's fastest speed was 37.69 m/s on segment 1."`

```js
assert.equal(getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.");
```

`getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508])` dovrebbe restituire `"The luger's fastest speed was 39.24 m/s on segment 4."`

```js
assert.equal(getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.");
```

# --seed--

## --seed-contents--

```js
function getFastestSpeed(times) {

  return times;
}
```

# --solutions--

```js
function getFastestSpeed(times) {
  const distances = [320, 280, 350, 300, 250];
  let fastestSpeed = 0;
  let fastestSegment = 0;

  for (let i = 0; i < times.length; i++) {
    const speed = distances[i] / times[i];
    if (speed > fastestSpeed) {
      fastestSpeed = speed;
      fastestSegment = i + 1;
    }
  }

  return `The luger's fastest speed was ${fastestSpeed.toFixed(2)} m/s on segment ${fastestSegment}.`;
}
```
