---
id: 697a49e6ff50d756c9b69360
title: "Sfida 183: giochi invernali 2026 giorno 4: salto con gli sci"
challengeType: 28
dashedName: challenge-183
---

# --description--

Dati i punti distanza, i punti stile, un valore di compensazione del vento e un bonus per il punto K, calcola il tuo punteggio per il salto con gli sci e determina se hai vinto una medaglia o no.

- Il tuo punteggio si calcola sommando i quattro valori sopra indicati.

I punteggi totali attuali degli altri saltatori sono:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- Se il tuo punteggio è il migliore, restituisci `"Gold"`
- Se è il secondo migliore, restituisci `"Silver"`
- Se è il terzo migliore, restituisci `"Bronze"`
- Altrimenti, restituisci `"No Medal"`

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)` dovrebbe restituire `"Gold"`.

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)` dovrebbe restituire `"Bronze"`.

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)` dovrebbe restituire `"Silver"`.

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)` dovrebbe restituire `"No Medal"`.

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)` dovrebbe restituire `"Gold"`.

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)` dovrebbe restituire `"No Medal"`.

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```
