---
id: 697a49e6ff50d756c9b6935f
title: "Sfida 182: giochi invernali 2026 giorno 3: biathlon"
challengeType: 28
dashedName: challenge-182
---

# --description--

Dato un array di numeri interi, dove ogni valore rappresenta il numero di obiettivi colpiti in un singolo round di biathlon, restituisci la distanza totale della penalità che l'atleta deve sciare.

- Ogni round consiste in 5 obiettivi.
- Ogni obiettivo mancato comporta un loop di penalità di 150 metri.

# --hints--

`calculatePenaltyDistance([4, 4])` dovrebbe restituire `300`.

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` dovrebbe restituire `0`.

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` dovrebbe restituire `450`.

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` dovrebbe restituire `150`.

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` dovrebbe restituire `1500`.

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
