---
id: 697a49e6ff50d756c9b6935d
title: "Sfida 180: giochi invernali 2026 giorno 1: giorno di apertura"
challengeType: 28
dashedName: challenge-180
---

# --description--

Oggi inizia la 2026 Winter Games. I prossimi 17 giorni ti proporranno sfide di programmazione ispirate a questi giochi.

Per la prima, ti viene dato un codice paese di due lettere e devi restituire l'emoji della bandiera di quel paese.

Usa questa lista:

| Paese                     | Codice     | Bandiera |
| -                         | -:         | :--:     |
| Albania                   |            `"AL"` | `"🇦🇱"` |
| Andorra                   |            `"AD"` | `"🇦🇩"` |
| Argentina                 |            `"AR"` | `"🇦🇷"` |
| Armenia                   |            `"AM"` | `"🇦🇲"` |
| Australia                 |            `"AU"` | `"🇦🇺"` |
| Austria                   |            `"AT"` | `"🇦🇹"` |
| Azerbaijan                |            `"AZ"` | `"🇦🇿"` |
| Belgio                    |            `"BE"` | `"🇧🇪"` |
| Benin                     |            `"BJ"` | `"🇧🇯"` |
| Bolivia                   |            `"BO"` | `"🇧🇴"` |
| Bosnia ed Erzegovina      |            `"BA"` | `"🇧🇦"` |
| Brasile                   |            `"BR"` | `"🇧🇷"` |
| Bulgaria                  |            `"BG"` | `"🇧🇬"` |
| Canada                    |            `"CA"` | `"🇨🇦"` |
| Cile                      |            `"CL"` | `"🇨🇱"` |
| Cina                      |            `"CN"` | `"🇨🇳"` |
| Colombia                  |            `"CO"` | `"🇨🇴"` |
| Croazia                   |            `"HR"` | `"🇭🇷"` |
| Cipro                     |            `"CY"` | `"🇨🇾"` |
| Repubblica Ceca           |            `"CZ"` | `"🇨🇿"` |
| Danimarca                 |            `"DK"` | `"🇩🇰"` |
| Ecuador                   |            `"EC"` | `"🇪🇨"` |
| Eritrea                   |            `"ER"` | `"🇪🇷"` |
| Estonia                   |            `"EE"` | `"🇪🇪"` |
| Finlandia                 |            `"FI"` | `"🇫🇮"` |
| Francia                   |            `"FR"` | `"🇫🇷"` |
| Georgia                   |            `"GE"` | `"🇬🇪"` |
| Germania                  |            `"DE"` | `"🇩🇪"` |
| Gran Bretagna             |            `"GB"` | `"🇬🇧"` |
| Grecia                    |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau             |            `"GW"` | `"🇬🇼"` |
| Haiti                     |            `"HT"` | `"🇭🇹"` |
| Hong Kong                 |            `"HK"` | `"🇭🇰"` |
| Ungheria                  |            `"HU"` | `"🇭🇺"` |
| Islanda                   |            `"IS"` | `"🇮🇸"` |
| India                     |            `"IN"` | `"🇮🇳"` |
| Iran                      |            `"IR"` | `"🇮🇷"` |
| Irlanda                   |            `"IE"` | `"🇮🇪"` |
| Israele                   |            `"IL"` | `"🇮🇱"` |
| Italia                    |            `"IT"` | `"🇮🇹"` |
| Giamaica                  |            `"JM"` | `"🇯🇲"` |
| Giappone                  |            `"JP"` | `"🇯🇵"` |
| Kazakistan                |            `"KZ"` | `"🇰🇿"` |
| Kenya                     |            `"KE"` | `"🇰🇪"` |
| Kosovo                    |            `"XK"` | `"🇽🇰"` |
| Kirghizistan              |            `"KG"` | `"🇰🇬"` |
| Lettonia                  |            `"LV"` | `"🇱🇻"` |
| Libano                    |            `"LB"` | `"🇱🇧"` |
| Liechtenstein             |            `"LI"` | `"🇱🇮"` |
| Lituania                  |            `"LT"` | `"🇱🇹"` |
| Lussemburgo               |            `"LU"` | `"🇱🇺"` |
| Madagascar                |            `"MG"` | `"🇲🇬"` |
| Malesia                   |            `"MY"` | `"🇲🇾"` |
| Malta                     |            `"MT"` | `"🇲🇹"` |
| Messico                   |            `"MX"` | `"🇲🇽"` |
| Moldavia                  |            `"MD"` | `"🇲🇩"` |
| Monaco                    |            `"MC"` | `"🇲🇨"` |
| Mongolia                  |            `"MN"` | `"🇲🇳"` |
| Montenegro                |            `"ME"` | `"🇲🇪"` |
| Marocco                   |            `"MA"` | `"🇲🇦"` |
| Paesi Bassi               |            `"NL"` | `"🇳🇱"` |
| Nuova Zelanda             |            `"NZ"` | `"🇳🇿"` |
| Nigeria                   |            `"NG"` | `"🇳🇬"` |
| Macedonia del Nord        |            `"MK"` | `"🇲🇰"` |
| Norvegia                  |            `"NO"` | `"🇳🇴"` |
| Pakistan                  |            `"PK"` | `"🇵🇰"` |
| Filippine                 |            `"PH"` | `"🇵🇭"` |
| Polonia                   |            `"PL"` | `"🇵🇱"` |
| Portogallo                |            `"PT"` | `"🇵🇹"` |
| Porto Rico                |            `"PR"` | `"🇵🇷"` |
| Romania                   |            `"RO"` | `"🇷🇴"` |
| San Marino                |            `"SM"` | `"🇸🇲"` |
| Arabia Saudita            |            `"SA"` | `"🇸🇦"` |
| Serbia                    |            `"RS"` | `"🇷🇸"` |
| Singapore                 |            `"SG"` | `"🇸🇬"` |
| Slovacchia                |            `"SK"` | `"🇸🇰"` |
| Slovenia                  |            `"SI"` | `"🇸🇮"` |
| Sudafrica                 |            `"ZA"` | `"🇿🇦"` |
| Corea del Sud             |            `"KR"` | `"🇰🇷"` |
| Spagna                    |            `"ES"` | `"🇪🇸"` |
| Svezia                    |            `"SE"` | `"🇸🇪"` |
| Svizzera                  |            `"CH"` | `"🇨🇭"` |
| Thailandia                |            `"TH"` | `"🇹🇭"` |
| Trinidad e Tobago         |            `"TT"` | `"🇹🇹"` |
| Turchia                   |            `"TR"` | `"🇹🇷"` |
| Ucraina                   |            `"UA"` | `"🇺🇦"` |
| Emirati Arabi Uniti       |            `"AE"` | `"🇦🇪"` |
| Stati Uniti               |            `"US"` | `"🇺🇸"` |
| Uruguay                   |            `"UY"` | `"🇺🇾"` |
| Uzbekistan                |            `"UZ"` | `"🇺🇿"` |
| Venezuela                 |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` dovrebbe restituire `"🇦🇱"`.

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` dovrebbe restituire `"🇦🇩"`.

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` dovrebbe restituire `"🇦🇷"`.

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` dovrebbe restituire `"🇦🇲"`.

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` dovrebbe restituire `"🇦🇺"`.

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` dovrebbe restituire `"🇦🇹"`.

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` dovrebbe restituire `"🇦🇿"`.

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` dovrebbe restituire `"🇧🇪"`.

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` dovrebbe restituire `"🇧🇯"`.

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` dovrebbe restituire `"🇧🇴"`.

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` dovrebbe restituire `"🇧🇦"`.

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` dovrebbe restituire `"🇧🇷"`.

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` dovrebbe restituire `"🇧🇬"`.

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` dovrebbe restituire `"🇨🇦"`.

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` dovrebbe restituire `"🇨🇱"`.

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` dovrebbe restituire `"🇨🇳"`.

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` dovrebbe restituire `"🇨🇴"`.

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` dovrebbe restituire `"🇭🇷"`.

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` dovrebbe restituire `"🇨🇾"`.

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` dovrebbe restituire `"🇨🇿"`.

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` dovrebbe restituire `"🇩🇰"`.

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` dovrebbe restituire `"🇪🇨"`.

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` dovrebbe restituire `"🇪🇷"`.

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` dovrebbe restituire `"🇪🇪"`.

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` dovrebbe restituire `"🇫🇮"`.

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` dovrebbe restituire `"🇫🇷"`.

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` dovrebbe restituire `"🇬🇪"`.

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` dovrebbe restituire `"🇩🇪"`.

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` dovrebbe restituire `"🇬🇧"`.

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` dovrebbe restituire `"🇬🇷"`.

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` dovrebbe restituire `"🇬🇼"`.

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` dovrebbe restituire `"🇭🇹"`.

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` dovrebbe restituire `"🇭🇰"`.

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` dovrebbe restituire `"🇭🇺"`.

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` dovrebbe restituire `"🇮🇸"`.

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` dovrebbe restituire `"🇮🇳"`.

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` dovrebbe restituire `"🇮🇷"`.

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` dovrebbe restituire `"🇮🇪"`.

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` dovrebbe restituire `"🇮🇱"`.

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` dovrebbe restituire `"🇮🇹"`.

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` dovrebbe restituire `"🇯🇲"`.

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` dovrebbe restituire `"🇯🇵"`.

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` dovrebbe restituire `"🇰🇿"`.

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` dovrebbe restituire `"🇰🇪"`.

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` dovrebbe restituire `"🇽🇰"`.

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` dovrebbe restituire `"🇰🇬"`.

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` dovrebbe restituire `"🇱🇻"`.

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` dovrebbe restituire `"🇱🇧"`.

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` dovrebbe restituire `"🇱🇮"`.

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` dovrebbe restituire `"🇱🇹"`.

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` dovrebbe restituire `"🇱🇺"`.

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` dovrebbe restituire `"🇲🇬"`.

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` dovrebbe restituire `"🇲🇾"`.

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` dovrebbe restituire `"🇲🇹"`.

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` dovrebbe restituire `"🇲🇽"`.

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` dovrebbe restituire `"🇲🇩"`.

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` dovrebbe restituire `"🇲🇨"`.

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` dovrebbe restituire `"🇲🇳"`.

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` dovrebbe restituire `"🇲🇪"`.

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` dovrebbe restituire `"🇲🇦"`.

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` dovrebbe restituire `"🇳🇱"`.

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` dovrebbe restituire `"🇳🇿"`.

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` dovrebbe restituire `"🇳🇬"`.

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` dovrebbe restituire `"🇲🇰"`.

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` dovrebbe restituire `"🇳🇴"`.

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` dovrebbe restituire `"🇵🇰"`.

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` dovrebbe restituire `"🇵🇭"`.

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` dovrebbe restituire `"🇵🇱"`.

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` dovrebbe restituire `"🇵🇹"`.

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` dovrebbe restituire `"🇵🇷"`.

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` dovrebbe restituire `"🇷🇴"`.

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` dovrebbe restituire `"🇸🇲"`.

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` dovrebbe restituire `"🇸🇦"`.

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` dovrebbe restituire `"🇷🇸"`.

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` dovrebbe restituire `"🇸🇬"`.

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` dovrebbe restituire `"🇸🇰"`.

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` dovrebbe restituire `"🇸🇮"`.

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` dovrebbe restituire `"🇿🇦"`.

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` dovrebbe restituire `"🇰🇷"`.

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` dovrebbe restituire `"🇪🇸"`.

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` dovrebbe restituire `"🇸🇪"`.

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` dovrebbe restituire `"🇨🇭"`.

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` dovrebbe restituire `"🇹🇭"`.

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` dovrebbe restituire `"🇹🇹"`.

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` dovrebbe restituire `"🇹🇷"`.

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` dovrebbe restituire `"🇺🇦"`.

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` dovrebbe restituire `"🇦🇪"`.

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` dovrebbe restituire `"🇺🇸"`.

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` dovrebbe restituire `"🇺🇾"`.

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` dovrebbe restituire `"🇺🇿"`.

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` dovrebbe restituire `"🇻🇪"`.

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
