---
id: 698a1a73ade5ac0e19180fa0
title: "Sfida 197: compatibilità dei gruppi sanguigni"
challengeType: 28
dashedName: challenge-197
---

# --description--

Dato un gruppo sanguigno del donatore e un gruppo sanguigno del ricevente, determina se il donatore può donare sangue al ricevente.

Ogni gruppo sanguigno è composto da:

- Una lettera: `"A"`, `"B"`, `"AB"` o `"O"`
- E un fattore Rh: `"+"` o `"-"`

I gruppi sanguigni saranno una delle lettere valide seguite da un fattore Rh. Per esempio, `"AB+"` e `"O-"` sono gruppi sanguigni validi.

Regole per le lettere:

- `"O"` può donare ad altri tipi di lettera.
- `"A"` può donare a `"A"` e `"AB"`.
- `"B"` può donare a `"B"` e `"AB"`.
- `"AB"` può donare solo a `"AB"`.

Regole per il fattore Rh:

- Negativo (`"-"`) può donare sia a `"-"` che a `"+"`.
- Positivo (`"+"`) può donare solo a `"+"`.

Sia la regola della lettera sia quella del fattore Rh devono essere rispettate affinché un donatore possa donare al ricevente.

# --hints--

`canDonate("B+", "B+")` dovrebbe restituire `true`.

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` dovrebbe restituire `true`.

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` dovrebbe restituire `false`.

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` dovrebbe restituire `true`.

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` dovrebbe restituire `false`.

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` dovrebbe restituire `true`.

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` dovrebbe restituire `false`.

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` dovrebbe restituire `true`.

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` dovrebbe restituire `false`.

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` dovrebbe restituire `false`.

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
