---
id: 698a1a863194f1f4e63f645e
title: "Sfida 207: spazio più piccolo"
challengeType: 28
dashedName: challenge-207
---

# --description--

Data una stringa, restituisci la sottostringa compresa tra due caratteri identici che hanno il numero più piccolo di caratteri tra loro (spazio più piccolo).

- Ci sarà sempre almeno una coppia di caratteri corrispondenti.
- La sottostringa restituita deve escludere i caratteri corrispondenti.
- Se due o più spazi hanno la stessa lunghezza, restituisci i caratteri del primo.

Per esempio, data `"ABCDAC"`, restituisci `"DA"` perché:

- Solo `"A"` e `"C"` si ripetono nella stringa.
- Il numero di caratteri tra i due caratteri `"A"` è 3, e tra i caratteri `"C"` è 2.
- Quindi restituisci la stringa tra i due caratteri `"C"`.

# --hints--

`smallestGap("ABCDAC")` dovrebbe restituire `"DA"`.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` dovrebbe restituire `"e"`.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` dovrebbe restituire `"#q6e&rkf(p"`.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` dovrebbe restituire `""`.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` dovrebbe restituire `"fox"`.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
