---
id: 698a1a73ade5ac0e19180fa1
title: "Sfida 198: conteggio dei giorni lavorativi"
challengeType: 28
dashedName: challenge-198
---

# --description--

Dato un giorno di inizio e un giorno di fine, restituisci il numero di giorni lavorativi compresi tra i due.

- Le date fornite sono nel formato `"YYYY-MM-DD"`.
- I giorni feriali sono giorni lavorativi (da lunedì a venerdì).
- I fine settimana non sono giorni lavorativi (sabato e domenica).
- Conta sia il giorno di inizio sia quello di fine.

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` dovrebbe restituire `3`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` dovrebbe restituire `4`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` dovrebbe restituire `5`.

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` dovrebbe restituire `7`.

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` dovrebbe restituire `262`.

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
