---
id: 698a1a73ade5ac0e19180fa5
title: "Sfida 202: aggiungere la punteggiatura"
challengeType: 28
dashedName: challenge-202
---

# --description--

Data una stringa di frasi senza punti, aggiungi un punto (`"."`) nei seguenti casi:

- Prima di ogni spazio che precede immediatamente una lettera maiuscola
- E alla fine della stringa

Restituisci la stringa risultante.

# --hints--

`addPunctuation("Hello world")` should return `"Hello world."`

```js
assert.equal(addPunctuation("Hello world"), "Hello world.");
```

`addPunctuation("Hello world It's nice today")` should return `"Hello world. It's nice today."`

```js
assert.equal(addPunctuation("Hello world It's nice today"), "Hello world. It's nice today.");
```

`addPunctuation("JavaScript is great Sometimes")` should return `"JavaScript is great. Sometimes."`

```js
assert.equal(addPunctuation("JavaScript is great Sometimes"), "JavaScript is great. Sometimes.");
```

`addPunctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z")` should return `"A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z."`

```js
assert.equal(addPunctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z"), "A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z.");
```

`addPunctuation("Wait.. For it")` should return `"Wait... For it."`

```js
assert.equal(addPunctuation("Wait.. For it"), "Wait... For it.");
```

# --seed--

## --seed-contents--

```js
function addPunctuation(sentences) {

  return sentences;
}
```

# --solutions--

```js
function addPunctuation(sentences) {
  let result = "";

  for (let i = 0; i < sentences.length; i++) {
    if (sentences[i] === " " && i + 1 < sentences.length && /[A-Z]/.test(sentences[i + 1])) {
      result += ".";
    }
    result += sentences[i];
  }
  result += ".";

  return result;
}
```
