---
id: 698a1a863194f1f4e63f6460
title: "Sfida 209: dimensione dell'elemento"
challengeType: 28
dashedName: challenge-209
---

# --description--

Dato un valore per la dimensione della finestra, la larghezza di un elemento in unità `"vw"` di viewport width e l'altezza di un elemento in unità `"vh"` di viewport height, determina la dimensione dell'elemento in pixel.

- La dimensione della finestra fornita e la dimensione dell'elemento restituita sono stringhe nel formato `"width x height"`, `"1200 x 800"` per esempio.

- Le unità `"vw"` rappresentano la percentuale della larghezza della finestra. `"50vw"`, per esempio, è il 50% della larghezza della finestra.
- Le unità `"vh"` rappresentano la percentuale dell'altezza della finestra. `"50vh"`, per esempio, è il 50% dell'altezza della finestra.

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` dovrebbe restituire `"600 x 400"`.

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` dovrebbe restituire `"80 x 240"`.

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` dovrebbe restituire `"70 x 15"`.

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` dovrebbe restituire `"1824 x 1080"`.

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` dovrebbe restituire `"0 x 0"`.

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` dovrebbe restituire `"1440 x 1026"`.

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
