---
id: 698a1a73ade5ac0e19180fa6
title: "Sfida 203: array appiattito"
challengeType: 28
dashedName: challenge-203
---

# --description--

Dato un array, determina se è appiattito.

- Un array è appiattito se nessuno dei suoi elementi è un array.

# --hints--

`isFlat([1, 2, 3, 4])` dovrebbe restituire `true`.

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` dovrebbe restituire `false`.

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` dovrebbe restituire `true`.

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` dovrebbe restituire `false`.

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` dovrebbe restituire `false`.

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
