---
id: 698a1a73ade5ac0e19180fa9
title: "Sfida 206: valori delle carte da gioco"
challengeType: 28
dashedName: challenge-206
---

# --description--

Dato un array di carte da gioco, restituisci un nuovo array con il valore numerico di ogni carta.

Valori delle carte:

- Un asso (`"A"`) ha valore 1.
- Le carte numerate (`"2"` - `"10"`) hanno il loro valore nominale: rispettivamente da 2 a 10.
- Le figure: fante (`"J"`), donna (`"Q"`) e re (`"K"`) valgono ciascuna 10.

Semi:

- Ogni carta ha un seme: picche (`"S"`), fiori (`"C"`), quadri (`"D"`) o cuori (`"H"`).

Formato della carta:

- Ogni carta è rappresentata come una stringa: `"valueSuit"`. Per esempio: `"AS"` è l'asso di picche, `"10H"` è il dieci di cuori e `"QC"` è la donna di fiori.

# --hints--

`cardValues(["3H", "4D", "5S"])` dovrebbe restituire `[3, 4, 5]`.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` dovrebbe restituire `[1, 10, 10, 6, 7]`.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` dovrebbe restituire `[8, 10, 2, 10, 9]`.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` dovrebbe restituire `[1, 10]`.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` dovrebbe restituire `[10, 10, 10, 10, 1]`.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
