---
id: 698a1a863194f1f4e63f645f
title: "Sfida 208: attraversamento del sentiero"
challengeType: 28
dashedName: challenge-208
---

# --description--

Dato un array di stringhe che rappresentano la mappa del tuo sentiero, restituisci una stringa con le mosse necessarie per raggiungere il tuo obiettivo.

Le stringhe fornite conterranno i valori:

- `"C"`: la tua posizione attuale
- `"G"`: il tuo obiettivo
- `"T"`: parti del sentiero percorribili
- `"-"`: parti della mappa non percorribili

Restituisci una stringa con le mosse necessarie per seguire il sentiero dalla tua posizione all'obiettivo dove:

- `"R"` è una mossa a destra
- `"D"` è una mossa verso il basso
- `"L"` è una mossa a sinistra
- `"U"` è una mossa verso l'alto

- Ci sarà sempre un sentiero continuo unico, senza diramazioni, dalla tua posizione attuale al tuo obiettivo.
- Ogni posizione del sentiero avrà al massimo due posizioni percorribili adiacenti.

Per esempio, dato:

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

Restituisci `"RDDRDD"`.

# --hints--

`navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` dovrebbe restituire `"RDDRDD"`.

```js
assert.equal(navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD");
```

`navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` dovrebbe restituire `"RRUUURR"`.

```js
assert.equal(navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR");
```

`navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` dovrebbe restituire `"DLDDRRRRD"`.

```js
assert.equal(navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD");
```

`navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` dovrebbe restituire `"RRRDDL"`.

```js
assert.equal(navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL");
```

`navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` dovrebbe restituire `"ULLLUUURRRRRRDDDR"`.

```js
assert.equal(navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR");
```

# --seed--

## --seed-contents--

```js
function navigateTrail(map) {

  return map;
}
```

# --solutions--

```js
function navigateTrail(map) {
  const rows = map.length;
  const cols = map[0].length;

  const directions = [
    [0, 1, "R"],
    [1, 0, "D"],
    [0, -1, "L"],
    [-1, 0, "U"],
  ];

  let row, col;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (map[r][c] === "C") {
        row = r;
        col = c;
      }
    }
  }

  let result = "";
  let prev = null;

  while (map[row][col] !== "G") {
    for (let [dr, dc, move] of directions) {
      const newRow = row + dr;
      const newCol = col + dc;

      if (
        newRow >= 0 &&
        newRow < rows &&
        newCol >= 0 &&
        newCol < cols &&
        (map[newRow][newCol] === "T" ||
          map[newRow][newCol] === "G") &&
        (!prev || newRow !== prev[0] || newCol !== prev[1])
      ) {
        result += move;
        prev = [row, col];
        row = newRow;
        col = newCol;
        break;
      }
    }
  }

  return result;
}
```
