---
id: 698a1a73ade5ac0e19180fa4
title: "Sfida 201: spostamento di una matrice"
challengeType: 28
dashedName: challenge-201
---

# --description--

Data una matrice (array di array) di numeri e un intero, sposta tutti i valori nella matrice della quantità indicata.

- Uno spostamento positivo muove i valori verso destra.
- Uno spostamento negativo muove i valori verso sinistra.

I valori devono avvolgersi:

- Considera la matrice come una sequenza continua di valori.
- Quando un valore supera la fine di una riga, continua all'inizio della riga successiva.
- Quando un valore supera l'ultima posizione nella matrice, si avvolge alla prima posizione.
- Lo stesso vale al contrario quando si sposta a sinistra.

Per esempio, dato:

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

con uno spostamento di `1`, sposta tutti i numeri di uno verso destra:

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shiftMatrix([[1, 2, 3], [4, 5, 6]], 1)` dovrebbe restituire `[[6, 1, 2], [3, 4, 5]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6]], -1)` dovrebbe restituire `[[2, 3, 4], [5, 6, 1]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` dovrebbe restituire `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` dovrebbe restituire `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` dovrebbe restituire `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` dovrebbe restituire `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]);
```

# --seed--

## --seed-contents--

```js
function shiftMatrix(matrix, shift) {

  return matrix;
}
```

# --solutions--

```js
function shiftMatrix(matrix, shift) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const flat = matrix.flat();
  const n = flat.length;
  const s = ((shift % n) + n) % n;
  const shifted = flat.slice(-s).concat(flat.slice(0, n - s));
  const result = [];
  for (let i = 0; i < rows; i++) {
    result.push(shifted.slice(i * cols, (i + 1) * cols));
  }

  return result;
}
```
