---
id: 698a1a73ade5ac0e19180fa8
title: "Sfida 205: conteggio dei cubi perfetti"
challengeType: 28
dashedName: challenge-205
---

# --description--

Dati due numeri interi, determina quanti cubi perfetti esistono nell'intervallo compreso tra i due numeri inclusi.

- Il numero più basso non è garantito essere il primo argomento.
- Un numero è un cubo perfetto se esiste un numero intero (`n`) tale che `n * n * n = number`. Per esempio, 27 è un cubo perfetto perché `3 * 3 * 3 = 27`.

# --hints--

`countPerfectCubes(3, 30)` dovrebbe restituire `2`.

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` dovrebbe restituire `3`.

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` dovrebbe restituire `4`.

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` dovrebbe restituire `9`.

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` dovrebbe restituire `41`.

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
