---
id: 698a1a73ade5ac0e19180fa3
title: "Sfida 200: conteggio di lettere e numeri"
challengeType: 28
dashedName: challenge-200
---

# --description--

Data una stringa, restituisci un messaggio con il conteggio di quante lettere e numeri contiene.

- Le lettere sono `A-Z` e `a-z`.
- I numeri sono `0-9`.
- Ignora tutti gli altri caratteri.

Restituisci `"The string has X letters and Y numbers."`, dove `"X"` è il conteggio delle lettere e `"Y"` è il conteggio dei numeri. Se uno dei due conteggi è 1, usa la forma singolare per quell’elemento. Per esempio: `"1 letter"` invece di `"1 letters"` e `"1 number"` invece di `"1 numbers"`.

# --hints--

`countLettersAndNumbers("helloworld123")` dovrebbe restituire `"The string has 10 letters and 3 numbers."`.

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")` dovrebbe restituire `"The string has 5 letters and 2 numbers."`.

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")` dovrebbe restituire `"The string has 1 letter and 1 number."`.

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")` dovrebbe restituire `"The string has 0 letters and 5 numbers."`.

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")` dovrebbe restituire `"The string has 8 letters and 0 numbers."`.

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
