---
id: 698a1a73ade5ac0e19180fa7
title: "Sfida 204: somma le lettere"
challengeType: 28
dashedName: challenge-204
---

# --description--

Data una stringa, restituisci la somma delle sue lettere.

- Le lettere sono da A a Z in maiuscolo o minuscolo
- I valori delle lettere sono: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26
- Le lettere maiuscole e minuscole hanno lo stesso valore.
- Ignora tutti i caratteri che non sono lettere.

# --hints--

`sumLetters("Hello")` dovrebbe restituire `52`.

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")` dovrebbe restituire `94`.

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")` dovrebbe restituire `473`.

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` dovrebbe restituire `1681`.

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")` dovrebbe restituire `0`.

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```
