---
id: 6994cff2290543b3aec9f50f
title: "Sfida 210: validatore HSL"
challengeType: 28
dashedName: challenge-210
---

# --description--

Data una stringa, determina se è un valore colore CSS `hsl()` valido.

- Un valore HSL valido deve essere nel formato `"hsl(h, s%, l%)"`, dove:
  - `h` (tonalità) deve essere un numero compreso tra 0 e 360 (inclusi).
  - `s` (saturazione) deve essere una percentuale tra 0% e 100%.
  - `l` (luminosità) deve essere una percentuale tra 0% e 100%.

- Gli spazi sono consentiti solo:
  - Dopo la parentesi di apertura
  - Prima e/o dopo le virgole
  - Prima e/o dopo la parentesi di chiusura

- Opzionalmente, il valore può terminare con un punto e virgola (`";"`).

Per esempio, `"hsl(240, 50%, 50%)"` è un valore HSL valido.

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
