---
id: 699c8e045ee7cb94ed2322d6
title: "Sfida 215: calcolatore della tariffa del parcheggio"
challengeType: 28
dashedName: challenge-215
---

# --description--

Date due stringhe che rappresentano l'orario in cui hai parcheggiato la tua auto e l'orario in cui l'hai ritirata, calcola la tariffa del parcheggio.

- Le stringhe date saranno nel formato `"HH:MM"` usando un orologio a 24 ore. Quindi `"14:00"` è per esempio le 14:00.
- La prima stringa sarà l'orario in cui hai parcheggiato l'auto, la seconda sarà l'orario in cui l'hai ritirata.
- Se l'orario di ritiro è precedente a quello di entrata, significa che il parcheggio è durato oltre la mezzanotte nel giorno successivo.

Regole della tariffa:

- Ogni ora di parcheggio costa 3$.
- Le ore parziali si arrotondano per eccesso all'ora intera successiva.
- Se il parcheggio dura tutta la notte (oltre la mezzanotte), si applica una tariffa aggiuntiva di 10$ per la notte.
- C'è una tariffa minima di 5$ (usata solo se il totale sarebbe inferiore a 5$).

Restituisci il costo totale nel formato `"$cost"`, `"$5"` per esempio.

# --hints--

`calculateParkingFee("09:00", "11:00")` dovrebbe restituire `"$6"`.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` dovrebbe restituire `"$5"`.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` dovrebbe restituire `"$9"`.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` dovrebbe restituire `"$27"`.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` dovrebbe restituire `"$34"`.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` dovrebbe restituire `"$82"`.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
