---
id: 699c8e045ee7cb94ed2322d9
title: "Sfida 218: divisibile senza resto"
challengeType: 28
dashedName: challenge-218
---

# --description--

Dato due numeri interi, determina se il primo può essere diviso senza resto per il secondo.

# --hints--

`isEvenlyDivisible(4, 2)` dovrebbe restituire `true`.

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)` dovrebbe restituire `false`.

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)` dovrebbe restituire `false`.

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)` dovrebbe restituire `true`.

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)` dovrebbe restituire `true`.

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)` dovrebbe restituire `false`.

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
