---
id: 699c8e045ee7cb94ed2322d4
title: "Sfida 213: convertitore di lunghezza delle parole"
challengeType: 28
dashedName: challenge-213
---

# --description--

Data una stringa di parole, restituisci una nuova stringa in cui ogni parola è sostituita dalla sua lunghezza.

- Le parole nella stringa data saranno separate da uno spazio singolo
- Mantieni gli spazi nella stringa restituita.

Ad esempio, data `"hello world"`, restituisci `"5 5"`.

# --hints--

`convertWords("hello world")` dovrebbe restituire `"5 5"`.

```js
assert.equal(convertWords("hello world"), "5 5");
```

`convertWords("Thanks and happy coding")` dovrebbe restituire `"6 3 5 6"`.

```js
assert.equal(convertWords("Thanks and happy coding"), "6 3 5 6");
```

`convertWords("The quick brown fox jumps over the lazy dog")` dovrebbe restituire `"3 5 5 3 5 4 3 4 3"`.

```js
assert.equal(convertWords("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3");
```

`convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` dovrebbe restituire `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`.

```js
assert.equal(convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4");
```

# --seed--

## --seed-contents--

```js
function convertWords(str) {

  return str;
}
```

# --solutions--

```js
function convertWords(str) {
  return str
    .split(" ")
    .map(word => word.length)
    .join(" ");
}
```
