---
id: 699c8e045ee7cb94ed2322db
title: "Sfida 220: numero più grande"
challengeType: 28
dashedName: challenge-220
---

# --description--

Data una stringa di numeri separati da vari segni di punteggiatura, restituisci il numero più grande.

- La stringa fornita conterrà solo numeri e separatori.
- I separatori possono essere virgole (`","`), punti esclamativi (`"!"`), punti interrogativi (`"?"`), due punti (`":"`) o punti e virgola (`";"`).

# --hints--

`largestNumber("1,2")` dovrebbe restituire `2`.

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` dovrebbe restituire `60`.

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` dovrebbe restituire `-402`.

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` dovrebbe restituire `99.9`.

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
