---
id: 699c8e045ee7cb94ed2322dd
title: "Sfida 222: ombre all'equinozio"
challengeType: 28
dashedName: challenge-222
---

# --description--

Oggi è l'equinozio, quando il sole è direttamente sopra l'equatore e perfettamente allo zenit a mezzogiorno. Data un'ora, determina l'ombra proiettata da un palo verticale di 4 piedi.

- L'ora sarà una stringa in formato `"HH:MM"` 24 ore (per esempio, `"15:00"` corrisponde alle 15:00).
- Ti verrà fornita solo un'ora con incrementi di 30 minuti.

Regole:

- Il sole sorge alle 6:00 direttamente `"east"` e tramonta alle 18:00 direttamente `"west"`.
- Un'ombra punta sempre nella direzione opposta al sole.
- La lunghezza dell'ombra (in piedi) è il numero di ore di distanza da mezzogiorno, elevato al cubo.
- Non c'è ombra prima dell'alba (prima delle 6:00), dopo il tramonto (18:00 o oltre) o a mezzogiorno.

Restituisci:

- Se esiste un'ombra, restituisci `"(length)ft (direction)"`. Per esempio, `"8ft west"`.
- Altrimenti, restituisci `"No shadow"`.

Per esempio, data `"10:00"`, restituisci `"8ft west"` perché le 10:00 sono 2 ore da mezzogiorno, quindi 2<sup>3</sup> = 8 piedi, e l'ombra punta a ovest perché il sole è a est alle 10:00.

# --hints--

`getShadow("10:00")` dovrebbe restituire `"8ft west"`.

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")` dovrebbe restituire `"27ft east"`.

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")` dovrebbe restituire `"No shadow"`.

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")` dovrebbe restituire `"166.375ft east"`.

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")` dovrebbe restituire `"No shadow"`.

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")` dovrebbe restituire `"216ft west"`.

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")` dovrebbe restituire `"No shadow"`.

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")` dovrebbe restituire `"91.125ft west"`.

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")` dovrebbe restituire `"No shadow"`.

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
