---
id: 699c8e045ee7cb94ed2322d8
title: "Sfida 217: pezzi degli scacchi catturati"
challengeType: 28
dashedName: challenge-217
---

# --description--

Dato un array di stringhe che rappresentano i pezzi degli scacchi che hai ancora sulla scacchiera, calcola il valore dei pezzi che il tuo avversario ha catturato.

Negli scacchi, inizi con 16 pezzi:

| Pezzo  | Abbreviazione | Quantità | Valore |
| ------ | ------------- | -------- | ------ |
| Pedone | `"P"`     | 8        | 1      |
| Torre  | `"R"`     | 2        | 5      |
| Cavallo| `"N"`     | 2        | 3      |
| Alfiere| `"B"`     | 2        | 3      |
| Regina | `"Q"`     | 1        | 9      |
| Re     | `"K"`     | 1        | 0      |

- L'array dato conterrà solo le abbreviazioni sopra.
- Qualsiasi dei 16 pezzi non incluso nell'array dato è stato catturato.
- Restituisci il valore totale di tutti i pezzi catturati, a meno che...
- Se il Re è stato catturato, restituisci `"Checkmate"`.

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` dovrebbe restituire `8`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])` dovrebbe restituire `26`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` dovrebbe restituire `16`.

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` dovrebbe restituire `4`.

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])` dovrebbe restituire `38`.

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` dovrebbe restituire `0`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` dovrebbe restituire `"Checkmate"`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```
