---
id: 699c8e045ee7cb94ed2322da
title: "Sfida 219: traguardi anniversari"
challengeType: 28
dashedName: challenge-219
---

# --description--

Dato un numero intero che rappresenta gli anni di matrimonio di una coppia, restituisci il loro traguardo anniversario più recente secondo questo grafico:

| Anni di matrimonio | Traguardo    |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- Se non hanno raggiunto il primo traguardo, restituisci `"Newlyweds"`.

# --hints--

`getMilestone(0)` dovrebbe restituire `"Newlyweds"`.

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` dovrebbe restituire `"Paper"`.

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` dovrebbe restituire `"Wood"`.

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` dovrebbe restituire `"Tin"`.

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` dovrebbe restituire `"Silver"`.

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` dovrebbe restituire `"Ruby"`.

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` dovrebbe restituire `"Gold"`.

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` dovrebbe restituire `"Diamond"`.

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` dovrebbe restituire `"Platinum"`.

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
