---
id: 69a890af247de743333bd4d1
title: "Sfida 228: serata al cinema"
challengeType: 28
dashedName: challenge-228
---

# --description--

Dato un giorno della settimana come stringa, un'altra stringa per l'orario dello spettacolo e un numero intero di biglietti, restituisci il costo totale dei biglietti per quella proiezione.

Il giorno fornito sarà uno tra:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

L'orario dello spettacolo sarà nel formato `"H:MMam"` o `"H:MMpm"`. Per esempio `"10:00am"` o `"10:00pm"`.

Restituisci il costo totale nel formato `"$D.CC"` usando queste regole:

- Weekend (venerdì - domenica): 12,00 $ per biglietto.
- Giorni feriali (lunedì - giovedì): 10,00 $ per biglietto.
- Matinée (prima delle 17:00): sconto di 2,00 $ per biglietto (tranne il martedì).
- Martedì: tutti i biglietti costano 5,00 $ ciascuno.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` dovrebbe restituire `"$12.00"`.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` dovrebbe restituire `"$10.00"`.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` dovrebbe restituire `"$10.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` dovrebbe restituire `"$30.00"`.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` dovrebbe restituire `"$32.00"`.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` dovrebbe restituire `"$50.00"`.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` dovrebbe restituire `"$5.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
