---
id: 69a890af247de743333bd4cd
title: "Sfida 224: rilevatore di tostatura del caffè"
challengeType: 28
dashedName: challenge-224
---

# --description--

Data una stringa che rappresenta i chicchi usati per preparare una tazza di caffè, determina la tostatura della tazza.

- La stringa fornita conterrà i seguenti caratteri, ognuno rappresentante un tipo di chicco:

  - Un apostrofo (`'`) è un chicco a tostatura chiara che vale 1 punto ciascuno.
  - Un trattino (`-`) è un chicco a tostatura media che vale 2 punti ciascuno.
  - Un punto (`.`) è un chicco a tostatura scura che vale 3 punti ciascuno.

- Il livello di tostatura si determina dalla media di tutti i chicchi.

Restituisci:

- `"Light"` se la media è inferiore a 1,75.
- `"Medium"` se la media è tra 1,75 e 2,5.
- `"Dark"` se la media è superiore a 2,5.

# --hints--

`detectRoast("''-''''''-'-''--''''")` dovrebbe restituire `"Light"`.

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` dovrebbe restituire `"Medium"`.

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` dovrebbe restituire `"Medium"`.

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` dovrebbe restituire `"Dark"`.

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` dovrebbe restituire `"Medium"`.

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` dovrebbe restituire `"Dark"`.

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` dovrebbe restituire `"Light"`.

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
