---
id: 69a890af247de743333bd4d0
title: "Sfida 227: tempo di raffreddamento"
challengeType: 28
dashedName: challenge-227
---

# --description--

Date due date temporali, la prima che indica quando un utente ha terminato un esame e la seconda che indica l'ora attuale, determina se l'utente può sostenere di nuovo un esame.

- Entrambe le date temporali saranno fornite nel formato: `"YYYY-MM-DDTHH:MM:SS"`, per esempio `"2026-03-25T14:00:00"`. Nota che l'ora è in formato 24 ore.
- Un utente deve aspettare almeno 48 ore prima di ripetere un esame.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` dovrebbe restituire `true`.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` dovrebbe restituire `false`.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` dovrebbe restituire `true`.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` dovrebbe restituire `false`.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
