---
id: 69a890af247de743333bd4cc
title: "Sfida 223: decodificatore QR"
challengeType: 28
dashedName: challenge-223
---

# --description--

Data una matrice 6x6 (array di array), che rappresenta un codice QR, restituisci la stringa di dati binari contenuti nel codice.

- Il codice QR può essere fornito in qualsiasi rotazione a incrementi di 90 gradi.
- Un codice orientato correttamente ha un gruppo 2x2 di `1` (marcatori di orientamento) negli angoli in basso a sinistra, in alto a sinistra e in alto a destra.
- I tre marcatori di orientamento 2x2 non fanno parte dei dati binari.
- I dati binari si leggono da sinistra a destra, dall’alto verso il basso (come un libro) quando il codice QR è orientato correttamente.
- Un codice avrà sempre esattamente un’orientazione valida.

Per esempio, dato:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

o dato lo stesso codice con un’orientazione diversa:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

Restituisci `"000000000000000000000001"`, tutti i dati binari esclusi i tre marcatori di orientamento 2x2.

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])` dovrebbe restituire `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])` dovrebbe restituire `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])` dovrebbe restituire `"001101000011000000110100"`.

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])` dovrebbe restituire `"010001000100010101010110"`.

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])` dovrebbe restituire `"010000100100100101001110"`.

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
