---
id: 69a890af247de743333bd4d2
title: "Sfida 229: troncare il testo 2"
challengeType: 28
dashedName: challenge-229
---

# --description--

Data una stringa, restituisci una nuova stringa troncata in modo che la larghezza totale dei caratteri non superi 50 unità.

Ogni carattere ha una larghezza specifica:

| Lettere | Larghezza |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

La tabella sopra include tutte le lettere maiuscole e minuscole. Inoltre:

- Gli spazi (`" "`) hanno una larghezza di 2
- I punti (`"."`) hanno una larghezza di 1

- Se la stringa data è di 50 unità o meno, restituisci la stringa così com'è, altrimenti
- Tronca la stringa e aggiungi tre punti alla fine (`"..."`) in modo che la larghezza totale, inclusi i tre punti, sia il più possibile vicina a 60 unità senza superarla.

# --hints--

`truncateText("The quick brown fox")` dovrebbe restituire `"The quick brown f..."`.

```js
assert.equal(truncateText("The quick brown fox"), "The quick brown f...");
```

`truncateText("The silky smooth sloth")` dovrebbe restituire `"The silky smooth s..."`.

```js
assert.equal(truncateText("The silky smooth sloth"), "The silky smooth s...");
```

`truncateText("THE LOUD BRIGHT BIRD")` dovrebbe restituire `"THE LOUD BRIG..."`.

```js
assert.equal(truncateText("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...");
```

`truncateText("The fast striped zebra")` dovrebbe restituire `"The fast striped z..."`.

```js
assert.equal(truncateText("The fast striped zebra"), "The fast striped z...");
```

`truncateText("The big black bear")` dovrebbe restituire `"The big black bear"`.

```js
assert.equal(truncateText("The big black bear"), "The big black bear");
```

# --seed--

## --seed-contents--

```js
function truncateText(str) {

  return str;
}
```

# --solutions--

```js
function truncateText(str) {
  const MAX_WIDTH = 50;
  const ELLIPSIS = "...";
  const ELLIPSIS_WIDTH = 3;
  const TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH;

  const charGroups = {
    'ilI.': 1,
    'fjrt ': 2,
    'abcdeghkmnopqrstuvwxyzJL': 3,
    'ABCDEFGHKMNOPQRSTUVWXYZ': 4
  };

  function getCharWidth(char) {
    for (const key in charGroups) {
      if (key.includes(char)) return charGroups[key];
    }
    return 3;
  }

  function stringWidth(str) {
    let totalWidth = 0;
    for (const char of str) {
      totalWidth += getCharWidth(char);
    }
    return totalWidth;
  }

  if (stringWidth(str) <= MAX_WIDTH) return str;

  let result = "";
  let totalWidth = 0;

  for (const char of str) {
    const charWidth = getCharWidth(char);

    if (totalWidth + charWidth > TRUNCATE_LIMIT) break;

    result += char;
    totalWidth += charWidth;
  }

  return result + ELLIPSIS;
}
```
