---
id: 69a890af247de743333bd4ce
title: "Sfida 225: nessuna ripetizione consecutiva"
challengeType: 28
dashedName: challenge-225
---

# --description--

Data una stringa, determina se non contiene caratteri ripetuti.

- Una stringa non ha ripetizioni se non presenta lo stesso carattere due o più volte di seguito.

# --hints--

`hasNoRepeats("hi world")` dovrebbe restituire `true`.

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")` dovrebbe restituire `false`.

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")` dovrebbe restituire `true`.

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")` dovrebbe restituire `false`.

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")` dovrebbe restituire `true`.

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")` dovrebbe restituire `false`.

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```
