---
id: 69b1028d6e265413d0198a2b
title: "Sfida 232: data di scadenza"
challengeType: 28
dashedName: challenge-232
---

# --description--

Data una stringa che rappresenta una data, restituisci la data 9 mesi dopo.

- Le stringhe di input e output hanno il formato `"YYYY-MM-DD"`.
- Se il mese nove mesi dopo non contiene il giorno originale, restituisci l'ultimo giorno di quel mese.

# --hints--

`getDueDate("2025-03-30")` dovrebbe restituire `"2025-12-30"`.

```js
assert.equal(getDueDate("2025-03-30"), "2025-12-30");
```

`getDueDate("2025-04-27")` dovrebbe restituire `"2026-01-27"`.

```js
assert.equal(getDueDate("2025-04-27"), "2026-01-27");
```

`getDueDate("2025-05-29")` dovrebbe restituire `"2026-02-28"`.

```js
assert.equal(getDueDate("2025-05-29"), "2026-02-28");
```

`getDueDate("2026-06-30")` dovrebbe restituire `"2027-03-30"`.

```js
assert.equal(getDueDate("2026-06-30"), "2027-03-30");
```

`getDueDate("2026-10-11")` dovrebbe restituire `"2027-07-11"`.

```js
assert.equal(getDueDate("2026-10-11"), "2027-07-11");
```

# --seed--

## --seed-contents--

```js
function getDueDate(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getDueDate(dateStr) {
  let [year, month, day] = dateStr.split("-").map(Number);

  month += 9;

  year += Math.floor((month - 1) / 12);
  month = ((month - 1) % 12) + 1;

  const daysInMonth = new Date(year, month, 0).getDate();

  if (day > daysInMonth) {
    day = daysInMonth;
  }

  const mm = String(month).padStart(2, "0");
  const dd = String(day).padStart(2, "0");

  return `${year}-${mm}-${dd}`;
}
```
