---
id: 69b1028d6e265413d0198a2f
title: "Sfida 236: cronologia del browser"
challengeType: 28
dashedName: challenge-236
---

# --description--

Dato un array di comandi del browser, restituisci un array con due valori: la cronologia come array di URL e l'indice della pagina corrente.

I comandi validi sono:

- `"URL"` - Dove URL è un indirizzo web (ad esempio `"freecodecamp.org"`). Naviga all'URL indicato, lo aggiunge alla cronologia nella posizione successiva e scarta qualsiasi cronologia futura.
- `"Back"` - si sposta alla pagina precedente nella cronologia, oppure rimane sulla pagina corrente se non ce n'è una.
- `"Forward"` - si sposta alla pagina successiva nella cronologia, oppure rimane sulla pagina corrente se non ce n'è una.

Per esempio, dato `["freecodecamp.org", "freecodecamp.org/learn", "Back"]`, restituisci `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

# --hints--

`getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` dovrebbe restituire `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` dovrebbe restituire `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact",  "example.com/blog", "Back", "Back", "Forward"])` dovrebbe restituire `[["example.com", "example.com/contact", "example.com/blog"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` dovrebbe restituire `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "Back"])` dovrebbe restituire `[["example.com", "example.com/about"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "Forward"])` dovrebbe restituire `[["example.com", "example.com/about"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1]);
```

# --seed--

## --seed-contents--

```js
function getBrowserHistory(commands) {

  return commands;
}
```

# --solutions--

```js
function getBrowserHistory(commands) {
  const stack = [];
  let index = -1;

  for (const command of commands) {
    if (command === "Back") {
      if (index > 0) index--;
    } else if (command === "Forward") {
      if (index < stack.length - 1) index++;
    } else {
      stack.splice(index + 1);
      stack.push(command);
      index++;
    }
  }

  return [stack, index];
}
```
