---
id: 69b1028d6e265413d0198a29
title: "Sfida 230: riga del triangolo di Pascal"
challengeType: 28
dashedName: challenge-230
---

# --description--

Dato un intero `n`, restituisci la `n`-esima riga del triangolo di Pascal come array.

Nel triangolo di Pascal, ogni riga inizia e finisce con 1, e ogni valore interno è la somma dei due valori direttamente sopra di esso.

Ecco le prime 5 righe del triangolo:

```js
    1
   1 1
  1 2 1
 1 3 3 1
1 4 6 4 1
```

# --hints--

`pascalRow(5)` dovrebbe restituire `[1, 4, 6, 4, 1]`.

```js
assert.deepEqual(pascalRow(5), [1, 4, 6, 4, 1]);
```

`pascalRow(3)` dovrebbe restituire `[1, 2, 1]`.

```js
assert.deepEqual(pascalRow(3), [1, 2, 1]);
```

`pascalRow(1)` dovrebbe restituire `[1]`.

```js
assert.deepEqual(pascalRow(1), [1]);
```

`pascalRow(10)` dovrebbe restituire `[1, 9, 36, 84, 126, 126, 84, 36, 9, 1]`.

```js
assert.deepEqual(pascalRow(10), [1, 9, 36, 84, 126, 126, 84, 36, 9, 1]);
```

`pascalRow(15)` dovrebbe restituire `[1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]`.

```js
assert.deepEqual(pascalRow(15), [1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]);
```

# --seed--

## --seed-contents--

```js
function pascalRow(n) {

  return n;
}
```

# --solutions--

```js
function pascalRow(n) {
  if (n === 1) return [1];

  let row = [1];

  for (let i = 2; i <= n; i++) {
    const next = [1];

    for (let j = 1; j < row.length; j++) {
      next.push(row[j - 1] + row[j]);
    }

    next.push(1);
    row = next;
  }

  return row;
}
```
