---
id: 69b1028d6e265413d0198a2a
title: "Sfida 231: validatore ISBN-10"
challengeType: 28
dashedName: challenge-231
---

# --description--

Data una stringa, determina se è un ISBN-10 valido.

Un ISBN-10 è composto da trattini (`"-"`) e 10 altri caratteri. Dopo aver rimosso i trattini (`"-"`):

- I primi 9 caratteri devono essere cifre, e
- L'ultimo carattere può essere una cifra o la lettera `"X"`, che rappresenta il numero 10.

Per validarlo:

- Moltiplica ogni cifra (o valore) per la sua posizione (moltiplica la prima cifra per 1, la seconda per 2 e così via).
- Somma tutti i risultati.
- Se il totale è divisibile per 11, è valido.

# --hints--

`isValidIsbn10("0-306-40615-2")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
