---
id: 69b1028d6e265413d0198a2c
title: "Sfida 233: sveglia"
challengeType: 28
dashedName: challenge-233
---

# --description--

Data una stringa che rappresenta l'ora in cui hai impostato la sveglia e una stringa che rappresenta l'ora in cui ti sei effettivamente svegliato, determina se ti sei svegliato presto, puntuale o in ritardo.

- Entrambi gli orari saranno forniti nel formato 24 ore `"HH:MM"`.

Restituisci:

- `"early"` se ti sei svegliato prima dell'ora della sveglia.
- `"on time"` se ti sei svegliato all'ora della sveglia o entro la finestra di 10 minuti di snooze dopo l'ora della sveglia.
- `"late"` se ti sei svegliato più di 10 minuti dopo l'ora della sveglia.

Entrambi gli orari sono nello stesso giorno.

# --hints--

`alarmCheck("07:00", "06:45")` dovrebbe restituire `"early"`.

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")` dovrebbe restituire `"on time"`.

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")` dovrebbe restituire `"on time"`.

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")` dovrebbe restituire `"late"`.

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")` dovrebbe restituire `"on time"`.

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")` dovrebbe restituire `"late"`.

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")` dovrebbe restituire `"early"`.

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
