---
id: 69b58ce40693f140c84c855a
title: "Sfida 241: validatore fizzbuzz"
challengeType: 28
dashedName: challenge-241
---

# --description--

Dato un array di numeri interi sequenziali, con i multipli di 3 e 5 sostituiti, determina se è una sequenza FizzBuzz valida.

In una sequenza FizzBuzz valida:

- I multipli di 3 sono sostituiti con "Fizz".
- I multipli di 5 sono sostituiti con "Buzz".
- I multipli sia di 3 che di 5 sono sostituiti con "FizzBuzz".
- Tutti gli altri numeri rimangono numeri interi.

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz"])` dovrebbe restituire `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz"]));
```

`isFizzBuzz([13, 14, "FizzBuzz", 16, 17])` dovrebbe restituire `true`.

```js
assert.isTrue(isFizzBuzz([13, 14, "FizzBuzz", 16, 17]));
```

`isFizzBuzz([1, 2, "Fizz", 4, 5])` dovrebbe restituire `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, 5]));
```

`isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` dovrebbe restituire `true`.

```js
assert.isTrue(isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", "Buzz", 5])` dovrebbe restituire `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", "Buzz", 5]));
```

`isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` dovrebbe restituire `false`.

```js
assert.isFalse(isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]));
```

`isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` dovrebbe restituire `true`.

```js
assert.isTrue(isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(arr) {

  return arr;
}
```

# --solutions--

```js
function isFizzBuzz(arr) {
  let start = null;
  for (let i = 0; i < arr.length; i++) {
    if (typeof arr[i] === "number") {
      start = arr[i] - i;
      break;
    }
  }

  if (start === null) return false;

  for (let i = 0; i < arr.length; i++) {
    const n = start + i;
    const expected = n % 15 === 0 ? "FizzBuzz" : n % 3 === 0 ? "Fizz" : n % 5 === 0 ? "Buzz" : n;
    if (arr[i] !== expected) return false;
  }

  return true;
}
```
