---
id: 69b559d2903b9e4afe9075f8
title: "Sfida 238: escape della rotazione delle cifre"
challengeType: 28
dashedName: challenge-238
---

# --description--

Dato un numero intero positivo, determina se esso, o una qualsiasi delle sue rotazioni, è divisibile esattamente per il numero delle sue cifre.

Una rotazione significa spostare la prima cifra alla fine. Per esempio, dopo una rotazione, 123 diventa 231.

- Controlla prima la rotazione `0` (il numero dato).
- I numeri dati non conterranno zeri.
- Restituisci il primo numero di rotazione se ne trovi uno, altrimenti `"none"`.

# --hints--

`getRotation(123)` dovrebbe restituire `0`.

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` dovrebbe restituire `3`.

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` dovrebbe restituire `"none"`.

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` dovrebbe restituire `6`.

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
