---
id: 69b58ce40693f140c84c8559
title: "Sfida 240: caratteri palindromi"
challengeType: 28
dashedName: challenge-240
---

# --description--

Data una stringa, determina se è un palindromo e restituisci il carattere centrale (se la lunghezza è dispari) o i due caratteri centrali (se la lunghezza è pari).

- Un palindromo è una stringa che si legge uguale da sinistra a destra e da destra a sinistra.
- Se non è un palindromo, restituisci `"none"`.

# --hints--

`palindromeLocator("racecar")` dovrebbe restituire `"e"`.

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` dovrebbe restituire `"v"`.

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` dovrebbe restituire `"none"`.

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` dovrebbe restituire `"oo"`.

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` dovrebbe restituire `"00"`.

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
