---
id: 69b5b2be76ec8135a7fbe973
title: "Sfida 244: attacco di torre e alfiere"
challengeType: 28
dashedName: challenge-244
---

# --description--

Dato un valore stringa per la posizione di una torre su una scacchiera e un altro per la posizione di un alfiere, determina se un pezzo può attaccare l'altro.

Una scacchiera standard è 8x8, con colonne etichettate `A` fino a `H` (da sinistra a destra) e righe etichettate `1` fino a `8` (dal basso verso l'alto). Si presenta così:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- Le torri possono muoversi di quante caselle vogliono in orizzontale o verticale.
- Gli alfieri possono muoversi di quante caselle vogliono in qualsiasi direzione diagonale.
- Un pezzo può attaccare un altro se può muoversi nella posizione occupata da quell'altro pezzo.

Restituisci:

- `"rook"` se la torre può attaccare l'alfiere.
- `"bishop"` se l'alfiere può attaccare la torre.
- `"neither"` se nessuno dei due pezzi può attaccare l'altro.

# --hints--

`rookBishopAttack("A1", "A5")` dovrebbe restituire `"rook"`.

```js
assert.equal(rookBishopAttack("A1", "A5"), "rook");
```

`rookBishopAttack("C3", "F6")` dovrebbe restituire `"bishop"`.

```js
assert.equal(rookBishopAttack("C3", "F6"), "bishop");
```

`rookBishopAttack("D4", "D7")` dovrebbe restituire `"rook"`.

```js
assert.equal(rookBishopAttack("D4", "D7"), "rook");
```

`rookBishopAttack("B7", "H1")` dovrebbe restituire `"bishop"`.

```js
assert.equal(rookBishopAttack("B7", "H1"), "bishop");
```

`rookBishopAttack("B3", "C5")` dovrebbe restituire `"neither"`.

```js
assert.equal(rookBishopAttack("B3", "C5"), "neither");
```

`rookBishopAttack("G3", "E8")` dovrebbe restituire `"neither"`.

```js
assert.equal(rookBishopAttack("G3", "E8"), "neither");
```

# --seed--

## --seed-contents--

```js
function rookBishopAttack(rook, bishop) {

  return rook;
}
```

# --solutions--

```js
function rookBishopAttack(rook, bishop) {
  const colDiff = Math.abs(rook.charCodeAt(0) - bishop.charCodeAt(0));
  const rowDiff = Math.abs(rook[1] - bishop[1]);

  if (rook[0] === bishop[0] || rook[1] === bishop[1]) return "rook";
  if (colDiff === rowDiff) return "bishop";
  return "neither";
}
```
