---
id: 69b559d2903b9e4afe9075f7
title: "Sfida 237: validazione di un'equazione"
challengeType: 28
dashedName: challenge-237
---

# --description--

Data una stringa che rappresenta un'equazione matematica, determina se è corretta.

- Il lato sinistro può contenere fino a tre numeri interi positivi e gli operatori `+`, `-`, `*` e `/`.
- L'equazione sarà fornita nel formato: `"number operator number = number"` (con due o tre numeri a sinistra). Per esempio: `"2 + 2 = 4"` o `"2 + 3 - 1 = 4"`.
- Il lato destro sarà sempre un singolo numero intero.

Segui l'ordine standard delle operazioni: moltiplicazione e divisione si valutano prima di addizione e sottrazione, da sinistra a destra.

# --hints--

`isValidEquation("2 + 2 = 4")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidEquation("2 + 2 = 4"));
```

`isValidEquation("2 + 3 - 1 = 4")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidEquation("2 + 3 - 1 = 4"));
```

`isValidEquation("8 / 2 = 4")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidEquation("8 / 2 = 4"));
```

`isValidEquation("10 * 5 = 50")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidEquation("10 * 5 = 50"));
```

`isValidEquation("2 - 2 = 0")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidEquation("2 - 2 = 0"));
```

`isValidEquation("2 + 9 / 3 = 5")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidEquation("2 + 9 / 3 = 5"));
```

`isValidEquation("20 - 2 * 3 = 14")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidEquation("20 - 2 * 3 = 14"));
```

`isValidEquation("2 + 5 = 6")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidEquation("2 + 5 = 6"));
```

`isValidEquation("10 - 2 * 3 = 24")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidEquation("10 - 2 * 3 = 24"));
```

`isValidEquation("3 + 9 / 3 = 4")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidEquation("3 + 9 / 3 = 4"));
```

# --seed--

## --seed-contents--

```js
function isValidEquation(equation) {

  return equation;
}
```

# --solutions--

```js
function isValidEquation(equation) {
  const [left, right] = equation.split(" = ");
  const tokens = left.split(" ");

  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    if (operator === "*" || operator === "/") {
      const leftOperand = Number(tokens[i - 1]);
      const rightOperand = Number(tokens[i + 1]);
      const computed = operator === "*" ? leftOperand * rightOperand : leftOperand / rightOperand;
      tokens.splice(i - 1, 3, String(computed));
      i -= 2;
    }
  }

  let result = Number(tokens[0]);
  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    const operand = Number(tokens[i + 1]);
    result = operator === "+" ? result + operand : result - operand;
  }

  return result === Number(right);
}
```
