---
id: 69b58ce40693f140c84c855b
title: "Sfida 242: numero bingo successivo"
challengeType: 28
dashedName: challenge-242
---

# --description--

Dato un numero bingo, restituisci il numero bingo successivo in ordine sequenziale.

Un numero bingo è una singola lettera seguita da un numero nel suo intervallo secondo questo grafico:

| Lettera | Intervallo numerico |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

Per esempio, dato `"B10"`, restituisci `"B11"`, il numero bingo successivo. Se ti viene dato l’ultimo numero bingo, restituisci `"B1"`.

# --hints--

`getNextBingoNumber("B10")` dovrebbe restituire `"B11"`.

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` dovrebbe restituire `"N34"`.

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` dovrebbe restituire `"N31"`.

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` dovrebbe restituire `"O61"`.

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` dovrebbe restituire `"B1"`.

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
