---
id: 69b559d2903b9e4afe9075f9
title: "Sfida 239: che giorno è?"
challengeType: 28
dashedName: challenge-239
---

# --description--

Dato un timestamp Unix in millisecondi, restituisci il giorno della settimana.

I giorni validi da restituire sono:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Assicurati di ignorare i fusi orari.

# --hints--

`getDayOfWeek(1775492249000)` dovrebbe restituire `"Monday"`.

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` dovrebbe restituire `"Saturday"`.

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` dovrebbe restituire `"Tuesday"`.

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` dovrebbe restituire `"Sunday"`.

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` dovrebbe restituire `"Thursday"`.

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
