---
id: 69bc6cb30c1d112a2e110a04
title: "Sfida 247: ultima lettera"
challengeType: 28
dashedName: challenge-247
---

# --description--

Data una stringa, restituisci la lettera della stringa che appare per ultima nell'alfabeto.

- Se due o più lettere sono a pari merito per l'ultima nell'alfabeto, restituisci la prima.
- Ignora tutti i caratteri che non sono lettere.

# --hints--

`getLastLetter("world")` dovrebbe restituire `"w"`.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` dovrebbe restituire `"W"`.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` dovrebbe restituire `"z"`.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` dovrebbe restituire `"L"`.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` dovrebbe restituire `"T"`.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
