---
id: 69bc6cb30c1d112a2e110a06
title: "Sfida 249: matematica con le stringhe"
challengeType: 28
dashedName: challenge-249
---

# --description--

Data una stringa con numeri e altri caratteri, esegui operazioni matematiche sui numeri in base al conteggio dei caratteri non numerici che li separano.

- Se il numero di caratteri che separano due numeri è pari, usa l'addizione.
- Se è dispari, usa la sottrazione.
- Le cifre consecutive formano un unico numero.
- Le operazioni si applicano da sinistra a destra.
- Ignora i caratteri iniziali e finali che non sono cifre.

Per esempio, data `"3ab10c8"`, restituisci `5`. Somma 3 e 10 per ottenere 13 perché tra di loro c'è un numero pari di caratteri. Poi sottrai 8 da 13 perché tra il risultato e 8 c'è un numero dispari di caratteri.

# --hints--

`doMath("3ab10c8")` dovrebbe restituire `5`.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` dovrebbe restituire `2`.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` dovrebbe restituire `12`.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` dovrebbe restituire `15`.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` dovrebbe restituire `67`.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
