---
id: 69bc6cb30c1d112a2e110a08
title: "Sfida 251: trovare la somma in un array"
challengeType: 28
dashedName: challenge-251
---

# --description--

Dato un array di numeri e un numero obiettivo, restituisci il primo sottoinsieme di due o più numeri che sommati danno l'obiettivo.

- Il "primo" sottoinsieme è quello i cui elementi hanno gli indici più bassi possibili, dando priorità all'indice più basso prima di tutto.
- Ogni numero nell'array può essere usato solo una volta.
- Se non esiste un sottoinsieme valido, restituisci `"Sum not found"`.

Restituisci i numeri corrispondenti come un array nell'ordine in cui appaiono nell'array originale.

# --hints--

`findSum([1, 3, 5, 7], 6)` dovrebbe restituire `[1, 5]`.

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)` dovrebbe restituire `[1, 4]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)` dovrebbe restituire `[1, 2, 3]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)` dovrebbe restituire `[-1, -2, 4]`.

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)` dovrebbe restituire `[3, 1, 4, 2]`.

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` dovrebbe restituire `[1, 2, 3, 5, 9]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)` dovrebbe restituire `"Sum not found"`.

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```
