---
id: 69bc6cb30c1d112a2e110a03
title: "Sfida 246: iniziali del nome"
challengeType: 28
dashedName: challenge-246
---

# --description--

Dato un nome completo come stringa, restituisci le sue iniziali.

- I nomi da inizializzare sono separati da uno spazio.
- Le iniziali devono essere in maiuscolo.
- Le iniziali devono essere separate da punti.

Per esempio, `"Tommy Millwood"` restituisce `"T.M."`.

# --hints--

`getInitials("Tommy Millwood")` dovrebbe restituire `"T.M."`.

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` dovrebbe restituire `"S.P."`.

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` dovrebbe restituire `"F.C.C."`.

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` dovrebbe restituire `"D."`.

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` dovrebbe restituire `"D.V.C.H.N."`.

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
