---
id: 69c5f3d787b1725d5f00c8ba
title: "Sfida 256: direzione del tempo più vicina"
challengeType: 28
dashedName: challenge-256
---

# --description--

Date due ore, determina se puoi passare dalla prima alla seconda più velocemente andando avanti o indietro.

- Le ore sono fornite nel formato 24 ore (`"HH:MM"`)
- L'orologio fa il giro completo (23:59 va a 00:00 muovendosi in avanti, e 00:00 va a 23:59 muovendosi indietro)

Restituisci:

- `"forward"` se muoversi in avanti è più breve
- `"backward"` se muoversi indietro è più breve
- `"equal"` se entrambe le direzioni richiedono lo stesso tempo

# --hints--

`getDirection("10:00", "12:00")` dovrebbe restituire `"forward"`.

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` dovrebbe restituire `"backward"`.

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` dovrebbe restituire `"equal"`.

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` dovrebbe restituire `"forward"`.

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` dovrebbe restituire `"backward"`.

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` dovrebbe restituire `"equal"`.

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
