---
id: 69c5f3d787b1725d5f00c8bd
title: "Sfida 259: esplosione fizzbuzz"
challengeType: 28
dashedName: challenge-259
---

# --description--

Dato un numero intero, restituisci il numero di passaggi necessari per trasformare la parola `"fizzbuzz"` in una stringa con almeno il numero dato di `"z"` usando le seguenti regole:

- Inizia con la stringa `"fizzbuzz"`.
- Ad ogni passaggio, applica le regole standard di FizzBuzz usando la posizione della lettera nella stringa (la prima `"f"` è in posizione 1).
  - Se la posizione della lettera è divisibile per 3, sostituisci la lettera con `"fizz"`
  - Se è divisibile per 5, sostituisci la lettera con `"buzz"`
  - Se è divisibile per 3 e 5, sostituisci la lettera con `"fizzbuzz"`

Quindi dopo 1 passaggio, `"fizzbuzz"` diventa `"fifizzzbuzzfizzzz"`, che contiene 9 `"z"`.

# --hints--

`explodeFizzbuzz(9)` dovrebbe restituire `1`.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` dovrebbe restituire `2`.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` dovrebbe restituire `3`.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` dovrebbe restituire `4`.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` dovrebbe restituire `5`.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` dovrebbe restituire `6`.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` dovrebbe restituire `11`.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` dovrebbe restituire `14`.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
