---
id: 69c5f3d787b1725d5f00c8b9
title: "Sfida 255: squadra di pulizia per la giornata della terra"
challengeType: 28
dashedName: challenge-255
---

# --description--

Oggi è la giornata della Terra. Dato un array di elementi che hai raccolto, restituisci il punteggio totale della pulizia basandoti sulle regole seguenti.

Gli elementi dati saranno uno dei seguenti:

| Elemento | Valore base |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

Un elemento raro è rappresentato come `["rare", value]`. Per esempio, `["rare", 80]`. Gli elementi rari non ricevono bonus per la serie.

- Bonus per la serie: se lo stesso elemento appare consecutivamente, ottiene punti bonus crescenti.
  - Prima occorrenza consecutiva: valore base
  - Seconda: valore base + 1
  - Terza: valore base + 2
  - ecc.
- Moltiplicatore per il quinto elemento: ogni quinto elemento raccolto riceve un moltiplicatore.
  - Quinto elemento: *2
  - Decimo elemento: *3
  - ecc.

- Applica il moltiplicatore dopo aver calcolato eventuali bonus.

# --hints--

`getCleanupScore(["bottle", "straw", "shoe", "battery"])` dovrebbe restituire `44`.

```js
assert.equal(getCleanupScore(["bottle", "straw", "shoe", "battery"]), 44);
```

`getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"])` dovrebbe restituire `58`.

```js
assert.equal(getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"]), 58);
```

`getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` dovrebbe restituire `79`.

```js
assert.equal(getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79);
```

`getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` dovrebbe restituire `358`.

```js
assert.equal(getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358);
```

`getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` dovrebbe restituire `383`.

```js
assert.equal(getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383);
```

# --seed--

## --seed-contents--

```js
function getCleanupScore(items) {

  return items;
}
```

# --solutions--

```js
function getCleanupScore(items) {
  const values = {
    bottle: 10,
    can: 6,
    bag: 8,
    tire: 35,
    straw: 4,
    cardboard: 3,
    newspaper: 3,
    shoe: 12,
    electronics: 25,
    battery: 18,
    mattress: 38
  };

  let total = 0;
  let prevItem = null;
  let streak = 0;

  for (let i = 0; i < items.length; i++) {
    let item = items[i];
    let base;
    let isRare = Array.isArray(item) && item[0] === "rare";

    if (isRare) {
      base = item[1];
      prevItem = null;
      streak = 0;
    } else {
      base = values[item];
      streak = item === prevItem ? streak + 1 : 0;
      prevItem = item;
    }

    let points = base;
    if (!isRare) points += streak;

    if ((i + 1) % 5 === 0) {
      const multiplier = (i + 1) / 5 + 1;
      points *= multiplier;
    }

    total += points;
  }

  return total;
}
```
