---
id: 69c5f3d787b1725d5f00c8b7
title: "Sfida 253: trovare l'acronimo"
challengeType: 28
dashedName: challenge-253
---

# --description--

Data una stringa che rappresenta un acronimo, restituisci il nome completo dell'organizzazione a cui appartiene dalla lista qui sotto:

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

Ogni lettera dell'acronimo dato deve corrispondere alla prima lettera di ogni parola dell'organizzazione a cui appartiene, nello stesso ordine.

# --hints--

`findOrg("NASA")` dovrebbe restituire `"National Avocado Storage Authority"`.

```js
assert.equal(findOrg("NASA"), "National Avocado Storage Authority");
```

`findOrg("CIA")` dovrebbe restituire `"Cats Infiltration Agency"`.

```js
assert.equal(findOrg("CIA"), "Cats Infiltration Agency");
```

`findOrg("FBI")` dovrebbe restituire `"Fluffy Beanbag Inspectors"`.

```js
assert.equal(findOrg("FBI"), "Fluffy Beanbag Inspectors");
```

`findOrg("DOJ")` dovrebbe restituire `"Department Of Jelly"`.

```js
assert.equal(findOrg("DOJ"), "Department Of Jelly");
```

`findOrg("WHO")` dovrebbe restituire `"Wild Honey Organization"`.

```js
assert.equal(findOrg("WHO"), "Wild Honey Organization");
```

`findOrg("EPA")` dovrebbe restituire `"Eating Pancakes Administration"`.

```js
assert.equal(findOrg("EPA"), "Eating Pancakes Administration");
```

# --seed--

## --seed-contents--

```js
function findOrg(acronym) {

  return acronym;
}
```

# --solutions--

```js
function findOrg(acronym) {
  const orgs = [
    "National Avocado Storage Authority",
    "Cats Infiltration Agency",
    "Fluffy Beanbag Inspectors",
    "Department Of Jelly",
    "Wild Honey Organization",
    "Eating Pancakes Administration"
  ];

  for (let org of orgs) {
    const initials = org
      .split(" ")
      .map(word => word[0])
      .join("");

    if (initials.toUpperCase() === acronym.toUpperCase()) {
      return org;
    }
  }

  return "not found";
}
```
