---
id: 69c5f3d787b1725d5f00c8b8
title: "Sfida 254: parole dispari"
challengeType: 28
dashedName: challenge-254
---

# --description--

Data una stringa di parole, restituisci solo le parole con un numero dispari di lettere.

- Le parole nella stringa data saranno separate da uno spazio singolo.
- Restituisci le parole separate da uno spazio singolo.

# --hints--

`getOddWords("This is a super good test")` dovrebbe restituire `"a super"`.

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` dovrebbe restituire `"one two three"`.

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` dovrebbe restituire `"split rainbow sprinkles top"`.

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` dovrebbe restituire `"The quick brown fox the river"`.

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
