---
id: 69c6ff713a52713463aa7928
title: "Sfida 262: parser di query URL"
challengeType: 28
dashedName: challenge-262
---

# --description--

Dato un URL che contiene una stringa di query, analizza la stringa di query in un oggetto (o dizionario) di coppie chiave-valore.

- La stringa di query inizia dopo il `"?"`,
- ogni parametro è separato da `"&"`,
- ogni coppia chiave/valore è separata da `"="`

Per esempio, dato `"https://example.com/search?name=Alice&age=30"`, restituisci:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Tutti i valori devono essere restituiti come stringhe.

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` dovrebbe restituire `{"name": "Alice", "age": "30"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` dovrebbe restituire `{"skill": "programming", "language": "python"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` dovrebbe restituire `{"category": "books", "sort": "asc", "page": "2"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` dovrebbe restituire `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
