---
id: 69c6ff713a52713463aa7927
title: "Sfida 261: parole dei numeri"
challengeType: 28
dashedName: challenge-261
---

# --description--

Dato un numero intero da 0 a 99, restituisci la sua rappresentazione in parole inglesi.

- `0` restituisce `"zero"`.
- I numeri da 1 a 19 hanno nomi unici (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- I multipli di 10 da 20 a 90 hanno nomi propri (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- I numeri da 21 a 99 che non sono multipli di 10 si scrivono come due parole unite da un trattino. Per esempio `"forty-two"` e `"fifty-three"`.

# --hints--

`getNumberWords(0)` dovrebbe restituire `"zero"`.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` dovrebbe restituire `"ten"`.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` dovrebbe restituire `"nineteen"`.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` dovrebbe restituire `"thirty"`.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` dovrebbe restituire `"fifty-three"`.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` dovrebbe restituire `"seven"`.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` dovrebbe restituire `"twelve"`.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` dovrebbe restituire `"sixty"`.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` dovrebbe restituire `"sixty-seven"`.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` dovrebbe restituire `"ninety-eight"`.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
