---
id: 69c6ff713a52713463aa7926
title: "Sfida 260: punteggio parola"
challengeType: 28
dashedName: challenge-260
---

# --description--

Data una parola, restituisci il suo punteggio usando una tabella standard di valori per lettera:

| Lettera | Valore |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- Le lettere maiuscole e minuscole hanno lo stesso valore.

# --hints--

`getWordScore("hi")` dovrebbe restituire `17`.

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` dovrebbe restituire `52`.

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` dovrebbe restituire `169`.

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` dovrebbe restituire `94`.

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
