---
id: 69c6ff713a52713463aa7929
title: "Sfida 263: cruciverba binario"
challengeType: 28
dashedName: challenge-263
---

# --description--

Dato un carattere, determina se la sua rappresentazione binaria a 8 bit può essere trovata nella seguente griglia, orizzontalmente o verticalmente in entrambe le direzioni:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

Per esempio, `"A"` ha la rappresentazione binaria `01000001`, che appare nella prima riga da sinistra a destra.

# --hints--

`isInCrossword("I")` dovrebbe restituire `true`.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` dovrebbe restituire `true`.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` dovrebbe restituire `true`.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` dovrebbe restituire `true`.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` dovrebbe restituire `false`.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` dovrebbe restituire `false`.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` dovrebbe restituire `false`.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` dovrebbe restituire `false`.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
