---
id: 69cfca90e8a0a6d4d6871c50
title: "Sfida 266: buona giornata"
challengeType: 28
dashedName: challenge-266
---

# --description--

Data una stringa oraria in formato `"HH:MM"` (orologio a 24 ore), restituisci:

- `"Good morning"` per orari da `05:00` a `11:59`
- `"Good afternoon"` per orari da `12:00` a `17:59`
- `"Good evening"` per orari da `18:00` a `21:59`
- `"Good night"` per orari da `22:00` a `04:59`

# --hints--

`getGreeting("06:30")` dovrebbe restituire `"Good morning"`.

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")` dovrebbe restituire `"Good afternoon"`.

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")` dovrebbe restituire `"Good evening"`.

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")` dovrebbe restituire `"Good night"`.

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")` dovrebbe restituire `"Good morning"`.

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
