---
id: 69cfca90e8a0a6d4d6871c4f
title: "Sfida 265: parentesi più profonde"
challengeType: 28
dashedName: challenge-265
---

# --description--

Data una stringa contenente parentesi bilanciate, restituisci il contenuto delle parentesi più profondamente annidate.

- Le parentesi possono essere di uno dei tre tipi: `()`, `[]` e `{}`.
- L'input avrà sempre un singolo gruppo più profondo.

Per esempio, dato `"(hello (world))"`, restituisci `"world"`.

# --hints--

`getDeepestBrackets("(hello (world))")` dovrebbe restituire `"world"`.

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` dovrebbe restituire `"inner"`.

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` dovrebbe restituire `"e"`.

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` dovrebbe restituire `"fox"`.

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` dovrebbe restituire `"C"`.

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
