---
id: 69cfca90e8a0a6d4d6871c53
title: "Sfida 269: pasti senza allergeni"
challengeType: 28
dashedName: challenge-269
---

# --description--

Dato un array di pasti e un array di allergeni da evitare, restituisci i nomi di tutti i pasti che non contengono nessuno degli allergeni indicati.

- Ogni pasto ha il formato `[meal, allergens]`, dove `meal` è il nome del pasto e `allergens` è un array degli allergeni contenuti nel pasto. Per esempio, `["pasta", ["wheat", "milk"]]`.
- Gli allergeni da evitare saranno un array di stringhe.

Restituisci i nomi dei pasti sicuri nello stesso ordine in cui sono dati. Se nessun pasto è sicuro, restituisci un array vuoto.

# --hints--

`getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"])` dovrebbe restituire `["salad"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"]), ["salad"]);
```

`getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"])` dovrebbe restituire `["fried rice", "chicken parmesan"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"]), ["fried rice", "chicken parmesan"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"])` dovrebbe restituire `["oatmeal", "granola", "toast"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"]), ["oatmeal", "granola", "toast"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"])` dovrebbe restituire `["granola", "yogurt", "eggs"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"]), ["granola", "yogurt", "eggs"]);
```

# --seed--

## --seed-contents--

```js
function getAllergenFriendlyMeals(meals, allergens) {

  return meals;
}
```

# --solutions--

```js
function getAllergenFriendlyMeals(meals, allergens) {
  return meals
    .filter(([name, mealAllergens]) =>
      !allergens.some(a => mealAllergens.includes(a)))
    .map(([name]) => name);
}
```
