---
id: 69cfca90e8a0a6d4d6871c52
title: "Sfida 268: numero narcisista"
challengeType: 28
dashedName: challenge-268
---

# --description--

Dato un numero intero positivo, determina se è un numero narcisista.

- Un numero è narcisista se la somma di ciascuna delle sue cifre elevata alla potenza del numero totale di cifre è uguale al numero stesso.

Per esempio, 153 ha 3 cifre, e 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153, quindi è narcisista.

# --hints--

`isNarcissistic(153)` dovrebbe restituire `true`.

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` dovrebbe restituire `false`.

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` dovrebbe restituire `true`.

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` dovrebbe restituire `false`.

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` dovrebbe restituire `true`.

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` dovrebbe restituire `false`.

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` dovrebbe restituire `true`.

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` dovrebbe restituire `false`.

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
