---
id: 69cfca90e8a0a6d4d6871c51
title: "Sfida 267: convertitore di parsec"
challengeType: 28
dashedName: challenge-267
---

# --description--

In una galassia lontana, i parsec sono usati per misurare sia il tempo sia la distanza. Dato un numero intero di parsec, restituisci il suo equivalente in tempo o distanza.

- Se l'intero dato è dispari, rappresenta il tempo. Se è pari, rappresenta la distanza.

Usa questi tassi di conversione:

| Parsec | Tempo/Distanza |
| - | - |
| 1 | 2 ore |
| 2 | 6 anni luce |

Restituisci il valore convertito come un intero.

# --hints--

`convertParsecs(1)` dovrebbe restituire `2`.

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` dovrebbe restituire `6`.

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` dovrebbe restituire `62`.

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` dovrebbe restituire `264`.

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` dovrebbe restituire `34`.

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` dovrebbe restituire `42`.

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
