---
id: 69e2383af7832c8032603b8f
title: "Sfida 273: validatore ISBN-13"
challengeType: 28
dashedName: challenge-273
---

# --description--

Data una stringa, determina se è un numero ISBN-13 valido.

Un ISBN-13 valido:

- Contiene solo cifre e trattini
- Ha esattamente 13 cifre dopo aver rimosso i trattini
- Supera il seguente controllo:
  1. Moltiplica ogni cifra per 1 o 3, alternando (moltiplica la prima cifra per 1, la seconda per 3, la terza per 1 e così via).
  2. La somma dei risultati deve essere divisibile per 10.

# --hints--

`isValidIsbn13("9780306406157")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
