---
id: 69e2383af7832c8032603b90
title: "Sfida 274: la persona più anziana"
challengeType: 28
dashedName: challenge-274
---

# --description--

Dato un array di oggetti, ciascuno con una proprietà `"name"` e `"age"`, restituisci un array contenente il nome della persona più anziana.

Se più persone hanno la stessa età più alta, restituisci tutti i loro nomi nell'ordine in cui appaiono nell'input.

# --hints--

`getOldest([{ name: "Brenda", age: 40 }])` dovrebbe restituire `["Brenda"]`.

```js
assert.deepEqual(getOldest([{ name: "Brenda", age: 40 }]), ["Brenda"]);
```

`getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }])` dovrebbe restituire `["Alice"]`.

```js
assert.deepEqual(getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }]), ["Alice"]);
```

`getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }])` dovrebbe restituire `["Bill", "Carol"]`.

```js
assert.deepEqual(getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }]), ["Bill", "Carol"]);
```

`getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }])` dovrebbe restituire `["George", "Holly", "Zach"]`.

```js
assert.deepEqual(getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }]), ["George", "Holly", "Zach"]);
```

# --seed--

## --seed-contents--

```js
function getOldest(people) {

  return people;
}
```

# --solutions--

```js
function getOldest(people) {
  const maxAge = Math.max(...people.map(p => p.age));
  return people.filter(p => p.age === maxAge).map(p => p.name);
}
```
