---
id: 69f35a5bb823ed620fcb7cbc
title: "Sfida 283: stringa a zip"
challengeType: 28
dashedName: challenge-283
---

# --description--

Date due stringhe, restituisci una nuova stringa che intreccia i loro caratteri uno alla volta. Se una stringa è più lunga, aggiungi i caratteri rimanenti alla fine.

Inizia con il primo carattere della prima stringa.

# --hints--

`zipStrings("abc", "123")` dovrebbe restituire `"a1b2c3"`.

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` dovrebbe restituire `"abcdefghijklmnopqrstuvwxyz"`.

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` dovrebbe restituire `"dnaiyght"`.

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` dovrebbe restituire `"pjyatvhaosncript"`.

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` dovrebbe restituire `"freeCodeCamp"`.

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
