---
id: 69f35a5bb823ed620fcb7cbe
title: "Sfida 285: orario dell'incontro"
challengeType: 28
dashedName: challenge-285
---

# --description--

Dato un array 3D che rappresenta le finestre di disponibilità di più persone, restituisci l'orario più presto in cui tutti hanno un'ora libera. Se non esiste un orario del genere, restituisci `"None"`.

- La disponibilità di ogni persona è un array di `[start, end]` coppie di interi in formato 24 ore. Per esempio, `[10, 12]` significherebbe che la persona è disponibile dalle 10 alle 12. Gli orari di inizio vanno da 0 a 23, e quelli di fine da 1 a 24.

Per esempio, dato:

```json
[
  [[10, 12], [15, 16]], // person 1
  [[11, 14], [15, 16]]  // person 2
]
```

Restituisci `11`, l'inizio della loro prima ora libera condivisa.

# --hints--

`getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]])` dovrebbe restituire `11`.

```js
assert.equal(getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]]), 11);
```

`getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]])` dovrebbe restituire `13`.

```js
assert.equal(getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]]), 13);
```

`getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]])` dovrebbe restituire `9`.

```js
assert.equal(getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]]), 9);
```

`getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]])` dovrebbe restituire `None`.

```js
assert.equal(getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]]), "None");
```

`getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]])` dovrebbe restituire `21`.

```js
assert.equal(getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]]), 21);
```

# --seed--

## --seed-contents--

```js
function getMeetingTime(availability) {

  return availability;
}
```

# --solutions--

```js
function getMeetingTime(availability) {
  for (let hour = 0; hour < 24; hour++) {
    const allFree = availability.every(person =>
      person.some(([start, end]) => start <= hour && hour + 1 <= end)
    );
    if (allFree) return hour;
  }
  return "None";
}
```
