---
id: 69f35a5bb823ed620fcb7cbb
title: "Sfida 282: debito di sonno"
challengeType: 28
dashedName: challenge-282
---

# --description--

Dato un array delle ore di sonno per ogni notte fino a oggi e un numero obiettivo di ore per notte, restituisci quante ore di sonno ti servono stanotte per eliminare il debito di sonno.

- Includi le ore di stanotte nel totale necessario per recuperare.
- Se hai dormito abbastanza per coprire l'obiettivo di stanotte o più, restituisci `0`.

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)` dovrebbe restituire `20`.

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)` dovrebbe restituire `10`.

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)` dovrebbe restituire `4`.

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)` dovrebbe restituire `0`.

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)` dovrebbe restituire `0`.

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
