---
id: 69f35a5bb823ed620fcb7cb9
title: "Sfida 280: data ID Mongo"
challengeType: 28
dashedName: challenge-280
---

# --description--

Dato un ID MongoDB come stringa, restituisci il suo tempo di creazione come stringa ISO 8601.

- Un ID MongoDB è una stringa esadecimale di 24 caratteri. I primi 8 caratteri rappresentano un timestamp Unix (in secondi) codificato come un intero in base 16.

Per esempio, `"6a094b50bcf86cd799439011"` ha un timestamp di `"6a094b50"` in esadecimale, che corrisponde a `1778994000` in decimale, rappresentando un tempo di creazione di `"2026-05-17T05:00:00.000Z"`.

# --hints--

`mongoIdToDate("6a094b50bcf86cd799439011")` dovrebbe restituire `"2026-05-17T05:00:00.000Z"`.

```js
assert.equal(mongoIdToDate("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z");
```

`mongoIdToDate("695344eb1f4a4c1123042128")` dovrebbe restituire `"2025-12-30T03:20:11.000Z"`.

```js
assert.equal(mongoIdToDate("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z");
```

`mongoIdToDate("386da62df34123ac54617e56")` dovrebbe restituire `"2000-01-01T07:01:01.000Z"`.

```js
assert.equal(mongoIdToDate("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z");
```

`mongoIdToDate("69f571c3d7711807afd3dd55")` dovrebbe restituire `"2026-05-02T03:38:43.000Z"`.

```js
assert.equal(mongoIdToDate("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z");
```

`mongoIdToDate("68adce01c0e1144d0a90295a")` dovrebbe restituire `"2025-08-26T15:08:49.000Z"`.

```js
assert.equal(mongoIdToDate("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z");
```

# --seed--

## --seed-contents--

```js
function mongoIdToDate(id) {

  return id;
}
```

# --solutions--

```js
function mongoIdToDate(id) {
  const timestamp = parseInt(id.slice(0, 8), 16);
  return new Date(timestamp * 1000).toISOString();
}
```
