---
id: 69f8c998d78ad3171a0713be
title: "Sfida 291: conteggio fizzbuzz"
challengeType: 28
dashedName: challenge-291
---

# --description--

Dato un numero di inizio e uno di fine, conta il numero di apparizioni di fizz e buzz nell'intervallo (inclusivo).

- I numeri divisibili per 3 contano come fizz.
- I numeri divisibili per 5 contano come buzz.
- I numeri divisibili sia per 3 che per 5 contano sia come fizz che come buzz.

Restituisci un oggetto o un dizionario con i conteggi nel formato: `{ fizz, buzz }`.

# --hints--

`fizzBuzzCount(1, 11)` dovrebbe restituire `{fizz: 3, buzz: 2}`.

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` dovrebbe restituire `{fizz: 9, buzz: 6}`.

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` dovrebbe restituire `{fizz: 26, buzz: 16}`.

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` dovrebbe restituire `{fizz: 207, buzz: 125}`.

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` dovrebbe restituire `{fizz: 4074, buzz: 2444}`.

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
