---
id: 69f8c998d78ad3171a0713bc
title: "Sfida 289: somma delle differenze"
challengeType: 28
dashedName: challenge-289
---

# --description--

Dato un array di numeri, restituisci la somma delle differenze tra ogni numero e quello che lo segue.

Per esempio, dato `[1, 3, 4]`, restituisci `3` (2 + 1).

# --hints--

`sumOfDifferences([1, 3, 4])` dovrebbe restituire `3`.

```js
assert.equal(sumOfDifferences([1, 3, 4]), 3);
```

`sumOfDifferences([5, -3, 3, 9, 10])` dovrebbe restituire `5`.

```js
assert.equal(sumOfDifferences([5, -3, 3, 9, 10]), 5);
```

`sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7])` dovrebbe restituire `-16`.

```js
assert.equal(sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16);
```

`sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])` dovrebbe restituire `25`.

```js
assert.equal(sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25);
```

# --seed--

## --seed-contents--

```js
function sumOfDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function sumOfDifferences(arr) {
  return arr.slice(0, -1).reduce((sum, n, i) => sum + (arr[i + 1] - n), 0);
}
```
