---
id: 69f8c998d78ad3171a0713bb
title: "Sfida 288: numero segreto"
challengeType: 28
dashedName: challenge-288
---

# --description--

Dato un numero segreto e un tentativo, determina se il tentativo è corretto.

Restituisci:

- `"higher"` se il numero segreto è maggiore del tentativo.
- `"lower"` se il numero segreto è minore del tentativo.
- `"you got it!"` se il tentativo è corretto.

# --hints--

`guessNumber(50, 30)` dovrebbe restituire `"higher"`.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` dovrebbe restituire `"lower"`.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` dovrebbe restituire `"you got it!"`.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` dovrebbe restituire `"higher"`.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` dovrebbe restituire `"lower"`.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` dovrebbe restituire `"you got it!"`.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
