---
id: 69f8c998d78ad3171a0713bf
title: "Sfida 292: rapporto d'aspetto più ampio"
challengeType: 28
dashedName: challenge-292
---

# --description--

Date due stringhe che indicano dimensioni diverse di un'immagine, restituisci il rapporto d'aspetto dell'immagine con il rapporto larghezza-altezza maggiore.

- Le stringhe fornite saranno nel formato `"WxH"`, per esempio `"1920x1080"`.
- Il rapporto d'aspetto è il rapporto tra larghezza e altezza, ridotto ai numeri interi più bassi. Per esempio, `"1920x1080"` si riduce a `"16:9"`.
- Restituisci una stringa nel formato `"W:H"`, per esempio `"16:9"`.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` dovrebbe restituire `"16:9"`.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` dovrebbe restituire `"4:3"`.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` dovrebbe restituire `"2:1"`.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` dovrebbe restituire `"9:16"`.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` dovrebbe restituire `"43:18"`.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` dovrebbe restituire `"2089:8539"`.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
