---
id: 69f90c1329a94b37e2a2086c
title: "Sfida 293: la mano migliore"
challengeType: 28
dashedName: challenge-293
---

# --description--

Dato un array di cinque stringhe che rappresentano carte da gioco, restituisci il nome della mano migliore.

- Ogni carta è rappresentata come una stringa di due caratteri: il valore seguito dal seme, `"2h"` per esempio.
  - I valori, dal più basso al più alto, sono: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"` e `"A"`.
  - I semi sono: `"h"`, `"d"`, `"c"` e `"s"`.
- Gli assi (`"A"`) possono essere usati come alti o bassi in una scala.

Le mani, in ordine dalla peggiore alla migliore, sono:

| Nome | Descrizione |
| - | - |
| `"High Card"` | Nessuna coppia o meglio |
| `"Pair"` | Due carte dello stesso valore |
| `"Two Pair"` | Due carte di un valore e due di un altro |
| `"Three of a Kind"` | Tre carte dello stesso valore |
| `"Straight"` | Cinque valori consecutivi |
| `"Flush"` | Cinque carte dello stesso seme |
| `"Full House"` | Tre carte di un valore e due di un altro |
| `"Four of a Kind"` | Quattro carte dello stesso valore |
| `"Straight Flush"` | Cinque valori consecutivi dello stesso seme |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` dello stesso seme |

Restituisci il nome della mano migliore.

# --hints--

`getBestHand(["7s", "7h", "7d", "2c", "5h"])` dovrebbe restituire `"Three of a Kind"`.

```js
assert.equal(getBestHand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind");
```

`getBestHand(["Ks", "Kh", "Kd", "4s", "4h"])` dovrebbe restituire `"Full House"`.

```js
assert.equal(getBestHand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House");
```

`getBestHand(["2h", "5h", "7h", "9h", "Jh"])` dovrebbe restituire `"Flush"`.

```js
assert.equal(getBestHand(["2h", "5h", "7h", "9h", "Jh"]), "Flush");
```

`getBestHand(["As", "Ah", "Ad", "Ac", "Kh"])` dovrebbe restituire `"Four of a Kind"`.

```js
assert.equal(getBestHand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind");
```

`getBestHand(["Ts", "Th", "9d", "9c", "8h"])` dovrebbe restituire `"Two Pair"`.

```js
assert.equal(getBestHand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair");
```

`getBestHand(["9c", "8c", "7c", "6c", "5c"])` dovrebbe restituire `"Straight Flush"`.

```js
assert.equal(getBestHand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush");
```

`getBestHand(["As", "Kh", "Jd", "8c", "5h"])` dovrebbe restituire `"High Card"`.

```js
assert.equal(getBestHand(["As", "Kh", "Jd", "8c", "5h"]), "High Card");
```

`getBestHand(["As", "2h", "3d", "4c", "5h"])` dovrebbe restituire `"Straight"`.

```js
assert.equal(getBestHand(["As", "2h", "3d", "4c", "5h"]), "Straight");
```

`getBestHand(["Ts", "Th", "7c", "6d", "5h"])` dovrebbe restituire `"Pair"`.

```js
assert.equal(getBestHand(["Ts", "Th", "7c", "6d", "5h"]), "Pair");
```

`getBestHand(["As", "Ks", "Qs", "Js", "Ts"])` dovrebbe restituire `"Royal Flush"`.

```js
assert.equal(getBestHand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush");
```

# --seed--

## --seed-contents--

```js
function getBestHand(cards) {

  return cards;
}
```

# --solutions--

```js
function getBestHand(cards) {
  const rankOrder = "23456789TJQKA";

  const ranks = cards.map(c => c[0]);
  const suits = cards.map(c => c[1]);

  const isFlush = suits.every(s => s === suits[0]);

  const rankIndices = ranks.map(r => rankOrder.indexOf(r)).sort((a, b) => a - b);

  function isStraight(indices) {
    const isRegular = indices.every((r, i) => i === 0 || r === indices[i - 1] + 1);
    const isLowAce = JSON.stringify(indices) === JSON.stringify([0, 1, 2, 3, 12]);
    return isRegular || isLowAce;
  }

  const straight = isStraight(rankIndices);

  const counts = {};
  ranks.forEach(r => counts[r] = (counts[r] || 0) + 1);
  const freq = Object.values(counts).sort((a, b) => b - a);

  if (isFlush && JSON.stringify(rankIndices) === JSON.stringify([8, 9, 10, 11, 12])) return "Royal Flush";
  if (isFlush && straight) return "Straight Flush";
  if (freq[0] === 4) return "Four of a Kind";
  if (freq[0] === 3 && freq[1] === 2) return "Full House";
  if (isFlush) return "Flush";
  if (straight) return "Straight";
  if (freq[0] === 3) return "Three of a Kind";
  if (freq[0] === 2 && freq[1] === 2) return "Two Pair";
  if (freq[0] === 2) return "Pair";
  return "High Card";
}
```
