---
id: 69f90c1329a94b37e2a2086d
title: "Sfida 294: combinazioni di parentesi tonde"
challengeType: 28
dashedName: challenge-294
---

# --description--

Dato un intero, `n`, restituisci il numero di combinazioni valide di `n` coppie di parentesi tonde.

- Una combinazione valida è una stringa in cui ogni parentesi tonda aperta ha una corrispondente parentesi tonda chiusa, e nessuna parentesi tonda chiusa appare prima della sua parentesi tonda aperta corrispondente.

Per esempio, dato `2`, ci sono 2 combinazioni valide:

```json
(())
()()
```

# --hints--

`getCombinations(2)` dovrebbe restituire `2`.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` dovrebbe restituire `5`.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` dovrebbe restituire `42`.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` dovrebbe restituire `1430`.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` dovrebbe restituire `742900`.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
