---
id: 6a0dcc730cb92a616f86f0c5
title: "Sfida 301: Ultimo carico"
challengeType: 28
dashedName: challenge-301
---

# --description--

Dato il numero di misurini di detersivo per bucato che ti restano e un array con quanti misurini hai usato nei giorni precedenti, restituisci il numero di giorni pieni di detersivo che ti restano.

Calcola la tua media di utilizzo giornaliero dalla cronologia degli utilizzi e considera quella quantità come consumo giornaliero per i giorni a venire.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` dovrebbe restituire `5`.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` dovrebbe restituire `7`.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` dovrebbe restituire `11`.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` dovrebbe restituire `12`.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` dovrebbe restituire `1`.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
