---
id: 6a0dcc730cb92a616f86f0c1
title: "Sfida 297: validatore di schema parte 3"
challengeType: 28
dashedName: challenge-297
---

# --description--

Dato un oggetto (JavaScript) o un dizionario (Python), determina se corrisponde al seguente schema:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles
}
```

- Il simbolo pipe (`|`) significa "o". `role` deve essere uno dei valori elencati in `Roles`.
- Sono consentite chiavi aggiuntive

# --hints--

`isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" })` dovrebbe restituire `true`.

```js
assert.isTrue(isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" }));
```

`isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 })` dovrebbe restituire `true`.

```js
assert.isTrue(isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 }));
```

`isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" })` dovrebbe restituire `true`.

```js
assert.isTrue(isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" }));
```

`isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" })` dovrebbe restituire `true`.

```js
assert.isTrue(isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" }));
```

`isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" })` dovrebbe restituire `true`.

```js
assert.isTrue(isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" }));
```

`isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" })` dovrebbe restituire `false`.

```js
assert.isFalse(isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" }));
```

`isValidSchema({ username: "wendy", posts: 10, verified: true })` dovrebbe restituire `false`.

```js
assert.isFalse(isValidSchema({ username: "wendy", posts: 10, verified: true }));
```

`isValidSchema({ username: "fabian", posts: 1, verified: true, role: true })` dovrebbe restituire `false`.

```js
assert.isFalse(isValidSchema({ username: "fabian", posts: 1, verified: true, role: true }));
```

`isValidSchema({ username: 8, posts: 1, verified: true, role: "user" })` dovrebbe restituire `false`.

```js
assert.isFalse(isValidSchema({ username: 8, posts: 1, verified: true, role: "user" }));
```

`isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" })` dovrebbe restituire `false`.

```js
assert.isFalse(isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" }));
```

`isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" })` dovrebbe restituire `false`.

```js
assert.isFalse(isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role)
  );
}
```
