---
id: 6a0dcd03ee4e68698080ef6b
title: "Sfida 307: regolamenti di zonizzazione"
challengeType: 28
dashedName: challenge-307
---

# --description--

Dato un array bidimensionale (array di array) che rappresenta la disposizione degli edifici di una città, restituisci le coordinate di tutti gli edifici che violano le regole di zonizzazione.

Ogni cella nella griglia contiene una delle etichette della tabella sottostante. Un edificio viola le regole se uno qualsiasi dei suoi (fino a) 4 vicini, orizzontali o verticali, è di un tipo a cui non può essere adiacente.

| Etichetta | Tipo | Non può essere adiacente a |
| - | - | - |
| `"i"` | industriale | `"R"`, `"I"` |
| `"A"` | agricolo | `"C"` |
| `"R"` | residenziale | `"i"`, `"C"` |
| `"I"` | istituzionale | `"i"` |
| `"C"` | commerciale | `"R"`, `"A"` |
| `""` (stringa vuota) | non sviluppato | nessuna restrizione |

Restituisci le coordinate di tutte le celle in violazione come un array di coppie `[row, col]`, in qualsiasi ordine. Se non ci sono violazioni, restituisci un array vuoto.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` dovrebbe restituire `[[0, 0], [0, 1]]`.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` dovrebbe restituire `[[0, 1], [1, 1]]`.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` dovrebbe restituire `[]`.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` dovrebbe restituire `[[0, 1], [0, 2], [0, 3]]`.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` dovrebbe restituire `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
