---
id: 6a0dcd03ee4e68698080ef6c
title: "Sfida 308: validatore di carte di credito"
challengeType: 28
dashedName: challenge-308
---

# --description--

Data una stringa di cifre che rappresenta un numero di carta di credito, determina se è un numero valido usando il seguente metodo:

- Partendo dalla penultima cifra, raddoppia ogni altra cifra procedendo verso sinistra.
- Se il raddoppio di una cifra dà un numero maggiore di 9, sottrai 9.
- Somma tutte le cifre (raddoppiate e non).
- Se il totale è divisibile per 10, il numero è valido.

# --hints--

`isValidCard("4532015112830366")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` dovrebbe restituire `true`.

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` dovrebbe restituire `false`.

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
