---
id: 6a0dcd03ee4e68698080ef68
title: "Sfida 304: disposizioni dell'itinerario"
challengeType: 28
dashedName: challenge-304
---

# --description--

Dato un array di almeno due fermate opzionali per una gita di un giorno, restituisci il numero di disposizioni valide dell'itinerario.

L'itinerario include sempre `"breakfast"`, `"lunch"` e `"dinner"`, che non saranno passati come argomenti. Le fermate opzionali possono essere collocate ovunque nell'itinerario, soggette alle seguenti regole:

- `"breakfast"` è sempre la prima, con almeno una fermata prima di `"lunch"`.
- `"lunch"` deve apparire prima di `"dinner"`, con almeno una fermata in mezzo.
- Al massimo, una fermata opzionale può apparire dopo `"dinner"`.

Restituisci il numero di disposizioni valide.

# --hints--

`getItineraryCount(["library", "park"])` dovrebbe restituire `2`.

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` dovrebbe restituire `18`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` dovrebbe restituire `120`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` dovrebbe restituire `840`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` dovrebbe restituire `55440`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
