---
id: 6a0dcc730cb92a616f86f0bf
title: "Sfida 295: validatore di schema parte 1"
challengeType: 28
dashedName: challenge-295
---

# --description--

Dato un oggetto (JavaScript) o un dizionario (Python), determina se corrisponde al seguente schema:

```json
{
  username: string
}
```

- Sono consentite chiavi extra

# --hints--

`isValidSchema({ username: "bob" })` dovrebbe restituire `true`.

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` dovrebbe restituire `true`.

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` dovrebbe restituire `true`.

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` dovrebbe restituire `false`.

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` dovrebbe restituire `false`.

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
