---
id: 6a0dcd03ee4e68698080ef66
title: "Sfida 302: jet lag"
challengeType: 28
dashedName: challenge-302
---

# --description--

Dato un città di partenza, una città di arrivo, una durata del volo in ore e una direzione di viaggio, restituisci il numero di ore di jet lag che il viaggiatore sta sperimentando.

Le città fornite saranno tratte dalla seguente lista che include il loro offset UTC:

| Città | Offset |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Per calcolare le ore di jet lag:

1. Trova la differenza di fuso orario in ore tra le due città.
2. Determina il moltiplicatore di direzione. Se si viaggia `"east"`, è 1.5, altrimenti è 1.0.
3. Ottieni le ore di jet lag con la formula: differenza di fuso orario + (durata del volo * 0.1) * moltiplicatore di direzione

Restituisci le ore di jet lag arrotondate a una cifra decimale.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` dovrebbe restituire `6.5`.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` dovrebbe restituire `5.8`.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` dovrebbe restituire `1.6`.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` dovrebbe restituire `4.7`.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` dovrebbe restituire `17.5`.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` dovrebbe restituire `5.9`.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` dovrebbe restituire `9.5`.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
