---
id: 6a0dcd03ee4e68698080ef67
title: "Sfida 303: coinquilini"
challengeType: 28
dashedName: challenge-303
---

# --description--

Dato un array di persone e il loro gruppo di coinquilini, restituisci le assegnazioni delle stanze per un soggiorno in hotel seguendo queste regole:

- Ogni persona ha una proprietà `name` e una `group`:

```json
[
  { "name": "Alice", "group": "A" },
  { "name": "Bob", "group": "B" },
  { "name": "Carol", "group": "A" }
]
```

- Le persone possono condividere la stanza solo con qualcuno dello stesso gruppo e sono abbinate nell'ordine in cui sono fornite.
- Restituisci un array di stringhe con i nomi separati da `" and "` per una stanza condivisa, e solo il nome per una stanza singola. I nomi devono apparire nell'ordine in cui sono stati abbinati. Per l'esempio sopra, restituisci `["Alice and Carol", "Bob"]`.

# --hints--

`getRoommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }])` dovrebbe restituire `["Alice and Carol", "Bob"]`.

```js
assert.deepEqual(getRoommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }]).sort(), ["Alice and Carol", "Bob"].sort());
```

`getRoommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }])` dovrebbe restituire `["John and Julia", "Jim"]`.

```js
assert.deepEqual(getRoommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }]).sort(), ["John and Julia", "Jim"].sort());
```

`getRoommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }])` dovrebbe restituire `["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"]`.

```js
assert.deepEqual(getRoommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }]).sort(), ["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"].sort());
```

`getRoommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }])` dovrebbe restituire `["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"]`.

```js
assert.deepEqual(getRoommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }]).sort(), ["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"].sort());
```

`getRoommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }])` dovrebbe restituire `["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"]`.

```js
assert.deepEqual(getRoommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }]).sort(), ["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"].sort());
```

# --seed--

## --seed-contents--

```js
function getRoommates(people) {

  return people;
}
```

# --solutions--

```js
function getRoommates(people) {
  const groups = {};

  for (const { name, group } of people) {
    (groups[group] = groups[group] || []).push(name);
  }

  return Object.values(groups).flatMap(group => {
    const rooms = [];
    for (let i = 0; i < group.length; i += 2) {
      rooms.push(group.slice(i, i + 2).join(" and "));
    }
    return rooms;
  });
}
```
