---
id: 6a15cadf5f240d05a264955a
title: "Sfida 316: 1337 speak"
challengeType: 28
dashedName: challenge-316
---

# --description--

Data una stringa minuscola, restituiscila tradotta in leet speak sostituendo le lettere seguenti con le loro corrispondenti in leet:

| Lettera | Leet |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- I caratteri senza sostituzione rimangono invariati.

# --hints--

`makeLeet("cool")` dovrebbe restituire `"c001"`.

```js
assert.equal(makeLeet("cool"), "c001");
```

`makeLeet("leet")` dovrebbe restituire `"1337"`.

```js
assert.equal(makeLeet("leet"), "1337");
```

`makeLeet("hacker")` dovrebbe restituire `"h4ck3r"`.

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")` dovrebbe restituire `"547311173"`.

```js
assert.equal(makeLeet("satellite"), "547311173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")` dovrebbe restituire `"4bcd3f9h1jk1mn0pqr57uvwxyz"`.

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return str;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", g: "9", i: "1", l: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
