---
id: 6a15cadf5f240d05a264955e
title: "Sfida 320: banca del sangue"
challengeType: 28
dashedName: challenge-320
---

# --description--

Dato un array dell'inventario di una banca del sangue e un array di richieste di tipo di sangue dei pazienti, restituisci una stringa nel formato `"X of Y patients served"`. Dove `X` è il numero massimo di pazienti che possono ricevere sangue dall'inventario della banca e `Y` è il numero totale di pazienti.

Ogni elemento in entrambi gli array è uno dei seguenti tipi di sangue: `"AB"`, `"A"`, `"B"` o `"O"`.

Regole di compatibilità:

- `"AB"` può ricevere da qualsiasi tipo di sangue.
- `"A"` può ricevere da `"A"` e `"O"`.
- `"B"` può ricevere da `"B"` e `"O"`.
- `"O"` può ricevere solo da `"O"`.

Le voci duplicate negli array dati rappresentano la quantità.

# --hints--

`triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` dovrebbe restituire `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` dovrebbe restituire `"3 of 5 patients served"`.

```js
assert.equal(triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served");
```

`triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` dovrebbe restituire `"4 of 5 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served");
```

`triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` dovrebbe restituire `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` dovrebbe restituire `"8 of 11 patients served"`.

```js
assert.equal(triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served");
```

`triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` dovrebbe restituire `"13 of 14 patients served"`.

```js
assert.equal(triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served");
```

# --seed--

## --seed-contents--

```js
function triageBlood(bank, patients) {

  return bank;
}
```

# --solutions--

```js
function triageBlood(bank, patients) {
  const b = {};
  const p = {};
  for (const t of bank) b[t] = (b[t] || 0) + 1;
  for (const t of patients) p[t] = (p[t] || 0) + 1;

  const serve = (patient, donors) => {
    for (const donor of donors) {
      const n = Math.min(p[patient] || 0, b[donor] || 0);
      p[patient] = (p[patient] || 0) - n;
      b[donor] = (b[donor] || 0) - n;
      saved += n;
    }
  };

  let saved = 0;
  serve("O",  ["O"]);
  serve("A",  ["A", "O"]);
  serve("B",  ["B", "O"]);
  serve("AB", ["AB", "A", "B", "O"]);
  return `${saved} of ${patients.length} patients served`;
}
```
