---
id: 6a151d32d772271dd2448c2d
title: "Sfida 311: incantatore"
challengeType: 28
dashedName: challenge-311
---

# --description--

Data una stringa di codici di incantesimi che stai lanciando, calcola il punteggio totale.

Ogni carattere nella stringa rappresenta un incantesimo:

| Codice | Incantesimo | Categoria | Punteggio base |
| - | - | - | - |
| `"f"` | Fuoco | Distruzione | 3 |
| `"l"` | Fulmine | Distruzione | 3 |
| `"i"` | Ghiaccio | Controllo | 2 |
| `"w"` | Vento | Controllo | 2 |
| `"h"` | Cura | Restauro | 1 |
| `"s"` | Scudo | Restauro | 1 |

Un moltiplicatore combo viene applicato in base a quanti incantesimi consecutivi sono stati lanciati da categorie diverse:

- Il primo incantesimo vale sempre il punteggio base.
- Ogni incantesimo consecutivo di categoria diversa dalla precedente aumenta il moltiplicatore di 1.
- Lanciare un incantesimo della stessa categoria del precedente azzera il moltiplicatore a 1.
- Il punteggio di ogni incantesimo è il suo punteggio base moltiplicato per il moltiplicatore corrente.

Restituisci il punteggio totale dalla sequenza di incantesimi.

# --hints--

`cast("fihwl")` dovrebbe restituire `33`.

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` dovrebbe restituire `45`.

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` dovrebbe restituire `37`.

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` dovrebbe restituire `50`.

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` dovrebbe restituire `101`.

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
