---
id: 6a15cadf5f240d05a264955d
title: "Sfida 319: parser frontmatter"
challengeType: 28
dashedName: challenge-319
---

# --description--

Dato una stringa che rappresenta un blocco frontmatter, analizzala e restituisci un oggetto (JavaScript) o un dizionario (Python) con le chiavi e i valori.

Il frontmatter è racchiuso tra i delimitatori `---` e contiene coppie `key: value`, una per riga. Per esempio:

```md
---
title: My Post
draft: false
views: 100
---
```

Dovrebbe restituire:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- Numeri, booleani e stringhe devono essere restituiti come il loro rispettivo tipo.
- La stringa fornita avrà le nuove righe separate dal carattere di nuova linea (`"\n"`). L'esempio sopra sarebbe fornito come: `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`.

# --hints--

`parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` dovrebbe restituire `{ title: "My Post", draft: false, views: 100 }`.

```js
assert.deepEqual(parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---"), { title: "My Post", draft: false, views: 100 });
```

`parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` dovrebbe restituire `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false }`.

```js
assert.deepEqual(parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---"), { id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false });
```

`parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` dovrebbe restituire `{ version: "1.0.0", url: "https://example.com", private: true }`.

```js
assert.deepEqual(parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---"), { version: "1.0.0", url: "https://example.com", private: true });
```

`parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---")` dovrebbe restituire `{ rating: 4.5, price: 9.99 }`.

```js
assert.deepEqual(parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---"), { rating: 4.5, price: 9.99 });
```

# --seed--

## --seed-contents--

```js
function parseFrontmatter(str) {

  return str;
}
```

# --solutions--

```js
function parseFrontmatter(str) {
  const lines = str.split("\n").filter(line => line !== "---" && line !== "");

  return lines.reduce((obj, line) => {
    const [key, ...rest] = line.split(":");
    const value = rest.join(":").trim();

    if (value === "true") obj[key.trim()] = true;
    else if (value === "false") obj[key.trim()] = false;
    else if (!isNaN(value)) obj[key.trim()] = Number(value);
    else obj[key.trim()] = value;

    return obj;
  }, {});
}
```
