---
id: 6a15cadf5f240d05a2649557
title: "Sfida 313: costo del noleggio"
challengeType: 28
dashedName: challenge-313
---

# --description--

Dato un timestamp di noleggio, un timestamp di restituzione e un livello di noleggio, restituisci il costo totale del noleggio inclusi eventuali costi per ritardo.

- I timestamp forniti sono stringhe ISO UTC, per esempio: `"2026-06-18T18:30:00Z"`.
- Il livello di noleggio è il numero di giorni prima della scadenza del noleggio: `1`, `3` o `7`.
- I noleggi devono essere restituiti entro le 12:00 PM UTC o prima nell'ultimo giorno del periodo di noleggio. Per esempio, un noleggio di 1 giorno preso in qualsiasi momento il 15 marzo deve essere restituito entro le 12:00 PM UTC del 16 marzo.
- Ogni giorno oltre la data e l'ora di scadenza comporta una penale per ritardo.

I prezzi sono i seguenti:

| Livello | Costo base | Penale per ritardo al giorno |
| - | - | - |
| 1 giorno | $4.99 | $3.99 |
| 3 giorni | $3.99 | $2.99 |
| 7 giorni | $2.99 | $0.99 |

Restituisci il costo totale arrotondato a due decimali nel formato `"$D.CC"`.

# --hints--

`getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` dovrebbe restituire `"$4.99"`.

```js
assert.equal(getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99");
```

`getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` dovrebbe restituire `"$12.97"`.

```js
assert.equal(getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97");
```

`getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` dovrebbe restituire `"$3.99"`.

```js
assert.equal(getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99");
```

`getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` dovrebbe restituire `"$9.97"`.

```js
assert.equal(getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97");
```

`getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` dovrebbe restituire `"$2.99"`.

```js
assert.equal(getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99");
```

`getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` dovrebbe restituire `"$358.40"`.

```js
assert.equal(getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40");
```

# --seed--

## --seed-contents--

```js
function getRentalCost(rented, returned, tier) {

  return rented;
}
```

# --solutions--

```js
function getRentalCost(rented, returned, tier) {
  const pricing = {
    1: { base: 4.99, late: 3.99 },
    3: { base: 3.99, late: 2.99 },
    7: { base: 2.99, late: 0.99 },
  };

  const rentedDate = new Date(rented);
  const returnDate = new Date(returned);
  const due = new Date(rentedDate);
  due.setUTCDate(due.getUTCDate() + tier);
  due.setUTCHours(12, 0, 0, 0);
  const { base, late } = pricing[tier];

  if (returnDate <= due) {
    return `$${base.toFixed(2)}`;
  }

  const msPerDay = 1000 * 60 * 60 * 24;
  const daysLate = Math.ceil((returnDate - due) / msPerDay);
  const total = base + daysLate * late;

  return `$${total.toFixed(2)}`;
}
```
