---
id: 6a15cadf5f240d05a2649558
title: "Sfida 314: fattorizzazione in fattori primi"
challengeType: 28
dashedName: challenge-314
---

# --description--

Dato un intero maggiore di 1, restituisci la sua fattorizzazione in fattori primi come un array di numeri in ordine crescente.

La fattorizzazione in fattori primi è l'insieme di numeri primi che moltiplicati insieme producono l'intero dato. Ogni numero ha esattamente un insieme. Per esempio, la fattorizzazione in fattori primi di 20 è `[2, 2, 5]` perché 2 * 2 * 5 = 20.

Se l'intero dato è esso stesso primo, restituiscilo in un array con un solo elemento.

# --hints--

`primeFactorization(20)` dovrebbe restituire `[2, 2, 5]`.

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` dovrebbe restituire `[17]`.

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` dovrebbe restituire `[3, 5]`.

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` dovrebbe restituire `[5, 7]`.

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` dovrebbe restituire `[3, 3, 3, 37]`.

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` dovrebbe restituire `[2, 2, 2, 3, 3, 5]`.

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` dovrebbe restituire `[2, 3, 5, 7, 11, 13, 17]`.

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
