---
id: 6a15cadf5f240d05a264955f
title: "Sfida 321: Ortografia periodica"
challengeType: 28
dashedName: challenge-321
---

# --description--

Data una parola, determina se può essere scritta usando i simboli degli elementi della tavola periodica.

- Ignora le maiuscole e minuscole quando scrivi una parola. `"neon"` può essere scritto con i simboli `"Ne"`, `"O"` e `"N"`.

Ecco un elenco completo dei simboli degli elementi:

```json
["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];
```

Restituisci un array degli elementi usati per scrivere la parola, mantenendo la loro maiuscola originale e nell'ordine per scrivere la parola. Oppure, un array vuoto se non può essere scritta.

# --hints--

`getPeriodicSpelling("neon")` dovrebbe restituire `["Ne", "O", "N"]`.

```js
assert.deepEqual(getPeriodicSpelling("neon"), ["Ne", "O", "N"]);
```

`getPeriodicSpelling("rational")` dovrebbe restituire `["Ra", "Ti", "O", "N", "Al"]`.

```js
assert.deepEqual(getPeriodicSpelling("rational"), ["Ra", "Ti", "O", "N", "Al"]);
```

`getPeriodicSpelling("yarn")` dovrebbe restituire `["Y", "Ar", "N"]`.

```js
assert.deepEqual(getPeriodicSpelling("yarn"), ["Y", "Ar", "N"]);
```

`getPeriodicSpelling("carbon")` dovrebbe restituire `["C", "Ar", "B", "O", "N"]` oppure `["Ca", "Rb", "O", "N"]`.

```js
const result = JSON.stringify(getPeriodicSpelling("carbon"));
const path1 = JSON.stringify(["C", "Ar", "B", "O", "N"]);
const path2 = JSON.stringify(["Ca", "Rb", "O", "N"]);
assert.isTrue(result === path1 || result === path2);
```


`getPeriodicSpelling("noisy")` dovrebbe restituire `["N", "O", "I", "S", "Y"]` oppure `["No", "I", "S", "Y"]`.

```js
const result = JSON.stringify(getPeriodicSpelling("noisy"));
const path1 = JSON.stringify(["N", "O", "I", "S", "Y"]);
const path2 = JSON.stringify(["No", "I", "S", "Y"]);
assert.isTrue(result === path1 || result === path2);
```


`getPeriodicSpelling("bicycles")` dovrebbe restituire `["B", "I", "C", "Y", "Cl", "Es"]` oppure `["Bi", "C", "Y", "Cl", "Es"]`.

```js
const result = JSON.stringify(getPeriodicSpelling("bicycles"));
const path1 = JSON.stringify(["B", "I", "C", "Y", "Cl", "Es"]);
const path2 = JSON.stringify(["Bi", "C", "Y", "Cl", "Es"]);
assert.isTrue(result === path1 || result === path2);
```

`getPeriodicSpelling("optics")` dovrebbe restituire `["O", "P", "Ti", "C", "S"]`, `["O", "P", "Ti", "Cs"]`, `["O", "Pt", "I", "C", "S"]` oppure `["O", "Pt", "I", "Cs"]`.

```js
const result = JSON.stringify(getPeriodicSpelling("optics"));
const path1 = JSON.stringify(["O", "P", "Ti", "C", "S"]);
const path2 = JSON.stringify(["O", "P", "Ti", "Cs"]);
const path3 = JSON.stringify(["O", "Pt", "I", "C", "S"]);
const path4 = JSON.stringify(["O", "Pt", "I", "Cs"]);
assert.isTrue(result === path1 || result === path2 || result === path3 || result === path4);
```

`getPeriodicSpelling("value")` dovrebbe restituire `[]`.

```js
assert.deepEqual(getPeriodicSpelling("value"), []);
```

# --seed--

## --seed-contents--

```js
function getPeriodicSpelling(word) {

  return word;
}
```

# --solutions--

```js
function getPeriodicSpelling(word) {
  const elements = ["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];

  const lower = new Map(elements.map(e => [e.toLowerCase(), e]));

  function spell(word, i) {
    if (i === word.length) return [];
    const one = word.slice(i, i + 1);
    const two = word.slice(i, i + 2);
    if (lower.has(two)) {
      const rest = spell(word, i + 2);
      if (rest !== null) return [lower.get(two), ...rest];
    }
    if (lower.has(one)) {
      const rest = spell(word, i + 1);
      if (rest !== null) return [lower.get(one), ...rest];
    }
    return null;
  }

  return spell(word.toLowerCase(), 0) ?? [];
}
```
