---
id: 6a19b1062d1b153d8ac76d73
title: "Sfida 325: numero fortunato"
challengeType: 28
dashedName: challenge-325
---

# --description--

Dato una stringa con il nome e il cognome di una persona, calcola il suo numero fortunato usando le seguenti regole:

- Il nome e il cognome sono separati da uno spazio
- Trova il conteggio delle vocali e delle consonanti per ogni nome
- Moltiplica tra loro i conteggi minori di vocali e consonanti e poi per la lunghezza del nome più corto
- Fai lo stesso per i due conteggi maggiori e per il nome più lungo
- Sottrai il valore minore da quello maggiore per ottenere il numero fortunato

Se il valore finale è zero (`0`), restituisci `13`.

# --hints--

`getLuckyNumber("John Doe")` dovrebbe restituire `21`.

```js
assert.equal(getLuckyNumber("John Doe"), 21);
```

`getLuckyNumber("Olivia Lewis")` dovrebbe restituire `52`.

```js
assert.equal(getLuckyNumber("Olivia Lewis"), 52);
```

`getLuckyNumber("James Wilson")` dovrebbe restituire `18`.

```js
assert.equal(getLuckyNumber("James Wilson"), 18);
```

`getLuckyNumber("Elizabeth Hernandez")` dovrebbe restituire `81`.

```js
assert.equal(getLuckyNumber("Elizabeth Hernandez"), 81);
```

`getLuckyNumber("Mike Walker")` dovrebbe restituire `32`.

```js
assert.equal(getLuckyNumber("Mike Walker"), 32);
```

`getLuckyNumber("Chloe Perez")` dovrebbe restituire `13`.

```js
assert.equal(getLuckyNumber("Chloe Perez"), 13);
```

# --seed--

## --seed-contents--

```js
function getLuckyNumber(name) {

  return name;
}
```

# --solutions--

```js
function getLuckyNumber(name) {
  const vowels = "aeiouAEIOU";

  function nameStats(str) {
    const vowelCount = str.split("").filter(c => vowels.includes(c)).length;
    const consonantCount = str.split("").filter(c => c.match(/[a-zA-Z]/) && !vowels.includes(c)).length;
    return { vowelCount, consonantCount, length: str.length };
  }

  const [first, last] = name.split(" ");
  const firstStats = nameStats(first);
  const lastStats = nameStats(last);

  const largerVowels = Math.max(firstStats.vowelCount, lastStats.vowelCount);
  const smallerVowels = Math.min(firstStats.vowelCount, lastStats.vowelCount);
  const largerConsonants = Math.max(firstStats.consonantCount, lastStats.consonantCount);
  const smallerConsonants = Math.min(firstStats.consonantCount, lastStats.consonantCount);
  const largerLength = Math.max(firstStats.length, lastStats.length);
  const smallerLength = Math.min(firstStats.length, lastStats.length);

  const largerValue = largerVowels * largerConsonants * largerLength;
  const smallerValue = smallerVowels * smallerConsonants * smallerLength;
  const result = largerValue - smallerValue;

  return result === 0 ? 13 : result;
}
```
