---
id: 6a19b1062d1b153d8ac76d72
title: "Sfida 324: conteggio dei caratteri duplicati"
challengeType: 28
dashedName: challenge-324
---

# --description--

Date due stringhe, restituisci il conteggio dei caratteri della seconda stringa che si trovano nella prima.

- I caratteri duplicati nella seconda stringa vengono contati separatamente.

# --hints--

`duplicateCharacterCount("aloha", "hei")` dovrebbe restituire `1`.

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` dovrebbe restituire `4`.

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` dovrebbe restituire `3`.

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` dovrebbe restituire `0`.

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` dovrebbe restituire `5`.

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` dovrebbe restituire `2`.

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` dovrebbe restituire `26`.

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
