---
id: 6a1d9f98e819ed70a0e994da
title: "Sfida 328: la routine di Kaprekar"
challengeType: 28
dashedName: challenge-328
---

# --description--

Dato un numero di 4 cifre, restituisci il numero di volte che devi applicare la routine di Kaprekar fino a raggiungere 6174.

La routine di Kaprekar funziona così:

- Disponi le cifre in ordine decrescente per formare il numero più grande
- Disponi le cifre in ordine crescente per formare il numero più piccolo (aggiungi zeri iniziali se necessario)
- Sottrai il numero più piccolo da quello più grande
- Ripeti con il nuovo numero

# --hints--

`kaprekar(1234)` dovrebbe restituire `3`.

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` dovrebbe restituire `6`.

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` dovrebbe restituire `4`.

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` dovrebbe restituire `7`.

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` dovrebbe restituire `2`.

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
