---
id: 6a1d9f98e819ed70a0e994de
title: "Sfida 332: smistamento dei problemi"
challengeType: 28
dashedName: challenge-332
---

# --description--

Dato un numero di millisecondi dall'ultimo post su un problema e l'ultimo messaggio pubblicato sul problema, determina cosa dovresti fare con il problema secondo queste regole:

- Se l'ultimo messaggio risale a meno di 7 giorni fa, restituisci `"leave it"`
- Se l'ultimo messaggio risale a 7 o più giorni fa e il suo contenuto contiene `"bump"` (non sensibile alle maiuscole), restituisci `"close it"`
- Altrimenti, restituisci `"bump it"`

# --hints--

`triageIssue(86400000, "Lets fix it")` dovrebbe restituire `"leave it"`.

```js
assert.equal(triageIssue(86400000, "Lets fix it"), "leave it");
```

`triageIssue(1209600000, "still waiting")` dovrebbe restituire `"bump it"`.

```js
assert.equal(triageIssue(1209600000, "still waiting"), "bump it");
```

`triageIssue(864000000, "bump")` dovrebbe restituire `"close it"`.

```js
assert.equal(triageIssue(864000000,  "bump"), "close it");
```

`triageIssue(604800000, "Do we still want this?")` dovrebbe restituire `"bump it"`.

```js
assert.equal(triageIssue(604800000, "Do we still want this?"), "bump it");
```

`triageIssue(604800000, "Bumping this")` dovrebbe restituire `"close it"`.

```js
assert.equal(triageIssue(604800000, "Bumping this"), "close it");
```

`triageIssue(345600000, "I'll make a PR")` dovrebbe restituire `"leave it"`.

```js
assert.equal(triageIssue(345600000, "I'll make a PR"), "leave it");
```

# --seed--

## --seed-contents--

```js
function triageIssue(ms, message) {

  return ms;
}
```

# --solutions--

```js
function triageIssue(ms, message) {
  const sevenDays = 7 * 24 * 60 * 60 * 1000;

  if (ms < sevenDays) return "leave it";
  if (message.toLowerCase().includes("bump")) return "close it";
  return "bump it";
}
```
