---
id: 6a1d9f98e819ed70a0e994df
title: "Sfida 333: Smistamento problemi 2"
challengeType: 28
dashedName: challenge-333
---

# --description--

Dato un titolo di un problema e un array di etichette attuali, restituisci un array aggiornato di etichette basato sulle seguenti regole:

Se il problema non ha etichette, aggiungi:

- `"bug"` e `"needs triage"` se il titolo contiene `"error"` o `"bug"`
- `"enhancement"` e `"discussing"` se il titolo contiene `"feature"` o `"add"`

Altrimenti, se le etichette date contengono:

- `"needs triage"` e il titolo contiene `"simple"` o `"easy"`, rimuovi `"needs triage"` e aggiungi `"good first issue"`
- `"discussing"` e il titolo contiene `"planned"` o `"next"`, rimuovi `"discussing"` e aggiungi `"on the roadmap"`
- Altrimenti, se è presente `"needs triage"` o `"discussing"`, rimuovila e aggiungi `"help wanted"`

Se il titolo contiene:

- `"security"`, aggiungi un'etichetta `"critical"`

# --hints--

`triageIssue("app crashes with error", [])` dovrebbe restituire `["bug", "needs triage"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", []).sort(), ["bug", "needs triage"].sort());
```

`triageIssue("app crashes with error", ["bug", "needs triage"])` dovrebbe restituire `["bug", "help wanted"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", ["bug", "needs triage"]).sort(), ["bug", "help wanted"].sort());
```

`triageIssue("add dark mode", [])` dovrebbe restituire `["enhancement", "discussing"]`.

```js
assert.deepEqual(triageIssue("add dark mode", []).sort(), ["enhancement", "discussing"].sort());
```

`triageIssue("add dark mode", ["enhancement", "discussing"])` dovrebbe restituire `["enhancement", "help wanted"]`.

```js
assert.deepEqual(triageIssue("add dark mode", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted"].sort());
```

`triageIssue("xss security bug", [])` dovrebbe restituire `["bug", "needs triage", "critical"]`.

```js
assert.deepEqual(triageIssue("xss security bug", []).sort(), ["bug", "needs triage", "critical"].sort());
```

`triageIssue("security vulnerability in auth", [])` dovrebbe restituire `["critical"]`.

```js
assert.deepEqual(triageIssue("security vulnerability in auth", []).sort(), ["critical"].sort());
```

`triageIssue("easy a11y fix", ["bug", "needs triage"])` dovrebbe restituire `["bug", "good first issue"]`.

```js
assert.deepEqual(triageIssue("easy a11y fix", ["bug", "needs triage"]).sort(), ["bug", "good first issue"].sort());
```

`triageIssue("planned api migration", ["enhancement", "discussing"])` dovrebbe restituire `["enhancement", "on the roadmap"]`.

```js
assert.deepEqual(triageIssue("planned api migration", ["enhancement", "discussing"]).sort(), ["enhancement", "on the roadmap"].sort());
```

`triageIssue("improve security", ["enhancement", "discussing"])` dovrebbe restituire `["enhancement", "help wanted", "critical"]`.

```js
assert.deepEqual(triageIssue("improve security", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted", "critical"].sort());
```

# --seed--

## --seed-contents--

```js
function triageIssue(title, labels) {

  return title;
}
```

# --solutions--

```js
function triageIssue(title, labels) {
  const result = [...labels];
  const t = title.toLowerCase();

  const has = label => result.includes(label);
  const add = label => { if (!has(label)) result.push(label); };
  const remove = label => { const i = result.indexOf(label); if (i > -1) result.splice(i, 1); };

  if (result.length === 0) {
    if (t.includes("error") || t.includes("bug")) {
      add("bug"); add("needs triage");
    } else if (t.includes("feature") || t.includes("add")) {
      add("enhancement"); add("discussing");
    }
  } else {
    if (has("needs triage") && (t.includes("simple") || t.includes("easy"))) {
      remove("needs triage"); add("good first issue");
    } else if (has("discussing") && (t.includes("planned") || t.includes("next"))) {
      remove("discussing"); add("on the roadmap");
    } else if (has("needs triage") || has("discussing")) {
      remove("needs triage"); remove("discussing"); add("help wanted");
    }
  }

  if (t.includes("security")) {
    add("critical");
  }

  return result;
}
```
