---
id: 6a1d9f98e819ed70a0e994dd
title: "Sfida 331: multiplo più vicino"
challengeType: 28
dashedName: challenge-331
---

# --description--

Dati due numeri interi, arrotonda il primo al multiplo più vicino del secondo.

# --hints--

`roundToNearestMultiple(5, 3)` dovrebbe restituire `6`.

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` dovrebbe restituire `16`.

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` dovrebbe restituire `45`.

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` dovrebbe restituire `33`.

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` dovrebbe restituire `96`.

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
