---
id: 6a1d9f98e819ed70a0e994e0
title: "Sfida 334: resto esatto"
challengeType: 28
dashedName: challenge-334
---

# --description--

Dato un importo intero in centesimi, restituisci il numero di modi distinti per dare il resto esatto usando pennies (1 cent), nickels (5 cent), dimes (10 cent) e quarters (25 cent).

# --hints--

`exactChange(3)` dovrebbe restituire `1`.

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` dovrebbe restituire `2`.

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` dovrebbe restituire `6`.

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` dovrebbe restituire `24`.

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` dovrebbe restituire `73`.

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` dovrebbe restituire `213`.

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
