---
id: 6a1d9f98e819ed70a0e994dc
title: "Sfida 330: parole minuscole"
challengeType: 28
dashedName: challenge-330
---

# --description--

Data una stringa, restituisci solo le parole che sono interamente minuscole, nel loro ordine originale e con uno spazio tra ogni parola.

# --hints--

`getLowercaseWords("hello GOOD world")` dovrebbe restituire `"hello world"`.

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` dovrebbe restituire `"these are all lowercase"`.

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` dovrebbe restituire `"less is more"`.

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` dovrebbe restituire `"eat pizza every day"`.

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` dovrebbe restituire `"the quick brown fox jumped over the lazy dog"`.

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
