---
id: 6a2037a68a0bc2aef0005fff
title: "Sfida 337: contatore di segni"
challengeType: 28
dashedName: challenge-337
---

# --description--

Data una stringa di segni di conteggio, restituisci il totale rappresentato.

- Ogni barra verticale `"|"` rappresenta un conteggio.
- Ogni quinto segno è rappresentato da una barra obliqua in avanti `"/"`, completando un gruppo di cinque (`"||||/"`).
- I gruppi sono separati da uno spazio.

# --hints--

`getTallyCount("||||")` dovrebbe restituire `4`.

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` dovrebbe restituire `5`.

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` dovrebbe restituire `8`.

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` dovrebbe restituire `17`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` dovrebbe restituire `41`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
