---
id: 6a2037a68a0bc2aef0006004
title: "Sfida 342: probabilità con i dadi"
challengeType: 28
dashedName: challenge-342
---

# --description--

Dato un numero di dadi a sei facce da lanciare e una somma obiettivo, restituisci la probabilità di ottenere quella somma come stringa nel formato `"1 in X"`.

- Il numero di dadi sarà tra 1 e 6.
- La somma obiettivo è sempre raggiungibile con il numero di dadi dato.
- Arrotonda `"X"` al numero intero più vicino.

# --hints--

`getOdds(1, 5)` dovrebbe restituire `"1 in 6"`.

```js
assert.equal(getOdds(1, 5), "1 in 6");
```

`getOdds(2, 4)` dovrebbe restituire `"1 in 12"`.

```js
assert.equal(getOdds(2, 4), "1 in 12");
```

`getOdds(3, 10)` dovrebbe restituire `"1 in 8"`.

```js
assert.equal(getOdds(3, 10), "1 in 8");
```

`getOdds(4, 7)` dovrebbe restituire `"1 in 65"`.

```js
assert.equal(getOdds(4, 7), "1 in 65");
```

`getOdds(5, 26)` dovrebbe restituire `"1 in 111"`.

```js
assert.equal(getOdds(5, 26), "1 in 111");
```

`getOdds(6, 35)` dovrebbe restituire `"1 in 7776"`.

```js
assert.equal(getOdds(6, 35), "1 in 7776");
```

# --seed--

## --seed-contents--

```js
function getOdds(dice, target) {

  return dice;
}
```

# --solutions--

```js
function getOdds(dice, target) {
  let validCombos = 0;
  const totalCombos = Math.pow(6, dice);

  function roll(diceLeft, sum) {
    if (diceLeft === 0) {
      if (sum === target) validCombos++;
      return;
    }
    for (let face = 1; face <= 6; face++) {
      roll(diceLeft - 1, sum + face);
    }
  }

  roll(dice, 0);
  return `1 in ${Math.round(totalCombos / validCombos)}`;
}
```
