---
id: 6a2037a68a0bc2aef0006007
title: "Sfida 345: fusione di parole"
challengeType: 28
dashedName: challenge-345
---

# --description--

Date due parole, restituisci una nuova parola combinando la prima metà della prima parola con la seconda metà della seconda parola.

- Per le parole di lunghezza dispari, la prima metà è la metà più corta.

# --hints--

`blendWords("turtle", "toucan")` dovrebbe restituire `"turcan"`.

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")` dovrebbe restituire `"chipingo"`.

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")` dovrebbe restituire `"falican"`.

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")` dovrebbe restituire `"hyana"`.

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")` dovrebbe restituire `"scorilla"`.

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")` dovrebbe restituire `"platerine"`.

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
