---
id: 6a2037a68a0bc2aef0006002
title: "Sfida 340: convertitore in Pig Latin"
challengeType: 28
dashedName: challenge-340
---

# --description--

Data una stringa, convertila in Pig Latin usando le seguenti regole:

- Se una parola inizia con una vocale (`"a"`, `"e"`, `"i"`, `"o"` o `"u"`), aggiungi `"way"` alla fine. Per esempio, `"universe"` si converte in `"universeway"`.
- Se una parola inizia con una o più consonanti, spostale alla fine e aggiungi `"ay"`. Per esempio, `"hello"` si converte in `"ellohay"`.
- Conserva il caso della prima lettera. Per esempio, `"Hello"` si converte in `"Ellohay"`.

# --hints--

`pigLatin("universe")` dovrebbe restituire `"universeway"`.

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` dovrebbe restituire `"ellohay"`.

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` dovrebbe restituire `"ellohay universeway"`.

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` dovrebbe restituire `"Ellohay universeway"`.

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` dovrebbe restituire `"Igpay Atinlay isway unfay"`.

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` dovrebbe restituire `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
