---
id: 6a2037a68a0bc2aef0006006
title: "Sfida 344: rapporto aureo"
challengeType: 28
dashedName: challenge-344
---

# --description--

Dato due numeri, determina se il loro rapporto si avvicina al rapporto aureo.

- Usa un rapporto aureo di `1.618`
- Consenti una tolleranza di `0.01`

# --hints--

`isGoldenRatio(21, 34)` dovrebbe restituire `true`.

```js
assert.isTrue(isGoldenRatio(21, 34));
```

`isGoldenRatio(15, 20)` dovrebbe restituire `false`.

```js
assert.isFalse(isGoldenRatio(15, 20));
```

`isGoldenRatio(8, 13)` dovrebbe restituire `true`.

```js
assert.isTrue(isGoldenRatio(8, 13));
```

`isGoldenRatio(10, 16)` dovrebbe restituire `false`.

```js
assert.isFalse(isGoldenRatio(10, 16));
```

`isGoldenRatio(1618, 1000)` dovrebbe restituire `true`.

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

`isGoldenRatio(88, 55)` dovrebbe restituire `false`.

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
