---
id: 6a2037a68a0bc2aef0006000
title: "Sfida 338: calcolatore dell'età degli animali domestici"
challengeType: 28
dashedName: challenge-338
---

# --description--

Dato un tipo di animale domestico e la sua età in anni umani, restituisci l'età equivalente in anni dell'animale usando la seguente tabella di conversione:

| Animale | Moltiplicatore |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` dovrebbe restituire `35`.

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` dovrebbe restituire `54`.

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` dovrebbe restituire `24`.

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` dovrebbe restituire `120`.

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` dovrebbe restituire `60`.

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` dovrebbe restituire `12`.

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` dovrebbe restituire `5`.

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
