---
id: 6a2037a68a0bc2aef0006001
title: "Sfida 339: sotto-array di un array"
challengeType: 28
dashedName: challenge-339
---

# --description--

Dato un array e una dimensione per i sotto-array, restituisci l'array diviso in sotto-array di quella dimensione.

- L'ultimo sotto-array può essere più piccolo se l'array non si divide in modo uniforme.

# --hints--

`chunkArray([1, 2, 3, 4, 5, 6], 3)` dovrebbe restituire `[[1, 2, 3], [4, 5, 6]]`.

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]]);
```

`chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` dovrebbe restituire `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]`.

```js
assert.deepEqual(chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]);
```

`chunkArray([1, 2, 3, 4, 5], 3)` dovrebbe restituire `[[1, 2, 3], [4, 5]]`.

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]]);
```

`chunkArray(["a", "b", "c", "d", "e"], 1)` dovrebbe restituire `[["a"], ["b"], ["c"], ["d"], ["e"]]`.

```js
assert.deepEqual(chunkArray(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]]);
```

`chunkArray([1, 2, 3], 5)` dovrebbe restituire `[[1, 2, 3]]`.

```js
assert.deepEqual(chunkArray([1, 2, 3], 5), [[1, 2, 3]]);
```

# --seed--

## --seed-contents--

```js
function chunkArray(arr, size) {

  return arr;
}
```

# --solutions--

```js
function chunkArray(arr, size) {
  const result = [];
  for (let i = 0; i < arr.length; i += size) {
    result.push(arr.slice(i, i + size));
  }
  return result;
}
```
