---
id: 6a22d77ddf034bc4e35b1d5b
title: "Sfida 354: valutazione del contrasto 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

Dato due array che rappresentano valori RGB e un booleano che indica se il testo è grande, restituisci la valutazione del contrasto WCAG usando il seguente metodo:

Per prima cosa, converti ogni valore RGB in luminanza relativa:

- Dividi ogni canale `[R, G, B]` per 255 per ottenere un valore tra `0` e `1`
- Applica la formula di correzione gamma a ogni canale:
  - Se il valore del canale è minore o uguale a `0.04045`: `channel / 12.92`
  - Altrimenti: `((channel + 0.055) / 1.055) ^ 2.4`
- Calcola la luminanza: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Poi, calcola il rapporto di contrasto aggiungendo 0.05 a ogni valore di luminanza, quindi dividendo quello più chiaro per quello più scuro. Quello più chiaro sarà sempre il primo argomento.

Restituisci la valutazione basata sul rapporto di contrasto usando la seguente tabella:

| Valutazione | Testo normale | Testo grande |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | sotto 4.5 | sotto 3.0 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)` dovrebbe restituire `"AAA"`.

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)` dovrebbe restituire `"AA"`.

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)` dovrebbe restituire `"Fail"`.

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)` dovrebbe restituire `"AAA"`.

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)` dovrebbe restituire `"AA"`.

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)` dovrebbe restituire `"Fail"`.

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
