---
id: 6a22d77ddf034bc4e35b1d58
title: "Sfida 351: numero prono"
challengeType: 28
dashedName: challenge-351
---

# --description--

Dato un numero, determina se è un numero prono.

Un numero prono è il prodotto di due numeri interi consecutivi. Per esempio, 6 è prono perché 2 * 3 = 6.

# --hints--

`isPronic(6)` dovrebbe restituire `true`.

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` dovrebbe restituire `false`.

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` dovrebbe restituire `true`.

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` dovrebbe restituire `true`.

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` dovrebbe restituire `false`.

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` dovrebbe restituire `true`.

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
