---
id: 6a22d77ddf034bc4e35b1d59
title: "Sfida 352: valutazione del contrasto 1"
challengeType: 28
dashedName: challenge-352
---

# --description--

Dato un rapporto di contrasto e un booleano che indica se il testo è grande, restituisci la valutazione WCAG usando la seguente tabella:

| Valutazione | Testo normale | Testo grande |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | sotto 4.5 | sotto 3.0 |

# --hints--

`getContrastRating("7.5", false)` dovrebbe restituire `"AAA"`.

```js
assert.equal(getContrastRating("7.5", false), "AAA");
```

`getContrastRating("4.8", false)` dovrebbe restituire `"AA"`.

```js
assert.equal(getContrastRating("4.8", false), "AA");
```

`getContrastRating("4.2", false)` dovrebbe restituire `"Fail"`.

```js
assert.equal(getContrastRating("4.2", false), "Fail");
```

`getContrastRating("4.5", true)` dovrebbe restituire `"AAA"`.

```js
assert.equal(getContrastRating("4.5", true), "AAA");
```

`getContrastRating("3.0", true)` dovrebbe restituire `"AA"`.

```js
assert.equal(getContrastRating("3.0", true), "AA");
```

`getContrastRating("2.7", false)` dovrebbe restituire `"Fail"`.

```js
assert.equal(getContrastRating("2.7", false), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(ratio, isLargeText) {

  return ratio;
}
```

# --solutions--

```js
function getContrastRating(ratio, isLargeText) {
  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
