---
id: 6a22d77ddf034bc4e35b1d57
title: "Sfida 350: distanza tra lettere"
challengeType: 28
dashedName: challenge-350
---

# --description--

Date due stringhe di uguale lunghezza, restituisci la somma delle distanze più brevi tra ogni coppia di caratteri.

- L'input conterrà solo lettere minuscole
- L'alfabeto è trattato come un cerchio, quindi la distanza tra `a` e `z` è 1.

# --hints--

`letterDistance("abc", "bcd")` dovrebbe restituire `3`.

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` dovrebbe restituire `9`.

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` dovrebbe restituire `10`.

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` dovrebbe restituire `43`.

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` dovrebbe restituire `47`.

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` dovrebbe restituire `55`.

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
