---
id: 6a22d77ddf034bc4e35b1d54
title: "Sfida 347: teoria dei giochi"
challengeType: 28
dashedName: challenge-347
---

# --description--

Date due stringhe di uguale lunghezza che rappresentano le strategie di due giocatori in un gioco, restituisci i punteggi come un array `[player1, player2]`.

- Le stringhe date conterranno solo una di due lettere: `"C"` (cooperare) o `"D"` (tradire).
- Ogni carattere rappresenta un turno, valutato come segue:
  - Se entrambi i giocatori cooperano, ciascuno ottiene 3 punti.
  - Se entrambi i giocatori tradiscono, ciascuno ottiene 1 punto.
  - Se un giocatore tradisce e l'altro coopera, il traditore ottiene 5 punti e il cooperante 0.

# --hints--

`playGame("CCCC", "CCCC")` dovrebbe restituire `[12, 12]`.

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` dovrebbe restituire `[4, 4]`.

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` dovrebbe restituire `[5, 10]`.

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` dovrebbe restituire `[24, 34]`.

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` dovrebbe restituire `[66, 21]`.

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
