---
id: 6a26df3e2988bcdded204897
title: "Sfida 361: tempo parlato"
challengeType: 28
dashedName: challenge-361
---

# --description--

Dato l'angolo delle lancette delle ore e dei minuti di un orologio analogico in gradi (in senso orario da 12), restituisci l'ora in inglese parlato.

Converti l'angolo della lancetta dei minuti in minuti (360° = 60 minuti), poi usa le seguenti regole:

| Minuti | Parlato |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (escluso 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (escluso 45) | `"X minutes to Z"` (dove `X` è 60 - minuti) |

Dove `Y` è l'ora corrente e `Z` è l'ora successiva, entrambe ricavate dall'angolo della lancetta delle ore (360° = 12 ore).

Nota: gli angoli delle lancette potrebbero non corrispondere esattamente a un numero, considera di arrotondarli in qualche modo.

# --hints--

`getSpokenTime(90, 0)` dovrebbe restituire `"3 o'clock"`.

```js
assert.equal(getSpokenTime(90, 0), "3 o'clock");
```

`getSpokenTime(160, 120)` dovrebbe restituire `"20 minutes past 5"`.

```js
assert.equal(getSpokenTime(160, 120), "20 minutes past 5");
```

`getSpokenTime(255, 180)` dovrebbe restituire `"half past 8"`.

```js
assert.equal(getSpokenTime(255, 180), "half past 8");
```

`getSpokenTime(67.5, 92)` dovrebbe restituire `"quarter past 2"`.

```js
assert.equal(getSpokenTime(67.5, 92), "quarter past 2");
```

`getSpokenTime(200, 240)` dovrebbe restituire `"20 minutes to 7"`.

```js
assert.equal(getSpokenTime(200, 240), "20 minutes to 7");
```

`getSpokenTime(322.5, 273)` dovrebbe restituire `"quarter to 11"`.

```js
assert.equal(getSpokenTime(322.5, 273), "quarter to 11");
```

`getSpokenTime(117.5, 335)` dovrebbe restituire `"5 minutes to 4"`.

```js
assert.equal(getSpokenTime(117.5, 335), "5 minutes to 4");
```

# --seed--

## --seed-contents--

```js
function getSpokenTime(hourAngle, minuteAngle) {

  return hourAngle;
}
```

# --solutions--

```js
function getSpokenTime(hourAngle, minuteAngle) {
  const minutes = Math.floor(minuteAngle / 6);
  const hour = Math.floor(hourAngle / 30) || 12;
  const nextHour = (hour % 12) + 1;

  if (minutes === 0) return `${hour} o'clock`;
  if (minutes === 15) return `quarter past ${hour}`;
  if (minutes === 30) return `half past ${hour}`;
  if (minutes === 45) return `quarter to ${nextHour}`;
  if (minutes < 30) return `${minutes} minutes past ${hour}`;
  return `${60 - minutes} minutes to ${nextHour}`;
}
```
