---
id: 6a26df95efa55a2524399745
title: "Sfida 364: tra due secchi"
challengeType: 28
dashedName: challenge-364
---

# --description--

Dati due secchi di vernice, ciascuno con un colore RGB e un livello di pienezza, restituisci il colore RGB miscelato come un array di tre numeri interi.

- Ogni secchio è un oggetto (JavaScript) o dizionario (Python) con una proprietà `color` (un array di tre numeri interi `[r, g, b]`) e una proprietà `fullness` (0–100).
- Il colore miscelato è una media ponderata di ogni canale nei due colori basata sul livello di pienezza, con ogni canale arrotondato al numero intero più vicino.

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })` dovrebbe restituire `[125, 125, 125]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })` dovrebbe restituire `[200, 200, 200]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })` dovrebbe restituire `[100, 150, 200]`.

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })` dovrebbe restituire `[114, 184, 170]`.

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })` dovrebbe restituire `[68, 163, 123]`.

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
