---
id: 6a26df3e2988bcdded204895
title: "Sfida 359: calcolatore di handicap nel golf"
challengeType: 28
dashedName: challenge-359
---

# --description--

Dato un array di punteggi di golf e un array corrispondente di valori par del campo, restituisci l'indice di handicap del golfista usando il seguente metodo:

- Calcola il differenziale per ogni turno sottraendo il par dal punteggio, poi restituisci la media di tutti i differenziali arrotondata a un decimale.

# --hints--

`calculateHandicap([72, 72, 72], [72, 72, 72])` dovrebbe restituire `0`.

```js
assert.equal(calculateHandicap([72, 72, 72], [72, 72, 72]), 0);
```

`calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72])` dovrebbe restituire `6`.

```js
assert.equal(calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72]), 6);
```

`calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36])` dovrebbe restituire `8.3`.

```js
assert.equal(calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3);
```

`calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` dovrebbe restituire `8.8`.

```js
assert.equal(calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8);
```

`calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` dovrebbe restituire `11.7`.

```js
assert.equal(calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7);
```

# --seed--

## --seed-contents--

```js
function calculateHandicap(scores, pars) {

  return scores;
}
```

# --solutions--

```js
function calculateHandicap(scores, pars) {
  const differentials = scores.map((score, i) => score - pars[i]);
  const average = differentials.reduce((sum, d) => sum + d, 0) / differentials.length;
  return Math.round(average * 10) / 10;
}
```
