---
id: 6a26df3e2988bcdded204893
title: "Sfida 357: catena alimentare"
challengeType: 28
dashedName: challenge-357
---

# --description--

Dato un array di coppie `[predator, prey]`, restituisci la catena alimentare dall'apice del predatore fino al fondo.

- L'apice del predatore è l'animale che non è mai preda di un altro animale.
- Restituisci la catena come un array di stringhe.

# --hints--

`getFoodChain([["cat", "mouse"]])` dovrebbe restituire `["cat", "mouse"]`.

```js
assert.deepEqual(getFoodChain([["cat", "mouse"]]), ["cat", "mouse"]);
```

`getFoodChain([["wolf", "deer"], ["deer", "grass"]])` dovrebbe restituire `["wolf", "deer", "grass"]`.

```js
assert.deepEqual(getFoodChain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"]);
```

`getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` dovrebbe restituire `["hawk", "snake", "frog", "fly"]`.

```js
assert.deepEqual(getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"]);
```

`getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` dovrebbe restituire `["eagle", "fox", "rabbit", "grass"]`.

```js
assert.deepEqual(getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"]);
```

`getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` dovrebbe restituire `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`.

```js
assert.deepEqual(getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"]);
```

# --seed--

## --seed-contents--

```js
function getFoodChain(pairs) {

  return pairs;
}
```

# --solutions--

```js
function getFoodChain(pairs) {
  const prey = new Set(pairs.map(([, p]) => p));
  const map = Object.fromEntries(pairs);
  const apex = pairs.map(([p]) => p).find(p => !prey.has(p));
  const chain = [apex];
  while (map[chain.at(-1)]) chain.push(map[chain.at(-1)]);
  return chain;
}
```
