---
id: 641da5462576784453146ec2
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Agora que você tem os dados que quer, você pode usá-los para preencher a interface do usuário. 

Mas os dados que você está buscando terão um total de 26 autores.Em vez de exibir todos os 26 autores de uma vez, você deve exibir 8 autores por vez.

Comece usando a palavra-chave `let` para criar duas variáveis chamadas `startingIndex` e `endingIndex` e atribua a elas os números `0` e `8`, respectivamente.Você precisará usar `let` aqui porque você irá reatribuir esses valores mais tarde.

Então, crie uma variável `authorDataArr` e atribua a ela um array vazio. 

Declare e inicialize cada variável em uma linha separada.

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

Você deve usar `let` para declarar uma variável chamada `startingIndex`.

```js
assert.match(code, /let\s+startingIndex/);
```

Sua variável `startingIndex` deve ser um número.

```js
assert.isNumber(startingIndex);
```

Sua variável `startingIndex` deve ser definida como `0`.

```js
assert.strictEqual(startingIndex, 0);
```

Você deve usar `let` para declarar uma variável chamada `endingIndex`.

```js
assert.match(code, /let\s+endingIndex/);
```

Sua variável `endingIndex` deve ser um número.
  
  ```js
assert.isNumber(endingIndex);
```

Você deve definir a variável `endingIndex` como `8`.

```js
assert.strictEqual(endingIndex, 8)
```

Você deve ter uma variável chamada `authorDataArr`.

```js
assert.isDefined(authorDataArr);
```

Sua variável `authorDataArr` deve ser um array vazio.

```js
assert.isArray(authorDataArr);
assert.isEmpty(authorDataArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
