---
id: 641d9a19bff38d34d5a5edb8
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você aprenderá como trabalhar com a Fetch API para carregar dados de uma fonte externa e exibi-los em uma página preenchida com autores do freeCodeCamp.

Todo o HTML e CSS para este workshop foi fornecido para você.Você pode dar uma olhada nos dois arquivos para se familiarizar com eles.

Quando estiver pronto, comece acessando os elementos `#author-container` e `#load-more-btn`.Lembre-se que você pode usar os métodos `getElementById` ou `querySelector` para isso.

Então atribua esses elementos às variáveis chamadas `authorContainer` e `loadMoreBtn`, respectivamente.

# --hints--

Você deve ter uma variável chamada `authorContainer`.

```js
assert.isDefined(authorContainer);
```

Você deve atribuir o elemento `#author-container` à variável `authorContainer`.

```js
assert.strictEqual(authorContainer, document.querySelector('#author-container'));
```

Você deve ter uma variável chamada `loadMoreBtn`.

```js
assert.isDefined(loadMoreBtn);
```

Você deve atribuir o elemento `#load-more-btn` à variável `loadMoreBtn`.


```js
assert.strictEqual(loadMoreBtn, document.querySelector('#load-more-btn'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
