---
id: a202eed8fc186c8434cb6d61
title: Construye una estructura String Inverter
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

En este laboratorio, construirás un inversor de cadenas simple que invierte los caracteres de una cadena dada.

Por ejemplo, `"hello"` debe convertirse `"olleh"`.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `reverseString` que tome una cadena como argumento.
2. La función debe devolver la cadena invertida.

# --hints--

Debes tener una función llamada `reverseString`.

```js
assert.isFunction(reverseString);
```

`reverseString` debe tomar una cadena como argumento.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` debe devolver una cadena.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` debe devolver la cadena `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` debe devolver la cadena `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` debe devolver la cadena `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
