---
id: 669e2f60e83c011754f711f9
title: Створіть сторінку туристичного агентства
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні мати оголошення `DOCTYPE`.
1. Ви повинні мати `html` елемент з атрибутом `lang`, встановленим у `en`.
1. Ви повинні мати `head` елемент, що містить порожній `meta` елемент з атрибутом `charset`, встановленим у `utf-8`, та `title` з текстом `Travel Agency Page`.
1. Ви повинні мати `meta` тег у вашому `head` елементі, який містить короткий опис вашого вебсайту для SEO.
1. Ви повинні мати `h1` елемент для представлення ваших туристичних напрямків. Це має бути єдиний `h1` елемент на вашій сторінці.
1. Ви повинні мати абзац під `h1` елементом, який вводить туристичні можливості.
1. Ви повинні мати `h2` елемент з текстом `Packages`.
1. Ви повинні мати `p` елемент, який коротко представляє різні пакети.
1. Ви повинні мати невпорядкований список з двома елементами списку. Два елементи списку повинні містити текст `Group Travels` та `Private Tours` відповідно. Текст кожного елемента списку повинен бути укладений в елемент посилання.
1. Ви повинні мати `h2` елемент з текстом `Top Itineraries`.
1. Ви повинні мати принаймні три `figure` елементи, кожен з яких містить елемент посилання та `figcaption` елемент, що описує зображення фігури.
1. Три елементи посилання повинні містити `img` елемент з відповідним атрибутом `alt` та атрибутом `src`, встановленим на дійсне зображення. Ви можете використати `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` та `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg`, якщо хочете.
1. Всі п’ять ваших елементів посилання повинні мати атрибут `href` зі значенням `https://www.freecodecamp.org/learn` та атрибут `target` зі значенням `_blank`.

# --hints--

Ваша сторінка туристичного агентства повинна мати оголошення `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Ви повинні мати елемент `html` з атрибутом `lang`, встановленим у `en`.

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

Ваш елемент `html` має мати кінцевий тег.

```js
assert.match(code, /<\/html\s*>/i);
```

Ви повинні мати елемент `head` всередині елемента `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

Ваш `head` елемент повинен мати кінцевий тег.

```js
assert.match(code, /<\/head\s*>/i);
```

Ви повинні мати `title` елемент всередині вашого `head` елемента.

```js
assert.isNotNull(document.querySelector('head title'));
```

Ваш `title` елемент повинен мати кінцевий тег.

```js
assert.match(code, /<\/title\s*>/i);
```

Ви повинні мати два елементи `meta` всередині вашого елемента `head`.

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

Один елемент `meta` повинен мати атрибут `name` зі значенням `description` та непорожній атрибут `content`.

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

Один `meta` елемент повинен мати атрибут `charset`, встановлений у `utf-8`.

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Ваш елемент `title` має містити назву вашого туристичного агентства.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Ви повинні мати елемент `body` всередині вашого елемента `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

Ваш `body` елемент повинен мати кінцевий тег.

```js
assert.match(code, /<\/body\s*>/i);
```

Ви повинні мати елемент `h1` для презентації ваших туристичних напрямків.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

У вас повинен бути лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Під елементом `h1` має бути елемент `p`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Ваш перший елемент `p` має коротко ознайомити з туристичними можливостями.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Ваш перший елемент `h2` має містити текст `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

Під вашим першим елементом `h2` має бути елемент `p`.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Ваш другий елемент `p` має коротко представити різні пакети послуг.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Під другим елементом `p` має бути елемент невпорядкованого списку.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

Ваш невпорядкований список має містити два елементи.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Обидва елементи списку мають містити елемент посилання.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

Елемент посилання першого елемента списку має охоплювати текст `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

Елемент посилання другого елемента списку має охоплювати текст `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

Після невпорядкованого списку має бути елемент `h2`.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Ваш другий елемент `h2` має містити текст `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

Ви повинні мати щонайменше три елементи `figure`.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Кожен елемент `figure` має містити елемент посилання як свій перший вкладений елемент.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Кожен елемент `figure` має містити елемент `figcaption` як другий вкладений елемент.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Кожен `figcaption` має містити текст.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Кожен з елементів `a`, вкладених у ваші елементи `figure`, має містити зображення.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Кожен елемент `img` має мати дійсний атрибут `src`.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Кожен елемент `img` має мати атрибут `alt` з відповідним значенням.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Кожен елемент `a` має мати атрибут `href` зі значенням `https://www.freecodecamp.org/learn`. Не забудьте про посилання в елементах списку.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Кожен елемент `a` має мати атрибут `target` зі значенням `_blank`. Не забудьте про посилання в елементах списку.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
