---
id: 686b9720ee1d032bd77a480a
title: Depurar um Validador de ISBN
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

O **ISBN (Número Padrão Internacional de Livro)** é um identificador único atribuído a livros comerciais. Pode ter **10** ou **13** dígitos, e o último dígito é um dígito de verificação calculado a partir dos outros dígitos.

Camperbot tentou construir seu próprio validador de ISBN. Entretanto, ele cometeu alguns erros ao longo do caminho.

Neste laboratório, você corrigirá o código existente e fará com que ele funcione corretamente.

Comportamento esperado:  
Quando o usuário executar o programa, ele exibirá o prompt `Digite o ISBN e o comprimento: `.</dfn> O usuário pode inserir o código ISBN que deseja validar no formato `ISBN,length`. O código ISBN não deve conter hífens, seguido pelo seu comprimento (`10` ou `13`), separados por vírgula.

Exemplos de entradas:  
`1530051126,10` para códigos ISBN-10.  
`9781530051120,13` para códigos ISBN-13.

Como encontrar o dígito de verificação:

- Você não precisa conhecer a lógica detalhada de cálculo neste laboratório. As funções `calculate_check_digit_10` e `calculate_check_digit_13` cuidarão do cálculo e retornarão o dígito verificador esperado em formato de string.
- O dígito verificador para códigos ISBN-10 pode ser um número de `0` a `9` ou uma letra maiúscula `X`.
- O dígito verificador para códigos ISBN-13 pode ser um número de `0` a `9`.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve corrigir o `IndentationError` no código atual.
1. Mesmo que o usuário não insira um valor separado por vírgulas, o programa deve tratar o `IndexError` sem travar.
1. Quando o usuário não inserir um valor separado por vírgulas, ele deve ver a mensagem `Digite valores separados por vírgulas.` no console, e o programa deve ser encerrado.
1. Mesmo que o usuário insira um valor não numérico para o comprimento, o programa deve tratar o `ValueError` sem travar.
1. Quando o usuário inserir um valor não numérico para o comprimento, ele deve ver a mensagem `O comprimento deve ser um número.` no console, e o programa deve ser encerrado.
1. Você deve corrigir o erro de deslocamento de um na função `validate_isbn`.
1. Você deve corrigir o `TypeError` no código atual que ocorre quando o usuário insere um código ISBN válido.
1. Você deve corrigir o `IndexError` no código atual quando o usuário inserir um código ISBN válido.
1. Mesmo que o usuário insira um código ISBN incorreto com caracteres diferentes de números, o programa deve tratar o `ValueError` que ocorre sem travar.
1. Quando o usuário inserir um código ISBN incorreto com caracteres diferentes de números, ele deverá ver a mensagem `Caractere inválido encontrado.` no console.
1. Quando o usuário inserir `1530051126,10`, ele deverá ver a mensagem `Código ISBN Válido.`
1. Quando o usuário inserir `9781530051120,13`, ele deverá ver a mensagem `Código ISBN válido.`

**Importante**: você precisará comentar a chamada `main()` no espaço global para que os testes sejam executados corretamente.

Quando você concluir o projeto, o usuário deverá ver as seguintes mensagens dependendo dos valores que ele inserir.

| ISBN Code | Length | Message | Example input |
|-----------|--------|---------| ------------- |
| Valid | Valid | `Valid ISBN Code.` | `1530051126,10` |
| Invalid Number | Valid | `Invalid ISBN Code.` | `1530051125,10` |
| Does not match specified length or left blank | Valid | `ISBN-10 code should be 10 digits long.` or `ISBN-13 code should be 13 digits long.`, depending on the length they entered. | `9781530051120,10` or `1530051126,13` |
| Contains non-numeric characters (except for the check digit) | Valid | `Invalid character was found.` | `15-0051126,10` |
| Any | Invalid Number | `Length should be 10 or 13.` | `1530051126,9` |
| Any | Contains non-numeric characters or left blank | `Length must be a number.` | `1530051125,A` |
| Not comma-separated | Not comma-separated | `Enter comma-separated values.` | `1530051125` |

Você pode usar os seguintes valores para testar seu código manualmente, se desejar.

Exemplos de entradas para códigos ISBN-10 válidos:  
`1530051126,10`  
`9971502100,10`  
`080442957X,10`

Exemplos de entradas para códigos ISBN-13 válidos:  
`9781530051120,13`  
`9781947172104,13`

# --hints--

Você deve comentar a chamada para a função `main` para permitir que o restante dos testes funcione corretamente.

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

Você deve ter uma função `validate_isbn`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

Você deve ter uma função `calculate_check_digit_10`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

Você deve ter uma função `calculate_check_digit_13`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

Você deve ter uma função `main`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

Quando o usuário insere um valor que não é um valor separado por vírgulas, você deve ver a mensagem `Digite valores separados por vírgulas.` no console.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Quando o usuário inserir um valor não numérico para o comprimento, você deverá ver a mensagem `O comprimento deve ser um número.` no console.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Quando o usuário inserir um código ISBN incorreto com caracteres diferentes de números, você deve ver a mensagem `Caractere inválido foi encontrado.` no console.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Quando o usuário inserir `1530051126,10`, você deverá ver a mensagem `Código ISBN Válido.` no console.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Quando o usuário digitar `9781530051120,13`, você deve ver a mensagem `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Quando o usuário digita `1530051125,10`, você deve ver a mensagem `Código ISBN Inválido.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

Quando o usuário digitar `9781530051120,10`, você deve ver a mensagem `ISBN-10 code should be 10 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

Quando o usuário digitar `1530051126,13`, você deve ver a mensagem `ISBN-13 code should be 13 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

Quando o usuário digitar `15-0051126,10`, você deve ver a mensagem `Invalid character was found.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Quando o usuário digitar `1530051126,9`, você deve ver a mensagem `Length should be 10 or 13.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

Quando o usuário digitar `1530051125,A`, você deve ver a mensagem `Length must be a number.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Quando o usuário digitar `1530051125`, você deve ver a mensagem `Enter comma-separated values.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Quando o usuário digitar `9971502100,10`, você deve ver a mensagem `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Quando o usuário digitar `080442957X,10`, você deve ver a mensagem `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Quando o usuário digitar `9781947172104,13`, você deve ver a mensagem `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
