---
id: 65f569725359e10d345bc52a
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

Vas a construir un programa que calcula y dibuja la trayectoria de un proyectil dado el ángulo, la velocidad y la altura del lanzamiento.

Comienza importando `math`, lo usarás mucho en este proyecto ya que tiene métodos útiles como `math.radians`, `math.cos`, `math.sin` y otros.

También crea estas variables para que tengan el valor de la aceleración gravitacional y algunos símbolos especiales que serán útiles más adelante (usa copiar y pegar para estos).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

Usa `import math` para importar el módulo `math`.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

Debes tener una variable `GRAVITATIONAL_ACCELERATION` con un valor de `9.81`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

Debes tener una variable `PROJECTILE` con valor `"∙"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

Debes tener una variable `x_axis_tick` con el valor de `"T"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

Debes tener una variable `y_axis_tick` con el valor de `"⊣"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
