---
id: 65fd4b3bedee044b4b957d33
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Crea una clase `Projectile` con un método `__init__` que inicialice la clase usando tres argumentos: la velocidad inicial, la altura inicial y el ángulo inicial del lanzamiento del proyectil, en este orden. 

Dentro del método `__init__`, asigna estos argumentos a tres atributos privados: `__speed`, `__height` y `__angle`. El ángulo proporcionado estará en grados; sin embargo, debe almacenarse internamente en radianes. Para lograr esto, usa la función `math.radians()` para convertir el ángulo de grados a radianes al asignarlo a `__angle`. 

El uso de dos guiones bajos antes del nombre de un atributo activa el name mangling en Python. Esto significa que los atributos no son accesibles directamente desde fuera de la clase usando sus nombres dados, y deben ser accedidos con los nombres mangled como `ball._Projectile__height` si se necesitan externamente, indicando que están destinados solo para uso interno.

# --hints--

Debes usar `class Projectile:` para crear la clase.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

Debes tener un método `__init__` dentro de la clase `Projectile`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

Una nueva instancia de la clase `Projectile` debe tener los valores de los tres atributos correctamente inicializados.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
