---
id: 65fd4bd84561a14d3e720061
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

La variable de clase `__slots__` tiene un uso especial en las clases de Python. Declarar `__slots__` y asignarle una secuencia de cadenas restringe la creación de atributos a los incluidos en esa secuencia. Además, evita la creación del atributo especial `__dict__` y permite un acceso más eficiente a los atributos.

Debes usar la variable `__slots__` dentro de la clase para definir qué atributos tiene la clase: asigna a `__slots__` una tupla que contenga 3 cadenas, cada una igual a uno de los nombres de los atributos definidos en el `__init__`.

# --hints--

Debes asignar a `__slots__` una tupla que contenga `'__height'`, `'__speed'` y `'__angle'`.

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
