---
id: 660400f02cf3a76795b82d4e
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Lo primero que debes configurar es un método que calcule el desplazamiento del proyectil, que es el espacio horizontal recorrido desde el lanzamiento hasta que el proyectil toca el suelo.

Crea un método `__calculate_displacement`, que tenga solo `self` como parámetro, y devuelve el desplazamiento del proyectil.

Usa la siguiente fórmula para calcular el desplazamiento del proyectil:
\[ d = \frac{v \cdot \cos\theta \cdot \left(v\cdot\sin\theta + \sqrt{v^2 \cdot \sin^2\theta + 2 \cdot g \cdot h}\right)}{g} \]

En donde $d$ es el desplazamiento, $v$ es la velocidad inicial, $\theta$ es el ángulo y $h$ es la altura inicial del proyectil.  
Para $g$ puedes usar la variable `GRAVITATIONAL_ACCELERATION`.

Debes usar los métodos `math.cos()` y `math.sin()` para las funciones trigonométricas y `math.sqrt()` para calcular la raíz cuadrada. También debes saber que $x^y$ se escribe como `x ** y` en python. Además, $\sin^2\theta$ significa que el valor resultante del seno se eleva al cuadrado.

Recuerda que con name mangling necesitas llamar al método como `_Projectile__calculate_displacement` si quieres probarlo, o usarlo desde fuera de la clase:

```py
ball = Projectile(10, 3, 45)
displacement_of_ball = ball._Projectile__calculate_displacement() # 12.6173996009878
```

# --hints--

Debes declarar un método llamado `__calculate_displacement` con `def __calculate_displacement(self):`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").has_function("__calculate_displacement")`))})
```

El método `__calculate_displacement` debe tener solo el argumento `self`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_args('self')`))})
```

El método `__calculate_displacement` debería devolver el valor correcto.

```js
({test: () => assert(runPython(`
p = Projectile(20, 21, 22)
disp = p._Projectile__calculate_displacement()
round(disp, 2) == 55.06 and round(disp, 2) != disp
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    
--fcc-editable-region--
```
