---
id: 660fcf3e1b9bb056b2edb567
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

En este punto estás listo para crear la representación en cadena.  
Comienza creando una instancia de la clase `Projectile`. Debajo de la definición de la clase, crea una variable `ball` y asígnale una llamada a `Projectile` usando `10, 3, 45` como argumentos.

Define un método `__str__` dentro de la clase, para darle a tu instancia una representación en cadena adecuada, y haz que retorne una cadena con el siguiente formato:

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

\nDebe comenzar y terminar con un salto de línea, el ángulo debe escribirse como un número entero en grados y el desplazamiento debe imprimirse con un dígito decimal.\n

Encontrarás útil `math.degrees` para convertir el ángulo de radianes a grados.

Cuando estés listo puedes `print(ball)` para probar tu función.

# --hints--

Debes tener una variable `ball`.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

La variable `ball` debe tener un valor de `Projectile(10, 3, 45)`.

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

La representación en cadena para `Projectile(45, 45, 45)` debe ser correcta.

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

La representación en cadena también debe ser correcta para otras instancias.

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

Debes tener una llamada a `print(ball)`.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
