---
id: 6616ec0bf9ee8548ce6b0f08
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

¡Ahora es momento de trabajar en calcular las coordenadas de la trayectoria! Crea un método de la clase `Projectile` llamado `__calculate_y_coordinate`, que debe tener, además de `self`, un parámetro `x`.

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

La fórmula anterior es para calcular la posición vertical $y$ para cualquier posición horizontal $x$, teniendo el ángulo inicial $\theta$, la velocidad $v_0$ y la altura $y_0$.

Necesitarás usar `math.tan()` y `math.cos()` y recordar que `x ** y` es la forma de escribir $x^y$, y que el valor de $g$ está en la variable `GRAVITATIONAL_ACCELERATION`.

Implementa el método para que retorne la coordenada $y$.

Si quieres probar el método, desde fuera de la clase, puedes llamar a `ball._Projectile__calculate_y_coordinate()` con un número como argumento después de que se declare `ball`.

# --hints--

Tu clase `Projectile` debe tener un método llamado `__calculate_y_coordinate`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

El método `__calculate_y_coordinate` debe tener `self, x` como argumentos.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

El método debe devolver el valor correcto. Para `Projectile(12, 12, 12)`, dadas las 4 coordenadas $x$ de 0, 1, 2 y 3, la coordenada y debe ser aproximadamente 12, 12.18, 12.28, 12.32.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

El método debe devolver el valor correcto. Para `Projectile(45, 12, 22)`, dadas las coordenadas $x$ de 6, 9, 10 y 12, la coordenada y debe ser aproximadamente 14.32, 15.41, 15.76, 16.44.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

El método `__calculate_y_coordinate` no debe redondear el resultado.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
