---
id: 699f3cb8d057f29954205d7d
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

如你從之前的課程所知，你可以像這樣為函式的參數添加型別註解：

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

在這個範例中，如果你嘗試用非數字的引數呼叫 `add` 函式，TypeScript 會顯示錯誤。

請更新你的 `square` 函式，為 `num` 參數加入型別註解。該型別應該是 `number`。

# --hints--

你應該為 `num` 函式中的 `square` 參數添加型別註解。該型別應該是 `number`。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
