---
id: 699f5063d2393c674b1f45dc
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

如你從之前的步驟所知，如果你用錯誤數量的引數呼叫函式，TypeScript 會丟出錯誤。

有一種方法可以將參數標記為選用，只要在參數名稱後面加上 `?`。這樣你就可以用比參數數量更少的引數來呼叫函式。

以下是一個範例：

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

在這個範例中，如果呼叫 `greeting` 函式時沒有提供 `greet` 參數，它會是 `undefined`，函式將傳回一個預設的問候語。

將你的 `raiseTo` 函式更新，讓 `exponent` 參數變成選用。

# --hints--

你的 `raiseTo` 函式應該將 `exponent` 參數標記為選用。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
