---
id: 699f52220de88c8091032076
title: 第 18 步
challengeType: 1
dashedName: step-18
---

# --description--

如果你打開主控台，結果是 `NaN`，因為 `exponent` 的第二個引數從未被提供。這個 `raiseTo` 函式需要幾個更新來修正這個問題。

第一個改動是要傳回 `base ** 2`。

在下一步，你將看到如何控制選用的 `exponent` 參數。

# --hints--

你的 `raiseTo` 函式應該傳回一個數字。

```js
assert.isNumber(raiseTo(4));
```

你的 `raiseTo` 函式應該傳回 `base` 的 `2` 次方的值。

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
