---
id: 699f412114c87221f16a99e3
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

使用 TypeScript 的另一個利益是，如果你在函式呼叫中傳入太多或太少的引數，它會捕捉到錯誤。

為了說明這點，請更新你的 `square` 函式呼叫，使其接受兩個引數 `5` 和 `10`。

# --hints--

你的 `square` 函式呼叫應該有兩個引數，分別是 `5` 和 `10`。

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
