---
id: 699f49669ea51a31d2ba89d8
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

如你從之前的課程所知，你可以像這樣明確指定函式的返回型別：

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

在這個範例中，`circleArea` 函式的返回型別是 `number`。如果你嘗試返回不同的型別，TypeScript 會顯示錯誤。

將你的 `getAverage` 函式更新為明確指定返回型別為 `number`。

# --hints--

你的 `getAverage` 函式應該有 `number` 的返回型別。

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
