---
id: 699f38937b053e4679878ff1
title: 第 3 步驟
challengeType: 1
dashedName: step-3
---

# --description--

在使用純 JavaScript 時，沒有任何機制能阻止你用字串或陣列來呼叫 `square` 函式，這會導致意料之外的結果。

為了說明這點，請將你的函式呼叫改成使用字串 `"something"`，而非數字 `5`。

在下一課中，你將學習如何使用 TypeScript 為你的程式碼添加型別安全，避免這類錯誤發生。

# --hints--

你的 `square` 函式呼叫應該帶有 `"something"` 作為引數。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
