---
id: 699f434e294a9e11cf4f65f0
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

現在是時候為你的套件專案添加另一個數學輔助函式了。

先從建立一個名為 `getAverage` 的函式開始，它接受一個參數。你的函式應該傳回一組數字陣列的平均值。

在接下來的幾個步驟中，你將為這個函式添加型別註解。

# --hints--

你應該有一個名為 `getAverage` 的函式。

```js
assert.isFunction(getAverage);
```

你的函式應該接受一個參數。

```js
assert.strictEqual(getAverage.length, 1);
```

你的函式應該傳回一個數字。

```js
assert.isNumber(getAverage([1, 2, 3]));
```

你的函式應該對任何數字陣列傳回正確結果。

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
