---
id: 699f4639ae0016caaf81b7c6
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

現在是時候測試你的函式（程式）了。

先建立一個名為 `avgResult` 的變數，並將 `getAverage` 的函式呼叫結果指定給它。該函式呼叫的引數應該是 `[2, 14, 26, 8]`。

接著將 `avgResult` 的值輸出到主控台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該有一個名為 `avgResult` 的變數。

```js
assert.isDefined(avgResult);
```

你應該將 `getAverage([2, 14, 26, 8])` 指定給你的 `avgResult` 變數。

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

你的 `avgResult` 變數應該是數字。

```js
assert.isNumber(avgResult);
```

你應該將 `avgResult` 輸出到主控台。

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
