---
id: 699f3f636aa840b7201aa37c
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

如果你打開主控台，應該會看到以下 TypeScript 的錯誤訊息：

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

現在參數已明確指定為 number 型別，TypeScript 就能偵測出嘗試傳入字串引數給 `square` 函式（程式）的錯誤。

將傳入 `square` 函式的引數改為數字 `5`，即可修正錯誤並在主控台看到正確輸出。

# --hints--

你的 `square` 函式呼叫應該傳入 `5` 作為引數，而不是 `"something"`。

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
