---
id: 699f52fac083214b6bf51de6
title: 第 19 步驟
challengeType: 1
dashedName: step-19
---

# --description--

在最後這一步，你需要為 `exponent` 函式中的 `raiseTo` 參數新增條件判斷。

如果提供了 `exponent` 引數，該函式應該傳回 `base` 的 `exponent` 次方。如果沒有提供 `exponent` 引數，則應該傳回 `base ** 2`。

完成這最後一步後，你的型別安全數學套件就完成了！

# --hints--

你的 `raiseTo` 函式在提供 `base ** exponent` 引數時，應該傳回 `exponent`。

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

你的 `raiseTo` 函式在未提供 `base ** 2` 引數時，應該傳回 `exponent`。

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
