---
id: 699f4861dd5b5f164e1bda4a
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

使用 TypeScript 的另一個利益是確保返回型別被強制執行，並且符合你的預期。有時 TypeScript 可以推斷函式的返回型別，但你也可以明確指定它。

為了說明這點，將 `getAverage` 函式的返回型別改為字串 `"something"`。

# --hints--

你的 `getAverage` 函式應該傳回一個字串。

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

你的 `getAverage` 函式應該傳回字串 `"something"`。

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
