---
id: 699f513e01eb1768cc5df726
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

現在 `exponent` 參數已標記為選用，你可以只用一個引數來呼叫這個函式。

更新你的 `raiseTo` 函式呼叫，只傳入數字 `3`。

# --hints--

你的 `raiseTo` 函式應該只用一個 `3` 引數來呼叫。

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
