---
id: 699f4b903bf5427a599f791d
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

在接下來的幾個步驟中，你將建置另一個函式，繼續練習函式的型別註記。

先建立一個名為 `raiseTo` 的函式，它接受兩個參數，分別叫做 `base` 和 `exponent`。這兩個參數都應該是 `number` 型別。此函式應該傳回將 `base` 次方於 `exponent` 的結果。函式的返回型別也應該註記為 `number`。

# --hints--

你應該有一個 `raiseTo` 函式。

```js
assert.isFunction(raiseTo);
```

你的 `raiseTo` 函式應該有兩個參數，分別是 `base` 和 `exponent`。

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*=>/
);
```

你的 `base` 和 `exponent` 參數都應該是 `number` 型別。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

你的 `raiseTo` 函式應該傳回一個數字。

```js
assert.isNumber(raiseTo(2, 3));
```

你的 `raiseTo` 函式應該傳回將 `base` 次方於 `exponent` 的結果。

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

你的 `raiseTo` 函式應該有 `number` 的返回型別註記。

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
