---
id: 699f1d72e906e4616155290d
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將繼續學習型別註解以及它們如何與函式（程式）搭配使用，透過建置一個型別安全的數學套件。

首先建立一個名為 `square` 的函式（程式），它接受一個參數並傳回該數字的平方。

# --hints--

你應該有一個名為 `square` 的函式（程式）。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

你的 `square` 函式（程式）應該有一個參數。

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

你的 `square` 函式（程式）應該傳回該數字的平方。

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
