---
id: 699f458e47fb88c756fdba81
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

目前，`numbers` 參數具有隱式的 `any` 型別。要修正這個問題，你可以將該參數明確地指定為數字陣列的型別。

回想之前的課程，你可以這樣為數字陣列指定型別：

```ts
const numbers: number[] = [1, 2, 3];
```

將你的 `getAverage` 函式更新，使 `numbers` 參數明確地指定為數字陣列的型別。

# --hints--

你的 `getAverage` 函式應該有一個名為 `numbers` 的參數，該參數明確地指定為數字陣列的型別。

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
