---
id: 699f3738921fb79a651ded56
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

現在是時候呼叫你的函式了。

創建一個名為 `result` 的變數，並將 `square` 的函式呼叫結果（引數為 `5`）指定給它。

在下面，將 `result` 的值輸出到主控台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該有一個名為 `result` 的變數。

```js
assert.isDefined(result);
```

你的 `result` 變數應該是數字。

```js
assert.isNumber(result);
```

你的 `result` 變數應該是 `5` 的平方。

```js
assert.strictEqual(result, 25);
```

你應該將 `result` 變數輸出到主控台。

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
