---
id: 699f4ec939f10da62fb56c3e
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

現在是時候測試你的函式了。

創建一個名為 `raisedResult` 的變數，並將呼叫 `raiseTo` 並傳入引數 `2` 和 `3` 的結果指定給它。

接著，將 `raisedResult` 輸出到主控台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該有一個名為 `raisedResult` 的變數。

```js
assert.isDefined(raisedResult);
```

你的 `raisedResult` 變數應該是數字型別。

```js
assert.isNumber(raisedResult);
```

你應該將 `raisedResult` 輸出到主控台。

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
