---
id: 699f41b8da20f816b5c7dece
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

如果你打開主控台，應該會看到以下錯誤訊息：

```md
Expected 1 arguments, but got 2.
```

為了解決這個錯誤，請更新你的 `square` 函式呼叫，使其只帶有一個 `5` 引數。

# --hints--

你的 `square` 函式呼叫應該只帶有一個 `5` 引數。

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
