---
id: a3566b1109230028080c9345
title: Algorithmus zum Summieren aller Zahlen entwerfen
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

In diesem Lab müssen Sie einen Algorithmus entwerfen, der alle Zahlen summiert.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `sumAll` haben, die ein Array mit zwei Zahlen akzeptiert.
1. `sumAll([n, m])` sollte die Summe von `n` und `m` sowie die Summe aller Zahlen dazwischen zurückgeben. Die kleinere Zahl steht nicht immer an erster Stelle. Zum Beispiel sollte `sumAll([4,1])` `10` zurückgeben, weil die Summe aller Zahlen zwischen `1` und `4` (beide eingeschlossen) `10` ist.

# --hints--

Sie sollten eine Funktion namens `sumAll` haben.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` sollte eine Zahl zurückgeben.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` sollte `10` zurückgeben.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` sollte `10` zurückgeben.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` sollte `45` zurückgeben.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` sollte `45` zurückgeben.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
