---
id: 612e96fc87fe8e44f69f7ec5
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Час стилізувати клавіші. Нижче правила `#piano` націльтеся на елемент `.keys` за допомогою селектора `class`. Додайте новому правилу властивість `background-color` зі значенням `#040404`, властивість `width` зі значенням `949px` та властивість `height` зі значенням `180px`.

# --hints--

У вас має бути селектор `.keys`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.keys'));
```

Ваш селектор `.keys` має мати властивість `background-color`, встановлену в `#040404`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor, 'rgb(4, 4, 4)');
```

Ваш селектор `.keys` має мати властивість `width`, встановлену в `949px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.keys')?.width, '949px');
```

Ваш селектор `.keys` має мати властивість `height`, встановлену в `180px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.keys')?.height, '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
