---
id: 612e89562043183c86df287c
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Браузери часто застосовують за замовчуванням значення полів і відступів до певних елементів. Щоб ваш макет піаніно відображався правильно, потрібно скинути модель блоку.

Додайте селектор `html` у свій CSS-файл і встановіть властивість `box-sizing` у значення `border-box`.

# --hints--

У вас має бути селектор `html`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Ваш селектор `html` має мати властивість `box-sizing`, встановлену в `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
