---
id: 612e804c54d5e7308d7ebe56
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У межах вашого елемента `.keys` додайте сім елементів `div`, кожен із класом `key`.

# --hints--

Ви маєте створити сім нових елементів `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert.lengthOf(divDivDiv, 9);
```

Ваші сім нових елементів `div` мають бути всередині вашого елемента `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf([...keys?.children], 7);
[...keys?.children].forEach(
  child => assert.equal(child.tagName, 'DIV')
)
```

Усі сім нових елементів `div` повинні мати множину `class`, встановлену в `key`.

```js
const keys = document.querySelector('.keys');
[...keys?.children].forEach(child => assert.isTrue(child?.classList.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
