---
id: 612ea4c4993aba52ab4aa32e
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Тепер час використати псевдоселектори, які ви налаштували раніше. Щоб створити чорні клавіші, додайте новий селектор `.key.black--key::after`. Він націлений на елементи з класами `key` і `black--key` та вибирає псевдоелемент, створений після цих елементів у HTML.

У новому селекторі встановіть `background-color` у `#1d1e22`. Також встановіть властивість `content` у `""`. Це зробить псевдоелементи порожніми.

Властивість `content` використовується для встановлення або заміщення вмісту елемента. За замовчуванням псевдоелементи, створені селекторами `::before` і `::after`, порожні, тобто вони не відтворюються на сторінці. Встановивши властивість `content` у порожній рядок `""`, ви забезпечуєте відтворення псевдоелементів, при цьому вони залишаються порожніми.

Якщо хочете поекспериментувати, спробуйте видалити властивість `background-color` і встановити різні значення для властивості `content`, наприклад `"♥"`. Не забудьте скасувати ці зміни після завершення, щоб тести пройшли успішно.

# --hints--

У вас має бути селектор `.key.black--key::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Ваш селектор `.key.black--key::after` має мати властивість `background-color`, встановлену в `#1d1e22`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor, 'rgb(29, 30, 34)');
```

Ваш селектор `.key.black--key::after` має мати властивість `content`, встановлену в `""`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content, '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
