---
id: 612e89d254fe5d3df7d6693d
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Тепер, коли ви скинули модель блоку html, потрібно застосувати це також до елементів всередині нього. Для цього потрібно встановити властивість `box-sizing` усіх інших елементів як успадковану, що гарантує, що цільові елементи отримають те саме значення, що й їхній батьківський елемент.

Також потрібно націлитися на псевдоелементи — це спеціальні ключові слова, які йдуть після селектора. Два псевдоелементи, які ви будете використовувати, — це псевдоелементи `::before` та `::after`.

Селектор `::before` створює псевдоелемент, який є першим дочірнім елементом вибраного елемента, тоді як селектор `::after` створює псевдоелемент, який є останнім дочірнім елементом вибраного елемента. Ці псевдоелементи часто використовують для створення косметичного вмісту, що ви побачите пізніше в цьому проєкті.

Поки що створіть CSS-селектор, щоб націлитися на всі елементи за допомогою `*`, і включіть псевдоелементи `::before` та `::after`. Встановіть властивість `box-sizing` у значення `inherit`.

# --hints--

Ви повинні мати селектор `html``.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Ваш селектор `html` має мати властивість `box-sizing`, встановлену в `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

Ви повинні мати селектор `*, ::before, ::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Ваш селектор `*, ::before, ::after` має мати властивість `box-sizing`, встановлену в `inherit`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing, 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
