---
id: 612e7d1c29fb872d6384379c
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Вкладіть другий `div` всередину вашого існуючого `div` і встановіть `class` у `keys`.

# --hints--

Ви маєте створити другий елемент `div`.

```js
const divDiv = document.querySelectorAll('div');
assert.lengthOf(divDiv, 2);
```

Ваш новий елемент `div` має бути всередині вашого існуючого елемента `div`.

```js
const div = document.querySelector('div');
assert.lengthOf(div?.children, 1);
assert.equal(div?.children?.[0]?.localName, 'div');
```

Ваш новий елемент `div` має мати встановлений `class` у `keys`.

```js
const div = document.querySelector('div');
assert.equal(div?.children?.[0]?.className, 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
