---
id: 612e8279827a28352ce83a72
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Тепер скопіюйте множину з семи `.key` елементів і вставте ще дві множини у div `.keys`.

# --hints--

У вас має бути 21 `.key` елемент.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

У вас має бути 15 `.black--key` елементів.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

Всі `.key` елементи мають бути всередині вашого `.keys` елемента.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

Всі `.black--key` елементи мають бути всередині вашого `.keys` елемента.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
