---
id: 612e813b3ba67633222cbe54
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Пам’ятайте, що атрибут `class` може містити кілька значень. Щоб розрізняти ваші білі та чорні клавіші, додайте друге значення `class` з `black--key` до другого, третього, п’ятого, шостого та сьомого елементів `.key`.

# --hints--

Ваш другий елемент `.key` також має мати атрибут `class` зі значенням `black--key`.

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

Ваш третій елемент `.key` має мати `black--key` у атрибуті `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

Ваш п’ятий елемент `.key` має мати `black--key` у атрибуті `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

Ваш шостий елемент `.key` має мати `black--key` у атрибуті `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

Ваш сьомий елемент `.key` має мати `black--key` у атрибуті `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

У вас має бути п’ять елементів `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

У вас має бути сім елементів `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
