---
id: 6690be42942f421648b119eb
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

You can add images to your website by using the `img` element. `img` elements have an opening tag without a closing tag. An element without a closing tag is known as a <dfn>void element</dfn>.

Aggiungi un elemento `img` sotto l'elemento `p`. A questo punto, nessuna immagine verrà visualizzata nel browser.

# --hints--

L'elemento `img` dovrebbe avere un tag di apertura. Le etichette di apertura hanno questa sintassi: `<nomeElemento>`.

```js
assert.exists(document.querySelector('img'));
```

Your `img` element should not have a closing tag `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

Dovresti avere un solo elemento `img`. Rimuovi gli altri.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

L'elemento `img` dovrebbe essere al di sotto dell'elemento `p`. Sono nell'ordine sbagliato.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```


