---
id: 66984e8ab54bee517207be42
title: Step 34
challengeType: 0
dashedName: step-34
---

# --description--

Nest a `p` element with the text `No Copyright - freeCodeCamp.org` within the `footer` element.

# --hints--

Hai eliminato l'elemento `footer` oppure manca un tag di apertura o di chiusura.

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

L'elemento `footer` dovrebbe avere un elemento `p`. Assicurati di aggiungere il tag di apertura e di chiusura per l'elemento `p`.

```js
assert.exists(document.querySelector('footer > p'));
```

L'elemento `footer` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere`/` subito dopo il carattere `<`.

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 2);
```

Il testo dell'elemento `p` dovrebbe essere `No Copyright - freeCodeCamp.org`. Hai omesso il testo, hai un refuso o il testo non è tra i tag di apertura e chiusura dell'elemento `p`.

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0].innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    <footer>

    </footer>
--fcc-editable-region--
  </body>
</html>
```


