---
id: 67cd4cf8b5fe391b324ad11f
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Ora che hai creato la variabile di stato `isVisible`, è il momento di iniziare a usarla.

Nelle lezioni precedenti, hai imparato il rendering condizionale in linea che ti permette di mostrare contenuti diversi in base a una certa condizione. È un pattern comune usare l'operatore logico AND (`&&`) per renderizzare condizionalmente un pezzo di testo in questo modo:

```jsx
function Notification({ message }) {
  return (
    <div>
      {message && <p>{message}</p>}
    </div>
  );
}
```

In questo esempio, l'elemento paragrafo verrà mostrato solo se message è truthy. Ricorda che "truthy" indica qualsiasi valore che in un contesto booleano viene valutato come `true`.

Per questo passaggio, dovrai renderizzare condizionalmente l'elemento paragrafo. Se fatto correttamente, vedrai che il paragrafo non sarà più visibile sulla pagina.

# --hints--

Dovresti renderizzare condizionalmente il paragrafo in modo che venga mostrato solo quando `isVisible` è `true`. Riferisciti all'esempio se hai bisogno di aiuto.

```js
  {
    React.useState = (arg) => ([false, () => {}]);
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.ToggleApp);
    const pEl = div.querySelector("p#message");
    assert.notExists(pEl);
  }
  {
    React.useState = (arg) => ([true, () => {}]);
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.ToggleApp);
    const pEl = div.querySelector("p#message");
    assert.exists(pEl);
  }
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(false);

  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      --fcc-editable-region--
      <p id="message">I love freeCodeCamp!</p>
      --fcc-editable-region--
    </div>
  );
};
```
