---
id: 685239f0cad43c7a909cbd98
title: ペットの里親募集ページのデバッグ
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

ペットの里親募集ストアのオーナーであるサリーは、初めてのウェブページを作成しましたが、いくつか問題があります。

サリーがページの作成を続けられるように、すべてのエラーを修正してください。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストを通過させてラボを完了させてください。

**ユーザーストーリー:**

1. サリーは猫の画像を使いたいのですが、正しく表示されていません。`img` 要素の以下を修正してください：
   - `href` 属性を画像のソースを指定する正しい属性に置き換えてください。
   - `att` 属性を画像の短く説明的なテキストを表す正しい属性に置き換えてください。
   - `</img>` の終了タグを削除してください。`img` 要素は空要素で終了タグがありません。
2. サリーはユーザーを犬と猫のページに誘導するリンクを使いたいのですが、リンクが正しく動作していません。`a` 要素の以下を修正してください：
   - 両方の `src` 属性をURLを指定する正しい属性に置き換えてください。

# --hints--

`img` 要素は `src` 属性ではなく `href` 属性を持つべきです。

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

`img` 要素は存在しない `alt` 属性ではなく `att` 属性を持つべきです。

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

`img` 要素は `</img>` の終了タグを持つべきではありません。

```js
assert.notMatch(code, /<\/img>/);
```

テキストが `a` の `Visit cats page` 要素は `href` 属性ではなく `src` 属性を持つべきです。

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

テキストが `a` の `Visit dogs page` 要素は `href` 属性ではなく `src` 属性を持つべきです。

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
