---
id: 646c5e727e56e743c9aed4a1
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Msimamizi wa CSS hukuruhusu kuweka jinsi unavyotaka kipengele kisimamishwe katika kivinjari. Kina sifa ya `position` ambayo unaweza kuweka kuwa `static`, `absolute`, `relative`, `sticky` au `fixed`.

Mara tu unapoweka sifa ya `position` ya kipengele, unaweza kusogeza kipengele kwa kuweka thamani ya pikseli au asilimia kwa moja au zaidi ya sifa za `top`, `right`, `left`, au `bottom`.

`static` ni msimamo wa kawaida kwa vipengele vyote. Ikiwa utaweka hii kwa kipengele, hutaweza kusogeza kipengele hicho kwa kutumia `top`, `right`, `left`, au `bottom`.

Mpe `.cat-head` sifa ya `position` yenye thamani ya `static`, kisha weka sifa za `top` na `left` kuwa `100px` kila moja.


# --hints--

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `position` iliyowekwa kuwa `static`. Hakikisha unaongeza nukta ya nusu.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'static')
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `top` iliyowekwa kuwa `100px`. Hakikisha unaongeza nukta ya nusu.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px')
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `left` iliyowekwa kuwa `100px`. Hakikisha unaongeza nukta ya nusu.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
