---
id: 646c48df8674cf2b91020ecc
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Umepatiwa muundo wa msingi. Anza kwa kuongeza kipengele cha `link` ndani ya kipengele chako cha `head`. Kwa kipengele hicho cha `link`, weka sifa ya kipengele `rel` kuwa `stylesheet` na `href` kuwa `./styles.css`.

# --hints--

Unapaswa kuwa na kipengele cha `link`.

```js
assert.exists(document.querySelector("head>link"));
```

Kipengele chako cha `link` kinapaswa kuwa na sifa za kipengele `rel` na `href` zilizowekwa kuwa `stylesheet` na `./styles.css`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
