---
id: 646ceb843412c74edee27a79
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

Kwa kutumia kichaguzi cha darasa la vitu, mpe kipengele cha `.cat-right-ear` sifa za `height` na `width` zilizo sawa na `100px`. Weka `background-color` kuwa `white`. Weka mipaka ya kushoto na kulia kuwa `35px solid blue`. Weka mipaka ya juu na chini kuwa `35px solid red`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `height` iliyowekwa kuwa `100px`. Usisahau kuongeza nukta ya nusu.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `width` iliyowekwa kuwa `100px`. Usisahau kuongeza nukta ya nusu.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `white`. Usisahau kuongeza nukta ya nusu.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `border-left` iliyowekwa kuwa `35px solid blue`. Usisahau kuongeza nukta ya nusu.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `border-right` iliyowekwa kuwa `35px solid blue`. Usisahau kuongeza nukta ya nusu.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `border-top` iliyowekwa kuwa `35px solid red`. Usisahau kuongeza nukta ya nusu.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `border-bottom` iliyowekwa kuwa `35px solid red`. Usisahau kuongeza nukta ya nusu.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
