---
id: 646c59e942f35541923104bf
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Rudi kwenye faili yako ya HTML, tengeneza kipengele cha `main`. Ndani ya kipengele hicho cha `main`, ongeza kipengele cha `div` chenye darasa la vitu la `cat-head`.

# --hints--

Unapaswa kuwa na kipengele cha `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Unapaswa kuwa na kipengele cha `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

Kipengele chako cha `div` kinapaswa kuwa na darasa la vitu la `cat-head`.

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

Kipengele chako cha `div` kinapaswa kuwa ndani ya lebo yako ya `main`.

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
