---
id: 6476f7a4827bcc61682f2347
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Unapotumia thamani ya `absolute` kwa sifa yako ya `position`, kipengele kinatolewa katika mtiririko wa kawaida wa hati, kisha nafasi yake huamuliwa na sifa za `top`, `right`, `bottom`, na `left`.

Weka sifa ya `position` ya kipengele chako cha `.cat-head` kuwa `absolute`, kisha weka sifa za `top` na `left` kwa thamani yoyote chanya ya pikseli.

<!-- **Kumbuka**: Unaweza kujaribu sifa za `top`, `left`, `bottom`, na `right` hapa, lakini mtihani utafaulu tu kwa `top` ya `300px`, na kushoto kwa `400px`. -->

# --hints--

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `position` iliyowekwa kuwa `absolute`. Hakikisha unaongeza nukta ya nusu.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `top` iliyowekwa kuwa nambari yoyote katika `px`. Hakikisha unaongeza nukta ya nusu.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `left` iliyowekwa kuwa nambari yoyote katika `px`. Hakikisha unaongeza nukta ya nusu.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
