---
id: 6476fd4213318f6ee211028a
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Sasa badilisha thamani ya sifa ya `position` ya `.cat-head` yako kuwa `fixed`.

Baada ya hapo, piga skrini juu na chini kuona jinsi thamani ya `fixed` inavyofanya kazi.

# --hints--

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `position` iliyowekwa kuwa `fixed`. Hakikisha unaongeza nukta ya nusu.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'fixed');
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `top` iliyowekwa kuwa `100px`. Hakikisha unaongeza nukta ya nusu.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `left` iliyowekwa kuwa `100px`. Hakikisha unaongeza nukta ya nusu.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>fCC Cat Painting</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class="cat-head"></div>
    <div class="box"></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: absolute;
--fcc-editable-region--
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
