---
id: 64770351e8586671ec0911f0
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Msimbo wa `sticky` ni mchanganyiko wa msimbo wa `relative` na `fixed`. Unaruhusu kipengele **kubana** katika nafasi maalum ndani ya kipengele kinachokijumuisha au eneo la kuona, kulingana na nafasi ya ku-sroll.

Badilisha thamani ya sifa ya `position` ya `.cat-head` kuwa `sticky`, weka `top` kuwa `0`, kisha ondoa `left` na thamani yake.

**Kumbuka**: Ili kuona jinsi `sticky` inavyofanya kazi, lazima uweke maandishi machache kabla na baada ya kipengele chako cha `.cat-head`. Ukisogeza chini baada ya hapo, utaona kuwa `.cat-head` inashikamana juu na kubaki hapo.

# --hints--

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `position` iliyowekwa kuwa `sticky`. Hakikisha unaongeza nukta ya nusu.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'sticky');
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `top` iliyowekwa kuwa `0`. Hakikisha unaongeza nukta ya nusu.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

Haupaswi kuwa na sifa ya `left` na thamani yake katika msimbo wako.

```js
assert.notMatch(code, /left:\s*100px;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
