---
id: 647d855ac12cd436059acd39
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Tumia kichaguzi cha huru kuongeza `box-sizing: border-box;` kwenye CSS yako. Hii inahakikisha vipengele vinajumuisha pedi na mpaka katika upana na urefu uliobainishwa.

# --hints--

Unapaswa kuwa na kichaguzi cha (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Kichaguzi chako cha zote kinapaswa kuwa na sifa ya `box-sizing` iliyowekwa kuwa `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
