---
id: 6676a8a8d9861319ee901a90
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Sasa unaweza kuanza kufanya kazi kwenye masikio ya paka wako.

Safisha msimbo wako wa ukaguzi kwa kuondoa kichaguzi cha `.cat-right-ear` na sifa zake zote.

Ukichagua kwa kutumia kichaguzi cha darasa la vitu, mpe kipengele cha `.cat-left-ear` mpaka mipaka ya kushoto na kulia ya `35px solid transparent` kila moja. Pia, weka mpaka wa chini kuwa `70px solid #5e5e5e`.

# --hints--


Kichaguzi chako cha `.cat-right-ear` hakipaswi kuwa na sifa ya `height`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
}
```

Kichaguzi chako cha `.cat-right-ear` hakipaswi kuwa na sifa ya `width`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
}
```


Kichaguzi chako cha `.cat-right-ear` hakipaswi kuwa na sifa ya `border-left`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
}
```

Kichaguzi chako cha `.cat-right-ear` hakipaswi kuwa na sifa ya `border-right`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
}
```

Kichaguzi chako cha `.cat-right-ear` hakipaswi kuwa na sifa ya `border-top`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
}
```

Kichaguzi chako cha `.cat-right-ear` hakipaswi kuwa na sifa ya `border-bottom`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
}
```

Haupasi kuwa na kichaguzi cha `.cat-right-ear`.

```js
assert.notExists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Unapaswa kuwa na kichaguzi cha `.cat-left-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-left-ear'))
```

Kichaguzi chako cha `.cat-left-ear` kinapaswa kuwa na sifa ya `border-left` iliyowekwa kuwa `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft, '35px solid transparent')
```

Kichaguzi chako cha `.cat-left-ear` kinapaswa kuwa na sifa ya `border-right` iliyowekwa kuwa `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight, '35px solid transparent')
```

Kichaguzi chako cha `.cat-left-ear` kinapaswa kuwa na sifa ya `border-bottom` iliyowekwa kuwa `70px solid #5e5e5e`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom, '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
