---
id: a3566b1109230028080c9345
title: Diseñar un algoritmo para sumar todos los números
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

En este laboratorio, necesitará diseñar un algoritmo que sume todos los números.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Historias de usuario:**

1. Debería tener una función llamada `sumAll` que acepte un array de dos números.
1. `sumAll([n, m])` debería devolver la suma de `n` y `m` más la suma de todos los números entre ellos. El número más bajo no siempre irá primero. Por ejemplo, `sumAll([4,1])` debería devolver `10` porque la suma de todos los números entre `1` y `4` (ambos inclusive) es `10`.

# --hints--

Debería tener una función llamada `sumAll`.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` debería devolver un número.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` debería devolver `10`.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` debería devolver `10`.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` debería devolver `45`.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` debería devolver `45`.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--
```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;

  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
