---
id: 673339d96c56d36e963e888e
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите каталог розплідника рослин, одночасно практикуючись у використанні об’єктів Map і Set у JavaScript.

Щоб представити кожну рослину у вашому каталозі, ви використаєте об’єкти, кожен з яких має три властивості:

- `commonName`, загальна назва рослини
- `scientificName`, наукова (або латинська) назва рослини
- `cultivar`, сорт рослини

Почніть зі створення об’єкта з трьома вказаними властивостями для `Spanish lavender` (наукова назва `Lavandula stoechas`) сорту під назвою `Ballerina` і призначте його змінній з ім’ям `ballerina`.

# --hints--

У вас має бути змінна з ім’ям `ballerina`.

```js
assert.exists(ballerina);
```

Ви маєте призначити об’єкт змінній `ballerina`.

```js
assert.isObject(ballerina);
```

Ваш об’єкт `ballerina` має мати властивість `commonName`, призначену `"Spanish lavender"`.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

Ваш об’єкт `ballerina` має мати властивість `scientificName`, призначену `"Lavandula stoechas"`.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

Ваш об’єкт `ballerina` має мати властивість `cultivar`, призначену `"Ballerina"`.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

Ваш об’єкт `ballerina` має мати три ключі.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
