---
id: 6733b5632fb854bc6c58dc0d
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Тепер створіть інший об’єкт із тими самими трьома властивостями, щоб представити інший сорт іспанської лаванди під назвою `Pretty Polly`, і призначте його змінній з ім’ям `prettyPolly`.

# --hints--

У вас має бути змінна з ім’ям `prettyPolly`.

```js
assert.exists(prettyPolly);
```

Ви повинні призначити об’єкт змінній `prettyPolly`.

```js
assert.isObject(prettyPolly);
```

Ваш об’єкт `prettyPolly` має мати властивість `commonName`, призначену `"Spanish lavender"`.

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

Ваш об’єкт `prettyPolly` має мати властивість `scientificName`, призначену `"Lavandula stoechas"`.

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

Ваш об’єкт `prettyPolly` має мати властивість `cultivar`, призначену `"Pretty Polly"`.

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

Ваш об’єкт `prettyPolly` має містити три ключі.

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
