---
id: 6733b5e6f210c1bd1afdf8b3
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Залишилися об’єкти рослин уже додані для вас. Тепер настав час створити свій каталог.

Як і звичайні об’єкти JavaScript, об’єкти `Map` містять колекцію пар ключ-значення. Щоб створити об’єкт `Map`, ви можете використати конструктор `Map`:

```js
const myMap = new Map();
```

Оголосіть змінну `catalog` і призначте їй об’єкт `Map`.

# --hints--

У вас має бути змінна з ім’ям `catalog`.

```js
assert.exists(catalog);
```

Ви маєте призначити об’єкт `Map` змінній `catalog`.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
