---
id: 6734ab61e916486c80c490d1
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Коли значення, пов’язане з ключем, є об’єктом, метод `get` поверне референс на цей об’єкт. Ви можете використати цей референс, щоб змінити об’єкт всередині множини.

Щоб перевірити це, встановіть властивість `small` об’єкта `getBallerina` у ціле число, відмінне від `20`. Потім подивіться в консоль браузера, щоб переконатися, що значення змінилося у вашому каталозі.

# --hints--

Ви повинні встановити властивість `small` об’єкта `getBallerina` у ціле число, відмінне від `20`.

```js
const small = catalog.get(ballerina).small;
assert.match(__helpers.removeJSComments(code), /getBallerina\s*(?:\.\s*small|\[\s*("|')small\1\s*\])\s*=\s*\d+/);
assert.notEqual(small, 20);
assert.isTrue(Number.isInteger(small));
assert.isAtLeast(small, 0);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const getBallerina = catalog.get(ballerina);
console.log(getBallerina);
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
```
