---
id: 6734a772196245601e4b80e0
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Коли потрібно отримати елемент зі спроєктованої структури, ви можете використати метод `get`. Він приймає ключ як аргумент і повертає пов’язане значення.

Використайте метод `get`, щоб отримати значення ключа `ballerina` зі свого каталогу, призначте його змінній з ім’ям `getBallerina`. Потім виведіть цю нову змінну в консоль.

# --hints--

У вас має бути змінна з ім’ям `getBallerina`.

```js
assert.exists(getBallerina);
```

Значення вашої змінної `getBallerina` має бути `catalog.get(ballerina)`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+getBallerina\s*=\s*catalog\s*\.\s*get\s*\(\s*ballerina\s*\)/);
```

Виведіть у консоль змінну `getBallerina`.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*getBallerina\s*\)/);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
```
