---
id: 6734a606e5968a5a19b77171
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Продовжуйте додавати решту рослин до свого каталогу так само, як ви робили в попередніх кроках.

# --hints--

Ви повинні встановити свій об’єкт `willowVale` як ключ у множині `catalog`.

```js
assert.isTrue(catalog.has(willowVale));
```

Значення ключа `willowVale` у множині `catalog` має бути об’єктом із пропами `small`, `medium` і `large`, кожен із значенням додатного цілого числа.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Ви повинні встановити свій об’єкт `hidcote` як ключ у множині `catalog`.

```js
assert.isTrue(catalog.has(hidcote));
```

Значення ключа `hidcote` у множині `catalog` має бути об’єктом із пропами `small`, `medium` і `large`, кожен із значенням додатного цілого числа.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Ви повинні встановити свій об’єкт `imperialGem` як ключ у множині `catalog`.

```js
assert.isTrue(catalog.has(imperialGem));
```

Значення ключа `imperialGem` у множині `catalog` має бути об’єктом із пропами `small`, `medium` і `large`, кожен із значенням додатного цілого числа.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Ви повинні встановити свій об’єкт `royalCrown` як ключ у множині `catalog`.

```js
assert.isTrue(catalog.has(royalCrown));
```

Значення ключа `royalCrown` у множині `catalog` має бути об’єктом із пропами `small`, `medium` і `large`, кожен із значенням додатного цілого числа.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

У вас все ще має бути ключ `ballerina` у множині `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

У вас все ще має бути ключ `prettyPolly` у множині `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Ваша множина `catalog` має містити шість елементів.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
