---
id: 6734a13a29acdf47f2b8e161
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

На відміну від звичайних об’єктів, для яких ключі завжди мають бути рядками, об’єкти map дозволяють використовувати будь-який тип даних як ключі. У вашому каталозі ви будете використовувати пари ключ-значення, де ключем буде об’єкт, а значенням — об’єкт, що містить кількість доступних горщиків різних розмірів.

Змініть свій виклик `set`, щоб використати об’єкт `ballerina` як ключ. Для значення використайте об’єкт із трьома властивостями `small`, `medium` і `large`, кожна з яких має позитивне ціле число.

# --hints--

Ви повинні встановити об’єкт `ballerina` як ключ у вашій мапі `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Значення ключа `ballerina` у мапі `catalog` має бути об’єктом із властивостями `small`, `medium` і `large`.

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

Властивість `small` об’єкта, пов’язаного з ключем `ballerina` у мапі `catalog`, має бути позитивним цілим числом.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

Властивість `medium` об’єкта, пов’язаного з ключем `ballerina` у мапі `catalog`, має бути позитивним цілим числом.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

Властивість `large` об’єкта, пов’язаного з ключем `ballerina` у мапі `catalog`, має бути позитивним цілим числом.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Ваша мапа `catalog` має містити одну пару ключ-значення.

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
