---
id: 6734a2385c52c54d916a8b86
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Додайте ще одну пару ключ-значення до свого каталогу. Використайте змінну `prettyPolly` як ключ, а значення зробіть об’єктом із властивостями `small`, `medium` і `large`, значення яких знову є цілими числами.

# --hints--

Ви повинні встановити свій об’єкт `prettyPolly` як ключ у вашій множині `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Значення ключа `prettyPolly` у множині `catalog` має бути об’єктом із властивостями `small`, `medium` і `large`.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

Властивість `small` об’єкта, пов’язаного з ключем `prettyPolly` у множині `catalog`, має бути додатним цілим числом.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

Властивість `medium` об’єкта, пов’язаного з ключем `prettyPolly` у множині `catalog`, має бути додатним цілим числом.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

Властивість `large` об’єкта, пов’язаного з ключем `prettyPolly` у множині `catalog`, має бути додатним цілим числом.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

У вашій множині `ballerina` має залишатися ключ `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Ваша множина `catalog` має містити два елементи.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
