---
id: 67349ba230cff237c0f2bdf2
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Об’єкти map запам’ятовують порядок, у якому ключі додаються до множини. Щоб встановити пару ключ-значення у вашому об’єкті map, ви можете використати метод `set`. Наведений нижче приклад встановлює ключ `"name"` зі значенням `"Jerry"` в об’єкті map `myMap`.

```js
myMap.set("name", "Jerry");
```

Встановіть ключ `"ballerina"` зі значенням вашого об’єкта `ballerina` у вашій множині `catalog`. Потім виведіть ваш каталог у консоль.

Зверніть увагу, що щоб побачити ваш об’єкт map правильно, вам потрібно буде відкрити консоль браузера.

# --hints--

Ви повинні встановити ключ `"ballerina"` зі значенням вашого об’єкта `ballerina` у вашій множині `catalog`.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

Ваша множина `catalog` має містити одну пару ключ-значення.

```js
assert.lengthOf(catalog, 1);
```

Ви повинні вивести вашу множину `catalog` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
