---
id: 673521668be7905059b2d555
title: Крок 22
challengeType: 1
dashedName: step-22
---

# --description--

Нарешті, ви можете використати метод `forEach` для запуску функції на кожній парі ключ-значення у вашій мапі.

```js
myMap.forEach((value, key) => {console.log(key, value)});
```

Зверніть увагу, що функція зворотного виклику, передана в `forEach`, отримує значення через свій перший аргумент і ключ через другий аргумент для кожного запису.

Змініть свою функцію так, щоб вона повертала рядок, що містить увесь ваш каталог, і виведіть результат у консоль. Побудуйте свій рядок, використовуючи таку структуру для кожного запису в `catalog`:

```md
<scientific-name> '<cultivar>': <small-pots> S, <medium-pots> M, <large-pots> L

```

Зверніть увагу, що представлення кожного запису має завершуватися символом нового рядка.

# --hints--

Ви повинні вивести результат виклику `displayCatalog` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*displayCatalog\s*\(\s*\)\s*\)/)
```

Ваша функція `displayCatalog` має використовувати метод `forEach` на `catalog`.

```js
let flag = false;
const temp = catalog.forEach;
try {
    catalog.forEach = obj => {flag = true};
    displayCatalog();    
    assert.isTrue(flag);    
} finally {
    catalog.forEach = temp;
}
```

`displayCatalog` має повертати `--fcc-expected--`, але знайдено `--fcc-actual--`.

```js
const displayCatalogTest = () => {
    let catalogString = "";
    catalog.forEach((val, key) => {
        catalogString += `${key.scientificName} '${key.cultivar}': ${val.small} S, ${val.medium} M, ${val.large} L
`   })
    return catalogString
}
assert.equal(displayCatalog(), displayCatalogTest())
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const removePlant = plant => catalog.delete(plant);

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}
--fcc-editable-region--
const displayCatalog = () => {
    for (const entry of catalog.entries()) {
        console.log(entry);
    }
}

displayCatalog()
--fcc-editable-region--
```
