---
id: 6852f11d412c29665985a915
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

ユーザーがメールを送信できるようになったので、受信したメールを保存する方法が必要です。各ユーザーは自分のメールを集める受信箱を持つべきです。

`inbox`クラスの`User`メソッドに`__init__`属性を追加し、空のリスト`[]`として初期化してください。これでユーザーが受信したすべてのメールを保存できます。

# --hints--

`self.inbox = []`クラスの`User`メソッドに`__init__`を追加してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
