---
id: 6852f0c3fcedb964ab959921
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`read` 属性が `Email` クラスに正しく追加されたかどうかをテストしましょう。前のステップで既に email オブジェクトがあるので、`read` 属性を出力して、デフォルトでは `False` になっていることを確認してください。

# --hints--

読み取り用の属性をテストするには `email_obj.read` を出力してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
