---
id: 6852f06fdde1b4634a263f03
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

`Email` クラスをテストするために、emailオブジェクトを作成し、情報が正しくストアされているか確認しましょう。いくつかの属性を出力して、すべてが正しく動作しているか検証します。

`email_obj` という名前のemailオブジェクトを作成してください。送信者には `alice@example.com` 、受信者には `bob@example.com` 、件名には `Hello` 、本文には `Hi Bob!` を指定します。その後、`sender` と `subject` の属性を別々の出力文で出力し、正しくストアされているか確認してください。

# --hints--

指定されたパラメータで `email_obj` という名前のemailオブジェクトを作成してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

`email_obj.sender` を出力してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

`email_obj.subject` を出力してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
