---
id: 6852ef95ccc27660275692cd
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

あなたの`Email`クラスは、作成された各メールの情報を保存する必要があります。

`Email`クラスの中で、既存の`pass`キーワードを削除し、`__init__`、`self`、`sender`、`receiver`、`subject`をパラメータとして受け取る`body`メソッドを作成してください。

# --hints--

`__init__`クラスの中に`Email`メソッドを定義する必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

`__init__`メソッドは、`self`、`sender`、`receiver`、`subject`、`body`のパラメータを持つべきです。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
