---
id: 6852ea0c3cadb54e897acdac
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、異なるユーザー間でメールの送受信や管理をシミュレートする*Email Simulator*を作成します。クラスやオブジェクト、オブジェクト指向でコードを整理する方法について学びます。

まず、`Email`キーワードを使って`class`という名前のクラスを作成してください。

# --hints--

`Email`という名前のクラスを作成してください。

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
