---
id: 6852f01116912861c7d1b0e6
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

`__init__` メソッドの中で、各メールが情報をストアできるように、パラメータの値をオブジェクトの属性に代入する必要があります。

`__init__` メソッドの中で、`pass` キーワードを削除し、パラメータを `self.sender`、`self.receiver`、`self.subject`、および `self.body` に代入してください。

# --hints--

`sender` パラメータを `self.sender` に代入してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

`receiver` パラメータを `self.receiver` に代入してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

`subject` パラメータを `self.subject` に代入してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

`body` パラメータを `self.body` に代入してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
