---
id: 6853f90fd03acc9fd75ff860
title: ステップ 46
challengeType: 20
dashedName: step-46
---

# --description--

ユーザーは自分の受信トレイを確認し、メールを読み、Userオブジェクトを通じて直接メールを削除できるようにする必要があります。

受信トレイを確認するために、`check_inbox`クラスに`User`というメソッドを追加してください。このメソッドはユーザーの名前を使ったパーソナライズされたヘッダーを出力し、その後にすべてのメールを表示します。

ヘッダーは`\nName's Inbox:`の形式で、`Name`はユーザーの名前に置き換えられます。

ヘッダーを出力した後、ユーザーの受信トレイに対して`list_emails()`メソッドを呼び出してください。

# --hints--

`check_inbox`クラスに`User`という名前のメソッドを定義し、`self`をパラメータとして受け取る必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").has_args("self")`))
    }

})
```

`check_inbox`内で、ユーザーの名前を含む`\n[user]'s Inbox:`の形式のヘッダーを出力してください。

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").find_body()[0].is_equivalent('print(f"\\\\n{self.name}\\\'s Inbox:")')`))
    }
})
```

`check_inbox`内で、メールを表示するために`self.inbox.list_emails()`を呼び出してください。

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").find_body()[1].is_equivalent("self.inbox.list_emails()")`))
    }
})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

--fcc-editable-region--
    
--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

```
