---
id: 6853f2da9b8ed891e9fdc5fc
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

次に送信者と受信者の情報を表示しましょう。メールを送った人と受け取った人を表示するには、`display_full_email` メソッドに次の形式で2つの出力文を追加してください。

- `From: sender` ただし `sender` は送信者の名前に置き換えられます
- `To: receiver` ただし `receiver` は受信者の名前に置き換えられます

# --hints--

送信者の名前を出力するには `From: {self.sender.name}` を使った f-strings を使ってください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender.name}')")`));
    }
})
```

受信者の名前を出力するには `To: {self.receiver.name}` を使った f-strings を使ってください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver.name}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
