---
id: 6853f4584f508e95cc7a9f4d
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

これから受信トレイのすべてのメールを一覧表示するメソッドを作成します。このメソッドは、受信トレイが空の場合の処理と、メールがある場合に番号付きリストを表示する処理を行う必要があります。

`list_emails` クラスに `Inbox` という名前のメソッドを追加してください。このメソッドは `self` をパラメータとして受け取ります。まず、`if` リストが空かどうかを確認する `emails` 文を作成してください。空の場合は、メッセージ `Your inbox is empty.\n` を出力してください。メソッドを終了させるために `return` 文を追加してください。

# --hints--

`list_emails` クラスに `Inbox` という名前のメソッドを定義し、`self` をパラメータとして受け取る必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

`self.emails` が空かどうかを `if not self.emails` を使って確認してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

メールがない場合はメッセージ `Your inbox is empty.\n` を出力してください。

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

`return` 文の最後に `if` 文を入れてください。

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
