---
id: 6853f4ded9f84896ff182dd8
title: ステップ 31
challengeType: 20
dashedName: step-31
---

# --description--

空の受信トレイのチェックの後、メッセージ `\nYour Emails:` を出力してください。その後、`for` ループを使って受信トレイ内のすべてのメールを列挙しながら繰り返してください。これは、開始番号を指定した列挙の構文です：`enumerate(iterable, start=start_number)`。

列挙を使ってメールに `1` から番号を付けてください。反復変数として `i`（インデックス）と `email`（メールオブジェクト）を使います。

ループ内では、反復インデックスの後に `.` を付け、その後にスペースとメールオブジェクトの文字列表現を次の形式で出力する f-string を出力してください：

```md
i. string_representation
```

ここで `i` はインデックス、`string_representation` はメールオブジェクトの表現です。

# --hints--

`\nYour Emails:` のヘッダーは `if` 文の外で出力してください。

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

`for` ループで `enumerate(self.emails, start=1)` を繰り返す必要があります。

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

`for` ループは `i` と `email` を使って `enumerate(self.emails, start=1)` を繰り返すべきです。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

ループ内では、番号付きの接頭辞を付けて各メールの要約を出力し、その後 `__str__` メソッドのフォーマット済み文字列を出力してください。番号の後に `.` を追加するのを忘れないでください。

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        
--fcc-editable-region--
```
