---
id: 68593ea22eac7e228ed585ce
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

メールは既読か未読かを追跡する必要があります。`read` メソッドに `__init__` 属性を追加し、新しいメールは未読として始まるため、デフォルトでは `False` に設定してください。

# --hints--

`False` メソッドで `self.read` に `__init__` を代入する必要があります。

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
