---
id: 6859400f8f7f872e50e1238c
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

`User` クラスには、他のユーザーにメールを送信するメソッドが必要です。メールを送信するときは、新しい `Email` オブジェクトを作成し、受信者の受信トレイに追加します。

`send_email` クラスに `User` という名前のメソッドを作成してください。このメソッドは `receiver`、`subject`、`body` のパラメータを受け取る必要があります。

メソッド内で、次のパラメータ値を使って新しい `Email` オブジェクトを作成し、`email` という変数に代入してください。

- `sender`: `self`（現在のユーザー）
- `receiver`: `receiver`（メールを受け取るユーザー）
- `subject`: `subject`（メールの件名）
- `body`: `body`（メールの本文）

# --hints--

`send_email` クラスに `User` メソッドを定義してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

`send_email` メソッドは `self`、`receiver`、`subject`、`body` のパラメータを持つ必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

メソッド内で、正しいパラメータを使って `Email` オブジェクトを `email` という名前で作成してください。

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
