---
id: 68594df7779a164e6a61e8ae
title: ステップ 32
challengeType: 20
dashedName: step-32
---

# --description--

受信トレイには特定のメールを読むためのメソッドが必要です。ユーザーが読みたいメール番号（インデックス0から始まる）を指定すると、そのメソッドがメールの全文を表示します。

`read_email` クラスに `Inbox` という名前のメソッドを追加してください。このメソッドは `index` パラメータを受け取ります。まず受信トレイが空かどうかを確認し、空の場合はメッセージ `Inbox is empty.\n` を出力してください。その後、メソッドを終了するために `return` 文を追加してください。

# --hints--

`read_email` クラスに `Inbox` という名前のメソッドを定義し、`self` と `index` をパラメータとして受け取る必要があります。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index")`));
    }
})
```

`self.emails` が空かどうかを `if not self.emails` を使って確認してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`));
    }
})
```

メールがない場合は `Inbox is empty.\n` を出力してから戻ってください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\\\n')")`));
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[1].is_equivalent("return")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    
--fcc-editable-region--
```
