---
id: 685c1e32254e5993a5c8823e
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`Email` クラスの初期設定ができたので、メールの送受信ができるユーザーを作成しましょう。各ユーザーは名前を持ち、メール操作を行うことができます。

`User` クラスを作成し、`__init__` メソッドを定義してください。このメソッドは `self` と `name` をパラメータとして受け取ります。クラス内で `name` パラメータを `self.name` に代入してください。

# --hints--

`User` クラスを作成する必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

`User` クラスには、`__init__` メソッドがあり、`self` と `name` をパラメータとして受け取ります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

`name` パラメータを `self.name` に代入してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
