---
id: 685c1eba966c6396850736ea
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

ユーザーはメールを送信したり受信トレイを持ったりできますが、受信トレイの操作を効率的に管理する専用のクラスが必要です。

`Inbox` クラスはメールのリストを保持し、新しいメールを追加したり、すべてのメールを一覧表示したり、管理するためのメソッドを提供します。

`Inbox` という空のリストを初期化する `__init__` メソッドを持つ `emails` という新しいクラスを作成してください。

# --hints--

`Inbox` クラスを作成してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

`Inbox` クラスは `__init__` をパラメータとして受け取る `self` メソッドを持つべきです。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

`self.emails = []` は `Inbox` クラスの `__init__` メソッドで初期化してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
