---
id: 685d2fd1f309b9191f257c55
title: ステップ 40
challengeType: 20
dashedName: step-40
---

# --description--

メールシステムにタイムスタンプを組み込む前に、datetimeのフォーマット操作を練習しましょう。`datetime.datetime.now()` 関数は現在の日付と時刻を取得し、`strftime()` メソッドを使ってさまざまな形式にフォーマットできます。

`strftime()` がフォーマットコードとどのように連携するかは以下の通りです。

```python
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # Output: 2024-03-15 (year-month-day with - separator)
```

`%Y`（年）、`%m`（月）、`%d`（日）などのフォーマットコードは`strftime()`に何を含めるかを指示し、`-`のような区切り文字を間に入れることもできます。

コードの最後に、`current_time` という変数を作成し、`datetime.datetime.now()` を代入してください。そして`strftime()` を使って、`:` を区切り文字にして「時:分:秒」形式の時刻を出力してください。

使うフォーマットコードは、24時間表記の時間に`%H`、分に`%M`、秒に`%S`です。

# --hints--

`current_time` という変数を作成し、`datetime.datetime.now()` を代入してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("current_time = datetime.datetime.now()")`));
    }
})
```

`strftime()` を使って、`%H:%M:%S` 形式の時刻を出力してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(current_time.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

--fcc-editable-region--

--fcc-editable-region--
```
