---
id: 68875554c6c32de788afc64e
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

ユーザーはメールを正しく読む方法が必要です。

`display_full_email` クラスに `Email` だけをパラメータとして受け取る `self` メソッドを追加してください。このメソッドの中で、誰かがメールを閲覧したときにメールを既読としてマークするために `mark_as_read()` メソッドを呼び出してください。

このメソッドは、メールの全文をきれいにフォーマットして表示するために使われます。

# --hints--

`display_full_email` クラスに `Email` メソッドを定義してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

`display_full_email` メソッドは `self` だけをパラメータとして受け取るべきです。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

メソッドはメールを既読としてマークするために `self.mark_as_read()` を呼び出すべきです。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
