---
id: 688754d6d59578df1313cb3c
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

いよいよ User クラスと Inbox クラスをつなげて、メールを実際に配信できるようにしましょう！

ユーザーは単なるリストではなく適切な Inbox オブジェクトを持つ必要があり、`send_email` メソッドは受信者の Inbox にメールを届けるべきです。

`User` クラス `__init__` を更新して、各ユーザーが先ほど作成した空のリストではなく実際の Inbox オブジェクトを持つようにしてください。

# --hints--

User クラスを更新して `self.inbox = Inbox()` を初期化する必要があります。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
