---
id: 688754f3494b17e0b352d67d
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

いよいよ完成したメールシステムをテストする時です！ユーザーを作成して、メール機能が実際に動く様子を見てみましょう。

2つのUserオブジェクトを作成してください。`alice` は名前が `"Alice"`、`bob` は名前が `"Bob"` です。これでユーザーがどのように作成され、互いにやり取りできるかを示します。

# --hints--

`alice` という名前のUserオブジェクトを、名前 `"Alice"` で作成してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

`bob` という名前のUserオブジェクトを、名前 `"Bob"` で作成してください。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
