---
id: 68b9366fae8b13c1278761b1
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

受信トレイには新しいメールを受け取る仕組みが必要です。誰かがユーザーにメールを送信すると、そのメールは受信トレイに追加されるべきです。

`receive_email` クラスに `Inbox` という名前のメソッドを追加してください。このメソッドは `self` とメールオブジェクト `email` をパラメータとして受け取ります。メソッドの中で、`email` メソッドを使って `emails` を `append` リストに追加してください。

# --hints--

`receive_email` クラスに `Inbox` という名前のメソッドを定義し、`self` と `email` をパラメータとして受け取る必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

`email` を `self.emails` に追加してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
