---
id: 68b9350183ff05bd106c79e4
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

作成した `mark_as_read` メソッドを `Email` クラスでテストしてください。既存の `Email` オブジェクト `email_obj` に対してそのメソッドを使い、ステータスを変更しましょう。メールを既読にした後、変更を確認するために `read` の `email_obj` 属性を出力してください。

# --hints--

メールを既読にするには、`mark_as_read` の `email_obj` メソッドを呼び出す必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

`read` メソッドを呼び出した後、`email_obj` の `mark_as_read` 属性を出力してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
