---
id: 68b96d6a290d52d5e57a7485
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

`email_obj` とそれに続く print 文を削除してください。

# --hints--

コードに `email_obj` の作成行があってはいけません。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`));
    }
})
```

コードに `print(email_obj.sender)` があってはいけません。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.sender)")`));
    }
})
```

コードに `print(email_obj.subject)` があってはいけません。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.subject)")`));
    }
})
```

コードに `print(email_obj.read)` があってはいけません。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

コードに `email_obj.mark_as_read()` があってはいけません。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
email_obj.mark_as_read()
print(email_obj.read)
--fcc-editable-region--
```
