---
id: 68b930ee0e7b8fb81895e59a
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

今度はメールを既読にするメソッドを追加します。誰かがメールを開いたとき、そのステータスを未読から既読に変更したいです。このメソッドはシンプルで、`read` 属性を `True` に設定するだけで十分です。

`mark_as_read` クラスに `Email` という名前のメソッドを追加してください。このメソッドは `self` だけをパラメータとして受け取り、`read` 属性を `True` に設定する必要があります。

# --hints--

`mark_as_read` クラスに `Email` という名前のメソッドを定義してください。
このメソッドは `self` だけをパラメータとして受け取るべきです。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

メソッドは `self.read` を `True` に設定する必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
