---
id: 68be6a4447e61c4395ef2359
title: ステップ 36
challengeType: 20
dashedName: step-36
---

# --description--

受信トレイにはメールを削除する方法も必要です。

Pythonの`del`キーワードは、リストのインデックスでアイテムを削除するために使えます。

`delete_email`クラスに`Inbox`という名前のメソッドを追加してください。`read_email`と同様に、以下のことを行う必要があります：

- `index`パラメータを受け取り、受信トレイが空かどうかを確認し、空の場合はメッセージ`Inbox is empty.\n`を出力してください。メソッドを終了するには`return`を使ってください。

# --hints--

`delete_email`クラスに`Inbox`という名前のメソッドを定義し、`self`と`index`をパラメータとして受け取る必要があります。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").has_args("self,index")`))
    }
})

```

まず`if not self.emails:`を使って受信トレイが空かどうかを確認してください。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})

```

受信トレイが空の場合は、メッセージ`Inbox is empty.\n`を出力し、メソッドを終了するために`return`を使ってください。

```js
({
    test: () => 
    {
    assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_body().is_equivalent("print('Inbox is empty.\\\\n')\\nreturn")`))
}
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()
--fcc-editable-region--

--fcc-editable-region--
```
