---
id: 691f7bbf76229bb97c60827d
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将通过创建一个函数来练习使用 `for...of` 循环，该函数用于计算一个字串在字串数组中出现的次数。

在编写该函数之前，你将先创建一个更简单的函数，将字串的每个字符记录到控制台。

首先定义一个名为 `printCharacters` 的空函数，带有参数 `str`。

# --hints--

你应该创建一个 `printCharacters` 函数。

```js
assert.isFunction(printCharacters);
```

你的 `printCharacters` 函数应有一个名为 `str` 的参数。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
