---
id: 6920d189e36b6b60a6330fee
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

现在，在循环内部，将 `char` 变量记录到控制台。

# --hints--

你应该在循环内部有一个 `console.log` 语句，并且 `char` 是它的参数。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

当 `printCharacters` 函数被调用且参数为 `'test'` 时，函数内部的 `console.log()` 语句应该依次用以下参数调用四次：`'t'`、`'e'`、`'s'` 和 `'t'`。但实际上它被调用时使用了以下参数：`--fcc-actual--`。

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
