---
id: 6920c9cda1c7f056f75e4b1a
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

现在是时候编写你的第一个 `for...of` 循环了。在前述的课程中，你已经学会了可以使用 `for...of` 循环来迭代来自可迭代对象（例如字串或数组）的值。

这是一个 `for...of` 循环的示例：

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

为你的函数添加一个带有空代码块的 `for...of` 循环。它应该迭代 `str` 参数中的每个字符。

# --hints--

你的 `printCharacters` 函数的代码块内应该有一个 `for...of` 循环。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

你的 `for...of` 循环应该迭代函数的 `str` 参数。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
