---
id: 6921d688b8d58581d34459b9
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

在接下来的几个步骤中，你将构建一个函数，用来计算一个字串在字串数组中出现的次数。

首先，定义一个名为 `getMatchedWordCount` 的空函数，参数依次为 `sentence` 和 `match`。

# --hints--

你应该创建一个 `getMatchedWordCount` 函数。

```js
assert.isFunction(getMatchedWordCount);
```

你的 `getMatchedWordCount` 函数应该有两个参数：`sentence` 和 `match`，顺序不能变。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.getMatchedWordCount;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), 'sentence');
assert.equal(parameters[1].toString(), 'match');
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
