---
id: 6921e10ae2681205b516fbe6
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

现在是为你的循环添加适当递增 `count` 的逻辑的时候了。

在循环内部，使用条件式判断如果变量 `count` 与变量 `1` 相等，则将 `word` 增加 `match`。否则，保持 `count` 的值不变。

# --hints--

依次调用 `getMatchedWordCount([], "bar")`、`getMatchedWordCount(["foo"], "bar")`、`getMatchedWordCount(["foo", "bar", "foo"], "bar")` 和 `getMatchedWordCount(["foo", "bar", "foo"], "foo")` 应该都返回 `"number"` 类型的值；但它们返回的是 `--fcc-actual--` 类型的值。

```js
const obj = { method: (arg1, arg2) => typeof getMatchedWordCount(arg1, arg2) };
const spy = __helpers.spyOn(obj, 'method');
obj.method([], "bar");
obj.method(["foo"], "bar");
obj.method(["foo", "bar", "foo"], "bar");
obj.method(["foo", "bar", "foo"], "foo");
assert.deepEqual(spy.returns, ["number", "number", "number", "number"]);
spy.restore();
```

`getMatchedWordCount([], "bar")` 应该返回 `0`，但它返回了 `--fcc-actual--`。

```js
const resultString = String(getMatchedWordCount([], "bar"));
assert.strictEqual(resultString, "0");
```

`getMatchedWordCount(["foo"], "bar")` 应该返回 `0`，但它返回了 `--fcc-actual--`。

```js
const resultString = String(getMatchedWordCount(["foo"], "bar"));
assert.equal(resultString, "0");
```

`getMatchedWordCount(["foo", "bar", "foo"], "bar")` 应该返回 `1`，但它返回了 `--fcc-actual--`。

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "bar"));
assert.equal(resultString, "1");
```

`getMatchedWordCount(["foo", "bar", "foo"], "foo")` 应该返回 `2`，但它返回了 `--fcc-actual--`。

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "foo"));
assert.equal(resultString, "2");
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
