---
id: 6921de8d2a111d0341cff702
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

太好了，你的函数返回了一个 `0`，但它还没有计数任何内容。

为了解决这个问题，在 `getMatchedWordCount` 内，创建一个带有空代码块的 `for...of` 循环，用来迭代 `sentence` 中的每个单词。

# --hints--

你的 `getMatchedWordCount` 函数的代码块内应该有一个 `for...of` 循环。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

你的 `for...of` 循环应该迭代函数的 `sentence` 参数。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
--fcc-editable-region--
  
--fcc-editable-region--
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
