---
id: 6921d90b5ce50185505d2004
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

由于 `getMatchedWordCount` 函数需要返回一个数字计数，你将从设置一个计数器变量开始。

使用 `count` 关键字创建一个名为 `let` 的变量，初始化为 `0`，并添加一个 `return` 语句来返回 `count` 变量。

# --hints--

你应该在函数 `count` 内声明 `getMatchedWordCount` 变量。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getMatchedWordCount?.variables.count);
```

为了让 `count` 变量以后可以递增，使用 `let` 关键字声明它。

```js
const explorer = await __helpers.Explorer(code);
const countVar = explorer.allFunctions.getMatchedWordCount?.variables.count;
assert.isTrue(countVar?.toString().startsWith('let'));
```

你的函数应包含将值 `0` 赋值给 `count` 变量的赋值语句。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.variables.count?.value.matches('0'));
```

你的函数应包含一个 `return` 语句来返回 `count` 变量。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.hasReturn('count'));
```

你的 `getMatchedWordCount` 函数应该返回值 `0`。

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
