---
id: 6921dae10700e300f26994e3
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

要查看你当前状态下的 `getMatchedWordCount` 函数的输出，请在它下面添加一个 `console.log` 语句。在 `console.log()` 中，使用参数 `["I", "really", "really", "really", "like", "to", "code"]` 调用该函数作为 `sentence` 参数，使用 `"really"` 作为 `match` 参数。

# --hints--

你应该在 `getMatchedWordCount` 语句中调用函数 `console.log`。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

你应该调用函数 `getMatchedWordCount`，传入参数 `["I", "really", "really", "really", "like", "to", "code"]` 作为 `sentence` 参数，`"really"` 作为 `match` 参数。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
