---
id: 5dc17d3bf86c76b9248c6eb4
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`p` 要素の下に `h2` 要素を作成し、次のテキストを入れてください。

`Everyone loves cute cats online!`

# --hints--

`p` 要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert.exists(document.querySelector('p'));
```

`p` タグは小文字で書く必要があります。慣例として、すべての HTML タグは小文字で記述されます。

```js
assert.notMatch(code, /<\/?P>/);
```

`p` 要素には終了タグが必要です。終了タグは `/` 文字の直後に `<` があります。

```js
assert.match(code, /<\/p\>/);
```

`p` 要素のテキストは `Everyone loves cute cats online!` である必要があります。テキストが抜けているか、タイプミスがあります。

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

`p` 要素は `h2` 要素の下にある必要があります。順序が間違っています。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
