---
id: 5dc174fcf86c76b9248c6eb2
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、見出し、段落、リストなどの基本的なHTML要素を使って、猫の写真アプリを作りながら学習を続けます。

ワークショップは、`h1` 要素を追加し、そのテキストに `CatPhotoApp` を入れることから始めてください。

# --hints--

コード内に `CatPhotoApp` というテキストがあることを確認してください。スペルミスがないかもチェックしましょう。

```js
assert.match(code, /catphotoapp/i);
```

`h1` 要素には開始タグが必要です。開始タグの構文は `<elementName>` です。

```js
assert.exists(document.querySelector('h1'));
```

`h1` タグは小文字で書く必要があります。慣例として、すべてのHTMLタグは小文字で記述します。

```js
assert.notMatch(code, /<\/?H1>/);
```

`h1` 要素には終了タグが必要です。終了タグの構文は `</elementName>` です。

```js
assert.match(code, /<\/h1\>/);
```

`h1` 要素のテキストは `CatPhotoApp` であるべきです。テキストが抜けているか、タイプミスがあるか、`h1` 要素の開始タグと終了タグの間にテキストがありません。

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
