---
id: 5dc24614f86c76b9248c6ebd
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

アンカー（`a`）要素を使って別のページにリンクすることができます。

こちらは`https://www.freecodecamp.org`にリンクする例です：

```html
<a href="https://www.freecodecamp.org"></a>
```

段落の後に、`https://freecatphotoapp.com`にリンクするアンカー要素を追加してください。この時点では、プレビューにリンクは表示されません。

# --hints--

アンカー（`a`）要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert.exists(document.querySelector('a'));
```

アンカー（`a`）要素には終了タグが必要です。終了タグは`/`文字の直後に`<`があります。

```js
assert.match(code, /<\/a\>/);
```

アンカー（`a`）要素は`p`要素の下にあるべきです。順番が逆になっています。

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

アンカー（`a`）要素に`href`属性がありません。開始タグの名前の後にスペースがあるか、すべての属性名の前にスペースがあるか確認してください。

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

アンカー（`a`）要素は`https://freecatphotoapp.com`にリンクするべきです。URLを省略しているか、タイプミスがあります。

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

アンカー（`a`）要素の`href`属性に正しいリンクを設定していますが、属性値は常に引用符で囲むことをお勧めします。

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
