---
id: 5dc2385ff86c76b9248c6eb7
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 には、さまざまなコンテンツの領域を表す要素があります。 これらの要素によって HTML が読みやすくなり、検索エンジン最適化 (SEO) やアクセシビリティの向上に役立ちます。

`main` 要素は、HTML ドキュメント本文のメインとなるコンテンツを表すために使われます。 `main` 要素内のコンテンツはそのドキュメント特有のコンテンツであり、また、他の部分と重複しない内容であるべきです。

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

このページのメインとなる領域を表すため、`h1` 要素の前に開始タグ `<main>`、`p` 要素の後に終了タグ `</main>` を追加しましょう。

# --hints--

`main` 要素には開始タグが必要です。 Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

Your `main` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

`main` 要素には終了タグが必要です。 Closing tags have a `/` just after the `<` character.

```js
assert.match(code, /<\/main\>/);
```

`main` 要素の開始タグは `body` 要素の開始タグの下に置く必要があります。 順番が誤っているようです。

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

`main` 要素の開始タグは `h1` 要素の上に置く必要があります。 順番に誤りがあるようです。

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

`main` 要素の終了タグは `p` 要素の下に置く必要があります。 順番に誤りがあるようです。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```
