---
id: 5dc2385ff86c76b9248c6eb7
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5には異なるコンテンツエリアを識別する要素があります。これらの要素はHTMLを読みやすくし、検索エンジン最適化（SEO）やアクセシビリティに役立ちます。

`main` 要素はHTMLドキュメントの本文の主要なコンテンツを表すために使われます。`main` 要素内のコンテンツはドキュメント固有であり、他の部分で繰り返されるべきではありません。

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

このページのメインセクションを特定するために、`<main>` 要素の前に`h1` の開始タグを追加し、`</main>` 要素の後に`p` の終了タグを追加してください。

# --hints--

`main` 要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert.exists(document.querySelector('main'));
```

`main` タグは小文字である必要があります。慣例として、すべてのHTMLタグは小文字で書かれます。

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

`main` 要素には終了タグが必要です。終了タグは`/` の直後に`<` 文字が続きます。

```js
assert.match(code, /<\/main\>/);
```

`main` 要素の開始タグは`body` 要素の開始タグの下にあるべきです。順序が逆になっています。

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

`main` 要素の開始タグは`h1` 要素の上にあるべきです。順序が逆になっています。

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

`main` 要素の終了タグは`p` 要素の下にあるべきです。順序が逆になっています。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
