---
id: 5dc23f9bf86c76b9248c6eba
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

ウェブサイトに画像を追加するには、`img` 要素を使います。`img` 要素は終了タグを持たない開始タグだけの要素です。終了タグを持たない要素は <dfn>void element</dfn>（空要素）と呼ばれます。

`img` 要素の下に `p` 要素を追加してください。この時点ではブラウザに画像は表示されません。

# --hints--

`img` 要素は開始タグを持つ必要があります。開始タグの構文は次の通りです：`<elementName>`。

```js
assert.exists(document.querySelector('img'));
```

`img` 要素は終了タグ `</img>` を持ってはいけません。

```js
assert.notMatch(code, /<\/img\>/);
```

`img` 要素は1つだけにしてください。余分なものは削除してください。

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

`img` 要素は `p` 要素の下にあるべきです。順番が間違っています。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
