---
id: 5ddbd81294d8ddc1510a8e56
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

リンクのテキストは、アンカー要素（`a`）の開始タグと終了タグの間に配置する必要があります。

以下は、テキストが `click here to go to freeCodeCamp.org` のリンクの例です。

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

アンカー要素にアンカーテキスト `cat photos` を追加してください。これがリンクのテキストになります。

# --hints--

アンカー（`a`）要素には開始タグが必要です。開始タグの構文は`<elementName>`です。

```js
assert.exists(document.querySelector('a'));
```

アンカー（`a`）要素には終了タグが必要です。終了タグは`/`文字の直後に`<`があります。

```js
assert.match(code, /<\/a\>/);
```

アンカー要素（`a`）のテキストは `cat photos` にしてください。リンクテキストはアンカー要素（`a`）の開始タグと終了タグの間に置くことを確認してください。

```js
assert.equal(
  document.querySelector('a')?.innerText?.trim().toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
