---
id: 5dfb6a35eacea3f48c6300b4
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

`figcaption` 要素は、`figure` 要素内に含まれる画像を説明するキャプションを追加するために使います。

以下は、キャプションが `A cute cat` の `figcaption` 要素の例です。

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

`figure` 要素にネストされた画像の後に、テキストを次のように設定した `figcaption` 要素を追加してください。

`Cats love lasagna.`

# --hints--

`figcaption` 要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert.exists(document.querySelector('figcaption'));
```

`figcaption` 要素には終了タグが必要です。終了タグは `/` 文字の直後に `<` があります。

```js
assert.match(code, /<\/figcaption\>/);
```

`figcaption` 要素は `figure` 要素の中にネストされている必要があります。

```js
assert.exists(document.querySelector('figure > figcaption'));
```

ラザニアの `img` 要素は `figure` 要素にネストされている必要があります。

```js
assert.exists(document.querySelector('figure > img'))
assert.equal(
    document.querySelector('figure > img')?.getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

`figure` 要素にネストされた `figcaption` 要素は、`img` 要素の下に配置してください。順序が逆になっています。

```js
assert.equal(
  document.querySelector('figcaption')?.previousElementSibling.nodeName, 'IMG'
);
```

`figcaption` 要素のテキストは `Cats love lasagna.` である必要があります。テキストが抜けているか、タイプミスがあります。

```js
assert.match(
  document.querySelector('figcaption')?.innerText?.trim(), /Cats love lasagna.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          
--fcc-editable-region--          
        </figure>
      </section>
    </main>
  </body>
</html>
```
