---
id: 5ef9b03c81a63668521804d0
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--

特定の単語やフレーズを強調するには、`em` 要素を使うことができます。

`love` 要素内の単語 `figcaption` を、強調用の `em` 要素で囲んで強調してください。

# --hints--

強調用の `em` 要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert.exists(document.querySelector('em'));
```

強調用の `em` 要素には終了タグが必要です。終了タグは `/` の直後に `<` があります。

```js
assert.match(code, /<\/em\>/);
```

`figcaption` 要素を削除したか、開始タグまたは終了タグが欠けています。

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

強調用の `em` 要素はテキスト `love` を囲む必要があります。テキストを省略したか、タイプミスがあります。

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.trim().toLowerCase(), 'love'
);
```

`figcaption` のテキストは `Cats love lasagna` であるべきです。タイプミスがないか、また `em` 要素の開始タグと終了タグの周囲に必要なスペースがあるか確認してください。

```js
assert.match(
  document
    .querySelector('figcaption')
    ?.innerText?.trim().replace(/\s+/gi, ' '),
    /^cats love lasagna\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
