---
id: 5ef9b03c81a63668521804d2
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

番号付きリスト（`ol`）のコードは番号なしリストに似ていますが、番号付きリストのリスト項目は表示されると番号が付けられます。

`h3` 要素の下に、次の3つのリスト項目を持つ番号付きリストを追加してください。

`flea treatment`
`thunder`
`other cats`

# --hints--

`ol` 要素には開始タグが必要です。開始タグの構文は `<elementName>` のようになります。

```js
assert.exists(document.querySelector('ol'));
```

`ol` 要素には終了タグが必要です。終了タグは `/` 文字の直後に `<` があります。

```js
assert.match(code, /<\/ol>/);
```

`ol` 要素は2番目の `section` 要素の終了タグの上にあるべきです。順序が間違っています。

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'OL');
```

3つの `li` 要素は `ol` 要素の中にネストされている必要があります。

```js
assert.lengthOf(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ), 3
);
```

テキストが `li`、`flea treatment`、`thunder` の3つの `other cats` 要素があれば順序は問いません。

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText?.trim().replace(/\s+/g, ' ').toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

`ol` 要素は1つだけにしてください。

```js 
assert.lengthOf([...document.querySelectorAll('ol')], 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
