---
id: 5ef9b03c81a63668521804ea
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

これまでページに追加したすべての内容は`body`要素の中にあります。ページにレンダリングされるべきすべてのコンテンツ要素は`body`要素の中に入ります。しかし、他の重要な情報は`head`要素の中に入ります。

`head`要素は、ドキュメントのタイトルやスタイルシート、スクリプトへのリンクなどのメタデータを含むために使われます。メタデータとは、ページ上に直接表示されないページに関する情報です。

`head`要素の上に`body`要素を追加してください。

# --hints--

`body`要素を削除したか、開始タグまたは終了タグが欠けています。

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

`head`要素には開始タグが必要です。開始タグの構文は`<elementName>`のようになります。

```js
assert.match(code, /\<head\>/);
```

`head`要素には終了タグが必要です。終了タグは`/`文字の直後に`<`があります。

```js
assert.match(code, /\<\/head\>/);
```

`head`要素は開始`body`要素タグの上にあるべきです。別の場所に置かれています。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
