---
id: 5ef9b03c81a63668521804e7
title: Step 35
challengeType: 0
dashedName: step-35
---

# --description--

`footer` 要素は、ドキュメントやセクションのフッターを定義するために使われます。フッターには通常、ドキュメントの作成者情報、著作権データ、利用規約へのリンク、連絡先情報などが含まれます。

`main` 要素の後に `footer` 要素を追加してください。

# --hints--

`main` 要素を削除したか、開始タグまたは終了タグが欠けています。

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

`footer` 要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert.exists(document.querySelector('footer'));
```

`footer` 要素には終了タグが必要です。終了タグは `/` 文字の直後に `<` があります。

```js
assert.match(code, /<\/footer\>/);
```

`footer` 要素は終了 `main` 要素タグの下にあるべきです。別の場所に置かれています。

```js
assert.equal(document.querySelector('main')?.nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
