---
id: 5ef9b03c81a63668521804eb
title: Step 39
challengeType: 0
dashedName: step-39
---

# --description--

`title` 要素は、ブラウザーがページのタイトルバーやタブに表示する内容を決定します。

以下のテキストを使って、`title` 要素の中に `head` 要素を追加してください：

`CatPhotoApp`

# --hints--

`head` 要素を削除したか、開始タグまたは終了タグが欠けています。

```js
assert.match(code, /\<head\>/);
assert.match(code, /\<\/head\>/);
```

`title` 要素は `head` 要素の中にネストする必要があります。`title` 要素の開始タグと終了タグを必ず追加してください。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<head\>\<title\>.*\<\/title\>\<\/head\>/);
```

`title` 要素には終了タグが必要です。終了タグは `/` 文字の直後に `<` があります。

```js
assert.match(code, /\<\/title\>/);
```

`title` 要素のテキストは `CatPhotoApp` であるべきです。テキストを省略しているか、タイプミスがあります。

```js
assert.equal(document.title?.toLowerCase(), 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <head>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
