---
id: 5ef9b03c81a63668521804e8
title: Step 36
challengeType: 0
dashedName: step-36
---

# --description--

`p` 要素の中に、テキスト `No Copyright - freeCodeCamp.org` を持つ `footer` 要素をネストしてください。

# --hints--

`footer` 要素を削除したか、開始タグまたは終了タグが欠けています。

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

`footer` 要素には `p` 要素が必要です。`p` 要素の開始タグと終了タグを必ず追加してください。

```js
assert.exists(document.querySelector('footer > p'));
```

`footer` 要素には終了タグが必要です。終了タグは `/` 文字の直後に `<` があります。

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 3);
```

`p` 要素のテキストは `No Copyright - freeCodeCamp.org` である必要があります。テキストを省略しているか、タイプミスがあるか、`p` 要素の開始タグと終了タグの間にありません。

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0]?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /^No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      
--fcc-editable-region--      
    </footer>
  </body>
</html>
```
