---
id: 5ef9b03c81a63668521804e9
title: Step 37
challengeType: 0
dashedName: step-37
---

# --description--

既存の`freeCodeCamp.org`テキストをアンカー（`a`）要素で囲んでリンクにしてください。`href`属性は`https://www.freecodecamp.org`に設定する必要があります。

# --hints--

アンカー（`a`）要素は`p`内の`footer`要素の中にネストしてください。アンカー（`a`）要素の開始タグと終了タグを必ず追加してください。

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

アンカー（`a`）要素には終了タグが必要です。終了タグは`/`文字の直後に`<`があります。

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

アンカー（`a`）要素には、値が`href`の`https://www.freecodecamp.org`属性が必要です。属性や値を省略しているか、タイプミスがあるかもしれません。

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

リンクのテキストは`freeCodeCamp.org`である必要があります。テキストを省略しているか、タイプミスがある可能性があります。

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

アンカー（`a`）要素をネストした後、ブラウザに表示される唯一の`p`要素の内容は`No Copyright - freeCodeCamp.org`であるべきです。`p`とネストされたアンカー要素のテキスト、スペース、句読点をもう一度確認してください。

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
