---
id: 5f07be6ef7412fbad0c5626b
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

新しいコンテンツを追加する前に、猫の写真のコンテンツと今後のコンテンツを分けるために `section` 要素を使うべきです。

`section` 要素は、章、ヘッダー、フッター、またはその他のドキュメントのセクションなど、ドキュメント内のセクションを定義するために使われます。これは SEO とアクセシビリティに役立つセマンティックな要素です。

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

`h2` 要素、2つの `p` 要素、およびアンカー（`a`）要素を `section` 要素の中にネストしてください。

# --hints--

`section` 要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert.exists(document.querySelector('section'));
```

`section` 要素には終了タグが必要です。終了タグは `/` 文字の直後に `<` があります。

```js
assert.match(code, /<\/section\s*>/i);
```

`section` 要素全体は `main` 要素の開始タグと終了タグの間にあるべきです。

```js
assert.equal(document.querySelector('section')?.parentNode.nodeName, 'MAIN');
```

既存の `h2`、`p` 要素、およびアンカー（`a`）要素は `section` 要素の開始タグと終了タグの間にあるべきです。

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4)
```

`h1` 要素は `section` 要素の中にネストしてはいけません。

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      
      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
