---
id: 5f07c98cdb9413cbd4b16750
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

では新しいセクションを追加しましょう。 既存の `section` 要素の下に、2 つ目の `section` 要素を追加してください。

# --hints--

`section` 要素には開始タグが必要です。 Opening tags have this syntax: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

`section` の開始タグは 1 つだけ追加してください。 余分なものは削除してください。

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

`section` 要素には終了タグが必要です。 Closing tags have a `/` just after the `<` character.

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

`section` の終了タグは 1 つだけ追加してください。 余分なものは削除してください。

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

2 つ目の `section` 要素を 1 つ目の `section` 要素の中にネストしないでください。

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

両方の `section` 要素が、`main` 要素の開始タグと終了タグの間にあるようにしてください。

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```
