---
id: 671141feba228a35cefba82d
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

既存のテキスト `cute cats` を、次にリンクするアンカー要素に変えてください：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

最初の `p` 要素の中に新しいアンカー要素があるはずです。

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

アンカーテキストは `cute cats` であるべきです。

```js
assert.strictEqual(document.querySelector('a')?.innerText?.trim(), "cute cats");
```

新しいアンカー要素の `href` 属性は `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg` であるべきです。

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

コード内にフレーズ `cute cats` は1回だけ存在するはずです。

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

`p` 要素のテキストはまだ `Everyone loves cute cats online!` であるべきです。

```js
assert.strictEqual(document.querySelector('p')?.innerText?.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
