---
id: 68da6138d70960c197ec5ef2
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Se a string atual ainda não estiver completa, você precisa explorar os próximos estados possíveis. Adicione uma cláusula `else` ao seu comando `if`.

Dentro do bloco `else`, você vai lidar com a adição de parênteses de abertura. Adicione uma declaração `if` para verificar se `opens_used < pairs`. Isso garante que você não use mais parênteses de abertura do que o permitido.

Se esta condição for verdadeira, adicione uma nova tupla à `queue`: `(current + '(', opens_used + 1, closes_used)`. Isso representa o estado após adicionar um parêntese de abertura.

# --hints--

Você deve ter uma cláusula `else` para sua declaração `if` dentro do loop `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

Você deve ter uma declaração `if` dentro do bloco `else` verificando se `opens_used < pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

Você deve adicionar `(current + '(', opens_used + 1, closes_used)` à fila quando a condição for verdadeira.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
