---
id: 68da6138d70960c197ec5ef0
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Agora você vai implementar o loop principal do BFS. Crie um loop `while` que continue enquanto a `queue` não estiver vazia (enquanto `queue` for avaliada como `True`).

Dentro do loop, imprima `queue`.

# --hints--

Você deve ter um loop `while` na sua função `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

Seu loop `while` deve ter a condição `queue`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

Você deve imprimir `queue` dentro do seu loop `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
