---
id: 68da6138d70960c197ec5eee
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora você vai configurar a estrutura de dados para armazenar seus resultados. Crie uma variável chamada `result` e inicialize-a como uma lista vazia. Essa lista armazenará todas as combinações válidas de parênteses que você gerar.

Atualize sua instrução de retorno para retornar `result` em vez de uma lista vazia.

# --hints--

Você deve declarar uma variável chamada `result` na sua função `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

Você deve inicializar `result` como uma lista vazia.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("result").is_equivalent("result = []")
`) })
```

Você deve retornar `result` no final da sua função.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    return []
--fcc-editable-region--
```
