---
id: 68da6138d70960c197ec5ef3
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Agora você precisa lidar com a adição de parênteses de fechamento. A regra principal é que você só pode adicionar um parêntese de fechamento se isso mantiver o equilíbrio, ou seja, deve haver mais parênteses de abertura usados do que parênteses de fechamento.

Ainda dentro do bloco `else`, adicione a segunda instrução `if` para verificar se `closes_used < opens_used`.

Se esta condição for verdadeira, adicione outra nova tupla à `queue`: `(current + ')', opens_used, closes_used + 1)`. Isso representa o estado após adicionar um parêntese de fechamento.

# --hints--

Você deve ter uma segunda declaração `if` dentro do seu bloco `else`.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

Sua segunda declaração `if` aninhada deve verificar se `closes_used < opens_used`.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

Você deve adicionar `(current + ')', opens_used, closes_used + 1)` à fila quando a condição for verdadeira.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
