---
id: 68da6138d70960c197ec5eef
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Para a abordagem de busca em largura, você usará uma fila para rastrear diferentes estados enquanto constrói as combinações de parênteses. Cada estado será representado como uma tupla contendo três elementos:

- A string atual sendo construída
- O número de parênteses de abertura usados até agora
- O número de parênteses de fechamento usados até agora

Crie uma variável chamada `queue` e inicialize-a com uma lista contendo uma tupla: `('', 0, 0)`. Isso representa o estado inicial com uma string vazia e zero parênteses usados.

# --hints--

Você deve declarar uma variável chamada `queue` na sua função `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

Você deve inicializar `queue` com uma lista contendo a tupla `('', 0, 0)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    result = []
    return result
--fcc-editable-region--
```
