---
id: 68da6138d70960c197ec5ef1
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Agora você precisa identificar quando construiu uma combinação completa de parênteses. Uma combinação completa tem um comprimento igual ao dobro do número de pares (já que cada par contribui com um parêntese de abertura e um de fechamento).

Dentro do seu loop `while`, adicione uma instrução `if` para verificar se `len(current) == 2 * pairs`. Quando essa condição for verdadeira, adicione `current` à lista `result`.

# --hints--

Você deve ter uma declaração `if` dentro do seu loop `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

Sua declaração `if` deve verificar se `len(current) == 2 * pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

Você deve adicionar `current` à lista `result` dentro do seu novo comando `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
--fcc-editable-region--    
    return result

print(gen_parentheses(1))
```
