---
id: 68da604b8b6012a33e2344fb
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você implementará uma função que gera todas as combinações válidas de parênteses usando uma abordagem de busca em largura (BFS). Por exemplo, as combinações válidas de dois pares de parênteses são `(())` e `()()`.

Comece criando uma função chamada `gen_parentheses` com um único parâmetro `pairs`. Por enquanto, retorne uma lista vazia da função.

# --hints--

Você deve definir uma função chamada `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

Sua função `gen_parentheses` deve ter um único parâmetro chamado `pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

Sua função deve retornar uma lista vazia.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
