---
id: 68da6138d70960c197ec5eec
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Antes de implementar o algoritmo principal, você precisa validar a entrada. O parâmetro `pairs` deve ser um inteiro, pois representa o número de pares de parênteses a serem gerados.

Adicione uma declaração `if` no início da sua função para verificar se `pairs` não é um inteiro. Use a função `isinstance()` para isso.

Se a condição for verdadeira, retorne a string `The number of pairs should be an integer`.

# --hints--

Você deve ter uma declaração `if` na sua função `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

Sua declaração `if` deve verificar se `pairs` não é um inteiro usando `not isinstance(pairs, int)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

Você deve retornar a string `The number of pairs should be an integer` quando a entrada não for um inteiro.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    
    return []
--fcc-editable-region--
```
