---
id: 68da6138d70960c197ec5eed
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Em seguida, você precisa validar que o número de pares é pelo menos um, já que não é possível gerar combinações de parênteses com zero ou pares negativos.

Adicione outra declaração `if` para verificar se `pairs` é menor que `1`. Se essa condição for verdadeira, retorne a string `The number of pairs should be at least 1`.

# --hints--

Você deve ter uma segunda declaração `if` na sua função `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

Sua segunda declaração `if` deve verificar se `pairs` é menor que `1`.

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

Você deve retornar a string `The number of pairs should be at least 1` quando `pairs` for menor que `1`.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    
    return []
--fcc-editable-region--
```
