---
id: 68e00c5349cbddfef6dbc053
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Dentro do seu loop `while`, use `queue.pop(0)` para remover e obter o primeiro elemento da fila. Isso implementa o comportamento first-in-first-out (FIFO) característico do BFS.

Descompacte esta tupla em três variáveis: `current`, `opens_used` e `closes_used`.

# --hints--

Você deve usar `queue.pop(0)` para obter o primeiro elemento da fila.

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

Você deve desempacotar o resultado em três variáveis: `current`, `opens_used` e `closes_used`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        
--fcc-editable-region--
    return result
```
