---
id: 66be24cb4144f955b6bcc550
title: トランプのページを作成する
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. 少なくとも3枚のトランプを表示するウェブページを作成してください。
1. `main` 要素に `playing-cards` のIDを付けてください。
1. `#playing-cards` 要素の中に、クラスが `card` の`div` 要素を少なくとも3つ含めてください。
1. 各`.card` 要素の中に、クラスが `left` の`div` 要素を1つ目、クラスが `middle` の`div` 要素を2つ目、クラスが `right` の`div` 要素を3つ目として3つ含めてください。
1. `#playing-cards` 要素はフレックスボックスを使い、子要素を横方向に中央揃えにし、折り返しを許可し、子要素間に `20px` の間隔を設けてください。
1. 各`.card` 要素はフレックスボックスを使い、子要素を `space-between` で配置してください。
1. 各`.left` 要素はフレックスボックスのアイテムプロパティを使い、親コンテナの開始位置に揃えてください。
1. 各`.middle` 要素はフレックスボックスのアイテムプロパティを使い、親コンテナの中央に揃えてください。
1. 各`.right` 要素はフレックスボックスのアイテムプロパティを使い、親コンテナの終了位置に揃えてください。
1. 各`.middle` 要素はフレックスボックスを使い、子要素を縦方向の列に表示してください。

カードに使いたい場合は、以下の文字をコピー＆ペーストして使えます：`♠`、`♣`、`♥`、および `♦`。

# --hints--

`main` 要素に `id` が `playing-cards` のものを用意してください。

```js
assert.exists(document.querySelector('main#playing-cards'));
```

`#playing-cards` 要素の中に、クラスが `card` の`div` 要素を少なくとも3つ含めてください。

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

`.card` 要素には幅と高さを設定してください。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

`.card` 要素には、子要素として正確に3つの `div` 要素を含めてください。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

各`.card` 要素の中に、クラスが `left` の `div` を含めてください。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

各`.card` 要素の中に、クラスが `middle` の `div` を含めてください。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

各`.card` 要素の中に、クラスが `right` の `div` を含めてください。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

`#playing-cards` セレクターで、その要素の `display` を `flex` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

`#playing-cards` セレクターで、その要素の `justify-content` を `center` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

`#playing-cards` セレクターで、その要素の `flex-wrap` を `wrap` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

`#playing-cards` セレクターで、その要素の `gap` を `20px` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

`.card` セレクターで、その要素の `display` を `flex` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

`.card` セレクターで、その要素の `justify-content` を `space-between` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

`.left` セレクターで、その要素の `align-self` を `flex-start` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

`.middle` セレクターで、その要素の `align-self` を `center` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

`.right` セレクターで、その要素の `align-self` を `flex-end` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

`.middle` セレクターで、その要素の `display` を `flex` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

`.middle` クラスの中で `flex-direction` を `column` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
