---
id: 699d20623e57e8ffee885454
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

الخطوة الأولى في تحليل بطاقة الكتالوج هي فصل الحقول الأربعة. طريقة `.split()` تُقسّم سلسلة نصية إلى مصفوفة من السلاسل النصية الفرعية عند كل ظهور للفاصل الذي تحدده.

على سبيل المثال:

```js
"The Shining | King, Stephen | 1977 | Shelf K1".split("|")
// ["The Shining ", " King, Stephen ", " 1977 ", " Shelf K1"]
```

داخل `parseCard`، استدعِ `.split("|")` على `rawString` وأعد النتيجة. بعد هذه الخطوة، يجب أن تعرض وحدة التحكم مصفوفة تحتوي على أربع سلاسل نصية.

**ملاحظة:** قد تلاحظ وجود فراغات زائدة حول القيم. ستقوم بتنظيفها في الخطوة التالية.

# --hints--

يجب أن تستدعي `.split("|")` على `rawString` داخل دالة `parseCard` الخاصة بك.

```js
assert.match(parseCard.toString(), /rawString\s*\.\s*split\s*\(\s*('|"|`)\|\1\s*\)\s*;?/);
```

يجب أن تُعيد `parseCard` مصفوفة.

```js
assert.isArray(parseCard("A | B | C | D"));
```

عند استدعائها مع `"A | B | C | D"`، يجب أن تُعيد `parseCard` مصفوفة تحتوي على أربعة عناصر.

```js
assert.strictEqual(parseCard("A | B | C | D").length, 4);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
--fcc-editable-region--
  
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
