---
id: 699d20623e57e8ffee885455
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

كما ترى من الناتج في وحدة التحكم، بعد التقسيم، لا تزال كل جزء يحتوي على مسافات بيضاء زائدة في البداية و/أو النهاية. لتنظيف كل عنصر، تحتاج إلى معالجة كل واحد على حدة.

هذه حالة استخدام مثالية لـ `for` loop، التي تكرر كتلة شفرة مرة واحدة لكل عنصر في قابل للتكرار، مثل المصفوفة:

```js
for (let i = 0; i < array.length; i++) {
  // code to run on each iteration
}
```

تبدأ `i` عند `0` (الفهرس الأول)، وتستمر الحلقة طالما أن `i` أقل من طول المصفوفة، و`i++` يزيد `i` بواحد بعد كل تكرار.

داخل `parseCard`، أعلن عن مصفوفة فارغة باسم `trimmedParts`. ثم أضف حلقة `for` تتكرر على `parts`. اترك جسم الحلقة فارغًا الآن، لأنك ستملأه في الخطوة التالية.

# --hints--

يجب أن تعلن عن متغير باسم `trimmedParts` وتعيّنه إلى مصفوفة فارغة.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+trimmedParts\s*=\s*\[\s*\]/);
```

يجب أن تستخدم حلقة `for` للتكرار على `parts`.

```js
assert.match(parseCard.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*parts\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--

--fcc-editable-region--
  return parts;
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
