---
id: 699d1f8e51fcdc8230f779fb
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستبني نظام فهرسة لمكتبة تراثية. ستُحلل بيانات الكتب الخام، وتبحث وتُجمّع المدخلات، وتتحقق من جودة البيانات، وتصدر النتائج إلى صيغ مختلفة.

يُخزن كل بطاقة فهرسة كسلسلة نصية واحدة بالتنسيق `"Title | Author | Year | Location"`، مع فصل الحقول بواسطة رمز الأنبوب (`|`).

ستستخدم مصفوفة من هذه السلاسل النصية طوال الورشة لبناء واختبار دوال الفهرسة. أنشئ مصفوفة باسم `rawCatalogCards` وأضف إليها السلسلتين التاليتين:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

يجب إعلان متغير باسم `rawCatalogCards`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

يجب أن تكون `rawCatalogCards` مصفوفة.

```js
assert.isArray(rawCatalogCards);
```

العنصر الأول في `rawCatalogCards` يجب أن يكون `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

العنصر الثاني في `rawCatalogCards` يجب أن يكون `"The Shining | King, Stephen | 1977 | Shelf K1"`.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
