---
id: 699d20623e57e8ffee88545a
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

بعد تحليل الكتالوج، يمكنك البحث فيه حسب الحقل. ستقوم `findByAuthor` بتصفية المدخلات التي يحتوي مؤلفها على مصطلح البحث.

احذف عبارة `console.log(catalog.length)`.

ثم أنشئ دالة باسم `findByAuthor` تأخذ `catalog` و`author` كمعلمات. بداخلها، أعلن عن متغير باسم `searchTerm` وضعه مساويًا للإصدار بالأحرف الصغيرة من `author` لتمكين المطابقة غير الحساسة لحالة الأحرف لاحقًا.

أعلن أيضًا متغيرًا باسم `results` وعيّنه إلى مصفوفة فارغة للاحتفاظ بالتطابقات.

# --hints--

يجب أن تحذف عبارة `console.log(catalog.length)`.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /console\.log\(catalog\.length\)/);
```

يجب أن يكون لديك دالة باسم `findByAuthor` مع `catalog` و`author` كمعلمات.

```js
const explorer = await __helpers.Explorer(code);
const { findByAuthor } = explorer.allFunctions;
assert.exists(findByAuthor);
assert.equal(findByAuthor.parameters[0]?.toString(), "catalog");
assert.equal(findByAuthor.parameters[1]?.toString(), "author");
```

يجب أن تعلن متغيرًا باسم `searchTerm` وتعيّنه إلى `author.toLowerCase()`.

```js
const explorer = await __helpers.Explorer(code);
const { findByAuthor } = explorer.allFunctions;
assert.exists(findByAuthor);
const { searchTerm } = findByAuthor.variables;
assert.exists(searchTerm);
assert.isTrue(searchTerm.value.matches('author.toLowerCase()'));
```

يجب أن تعلن متغيرًا باسم `results` وتعيّنه إلى مصفوفة فارغة.

```js
const explorer = await __helpers.Explorer(code);
const { findByAuthor } = explorer.allFunctions;
assert.exists(findByAuthor);
const { results } = findByAuthor.variables;
assert.exists(results);
const resultsValue = eval(results.value.toString());
assert.isArray(resultsValue);
assert.isEmpty(resultsValue);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

--fcc-editable-region--
console.log(catalog.length);
--fcc-editable-region--
```
