---
id: 699d20623e57e8ffee885451
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

تأخذ الدالة `parseCard` سلسلة نصية خام من الكتالوج وتُرجع كائنًا منظمًا. قبل بناء دالتك، جهز طريقة لرؤية ناتجها أثناء العمل.

استدعِ `parseCard` مع `rawCatalogCards[2]` كمعلمة، عيّن النتيجة إلى متغير اسمه `cardResult`، وسجل `cardResult` في وحدة التحكم.

حاليًا لا تحتوي `parseCard` على جسم، لذا تُرجع `undefined`. مع إضافة شفرة إلى الدالة في الخطوات التالية، سيُظهر الناتج المسجل بالضبط ما تُرجعه الدالة في كل مرحلة.

# --hints--

يجب استدعاء `parseCard` مع `rawCatalogCards[2]` كمعلمة.

```js
assert.match(__helpers.removeJSComments(code), /parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)/);
```

يجب تعيين نتيجة استدعاء `parseCard` مع `rawCatalogCards[2]` إلى متغير اسمه `cardResult`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+cardResult\s*=\s*parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)\s*;?/);
```

يجب تسجيل `cardResult` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*cardResult\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {}

--fcc-editable-region--

--fcc-editable-region--
```
