---
id: 69aa11c2584523e6e991330c
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

تُزيل الطريقة `.trim()` الفراغات من كلا طرفي السلسلة النصية. على سبيل المثال، تعيد `" hello ".trim()` القيمة `"hello"`.

داخل جسم الحلقة، أضف `parts[i].trim()` إلى `trimmedParts`. ثم حدّث جملة `return` الحالية لتُعيد `trimmedParts` بدلًا من `parts`.

بعد هذه الخطوة، يجب أن تُظهر النتيجة المسجلة قيمًا نظيفة بدون فراغات زائدة.

# --hints--

يجب أن تضيف `parts[i].trim()` إلى `trimmedParts` داخل الحلقة.

```js
assert.match(parseCard.toString(), /trimmedParts\s*\.\s*push\s*\(\s*parts\s*\[\s*\w+\s*\]\s*\.\s*trim\s*\(\s*\)\s*\)/);
```

يجب أن تُعيد `parseCard` القيمة `trimmedParts`.

```js
const explorer = await __helpers.Explorer(code);
const { parseCard } = explorer.allFunctions;
assert.isTrue(parseCard.hasReturn("trimmedParts"));
```

يجب أن تُعيد `parseCard` مصفوفة حيث لا يحتوي أي عنصر على فراغات بادئة أو لاحقة.

```js
assert.isArray(parseCard("  A  |  B  |  C  |  D  "));
```

عند استدعائها مع `"  A  |  B  |  C  |  D  "`، يجب أن يكون العنصر الأول في المصفوفة المعادة هو `"A"`.

```js
assert.strictEqual(parseCard("  A  |  B  |  C  |  D  ")[0], "A");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
--fcc-editable-region--
    
  }
  return parts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
