---
id: 68e2c945cc1d8e778152be31
title: 플레이어 인터페이스 만들기
challengeType: 27
dashedName: build-a-player-interface
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `Player`라는 이름의 추상 클래스를 정의하고 `abc.ABC` 클래스를 상속받아야 합니다.
1. `Player` 클래스는 다음을 설정하는 `__init__` 메서드를 가져야 합니다:

   - `moves` 속성을 빈 리스트로 설정합니다.
   - `position` 속성을 `(0, 0)`로 설정합니다.
   - `path` 속성을 초기 위치를 포함하는 리스트로 설정합니다.

1. `Player` 클래스는 `make_move`라는 메서드를 가져야 합니다:

   - `random.choice`을 사용해 `moves` 속성(구체 클래스에서 정의됨)에서 무작위 이동을 가져옵니다.
   - 선택된 이동의 값을 현재 위치에 더해 `position` 속성을 업데이트합니다.
   - 새로운 `position` 튜플을 `path` 속성에 추가합니다.
   - 새로운 `position`를 반환합니다.

1. `Player` 클래스는 구체 클래스에서 구현할 추상 메서드 `level_up`를 가져야 합니다.
1. `Pawn` 클래스를 정의하고 `Player` 클래스를 상속받아야 합니다.
1. `Pawn` 클래스는 부모의 `super()` 메서드를 `__init__`을 사용해 호출한 후 `moves` 속성을 `x, y` 좌표를 나타내는 튜플 리스트로 설정해야 합니다.
1. 각 좌표 튜플은 다음 방향으로 `1` 단위 이동을 나타내야 합니다: 위, 아래, 왼쪽, 오른쪽.
1. `Pawn` 클래스는 `level_up` 메서드를 구체적으로 구현해 `moves` 속성에 더 많은 이동을 추가해야 합니다. 추가된 이동은 네 개의 대각선 이동을 나타내야 합니다 (예: `1` 단위 아래로 이동하고 `1` 단위 왼쪽으로 이동).

**참고:** 표준 라이브러리 모듈은 별칭 없이 가져와야 합니다. `Player` 클래스 관련 테스트는 `Pawn` 클래스가 인스턴스화 가능해질 때까지 실패합니다.

# --hints--

`Player`라는 이름의 클래스가 있어야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Player")`) })
```

`Player` 클래스는 `ABC` 모듈의 `abc` 클래스를 상속해야 합니다.

```js
({ test: () => runPython(`
from abc import ABC
import ast
cls = _Node(_code).find_class("Player")
assert cls.inherits_from("ABC") or cls.inherits_from("abc.ABC")
`) })
```

`Player` 클래스는 `__init__` 메서드를 가져야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("__init__")`) })
```

`Player`의 `__init__` 메서드는 단일 매개변수 `self`를 가져야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("__init__").has_args("self")`) })
```

`Player`의 `__init__` 메서드는 `self.moves`를 빈 리스트로 설정해야 합니다.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.moves = []") or init.has_stmt("self.moves = list()")
`) })
```

`Player`의 `__init__` 메서드는 `self.position`를 `(0, 0)`로 설정해야 합니다.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.position = (0, 0)")
`) })
```

`Player`의 `__init__` 메서드는 초기 위치를 포함하는 리스트로 `self.path`를 설정해야 합니다.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.path = [self.position]")
`) })
```

`Player` 클래스는 `make_move` 메서드를 가져야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("make_move")`) })
```

`Player`의 `make_move` 메서드는 단일 매개변수 `self`를 가져야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("make_move").has_args("self")`) })
```

`Player`의 `make_move` 메서드는 `random.choice`를 사용해 `moves` 속성에서 무작위 이동을 가져와야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").block_has_call("choice", "make_move")`) })
```

`Player`의 `make_move` 메서드는 무작위로 선택된 이동 좌표를 더해 `position` 속성을 업데이트해야 합니다.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.position == (1, 2)
  p.make_move()
  assert p.position == (2, 4)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

`Player`의 `make_move` 메서드는 새로운 `position` 튜플을 `path` 속성에 추가해야 합니다.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.path == [(0, 0), (1, 2)]
  p.make_move()
  assert p.path == [(0, 0), (1, 2), (2, 4)]

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

`Player`의 `make_move` 메서드는 업데이트된 `position` 속성을 반환해야 합니다.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  assert p.make_move() == (1, 2)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

`Player` 클래스는 `level_up` 메서드를 가져야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("level_up")`) })
```

`Player`의 `level_up` 메서드는 단일 매개변수 `self`를 가져야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("level_up").has_args("self")`) })
```

`Player`의 `level_up` 메서드는 추상 메서드여야 합니다.

```js
({ test: () => runPython(`
target = _Node(_code).find_class("Player").find_function("level_up")
assert target.has_decorators("abstractmethod") or target.has_decorators("abc.abstractmethod")
`) })
```

`Player` 클래스는 추상 클래스여야 합니다.

```js
({ test: () => runPython(`
try:
  Player()
except TypeError as e:
  assert str(e) == "Can't instantiate abstract class Player with abstract method level_up"
else:
  assert False, "Player class should not be instantiable"
`) })
```

`Pawn`라는 이름의 클래스가 있어야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Pawn")`) })
```

`Pawn` 클래스는 `Player` 클래스를 상속해야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").inherits_from("Player")`) })
```

`Pawn` 클래스는 `__init__` 메서드를 가져야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("__init__")`) })
```

`Pawn`의 `__init__` 메서드는 단일 매개변수 `self`를 가져야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_args("self")`) })
```

`Pawn`의 `__init__` 메서드는 `__init__` 함수를 사용해 부모의 `super`를 호출해야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_stmt("super().__init__()")`) })
```

`Pawn`의 `__init__` 메서드는 `moves` 속성을 `x, y` 좌표를 나타내는 튜플 리스트로 설정해야 하며, 각 좌표 튜플은 다음 방향으로 `1` 단위 이동을 나타내야 합니다: 위, 아래, 왼쪽, 오른쪽.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]
assert isinstance(p.moves, list)
assert len(p.moves) == 4
assert all(move in p.moves for move in moves)
`) })
```

`Pawn` 클래스는 `level_up` 메서드를 가져야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("level_up")`) })
```

`Pawn`의 `level_up` 메서드는 `1` 단위의 네 개 대각선 이동을 `moves` 속성에 추가해야 합니다.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1), (1, 1), (1, -1), (-1, 1), (-1, -1)]
p.level_up()
assert isinstance(p.moves, list)
assert len(p.moves) == 8
assert all(move in p.moves for move in moves)
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
from abc import ABC, abstractmethod
from random import choice


class Player(ABC):
    def __init__(self):
        self.moves = []
        self.position = (0, 0)
        self.path = [self.position]

    def make_move(self):
        x, y = self.position
        move = choice(self.moves)
        new_x = move[0] + x
        new_y = move[1] + y
        self.position = (new_x, new_y)
        self.path.append(self.position)
        return self.position

    @abstractmethod
    def level_up(self):
        pass


class Pawn(Player):
    def __init__(self):
        super().__init__()
        self.moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]

    def level_up(self):
        self.moves.extend([(1, 1), (1, -1), (-1, 1), (-1, -1)])
```
