---
id: 66a207974c806a19d6607073
title: Створіть застосунок для обліку книг
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути елемент `h1` з текстом `Book Inventory`.
1. У вас має бути елемент `table` з колонками з назвами `Title`, `Author`, `Category`, `Status` та `Rate`.
1. У вашій таблиці має бути щонайменше чотири рядки: перший для заголовків колонок, а решта заповнені інформацією про книги.
1. Кожен рядок таблиці всередині тіла таблиці має мати клас `read`, `to-read` або `in-progress`.
1. Елементи `td` колонки `Status` повинні містити елемент `span` з властивістю `class` `status`, що оточує текст `Read`, `To Read` або `In Progress`, залежно від класу цього рядка.
1. Елементи `td` колонки `Rate` повинні містити елемент `span` з властивістю `class` `rate`, що обгортає три порожні елементи `span`.
1. Елементи `.rate`, розміщені всередині рядків `.read`, повинні мати додатковий клас зі значенням `one`, `two` або `three`, залежно від особистої оцінки. Це значення має йти після `rate`.
1. Ви маєте створити три селектори атрибутів, щоб націлити елементи з класами `read`, `to-read` та `in-progress` і встановити їх властивість `background-image` з використанням `linear-gradient` на ваш вибір.
1. Ви маєте встановити властивість `display` кожного елемента `span` у значення `inline-block`.
1. Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `to-read`, і встановіть їх властивості `border` та `background-image`.
1. Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `read`, і встановіть їх властивості `border` та `background-image`.
1. Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `in-progress`, і встановіть їх властивості `border` та `background-image`.
1. Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status` та елементи `span` зі значенням класу, що починається з `rate`, і встановіть їх властивості `height`, `width` та `padding`.
1. Використайте селектор атрибутів, щоб націлити елементи `span`, які є прямими дочірніми елементами `span` зі значенням `class`, що починається з `rate`, і встановіть їх властивості `border`, `border-radius`, `margin`, `height`, `width` та `background-color`.
1. Використайте селектор атрибутів, щоб націлити перший нащадок елементів `span`, які мають слово `one` як частину свого значення `class`, і встановіть його властивість `background-image` з використанням `linear-gradient`.
1. Використайте селектор атрибутів, щоб націлити перших двох нащадків елементів `span`, які мають слово `two` як частину свого значення `class`, і встановіть їх властивість `background-image` з використанням `linear-gradient`.
1. Використайте селектор атрибутів, щоб націлити три елементи `span`, які є нащадками елементів `span`, що мають слово `three` як частину свого значення `class`, і встановіть їх властивість `background-image` з використанням `linear-gradient`.

# --hints--

У вас має бути елемент `h1` з текстом `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

У вас має бути лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

У вас має бути елемент `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

У вас має бути один елемент `thead` та один елемент `tbody` всередині `table`.

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

Всередині `thead` має бути один `tr` з 5 елементами `th`.

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

Ваш перший стовпець має мати текст `Title` як заголовок.

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

Ваш другий стовпець має мати текст `Author` як заголовок.

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

Ваш третій стовпець має мати текст `Category` як заголовок.

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

Ваш четвертий стовпець має мати текст `Status` як заголовок.

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

Ваш п’ятий стовпець має мати текст `Rate` як заголовок.

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

У вашій таблиці має бути щонайменше чотири рядки.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Кожен рядок має завжди містити 5 стовпців.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

Кожен рядок таблиці, крім заголовка, має мати клас `read`, `to-read` або `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

Елементи `td` колонки `Status` повинні містити елемент `span`.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Кожен елемент `span` колонки `Status` має мати клас `status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Кожен елемент `.status` має містити текст `Read`, `To Read` або `In Progress`, залежно від класу свого рядка.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

Елементи `td` колонки `Rate` повинні містити елемент `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Кожен елемент `span`, який є прямим дочірнім елементом `td` колонки `Rate`, має мати клас `rate` як перший клас.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Кожен елемент `.rate` має містити три порожні елементи `span`.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

Елементи `.rate`, розміщені всередині рядків `.read`, мають мати додатковий клас після класу `rate` зі значенням `one`, `two` або `three`, залежно від особистої оцінки.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

У вас має бути селектор атрибутів, щоб націлити рядки з класом `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Використайте селектор атрибутів, щоб націлити рядки з класом `read` і встановіть їх властивість `background-image` на лінійний градієнт на ваш вибір.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

У вас має бути селектор атрибутів, щоб націлити рядки з класом `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Використайте селектор атрибутів, щоб націлити рядки з класом `to-read` і встановіть їх властивість `background-image` на лінійний градієнт на ваш вибір.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

У вас має бути селектор атрибутів, щоб націлити рядки з класом `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Використайте селектор атрибутів, щоб націлити рядки з класом `in-progress` і встановіть їх властивість `background-image` на лінійний градієнт на ваш вибір.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Ви маєте встановити властивість `display` кожного елемента `span` у значення `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

У вас має бути селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `to-read`, і встановіть їх властивість `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `to-read`, і встановіть їх властивість `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

У вас має бути селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `read`, і встановіть їх властивість `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `read`, і встановіть їх властивість `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

У вас має бути селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `in-progress`, і встановіть їх властивість `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status`, які є нащадками елементів `tr` з класом `in-progress`, і встановіть їх властивість `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

У вас має бути селектор атрибутів, щоб націлити елементи `span` з класом `status` та елементи `span` зі значенням класу, що починається з `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status` та елементи `span` зі значенням класу, що починається з `rate`, і встановіть їх властивість `height`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status` та елементи `span` зі значенням класу, що починається з `rate`, і встановіть їх властивість `width`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Використайте селектор атрибутів, щоб націлити елементи `span` з класом `status` та елементи `span` зі значенням класу, що починається з `rate`, і встановіть їх властивість `padding`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

У вас має бути селектор атрибутів, щоб націлити елементи `span`, які є прямими дочірніми елементами `span` зі значенням `class`, що починається з `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Використайте селектор атрибутів, щоб націлити елементи `span`, які є прямими дочірніми елементами `span` зі значенням `class`, що починається з `rate`, і встановіть їх властивість `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Використайте селектор атрибутів, щоб націлити елементи `span`, які є прямими дочірніми елементами `span` зі значенням `class`, що починається з `rate`, і встановіть їх властивість `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Використайте селектор атрибутів, щоб націлити елементи `span`, які є прямими дочірніми елементами `span` зі значенням `class`, що починається з `rate`, і встановіть їх властивість `margin`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Використайте селектор атрибутів, щоб націлити елементи `span`, які є прямими дочірніми елементами `span` зі значенням `class`, що починається з `rate`, і встановіть їх властивість `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Використайте селектор атрибутів, щоб націлити елементи `span`, які є прямими дочірніми елементами `span` зі значенням `class`, що починається з `rate`, і встановіть їх властивість `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Використайте селектор атрибутів, щоб націлити елементи `span`, які є прямими дочірніми елементами `span` зі значенням `class`, що починається з `rate`, і встановіть їх властивість `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

У вас має бути селектор атрибутів, щоб націлити перший нащадок елементів `span`, які мають слово `one` як частину свого значення `class`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Використайте селектор атрибутів, щоб націлити перший нащадок елементів `span`, які мають слово `one` як частину свого значення `class`, і встановіть його властивість `background-image` з використанням `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

У вас має бути селектор атрибутів, щоб націлити перших двох нащадків елементів `span`, які мають слово `two` як частину свого значення `class`.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Використайте селектор атрибутів, щоб націлити перших двох нащадків елементів `span`, які мають слово `two` як частину свого значення `class`, і встановіть їх властивість `background-image` з використанням `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

У вас має бути селектор атрибутів, щоб націлити елементи `span`, які є нащадками елементів `span`, що мають слово `three` як частину свого значення `class`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Використайте селектор атрибутів, щоб націлити елементи `span`, які є нащадками елементів `span`, що мають слово `three` як частину свого значення `class`, і встановіть їх властивість `background-image` з використанням `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
