---
id: a0b5010f579e69b815e7c5d6
title: crea una funzione intelligente per la sostituzione di parole
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

In questo laboratorio, creerai una funzione che esegue una ricerca e sostituzione nella frase usando gli argomenti forniti e poi restituisce la nuova frase.

**Nota:** conserva il caso del primo carattere nella parola originale quando la sostituisci. Per esempio, se devi sostituire la parola `Book` con la parola `dog`, dovrebbe essere sostituita come `Dog`

**Obiettivo**: soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi creare una funzione chiamata `myReplace`.
1. La funzione `myReplace` deve prendere tre argomenti: una stringa, una parola da sostituire e la parola con cui sostituirla.
1. La funzione `myReplace` deve restituire la nuova stringa con la sostituzione.
1. Devi conservare il caso del primo carattere nella parola originale quando la sostituisci.

# --hints--

Devi creare una funzione chiamata `myReplace`.

```js
assert.isFunction(myReplace);
```

`myReplace` deve prendere tre argomenti.

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` deve restituire la stringa `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` deve restituire la stringa `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` deve restituire la stringa `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` deve restituire la stringa `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` deve restituire la stringa `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` deve restituire la stringa `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
