---
id: 66eaddd04a9e533fba689001
title: Projetar um Cartão de Postagem de Blog
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

Neste laboratório, você praticará como estilizar **fundos** e *bordas* criando um cartão de postagem de blog.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um elemento `div` com uma classe `blog-post-card` para conter todos os seus elementos de cartão.
2. Dentro do elemento `.blog-post-card`, você deve ter um elemento de imagem com um atributo `alt` válido e texto, e uma classe `post-img`. Você pode usar `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` para o atributo `src` da imagem.
3. Você deve ter um `div` com uma classe `post-content` dentro do elemento `.blog-post-card` com o seguinte:
    - Um elemento `h2` com uma classe `post-title` e texto para o título do seu post no blog.
    - Um elemento `p` com uma classe `post-excerpt` e texto para resumir seu post no blog.
    - Um elemento `a` com uma classe `read-more` com o texto `Read More`.
4. Você deve definir a cor de fundo do elemento `body` para um valor diferente de branco.
5. Você deve aplicar os seguintes estilos ao elemento `.blog-post-card`:
    - Um fundo branco.
    - Cantos arredondados.
    - Uma largura de sua escolha.
    - O alinhamento de texto da sua escolha.
6. O elemento `.post-img` deve ser estilizado para que a imagem preencha toda a largura do cartão e tenha uma borda inferior.
7. O elemento `.post-content` deve ser estilizado para que haja espaçamento interno dentro do cartão.
8. Os elementos `.post-title` e `.post-excerpt` devem ter uma cor de texto diferente da padrão e margens em todos os lados.
9. O elemento `.read-more` deve ser estilizado como um botão e possuir:
    - Uma cor de texto diferente da padrão.
    - Uma cor de fundo.
    - Margens em todos os lados.
    - Preenchimento em todos os lados.
    - Cantos arredondados.
    - Uma propriedade `display` definida como `inline-block`.
    - Um efeito hover que altera sua cor de fundo.

**Nota:** Certifique-se de vincular sua folha de estilos no seu HTML e aplicar seu CSS.

# --hints--

Você deve ter um `div` com uma classe `blog-post-card`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

Você deve ter um elemento `img` com uma `class` de `post-img`. Certifique-se de que sua imagem possua um atributo `alt` com texto e um atributo `src` com um valor.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

Você deve ter um `div` com a classe `post-content`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

Você deve ter um elemento `h2` com uma classe `post-title`. Certifique-se de que contenha algum conteúdo de texto para o título do seu post no blog.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

Você deve ter um `p` com uma classe `post-excerpt`. Certifique-se de que ele contenha algum conteúdo de texto para resumir sua postagem no blog.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

Você deve ter um elemento `a` com uma `class` de `read-more`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

Seu elemento `.read-more` deve conter o texto `Leia Mais`.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

Seu elemento `.blog-post-card` deve ter uma propriedade `border-radius` com um valor (não deve ser `0` nem um valor negativo).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

Seu elemento `.blog-post-card` deve ter um fundo branco.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

Seu elemento `body` deve ter uma cor de fundo diferente de branco.

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

Você deve direcionar `.blog-post-card` e definir sua propriedade `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

Você deve direcionar `.blog-post-card` e definir sua propriedade `text-align`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

Você deve direcionar `.post-content` e definir sua propriedade `padding`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

Seu elemento `.read-more` deve ter um efeito de hover.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

Você deve direcionar `.read-more` e definir sua propriedade `color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

Você deve direcionar `.read-more` e definir sua propriedade `background-color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

Você deve direcionar `.read-more` e definir sua propriedade `margin`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

Você deve direcionar `.read-more` e definir sua propriedade `display`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

Você deve direcionar `.read-more` e definir sua propriedade `border-radius`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

Você deve direcionar `.read-more` e definir sua propriedade `padding`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

Seu elemento `.post-img` deve preencher a `width` do cartão.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

Seu elemento `.post-img` deve ter um valor de `border-bottom`.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

Seus elementos `.post-title` e `.post-excerpt` devem ter margens e cores de texto não padrão.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
