---
id: 6a57c331ae94fa4631b60867
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

프런트 데스크는 `102`에 보관된 유물이 무엇인지 알아야 합니다. 반환된 제목을 기록하려면 `console.log()` 안에서 `getArtifactTitle(102)`를 호출하세요.

# --hints--

`console.log()` 안에서 `getArtifactTitle(102)`를 호출하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*getArtifactTitle\s*\(\s*102\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

--fcc-editable-region--

--fcc-editable-region--
```
