---
id: 6a57c331ae94fa4631b60871
title: 18단계
challengeType: 1
dashedName: step-18
---

# --description--

아티팩트의 현재 갤러리는 가장 최근 위치입니다. `locations` 배열의 마지막 항목을 확인하면 이 위치를 찾을 수 있습니다.

`currentLocation` 변수를 선언하고 `artifact.locations[artifact.locations.length - 1]`을 할당하세요. 이 표현식은 배열의 `length` 속성을 사용해 마지막 요소에 접근합니다.

# --hints--

`buildSummary` 안에 `currentLocation`이라는 변수를 선언하세요.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary?.variables.currentLocation);
```

아티팩트의 `locations` 배열 마지막 요소를 `currentLocation`에 할당하세요.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.buildSummary?.variables.currentLocation?.value.matches("artifact.locations[artifact.locations.length - 1]"));
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

function buildSummary(id) {
  const artifact = collection[id];
  if (!artifact) {
    return "Artifact not found";
  }
--fcc-editable-region--
  
--fcc-editable-region--
}
```
