---
id: 6a57c331ae94fa4631b6086e
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

박물관이 커지면서 유물의 담당자가 바뀔 수 있습니다.

`id`와 `name`을 매개변수로 받는 `updateCurator` 함수를 만드세요. `collection[id]`를 사용해 유물을 찾아보세요. 유물이 존재하면, 그 유물의 `curator` 객체의 `name` 속성을 전달받은 `name`으로 업데이트하세요.

# --hints--

`id`와 `name`이라는 매개변수를 가진 `updateCurator`라는 함수를 만들어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.updateCurator);
const params = explorer.allFunctions.updateCurator.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "name");
```

`updateCurator`를 호출하면 유물의 `curator` 객체의 `name`이 변경되어야 합니다.

```js
const curatorBefore = { ...collection[102].curator };
updateCurator(102, "Test Curator");
assert.strictEqual(collection[102].curator.name, "Test Curator");
assert.deepEqual(collection[102].curator, {
  ...curatorBefore,
  name: "Test Curator",
});
```

`updateCurator`는 `id`가 가리키는 유물의 담당자를 변경해야 합니다.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
updateCurator(457, "Fresh Curator");
assert.strictEqual(collection[457].curator.name, "Fresh Curator");
```

`updateCurator`는 `id`와 일치하는 유물이 없을 때 오류를 발생시키거나 컬렉션을 변경하지 않아야 합니다.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => updateCurator(999, "Test Curator"));
assert.strictEqual(JSON.stringify(collection), before);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

--fcc-editable-region--

--fcc-editable-region--
```
