---
id: 6a57c331ae94fa4631b60861
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

모든 유물에는 책임자가 있습니다.

키가 `101`인 저장소에 있는 아티팩트에 `curator` 프로프를 추가하고, 그 값을 `id`가 `201`이고 `name`이 `"Earl Sinclair"`인 객체로 설정하세요.

# --hints--

`101`에 저장된 유물에는 `curator` 속성이 있어야 합니다.

```js
assert.property(collection[101], "curator");
```

`curator`는 객체여야 합니다.

```js
assert.isObject(collection[101].curator);
```

`curator`는 `id`가 `201`이어야 합니다.

```js
assert.strictEqual(collection[101].curator.id, 201);
```

`curator`는 `name`이 `"Earl Sinclair"`여야 합니다.

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
