---
id: 6a57c265abbc95dfebc05098
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 유물 출처 감사 도구를 만들 거예요. 박물관 유물을 객체에 저장하고, 이를 조회하고, 태그를 붙이고, 전시실 간에 이동시키고, 역사를 요약하는 함수를 작성할 거예요.

각 유물은 ID 아래에 저장됩니다. `collection`이라는 객체를 만들고 항목을 추가하세요. 키는 `101`로 설정하고 값은 `title`이 `"Golden Mask"`이고 `category`가 `"Ceremonial"`인 객체로 만드세요.

# --hints--

`collection`이라는 변수를 선언하세요.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection`은 객체여야 합니다.

```js
assert.isObject(collection);
```

`collection`에는 하나의 항목이 있어야 합니다.

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

`101` 아래에 저장된 유물은 객체여야 합니다.

```js
assert.isObject(collection[101]);
```

`101` 아래에 저장된 유물은 `title`이 `"Golden Mask"`여야 합니다.

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

`101` 아래에 저장된 유물은 `category`가 `"Ceremonial"`이어야 합니다.

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
