---
id: 6a57c331ae94fa4631b60872
title: 19단계
challengeType: 1
dashedName: step-19
---

# --description--

다음으로, 유물을 요약하는 템플릿 리터럴을 반환하세요. 첫 번째 줄에는 제목을 넣고, 그다음 줄부터는 카테고리, 큐레이터 이름, 현재 갤러리, 전시 상태를 각각 한 줄씩 넣으세요. 갤러리는 `currentLocation.gallery`를 사용하세요.

예를 들어, `buildSummary(101)`을 호출하면 다음을 반환해야 합니다:

```md
Golden Mask
Category: Ceremonial
Curator: Fran Sinclair
Current Gallery: Hall C
On Display: true
```

# --hints--

`buildSummary(101)`에는 유물의 제목인 `Golden Mask`가 포함되어야 합니다.

```js
assert.include(buildSummary(101), "Golden Mask");
```

`buildSummary(101)`에는 `Category: Ceremonial` 줄이 포함되어야 합니다.

```js
assert.include(buildSummary(101), "Category: Ceremonial");
```

`buildSummary(101)`에는 `Curator: Fran Sinclair` 줄이 포함되어야 합니다.

```js
assert.include(buildSummary(101), "Curator: Fran Sinclair");
```

`buildSummary(101)`에는 `Current Gallery: Hall C` 줄이 포함되어야 하며, `currentLocation.gallery`를 사용해야 합니다.

```js
assert.include(buildSummary(101), "Current Gallery: Hall C");
```

`buildSummary(101)`에는 `On Display: true` 줄이 포함되어야 합니다.

```js
assert.include(buildSummary(101), "On Display: true");
```

`buildSummary(101)`는 각 필드를 각기 다른 줄에 넣어야 합니다.

```js
assert.isAtLeast(buildSummary(101).split("\n").length, 5);
```

`buildSummary(102)`는 가장 최근 위치를 사용하여 `Current Gallery: Hall B`를 표시해야 합니다.

```js
assert.include(buildSummary(102), "Current Gallery: Hall B");
```

`buildSummary`는 어떤 유물이든 각 필드별로 한 줄씩 순서대로 요약을 만들어야 합니다.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [
    { gallery: "Hall Y", year: 2001 },
    { gallery: "Hall Z", year: 2002 },
  ],
  tags: [],
  onDisplay: false,
};
const summary = buildSummary(457);
assert.deepEqual(
  summary.split("\n").map((line) => line.trim()),
  [
    "Test Artifact",
    "Category: Test Category",
    "Curator: Test Curator",
    "Current Gallery: Hall Z",
    "On Display: false",
  ]
);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

function buildSummary(id) {
  const artifact = collection[id];
  if (!artifact) {
    return "Artifact not found";
  }
  const currentLocation = artifact.locations[artifact.locations.length - 1];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
