---
id: 6a57c331ae94fa4631b6086c
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

시간이 지나면서 유물은 자주 전시 상태가 바뀝니다.

`id` 매개변수를 받는 `toggleDisplayStatus` 함수를 만드세요. `collection[id]`를 사용해 유물을 찾아보세요. 유물이 존재하면, 그 유물의 `onDisplay` 속성을 현재 값의 반대로 설정하세요.

# --hints--

`toggleDisplayStatus`라는 이름의 함수가 있고, `id`라는 매개변수를 가져야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.toggleDisplayStatus);
const params = explorer.allFunctions.toggleDisplayStatus.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

`toggleDisplayStatus`를 호출하면 유물의 `onDisplay` 값이 양방향으로 바뀌어야 합니다.

```js
const before101 = collection[101].onDisplay;
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, !before101);
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, before101);
const before102 = collection[102].onDisplay;
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, !before102);
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, before102);
```

`toggleDisplayStatus`는 `id`와 일치하는 유물이 없을 때 오류를 발생시키거나 컬렉션을 변경하지 않아야 합니다.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => toggleDisplayStatus(999));
assert.strictEqual(JSON.stringify(collection), before);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

--fcc-editable-region--

--fcc-editable-region--
```
