---
id: 6a57c331ae94fa4631b60862
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

각 유물은 어디에 있었는지, 어떻게 설명되는지, 전시 중인지도 추적합니다.

`101`에 저장된 유물에 세 가지 프로프를 더 추가하세요:

- `locations`: 두 개의 객체 `{ gallery: "Hall A", year: 2020 }`와 `{ gallery: "Hall C", year: 2024 }`를 담은 배열
- `tags`: 문자열 `"gold"`와 `"egypt"`를 담은 배열
- `onDisplay`: 불리언 `true`

# --hints--

유물의 `locations`는 주어진 두 위치 객체를 담은 배열이어야 합니다.

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

유물의 `tags`는 `["gold", "egypt"]`이어야 합니다.

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

유물의 `onDisplay`는 `true`이어야 합니다.

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
