---
id: 6a57c331ae94fa4631b6086b
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

`id`가 `102`인 유물을 `2026`년에 `"Hall B"`로 옮기세요. 그런 다음 `console.log()`를 `collection[102].locations`와 함께 호출해 그 유물의 위치 배열을 기록하세요.

# --hints--

`moveArtifact`를 `102`, `"Hall B"`, `2026`과 함께 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)/);
```

두 번째 유물의 `locations`를 콘솔에 기록해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

유물을 옮긴 후 `locations`를 기록해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

두 번째 유물의 가장 최근 위치는 `{ gallery: "Hall B", year: 2026 }`이어야 합니다.

```js
assert.deepEqual(collection[102].locations[collection[102].locations.length - 1], { gallery: "Hall B", year: 2026 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
