---
id: 6a57c331ae94fa4631b60866
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

가끔 직원들이 박물관에 없는 유물을 요청할 때가 있습니다. 만약 해당 `id`에 유물이 저장되어 있지 않으면, `collection[id]`는 `undefined`가 되고, 여기서 `.title`을 읽으면 오류가 발생합니다.

return 문을 조건문을 사용하도록 업데이트하세요. 유물이 발견되면 `artifact.title`을 반환하고, 그렇지 않으면 문자열 `"Artifact not found"`를 반환하세요.

# --hints--

`getArtifactTitle(101)`은 여전히 `"Golden Mask"`를 반환해야 합니다.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle`은 `id`와 일치하는 유물이 없으면 `"Artifact not found"`를 반환해야 합니다.

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle`은 `101`뿐만 아니라 찾은 모든 유물의 제목을 반환해야 합니다.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
