---
id: 6a57c331ae94fa4631b6086f
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

아티팩트 `101`의 큐레이터를 `"Fran Sinclair"`로 업데이트하세요. 그런 다음 `console.log()`를 `collection[101].curator.name`과 함께 호출하여 해당 아티팩트의 큐레이터 이름을 기록하세요.

# --hints--

`updateCurator`를 `101`과 `"Fran Sinclair"`로 호출하세요.

```js
assert.match(__helpers.removeJSComments(code), /updateCurator\s*\(\s*101\s*,\s*["']Fran Sinclair["']\s*\)/);
```

업데이트한 후 첫 번째 아티팩트의 큐레이터 `name`을 다시 기록하세요.

```js
assert.match(__helpers.removeJSComments(code), /updateCurator\s*\(\s*101\s*,\s*["']Fran Sinclair["']\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

첫 번째 아티팩트의 큐레이터 `name`은 `"Fran Sinclair"`여야 합니다.

```js
assert.strictEqual(collection[101].curator.name, "Fran Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
