---
id: 6a57c331ae94fa4631b60863
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

박물관에는 기록된 유물이 한 점 이상 있습니다. 첫 번째와 같은 형태로 키 `102` 아래에 두 번째 유물을 `collection`에 추가하세요.

다음 값을 사용하세요:

```js
{
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
}
```

# --hints--

`collection`에는 두 개의 항목이 있어야 합니다.

```js
assert.strictEqual(Object.keys(collection).length, 2);
```

`102` 아래에 저장된 유물은 설명된 객체와 일치해야 합니다.

```js
assert.deepEqual(collection[102], {
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
});
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
--fcc-editable-region--
  
--fcc-editable-region--
};
```
