---
id: 6a57c331ae94fa4631b60869
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

새로운 연구 결과 골든 마스크가 왕실에 속했음을 확인했습니다. `id`가 `101`인 유물에 태그 `"royal"`을 추가하세요. 그런 다음 `console.log()`를 `collection[101].tags`와 함께 호출하여 해당 유물의 태그를 기록하세요.

# --hints--

`101`과 `"royal"`을 인수로 하여 `addTag`를 호출하세요.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)/);
```

첫 번째 유물의 `tags`를 콘솔에 기록하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

새 태그를 추가한 후 `tags`를 기록하세요.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

첫 번째 유물의 `tags`에 `"royal"`이 포함되어야 합니다.

```js
assert.include(collection[101].tags, "royal");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
