---
id: 6a57c331ae94fa4631b6086a
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

유물들이 시간이 지나면서 갤러리 사이를 이동할 때마다, 각 이동은 `locations` 배열에 새로운 항목으로 기록됩니다.

`id`, `gallery`, `year` 매개변수를 가진 `moveArtifact` 함수를 만드세요. `collection[id]`를 사용해 유물을 찾으세요. 존재한다면, `gallery`와 `year`를 포함한 새 객체를 그 유물의 `locations` 배열에 추가하세요.

# --hints--

`id`, `gallery`, `year`라는 매개변수를 가진 `moveArtifact`라는 이름의 함수가 있어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.moveArtifact);
const params = explorer.allFunctions.moveArtifact.parameters;
assert.equal(params.length, 3);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "gallery");
assert.equal(params[2].toString(), "year");
```

`moveArtifact`를 호출하면 유물의 `locations` 배열에 새 `{ gallery, year }` 객체가 추가되어야 합니다.

```js
const start = collection[101].locations.length;
moveArtifact(101, "Test Hall", 2030);
assert.lengthOf(collection[101].locations, start + 1);
assert.deepEqual(collection[101].locations[collection[101].locations.length - 1], { gallery: "Test Hall", year: 2030 });
```

`moveArtifact`는 `id`가 가리키는 어떤 유물에든 새 위치를 추가해야 합니다.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
moveArtifact(457, "Test Vault", 2032);
assert.lengthOf(collection[457].locations, 2);
assert.deepEqual(collection[457].locations[1], { gallery: "Test Vault", year: 2032 });
```

`moveArtifact`는 `id`와 일치하는 유물이 없을 때 오류를 발생시키거나 컬렉션을 변경하지 않아야 합니다.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => moveArtifact(999, "Test Hall", 2030));
assert.strictEqual(JSON.stringify(collection), before);
```

배열을 교체하지 말고, 유물의 기존 `locations` 배열에 새 위치를 추가해야 합니다.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const locationsBefore = collection[458].locations;
const spy = __helpers.spyOn(locationsBefore, "push");
try {
  moveArtifact(458, "Push Hall", 2033);
  assert.deepEqual(spy.calls, [[{ gallery: "Push Hall", year: 2033 }]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].locations, locationsBefore);
assert.deepEqual(collection[458].locations[1], { gallery: "Push Hall", year: 2033 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

--fcc-editable-region--

--fcc-editable-region--
```
