---
id: 6a57c331ae94fa4631b6086d
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

토글이 제대로 작동하는지 확인하세요. 먼저 `console.log()`를 `collection[102].onDisplay`와 함께 호출하세요. 다음으로 `toggleDisplayStatus(102)`를 호출하세요. 마지막으로 다시 `console.log()`를 `collection[102].onDisplay`와 함께 호출하세요.

# --hints--

`toggleDisplayStatus`를 `102`와 함께 호출하세요.

```js
assert.match(__helpers.removeJSComments(code), /toggleDisplayStatus\s*\(\s*102\s*\)/);
```

두 번째 아티팩트의 `onDisplay` 값을 로그로 출력한 다음 `toggleDisplayStatus`를 호출하고, 다시 로그로 출력하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*onDisplay|\[\s*["']onDisplay["']\s*\])\s*\)[\s\S]*toggleDisplayStatus\s*\(\s*102\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*onDisplay|\[\s*["']onDisplay["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
