---
id: 6a57c331ae94fa4631b60865
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

직원들은 하루에도 수십 번 유물을 조회하므로, 감사자가 재사용 가능한 조회 기능이 필요합니다. 대괄호는 변수도 받을 수 있어서 `collection[id]`는 `id`가 가리키는 유물을 반환합니다.

`id` 매개변수를 받는 `getArtifactTitle` 함수를 만드세요. `collection[id]`를 사용해 유물을 가져와 `artifact`라는 변수에 저장하고, `artifact.title`을 반환하세요.

# --hints--

`id`라는 매개변수를 가진 `getArtifactTitle` 함수를 만들어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

`collection[id]`를 `artifact`라는 변수에 할당해야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle?.variables.artifact);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.variables.artifact?.value.matches("collection[id]"));
```

`getArtifactTitle`은 `artifact.title`을 반환해야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.hasReturn("artifact.title"));
```

`getArtifactTitle(101)`은 `"Golden Mask"`를 반환해야 합니다.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle(102)`은 `"Bronze Tablet"`를 반환해야 합니다.

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
