---
id: 6a57c331ae94fa4631b60868
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

큐레이터는 유물을 더 쉽게 검색할 수 있도록 태그를 붙이며, 태그는 절대 두 번 추가해서는 안 됩니다. 새 태그를 추가하기 전에 `includes()` 메서드를 사용해 이미 존재하는지 확인할 수 있는데, 배열에 값이 이미 포함되어 있으면 `true`를 반환합니다.

`id`와 `tag` 매개변수를 받는 `addTag`라는 함수를 만드세요. `collection[id]`를 사용해 유물을 찾아보세요. 유물이 존재하고 그 `tags` 배열에 `tag`가 아직 포함되어 있지 않으면, `tag`를 `artifact.tags`에 추가하세요.

# --hints--

`id`와 `tag`라는 매개변수를 가진 `addTag`라는 함수가 있어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

`addTag`를 새 태그와 함께 호출하면 그 태그가 유물의 `tags` 배열에 추가되어야 합니다.

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

`addTag`를 유물이 이미 가진 태그와 함께 호출하면 중복 태그가 추가되지 않아야 합니다.

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

`addTag`는 `id`가 가리키는 유물에 태그를 추가해야 합니다.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

`id`와 일치하는 유물이 없을 때 `addTag`는 오류를 발생시키거나 컬렉션을 변경하지 않아야 합니다.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

`tag`를 유물의 기존 `tags` 배열에 추가해야 하며, 배열을 교체해서는 안 됩니다.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
