---
id: 6a57c331ae94fa4631b60864
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

감사는 이미 기록된 내용을 읽는 것부터 시작하세요. 각 유물은 ID 아래 저장소에 저장되므로, `collection[101]`처럼 그 ID를 대괄호 안에 넣어 접근할 수 있습니다.

첫 번째 유물의 제목을 기록하려면 `collection[101].title`과 함께 `console.log()`를 호출하세요. 그런 다음, 그 유물의 큐레이터 이름을 기록하려면 `collection[101].curator.name`과 함께 `console.log()`를 호출하세요.

# --hints--

첫 번째 유물의 `title`을 콘솔에 기록하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)/);
```

첫 번째 유물의 큐레이터 `name`을 콘솔에 기록하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

`title`을 먼저, 큐레이터 `name`을 두 번째로 기록하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
