---
id: 68a47dcdbce12193caf77012
title: Implementar um Localizador de Índice Ordenado
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

Neste laboratório, você criará uma função que retorna o índice mais baixo no qual um valor deve ser inserido em um array depois que ele tiver sido ordenado em ordem crescente.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função `getIndexToIns` que recebe dois argumentos: um array e um número.
2. Você deve usar o método `sort` para ordenar o array em ordem crescente.
3. Sua função `getIndexToIns` deve retornar o índice mais baixo no qual o número deve ser inserido utilizando o método `findIndex`.
4. Sua função `getIndexToIns` deve sempre retornar um número.

**Dica:** O método `findIndex` é um método embutido de array em JavaScript. Ele recebe uma função de callback e retorna o índice do primeiro elemento que satisfaz a condição. Tanto `findIndex` quanto `sort` são funções de ordem superior.

**Exemplos:**

- `getIndexToIns([1, 2, 3, 4], 1.5)` deve retornar `1` porque `1.5` é maior que `1` (índice `0`) e menor que `2` (índice `1`).
- `getIndexToIns([20, 3, 5], 19)` deve retornar `2` porque após ordenar para `[3, 5, 20]`, `19` é menor que `20` (índice `2`) e maior que `5` (índice `1`).


# --hints--

Você deve ter uma função `getIndexToIns`.

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` deve sempre retornar um número.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

Sua função deve usar o método `sort`.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


Sua função deve fazer uso do método `findIndex`.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` deve retornar `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` deve retornar `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` deve retornar `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` deve retornar `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` deve retornar `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` deve retornar `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` deve retornar `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` deve retornar `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
``` 

