---
id: 66cc1a3a39aef47d6473cb2f
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

El siguiente paso es agregar la funcionalidad de división a la calculadora.

Declare una función llamada `calculateQuotient` que tome dos parámetros: `num1` y `num2`.

Dentro de la función, devuelva el cociente de los dos parámetros.

Luego, agregue un `console.log` que llame a la función `calculateQuotient` con los argumentos `7` y `11`.

# --hints--

Debe tener una función llamada `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

Su función `calculateQuotient` debe tener dos parámetros llamados `num1` y `num2`.

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Su función `calculateQuotient` debe devolver el cociente de los dos parámetros.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

Debe tener un `console.log` que llame a la función `calculateQuotient` con los argumentos `7` y `11`.

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
