---
id: 66cc21d23238dc8240a8a182
title: Paso 16
challengeType: 1
dashedName: step-16
---

# --description--

Sería agradable que la calculadora calcule el cuadrado de un número. El cuadrado de un número es el número multiplicado por sí mismo.

Para calcular el cuadrado de un número en JavaScript, puedes usar el método <dfn>Math.pow()</dfn>. El método `Math.pow()` toma dos argumentos: el número base y el exponente. 

También puedes usar el operador de potenciación (`**`) para calcular el cuadrado de un número.

Aquí hay un ejemplo:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

Declara una función llamada `calculateSquare` que tome un parámetro `num` y devuelva el cuadrado de `num`.

# --hints--

Debes tener una función llamada `calculateSquare`.

```js
assert.isFunction(calculateSquare);
```

Tu función `calculateSquare` debe tomar un parámetro `num`.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Tu función `calculateSquare` debe devolver el cuadrado de `num`.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
