---
id: 66cc095dfe1682753d2ab030
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Uno de los conceptos que aprendiste en lecciones anteriores fue cómo devolver valores desde una función.

Aquí hay un recordatorio de cómo devolver un valor de una función:

```js
function myFunction() {
  return "Hello World";
}
```

Dentro de tu función `addTwoAndSeven`, devuelve la suma de `2` y `7`.

# --hints--

Tu función `addTwoAndSeven` debería devolver la suma de `2` y `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Tu función `addTwoAndSeven` debería devolver la suma usando el operador `+`.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  
}
--fcc-editable-region--
```
