---
id: 66cc0a9e06e00b75d6782be9
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

En lecciones anteriores, aprendiste cómo llamar (invocar) a una función. Llamar a una función significa ejecutar el código dentro de la función.

Aquí tienes un recordatorio de cómo llamar a una función:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Agrega una declaración `console.log`, y dentro de ella, llama a la función `addTwoAndSeven`.

Ahora deberías ver la suma de `2` y `7` en el registro de la consola.

# --hints--

Debes llamar a la función `addTwoAndSeven` dentro de un `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addTwoAndSeven())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
