---
id: 66cc10958d72e57907f898b2
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Dentro de tu función `calculateSum`, necesitarás devolver la suma de los dos parámetros.

# --hints--

Tu función `calculateSum` debe devolver la suma de los dos parámetros.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function calculateSum(num1, num2) {
  
}
--fcc-editable-region--
```
