---
id: 66cc0de2920bec775f610424
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Aunque las funciones funcionan como se espera, hay mucha repetición en tu código.

Dado que estás construyendo una calculadora, no deseas tener que crear una función para cada posible combinación de números que quieras sumar.

Sería mejor crear una única función reutilizable que pueda sumar cualquier par de números.

Elimina todo el código que has escrito hasta ahora.

En el siguiente paso, repasarás cómo trabajar con parámetros y argumentos en funciones.

# --hints--

No debes tener una función llamada `addTwoAndSeven`.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
assert.notExists(addTwoAndSeven);
```

No debes tener una llamada a la función `addTwoAndSeven`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addTwoAndSeven())";
assert.notInclude(strippedCode, unexpectedStatement);
```

No debes tener una función llamada `addThreeAndFour`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
assert.notExists(addThreeAndFour);
```

No debes tener una llamada a la función `addThreeAndFour`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addThreeAndFour())";
assert.notInclude(strippedCode, unexpectedStatement);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
