---
id: 66cc0ba5881acb7692cfc4de
title: Pasos 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora, es momento de añadir otra función.

Declare una función llamada `addThreeAndFour` que devuelva la suma de `3` y `4`.

Luego, llame a la función `addThreeAndFour` dentro de un `console.log` para ver el resultado.

# --hints--

Debería declarar una función llamada `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

Su función `addThreeAndFour` debería devolver la suma de `3` y `4`.

```js
assert.equal(addThreeAndFour(), 7);
```

Su función `addThreeAndFour` debería devolver la suma usando el operador `+`.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

Debería llamar a la función `addThreeAndFour` dentro de un `console.log`.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());
--fcc-editable-region--

--fcc-editable-region--
```
