---
id: 66cc26907ff6908402af0149
title: Paso 17
challengeType: 1
dashedName: step-17
---

# --description--

Para probar tu función `calculateSquare`, necesitarás llamar a la función un par de veces.

Agrega un `console.log` que llame a la función `calculateSquare` con el argumento de `2`.

Luego, agrega otro `console.log` que llame a la función `calculateSquare` con el argumento de `9`.

# --hints--

Debes tener un `console.log` que llame a la función `calculateSquare` con el argumento de `2`.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*2\s*\)\s*\)/);
```

Debes tener un `console.log` que llame a la función `calculateSquare` con el argumento de `9`.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*9\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

--fcc-editable-region--

--fcc-editable-region--
```
