---
id: 66cc0f1ae40802781b2ea972
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

In previous lessons, you learned how to work with function parameters.

Un parámetro de función es una variable que se define en la declaración de la función y actúa como un marcador de posición.

Aquí tienes un ejemplo de una función que tiene un parámetro:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Declara una función llamada `calculateSum` que tome dos parámetros, `num1` y `num2`.

# --hints--

Deberías declarar una función llamada `calculateSum`.

```js
assert.isFunction(calculateSum);
```

Tu función `calculateSum` debería tomar dos parámetros, `num1` y `num2`.

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
