---
id: 66cc0ba5881acb7692cfc4de
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora es momento de agregar otra función.

Declare una función llamada `addThreeAndFour` que devuelva la suma de `3` y `4`.

Luego, llame a la función `addThreeAndFour` dentro de un `console.log` para ver el resultado.

# --hints--

Debería declarar una función llamada `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

Su función `addThreeAndFour` debería devolver la suma de `3` y `4`.

```js
assert.strictEqual(addThreeAndFour(), 7);
```

Su función `addThreeAndFour` debería devolver la suma usando el operador `+`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

Debería llamar a la función `addThreeAndFour` dentro de un `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
