---
id: 66cc16abc30b367ba1d1f9ef
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Ahora es el momento de probar tu función `calculateDifference`.

Comienza añadiendo un `console.log` que llame a la función `calculateDifference` con los argumentos `22` y `5`.

Luego, añade un segundo `console.log` que llame a la función `calculateDifference` con los argumentos `12` y `1`.

Finalmente, añade un tercer `console.log` que llame a la función `calculateDifference` con los argumentos `17` y `9`.

# --hints--

Debes tener un `console.log` que llame a la función `calculateDifference` con los argumentos `22` y `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(22,5))";
assert.include(strippedCode, expectedStatement);
```

Debes tener un `console.log` que llame a la función `calculateDifference` con los argumentos `12` y `1`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(12,1))";
assert.include(strippedCode, expectedStatement);
```

Debes tener un `console.log` que llame a la función `calculateDifference` con los argumentos `17` y `9`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(17,9))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
