---
id: 66cc12fa504b0479dac479a0
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

La ventaja de las funciones es que pueden ser llamadas con diferentes argumentos, permitiéndote reutilizar el mismo código con varios valores.

Aquí tienes algunos ejemplos de llamadas a funciones con diferentes argumentos de cadena:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

Añade un `console.log` que llame a la función `calculateSum` con los argumentos `2` y `5`.

# --hints--

Deberías tener un `console.log` que llame a la función `calculateSum` con los argumentos `2` y `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateSum(2,5))", "console.log(calculateSum(5,2))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
