---
id: 66cc172af7159a7c67804544
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

Ahora es el momento de agregar la funcionalidad de multiplicación a la calculadora.

Declara una función llamada `calculateProduct` que tome dos parámetros, `num1` y `num2`.

Dentro de la función, devuelve el producto de los dos parámetros.

Luego, agrega un `console.log` que llame a la función `calculateProduct` con los argumentos `13` y `5`.

# --hints--

Debes declarar una función llamada `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

La función `calculateProduct` debería tener dos parámetros llamados `num1` y `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateProduct } = explorer.allFunctions;
assert.lengthOf(calculateProduct?.parameters, 2);
assert.equal(calculateProduct?.parameters[0].toString(), "num1");
assert.equal(calculateProduct?.parameters[1].toString(), "num2");
```

La función `calculateProduct` debería devolver el producto de los dos parámetros.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

Debería tener un `console.log` que llame a la función `calculateProduct` con los argumentos `13` y `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateProduct(13,5))", "console.log(calculateProduct(5,13))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
