---
id: 66cc14ca14e65e7a9ebecb80
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Ahora que su calculadora puede sumar dos números, es hora de crear una función que reste dos números.

Declare una función llamada `calculateDifference` que tome dos parámetros, `num1` y `num2`.

Dentro de la función, retorne la diferencia entre los dos parámetros.

# --hints--

Debería tener una función llamada `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

Su función `calculateDifference` debe tener dos parámetros llamados `num1` y `num2`.

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Su función `calculateDifference` debe retornar la diferencia entre los dos parámetros.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
