---
id: 66cc095dfe1682753d2ab030
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Uno de los conceptos que aprendiste en lecciones anteriores fue cómo devolver valores desde una función.

Aquí hay un recordatorio de cómo devolver un valor de una función:

```js
function myFunction() {
  return "Hello World";
}
```

Dentro de tu función `addTwoAndSeven`, devuelve la suma de `2` y `7`.

# --hints--

Tu función `addTwoAndSeven` debería devolver la suma de `2` y `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Tu función `addTwoAndSeven` debería devolver la suma usando el operador `+`.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
const validExpressions = ["2 + 7", "7 + 2", "(2 + 7)", "(7 + 2)"];
assert.isTrue(validExpressions.some(expr => addTwoAndSeven?.hasReturn(expr)));
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
