---
id: 66cdf76685e4cb5a8726e27b
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Para seguir probando tu función `calculateSum`, necesitarás llamarla con diferentes argumentos.

Añade otro `console.log` que llame a la función `calculateSum` con los argumentos `10` y `10`.

Debajo de ese `console.log`, añade otro `console.log` que llame a la función `calculateSum` con los argumentos `5` y `5`.

# --hints--

Deberías tener un `console.log` que llame a la función `calculateSum` con los argumentos `10` y `10`.

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

Deberías tener un `console.log` que llame a la función `calculateSum` con los argumentos `5` y `5`.

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--

--fcc-editable-region--
```
