---
id: 67d83df6f82eda3868dd2a84
title: Створіть персонажа RPG
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

У цій лабораторній роботі ви попрактикуєтеся у базових навичках Python, створивши невеликий застосунок, який створює персонажа для пригоди в RPG.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з назвою `create_character`.
1. Функція має приймати, у такому порядку, ім’я персонажа, а потім три характеристики: силу, інтелект і харизму.
1. Ім’я персонажа має бути перевірене:
   - Якщо ім’я персонажа не є рядком, функція має повертати `The character name should be a string`.
   - Якщо ім’я персонажа є порожнім рядком, функція має повертати `The character should have a name`.
   - Якщо ім’я персонажа довше за 10 символів, функція має повертати `The character name is too long`.
   - Якщо ім’я персонажа містить пробіли, функція має повертати `The character name should not contain spaces`.
1. Характеристики також мають бути перевірені:
   - Якщо одна або кілька характеристик не є цілими числами, функція має повертати `All stats should be integers`.
   - Якщо одна або кілька характеристик менші за 1, функція має повертати `All stats should be no less than 1`.
   - Якщо одна або кілька характеристик більші за 4, функція має повертати `All stats should be no more than 4`.
   - Якщо сума всіх характеристик відрізняється від 7, функція має повертати `The character should start with 7 points`.
1. Якщо всі значення проходять перевірку, функція має повертати рядок із чотирьох ліній:
   - перша лінія має містити ім’я персонажа
   - лінії 2-4 мають починатися з абревіатури характеристики `STR`, `INT` або `CHA` (у цьому порядку), потім пробіл, а потім кількість повних крапок (`●`) відповідно до значення характеристики, і кількість порожніх крапок (`○`) до 10. Приклад: якщо значення сили дорівнює 3, має бути 3 повні крапки, за якими йдуть 7 порожніх. Крапки надані в редакторі.

Ось рядок, який має повертати `create_character('ren', 4, 2, 1)`:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

ПРИМІТКА: хоча `str` і `int` є загальними абревіатурами для характеристик, пам’ятайте, що це зарезервовані ключові слова в Python і їх не слід використовувати як імена змінних.

# --hints--

У вас має бути функція з назвою `create_character`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

Коли `create_character` викликається з першим аргументом, який не є рядком, вона має повертати `The character name should be a string`.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

Коли `create_character` викликається з першим аргументом, який є рядком, вона не має повертати `The character name should be a string`.

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

Коли `create_character` викликається з першим аргументом, який є порожнім рядком, вона має повертати `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

Коли `create_character` викликається з першим аргументом, який не є порожнім рядком, вона не має повертати `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

Коли `create_character` викликається з першим аргументом, довшим за 10 символів, вона має повертати `The character name is too long`.

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

Функція `create_character` не має повідомляти, що персонаж занадто довгий, якщо він не довший за 10 символів.

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

Коли `create_character` викликається з першим аргументом, який містить пробіл, вона має повертати `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

Коли `create_character` викликається з першим аргументом, який не містить пробілів, вона не має повертати `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

Коли `create_character` викликається з другим, третім або четвертим аргументом, який не є цілим числом, вона має повертати `All stats should be integers`.

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

Коли `create_character` викликається з другим, третім і четвертим аргументами, які всі є цілими числами, вона не має повертати `All stats should be integers`.

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

Коли `create_character` викликається з другим, третім або четвертим аргументом, який менший за `1`, вона має повертати `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

Коли `create_character` викликається з другим, третім і четвертим аргументами, які всі не менші за `1`, вона не має повертати `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

Коли `create_character` викликається з другим, третім або четвертим аргументом, який більший за `4`, вона має повертати `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

Коли `create_character` викликається з другим, третім і четвертим аргументами, які всі не більші за `4`, вона не має повертати `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

Коли `create_character` викликається з другим, третім або четвертим аргументом, сума яких не дорівнює `7`, вона має повертати `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

Коли `create_character` викликається з другим, третім і четвертим аргументами, сума яких дорівнює `7`, вона не має повертати `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` має повертати `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

Коли `create_character` викликається з дійсними значеннями, вона має виводити характеристики персонажа, як потрібно.

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
