---
id: 66e45c8140f9fda5c105ae26
title: 이벤트 전단지 페이지 만들기
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 본문에 `header` 요소가 있어야 합니다.
1. `header` 요소에는 이벤트 이미지를 포함하고, 그 안에 이벤트 제목이 들어간 `h1`가 순서대로 있어야 합니다. 원한다면 이 이미지를 사용할 수 있습니다: `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. 본문에 `main` 요소가 있어야 합니다.
1. `main` 요소 안에는 이벤트 기능을 보여주는 `section` 요소가 최소 두 개 있어야 합니다.
1. `section` 요소 각각에는 `h2`가 포함되어 있어야 합니다.
1. 본문에는 위아래 패딩이 `50px`로 설정되어야 합니다.
1. 본문에는 위아래 여백이 `0`로 설정되고, 좌우 여백은 중앙 정렬되도록 해야 합니다.
1. 본문의 너비는 뷰포트 너비에 상대적으로 설정되어야 합니다.
1. 본문은 `calc` 함수를 사용하여 `min-height` 속성을 뷰포트 높이의 100%에서 본문의 위아래 패딩을 뺀 값으로 설정해야 합니다.
1. 전단지 안에 `hr`가 최소 하나 있어야 합니다.
1. 모든 `width`와 `hr` 요소의 `section`는 부모에 상대적인 백분율 값으로 설정해야 합니다.

**참고:** HTML에 스타일시트를 연결하고 CSS를 적용했는지 확인하세요.

# --hints--

본문에 `header` 요소가 `body` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector('body > header'));
```

`img`의 첫 번째 자식으로 `header` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

`h1`의 두 번째 자식으로 `header` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

`h1` 요소는 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`h1`에는 이벤트 제목이 있어야 합니다.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

`main` 요소 안에 `body` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('body > main'));
```

`section` 요소 안에 `main` 요소가 최소 두 개 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

`section` 요소 각각에는 `h2`가 포함되어 있어야 합니다.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

`h2` 요소는 비어 있으면 안 됩니다.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

`body` 요소에는 위아래 패딩이 `50px`로 설정되어야 합니다.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

`body` 요소에는 위아래 여백이 `0`로 설정되어야 합니다.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

`body` 요소에는 좌우 여백이 `auto`로 설정되어야 합니다.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

`body` 요소는 `vw`를 사용해 너비를 설정해야 합니다.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

`body`는 `calc`를 사용해 `min-height`를 `100vh - 100px`로 설정해야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

`hr` 요소가 최소 하나 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

`hr` 요소들의 너비는 백분율 값으로 설정해야 합니다.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

`section` 요소들의 너비는 백분율 값으로 설정해야 합니다.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
