---
id: 65ef19425d1b27b6c930bae6
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Ikiwa nambari unayotaka kupata mzizi wa mraba ni hasi, msimbo unapaswa kutoa kosa kwa sababu mzizi wa mraba wa nambari hasi haujafafanuliwa katika nambari halisi.

Ondoa tamko la `pass` na unda tamko la `if` ili ukague kama `square_target` ni ndogo kuliko `0`.

# --hints--

Unapaswa kuondoa neno kuu la `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

Unapaswa kuunda tamko la `if` ili ukague kama `square_target < 0`.

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
