---
id: 65ef181dad3d96b14a21a03e
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika mradi huu, utapata mzizi wa mraba wa takriban wa nambari iliyotolewa kwa kutumia njia ya bisection.

Njia ya bisection ni mbinu ya kupata mizizi ya kitendakazi chenye thamani halisi. Inafanya kazi kwa kupunguza eneo ambalo mzizi wa mraba uko ndani yake hadi itakapokaribia thamani ndani ya uvumilivu uliowekwa.

Anza kwa kuunda kitendakazi kinachoitwa `square_root_bisection`. Acha vigezo viwe tupu kwa sasa.

# --hints--

Unapaswa kufafanua kitendakazi `square_root_bisection` bila vigezo vyovyote.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
