---
id: 65ef1be2b958cabf712296cd
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Sasa utapunguza mara kwa mara kipindi kwa kupata nambari ya katikati ya kipindi cha sasa na kulinganisha mraba wa nambari ya katikati na thamani lengwa.

Kwa ajili ya hilo, ndani ya blokki ya `else`, tengeneza mzunguko wa `for` unaoendesha hadi mara `max_iterations`.

Kwa mzunguko wako, tumia kitendakazi cha `range`, ambacho kinazalisha mfululizo wa nambari unazoweza kurudia juu yake. Sintaksia ni `range(start, stop, step)`, ambapo `start` ni nambari ya kuanzia (ikiwa ni pamoja na), `stop` ni nambari ya mwisho (bila kujumuisha), na `step` ni tofauti kati ya nambari na ile iliyotangulia katika mfululizo.

Pia, tumia `_` kama kigezo cha mzunguko. `_` hufanya kazi kama kiwakilishi na ni muhimu wakati unahitaji kutumia kigezo lakini hauna haja ya thamani yake halisi.

# --hints--

Unapaswa kutengeneza mzunguko wa `for` kurudia juu ya `range(max_iterations)`. Tumia `_` kama kigezo cha mzunguko.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
