---
id: 65ef1a50049cf9bada13266f
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Tengeneza tamko la `elif` ili kukagua kama `square_target` ni sawa na `0`. Ikiwa ni sawa, weka thamani `0` kwa kigezo `root`. Pia, chapisha ujumbe `'The square root of {square_target} is 0'`. Kumbuka kuunda ujumbe huo kwa kutumia mfuatano wa herufi wa f-string.

# --hints--

Unapaswa kuwa na tamko la `elif` ili kukagua sharti la `square_target == 0`.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

Unapaswa kuweka thamani `0` kwa kigezo `root` na kupitisha hoja `f'The square root of {square_target} is 0'` kwa wito wa `print`.


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
