---
id: 65ef190c6b51e9b5a5f7ed29
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Toa kitendakazi cha `square_root_bisection` vigezo vifuatavyo:

- `square_target`: Nambari unayotaka kupata mzizi wa mraba wake.
- `tolerance` (hiari): Tofauti inayokubalika kati ya mraba wa thamani ya mzizi wa makadirio na thamani halisi ya lengo (chaguo-msingi ni `1e-7`). Kuvumilika kwa `1e-7` kunamaanisha kuwa suluhisho litakuwa sahihi ndani ya `0.0000001` ya thamani halisi na ni chaguo nzuri cha msingi kinacholinganisha usahihi na utendaji.
- `max_iterations` (hiari): Idadi kubwa ya marudio ya kufanya (chaguo-msingi ni `100`). Ikiwa njia haitafikia suluhisho ndani ya kikomo hiki, utachukulia kuwa suluhisho halijapatikana.
  
# --hints--

Kitendakazi chako kinapaswa kuwa na vigezo hivi: `square_target`, `tolerance = 1e-7`, na `max_iterations = 100`. Mpangilio ni muhimu.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```
