---
id: 65ef19ec318e2fb9066381e3
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Utatengeneza kesi tofauti kwa wakati `square_target` ni `0` au `1`. 

Anza kwa kutengeneza usemi wa `if` ili ukague kama `square_target` ni sawa na `1`. 

# --hints--

Unapaswa kutengeneza usemi wa `if` ili ukague kwamba `square_target == 1`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[0].is_equivalent("square_target == 1")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')

--fcc-editable-region--
```
