---
id: 65ef198fde24dfb7ff675b42
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Sentensi ya `raise` inakuwezesha kulazimisha kutokea kwa aina fulani ya ubaguzi. Inajumuisha neno kuu la `raise` linalofuatiwa na aina ya ubaguzi, na inakuwezesha kutoa ujumbe wa makosa uliobinafsishwa:

```py
raise ValueError("Invalid value")
```

Wakati msimbo ulio hapo juu unapoendeshwa, `ValueError` hutolewa na ujumbe wa `"Invalid value"` unaonyeshwa kwa mtumizi.

Ikiwa `square_target` ni ndogo kuliko `0`, mzizi wa mraba wa thamani halisi hauwezi kuhesabiwa. Kwa hiyo, inua `ValueError` ukiambatanisha na ujumbe wa `'Square root of negative number is not defined in real numbers'`. Usisahau kuondoa neno kuu la `pass`.

# --hints--

Unapaswa kuondoa neno kuu la `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

Unapaswa kuinua `ValueError` ukiambatanisha na ujumbe wa `'Square root of negative number is not defined in real numbers'` ndani ya mwili wa `if`.

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
