---
id: 65ef1c34a541afc12dbb849d
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Kitendakazi cha `abs()` hurudisha thamani kamili ya nambari, ambayo daima ni chanya, bila kujali alama ya nambari hiyo. Utaikitumia kukagua kama mzizi wa mraba unaokadiriwa uko karibu vya kutosha na thamani halisi.

Sasa, tengeneza usemi wa `if` kukagua kama thamani kamili ya tofauti kati ya `square_mid` na `square_target` iko ndani ya `tolerance` iliyobainishwa.

# --hints--

Unapaswa kuwa na usemi wa `if` wenye sharti la `abs(square_mid - square_target) < tolerance:` ndani ya mwili wa mzunguko wa `for`.

```js
({
    test: () =>
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("abs(square_mid - square_target) < tolerance")`))
    } 
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2
            
--fcc-editable-region--
```
