---
id: 65ef1d5e3d2927c5e0f4997b
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Katika Python, neno kuu la `is` linakagua utambulisho wa kitu. Linatumika kubaini kama vigezo viwili vinaelekeza kwenye kitu kimoja katika kumbukumbu. Tofauti na `is`, opereta ya usawa (`==`) huthibitisha kama thamani za vitu viwili ni sawa, bila kujali kama ni kitu kimoja katika kumbukumbu.

Nje ya mzunguko wa for, tengeneza tamko la `if` la kukagua kama `root is None`. Ikiwa ni kweli, chapisha ujumbe wa `'Failed to converge within {max_iterations} iterations.'`. Kumbuka kuunda ujumbe huo kwa kutumia mfuatano wa herufi wa f-string.

# --hints--

Unapaswa kutengeneza tamko la `if` linalokagua kama `root is None`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

Unapaswa kupitisha hoja ya `f'Failed to converge within {max_iterations} iterations.'` kwa kitendakazi cha `print`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
