---
id: 69aaafeaa4037197c75587cd
title: "实现一个 `Queue`"
challengeType: 26
dashedName: implement-a-queue
---

# --description--

在本实验中，你将使用 `function` 实现一个 `queue` 数据结构。`queue` 是一种先进先出（FIFO）的数据结构，其中 `Element` 被添加到 `queue` 的末尾，并从前端移除。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `enqueue` 函数，将元素添加到队列的末尾。
1. 你应该有一个 `dequeue` 函数，用于移除并返回队列的前端元素。
1. 你应该有一个 `front` 函数，它返回队列的前端元素但不移除它。
1. 你应该有一个 `size` 函数，返回队列中元素的数量。
1. 你应该有一个 `isEmpty` 函数，如果队列为空则返回 `true`，否则返回 `false`。

# --hints--

你应该有一个 `enqueue` 函数。

```js
assert.isFunction(enqueue);
```

`enqueue` 函数应将一个元素添加到队列的末尾。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

你应该有一个 `dequeue` 函数。

```js
assert.isFunction(dequeue);
```

`dequeue` 函数应移除并返回队列的前端元素。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

当在空队列上调用 `dequeue` 时，应返回 `undefined`。

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` 应该在 `dequeues` 之后继续将元素添加到队列的末尾。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

不同的队列实例不应共享状态。

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

你应该有一个 `front` 函数。

```js
assert.isFunction(front);
```

`front` 函数应返回队列的前端元素且不移除它。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

当在空 `queue` 上调用时，`front` 应该返回 `undefined`。

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

你应该有一个 `size` 函数。

```js
assert.isFunction(size);
```

`size` 函数应返回队列中正确的元素数量。

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

你应该有一个 `isEmpty` 函数。

```js
assert.isFunction(isEmpty);
```

`isEmpty` 函数应为一个空队列 `return` `true`。

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

`isEmpty` 函数对于非空队列应返回 `false`。

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

`isEmpty` 函数应在队列中所有元素被移除后 `return` `true`。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
