---
id: 67ef8f695db52583424bd118
title: 設計一個電影評論頁面
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `main` 元素。
2. 在 `main` 元素中，你應該有一個用於電影標題的 `h1` 元素。
3. 在 `h1` 元素下方，你應該有一個用於顯示電影封面的 `img` 元素。你的 `img` 元素應該有描述圖片的 `alt` 文字。如果你願意，可以使用以下圖片：`https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`。
4. 你應該有一個包含簡短電影描述的 `p` 元素。
5. 你應該有另一個用於顯示電影評分的 `p` 元素。在其中，你應該按所列順序包含以下內容：
   - 一個帶有文字 `Movie Rating:` 的 `b` 元素。
   - 一個 `span` 元素，其 `aria-hidden` 屬性設定為 `true`，並使用星星 `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆` 直觀地表示評分。
   - 在該 span 之後，用括號表示評分的數值（例如 `9.2/10`）。
6. 你應該有一個帶有文字 `Cast Members` 的 `h2` 元素。
7. 你應該有一個 `ul` 元素。
8. 在 `ul` 元素中，你應該有多個 `li` 元素，每個元素都包含一個表示演員名字的 `b` 元素，後面接著以文字 `as` 為前綴的對應角色名。（例如，`James Holloway as Ethan Carter`）。

# --hints--

你應該有一個 `main` 元素。

```js
assert.isNotNull(document.querySelector("main"));
```

在 `main` 元素內，你應該有一個包含電影標題的 `h1` 元素。

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

在 `h1` 元素下方，你應該有一個帶有描述性 `alt` 屬性的 `img` 元素。
  
```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

你應該在包含電影描述的影像後面放置一個 `p` 元素。

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

你應該有另一個包含評分的 `p` 元素。

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

在第二個 `p` 元素內，你應該有一個 `b` 元素，文字為 `Movie Rating: `。

```js
const boldEl = document.querySelector("main p:nth-of-type(2) b");
assert.match(boldEl?.innerText.trim(), /^Movie Rating:?$/);
```

評分段落內應該有一個 `span` 元素。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

評分段落內的 `span` 元素應該有十顆星星，填滿的（⭐）或空的（☆），接著在 `span` 後面以小括弧顯示數值。

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

評分段落內的 `span` 元素應該設定 `aria-hidden` 屬性為 `true`。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

你應該有一個 `h2` 元素。

```js
assert.exists(document.querySelector("h2"));
```

你的 `h2` 應該位於第二個 `p` 元素之後。

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

你的 `h2` 應該有文字 `Cast Members`。

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

你應該在 `h2` 之後有一個 `ul` 元素。

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

`ul` 應該包含多個 `li` 元素。

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

每個 `li` 應該包含一個 `b` 元素表示演員名字，後面接著以 `as` 為前綴的角色名稱。（例如 `James Holloway as Ethan Carter`）。

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const bold = li.querySelector("b");
  assert.isNotNull(bold);
  const fullText = li.textContent.trim();
  const actorName = bold.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <b>Movie Rating:</b>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><b>James Holloway</b> as Ethan Carter</li>
        <li><b>Olivia Sterling</b> as Lena Mitchell</li>
        <li><b>William Lancaster</b> as Professor Adams</li>
        <li>
          <b>Maria Collins</b> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
