---
id: 67ef8f695db52583424bd118
title: 設計一個電影評價頁面
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. 你應該有一個 `main` 元素。
2. 在 `main` 元素內部，你應該有一個 `h1` 元素，用於顯示電影標題。
3. 在 `h1` 元素下面，你應該有一個 `img` 元素，用於顯示電影封面。 你的 `img` 元素應包含描述性 `alt` 文本描述圖像。 你可以隨意使用此圖片：`https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`。
4. 你應該有一個 `p` 元素，其中包含簡短的電影描述。
5. 你應該使用另一個 `p` 元素來顯示電影評分。 在其中，你應該按照所列順序排列這些項目：
   - 一個 `strong` 元素，文本爲 `Movie Rating:`。
   - 一個 `span` 元素，其 `aria-hidden` 屬性設爲 `true`，並且包含評分的可視化表示（例如星星 `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`）。
   - 在 span 之後，顯示一個括號括起來的數字評分值，（例如 `9.2/10`）。
6. 你應該有一個 `h2` 元素，文本爲 `Cast Members` 。
7. 你應該有一個 `ul` 元素。
8. 在 `ul` 元素內部，應該包含多個 `li` 元素，每個都應包含一個 `strong` 元素表示演員姓名，之後是對應的角色名，並且以 `as` 連接。 （例如，`James Holloway as Ethan Carter`）。

# --hints--

你應該有一個 `main` 元素。

```js
assert.isNotNull(document.querySelector("main"));
```

在 `main` 元素內部，你應該有一個包含電影標題的 `h1` 元素。

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

在 `h1` 元素下面，你應該有一個 `img` 元素，並且帶有描述性的 `alt` 屬性。

```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

你應該在圖片後添加一個 `p` 元素，其中包含電影描述。

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

你應該有另一個包含評分的 `p` 元素。

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

在第二個 `p` 元素中，你應該有一個 `strong` 元素，文本爲 `Movie Rating:`。

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

評分段落中應該有一個 `span` 元素。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

評分段落元素中的 `span` 元素，應當包含 10 個星星，可以是實心（⭐）或空心（☆），並且在 `span` 後顯示一個括號括起來的數字評分值。

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

評分段落元素中的 `span` 元素應該有一個 `aria-hidden` 屬性，設置爲 `true`。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

你應該有一個 `h2` 元素。

```js
assert.exists(document.querySelector("h2"));
```

你的 `h2` 元素應當位於第二個 `p` 元素之後。

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

你的 `h2` 元素的文本應爲 `Cast Members`。

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

你應該在 `h2` 之後有一個 `ul` 元素。

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

`ul` 元素應包含多個 `li` 元素。

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

每個 `li` 元素應包含一個 `strong` 元素表示演員姓名，並在其後用 `as` 連接角色名。 （例如，`James Holloway as Ethan Carter`）。

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
