---
id: 66e2e99caa0cb3570f4822b4
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora es el momento de probar tu función `getConsonantCount`.

Crea una variable `consonantCount` y asígnale el resultado de llamar a la función `getConsonantCount` con el argumento de `"Coding is fun"`.

Después de eso, registra lo siguiente en la consola: `"Consonant Count: [Consonant count goes here]"`. Reemplaza `[Consonant count goes here]` con el nombre real de la variable. Puede elegir usar cadenas de plantillas o concatenación de cadenas con el operador `+` aquí.

# --hints--

Debes crear una variable `consonantCount`.

```js
assert.isNotNull(consonantCount)
```

Tu variable `consonantCount` debe establecerse en el resultado de `getConsonantCount("Coding is fun")`.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

Debes registrar la variable `consonantCount` en la consola con el parámetro `"Consonant Count: ${consonantCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
