---
id: 66e2e5d890ac8a54c9d703bd
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Es hora de contar las consonantes. Crea una función de `getConsonantCount` con un parámetro de `sentence`.

Dentro de la función, usa un bucle para contar el número de consonantes en la `sentence` que se pasará a la función cuando sea llamada. Una consonante es cualquier letra que no sea uno de los siguientes caracteres: `"aeiou"`.

Su función de `getConsonantCount` debe devolver un número.

# --hints--

Deberías crear una función de `getConsonantCount`.

```js
assert.isFunction(getConsonantCount);
```

Su función de `getConsonantCount` debe tener un parámetro de `sentence`.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

Su función de `getConsonantCount` debería devolver un número.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

Cuando la oración es `"Coding is fun"`, la función `getConsonantCount` debería devolver `7`.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

Cuando la oración es `"Hello, World!"`, la función `getConsonantCount` debería devolver `7`.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

El conteo de consonantes debe ser insensible a las mayúsculas.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

Su función de `getConsonantCount` debería devolver el conteo correcto de consonantes para cualquier oración.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
