---
id: 66e2f06b191e305c00574e4d
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Finalmente, cuenta el número de palabras creando una función `getWordCount` con un parámetro `sentence`. La función debe devolver el número total de palabras en la oración cuando se llame.

# --hints--

Debes crear una función `getWordCount`.

```js
assert.isFunction(getWordCount);
```

Tu función `getWordCount` debe tener un parámetro `sentence`.

```js
assert.match(getWordCount.toString(), /sentence/);
```

Tu función `getWordCount` debe devolver un número.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

Cuando la oración sea `"When are you gonna start learning to code?"`, la función `getWordCount` debe devolver `8`.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

Cuando la oración sea `"What's going on?"`, la función `getWordCount` debe devolver `3`.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

El conteo de palabras debe ser insensible a mayúsculas y minúsculas.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

Tu función `getWordCount` debe devolver el conteo correcto de palabras para cualquier oración.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both frontend and backend development"), 7);
```

Tu función `getWordCount` debe devolver el conteo correcto de palabras para una cadena vacía o una cadena con solo espacios.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
