---
id: 66e2df8c5676fe51cead8ec8
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Ahora es el momento de probar tu función `getVowelCount`.

Crea una variable `vowelCount` y asígnale el resultado de llamar a la función `getVowelCount` con el argumento `"Apples are tasty fruits"`.

Después de eso, regístralo en la consola: `"Vowel Count: [vowel count goes here]"`. Reemplaza `[vowel count goes here]` con el nombre real de la variable. Puede elegir usar cadenas de plantillas o concatenación de cadenas con el operador `+` aquí.

# --hints--

Debes crear una variable `vowelCount`.

```js
assert.isNotNull(vowelCount)
```

Tu variable `vowelCount` debe establecerse en el resultado de `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

Debes registrar la variable `vowelCount` en la consola con el parámetro `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
