---
id: 66e2eab8a5638f57b637b7cc
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Debe contar el número de puntuaciones ahora.

Crea una función `getPunctuationCount` con un parámetro `sentence`.

Dentro de la función, cree un ciclo para contar el número de caracteres de puntuación en `sentence` que se pasará a la función cuando se llame. Para nuestros propósitos, un carácter de puntuación es cualquier carácter que no sea un espacio (`" "`) o una letra.

Su función `getPunctuationCount` debe devolver un número.

# --hints--

Debe crear una función `getPunctuationCount`.

```js
assert.isFunction(getPunctuationCount);
```

Su función `getPunctuationCount` debería tener un parámetro `sentence`.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

Su función `getPunctuationCount` debe devolver un número.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

Cuando la oración sea `"What's going on here?"`, la función `getPunctuationCount` debe devolver `2`.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

Cuando la oración sea `"What????!"`, la función `getPunctuationCount` debe devolver `5`.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

Su función `getPunctuationCount` debe devolver el recuento correcto de puntuación para cualquier oración.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
