---
id: 5dc1798ff86c76b9248c6eb3
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Під елементом `h1` додайте елемент `h2` з таким текстом:

`Cat Photos`

# --hints--

Ваш елемент `h1` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Ваш елемент `h1` має мати кінцевий тег. Кінцеві теги мають такий синтаксис: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

У вас має бути лише один елемент `h1`. Видаліть зайвий.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Текст вашого елемента `h1` має бути 'CatPhotoApp'. Ви або пропустили текст, або зробили помилку.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

Ваші теги `h2` мають бути в нижньому регістрі. За конвенцією всі HTML-теги пишуться в нижньому регістрі.

```js
assert.notMatch(code, /<\/?H2>/);
```

Ваш елемент `h2` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Ваш елемент `h2` має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/h2\>/);
```

Текст вашого елемента `h2` має бути `Cat Photos`. Поміщайте текст `Cat Photos` лише між початковим і кінцевим тегами `h2`.

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

Ваш елемент `h2` має бути під елементом `h1`. Елемент `h1` має більший пріоритет і повинен бути над елементом `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
