---
id: 5dc17dc8f86c76b9248c6eb5
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Коментування дозволяє залишати повідомлення, не впливаючи на відображення в браузері. Це також дає змогу зробити код неактивним. Коментар у HTML починається з `<!--`, містить будь-яку кількість рядків тексту і закінчується `-->`.

Ось приклад коментаря з `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Додайте коментар над елементом `p` з таким текстом:

`TODO: Add link to cat photos`

# --hints--

Ваш коментар має починатися з `<!--`. У вас відсутній один або кілька символів, які визначають початок коментаря.

```js
assert.match(code, /<!--/);
```

Ваш коментар має закінчуватися `-->`. У вас відсутній один або кілька символів, які визначають кінець коментаря.

```js
assert.match(code, /-->/);
```

У вашому коді не повинно бути зайвих символів відкриття/закриття коментаря. У браузері відображається зайвий `<!--` або `-->`.

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g)?.length, 2)
assert.isBelow(noSpaces.match(/-->/g)?.length, 2);
```

Ваш коментар має бути над елементом `p`. У вас вони розташовані в неправильному порядку.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

Ваш коментар має містити текст `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
  </body>
</html>
```
