---
id: 5dc17d3bf86c76b9248c6eb4
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Створіть елемент `p` під вашим елементом `h2` і додайте йому такий текст:

`Everyone loves cute cats online!`

# --hints--

Ваш елемент `p` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

Ваші теги `p` мають бути в нижньому регістрі. За конвенцією всі HTML-теги пишуться в нижньому регістрі.

```js
assert.notMatch(code, /<\/?P>/);
```

Ваш елемент `p` має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/p\>/);
```

Текст вашого елемента `p` має бути `Everyone loves cute cats online!`. Ви або пропустили текст, або зробили помилку.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Ваш елемент `p` має бути розташований під елементом `h2`. Ви розташували їх у неправильному порядку.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
