---
id: 5dc174fcf86c76b9248c6eb2
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви продовжите працювати з базовими HTML-елементами, такими як заголовки, абзаци та списки, створюючи застосунок для фото котів.

Почніть практичне заняття з додавання елемента `h1` з текстом `CatPhotoApp`.

# --hints--

Текст `CatPhotoApp` має бути присутнім у коді. Варто перевірити орфографію.

```js
assert.match(code, /catphotoapp/i);
```

Ваш елемент `h1` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Ваші теги `h1` мають бути в нижньому регістрі. За конвенцією всі HTML-теги пишуться в нижньому регістрі.

```js
assert.notMatch(code, /<\/?H1>/);
```

Ваш елемент `h1` повинен мати кінцевий тег. Кінцеві теги мають такий синтаксис: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Текст елемента `h1` має бути `CatPhotoApp`. Ви або пропустили текст, або зробили помилку, або він не знаходиться між початковим і кінцевим тегами елемента `h1`.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
