---
id: 5dc24073f86c76b9248c6ebb
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

<dfn>HTML-атрибути</dfn> — це спеціальні слова, які використовують в межах початкового тега елемента, щоб керувати його поведінкою. Атрибут `src` в елементі `img` вказує посилання зображення (де розташоване зображення).

Ось приклад елемента `img` з атрибутом `src`, який направляє на логотип freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

До наявного елемента `img` додайте атрибут `src` з таким посиланням:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Код повинен мати елемент `img`. Можливо, ви видалили елемент `img` або значення атрибута `src` не в лапках.

```js
assert.exists(document.querySelector('img'));
```

Елемент `img` повинен мати атрибут `src`. Ви або не додали атрибут, або маєте друкарську помилку. Переконайтеся, що між назвою елемента та назвою атрибута є пробіл.

```js
assert.exists(document.querySelector('img')?.src);
```

Значенням атрибута `src` елемента `img` має бути '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Ви або не написали посилання, або маєте друкарську помилку. Зверніть увагу на регістр.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Хоча ви встановили значення атрибута `src` елемента `img` як правильне посилання, значення атрибута краще писати в лапках.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Елемент `img` має бути закритим за допомогою `>` або `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
