---
id: 5dc24165f86c76b9248c6ebc
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Усі елементи `img` повинні мати атрибут `alt`. Текст атрибута `alt` використовується читачами екрана для покращення доступності та відображається, якщо зображення не завантажується.

Ось приклад елемента `img` з атрибутом `alt`:

```html
<img src="cat.jpg" alt="A cat">
```

Всередині елемента `img` додайте атрибут `alt` з таким текстом:

`A cute orange cat lying on its back`

# --hints--

Ваш код повинен містити елемент `img`. Ви видалили елемент `img` на попередньому кроці.

```js
assert.exists(document.querySelector('img'));
```

Ваш елемент `img` не має атрибута `alt`. Перевірте, чи є пробіл після назви початкового тегу та/або чи є пробіли перед усіма назвами атрибутів.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

Значення атрибута `img` елемента `alt` встановлено не як `A cute orange cat lying on its back`. Переконайтеся, що значення атрибута `alt` взято в лапки.

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
