---
id: 5dc23991f86c76b9248c6eb8
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

У попередньому кроці ви розмістили елементи `h1`, `h2`, коментар і `p` всередині елемента `main`. Це називається *вкладенням*. Вкладені елементи слід розміщувати на два пробіли правіше від елемента, у який вони вкладені. Такий відступ називається індентацією і використовується, щоб зробити HTML легшим для читання.

Ось приклад вкладення та індентації:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Елементи `h1`, `h2` і коментар відступають на два пробіли більше, ніж елемент `main` у наведеному нижче коді. Використайте пробіл на клавіатурі, щоб додати два пробіли перед елементом `p`, щоб він також мав правильний відступ.

# --hints--

Ваш код має містити елемент `h2` з текстом `Cat Photos`. Можливо, ви випадково його видалили, він відсутній як початковий, так і кінцевий тег, або текст змінився.

```js
assert.exists(document.querySelector('h2'));
assert.match(code, /<\/h2\>/);
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(),'cat photos')
```

Вам не слід додавати елементи `ul` або `li` з прикладу.

```js
assert.notExists(document.querySelector('ul'));
assert.notExists(document.querySelector('li'));
```

Вам не слід змінювати індентацію на рядку з вашим елементом `h2`. Його початковий тег має починатися на 6 пробілів від початку рядка. Ви можете перезапустити крок, щоб відновити початкову індентацію.

```js
assert.match(code.toLowerCase(), /<\/h1\>\s*\n\s{6}<h2>/);
```

Ваш код має містити коментар. Ви видалили коментар з попереднього кроку.

```js
assert.match(code, /<!--.*-->/);
```

Текст коментаря має бути `TODO: Add link to cat photos`. Не змінюйте текст або відступ коментаря.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i);
```

Вам не слід змінювати індентацію на рядку з вашим коментарем. Його початковий тег має починатися на 6 пробілів від початку рядка. Ви можете перезапустити крок, щоб відновити початкову індентацію.

```js
assert.match(
  code
    .toLowerCase(),
    /<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/
);
```

Ваш код має містити елемент `p`. Ви видалили елемент `p` з попереднього кроку.

```js
assert.exists(document.querySelector('p'));
```

Текст елемента `p` має бути `Everyone loves cute cats online!`. Не змінюйте текст, відступи або пунктуацію елемента `p`.

```js
assert.match(
  document
    .querySelector('p')
    ?.innerText?.trim().toLowerCase(),
    /^everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

Початковий тег `p` має мати індентацію, що відповідає вашим елементам `h2` і коментарю. Його початковий тег має починатися на 6 пробілів від початку рядка.

```js
assert.match(code.toLowerCase(), /-->\s*\n\s{6}<p>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
    </main>
  </body>
</html>
```
