---
id: 5dc2385ff86c76b9248c6eb7
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 має кілька елементів, які ідентифікують різні області вмісту. Ці елементи роблять ваш HTML легшим для читання та допомагають із пошуковою оптимізацією (SEO) і доступністю.

Елемент `main` використовується, щоб представити основний вміст тіла HTML-документа. Вміст всередині елемента `main` має бути унікальним для документа і не повинен повторюватися в інших частинах документа.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Визначте основний розділ цієї сторінки, додавши початковий тег `<main>` перед елементом `h1` і кінцевий тег `</main>` після елемента `p`.

# --hints--

Ваш елемент `main` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

Ваші теги `main` мають бути в нижньому регістрі. За конвенцією всі HTML-теги пишуться в нижньому регістрі.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

Ваш елемент `main` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/main\>/);
```

Початковий тег вашого елемента `main` має бути нижче початкового тегу елемента `body`. Ви розташували їх у неправильному порядку.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Початковий тег вашого елемента `main` має бути вище за елемент `h1`. Ви розташували їх у неправильному порядку.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

Кінцевий тег вашого елемента `main` має бути нижче за елемент `p`. Ви розташували їх у неправильному порядку.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
