---
id: 5dc2385ff86c76b9248c6eb7
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 має деякі елементи, що позначають різні області вмісту. Ці елементи полегшують прочитання HTML, а також допомагають з оптимізацією пошукової системи (SEO) та доступністю.

Елемент `main` використовують, щоб позначити основний вміст тіла HTML-документа. Вміст в межах елемента `main` має бути унікальним та не повторюватися в інших частинах документа.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Визначте головний розділ цієї сторінки, додавши початковий тег `<main>` перед елементом `h1` та кінцевий тег `</main>` після елемента `p`.

# --hints--

Елемент `main` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<назваЕлемента>`.

```js
assert.exists(document.querySelector('main'));
```

Your `main` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

Елемент `main` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/main\>/);
```

Початковий тег елемента `main` має бути під початковим тегом елемента `body`. Вони знаходяться в неправильному порядку.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Початковий тег елемента `main` має бути над елементом `h1`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

Кінцевий тег елемента `main` має бути під елементом `p`. Вони знаходяться в неправильному порядку.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```
