---
id: 5dfa3589eacea3f48c6300ae
title: Крок 20
challengeType: 0
dashedName: step-20
---

# --description--

У другому `section` елементі додайте новий `h2` елемент із текстом `Cat Lists`.

# --hints--

Ваш `section` елемент має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Ваш `h2` елемент має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

Ваш `h2` елемент має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

Ваш другий `h2` елемент має бути безпосередньо над кінцевим тегом другого `section` елемента. Він розташований не на правильному місці.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection?.lastElementChild?.nodeName, 'H2');
```

Другий `h2` елемент має містити текст `Cat Lists`. Ви або пропустили текст, або зробили помилку.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
