---
id: 5dfb6250eacea3f48c6300b2
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Після невпорядкованого списку додайте нове зображення з атрибутом `src`, значення якого встановлено на:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

А його атрибут `alt` має мати значення:

`A slice of lasagna on a plate.`

# --hints--

Після закриваючого тегу `img` має бути елемент `</ul>`.

```js
assert.equal(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

Нове зображення не має атрибута `alt`. Перевірте, чи є пробіл після назви початкового тегу і/або чи є пробіли перед усіма назвами атрибутів.

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

Нове зображення має мати значення `alt` рівне `A slice of lasagna on a plate.`. Переконайтеся, що значення атрибута `alt` взяте в лапки.

```js
assert.match(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

Нове зображення не має атрибута `src`. Перевірте, чи є пробіл після назви початкового тегу і/або чи є пробіли перед усіма назвами атрибутів.

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

Нове зображення має мати значення `src` рівне `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. Переконайтеся, що значення атрибута `src` взяте в лапки.

```js
assert.equal(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Хоч ви й встановили атрибут `src` нового зображення на правильну URL-адресу, рекомендується завжди брати значення атрибута в лапки.

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
