---
id: 5ef9b03c81a63668521804ea
title: Крок 38
challengeType: 0
dashedName: step-38
---

# --description--

Зверніть увагу, що всі додані елементи сторінки знаходяться в межах елемента `body`. Всі елементи вмісту сторінки, які мають відтворюватися на сторінці, пишуться в елементі `body`. А ось інша важлива інформація пишеться в межах елемента `head`.

Елемент `head` використовують, щоб зберігати метадані про документ (серед яких його назва, посилання на таблиці стилів та скрипти). Метадані — це інформація про сторінку, яка не відтворюється на самій сторінці.

Додайте елемент `head` над елементом `body`.

# --hints--

Ви або видалили елемент `body`, або в ньому відсутній початковий чи кінцевий тег.

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

Елемент `head` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.match(code, /\<head\>/);
```

Елемент `head` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /\<\/head\>/);
```

Елемент `head` має бути над початковим тегом елемента `body`. Ви розташували його десь в іншому місці.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--

--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
