---
id: 5ef9b03c81a63668521804e9
title: Крок 37
challengeType: 0
dashedName: step-37
---

# --description--

Перетворіть існуючий текст `freeCodeCamp.org` на посилання, вклавши його в елемент посилання (`a`). Атрибут `href` має бути встановлений у значення `https://www.freecodecamp.org`.

# --hints--

Ваш елемент посилання (`a`) має бути вкладений у елемент `p` всередині `footer`. Обов’язково додайте початковий тег і кінцевий тег для елемента посилання (`a`).

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

Ваш елемент посилання (`a`) має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

Ваш елемент посилання (`a`) має мати атрибут `href` зі значенням `https://www.freecodecamp.org`. Можливо, ви пропустили атрибут/значення або допустили помилку.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

Текст посилання має бути `freeCodeCamp.org`. Ви або пропустили текст, або допустили помилку.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

Після вкладення елемента посилання (`a`) єдиним вмістом елемента `p`, видимим у браузері, має бути `No Copyright - freeCodeCamp.org`. Перевірте текст, пробіли або пунктуацію як у `p`, так і у вкладеному елементі посилання.

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
