---
id: 5f07c98cdb9413cbd4b16750
title: Крок 19
challengeType: 0
dashedName: step-19
---

# --description--

Час додати новий розділ. Додайте другий `section` елемент під існуючим `section` елементом.

# --hints--

Ваш `section` елемент має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

Ви повинні додати лише один початковий `section` тег. Будь ласка, видаліть усі зайві.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

Ваш `section` елемент має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

Ви повинні додати лише один кінцевий `section` тег. Будь ласка, видаліть усі зайві.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Другий `section` елемент не повинен бути вкладений у перший `section` елемент.

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Обидва `section` елементи мають бути між початковим і кінцевим тегами `main` елемента.

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
