---
id: 5f07be6ef7412fbad0c5626b
title: Крок 18
challengeType: 0
dashedName: step-18
---

# --description--

Перед тим як додати будь-який новий вміст, слід використати елемент `section`, щоб відокремити вміст із фотографіями котів від майбутнього вмісту.

Елемент `section` використовується для визначення розділів у документі, таких як глави, заголовки, нижні колонтитули або будь-які інші розділи документа. Це семантичний елемент, який допомагає з SEO та доступністю.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Візьміть свій елемент `h2`, два елементи `p` та елемент посилання (`a`) і вкладьте їх у елемент `section`.

# --hints--

Ваш елемент `section` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector('section'));
```

Ваш елемент `section` має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert.match(code, /<\/section\s*>/i);
```

Весь елемент `section` має бути між початковим і кінцевим тегами елемента `main`.

```js
assert.equal(document.querySelector('section')?.parentNode.nodeName, 'MAIN');
```

Існуючі елементи `h2`, `p` та елемент посилання (`a`) мають бути між початковим і кінцевим тегами елемента `section`.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4)
```

Елемент `h1` не повинен бути вкладеним у елемент `section`.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      
      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
