---
id: 66fe4f33a2cc9b33f4d5cd9b
title: بناء مولد التباديل
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

في هذا المختبر، ستبني مولد تباديل يأخذ سلسلة نصية ويُرجع كل التباديل الممكنة لأحرف السلسلة. على سبيل المثال، التباديل الممكنة للسلسلة النصية `cat` هي `cat`، `cta`، `act`، `atc`، `tac`، و `tca`.

تعمل الطريقة العودية لإنشاء التباديل لسلسلة نصية عن طريق تخزين الجزء الثابت في بداية السلسلة (البادئة)، وإنشاء التباديل لبقية السلسلة.

على سبيل المثال، اعتبر كلمة `machine`. الجولة الأولى من إنشاء التباديل تتم بتثبيت الحرف `m` كبادئة للسلسلة، ثم إنشاء التباديل لبقية السلسلة، وهي `achine`.

لبقية السلسلة، تستمر التباديل بنفس الطريقة. يُضاف حرف واحد إلى البادئة، ربما `c`، فتصبح البادئة `mc`. ثم تُلصق كل تباديل `ahine` إلى البادئة.

يستمر هذا حتى تحتوي البادئة على كل الأحرف، وتصبح بقية السلسلة فارغة، وهذا يعني أنه تم إنشاء تبادل واحد.

**الهدف:** حقق قصص المستخدم أدناه واجعل كل الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `permuteString`.
2. يجب أن تأخذ دالة `permuteString` معلمة واحدة، وهي سلسلة نصية، ثم معلمتين بقيمة افتراضية: قيمة بادئة ومصفوفة فارغة لتخزين النتائج وإرجاعها. يجب أن تكون قيمة البادئة سلسلة نصية تُستخدم لتجميع الأحرف لتكوين تبادل. ستُستدعى الدالة بمعلمة واحدة فقط، مثل `permuteString("cat")`.
3. داخل الدالة، تحقق مما إذا كان طول السلسلة الممررة هو `0`. إذا كان كذلك، أضف البادئة الحالية إلى النتائج وأرجع النتائج.
4. كرر على كل حرف في السلسلة المُدخلة ولكل تكرار، أزل الحرف الحالي من السلسلة واستدعِ دالة `permuteString` بشكل عودي مع المعلمات المحدثة لبناء التباديل المتبقية.
5. يجب أن تُرجع مصفوفة النتائج النهائية.
6. يجب أن تضمن أن التباديل فريدة عن طريق إزالة التكرارات.

# --hints--

يجب أن يكون لديك دالة `permuteString`.

```js
assert.isFunction(permuteString);
```

يجب أن تستخدم العودية في دالة `permuteString` الخاصة بك.

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` يجب أن تُرجع `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`.

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` يجب أن تُرجع `[ "fcc", "cfc", "ccf" ]`.

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` يجب أن تُرجع `[ "p" ]`.

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` يجب أن تُرجع `[""]`.

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` يجب أن تُرجع `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`.
`.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` يجب أن تُرجع النتائج الصحيحة.

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
