---
id: 66ec4c8e9878d8441956516f
title: 創建一個圖書目錄表格
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. 你應該創建一個用於列出圖書信息的 `table` 元素。
1. 你的表中應該有一個一行的表頭元素。
1. 你的表頭元素中的行應該有四個表標題元素，文本依次是 `Title`、`Author`、`Genre` 和 `Publication Year`。
1. 你的表中應該有一個至少是五行的表體元素。
1. 你的表體中的每一行都應該有四個表數據元素，顯示圖書的名稱、作者、體裁和出版年份。
1. 你的表格應該有一個一行的表腳元素。
1. 你的表腳元素中的行應當包含一個表格數據單元格元素，該元素需橫跨四列，並且顯示文本 `Total Books: N` ，其中 `N` 應替換爲你的表格中的圖書數量。

# --hints--

你應該有一個 `table` 元素。

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

在你的 `table` 元素中，應該有一個 `thead` 元素。

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

在你的 `thead` 元素內，應該有一個 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

在你的 `thead` 元素的行中，應該有四個 `th` 元素。

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

你這四個 `th` 元素應依次顯示文本 `Title`、`Author`、`Genre` 和 `Publication Year` 。

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent, 'Title');
assert.equal(ths[1]?.textContent, 'Author');
assert.equal(ths[2]?.textContent, 'Genre');
assert.equal(ths[3]?.textContent, 'Publication Year');
```

在你的 `table` 元素中，應該有一個 `tbody` 元素。

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

你的 `tbody` 元素至少應包含五行內容。

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

你的 `tbody` 元素中的每一行都應恰好有四個 `td` 元素作爲其子元素。

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

你的表體中的每個 `td` 元素都應包含有關圖書的信息。

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.length, 1);
});
```

在你的 `table` 元素內，應當有一個 `tfoot` 元素。

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

在你的 `tfoot` 元素中，應該恰好只有一個 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

你的 `tfoot` 元素行中的 `td` 元素應將其 `colspan` 屬性設置爲 `4` 。

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

The `td` element in your `tfoot` element's row should have the text `Total Books: N` where `N` is the number of books in your table.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
