---
id: 66ec4c8e9878d8441956516f
title: 建置書籍目錄表
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個 `table` 元素來列出書籍資訊。
1. 你的表格應該有一個包含一列的表頭元素。
1. 你的表格標題元素中的列應該有四個表格標題元素，且依序包含 `Title`、`Author`、`Genre` 和 `Publication Year` 的文字。
1. 你的表格應該有一個表格主體元素，且其中至少包含五列。
1. 你的表格主體中的每一列應該有四個表格資料元素，分別顯示書籍的 Title、Author、Genre 和 Publication Year。
1. 你的表格應該有一個包含一列的表格頁腳元素。
1. 你的表格頁腳元素中的列應該有一個跨越四個行的表格資料元素，並且包含文字 `Total Books: N`，其中 `N` 應該被你的表格中書籍的數量取代。

# --hints--

你應該有一個 `table` 元素。

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

你應該在你的 `table` 元素中有一個 `thead` 元素。

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

你的 `thead` 元素中應該有一個 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

你應該在你的 `thead` 元素的列中有四個 `th` 元素。

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

你的四個 `th` 元素應該依序有文字 `Title`、`Author`、`Genre` 和 `Publication Year`。

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent.trim(), 'Title');
assert.equal(ths[1]?.textContent.trim(), 'Author');
assert.equal(ths[2]?.textContent.trim(), 'Genre');
assert.equal(ths[3]?.textContent.trim(), 'Publication Year');
```

你應該在你的 `table` 元素中有一個 `tbody` 元素。

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

你的 `tbody` 元素應該至少有五列。

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

你的 `tbody` 元素中的每一列應該恰好有四個 `td` 元素作為子元素。

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

你的表格主體中的每個 `td` 元素都應該有包含書籍資訊的文字。

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.trim().length, 1);
});
```

你應該在你的 `table` 元素中有一個 `tfoot` 元素。

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

你的 `tfoot` 元素中應該恰好有一個 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

你在 `tfoot` 元素的列中的 `td` 元素應將其 `colspan` 屬性設為 `4`。

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

你在 `tfoot` 元素的 `td` 元素所在的列中，應該有文字 `Total Books: N`，其中 `N` 是你表格中的書籍數量。

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
