---
id: 65578cb031cd93b77a285db2
title: Paso 38
challengeType: 20
dashedName: step-38
---

# --description--

Tu función va a explorar todos los nodos conectados al nodo incial. Va a calcular el camino más corto para todos ellos. Luego, eliminará el nodo inicial de los nodos no visitados. 

Después, el nodo vecino más cercano será visitado y el proceso se repite hasta que todos los nodos son visitados.

A partir de ahora, vas a trabajar en el bucle principal que explora los nodos en el grafo. Para evitar problemas como causar un bucle infinito durante el desarrollo del algoritmo, comenta tu llamada a la función.

# --hints--

Debes comentar tu llamada a la función.

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
