---
id: 655774d01daeeaa1978b99d5
title: Paso 29
challengeType: 20
dashedName: step-29
---

# --description--

Al principio, se considera que todos los demás nodos en el grafo están a una distancia infinita del nodo fuente, porque la distancia aún no ha sido determinada.

Crea una cláusula `else` y asigna un valor infinito al nodo en el diccionario `distances`. Para eso, usa la función `float()` con la cadena `'inf'` como argumento para generar un número de punto flotante que representa el infinito positivo.

# --hints--

Debes tener una cláusula `else`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

Debes asignar `float('inf')` a `distances[node]` dentro de tu nueva cláusula `else`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
