---
id: 655773b0591c5f9f4045883e
title: Pasos 24
challengeType: 20
dashedName: step-24
---

# --description--

El algoritmo comenzará en un nodo especificado. Luego explorará el grafo para encontrar el camino más corto entre el nodo de inicio, o *source*, y todos los demás nodos.

Para eso, tu función necesita dos parámetros: `graph` y `start`. Agrégalos a la declaración de tu función.

# --hints--

Tu función debe tomar `graph` y `start` como parámetros, en este orden.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
