---
id: 65576ff7888f9e96f52a4be1
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Los diccionarios almacenan datos en forma de pares *clave*-*valor*. Una clave está separada del valor correspondiente por dos puntos. Y cada par clave-valor está separado del siguiente par por una coma:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Agrega un nuevo par clave-valor a tu diccionario. Usa la cadena `'species'` como clave y la cadena `'guinea pig'` como valor.

# --hints--

Debes tener una clave `'species'` con el valor `'guinea pig'` dentro de tu diccionario `copper`.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

Tu diccionario `copper` debe tener un solo par clave-valor.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
