---
id: 6557709b0aee699a6a00528c
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

Por ejemplo, un grafo se puede emplear para representar dos puntos en el espacio, `A` and `B`, conextados por un camino. Un grafo como este constará de dos nodos conectados por una arista.

Reemplaza la clave existente `'species'` con las cadenas `'A'`. Luego, reemplaza el valor correspondiente con la cadena `'B'` para representar la conexión entre los nodos `'A'` y `'B'`.

# --hints--

Tu diccionario debe tener una clave `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Tu tecla `'A'` debería tener `'B'` como valor.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Tu diccionario debería tener dos llaves.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
