---
id: 65577236b056379d5dbc7000
title: Pasos 23
challengeType: 20
dashedName: step-23
---

# --description--

Ahora vas a comenzar a desarrollar el algoritmo para calcular el camino más corto entre cada nodo en tu nuevo grafo.

Declara una función vacía llamada `shortest_path`. Usa la palabra clave `pass` para completar el cuerpo de la función. 

# --hints--

Debes definir una función llamada `shortest_path`.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
