---
id: 65577791ad8c26a7705e2919
title: Paso 34
challengeType: 20
dashedName: step-34
---

# --description--

Con una comprensión de diccionario, puedes crear un diccionario a partir de un diccionario existente:

```py
{key: val for key in dict}
```

En el ejemplo anterior, `val` es el valor que `key` tendrá en el nuevo diccionario, y `dict` es el diccionario existente.

Quieres hacer un seguimiento de los caminos entre el nodo inicial y cada otro nodo.

Después de la variable `distances`, crea una variable `paths` y asígnale un diccionario con todas las claves de `graph`. Asigna una lista vacía a cada clave y usa una comprensión de diccionario para construir tu diccionario.

# --hints--

Debes tener una variable llamada `paths`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

Tu variable `paths` debe usar la sintaxis de comprensión de diccionarios para asignar una lista vacía a cada nodo en el grafo.

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
