---
id: 65577a17564ce8a8e06c1460
title: Paso 35
challengeType: 20
dashedName: step-35
---

# --description--

La compresión de diccionarios soporta la sintaxis condicional `if`/`else` también:

```py
{key: val_1 if condition else val_2 for key in dict}
```

En el ejemplo anterior, `dict` es el diccionario existente. Cuando `condition` se evalúa como `True`, `key` tendrá el valor `val_1`, de lo contrario `val_2`.

Usa una comprensión de diccionario para crear un diccionario basado en `graph` y asígnalo a la variable `distances`. Dale a la clave un valor de cero si el nodo es igual al nodo inicial, e infinito en caso contrario. Usa `float('inf')` para lograr esto último.

# --hints--

Debe usar la sintaxis de compresión de diccionarios para dar un valor a la variable `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
