---
id: 6557712d77ce2d9bd7e63afd
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

Agrega otro nodo conectado a `B` a tu grafo y llámalo `C`.

Modifica tu diccionario existente para representar esta estructura: agrega otra clave `'C'` a `my_graph` y asígnale el valor de la cadena `'B'`.

Además, cambia el valor de la clave existente `'B'` a la lista `['A', 'C']` para representar las múltiples conexiones de tu nodo `'B'`.

# --hints--

Tu diccionario debe tener 3 claves — `'A'`, `'B'` y `'C'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

El valor de `my_graph['A']` debe ser la cadena `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` debería ser una lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

El valor de `my_graph['B']` debe ser una lista que contenga `'A'` y `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

El valor de `my_graph['C']` debe ser la cadena `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
