---
id: 6557746aad2844a0cd864e12
title: Paso 27
challengeType: 20
dashedName: step-27
---

# --description--

Mientras el algoritmo explora el grafo, debe mantener un registro de la distancia más corta conocida actualmente entre el nodo inicial y los otros nodos.

Antes de tu ciclo `for`, crea una nueva variable llamada `distances` y asígnale un diccionario vacío.

# --hints--

Debes tener una variable llamada `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

Tu variable `distances` debe ser un diccionario vacío.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
