---
id: 655774955b097ea14897db12
title: Paso 28
challengeType: 20
dashedName: step-28
---

# --description--

La distancia desde el nodo inicial es cero, porque el algoritmo comienza su evaluación desde ahí mismo.

Después de agregar `node` a `unvisited` en tu ciclo, crea una estructura `if` que se active si el nodo es igual al nodo inicial. Luego asigna `0` a ese nodo dentro del diccionario `distances`.

# --hints--

Debes crear una estructura `if` que se ejecute cuando `node` sea igual a `start`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

Dentro de tu nueva estructura `if` debes asignar `0` al nodo en el diccionario `distances`.

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
