---
id: 65579228c669fcbebffd01d5
title: Paso 52
challengeType: 20
dashedName: step-52
---

# --description--

El otro error es sutil. Cuando se encuentra una distancia más corta para un nodo vecino, `paths[current]` se asigna al camino del nodo vecino, `paths[node]`. 

Esto significa que ambas variables apuntan a la misma lista. Dado que las listas son mutables, cuando agregas el nodo vecino a su camino, tanto `paths[node]` como `paths[current]` se modifican porque son la misma lista. Esto resulta en caminos incorrectos, aunque las distancias son correctas.

Puedes corregir ese error asignando una copia de `paths[current]` a la ruta del nodo vecino. Para eso puedes usar la sintaxis slice:

```py
my_list[:]
```

Donde `my_list` es la lista que quieres copiar. Modifica la asignación existente `paths[node] = paths[current]` dentro de tu bloque `if` cortando `paths[current]`.

# --hints--

Debes asignar `paths[current][:]` a la ruta del nodo vecino.

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
