---
id: 6566195b0a021bb660b2b4b1
title: Paso 21
challengeType: 20
dashedName: step-21
---

# --description--

Ahora modifica `my_graph['B']` en una lista de tuplas, donde el primer elemento de la tupla es el nodo conectado y el segundo elemento es la distancia. La distancia `B-C` es `4`.

# --hints--

`my_graph['B']` debería ser una lista que contenga las tuplas `('A', 3)` y `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph` debería tener 4 claves llamadas `'A'`, `'B'`, `'C'` y `'D'`.

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
