---
id: 6578b57361f2f132a02e2a18
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Puedes acceder a los datos almacenados en un diccionario a través de sus claves:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

Después de tu diccionario, sigue el ejemplo anterior para acceder a la clave `'species'` de `copper` e imprimir el resultado.

# --hints--

No debes modificar la asignación de tu diccionario.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Debes usar `copper['species']` para acceder al valor de la clave `'species'`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

Debes llamar a `print()` pasando `copper['species']` como argumento.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
