---
id: 65789506b30453080f77470c
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

Hasta ahora, ya has conocido diferentes tipos de datos:

- Tipos de datos inmutables, como enteros, cadenas, tuplas y Booleanos.
- Tipos de datos mutables, como listas y diccionarios.

Un diccionario se identifica por un par de llaves, `{}`.

Comienza creando una variable llamada `copper` y asígnale un diccionario vacío usando un par de llaves, de la misma manera que crearías una lista vacía con un par de corchetes.

# --hints--

Deberías tener una variable llamada `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

Tu variable `copper` debe tener el valor de un diccionario vacío. Usa un par de llaves para eso.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
