---
id: 6578b13757611e2825beb8a5
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Las claves deben ser únicas dentro de un diccionario y sólo pueden ser de un tipos de dato inmuables. Esto significa que no puede utilizar una lista u otro diccionario como claves.

Agrega otra clave `'age'` a tu diccionario y asígnale el número entero `2` como valor.

# --hints--

Debes tener una clave `'age'` con el valor `2` justo después de `'species': 'guinea pig'`, dentro de tu diccionario `copper`. No olvides la coma.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Debes tener dos pares clave-valor en tu diccionario `copper`.

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
