---
id: 6579dd49fa8a8e1fd06b85a9
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Los grafos son estructuras de datos que representan relaciones entre pares de elementos. Estos elementos, llamados *nodos*, pueden ser objetos del mundo real, entidades, puntos en el espacio u otros. La conexiones entre nodos se conocen como *aristas*.

Aquí tienes una representación visual de un grafo:

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


Renombra el diccionario `copper` como `my_graph`. Esto representará el grafo para probar el algoritmo.

# --hints--

Debe renombrar el diccionario `copper` como `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

La variable `my_graph` debe ser un diccionario.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
