---
id: 6579ca0923cfa7162089d2f0
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

La misma sintaxis se puede utilizar para cambiar el valor de una clave existente.

Justo antes de la llamada a `print()`, accede a la clave `'species'` y reasigna su valor a `'Cavia porcellus'`.

# --hints--

Debes reasignar `copper['species']` a la cadena `'Cavia porcellus'` antes de la llamada a `print()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
