---
id: 65796fac81f983127558f3f4
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Ahora, modifica tu llamada existente a `print()` para imprimir el valor de la clave `age`.


# --hints--

Debes usar `copper['age']` para acceder al valor de la clave `age`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

Debes llamar a `print()` pasando `copper['age']` como argumento.

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
