---
id: 6579cbab9825b8170974c69a
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Puedes eliminar un par clave-valor de un diccionario usando la palabra clave `del`:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

Justo antes de tu ciclo `for`, usa la palabra clave `del` para eliminar la clave `'age'` y su valor de `copper`.

# --hints--

Debes usar la palabra clave `del` para eliminar `copper['age']` antes del ciclo `for`.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
