---
id: 6579717f0920131304286804
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Para agregar un nuevo par clave-valor después de declarar un diccionario, puedes indicar la clave de la misma manera que accederías a una clave existente, y establecer el valor de la nueva clave usando el operador de asignación:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

Elimina tu llamada a `print()`. Luego, después de la declaración de `copper`, agrega la clave `'food'` a tu diccionario y asigna su valor a `'hay'`.

# --hints--

No deberías tener `print(copper['age'])` en tu código.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

Debes agregar la clave `'food'` a `copper` después de declarar el diccionario.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

Debes establecer `copper['food']` a `'hay'` después de declarar el diccionario.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
