---
id: 6579cd5f6dd62c189e53ddbb
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

Ahora que revisaste los aspectos básicos de los diccionarios, puedes proceder a construir el algoritmo de camino más corto.

Elimina cada linea de código después de la declaración del diccionario `copper`.

# --hints--

Debes eliminar todas las líneas de código después de la declaración de tu diccionario.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
