---
id: 65affe602a2a828b0382a058
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

Reemplaza la clave `'age'` con la cadena `'B'` y establece su valor en la cadena `'A'` para representar la conexión entre los nodos en ambas direcciones.

# --hints--

Tu diccionario debe tener una clave `'B'`.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

La tecla `'B'` debería tener `'A'` como valor.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

Tu diccionario debe tener una clave `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Tu tecla `'A'` debería tener `'B'` como valor.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Tu diccionario debería tener dos llaves.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
