---
id: 65b7cf140d34058d7ea8935f
title: '--Paso 10--'
challengeType: 20
dashedName: step-10
---

# --description--

Si quieres iterar sobre los valores de las claves del diccionario, una forma es usar el método `.values()`.

Modifica tu bucle `for` para que itere sobre `copper.values()` en lugar de `copper` y observa la salida.

# --hints--

Debes modificar tu bucle `for` para que itere sobre `copper.values()` en lugar de `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
