---
id: 65b7cd2b7bd9a684ccf1dc16
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Para iterar sobre las claves de un diccionario, simplemente puedes colocar el diccionario dentro de un bucle `for`. El siguiente código imprimiría cada clave del diccionario `dict`:

```py
for i in dict:
   print(i)
```

Sustituye la llamada a `print()` por un bucle `for` que itere sobre `copper` e imprima cada clave.

# --hints--

No debes tener `print(copper)` en tu código.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

Debes crear un bucle `for` para iterar sobre el diccionario `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

Debes imprimir cada clave de `copper` dentro de tu bucle `for`.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
