---
id: 65b7d25992879698180e6a71
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Por último, si quieres recorrer los pares clave-valor, puedes usar el método `.items()`.

Modifica tu bucle `for` para que itere sobre `copper.items()` en lugar de `copper.values()`.

# --hints--

Debes modificar tu bucle `for` para que itere sobre `copper.items()` en lugar de `copper.values()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
