---
id: 65b7d3319ebec69b983fb91b
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Como puedes ver de la salida, `.items()` crea una estructura de datos que almacena cada par de clave-valor en una tupla distinta. Para iterar sobre los elementos en esas tuplas puedes añadir una segunda variable de ciclo:

```py
for i, j in dict.items():
    print(i, j)
```

Modifique su bucle `for` para tomar dos variables de bucle e imprima ambas dentro del cuerpo del bucle.

# --hints--

Debes modificar tu bucle `for` para usar dos variables.

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

Debes imprimir ambas variables de bucle dentro de tu bucle `for`.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
