---
id: 694a9ff65d4472dd451d8255
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستتدرب على كيفية استخدام القيم المنطقية والعبارات الشرطية في Python من خلال بناء حاسبة لحجز تذاكر الأفلام.

ابدأ بإعلان متغيرات تخزن معلومات عن المستخدم وعرض الفيلم.

أنشئ متغيرًا باسم `base_price` لتخزين السعر الأساسي لتذكرة الفيلم وحدد قيمته بـ `15`. أنشئ متغيرًا آخر باسم `age` لتخزين عمر المستخدم وحدد قيمته بـ `21`.

# --hints--

يجب إعلان متغير باسم `base_price`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

يجب تعيين القيمة الصحيحة `15` إلى `base_price`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

يجب إعلان متغير باسم `age`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

يجب تعيين القيمة الصحيحة `21` إلى `age`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
