---
id: 694cc1b515976889afd4cded
title: خطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

الآن عليك تخزين بعض القيم النصية المتعلقة بتذكرة الفيلم.

أنشئ متغيرًا باسم `seat_type` لتخزين نوع المقعد الذي اختاره المستخدم وعيّن قيمته إلى السلسلة النصية `Gold`. أنشئ متغيرًا آخر باسم `show_time` لتخزين وقت عرض الفيلم وعيّن قيمته إلى السلسلة النصية `Evening`.

# --hints--

يجب إعلان متغير باسم `seat_type`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

يجب تعيين السلسلة النصية `Gold` إلى `seat_type`. تذكر إحاطة القيمة إما بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

يجب إعلان متغير باسم `show_time`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

يجب تعيين السلسلة النصية `Evening` إلى `show_time`. تذكر إحاطة القيمة إما بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
