---
id: 6958e59845777391b86d7331
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

في بايثون، يمكن أيضًا وضع جملة `if` داخل جسم جملة `if` أخرى. يُسمى هذا جملة `if` مسكنة.

تسمح جملة `if` المسكنة بفحص شرط إضافي فقط بعد تحقق الشرط الأول. ستُنفذ جملة `if` الداخلية فقط إذا كان شرط جملة `if` الخارجية صحيحًا.

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

الآن ستحسب رسوم الخدمة بناءً على نوع المقعد الذي اختاره المستخدم.

داخل جسم آخر جملة `if`، أسفل السطر `print('Ticket booking condition satisfied')`، أنشئ متغيرًا باسم `service_charges` وعيّنه إلى `0`. تأكد من ترك مسافة بادئة بأربع فراغات للحفاظ على المتغير داخل جسم جملة `if` الخارجية.

بعد ذلك، أنشئ جملة `if` مسكنة للتحقق مما إذا كان `seat_type` يساوي `Premium`. داخل جسم جملة `if` المسكنة، حدّث قيمة `service_charges` إلى `5`.

# --hints--

يجب إعلان متغير باسم `service_charges` داخل جسم آخر جملة `if` لديك.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

يجب تعيين القيمة الصحيحة `0` إلى `service_charges`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

يجب أن تحتوي على جملة `if` داخل جسم آخر جملة `if` لديك.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

يجب أن يكون شرط جملة `if` الجديدة هو `seat_type == 'Premium'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

يجب تعيين القيمة الصحيحة `5` إلى متغير `service_charges` داخل جملة `if` الجديدة.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
