---
id: 6958e0ecc121b0ae3dd9c5e0
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

يجب أن تتعامل أيضًا مع الحالة التي لا يستحق فيها المستخدم خصم العضوية.

أضف بند `else` إلى جملة `if` واطبع `User does not qualify for membership discount` داخل جسم `else`.

يمكنك طباعة رسالة متبوعة بمتغير بهذه الطريقة:

```py
print('Total:', total)
```

تريد أيضًا عرض القيمة المحدثة لـ `discount`. أسفل جملة `if...else`، استخدم استدعاء `print()` لعرض رسالة تُظهر `Discount:` متبوعة بالقيمة المحدثة لـ `discount`. ثم تحقق من المخرجات في الطرفية.

# --hints--

يجب أن يكون لديك بند `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

يجب أن تطبع `User does not qualify for membership discount` داخل بند `else` الخاص بك.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

يجب أن يكون لديك `print('Discount:', discount)` أسفل جملة `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
