---
id: 6958e1875e592b58e8852757
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

أنشئ الآن متغيرًا باسم `extra_charges` وقم بتعيينه إلى `0`. هذا سيمثل الرسوم الإضافية التي تُطبق على تذكرة الفيلم في عطلات نهاية الأسبوع.

بعد ذلك، أنشئ جملة `if` للتحقق مما إذا كانت قيمة `is_weekend` صحيحة. داخل جسم جملة `if`، حدّث قيمة `extra_charges` إلى `2` واطبع في الطرفية الرسالة `Extra charges will be applied`.

# --hints--

يجب إعلان متغير باسم `extra_charges`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

يجب تعيين القيمة الصحيحة `0` إلى المتغير `extra_charges`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

يجب أن يكون لديك جملة `if` رابعة.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

شرط جملة `if` الرابعة يجب أن يكون `is_weekend`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

يجب تعيين القيمة الصحيحة `2` إلى المتغير `extra_charges` داخل جملة `if` الجديدة.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

يجب طباعة الرسالة `Extra charges will be applied` داخل جملة `if` الجديدة. تذكر أن تحيط الرسالة بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
