---
id: 6958e2377cacbc46a7f7383a
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

أضف الآن بند `else` إلى جملة `if` واطبع `No extra charges will be applied` داخل جسم `else`. هذا سيطبع الرسالة فقط عندما يكون شرط الرسوم الإضافية خاطئًا.

ثم، أسفل بند `else`، استخدم استدعاء `print()` لعرض رسالة تُظهر `Extra charges:` متبوعة بالقيمة المحدثة لـ `extra_charges` وتحقق من المخرجات في الطرفية.

# --hints--

يجب أن يكون لديك بند `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

يجب أن تطبع `No extra charges will be applied` داخل بند `else`. تذكر أن تحيط الرسالة بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

يجب أن يكون لديك `print('Extra charges:', extra_charges)` أسفل جملة `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
