---
id: 6958de0d773289a89515a6b8
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

الآن ستتحقق مما إذا كان يُسمح للمستخدم بحجز عرض مسائي بناءً على عمره.

أنشئ جملة `if` للتحقق مما إذا كانت قيمة `age` أكبر من أو تساوي `21`. داخل جسم جملة الـ `if`، اطبع `User is eligible for Evening shows`.

# --hints--

يجب أن تحتوي شفرتك على جملتي `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

يجب أن يكون شرط جملة الـ `if` الثانية هو `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

يجب أن تطبع `User is eligible for Evening shows` داخل جملة الـ `if` الجديدة. تذكر أن تحيط الرسالة بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
