---
id: 6958dd85587e10c99162dd81
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

كما تعلمت في الدروس السابقة، في Python يمكن استخدام جملة `if` لتشغيل شفرة بناءً على تحقق شرط ما.

تتكون جملة `if` من الكلمة المفتاحية `if`، تليها شرط ونقطتان. الشفرة التي تُنفذ عندما يكون الشرط صحيحًا، ويجب أن تكون مُزاحة، تُسمى جسم جملة `if`.

```py
if condition:
    # Code to execute if condition is True
```

في هذه الخطوة، ستتحقق مما إذا كان المستخدم مؤهلاً لحجز تذكرة فيلم بناءً على عمره.

أنشئ جملة `if` للتحقق مما إذا كانت قيمة `age` أكبر من `17`. داخل جسم جملة `if`، اطبع الرسالة `User is eligible to book a ticket`. ستُطبع هذه الرسالة فقط عندما يكون عمر المستخدم أكبر من `17`.

تذكر أن تُزاح جسم جملة `if` وأن تحيط الرسالة بعلامات اقتباس مفردة أو مزدوجة داخل استدعاء `print()`.

# --hints--

يجب أن تحتوي شفرتك على جملة `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

يجب أن يكون شرط جملة `if` هو `age > 17`.

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

يجب أن تحتوي جملة `if` على استدعاء `print()`. تذكر أن تُزاح جسم جملة `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

يجب أن تطبع الرسالة `User is eligible to book a ticket` داخل جملة `if`. تذكر أن تحيط الرسالة بعلامات اقتباس مفردة أو مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
