---
id: 696099c1ac72a4e585de1d5b
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

بعض المعلومات يمكن أن تكون صحيحة أو خاطئة فقط. كما تعلمت في الدروس السابقة، يمكن تمثيل ذلك باستخدام القيم المنطقية.

أنشئ متغيرًا باسم `is_member` للدلالة على ما إذا كان المستخدم عضوًا، وعيّن قيمته إلى `True`.

أسفل المتغير `is_member` أنشئ متغيرًا آخر باسم `is_weekend` للدلالة على ما إذا كان عرض الفيلم في عطلة نهاية الأسبوع، وعيّن قيمته إلى `False`. لا تحيط القيمة بعلامات اقتباس.

# --hints--

يجب إعلان متغير باسم `is_member`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

يجب تعيين القيمة المنطقية `True` إلى `is_member`. لا تحيط القيمة بعلامات اقتباس.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

يجب إعلان متغير باسم `is_weekend`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

يجب تعيين القيمة المنطقية `False` إلى `is_weekend`. لا تحيط القيمة بعلامات اقتباس.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
