---
id: 696099abb6acc539b27055a1
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

في الخطوة السابقة، تحققت من شرط باستخدام جملة `if`. يُستخدم بند `else` لمعالجة الحالة عندما يكون الشرط خاطئًا. هذه صيغة نحوية لجملة `if...else`:

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

الآن، أضف بند `else` إلى جملة `if` واطبع `User is not eligible for Evening shows` داخل جسم `else`. سيُطبع هذا فقط عندما يكون الشرط في جملة `if` خاطئًا.

# --hints--

يجب أن يكون لديك بند `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

يجب أن تطبع `User is not eligible for Evening shows` داخل بند `else`. تذكر أن تحيط الرسالة بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
