---
id: 6960a848aacf859f7bb15447
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

غيّر الآن قيمة المتغير `is_member` إلى `False` لأن المستخدم ليس عضوًا.

بعد ذلك، سترى أن قيمة `discount` تبقى الآن `0`، لأن الشرطين يجب أن يتحققا لتطبيق الخصم.

# --hints--

يجب تحديث قيمة المتغير `is_member` إلى `False`. لا تحيط القيمة بعلامات اقتباس.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--
is_member = True
--fcc-editable-region--
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
