---
id: 6960a5165f7efe52a0c0aa39
title: خطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

في Python، يُستخدم العامل `and` للتحقق مما إذا كانت عدة شروط صحيحة. إليك كيف تستخدمه لدمج شرطين في جملة `if`:

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

يجب أن ينطبق خصم العضوية فقط على الأعضاء إذا كان `age` أكبر من أو يساوي `21`.

حدّث شرط السطر `if is_member:` باستخدام العامل `and` لدمج الشرط الحالي مع شرط آخر يتحقق مما إذا كان `age` أكبر من أو يساوي `21`.

# --hints--

يجب أن يكون شرط جملة `if` الثالثة هو `is_member and age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
