---
id: 6960b8632b9544d1f57295f4
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

يمكن للمستخدمين الذين تتراوح أعمارهم بين `18` و `21` حجز التذاكر، ولكن فقط عندما لا يكون `show_time` هو `Evening`.

حدّث شرط السطر `if age >= 21:`. استخدم العامل `and` لبناء تعبير يتحقق مما إذا كان `age` أكبر من أو يساوي `18` و`show_time` ليس `Evening`. ثم استخدم العامل `or` لدمج هذا التعبير مع الشرط الحالي. لا تنشئ متغيرات جديدة.

# --hints--

يجب أن يكون شرط آخر جملة `if` هو `age >= 21 or age >= 18 and show_time != 'Evening'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
_first = _cond.is_equivalent("age >= 21 or age >= 18 and show_time != 'Evening'")
_second = _cond.is_equivalent("age >= 21 or 18 <= age and show_time != 'Evening'")
assert _first or _second
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21:
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
