---
id: 6960b6a6ea06e27dcc7eef37
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

الآن ستتحقق مما إذا كان المستخدم يستوفي الشروط لحجز تذكرة فيلم. يمكن للمستخدمين الذين تبلغ أعمارهم `21` سنة أو أكثر حجز التذاكر دائمًا دون أي قيود.

أنشئ عبارة `if` للتحقق مما إذا كانت قيمة `age` أكبر من أو تساوي `21`. داخل جسم عبارة `if`، اطبع `Ticket booking condition satisfied` على الطرفية.

ثم أضف بند `else` إلى عبارة `if` الخاصة بك واطبع `Ticket booking failed due to restrictions` داخل جسم `else`.

# --hints--

يجب أن يكون لديك عبارة `if` خامسة.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

يجب أن يكون شرط عبارة `if` الخامسة هو `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

يجب أن تطبع `Ticket booking condition satisfied` داخل عبارة `if` الجديدة. تذكر أن تحيط الرسالة بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

يجب أن تحتوي عبارة `if` الجديدة على بند `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

يجب أن تطبع `Ticket booking failed due to restrictions` داخل بند `else`. تذكر أن تحيط الرسالة بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
