---
id: 6960ba6e6c95ed4ad896c6f2
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

هناك استثناء آخر لقواعد الحجز. يمكن للمستخدمين الذين تتراوح أعمارهم بين `18` و`21` حجز العروض المسائية إذا كانوا أعضاءً.

عندما تُستخدم عدة معاملات منطقية في جملة `if`، تُقيّم الشروط المرتبطة بـ `and` قبل الشروط المرتبطة بـ `or`. تُستخدم الأقواس العادية `()` في بايثون لتجميع الشروط والتحكم في ترتيب تقييمها.

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

حدّث شرط السطر `if age >= 21 or age >= 18 and show_time != 'Evening':` لإضافة شرط آخر باستخدام معامل `or` للتحقق مما إذا كانت قيمة `is_member` صحيحة. استخدم الأقواس العادية `()` لتجميع الشرطين `show_time != 'Evening'` و`is_member` معًا كما في المثال أعلاه.

# --hints--

يجب أن يكون شرط جملة `if` الأخيرة هو `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
