---
id: 696cb889fc4fdeb337e52864
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

في Python، لا يجب أن تقارن شروط `if` القيم صراحةً بـ `True` أو `False`. بدلاً من ذلك، تعتمد على صحة القيم. القيم الصحيحة تُقيَّم كـ `True` في سياق القيمة المنطقية، مثل شرط جملة `if`. تشمل هذه الأعداد غير الصفرية، والسلاسل النصية غير الفارغة، والقيمة المنطقية `True`.

```py
if value:
   print('value is truthy')
```

أنشئ جملة `if` للتحقق مما إذا كانت `is_member` قيمة صحيحة. داخل جسم جملة `if`، حدّث قيمة `discount` إلى `3` واطبع الرسالة `User qualifies for membership discount` على الطرفية.

# --hints--

يجب أن تحتوي شفرتك على ثلاث جمل `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

يجب أن يكون شرط جملة `if` الثالثة هو `is_member`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

يجب تعيين القيمة الصحيحة `3` إلى المتغير `discount` داخل جملة `if` الجديدة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

يجب طباعة `User qualifies for membership discount` داخل جملة `if` الجديدة. تذكر أن تحيط الرسالة بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
