---
id: 6a57c331ae94fa4631b60869
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Nuevas investigaciones confirman que la máscara dorada pertenecía a la realeza. Agrega la etiqueta `"royal"` al artefacto con un `id` de `101`. Luego, llama a `console.log()` con `collection[101].tags` para registrar las etiquetas de ese artefacto.

# --hints--

Debes llamar a `addTag` con `101` y `"royal"`.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)/);
```

Debes registrar las `tags` del primer artefacto en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

Debes registrar las `tags` después de agregar la nueva etiqueta.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

Las `tags` del primer artefacto deben incluir `"royal"`.

```js
assert.include(collection[101].tags, "royal");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
