---
id: 6a57c331ae94fa4631b60864
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

La auditoría comienza con leer lo que ya está registrado. Como cada artefacto se almacena bajo su id, puedes acceder a uno poniendo ese id entre corchetes, como `collection[101]`.

Haz una llamada a `console.log()` con `collection[101].title` para registrar el título del primer artefacto. Luego, haz una llamada a `console.log()` con `collection[101].curator.name` para registrar el nombre de su curador.

# --hints--

Debes registrar el `title` del primer artefacto en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)/);
```

Debes registrar el `name` del curador del primer artefacto en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

Debes registrar primero el `title` y luego el `name` del curador.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
