---
id: 6a57c331ae94fa4631b60866
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

A veces el personal pedirá un artefacto que el museo no tiene. Si no hay ningún artefacto almacenado bajo ese `id`, `collection[id]` es `undefined`, y leer `.title` de él causaría un error.

Actualiza la sentencia return para usar una condicional. Devuelve `artifact.title` si se encuentra un artefacto, y la cadena `"Artifact not found"` si no es así.

# --hints--

`getArtifactTitle(101)` aún debería devolver `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle` debería devolver `"Artifact not found"` si ningún artefacto coincide con el `id`.

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle` debería devolver el título de cualquier artefacto que encuentre, no solo de `101`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
