---
id: 6a57c331ae94fa4631b60861
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Cada artefacto tiene una persona responsable.

Agrega una propiedad `curator` al artefacto almacenado bajo la clave `101` y asigna su valor a un objeto con un `id` de `201` y un `name` de `"Earl Sinclair"`.

# --hints--

El artefacto almacenado bajo `101` debe tener una propiedad `curator`.

```js
assert.property(collection[101], "curator");
```

El `curator` debe ser un objeto.

```js
assert.isObject(collection[101].curator);
```

El `curator` debe tener un `id` de `201`.

```js
assert.strictEqual(collection[101].curator.id, 201);
```

El `curator` debe tener un `name` de `"Earl Sinclair"`.

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
