---
id: 6a57c331ae94fa4631b60871
title: Paso 18
challengeType: 1
dashedName: step-18
---

# --description--

La galería actual de un artefacto es su ubicación más reciente. Puedes encontrarla revisando el último elemento en su arreglo `locations`.

Declara una variable `currentLocation` y asígnale `artifact.locations[artifact.locations.length - 1]`. Esta expresión usa la propiedad `length` del arreglo para acceder a su último elemento.

# --hints--

Debes declarar una variable llamada `currentLocation` dentro de `buildSummary`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary?.variables.currentLocation);
```

Debes asignar el último elemento del arreglo `locations` del artefacto a `currentLocation`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.buildSummary?.variables.currentLocation?.value.matches("artifact.locations[artifact.locations.length - 1]"));
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

function buildSummary(id) {
  const artifact = collection[id];
  if (!artifact) {
    return "Artifact not found";
  }
--fcc-editable-region--
  
--fcc-editable-region--
}
```
