---
id: 6a57c331ae94fa4631b6086a
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

A medida que los artefactos se trasladan entre galerías con el tiempo, cada movimiento se registra como una nueva entrada en el arreglo `locations`.

Crea una función `moveArtifact` con los parámetros `id`, `gallery` y `year`. Encuentra el artefacto usando `collection[id]`. Si existe, agrega un nuevo objeto con `gallery` y `year` a su arreglo `locations`.

# --hints--

Debes tener una función llamada `moveArtifact` con parámetros llamados `id`, `gallery` y `year`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.moveArtifact);
const params = explorer.allFunctions.moveArtifact.parameters;
assert.equal(params.length, 3);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "gallery");
assert.equal(params[2].toString(), "year");
```

Llamar a `moveArtifact` debe agregar un nuevo objeto `{ gallery, year }` al arreglo `locations` del artefacto.

```js
const start = collection[101].locations.length;
moveArtifact(101, "Test Hall", 2030);
assert.lengthOf(collection[101].locations, start + 1);
assert.deepEqual(collection[101].locations[collection[101].locations.length - 1], { gallery: "Test Hall", year: 2030 });
```

`moveArtifact` debe agregar la nueva ubicación al artefacto al que se refiere el `id`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
moveArtifact(457, "Test Vault", 2032);
assert.lengthOf(collection[457].locations, 2);
assert.deepEqual(collection[457].locations[1], { gallery: "Test Vault", year: 2032 });
```

`moveArtifact` no debe lanzar un error ni modificar la colección cuando ningún artefacto coincide con el `id`.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => moveArtifact(999, "Test Hall", 2030));
assert.strictEqual(JSON.stringify(collection), before);
```

Debes agregar la nueva ubicación al arreglo `locations` existente del artefacto, en lugar de reemplazar el arreglo.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const locationsBefore = collection[458].locations;
const spy = __helpers.spyOn(locationsBefore, "push");
try {
  moveArtifact(458, "Push Hall", 2033);
  assert.deepEqual(spy.calls, [[{ gallery: "Push Hall", year: 2033 }]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].locations, locationsBefore);
assert.deepEqual(collection[458].locations[1], { gallery: "Push Hall", year: 2033 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

--fcc-editable-region--

--fcc-editable-region--
```
