---
id: 6a57c331ae94fa4631b60868
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Los curadores etiquetan los artefactos para facilitar su búsqueda, y una etiqueta nunca debe agregarse dos veces. Antes de agregar una nueva etiqueta, puedes usar el método `includes()` para verificar si ya existe, ya que devuelve `true` cuando un arreglo ya contiene un valor.

Crea una función llamada `addTag` que reciba los parámetros `id` y `tag`. Busca el artefacto usando `collection[id]`. Si existe y su arreglo `tags` no incluye ya `tag`, agrega `tag` a `artifact.tags`.

# --hints--

Debes tener una función llamada `addTag` con parámetros llamados `id` y `tag`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

Llamar a `addTag` con una etiqueta nueva debe agregar esa etiqueta al arreglo `tags` del artefacto.

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

Llamar a `addTag` con una etiqueta que el artefacto ya tiene no debe agregar un duplicado.

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

`addTag` debe agregar la etiqueta al artefacto al que se refiere el `id`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

`addTag` no debe lanzar un error ni cambiar la colección cuando ningún artefacto coincide con el `id`.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

Debes agregar `tag` al arreglo `tags` existente del artefacto, en lugar de reemplazar el arreglo.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
