---
id: 6a57c331ae94fa4631b60862
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Cada objeto también registra dónde ha estado, cómo se describe y si está en exhibición.

Agrega tres propiedades más al objeto almacenado bajo `101`:

- `locations`: un arreglo que contiene dos objetos, `{ gallery: "Hall A", year: 2020 }` y `{ gallery: "Hall C", year: 2024 }`
- `tags`: un arreglo que contiene las cadenas `"gold"` y `"egypt"`
- `onDisplay`: el booleano `true`

# --hints--

El `locations` del objeto debe ser un arreglo con los dos objetos de ubicación dados.

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

El `tags` del objeto debe ser `["gold", "egypt"]`.

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

El `onDisplay` del objeto debe ser `true`.

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
