---
id: 6a57c265abbc95dfebc05098
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, vas a construir un auditor de procedencia de artefactos. Almacenarás artefactos de museo en un objeto y escribirás funciones para buscarlos, etiquetarlos, moverlos entre galerías y resumir su historia.

Cada artefacto se almacena bajo su id. Crea un objeto llamado `collection` y agrega una entrada. Establece la clave en `101` y el valor en un objeto con un `title` de `"Golden Mask"` y una `category` de `"Ceremonial"`.

# --hints--

Debes declarar una variable llamada `collection`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` debe ser un objeto.

```js
assert.isObject(collection);
```

`collection` debe tener una entrada.

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

El artefacto almacenado bajo `101` debe ser un objeto.

```js
assert.isObject(collection[101]);
```

El artefacto almacenado bajo `101` debe tener un `title` de `"Golden Mask"`.

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

El artefacto almacenado bajo `101` debe tener una `category` de `"Ceremonial"`.

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
