---
id: 6808fe91a5cb0c3355f23027
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將學習如何透過建置一個超級英雄申請表來操作 React 中的表單。

已提供部分樣板程式碼，你只需要創建剩餘的 React 組件。

首先傳回一個 `div` 元素，其 `className` 為 `form-wrap`。在該 `div` 內，創建一個 `h2` 元素，文字為 `Superhero Application Form`。在其下方，創建一個 `p` 元素，文字為 `Please complete all fields`。

# --hints--

你應該傳回一個 `div` 元素，其 `className` 為 `form-wrap`。

```js
assert.exists(document?.querySelector("div.form-wrap"));
```

你應該在你的 `div` 元素內創建一個 `h2` 元素。

```js
assert.exists(document?.querySelector(".form-wrap > h2"));
```

你的 `h2` 元素應該有文字 `Superhero Application Form`。

```js
const h2El = document?.querySelector("h2")
assert.equal(h2El?.textContent, "Superhero Application Form");
```

你應該在你的 `div` 元素內創建一個 `p` 元素。確保它位於 `h2` 之後。

```js
assert.exists(document?.querySelector(".form-wrap > h2 + p"));
```

你的 `p` 元素應該有文字 `Please complete all fields`。

```js
const pEl = document?.querySelector("p")
assert.equal(pEl?.textContent, "Please complete all fields");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  --fcc-editable-region--
  
  --fcc-editable-region--

};
```
