---
id: 68e76b1ce8216249d04df670
title: "Planet クラスを作成する"
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `Planet` という名前のクラスを作成してください。
1. `Planet` クラスには、次の条件を満たす `__init__` メソッドが必要です：
   - 4つのパラメータ `self`、`name`、`planet_type`、`star` を持つこと。
   - 渡された引数のいずれかが文字列型でない場合、メッセージ `TypeError` を持つ `name, planet type, and star must be strings` を発生させること。
   - 渡された引数のいずれかが空の文字列の場合、メッセージ `ValueError` を持つ `name, planet_type, and star must be non-empty strings` を発生させること。
   - 渡された値をインスタンス属性 `name`、`planet_type`、`star` に代入すること。
1. `Planet` クラスには、`orbit` メソッドがあり、`{name} is orbiting around {star}...` の形式の文字列を返すこと。
1. `Planet` クラスには、`__str__` メソッドがあり、`Planet: {name} | Type: {planet_type} | Star: {star}` の形式の文字列を返すこと。
1. `Planet` クラスのインスタンスを3つ作成し、それぞれ `planet_1`、`planet_2`、`planet_3` と名付けてください。
1. 各惑星オブジェクトを出力して、`__str__` メソッドの出力を確認してください。
1. 各惑星オブジェクトで `orbit` メソッドを呼び出し、その結果を出力してください。

# --hints--

`Planet` という名前のクラスを作成してください。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

`Planet` クラスには `__init__` メソッドが必要です。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

`__init__` メソッドは、順に `self`、`name`、`planet_type`、`star` の4つのパラメータを持つ必要があります。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

`__init__` メソッドは、最初の引数が文字列型でない場合、メッセージ `TypeError` を持つ `name, planet type, and star must be strings` を発生させる必要があります。

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` メソッドは、2番目の引数が文字列型でない場合、メッセージ `TypeError` を持つ `name, planet type, and star must be strings` を発生させる必要があります。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` メソッドは、3番目の引数が文字列型でない場合、メッセージ `TypeError` を持つ `name, planet type, and star must be strings` を発生させる必要があります。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` メソッドは、最初の引数が空の文字列の場合、メッセージ `ValueError` を持つ `name, planet_type, and star must be non-empty strings` を発生させる必要があります。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` メソッドは、2番目の引数が空の文字列の場合、メッセージ `ValueError` を持つ `name, planet_type, and star must be non-empty strings` を発生させる必要があります。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` メソッドは、3番目の引数が空の文字列の場合、メッセージ `ValueError` を持つ `name, planet_type, and star must be non-empty strings` を発生させる必要があります。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`name` を `self.name` に `__init__` メソッド内で代入してください。

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

`planet_type` を `self.planet_type` に `__init__` メソッド内で代入してください。

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

`star` を `self.star` に `__init__` メソッド内で代入してください。

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

`Planet` クラスには `orbit` メソッドが必要です。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

`orbit` メソッドは `{name} is orbiting around {star}...` の形式の文字列を返す必要があります。

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

`Planet` クラスには `__str__` メソッドが必要です。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

`__str__` メソッドは `Planet: {name} | Type: {planet_type} | Star: {star}` の形式の文字列を返す必要があります。

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

`Planet` クラスのインスタンスを3つ作成し、それぞれ `planet_1`、`planet_2`、`planet_3` と名付けてください。

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

各惑星オブジェクトを出力して、`__str__` メソッドの出力を確認してください。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

各惑星オブジェクトで `orbit` メソッドを呼び出し、その結果を出力してください。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
