---
id: 6a89c69c99298438d6fbe60b
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

أضف الآن الحالة العودية. بعد جملة `if`، أنشئ متغيرًا باسم `count_list`. عيّن له نتيجة استدعاء `countup` مع `number - 1`.

كل استدعاء يستخدم رقمًا أصغر، لذا تتحرك العودية نحو الحالة الأساسية.

# --hints--

يجب أن تعيّن `countup(number - 1)` إلى متغير باسم `count_list`.

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
