---
id: 6a89c69c8dbfda74151fa0b6
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستنشئ دالة `countup` تستخدم العودية لإنشاء قائمة أرقام من `1` إلى رقم معين.

ابدأ بتعريف دالة باسم `countup` تحتوي على معلمة واحدة باسم `number`. أضف الكلمة المفتاحية `pass` إلى جسم الدالة لتكون الدالة صحيحة أثناء فراغها.

# --hints--

يجب أن تحتوي على دالة باسم `countup` مع معلمة باسم `number`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

يجب أن تحتوي دالة `countup` على الكلمة المفتاحية `pass`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
