---
id: 6a89c69c6af682856558d78b
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

يرجع الاستدعاء التكراري قائمة تحتوي على الأعداد الأصغر. استخدم `append()` لإضافة `number` إلى `count_list`، ثم أرجع `count_list`.

# --hints--

يجب أن تضيف `number` إلى `count_list`.

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

يجب أن تُرجع دالة `countup` الخاص بك `count_list`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

يجب أن تُرجع دالة `countup` الخاص بك `[1, 2, 3]` عند استدعائها مع `3`.

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
