---
id: 6a89c69c111694edc6d1c628
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

اختبر دالة `countup` بتمرير `5` إليها داخل استدعاء `print()`. يجب أن يكون الناتج `[1, 2, 3, 4, 5]`.

بهذا، تكتمل ورشة عمل دالة countup.

# --hints--

يجب استدعاء `countup(5)` داخل دالة `print()`.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
