---
id: 6a89c69c91bc7365fbd4b733
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

أضف الحالة الأساسية إلى دالة `countup` الخاصة بك. استبدل `pass` بعبارة `if` تتحقق مما إذا كان `number` أقل من `1`. داخل عبارة `if`، أرجع قائمة فارغة.

# --hints--

يجب أن تُرجع دالة `countup` قائمة فارغة عندما يكون `number` أقل من `1`.

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

يجب ألا تُرجع دالة `countup` قائمة فارغة عندما يكون `number` أكبر من أو يساوي `1`.

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
