---
id: 587d8249367417b2b2512c41
title: Metrisch-Imperial-Konverter
challengeType: 4
forumTopicId: 301570
dashedName: metric-imperial-converter
---

# --description--

Erstellen Sie eine Full-Stack-JavaScript-App, die funktional ähnlich ist wie diese: <a href="https://metric-imperial-converter.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://metric-imperial-converter.freecodecamp.rocks/</a>. Die Arbeit an diesem Projekt beinhaltet, dass Sie Ihren Code mit einer der folgenden Methoden schreiben:

- Klonen Sie <a href="https://github.com/freeCodeCamp/boilerplate-project-metricimpconverter/" target="_blank" rel="noopener noreferrer nofollow">dieses GitHub-Repository</a> und schließen Sie Ihr Projekt lokal ab.
- Verwenden Sie einen Website-Builder Ihrer Wahl, um das Projekt abzuschließen. Stellen Sie sicher, dass Sie alle Dateien aus unserem GitHub-Repository einbinden.

**Hinweis:** Die Tests dieses Projekts funktionieren nicht, wenn Sie `glitch.com` verwenden.

# --instructions--

- Vervollständigen Sie die notwendige Umrechnungslogik in `/controllers/convertHandler.js`
- Vervollständigen Sie die notwendigen Routen in `/routes/api.js`
- Kopieren Sie die `sample.env`-Datei nach `.env` und setzen Sie die Variablen entsprechend
- Um die Tests automatisch auszuführen, fügen Sie `NODE_ENV=test` in Ihre `.env`-Datei ein
- Um die Tests in der Konsole auszuführen, verwenden Sie den Befehl `npm run test`.

Schreiben Sie die folgenden Tests in `tests/1_unit-tests.js`:

- `convertHandler` sollte eine ganze Zahl als Eingabe korrekt lesen.
- `convertHandler` sollte eine Dezimalzahl als Eingabe korrekt lesen.
- `convertHandler` sollte eine Bruch-Eingabe korrekt lesen.
- `convertHandler` sollte eine Bruch-Eingabe mit Dezimalzahl korrekt lesen.
- `convertHandler` sollte bei einem Doppelbruch (d.h. `3/2/3`) korrekt einen Fehler zurückgeben.
- `convertHandler` sollte standardmäßig eine numerische Eingabe von `1` verwenden, wenn keine numerische Eingabe bereitgestellt wird.
- `convertHandler` sollte jede gültige Eingabeeinheit korrekt lesen.
- `convertHandler` sollte für eine ungültige Eingabeeinheit korrekt einen Fehler zurückgeben.
- `convertHandler` sollte für jede gültige Eingabeeinheit die korrekte Rückgabe-Einheit liefern.
- `convertHandler` sollte für jede gültige Eingabeeinheit die ausgeschriebene String-Einheit korrekt zurückgeben.
- `convertHandler` sollte `gal` korrekt in `L` umrechnen.
- `convertHandler` sollte `L` korrekt in `gal` umrechnen.
- `convertHandler` sollte `mi` korrekt in `km` umrechnen.
- `convertHandler` sollte `km` korrekt in `mi` umrechnen.
- `convertHandler` sollte `lbs` korrekt in `kg` umrechnen.
- `convertHandler` sollte `kg` korrekt in `lbs` umrechnen.

Schreiben Sie die folgenden Tests in `tests/2_functional-tests.js`:

- Konvertieren Sie eine gültige Eingabe wie `10L`: `GET`-Anfrage zu `/api/convert`.
- Konvertieren Sie eine ungültige Eingabe wie `32g`: `GET`-Anfrage zu `/api/convert`.
- Konvertieren Sie eine ungültige Zahl wie `3/7.2/4kg`: `GET`-Anfrage zu `/api/convert`.
- Konvertieren Sie eine ungültige Zahl UND Einheit wie `3/7.2/4kilomegagram`: `GET`-Anfrage zu `/api/convert`.
- Konvertieren Sie ohne Zahl wie `kg`: `GET`-Anfrage zu `/api/convert`.

# --hints--

Sie können Ihr eigenes Projekt bereitstellen, nicht die Beispiel-URL.

```js
  assert(
    !/.*\/metric-imperial-converter\.freecodecamp\.rocks/.test(
      code
    )
  );
```

Sie können `GET` `/api/convert` mit einem einzigen Parameter aufrufen, der eine akzeptierte Zahl und Einheit enthält, und diese umrechnen lassen. (Hinweis: Teilen Sie die Eingabe, indem Sie den Index des ersten Zeichens suchen, der den Beginn der Einheit markiert)

```js

```

Sie können `'gal'` in `'L'` und umgekehrt umrechnen. (1 gal zu 3.78541 L)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 3.78541);
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10gal');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 37.8541);
    assert.equal(data2.returnUnit, 'L');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.26417);
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10l');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 2.64172);
    assert.equal(data4.returnUnit, 'gal');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Sie können `'lbs'` in `'kg'` und umgekehrt umrechnen. (1 lbs zu 0.453592 kg)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1lbs');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 0.45359);
    assert.equal(data1.returnUnit, 'kg');
    const response2 = await fetch(code + '/api/convert?input=10lbs');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 4.53592);
    assert.equal(data2.returnUnit, 'kg');
    const response3 = await fetch(code + '/api/convert?input=1kg');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 2.20462);
    assert.equal(data3.returnUnit, 'lbs');
    const response4 = await fetch(code + '/api/convert?input=10kg');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 22.04624);
    assert.equal(data4.returnUnit, 'lbs');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Sie können `'mi'` in `'km'` und umgekehrt umrechnen. (1 mi zu 1.60934 km)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 1.60934);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=10mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 16.0934);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(code + '/api/convert?input=1km');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.62137);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(code + '/api/convert?input=10km');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 6.21373);
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Alle eingehenden Einheiten sollten sowohl in Groß- als auch Kleinbuchstaben akzeptiert werden, aber sowohl in `initUnit` als auch `returnUnit` in Kleinbuchstaben zurückgegeben werden, außer Liter, das als Großbuchstabe `'L'` dargestellt werden sollte.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.initUnit, 'gal');
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10L');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.initUnit, 'L');
    assert.equal(data2.returnUnit, 'gal');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.initUnit, 'L');
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10KM');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.initUnit, 'km');
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Wenn die Maßeinheit ungültig ist, wird `'invalid unit'` zurückgegeben.

```js
  try {
    const response = await fetch(code + '/api/convert?input=1min');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid unit' || data === 'invalid unit');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Wenn die Zahl ungültig ist, wird `'invalid number'` zurückgegeben.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2gal'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid number' || data === 'invalid number');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Wenn sowohl die Einheit als auch die Zahl ungültig sind, wird `'invalid number and unit'` zurückgegeben.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2min'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(
      data.error === 'invalid number and unit' ||
        data === 'invalid number and unit'
    );
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Sie können Brüche, Dezimalzahlen oder beides im Parameter verwenden (z.B. 5, 1/2, 2.5/6), aber wenn nichts angegeben ist, wird standardmäßig 1 verwendet.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.approximately(data1.initNum, 1, 0.001);
    assert.approximately(data1.returnNum, 1.60934, 0.001);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=1/5mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.approximately(data2.initNum, 1 / 5, 0.1);
    assert.approximately(data2.returnNum, 0.32187, 0.001);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(
      code + '/api/convert?input=1.5/7km'
    );
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.approximately(data3.initNum, 1.5 / 7, 0.001);
    assert.approximately(data3.returnNum, 0.13315, 0.001);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(
      code + '/api/convert?input=3/2.7km'
    );
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.approximately(data4.initNum, 3 / 2.7, 0.001);
    assert.approximately(data4.returnNum, 0.69041, 0.001);
    assert.equal(data4.returnUnit, 'mi');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Ihre Rückgabe besteht aus `initNum`, `initUnit`, `returnNum`, `returnUnit` und `string`, die Einheiten im Format `'{initNum} {initUnitString} converts to {returnNum} {returnUnitString}'` ausschreiben, mit dem Ergebnis auf 5 Dezimalstellen gerundet.

```js
  try {
    const response = await fetch(code + '/api/convert?input=2mi');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.json();
    assert.equal(data.initNum, 2);
    assert.equal(data.initUnit, 'mi');
    assert.approximately(data.returnNum, 3.21868, 0.001);
    assert.equal(data.returnUnit, 'km', 'returnUnit did not match');
    assert.equal(data.string, '2 miles converts to 3.21868 kilometers');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Alle 16 Einheitentests sind abgeschlossen und bestanden.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const unitTests = getTests.filter(test => {
      return !!test.context.match(/Unit Tests/gi);
    });
    assert.isAtLeast(unitTests.length, 16, 'At least 16 tests passed');
    unitTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Alle 5 Funktionstests sind abgeschlossen und bestanden.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const functTests = getTests.filter(test => {
      return !!test.context.match(/Functional Tests/gi);
    });
    assert.isAtLeast(functTests.length, 5, 'At least 5 tests passed');
    functTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

