---
id: 69008059d43d6751385883c1
title: 调试有错误的应用中的递增和递减操作符错误
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

你的朋友请求你帮助他们调试代码。他们正在练习递增和递减操作符，但在控制台中得到意外的结果。

你的任务是帮助他们识别错误并更新代码，以使控制台输出与预期结果匹配。

你可以参考之前关于递增和递减操作符以及前缀和后序符号的课程。

**用户故事：**

1. 你应该有一个名为 `updatedCounter` 的变量，通过使用递增操作符正确生成值 `11`。更新表达式以使用正确的前缀或后序表示法，而不是硬编码该值。
2. 你应该有一个名为 `finalScore` 的变量，通过使用递增操作符正确生成值 `8`。更新表达式以使用正确的前缀或后序表示法，而不是硬编码该值。
3. 你应该有一个名为 `updatedCoins` 的变量，通过使用递减操作符正确生成值 `2`。更新表达式以使用正确的前缀或后序表示法，而不是硬编码该值。
4. 你应该有一个名为 `newHealth` 的变量，通过使用递减操作符正确生成值 `7`。更新表达式以使用正确的前缀或后序表示法，而不是硬编码该值。

# --hints--

你应该有一个名为 `updatedCounter` 的变量，使用前缀递增符号（`++counter`）。

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` 应该相等于 `11`。

```js
assert.strictEqual(updatedCounter, 11);
```

你应该有一个名为 `finalScore` 的变量，使用后序递增表示法（`score++`）。

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` 应该相等于 `8`。

```js
assert.strictEqual(finalScore, 8);
```

你应该有一个名为 `updatedCoins` 的变量，使用前缀递减符号（`--coins`）。

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` 应该相等于 `2`。

```js
assert.strictEqual(updatedCoins, 2);
```

你应该有一个名为 `newHealth` 的变量，使用后序递减符号（`health--`）。

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` 应该相等于 `7`。

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
