---
id: 69860157f7fce39f9df399a3
title: Nクイーンアルゴリズムの実装
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

Nクイーン問題は、N×NのチェスボードにN個のクイーンを配置し、どのクイーンも互いに攻撃し合わないようにする（同じ行、列、または斜めに2つのクイーンが存在しない）ことを求めます。

例えば、4x4のボードがある場合、有効な配置の一例は次の通りです。

```md
[1, 3, 0, 2]
```

これは、行0にはクイーンが列1に、行1には列3に、行2には列0に、行3には列2に配置されていることを意味します。

視覚的には、この配置は次のように見えます。

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

ここで`Q`はクイーンを、`.`は空のマスを表します。

このラボでは、深さ優先探索のアプローチを使ってNクイーン問題の解決プログラムを実装します。

**目的**：以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `dfsNQueens`という名前の関数を持っていること。
2. その関数は整数の`n`を1つだけ引数として受け取ること。
3. `n`が`1`より小さい場合、関数は空の配列（`[]`）を返すこと。
4. 関数は解のリストを返すこと。各解は長さ`n`のリストで、インデックス`i`の要素は行`i`のクイーンの列インデックス（0始まり）を表すこと。

# --hints--

`dfsNQueens`という名前の関数を1つの引数で持っている必要があります。

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

`n`が`1`より小さい場合、関数は空の配列を返す必要があります。

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

関数は解の配列を返し、各解は長さ`n`の配列です。

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)`は`[[0]]`を返すべきです。

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)`は`[]`を返すべきです。

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)`は`[]`を返すべきです。

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)`は`[[1, 3, 0, 2], [2, 0, 3, 1]]`を返すべきです。

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)`は`[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`を返すべきです。

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length`は`10`であるべきです。

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length`は`92`であるべきです。

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens`は正しい結果を返すべきです。

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
