---
id: bad82fee1322bd9aedf08721
title: Capire le unità assolute e relative
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN66JSL'
forumTopicId: 301089
dashedName: understand-absolute-versus-relative-units
---

# --description--

Le ultime sfide hanno impostato il margine o la spaziatura interna di un elemento usando i pixel (`px`). I pixel sono un tipo di unità di lunghezza, che indica al browser come dimensionare o distanziare un elemento. Oltre ai `px`, CSS offre diverse opzioni di unità di lunghezza che puoi usare.

I due tipi principali di unità di lunghezza sono assolute e relative. Le unità assolute si riferiscono a unità fisiche di lunghezza. Per esempio, `in` e `mm` indicano rispettivamente pollici e millimetri. Le unità di lunghezza assolute approssimano la misura reale su uno schermo, ma ci sono alcune differenze a seconda della risoluzione dello schermo.

Le unità relative, come `em` o `rem`, sono relative a un altro valore di lunghezza. Per esempio, `em` si basa sulla dimensione del carattere tipografico di un elemento. Se lo usi per impostare la proprietà `font-size` stessa, è relativo al carattere tipografico del genitore `font-size`.

**Nota:** Ci sono diverse unità relative legate alla dimensione del viewport. Sono trattate nella sezione dei principi del design web responsivo.

# --instructions--

Aggiungi una proprietà `padding` all'elemento con classe `red-box` e impostala su `1.5em`.

# --hints--

La tua classe `red-box` dovrebbe avere una proprietà `padding`.

```js
const redBox =document.querySelector('.red-box'); 
const style = window.getComputedStyle(redBox); 

assert.notEqual(style['padding-top'], '0px');
assert.notEqual(style['padding-right'], '0px');
assert.notEqual(style['padding-bottom'], '0px');
assert.notEqual(style['padding-left'], '0px');
```

La tua classe `red-box` dovrebbe dare agli elementi 1.5em di `padding`.

```js
assert.match(__helpers.removeCssComments(code), /\.red-box\s*?{[\s\S]*padding\s*:\s*?1\.5em/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;

  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;
    padding: 1.5em;
  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```
