---
id: bad87fee1348bd9aedf08826
title: usa la notazione oraria per specificare la spaziatura interna di un elemento
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB7mBS9'
forumTopicId: 18346
dashedName: use-clockwise-notation-to-specify-the-padding-of-an-element
---

# --description--

Invece di specificare singolarmente le proprietà `padding-top`, `padding-right`, `padding-bottom` e `padding-left` di un elemento, puoi specificarle tutte in una sola riga, così:

```css
padding: 10px 20px 10px 20px;
```

Questi quattro valori funzionano come un orologio: alto, destro, basso, sinistro, e producono lo stesso risultato di usare le istruzioni di spaziatura interna specifiche per lato.

# --instructions--

Usa la notazione oraria per dare alla classe `.blue-box` una spaziatura interna `padding` di `40px` sul lato alto e sinistro, ma solo `20px` sul lato basso e destro.

# --hints--

La tua classe `blue-box` dovrebbe dare agli elementi una spaziatura interna `40px` di `padding` sul lato alto.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '40px');
```

La tua classe `blue-box` dovrebbe dare agli elementi una spaziatura interna `20px` di `padding` sul lato destro.

```js
const blueBox = document.querySelector('.blue-box');
const paddingRight = window.getComputedStyle(blueBox)['padding-right'];
assert.strictEqual(paddingRight, '20px');
```

La tua classe `blue-box` dovrebbe dare agli elementi una spaziatura interna `20px` di `padding` sul lato basso.

```js
const blueBox = document.querySelector('.blue-box');
const paddingBottom = window.getComputedStyle(blueBox)['padding-bottom'];
assert.strictEqual(paddingBottom, '20px');
```

La tua classe `blue-box` dovrebbe dare agli elementi una spaziatura interna `40px` di `padding` sul lato sinistro.

```js
const blueBox = document.querySelector('.blue-box');
const paddingLeft = window.getComputedStyle(blueBox)['padding-left'];
assert.strictEqual(paddingLeft, '40px');
```

Dovresti usare la notazione oraria per impostare la spaziatura interna della classe `blue-box`.

```js
const css =  __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(css, /\.blue-box\s*{[\s\S]*padding\s*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
