---
id: bad87fee1348bd9aedf08746
title: "Ereditare gli stili dall'elemento body"
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bmdtR'
forumTopicId: 18204
dashedName: inherit-styles-from-the-body-element
---

# --description--

Ora abbiamo dimostrato che ogni pagina HTML ha un elemento `body`, e che anche il suo elemento `body` può essere stilizzato con CSS.

Ricorda, puoi definire lo stile del tuo elemento `body` proprio come qualsiasi altro elemento HTML, e tutti gli altri elementi erediteranno gli stili del tuo elemento `body`.

# --instructions--

Per prima cosa, crea un elemento `h1` con il testo `Hello World`.

Poi, diamo a tutti gli elementi della tua pagina il colore `green` aggiungendo `color: green;` alla dichiarazione di stile del tuo elemento `body`.

Infine, assegna al tuo elemento `body` il carattere tipografico `monospace` aggiungendo `font-family: monospace;` alla dichiarazione di stile del tuo elemento `body`.

# --hints--

Dovresti creare un elemento `h1`.

```js
assert.isNotEmpty(document.querySelectorAll('h1'));
```

Il tuo elemento `h1` dovrebbe avere il testo `Hello World`.

```js
assert.match(
  document.querySelector('h1').textContent,
  /hello world/i
);
```

Il tuo elemento `h1` dovrebbe avere un tag di chiusura.

```js
const commentlessCode = __helpers.removeHtmlComments(code);
assert.match(commentlessCode, /<\/h1>/g);
assert.match(commentlessCode, /<h1/g);
assert.lengthOf(commentlessCode.match(/<\/h1>/g), commentlessCode.match(/<h1/g).length);
```

Il tuo elemento `body` dovrebbe avere la proprietà `color` impostata su `green`.

```js
const bodyElement = document.querySelector('body');
const color = window.getComputedStyle(bodyElement)['color']; 
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Il tuo elemento `body` dovrebbe avere la proprietà `font-family` impostata su `monospace`.

```js
const bodyElement = document.querySelector('body');
const fontFamily = window.getComputedStyle(bodyElement)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

Il tuo elemento `h1` dovrebbe ereditare il carattere tipografico `monospace` dal tuo elemento `body`.

```js
const h1Element = document.querySelector('h1');
const fontFamily = window.getComputedStyle(h1Element)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

Il tuo elemento `h1` dovrebbe ereditare il colore `green` dal tuo elemento `body`.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }

</style>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }

</style>
<h1>Hello World!</h1>
```
