---
id: bad87fee1348bd9aedf04756
title: sovrascrivere gli stili nel CSS successivo
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDQug'
forumTopicId: 18253
dashedName: override-styles-in-subsequent-css
---

# --description--

La nostra classe `pink-text` ha sovrascritto la dichiarazione CSS del nostro elemento `body`!

Abbiamo appena dimostrato che le nostre classi sovrascrivono il CSS dell'elemento `body`. Quindi la domanda logica successiva è: cosa possiamo fare per sovrascrivere la nostra classe `pink-text`?

# --instructions--

Crea una classe CSS aggiuntiva chiamata `blue-text` che assegna il colore blu a un elemento. Assicurati che sia sotto la dichiarazione della tua classe `pink-text`.

Applica la classe `blue-text` al tuo elemento `h1` oltre alla classe `pink-text`, e vediamo quale prevale.

Applicare più attributi class a un elemento HTML si fa mettendo uno spazio tra di essi così:

```html
class="class1 class2"
```

**Nota:** Non importa in quale ordine le classi sono elencate nell'elemento HTML.

Tuttavia, l'ordine delle dichiarazioni `class` nella sezione `<style>` è ciò che conta. La seconda dichiarazione avrà sempre la precedenza sulla prima. Poiché `.blue-text` è dichiarata per seconda, sovrascrive gli attributi di `.pink-text`.

# --hints--

Il tuo elemento `h1` dovrebbe avere la classe `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Il tuo elemento `h1` dovrebbe avere la classe `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Sia `blue-text` che `pink-text` dovrebbero appartenere allo stesso elemento `h1`.

```js
assert.isTrue(document.querySelector('.pink-text').classList.contains('blue-text'));
```

Il tuo elemento `h1` dovrebbe essere blu.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }

  .blue-text {
    color: blue;
  }  
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```
