---
id: bad87fee1348bd9aede08718
title: Usa i valori RGB per colorare gli elementi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

Un altro modo per rappresentare i colori in CSS è usando i valori `RGB`.

Il valore `RGB` per il nero è così:

```css
rgb(0, 0, 0)
```

Il valore `RGB` per il bianco è così:

```css
rgb(255, 255, 255)
```

Invece di usare sei cifre esadecimali come fai con il codice esadecimale, con `RGB` specifichi la luminosità di ogni colore con un numero tra 0 e 255.

Se fai i calcoli, le due cifre per un colore equivalgono a 16 per 16, che ci dà un totale di 256 valori. Quindi `RGB`, che inizia a contare da zero, ha lo stesso numero esatto di valori possibili del codice esadecimale.

Ecco un esempio di come cambieresti lo sfondo `body` in arancione usando il suo codice RGB.

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

Sostituiamo il codice esadecimale nel colore di sfondo del nostro elemento `body` con il valore RGB per il nero: `rgb(0, 0, 0)`

# --hints--

Il tuo elemento `body` dovrebbe avere uno sfondo nero.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Dovresti usare `rgb` per dare al tuo elemento `body` uno sfondo nero.

```js
assert.match(__helpers.removeCssComments(code), /rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
