---
id: bad87fee1348bd9aedf08726
title: Usa il codice esadecimale per colori specifici
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

Sapevi che ci sono altri modi per rappresentare i colori in CSS? Uno di questi modi si chiama codice esadecimale, o codice hex in breve.

Di solito usiamo i <dfn>decimali</dfn>, o numeri in base 10, che usano i simboli da 0 a 9 per ogni cifra. Gli <dfn>esadecimali</dfn> (o <dfn>hex</dfn>) sono numeri in base 16. Questo significa che usano sedici simboli distinti. Come i decimali, i simboli da 0 a 9 rappresentano i valori da zero a nove. Poi A,B,C,D,E,F rappresentano i valori da dieci a quindici. In tutto, da 0 a F possono rappresentare una cifra in esadecimale, dandoci 16 valori possibili. Puoi trovare più informazioni sui <a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">numeri esadecimali qui</a>.

In CSS, possiamo usare 6 cifre esadecimali per rappresentare i colori, due per ciascuno dei componenti rosso (R), verde (G) e blu (B). Per esempio, `#000000` è nero ed è anche il valore più basso possibile. Puoi trovare più informazioni sul <a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">sistema di colori RGB qui</a>.

```css
body {
  color: #000000;
}
```

# --instructions--

Sostituisci la parola `black` nel colore di sfondo del nostro elemento `body` con la sua rappresentazione in codice esadecimale, `#000000`.

# --hints--

Il tuo elemento `body` dovrebbe avere il `background-color` nero.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Il codice esadecimale per il colore nero dovrebbe essere usato invece della parola `black`.

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
