---
id: bad87fee1348bd9aedf08721
title: Usa il codice esadecimale per mescolare i colori
challengeType: 0
videoUrl: 'https://scrimba.com/c/cK89PhP'
forumTopicId: 18359
dashedName: use-hex-code-to-mix-colors
---

# --description--

Per ripassare, i codici esadecimali usano 6 cifre esadecimali per rappresentare i colori, due per ciascun componente rosso (R), verde (G) e blu (B).

Da questi tre colori puri (rosso, verde e blu), possiamo variare le quantità di ciascuno per creare oltre 16 milioni di altri colori!

Per esempio, l'arancione è rosso puro, mescolato con un po' di verde e nessun blu. Nel codice esadecimale, questo si traduce in `#FFA500`.

La cifra `0` è il numero più basso nel codice esadecimale e rappresenta un'assenza completa di colore.

La cifra `F` è il numero più alto nel codice esadecimale e rappresenta la massima luminosità possibile.

# --instructions--

Sostituisci le parole dei colori nel nostro elemento `style` con i loro codici esadecimali corretti.

<table><tbody><tr><th>Colore</th><th>Codice esadecimale</th></tr><tr><td>Dodger Blue</td><td><code>#1E90FF</code></td></tr><tr><td>Verde</td><td><code>#00FF00</code></td></tr><tr><td>Arancione</td><td><code>#FFA500</code></td></tr><tr><td>Rosso</td><td><code>#FF0000</code></td></tr></tbody></table>

# --hints--

Il tuo elemento `h1` con il testo `I am red!` dovrebbe avere il rosso `color`.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Il `hex code` per il colore rosso dovrebbe essere usato invece della parola `red`.

```js
assert.match(code, /\.red-text\s*?{\s*?color\s*:\s*?(#FF0000|#F00)\s*?;?\s*?}/gi);
```

Il tuo elemento `h1` con il testo `I am green!` dovrebbe avere il verde `color`.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

Il `hex code` per il colore verde dovrebbe essere usato invece della parola `green`.

```js
assert.match(code, /\.green-text\s*?{\s*?color\s*:\s*?(#00FF00|#0F0)\s*?;?\s*?}/gi);
```

Il tuo elemento `h1` con il testo `I am dodger blue!` dovrebbe avere il dodger blue `color`.

```js
const blueText = document.querySelector('.dodger-blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(30, 144, 255)');
```

Il `hex code` per il colore dodger blue dovrebbe essere usato invece della parola `dodgerblue`.

```js
assert.match(code, /\.dodger-blue-text\s*?{\s*?color\s*:\s*?#1E90FF\s*?;?\s*?}/gi);
```

Il tuo elemento `h1` con il testo `I am orange!` dovrebbe avere l'arancione `color`.

```js
const orangeText = document.querySelector('.orange-text');
const color = window.getComputedStyle(orangeText)['color']; 

assert.strictEqual(color, 'rgb(255, 165, 0)');
```

Il `hex code` per il colore arancione dovrebbe essere usato invece della parola `orange`.

```js
assert.match(code, /\.orange-text\s*?{\s*?color\s*:\s*?#FFA500\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: black;
  }
  .green-text {
    color: black;
  }
  .dodger-blue-text {
    color: black;
  }
  .orange-text {
    color: black;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #FF0000;
  }
  .green-text {
    color: #00FF00;
  }
  .dodger-blue-text {
    color: #1E90FF;
  }
  .orange-text {
    color: #FFA500;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```
