---
id: bad87fee1348bd9aedf08805
title: Usa i selettori CSS per definire lo stile degli elementi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKMBT2'
forumTopicId: 18349
dashedName: use-css-selectors-to-style-elements
---

# --description--

Con CSS, ci sono centinaia di proprietà CSS che puoi usare per cambiare l'aspetto di un elemento nella tua pagina.

Quando hai inserito `<h2 style="color: red;">CatPhotoApp</h2>`, stavi definendo lo stile di quell'elemento `h2` individuale con CSS in linea, che sta per Cascading Style Sheets.

Questa è una modalità per specificare lo stile di un elemento, ma c'è un modo migliore per applicare il CSS.

In cima al tuo codice, crea un blocco `style` così:

```html
<style>
</style>
```

All'interno di quel blocco style, puoi creare un <dfn>selettore CSS</dfn> per tutti gli elementi `h2`. Per esempio, se vuoi che tutti gli elementi `h2` siano rossi, aggiungeresti una regola di stile che appare così:

```html
<style>
  h2 {
    color: red;
  }
</style>
```

Nota che è importante avere sia le parentesi graffe di apertura che di chiusura (`{` e `}`) attorno a ogni regola di stile per l'elemento. Devi anche assicurarti che la definizione dello stile del tuo elemento sia tra i tag di apertura e chiusura style. Infine, assicurati di aggiungere un punto e virgola alla fine di ogni regola di stile del tuo elemento.

# --instructions--

Elimina l'attributo style del tuo elemento `h2` e invece crea un blocco CSS `style`. Aggiungi il CSS necessario per rendere tutti gli elementi `h2` blu.

# --hints--

L'attributo `style` dovrebbe essere rimosso dal tuo elemento `h2`.

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

Dovresti creare un elemento `style`.

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
assert.isAtLeast(document.querySelectorAll('style:not(.fcc-hide-header)').length, 1);
```

Il tuo elemento `h2` dovrebbe essere blu.

```js
const h2Element = document.querySelector('h2');
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

La dichiarazione `h2` del tuo foglio di stile dovrebbe essere valida con un punto e virgola e una parentesi graffa di chiusura.

```js
assert.match(__helpers.removeCssComments(code), /h2\s*\{\s*color\s*:.*;\s*\}/g);
```

Tutti i tuoi elementi `style` dovrebbero essere validi e avere tag di chiusura.

```js
assert.match(__helpers.removeHtmlComments(code), /<\/style>/g);

const closingElementLength = __helpers.removeHtmlComments(code).match(/<\/style>/g).length;
const openingElementsLength = __helpers.removeHtmlComments(code).match(/<style((\s)*((type|media|scoped|title|disabled)="[^"]*")?(\s)*)*>/g).length;
assert.strictEqual(closingElementLength, openingElementsLength);
```

# --seed--

## --seed-contents--

```html
<h2 style="color: red;">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  h2 {
    color: blue;
  }
</style>
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
