---
id: bad87fee1348bd8aedf06756
title: sovrascrivere le dichiarazioni di classe definendo lo stile degli attributi id
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpDhB'
forumTopicId: 18251
dashedName: override-class-declarations-by-styling-id-attributes
---

# --description--

Abbiamo appena dimostrato che i browser leggono il CSS dall’alto verso il basso nell’ordine delle dichiarazioni. Questo significa che, in caso di conflitto, il browser userà la dichiarazione CSS che è stata scritta per ultima. Nota che anche se avessimo messo `blue-text` prima di `pink-text` nelle classi del nostro elemento `h1`, il browser guarderebbe comunque l’ordine delle dichiarazioni e non l’ordine del loro uso!

Ma non abbiamo ancora finito. Ci sono altri modi per sovrascrivere il CSS. Ti ricordi gli attributi id?

Sovrascriviamo le tue classi `pink-text` e `blue-text` e rendiamo arancione il tuo elemento `h1`, assegnando un id all’elemento `h1` e definendo lo stile di quell’id.

# --instructions--

Dai all’elemento `h1` l’attributo `id` con valore `orange-text`. Ricorda, gli stili per id si scrivono così:

```html
<h1 id="orange-text">
```

Lascia le classi `blue-text` e `pink-text` sull’elemento `h1`.

Crea una dichiarazione CSS per l’id `orange-text` nel tuo elemento `style`. Ecco un esempio di come si presenta:

```css
#brown-text {
  color: brown;
}
```

**Nota:** Non importa se dichiari questo CSS sopra o sotto la classe `pink-text`, perché l’attributo `id` avrà sempre la precedenza.

# --hints--

Il tuo elemento `h1` dovrebbe avere la classe `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Il tuo elemento `h1` dovrebbe avere la classe `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Il tuo elemento `h1` dovrebbe avere l’id `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'),'orange-text');
```

Dovrebbe esserci solo un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Il tuo id `orange-text` dovrebbe avere una dichiarazione CSS.

```js
assert.match(__helpers.removeCssComments(code), /#orange-text\s*{/gi);
```

Il tuo `h1` non dovrebbe avere attributi `style`.

```js
assert.notMatch(__helpers.removeHtmlComments(code), /<h1.*style.*>/gi);
```

Il tuo elemento `h1` dovrebbe essere arancione.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 165, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
  #orange-text {
    color: orange;
  }  
</style>
<h1 id="orange-text"  class="pink-text blue-text">Hello World!</h1>
```
