---
id: bad87fee1348bd9afdf08726
title: Usa la notazione oraria per specificare il margine di un elemento
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpybAd'
forumTopicId: 18345
dashedName: use-clockwise-notation-to-specify-the-margin-of-an-element
---

# --description--

Proviamo di nuovo, ma questa volta con `margin`.

Invece di specificare singolarmente le proprietĂ  `margin-top`, `margin-right`, `margin-bottom` e `margin-left` di un elemento, puoi specificarle tutte in una sola riga, cosĂ¬:

```css
margin: 10px 20px 10px 20px;
```

Questi quattro valori funzionano come un orologio: alto, destro, basso, sinistro, e produrranno lo stesso risultato di usare le istruzioni per il margine specifiche per lato.

# --instructions--

Usa la notazione oraria per dare all'elemento con la classe `blue-box` un margine di `40px` sul lato superiore e sinistro, ma solo `20px` sul lato inferiore e destro.

# --hints--

La tua classe `blue-box` dovrebbe dare agli elementi un margine superiore di `40px` di `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '40px');
```

La tua classe `blue-box` dovrebbe dare agli elementi un margine destro di `20px` di `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginRight = window.getComputedStyle(blueBox)['margin-right'];
assert.strictEqual(marginRight, '20px');
```

La tua classe `blue-box` dovrebbe dare agli elementi un margine inferiore di `20px` di `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginBottom = window.getComputedStyle(blueBox)['margin-bottom'];
assert.strictEqual(marginBottom, '20px');
```

La tua classe `blue-box` dovrebbe dare agli elementi un margine sinistro di `40px` di `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginLeft = window.getComputedStyle(blueBox)['margin-left'];
assert.strictEqual(marginLeft, '40px');
```

Dovresti usare la notazione oraria per impostare il margine della classe `blue-box`.

```js
const cssCode = __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(cssCode, /\.blue-box\s*{[\s\S]*margin[\s]*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    margin: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
