---
id: bad87fee1348bd9aedf08736
title: "definire lo stile dell'elemento body in HTML"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB77PHW'
forumTopicId: 18313
dashedName: style-the-html-body-element
---

# --description--

Ora ricominciamo da capo e parliamo di ereditarietà in CSS.

Ogni pagina HTML ha un elemento `body`.

# --instructions--

Possiamo dimostrare che l'elemento `body` esiste qui dandogli un `background-color` di nero.

Possiamo farlo aggiungendo quanto segue al nostro elemento `style`:

```css
body {
  background-color: black;
}
```

# --hints--

Il tuo elemento `body` dovrebbe avere il `background-color` nero.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color'];

assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

La tua regola CSS dovrebbe essere formattata correttamente con sia parentesi graffe di apertura che di chiusura.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

La tua regola CSS dovrebbe terminare con un punto e virgola.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

# --seed--

## --seed-contents--

```html
<style>

</style>
```

# --solutions--

```html
<style>
body {
  background-color: black;
}
</style>
```
