---
id: bad87fee1348bd9aedf08756
title: dare priorità a uno stile rispetto a un altro
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ8wnHv'
forumTopicId: 18258
dashedName: prioritize-one-style-over-another
---

# --description--

A volte i tuoi elementi HTML ricevono più stili che si contraddicono a vicenda.

Per esempio, il tuo elemento `h1` non può essere contemporaneamente verde e rosa.

Vediamo cosa succede quando creiamo una classe che rende il testo rosa e poi la applichiamo a un elemento. La nostra classe sovrascriverà la proprietà CSS `body` dell'elemento `color: green;`?

# --instructions--

Crea una classe CSS chiamata `pink-text` che assegna a un elemento il colore rosa.

Dai al tuo elemento `h1` la classe `pink-text`.

# --hints--

Il tuo elemento `h1` dovrebbe avere la classe `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Il tuo `<style>` dovrebbe avere una classe CSS `pink-text` che cambia il `color`.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*\{\s*color\s*:\s*.+\s*;?\s*\}/g);
```

Il tuo elemento `h1` dovrebbe essere rosa.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
</style>
<h1>Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```
