---
id: 68507cadbf36aa089a84ce2d
title: Tekeleza algoriti ya Luhn
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

Algoriti ya Luhn, inayojulikana pia kama algoriti ya "moduli 10" au "mod 10", ni fomula rahisi ya ukaguzi wa jumla inayotumika kuthibitisha nambari mbalimbali za utambulisho, kama nambari za kadi za mkopo. Hizi ni hatua za kuthibitisha nambari kwa kutumia algoriti ya Luhn:

- Kuanzia upande wa kulia, na kuondoa tarakimu ya mwisho kabisa (tarakimu ya ukaguzi), zidisha mara mbili thamani ya kila tarakimu nyingine.
- Ikiwa matokeo ya kuzidisha tarakimu ni zaidi ya `9`, jumlisha tarakimu hizo kupata tarakimu moja tu. Vinginevyo, unaweza kutoa `9` kutoka kwa matokeo.
- Chukua jumla ya tarakimu zote ikiwa ni pamoja na tarakimu ya ukaguzi.
- Ikiwa jumla ya tarakimu zote ni mara ya `10`, basi nambari ni halali; vinginevyo si halali.

Kwa mfano, tuchukulie nambari `453914881`. Hatua za kuithibitisha kwa kutumia algoriti ya Luhn ni:

 ```md 
 Nambari ya akaunti  4   5   3   9   1   4   8   8   1 
 Zidisha kila nyingine 4  10   3  18   1   8   8  16   1 
 Jumlisha tarakimu 4   1   3   9   1   8   8   7   1 
 ``` 

- Kisha jumlisha nambari zote, `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`.
- Kwa kuwa `42` si mara ya `10`, nambari si halali.

Katika maabara hii, utajenga kithibitishaji cha kadi za mkopo kwa kutumia algoriti ya Luhn.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kufafanua kitendakazi kinachoitwa `verify_card_number` kinachopokea mfuatano wa herufi za tarakimu (wakilisha nambari ya kadi) na kuthibitisha kama ni halali kulingana na algoriti ya Luhn.

2. Ndani ya kitendakazi cha `verify_card_number`:

   - Unapaswa kushughulikia dash au nafasi zozote zinazoweza kuwepo katika nambari ya kadi inayotumwa.
   - Rudisha `VALID!` ikiwa nambari ya kadi ni halali; vinginevyo, rudisha `INVALID!`.

Ukipata mradi umekamilika, unapaswa kuona ujumbe ufuatao kulingana na ingizo:

| Nambari ya kadi     | Ujumbe  |
| ------------------- | -------- |
| 453914889           | HALALI!  |
| 4111-1111-1111-1111 | HALALI!  |
| 1234 5678 9012 3456 | SI HALALI! |


# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `verify_card_number`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` inapaswa kurudisha `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` inapaswa kurudisha `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` inapaswa kurudisha `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` inapaswa kurudisha `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number` inapaswa kurudisha `VALID!` inapoitwa na nambari halali ya kadi ya mkopo.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number` inapaswa kurudisha `INVALID!` inapoitwa na nambari isiyo halali ya kadi ya mkopo.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
