---
id: 66ba762af611169359d9d369
title: 新聞記事を作成する
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. HTMLドキュメントのルート`font-size`を`24px`に設定してください。
1. すべての他の要素を含む`newspaper`クラスの要素を用意してください。
1. `.newspaper`要素には`font-size`が`16px`で、フォントが`Open Sans`、フォールバックフォントが`sans-serif`である必要があります。
1. `.newspaper`要素内には少なくとも7つの要素が必要です：新聞名用の`name`クラスの要素1つ、記事の日付用の`date`クラスの要素1つ、見出し用の`headline`クラスの要素1つ、副見出し用の`sub-headline`クラスの要素1つ、著者用の`author`クラスの要素1つ、そして`text`クラスの段落が2つ。それらすべてに記事の情報を入れてください。
1. `.name`要素には、ルート要素の`font-size`の2倍の`font-size`があり、`Times New Roman`フォントを使い、フォールバックフォントは`serif`にしてください。
1. `.name`と`.author`要素はCSSで全ての文字を大文字にしてください。
1. `.headline`要素には親要素の`font-size`の2倍の`font-size`があり、太字にしてください。
1. `.sub-headline`要素には`font-weight`が`100`で、`font-size`が親要素の`1.5`の`font-size`倍で、斜体にしてください。
1. `.author`はCSSで太字にしてください。
1. `.text`要素には`text-indent`が`20px`である必要があります。
1. `.text`要素には親要素の`line-height`の2倍の`font-size`が必要です。
1. `.text`要素の最初の文字は太字で親要素の`font-size`の2倍の大きさにしてください。これには`::first-letter`セレクターを使います。

# --hints--

`html`セレクターがあり、その`font-size`を`24px`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

`newspaper`クラスの要素を用意してください。

```js
assert.exists(document.querySelector('.newspaper'));
```

`.newspaper`セレクターがあり、その要素の`font-size`を`16px`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

`.newspaper`セレクターがあり、その要素の`font-family`を`'Open Sans', sans-serif`に設定してください。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

`name`要素内に`.newspaper`クラスの要素を用意してください。

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

`.name`要素には新聞名を入れてください。

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

`date`要素内に`.newspaper`クラスの要素を用意してください。

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

`.date`要素には日付を入れてください。

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

`headline`要素内に`.newspaper`クラスの要素を用意してください。

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

`.headline`要素には記事の見出しを入れてください。

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

`sub-headline`要素内に`.newspaper`クラスの要素を用意してください。

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

`.sub-headline`要素には記事の副見出しを入れてください。

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

`author`要素内に`.newspaper`クラスの要素を用意してください。

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

`.author`要素には記事の著者名を入れてください。

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

`text`要素内に`.newspaper`クラスの段落要素を少なくとも2つ用意してください。

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

`.text`要素には記事の本文を入れてください。

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

`.name`セレクターがあり、その要素の`font-size`を`2rem`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

`.name`セレクターがあり、その要素の`font-family`を`'Times New Roman', serif`に設定してください。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

`.name`要素には`text-transform`プロパティが`uppercase`に設定されています。

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

`.author`要素には`text-transform`プロパティが`uppercase`に設定されています。

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

`.headline`セレクターがあり、その要素の`font-size`を`2em`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

`.headline`セレクターがあり、その要素の`font-weight`を`bold`に設定してください。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

`.sub-headline`セレクターがあり、その要素の`font-weight`を`100`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

`.sub-headline`セレクターがあり、その要素の`font-size`を`1.5em`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

`.sub-headline`セレクターがあり、その要素の`font-style`を`italic`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

`.author`セレクターがあり、その要素の`font-weight`を`bold`に設定してください。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

`.text`セレクターがあり、その要素の`text-indent`を`20px`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

`.text`セレクターがあり、その要素の`line-height`を`2em`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

`.text::first-letter`セレクターがあり、その要素の`font-weight`を`bold`に設定してください。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

`.text::first-letter`セレクターがあり、その要素の`font-size`を`2em`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
