---
id: 699d20623e57e8ffee885451
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

La función `parseCard` tomará una cadena de catálogo sin procesar y devolverá un objeto estructurado. Antes de construir tu función, configura una forma de ver su salida mientras trabajas.

Llama a `parseCard` con `rawCatalogCards[2]` como argumento, asigna el resultado a una variable llamada `cardResult` y muestra `cardResult` en la consola.

En este momento `parseCard` no tiene cuerpo, por lo que devuelve `undefined`. A medida que agregues código a la función en los siguientes pasos, la salida mostrada te indicará exactamente qué está devolviendo la función en cada etapa.

# --hints--

Debes llamar a `parseCard` con `rawCatalogCards[2]` como argumento.

```js
assert.match(__helpers.removeJSComments(code), /parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)/);
```

Debes asignar el resultado de llamar a `parseCard` con `rawCatalogCards[2]` a una variable llamada `cardResult`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+cardResult\s*=\s*parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)\s*;?/);
```

Debes mostrar `cardResult` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*cardResult\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {}

--fcc-editable-region--

--fcc-editable-region--
```
