---
id: 699d20623e57e8ffee885457
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Algunas fichas de catálogo tienen campos faltantes. Por ejemplo, la entrada 16 tiene un autor en blanco, por lo que `parseCard` actualmente devuelve una cadena vacía para ese campo.

El operador `||` ofrece una alternativa: `a || b` evalúa a `a` cuando `a` es verdadero, y a `b` en caso contrario. Como una cadena vacía es falsa, `"" || "Unknown"` evalúa a `"Unknown"`.

Actualiza `title`, `author` y `location` en la sentencia return para usar `|| "Unknown"` como alternativa. Deja `year` para el siguiente paso.

# --hints--

Cuando se llame con una cadena que tenga el campo autor en blanco, la propiedad `author` del objeto devuelto debería ser `"Unknown"`.

```js
assert.strictEqual(parseCard("X |  | 2000 | Y").author, "Unknown");
```

Cuando se llame con una cadena que tenga el campo título en blanco, la propiedad `title` del objeto devuelto debería ser `"Unknown"`.

```js
assert.strictEqual(parseCard(" | A | 2000 | Y").title, "Unknown");
```

Cuando se llame con una cadena que tenga el campo ubicación en blanco, la propiedad `location` del objeto devuelto debería ser `"Unknown"`.

```js
assert.strictEqual(parseCard("X | A | 2000 | ").location, "Unknown");
```

Aún no deberías agregar una alternativa para `year`.

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "");
```

Cuando todos los campos estén presentes, `parseCard` debería devolver sus valores reales y no `"Unknown"`.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title,
    author: author,
    year: year,
    location: location
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
