---
id: 699d20623e57e8ffee885454
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

El primer paso para analizar una ficha de catálogo es separar los cuatro campos. El método `.split()` divide una cadena en un arreglo de subcadenas en cada ocurrencia del separador que proporciones.

Por ejemplo:

```js
"The Shining | King, Stephen | 1977 | Shelf K1".split("|")
// ["The Shining ", " King, Stephen ", " 1977 ", " Shelf K1"]
```

Dentro de `parseCard`, llama a `.split("|")` sobre `rawString` y devuelve el resultado. Después de este paso, la consola debería mostrar un arreglo con cuatro cadenas.

**Nota:** Puede que notes espacios en blanco adicionales alrededor de los valores. Los limpiarás en el siguiente paso.

# --hints--

Debes llamar a `.split("|")` sobre `rawString` dentro de tu función `parseCard`.

```js
assert.match(parseCard.toString(), /rawString\s*\.\s*split\s*\(\s*('|"|`)\|\1\s*\)\s*;?/);
```

`parseCard` debe devolver un arreglo.

```js
assert.isArray(parseCard("A | B | C | D"));
```

Cuando se llame con `"A | B | C | D"`, `parseCard` debe devolver un arreglo con cuatro elementos.

```js
assert.strictEqual(parseCard("A | B | C | D").length, 4);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
--fcc-editable-region--
  
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
