---
id: 699d1f8e51fcdc8230f779fb
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, construirás un sistema de catálogo para una biblioteca patrimonial. Analizarás datos sin procesar de libros, buscarás y agruparás entradas, validarás la calidad de los datos y exportarás resultados a diferentes formatos.

Cada ficha del catálogo se almacena como una cadena única en el formato `"Title | Author | Year | Location"`, con campos separados por el carácter de barra vertical (`|`).

Usarás un arreglo de estas cadenas a lo largo del taller para construir y probar tus funciones de catálogo. Crea un arreglo llamado `rawCatalogCards` y añade las siguientes dos cadenas:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

Debes declarar una variable llamada `rawCatalogCards`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` debe ser un arreglo.

```js
assert.isArray(rawCatalogCards);
```

El primer elemento de `rawCatalogCards` debe ser `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

El segundo elemento de `rawCatalogCards` debe ser `"The Shining | King, Stephen | 1977 | Shelf K1"`.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
