---
id: 69aa11c2584523e6e991330c
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

El método `.trim()` elimina los espacios en blanco de ambos extremos de una cadena. Por ejemplo, `" hello ".trim()` devuelve `"hello"`.

Dentro del cuerpo del ciclo, agrega `parts[i].trim()` a `trimmedParts`. Luego actualiza la instrucción `return` existente para que devuelva `trimmedParts` en lugar de `parts`.

Después de este paso, la salida registrada debería mostrar valores limpios sin espacios adicionales.

# --hints--

Debes agregar `parts[i].trim()` a `trimmedParts` dentro del ciclo.

```js
assert.match(parseCard.toString(), /trimmedParts\s*\.\s*push\s*\(\s*parts\s*\[\s*\w+\s*\]\s*\.\s*trim\s*\(\s*\)\s*\)/);
```

`parseCard` debería devolver `trimmedParts`.

```js
const explorer = await __helpers.Explorer(code);
const { parseCard } = explorer.allFunctions;
assert.isTrue(parseCard.hasReturn("trimmedParts"));
```

`parseCard` debería devolver un arreglo donde cada elemento no tenga espacios en blanco al inicio ni al final.

```js
assert.isArray(parseCard("  A  |  B  |  C  |  D  "));
```

Cuando se llame con `"  A  |  B  |  C  |  D  "`, el primer elemento del arreglo devuelto debería ser `"A"`.

```js
assert.strictEqual(parseCard("  A  |  B  |  C  |  D  ")[0], "A");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
--fcc-editable-region--

  }
  return parts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
