---
id: 69708d59833ceded7d00541f
title: تنفيذ عمليات القائمة المرتبطة
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

في هذا المختبر، ستنفذ عمليات إضافية لهيكل البيانات القائمة المرتبطة، بناءً على الوظائف الأساسية لإضافة العقد وإزالتها.

**الهدف**: نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة `contains` تقبل قائمة مرتبطة وعنصرًا. تُرجع `true` إذا كان العنصر المحدد موجودًا في القائمة المرتبطة، و`false` خلاف ذلك.
1. يجب أن تملك دالة `getAt` تقبل قائمة مرتبطة وفهرسًا. تُرجع العنصر عند الفهرس المعطى في القائمة المرتبطة. إذا كان الفهرس خارج النطاق، تُرجع `undefined`.
1. يجب أن تملك دالة `insertAt` تقبل قائمة مرتبطة، وفهرسًا، وعنصرًا. تُدرج العنصر المعطى في الموضع المحدد في القائمة المرتبطة. إذا كان الفهرس خارج النطاق، لا تُعدل القائمة.
1. يجب أن تملك دالة `removeAt` تقبل قائمة مرتبطة وفهرسًا. تزيل العقدة عند الفهرس المعطى في القائمة المرتبطة. إذا كان الفهرس خارج النطاق، لا تُعدل القائمة.
1. يجب أن تملك دالة `clear` تقبل قائمة مرتبطة. تزيل كل العناصر من القائمة المرتبطة، معيدة إياها إلى حالة فارغة.

**ملاحظة:** بعض الاختبارات اللاحقة تعتمد على الطرق السابقة. على سبيل المثال، إذا لم تُنفذ دالة `getAt` بشكل صحيح، قد تفشل اختبارات دوال مثل `insertAt` و`removeAt` حتى لو كانت تلك الدوال قريبة من الصحة.

# --hints--

يجب أن تُرجع `initList` كائنًا مع تعيين `head` إلى `null` و`length` إلى `0`.

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

يجب أن تزيد دالة `add` طول القائمة بمقدار واحد.

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

يجب أن تقلل دالة `remove` طول القائمة بمقدار واحد.

```js
const list = initList();  
add(list, 'apple');  
add(list, 'banana');  
remove(list, 'apple');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(list.head.element, 'banana');  
assert.isNull(list.head.next);
```

يجب أن تملك دالة `contains`.

```js
assert.isFunction(contains);
```

يجب أن تُرجع دالة `contains` `true` إذا كان العنصر موجودًا في القائمة.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

يجب أن تُرجع دالة `contains` `false` إذا لم يكن العنصر موجودًا في القائمة.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

يجب أن تُرجع دالة `contains` `false` للقائمة الفارغة.

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

يجب أن تملك دالة `getAt`.

```js
assert.isFunction(getAt);
```

يجب أن تُرجع دالة `getAt` العنصر عند فهرس معين.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

يجب أن تُرجع دالة `getAt` `undefined` للفهارس غير الصالحة.

```js
const list = initList();  
add(list, 'apple');  
assert.isUndefined(getAt(list, -1));  
assert.isUndefined(getAt(list, 1));  
assert.isUndefined(getAt(list, 5));
```

يجب أن تُرجع دالة `getAt` `undefined` للقائمة الفارغة.

```js
const list = initList();  
assert.isUndefined(getAt(list, 0));
```

يجب أن تملك دالة `insertAt`.

```js
assert.isFunction(insertAt);
```

يجب أن تُدرج دالة `insertAt` عنصرًا في بداية القائمة.

```js
const list = initList();  
add(list, 'second');  
add(list, 'third');  
insertAt(list, 0, 'first');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

يجب أن تُدرج دالة `insertAt` عنصرًا في منتصف القائمة.

```js
const list = initList();  
add(list, 'first');  
add(list, 'third');  
insertAt(list, 1, 'second');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

يجب أن تُدرج دالة `insertAt` عنصرًا في نهاية القائمة.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

يجب أن تُدرج دالة `insertAt` في قائمة فارغة عند الفهرس `0`.

```js  
const list = initList();  
insertAt(list, 0, 'first');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.isNull(list.head.next); 
```

يجب ألا تُعدل دالة `insertAt` القائمة للفهارس غير الصالحة.

```js
const list = initList();  
add(list, 'apple');  
insertAt(list, -1, 'invalid');  
insertAt(list, 10, 'invalid');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
assert.isUndefined(getAt(list, 1));
```

يجب أن تملك دالة `removeAt`.

```js
assert.isFunction(removeAt);
```

يجب أن تزيل دالة `removeAt` العنصر عند فهرس محدد.

```js
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'third');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

يجب أن تزيل دالة `removeAt` العنصر الأول بشكل صحيح.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

يجب أن تزيل دالة `removeAt` العنصر الأخير بشكل صحيح.

```js  
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 2);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

يجب أن تعيد دالة `removeAt` تعيين القائمة عند إزالة العنصر الوحيد.

```js
const list = initList();  
add(list, 'only');  
removeAt(list, 0);  
assert.strictEqual(list.length, 0);  
assert.isNull(list.head);  
assert.isTrue(isEmpty(list));
```

يجب ألا تُعدل دالة `removeAt` القائمة للفهارس غير الصالحة.

```js
const list = initList();  
add(list, 'apple');  
removeAt(list, -1);  
removeAt(list, 10);  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
```

يجب أن تملك دالة `clear`.

```js
assert.isFunction(clear);
```

يجب أن تزيل دالة `clear` كل العناصر من القائمة.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

يجب أن تسمح دالة `clear` بإعادة استخدام القائمة.

```js  
const list = initList();  
add(list, 'first');  
clear(list);  
clear(list);  
add(list, 'second');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'second');  
assert.isNull(list.head.next); 
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  
}

function getAt(list, index) {
 
}

function insertAt(list, index, element) {
  
}

function removeAt(list, index) {
 
}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
