---
id: 66f53dc2c5bd6a11d6c3282f
title: Створіть застосунок генератора паролів
challengeType: 26
dashedName: lab-password-generator
---

# --description--

У цій лабораторній роботі ви попрактикуєтеся у використанні функцій, створюючи випадковий генератор паролів.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви маєте створити функцію з назвою `generatePassword`, яка приймає параметр, що вказує довжину згенерованого пароля. Ви можете назвати параметр як завгодно.
2. Ваша функція має повертати рядок, який представляє випадково згенерований пароль. Ви маєте використати наступний рядок і різні методи `Math`, щоб допомогти повернути новий рядок із випадковими символами: `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`.
3. Ви маєте визначити змінну з назвою `password` і призначити їй результат виклику функції `generatePassword` з числовим аргументом, що представляє бажану довжину пароля.
4. Ви маєте мати `console.log`, який виводить у консоль один рядок, утворений конкатенацією повідомлення `Generated password:` і змінної `password`, розділених пробілом.

# --hints--

Ви маєте мати функцію `generatePassword` з параметром. Ви можете назвати параметр як завгодно.

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

Ваша функція `generatePassword` має повертати рядок.

```js
const result = generatePassword(5);
assert.isString(result);
```

Ваша функція `generatePassword` має повертати новий рядок правильної довжини.

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

Ваша функція має повертати випадково згенерований пароль із дійсними символами.

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

Ваша функція має повертати випадково згенерований пароль, який містить більше одного унікального символу.

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

Ваша функція має повертати новий випадковий рядок щоразу, коли її викликають.

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

Ви маєте мати змінну `password`.

```js
assert.isDefined(password);
```

Ваша змінна `password` має бути рядком.

```js
assert.isString(password);
```

Ви маєте викликати функцію `generatePassword` з числовим аргументом і зберегти повернений пароль у змінній `password`.

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

Ви маєте вивести в консоль один рядок, що поєднує `Generated password:` і `password`, розділені одним пробілом, використовуючи `+` або шаблонний літерал.

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
