---
id: 66a84111965a0c46df6bbd0a
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

في الدروس السابقة، تعلمت كيف تربط عنصر `label` مع `input` بهذا الشكل:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

يجب أن تتطابق الخاصية `for` في عنصر `label` مع الخاصية `id` في عنصر `input`. يُعرف هذا بالربط الصريح.

أسفل عنصر `legend` الخاص بك، أضف عنصر `label` يحتوي على النص `Name (required):`. عرّف خاصية `for` بقيمة `"full-name"`.

ثم أسفل عنصر `label`، أضف عنصر `input` بدون خصائص. في الخطوات التالية، ستضيف الخصائص اللازمة.

# --hints--

يجب أن يكون لديك وسم افتتاحي لـ `label`.

```js
assert.match(code, /<label.*>/);
```

يجب أن يكون لديك وسم إغلاق لـ `label`.

```js
assert.match(code, /<\/label>/);
```

يجب أن يكون عنصر `label` موضوعًا أسفل عنصر `legend` داخل عنصر `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

يجب أن يحتوي عنصر `label` على خاصية `for` بقيمة `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

يجب أن يحتوي عنصر `label` على النص `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

يجب أن يكون لديك عنصر `input` أسفل عنصر `label`.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
