---
id: 66a83e5e491625454b6f62c3
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

تتكون النماذج من `inputs` حيث يمكن للمستخدمين إدخال بياناتهم. يمكنك تجميع المُدخلات المرتبطة معًا باستخدام عنصر `fieldset`.

إليك مثال على استخدام عنصر `fieldset`:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

داخل عنصر `form` الخاص بك، أضف عنصر `fieldset`.

# --hints--

يجب أن يكون لديك وسم افتتاحي لـ `fieldset`.

```js
assert.match(code, /<fieldset>/i);
```

يجب أن يكون لديك وسم إغلاق لـ `fieldset`.

```js
assert.match(code, /<\/fieldset>/i);
```

يجب أن يكون عنصر `fieldset` داخل عنصر `form`.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
