---
id: 66a83fec026a7a4631e084d2
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

عند العمل مع عناصر `fieldset`، من الشائع استخدام تسمية توضيحية لوصف مجموعة المُدخَلات. يمكنك استخدام عنصر `legend` لهذا الغرض.

إليك مثالًا على استخدام عنصر `legend`:

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

داخل عنصر `fieldset` الخاص بك، أضف عنصر `legend` بالنص `Personal Information`.

# --hints--

يجب أن يكون لديك وسم افتتاحي لـ `legend`.

```js
assert.match(code, /<legend>/);
```

يجب أن يكون لديك وسم إغلاق لـ `legend`.

```js
assert.match(code, /<\/legend>/);
```

يجب أن يكون عنصر `legend` داخل عنصر `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

يجب أن يحتوي عنصر `legend` على النص `Personal Information`.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
