---
id: 66a83601cd819e37f0dccd14
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتدرب على العمل مع نماذج HTML من خلال بناء نموذج ملاحظات الفندق.

للنص التمهيدي، ستعرض العنوان الرئيسي متبوعًا بملاحظة قصيرة حول ترك الملاحظات.

داخل عنصر `body`، أضف عنصر `header`.

داخل عنصر `header`، أضف عنصر `h1`. يجب أن يكون نص عنصر `h1` هو `Hotel Feedback Form`.

أسفل عنصر `h1`، أضف عنصر `p`. يجب أن يكون نص عنصر `p` هو `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

يجب أن يكون لديك وسم افتتاحي `header`.

```js
assert.match(code, /<header>/i);
```

يجب أن يكون لديك وسم إغلاق `header`.

```js
assert.match(code, /<\/header>/i);
```

يجب أن يكون لديك وسم افتتاحي `h1`.

```js
assert.match(code, /<h1>/i);
```

يجب أن يكون لديك وسم إغلاق `h1`.

```js
assert.match(code, /<\/h1>/i);
```

يجب أن يحتوي عنصر `h1` على النص `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

يجب أن يكون عنصر `h1` داخل عنصر `header`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

يجب أن يكون لديك وسم افتتاحي `p`.

```js
assert.match(code, /<p>/i);
```

يجب أن يكون لديك وسم إغلاق `p`.

```js
assert.match(code, /<\/p>/i);
```

يجب أن يحتوي عنصر الفقرة على النص `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

يجب أن يكون عنصر الفقرة داخل `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

يجب أن يأتي عنصر الفقرة بعد عنصر `h1`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
