---
id: 66a83bdcf425e7446900b7c4
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

في الدروس السابقة، تعلمت كيف تتعامل مع عنصر `form` بهذا الشكل:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

تُستخدم الخاصية `action` لتحديد المكان الذي تُرسل إليه بيانات النموذج عند إرسال النموذج.

تُستخدم الخاصية `method` لتحديد طريقة HTTP التي تُستخدم عند إرسال بيانات النموذج. أكثر الطرق شيوعًا هي `GET` و `POST`.

**ملاحظة**: ستتعلم كيف تعمل طرق HTTP في الوحدات التالية.

داخل عنصر `main` الخاص بك، أضف عنصر `form` مع خاصية `action` مضبوطة على `"https://hotel-feedback.freecodecamp.org"` وخاصية `method` مضبوطة على `"POST"`.

# --hints--

يجب أن يكون لديك وسم افتتاحي `form`.

```js
assert.match(code, /<form.*>/i);
```

يجب أن يكون لديك وسم إغلاق `form`.

```js
assert.match(code, /<\/form>/i);
```

يجب أن يحتوي عنصر `form` الخاص بك على خاصية `action` مضبوطة على `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

يجب أن يحتوي عنصر `form` الخاص بك على خاصية `method` مضبوطة على `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
