---
id: 66a972137acd1179fa3fe8a0
title: الخطوة 26
challengeType: 0
dashedName: step-26
---

# --description--

داخل عنصر `select`، أضف العناصر الخمسة التالية من نوع `option` مع هذه القيم المقابلة لنص `option` و`value` الخاصية:

**قيم الخاصية:**

- ضعيف
- مرضٍ
- جيد
- جيد جدًا
- ممتاز

**نص عنصر الخيار:**

- ضعيف
- مرضٍ
- جيد
- جيد جدًا
- ممتاز

# --hints--

يجب أن يكون لديك عنصر `option` بقيمة مضبوطة على `"poor"`.

```js
assert.exists(document.querySelector('option[value="poor"]'));
```

يجب أن يحتوي عنصر `option` الخاص بك الذي يحمل `value` بقيمة `"poor"` على النص `Poor`.

```js
assert.strictEqual(document.querySelector('option[value="poor"]')?.textContent.trim(), 'Poor');
```

يجب أن يكون لديك عنصر `option` بقيمة `value` مضبوطة على `"satisfactory"`.

```js
assert.exists(document.querySelector('option[value="satisfactory"]'));
```

يجب أن يحتوي عنصر `option` الخاص بك الذي يحمل `value` بقيمة `"satisfactory"` على النص `Satisfactory`.

```js
assert.strictEqual(document.querySelector('option[value="satisfactory"]')?.textContent.trim(), 'Satisfactory');
```

يجب أن يكون لديك عنصر `option` بقيمة `value` مضبوطة على `"good"`.

```js
assert.exists(document.querySelector('option[value="good"]'));
```

يجب أن يحتوي عنصر `option` الخاص بك الذي يحمل `value` بقيمة `"good"` على النص `Good`.

```js

assert.strictEqual(document.querySelector('option[value="good"]')?.textContent.trim(), 'Good');
```

يجب أن يكون لديك عنصر `option` بقيمة مضبوطة على `"very-good"`.

```js
assert.exists(document.querySelector('option[value="very-good"]'));
```

عنصر `option` الذي يحمل الخاصية `value` بقيمة `"very-good"` يجب أن يحتوي على النص `Very Good`.

```js
assert.strictEqual(document.querySelector('option[value="very-good"]')?.textContent.trim(), 'Very Good');
```

يجب أن يكون لديك عنصر `option` بقيمة مضبوطة على `"excellent"`.

```js
assert.exists(document.querySelector('option[value="excellent"]'));
```

يجب أن يحتوي عنصر `option` الخاص بك الذي يحمل `value` بقيمة `"excellent"` على النص `Excellent`.

```js
assert.strictEqual(document.querySelector('option[value="excellent"]')?.textContent.trim(), 'Excellent');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
          --fcc-editable-region--
            
          --fcc-editable-region--
          </select>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
