---
id: 66a93bbe65a26169dbf3bc39
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

يجب أن يجمع نموذج تقييم الفندق عنوان بريد إلكتروني من المستخدم أيضًا.

ابدأ بإضافة عنصر `label` جديد يحتوي على النص `Email address (required):` إلى النموذج. يجب أن يحتوي عنصر `label` على خاصية `for` بقيمة `"email"`.

# --hints--

يجب أن يكون لديك عنصر `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

يجب أن يحتوي عنصر `label` على النص `Email address (required):`.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

يجب أن يحتوي عنصر `label` على خاصية `for` بقيمة `"email"`.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
