---
id: 66a962954f4e0d76223b37ed
title: الخطوة 21
challengeType: 0
dashedName: step-21
---

# --description--

بعد ذلك، أضف مربع اختيار `input` آخر مع تعيين خاصيتي `id` و `value` إلى `"location"`، وخاصية `name` إلى `"choice"`.

بالنسبة لعنصر `label`، يجب أن يكون نص `Location` وخاصية `for` معيينين إلى `"location"`.

أسفل ذلك العنصر `label`، أضف مربع اختيار `input` آخر مع تعيين خاصيتي `id` و `value` إلى `"reputation"`، وخاصية `name` إلى `"choice"`.

بالنسبة لعنصر `label`، يجب أن يكون نص `Reputation` وخاصية `for` معيينين إلى `"reputation"`.

# --hints--

يجب أن يكون لديك مربع اختيار `input` ثالث.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

يجب أن يحتوي مربع الاختيار `input` الثالث على خاصية `id` معينة إلى `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

يجب أن يحتوي مربع الاختيار `input` الثالث على خاصية `name` معينة إلى `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'choice');
```

يجب أن يحتوي مربع الاختيار `input` الثالث على خاصية `value` معينة إلى `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

يجب أن يكون لديك عنصر `label` بعد مربع الاختيار `input` مع النص `Location`.

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

يجب أن يحتوي `label` على خاصية `for` معينة إلى `"location"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

يجب أن يكون لديك مربع اختيار `input` رابع أسفل `label` الخاص بك.

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

يجب أن يحتوي مربع الاختيار `input` الرابع على خاصية `id` معينة إلى `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

يجب أن يحتوي مربع الاختيار `input` الرابع على خاصية `name` معينة إلى `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'choice');
```

يجب أن يحتوي مربع الاختيار `input` الرابع على خاصية `value` معينة إلى `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

يجب أن يكون لديك عنصر `label` مع النص `Reputation` أسفل مربع الاختيار الخاص بالسمعة.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

يجب أن يحتوي `label` على خاصية `for` معينة إلى `"reputation"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
