---
id: 66a9689b1bf24b7750898a1b
title: الخطوة 23
challengeType: 0
dashedName: step-23
---

# --description--

بالنسبة لآخر عنصر `input` و`label` داخل هذه الـ `fieldset`، أضف مربع اختيار `input` مع خاصيتي `id` و`value` مضبوطتين على `"price"`، وخاصية `name` مضبوطة على `"choice"`.

ثم أضف عنصر `label` يحتوي على النص `Price` وخاصية `for` مضبوطة على `"price"`.

الآن يمكنك اختبار الـ `form` الخاص بك باختيار مربعات الاختيار المختلفة.

# --hints--

يجب أن يكون لديك مربع اختيار خامس من نوع `input`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"]:nth-of-type(5)'));
```

يجب أن يحتوي مربع اختيار `input` الخامس على خاصية `id` مُعيّنة إلى `"price"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('id'), 'price');
```

يجب أن يحتوي مربع اختيار `input` الخامس على خاصية `name` مضبوطة على `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('name'), 'choice');
```

يجب أن يحتوي مربع اختيار `input` الخامس على خاصية `value` مضبوطة على `"price"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('value'), 'price');
```

يجب أن يكون لديك عنصر `label` بعد مربع اختيار `input` مع النص `Price`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.textContent?.trim(), 'Price');
```

يجب أن يحتوي الـ `label` الخاص بك على خاصية `for` مضبوطة على `"price"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.getAttribute('for'), 'price');
```   

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
