---
id: 66a93730719e1f68410cce54
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

تُستخدم خاصية `name` لتحديد بيانات النموذج بعد إرسالها إلى الخادم.

إليك مثالًا على كيفية استخدام خاصية `name`:

```html
<input type="email" name="email">
```

أضف خاصية `name` إلى عنصر `input` بقيمة `"name"`.

# --hints--

لا يجب إزالة عنصر `input` الموجود.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

لا يجب إنشاء عنصر `input` جديد.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

يجب أن يحتوي عنصر `input` على خاصية `name` بقيمة `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
