---
id: 66a954b2bcddba72076c1857
title: الخطوة 17
challengeType: 0
dashedName: step-17
---

# --description--

أسفل عنصر `label` الخاص بك، أضف زر اختيار `radio` مع تعيين `id` إلى `"no-option"`، وخاصية `name` إلى `"hotel-stay"`، وخاصية `value` إلى `"no"`.

أسفل زر الاختيار `radio` الجديد، أضف عنصر `label` آخر مع تعيين خاصية `for` إلى `"no-option"`. يجب أن يكون نص `label` هو `No`.

عند الانتهاء، يمكنك الآن تجربة أزرار الاختيار باختيار خيار واحد في كل مرة.

# --hints--

يجب أن يكون لديك عنصر `input` أسفل أول عنصر `label` لديك.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

يجب أن يحتوي عنصر `input` الخاص بك على خاصية `type` بقيمة `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

يجب أن يحتوي عنصر `input` الخاص بك على خاصية `id` بقيمة `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

يجب أن يحتوي عنصر `input` الخاص بك على خاصية `name` بقيمة `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

يجب أن يحتوي عنصر المُدخَل على خاصية `value` مضبوطة على `"no"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) label + input')?.value,
  'no'
);
```

يجب أن يكون لديك عنصر `label` ثاني أسفل عنصر `input` الثاني لديك.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

يجب أن يحتوي عنصر `label` الخاص بك على خاصية `for` بقيمة `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

يجب أن يحتوي عنصر `label` الخاص بك على النص `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
