---
id: 66a937e74920ba68ebe5e86d
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

في الدروس السابقة، تعلمت كيف تتعامل مع خاصيتي `placeholder` و `required` بهذه الطريقة:

```html
<input type="text" placeholder="e.g., John Doe" required>
```

أضف إلى عنصر `input` الموجود خاصية `placeholder` بقيمة `"e.g., John Doe"`.

أيضًا، أضف خاصية `required` إلى عنصر `input`.

# --hints--

لا يجب أن تزيل عنصر `input` الموجود.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

يجب أن يحتوي عنصر `input` على خاصية `placeholder` بقيمة `"e.g., John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

يجب أن يحتوي عنصر `input` على خاصية `required`.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
