---
id: 66a94ea5df66236ebed260e8
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

القسم التالي في النموذج سيكون مسؤولًا عن سؤال المستخدمين إذا كانوا قد أقاموا في الفندق من قبل.

ابدأ بإضافة عنصر `fieldset`.

داخل عنصر `fieldset`، أضف عنصر `legend` بالنص `Was this your first time at our hotel?`.

# --hints--

يجب أن يكون لديك عنصر `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 2);
```

يجب أن يكون لديك عنصر `legend` داخل عنصر `fieldset` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('fieldset legend'), 2);
```

يجب أن يحتوي عنصر `legend` الخاص بك على النص `Was this your first time at our hotel?`.

```js
assert.strictEqual(document.querySelectorAll('fieldset legend')[1]?.textContent.trim(), 'Was this your first time at our hotel?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
