---
id: 66a95d0eff8168747805f1f3
title: الخطوة 19
challengeType: 0
dashedName: step-19
---

# --description--

عندما تريد أن يختار المستخدم عدة خيارات من قائمة، يمكنك استخدام مربعات اختيار.

إليك مثالًا على كيفية العمل مع مربعات الاختيار المتعلقة بخيارات الطعام:

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

تُستخدم الخاصية `value` لتحديد القيمة التي سترسل إلى الخادم عند إرسال النموذج.

أسفل عنصر `legend`، أضف مُدخَل مربع اختيار `input` مع خاصيتي `id` و`value` مضبوطتين على `"ads"`، وخاصية `name` مضبوطة على `"choice"`.

أسفل مربع اختيار `input`، أضف عنصر `label` يحتوي على النص `Social Media Ads`. يجب تعيين خاصية `for` إلى `"ads"`.

# --hints--

يجب أن يكون لديك مُدخَل مربع اختيار `input` أسفل عنصر `legend`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

يجب أن يحتوي مربع اختيار `input` على خاصية `id` مضبوطة على `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

يجب أن يحتوي مربع اختيار `input` على خاصية `name` مضبوطة على `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "choice");
```

يجب أن يحتوي مربع اختيار `input` على خاصية `value` مضبوطة على `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

يجب أن يكون لديك عنصر `label` أسفل عنصر مربع اختيار `input`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

يجب أن يحتوي عنصر `label` على النص `Social Media Ads`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

يجب أن يحتوي عنصر `label` على خاصية `for` بقيمة `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
