---
id: 66a9521bc70162712caf118d
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

إذا أردت أن يختار المستخدمون خيارًا واحدًا من قائمة خيارات، يمكنك استخدام مجموعة من أزرار الاختيار.

إليك مثالًا على زري اختيار:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

في هذا المثال، تُجمَع أزرار الاختيار معًا باستخدام نفس قيمة خاصية `name`. هذا يعني أنه يمكن اختيار زر اختيار واحد فقط في كل مرة.

أسفل عنصر `legend`، أضف زر اختيار `radio` مع تعيين الخاصية `id` إلى `"yes-option"`، والخاصية `name` إلى `"hotel-stay"`، والخاصية `value` إلى `"yes"`.

أسفل زر الاختيار `radio`، أضف عنصر `label` يحتوي على النص `Yes` مع تعيين الخاصية `for` إلى `"yes-option"`.

# --hints--

يجب أن يكون لديك زر اختيار `radio` أسفل عنصر `legend`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

يجب أن يحتوي زر الاختيار `radio` على خاصية `id` مع تعيينها إلى `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

يجب أن يحتوي زر الاختيار `radio` على خاصية `name` مع تعيينها إلى `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

يجب أن يحتوي زر الاختيار على خاصية `value` مع تعيينها إلى `"yes"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

يجب أن يكون لديك عنصر `label` أسفل زر الاختيار `radio`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

يجب أن يحتوي عنصر `label` على خاصية `for` مع تعيينها إلى `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

يجب أن يحتوي عنصر `label` على النص `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
