---
id: 66a930b20f589b6664c51cb0
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

عندما يُدخل المستخدم اسمه الكامل، سيقبل الـ `input` النص العادي.

في الدروس السابقة، تعلمت كيف تتعامل مع الخاصية `type` بهذا الشكل:

```html
<input type="text">
```

أضف إلى عنصر الـ `input` الموجود خاصية `type` بقيمة `"text"`.

وأضف أيضًا إلى عنصر الـ `input` خاصية `id` بقيمة `"full-name"`.

# --hints--

لا يجب أن تزيل عنصر الـ `input` الموجود.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

لا يجب أن تنشئ عنصر `input` جديدًا.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

يجب أن يحتوي عنصر الـ `input` على خاصية `type` بقيمة `"text"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

يجب أن يحتوي عنصر الـ `input` على خاصية `id` بقيمة `"full-name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
