---
id: 66a93c95bc58e26a8fe95818
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

بعد ذلك، أضف `input` من النوع `"email"` أسفل `label` الخاص بالبريد الإلكتروني. يجب أن يحتوي هذا الـ `input` على خاصية `id` بقيمة `"email"`. كما عرّف خاصية `name` بقيمة `"email"`.

هذا الـ `input` مطلوب أيضًا، فتأكد من إضافة خاصية `required`.

أخيرًا، أضف خاصية `placeholder` بقيمة `"example@email.com"`.

# --hints--

يجب أن يكون لديك عنصر `input` أسفل `label` الخاص بالبريد الإلكتروني.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

يجب أن يكون لعنصر `input` الخاص بك `type` بقيمة `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

يجب أن يحتوي عنصر الـ `input` على خاصية `id` بقيمة `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

يجب أن يحتوي عنصر الـ `input` على خاصية `name` بقيمة `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

يجب أن يحتوي عنصر الـ `input` على خاصية `required`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

يجب أن يحتوي عنصر الـ `input` على خاصية `placeholder` بقيمة `"example@email.com"`.
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
