---
id: 66a9419e2d18476c645ce693
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

يُستخدم `input` من النوع number لإنشاء حقل إدخال رقمي.

إليك مثال على حقل إدخال رقم:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

تُستخدم الخاصيتان `min` و `max` لتحديد الحد الأدنى والحد الأقصى للقيم التي يمكن إدخالها في حقل الإدخال.

أسفل عنصر `label` الخاص بك، أضف مُدخَلًا مع خاصية `type` مضبوطة على `"number"` و`id` بقيمة `"age"`.

يجب ضبط خاصية `name` على `"age"`، وخاصية `min` على `"3"`، وخاصية `max` على `"100"`.

# --hints--

يجب أن يكون لديك عنصر `input` رقمي أسفل عنصر `label`.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

يجب أن يحتوي رقم الـ `input` على `id` بقيمة `age`.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

يجب أن يحتوي رقم الـ `input` على خاصية `name` مضبوطة على `age`.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

يجب أن يحتوي عنصر `input` الرقمي على خاصية `min` مضبوطة على `"3"`.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

يجب أن يحتوي عنصر `input` الرقمي على خاصية `max` مضبوطة على `"100"`.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
