---
id: 67a51d3a8a6fe123b77b0c6e
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

يمكن لعناصر `input` أن تحتوي على خاصية `size`. تحدد هذه الخاصية عدد الأحرف التي يجب أن تكون مرئية أثناء كتابة المستخدم في الـ `input`. يجب أن تكون قيمة `size` عددًا صحيحًا موجبًا. إذا لم تُحدد خاصية `size`، أو تم تحديدها بقيمة غير صالحة، فسيأخذ الـ `input` العرض الافتراضي الذي يحدده المتصفح.

```html
<label for="last-name">Last Name:</label>
<input id="last-name" name="last-name" type="text" size="10" />
```

أضف خاصية `size` بقيمة `"20"` إلى مُدخَلي الاسم والبريد الإلكتروني.

# --hints--

يجب أن تضيف خاصية `size` بقيمة `"20"` إلى عنصر `input` الخاص بالاسم.

```js
const nameInput = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(nameInput?.getAttribute('size'), '20');
```

يجب أن تضيف خاصية `size` بقيمة `"20"` إلى عنصر `input` الخاص بالبريد الإلكتروني.

```js
const emailInput = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(emailInput?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
