---
id: 7a97b79c096ecb000d03bd36
title: بناء حاسبة للعدد النوني في متتالية فيبوناتشي
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. أنشئ دالة باسم `fibonacci`.
2. عرّف مصفوفة باسم `sequence` داخل دالة `fibonacci`، وابدأها بالقيم `[0, 1]`.
3. يجب أن تستقبل دالة `fibonacci` معلمة واحدة، وهي عدد صحيح غير سالب `n`.
4. عند استدعاء `fibonacci(n)`، استخدم أسلوب البرمجة الديناميكية لحساب وإرجاع العدد رقم `n` من متتالية فيبوناتشي، حيث كل عدد هو مجموع العددين السابقين له.
5. أضف كل عدد فيبوناتشي يتم حسابه إلى مصفوفة `sequence`.

# --hints--

يجب أن يكون لديك دالة باسم `fibonacci`.

```js
assert.isFunction(fibonacci);
```

يجب أن تأخذ دالة `fibonacci` معلمة واحدة فقط.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

يجب أن يكون لديك مصفوفة باسم `sequence` داخل دالة `fibonacci` مهيأة إلى `[0, 1]`.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

يجب أن تُرجع `fibonacci(0)` القيمة `0`.

```js
assert.equal(fibonacci(0), 0);
```

يجب أن تُرجع `fibonacci(1)` القيمة `1`.

```js
assert.equal(fibonacci(1), 1);
```

يجب أن تُرجع `fibonacci(2)` القيمة `1`.

```js
assert.equal(fibonacci(2), 1);
```

يجب أن تُرجع `fibonacci(3)` القيمة `2`.

```js
assert.equal(fibonacci(3), 2);
```

يجب أن تُرجع `fibonacci(5)` القيمة `5`.

```js
assert.equal(fibonacci(5), 5);
```

يجب أن تُرجع `fibonacci(10)` القيمة `55`.

```js
assert.equal(fibonacci(10), 55);
```

يجب أن تُرجع `fibonacci(15)` القيمة `610`.

```js
assert.equal(fibonacci(15), 610);
```

يجب ألا تستخدم الاستدعاء الذاتي في شفرتك.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
