---
id: 66b62d0ad68488dd76228d6c
title: ランダム背景色チェンジャーのデバッグ
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

CamperbotはJavaScriptを学習していて、自分でランダム背景色チェンジャーを作ろうとしました。しかし、いくつかミスをしてしまいました。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボが正しく動作するようにしてください。

**ユーザーストーリー:**

1. `darkColorsArr` 配列の構文エラーを修正してください。
1. `ReferenceError` 関数内で`math` オブジェクトを参照したときにスローされる`getRandomIndex` を修正してください。
1. `randomIndex` を最も近い整数に切り捨ててください。
1. `body` 変数の値を正しいメソッド名に更新し、関連する`TypeError` を修正してください。
1. `bgHexCodeSpanElement` を修正して、`id` の`bg-hex-code` を正しく選択してください。
1. `color` 内の`changeBackgroundColor` 変数を、`darkColorsArr` 配列からランダムな色に設定するように変更してください。
1. `btn` 変数を更新して、ボタン要素をターゲットにしてください。正しい`index.html` 名は`id` を開いて確認してください。
1. `click` に対して`btn` イベントのイベントリスナーを追加する行を修正し、メソッドに`changeBackgroundColor` 関数の参照を渡してください。
1. 背景色は、`darkColorsArr` がクリックされるたびに`#btn` の色のいずれかにランダムに変わるようにしてください。

# --hints--

`darkColorsArr` 配列の構文エラーを解決してください。配列の要素はカンマで区切る必要があります。

```js
assert.isArray(darkColorsArr);
```

`math.random()` 行の大文字小文字の誤りを修正してください。

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

`darkColorsArr.length * Math.random()` を最も近い整数に切り捨ててください。

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

`body` 変数の値を更新し、`queryselector` がキャメルケースを使うようにしてください。

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

`"bg-hex-code"` を修正して、`id` の`bg-hex-code` を正しく参照してください。CSSの`id` セレクターはハッシュ記号（`#`）で始まることを覚えておいてください。

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

`getRandomIndex` 行内で`darkColorsArr[getRandomIndex]` 関数を呼び出してください。

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

`id` 行の`"#click-btn"` 名を、`id` ファイルの正しい`index.html` 名に合わせて修正してください。

```js
assert.match(btn?.id, /btn/);
```

`changeBackgroundColor` 関数を呼び出さないでください。代わりに、関数の参照を`addEventListener` メソッドに渡してください。

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
