---
id: 60a3e3396c7b40068ad6996a
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Agora, você já deve estar familiarizado com os elementos básicos que uma página em HTML deve ter.

Defina seu código com uma declaração `DOCTYPE` e com os elementos `html` (com o idioma definido para o inglês), `head` e `body`.

# --hints--

Seu código deve ter uma declaração `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Você deve ter uma tag de abertura `<html>` com um atributo `lang` igual a `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Seu código deve ter um elemento `head` dentro do elemento `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

Seu código deve ter um elemento `body` dentro do elemento `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

O elemento `head` deve vir antes do elemento `body`.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
