---
id: 67fccc2463253b213a000142
title: 構建一個多媒體播放器
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

In the prior lessons, you were introduced to working with `audio` and `video` elements. 在本實驗中，你將構建一個多媒體播放器，它會展示一個 `audio` 曲目和一個帶有字幕的 `video`。

對於 `audio` 元素，你需要包含一個 `source` 元素，用於指定所使用的媒體文件。

這是一個示例：

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

`source` 元素同樣可以在 `video` 元素中使用，如下所示：

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. 你應該有一個 `h1` 元素，作爲頁面的主標題。
2. 你應該有三個 `section` 元素。
3. 在第一個 `section` 元素中，應該包含一個 `h2` 元素，作爲正在播放歌曲的標題。
4. 在該 `h2` 元素下方，應該有一個 `audio` 元素，帶有 `controls` 屬性和 `aria-label` 屬性。
5. 在 `audio` 元素內部，應該有一個 `source` 元素，帶有指向音頻文件的 `src` 屬性和 `type` 屬性。 你可以使用此音頻文件：`https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. 在第二個 `section` 元素中，應該包含一個 `h2` 元素，作爲正在播放視頻的標題。
7. 在該 `h2` 元素下方，應該有一個 `video` 元素，帶有 `controls`、`width` 屬性和 `aria-label` 屬性。
8. 在 `video` 元素內部，應該有一個 `source` 元素，帶有指向視頻文件的 `src` 屬性和 `type` 屬性。 You are free to use this video URL if you like: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. 在 `source` 元素下方，應該有一個 `track` 元素，帶有指向字幕文件的 `src` 屬性、`kind` 屬性、`srclang` 屬性和 `label` 屬性。
10. 在第三個 `section` 元素中，應該包含一個 `h2` 元素，作爲該部分的標題，例如： "Transcript"。
11. 在該 `h2` 元素下方，應該有一個 `p` 元素，用來寫視頻的文稿。

# --hints--

你應該有一個 `h1` 元素。

```js
assert.isNotNull(document.querySelector("h1"));
```

你的 `h1` 應包含一些代表頁面主標題的文本。

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

你應該有三個 `section` 元素。

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

第一個 `section` 應該包含一個 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

你的 `h2` 應包含一些代表播放歌曲標題的文本。

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

第一個 `section` 應該包含一個 `audio` 元素。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

你的 `audio` 元素應該有 `controls` 屬性。

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

你的 `audio` 元素應該有一個帶描述性文本的 `aria-label` 屬性。

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

你的 `audio` 元素應該包含一個 `source` 元素。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

你的 `source` 元素應該有一個指向音頻文件的 `src` 屬性。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

你的 `source` 元素應該有一個指定音頻類型的 `type` 屬性。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

第二個 `section` 應該包含一個 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

你的 `h2` 應包含一些代表播放視頻標題的文本。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

第二個 `section` 應該包含一個 `video` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

你的 `video` 元素應該有 `controls` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

你的 `video` 元素應該有一個指定視頻寬度的 `width` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

該 `video` 元素應該包含一個 `source` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

該 `source` 元素應該有一個指向視頻文件的 `src` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

該 `source` 元素應該有一個指定視頻類型的 `type` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

你的 `video` 元素應該包含一個 `track` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

你的 `track` 元素應該有一個指向字幕文件的 `src` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

你的 `track` 元素應該有一個描述文件類型的 `kind` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

你的 `track` 元素應該有一個描述字幕語言的 `srclang` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

你的 `track` 元素應該有一個描述字幕語言標籤的 `label` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

第三個 `section` 應該包含一個 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

你的 `h2` 應包含一些代表該部分標題的文本。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

第三個 `section` 應該包含一個 `p` 元素，用來寫文稿。

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">

      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
