---
id: 67fccc2463253b213a000142
title: 建置多媒體播放器
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

在先前的課程中，你已經介紹了如何使用 `audio` 和 `video` 元素。在這個實驗中，你將建置一個多媒體播放器，該播放器會顯示 `audio` 音軌和帶有文字稿的 `video`。

對於 `audio` 元素，你需要包含一個 `source` 元素，用來指定所使用的媒體。 

這裡有一個範例：

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

`source` 元素也可以像這樣用在 `video` 元素中：

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `h1` 元素作為頁面的主標題。
2. 你應該有三個 `section` 元素。
3. 在第一個 `section` 元素內，你應該有一個 `h2` 元素作為正在播放歌曲的標題。
4. 在 `h2` 元素下方，你應該有一個帶有 `controls` 屬性和 `aria-label` 屬性的 `audio` 元素。
5. 在 `audio` 元素內，你應該有一個帶有指向音訊檔案的 `src` 屬性和 `type` 屬性的 `source` 元素。你可以自由使用此音訊 URL ：`https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. 在第二個 `section` 元素內，你應該有一個 `h2` 元素作為正在播放的影片標題。
7. 在 `h2` 元素下方，你應該有一個帶有 `controls`、`width` 屬性和 `aria-label` 屬性的 `video` 元素。
8. 在 `video` 元素內，你應該有一個帶有指向視訊檔案的 `src` 屬性和一個 `type` 屬性的 `source` 元素。你可以自由使用這個視訊 URL ：`https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. 在 `source` 元素下方，你應該有一個 `track` 元素，其 `src` 屬性指向字幕檔案，並且有一個 `kind` 屬性、一個 `srclang` 屬性和一個 `label` 屬性。如果你願意，可以自由使用這個字幕 URL：`https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. 在第三個 `section` 元素內，你應該有一個 `h2` 元素作為該區段的標題，例如 "Transcript"。
11. 在 `h2` 元素下方，你應該有一個包含影片文字稿的 `p` 元素。

# --hints--

你應該有一個 `h1` 元素。

```js
assert.isNotNull(document.querySelector("h1"));
```

你的 `h1` 應該包含一些表述頁面主要標題的文字。  

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```  

你應該有三個 `section` 元素。

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

第一個 `section` 應該包含一個 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

你的 `h2` 應該包含一些表述正在播放歌曲標題的文字。

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

第一個 `section` 應該包含一個 `audio` 元素。  

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

你的 `audio` 元素應該有一個 `controls` 屬性。  

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

你的 `audio` 元素應該有一個帶有描述音訊的文字的 `aria-label` 屬性。

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

你的 `audio` 元素應該有一個 `source` 元素。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

你的 `source` 元素應該有一個指向音訊檔案的 `src` 屬性。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

你的 `source` 元素應該有一個 `type` 屬性，其值指定音訊的型別。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

第二個 `section` 應該包含一個 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

你的 `h2` 應該包含表述正在播放影片標題的文字。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

第二個 `section` 應該包含一個 `video` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

你的 `video` 元素應該有 `controls` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

你的 `video` 元素應該有一個 `width` 屬性，其值指定影片的寬度。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

`video` 元素應該包含一個 `source` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

`source` 元素應該有一個指向影片檔案的 `src` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

`source` 元素應該有一個 `type` 屬性，其值指定影片的型別。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

你的 `video` 元素應該有一個 `track` 元素。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

你的 `track` 元素應該有一個指向字幕檔案的 `src` 屬性。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

你的 `track` 元素應該有一個 `kind` 屬性，該屬性包含描述檔案類型的文字。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

你的 `track` 元素應該有一個 `srclang` 屬性，該屬性包含描述字幕語言的文字。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

你的 `track` 元素應該有一個帶有描述字幕語言的 `label` 屬性文字。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

第三個 `section` 應該包含一個 `h2` 元素。

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

你的 `h2` 應該有一些表述該區段標題的文字。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

第三個 `section` 應該包含一個用於文字記錄的 `p` 元素。

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
