---
id: 660ee6e3a242da6bd579de69
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Moja ya dhana muhimu sana katika programu ni vigezo. <dfn>Kigezo</dfn> kinaelekeza kwenye anwani maalum ya kumbukumbu inayohifadhi thamani. Vigezo hupatiwa jina ambalo linaweza kutumika katika msimbo wako wote kupata thamani hiyo.

Kutangaza kigezo kunamaanisha kukipa jina. Katika JavaScript, mara nyingi hii hufanywa kwa kutumia neno la `let`. Kwa mfano, hapa ni jinsi unavyotangaza kigezo cha `hello`:

```js
let hello;
```

Jina la kigezo linafuata sheria maalum: majina yanaweza kujumuisha herufi, nambari, alama za dola, na alama za mstari chini, lakini hayawezi kuwa na nafasi na hayapaswi kuanza na nambari.

Tumia neno la `let` kutangaza kigezo kinachoitwa `character`.

_Kumbuka_: Ni desturi ya kawaida kumalizia tamko katika JavaScript kwa alama ya nukta ya nusu. `;`

# --hints--

Unapaswa kutumia `let` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Unapaswa kutumia `character` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

Unapaswa kutumia `let` kutangaza kigezo cha `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Tangazo lako linapaswa kumalizika kwa alama ya nukta ya nusu.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
