---
id: 660eebd83100d37862268781
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Kigezo chako cha `character` kwa sasa hakina thamani. Unaweza kuweka thamani kwa kutumia opereta ya <dfn>mgawo wa thamani</dfn> `=`. Kwa mfano:

```js
let hello = "Hello";
```

Kuweka thamani kwa kigezo wakati wa tamko lake huitwa <dfn>uanzishaji</dfn>.

Anzisha kigezo chako cha `character` kwa kumsadia thamani `"Hello"` wakati wa tamko lake.

# --hints--

Unapaswa kutumia opereta ya mgawo wa thamani `=`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Unapaswa kutumia mfuatano wa herufi `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Unapaswa kutumia alama za nukuu maradufu kwa mfuatano wako wa herufi `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Unapaswa kuweka `"Hello"` kwa kigezo chako cha `character`.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
