---
id: 660ef6355e8f5a9e67fe5f46
title: Hatua 16
challengeType: 1
dashedName: step-16
---

# --description--

Kwa aina ya data ya `number`, unaweza kufanya operesheni za kihesabu, kama vile kuongeza. Jaribu kuchapisha `count + 1` kwenye konsoli.

# --hints--

Unapaswa kupata upatikanaji wa `console` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Unapaswa kupata upatikanaji wa `.log()` wa `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Unapaswa kuandika `count + 1` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
