---
id: 660ef5c1904955978a986a5c
title: Hatua 15
challengeType: 1
dashedName: step-15
---

# --description--

Tumia `let` kutangaza variable ya `count`. Weka thamani ya <dfn>nambari</dfn> `8`. Unapotumia thamani ya nambari, hutumii nukuu. Kwa mfano:

```js
let money = 100;
```

# --hints--

Unapaswa kuwa na neno kuu la `let` la pili katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Unapaswa kutumia `let` kutangaza variable ya `count`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

Unapaswa kuweka thamani ya nambari `8` kwa variable yako ya `count`.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
