---
id: 660ef19b95d3308e7dd31bb6
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>Konsoli</dfn> inakuwezesha kuchapisha na kuona matokeo ya JavaScript. Unaweza kutuma taarifa kwa konsoli kwa kutumia `console.log()`. Kwa mfano, msimbo huu utachapisha `"Naomi"` kwenye konsoli:

```js
let developer = "Naomi";
console.log(developer);
```

Msimbo ulio juu unapata upatikanaji wa variable ya `developer` kwa jina lake katika `console.log()`. Kumbuka kuwa thamani iliyoko kati ya mabano ya kawaida ndiyo itakayochapishwa.

Chapisha thamani ya variable ya `character` kwenye konsoli. Kisha, bofya kitufe cha "Console" ili kuona konsoli ya JavaScript.

# --hints--

Uanzishaji wa variable ya `character` haupaswi kubadilishwa.

```js
assert.strictEqual(character, 'Hello');
```

Unapaswa kupata upatikanaji wa `console` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Unapaswa kupata sifa ya `log` ya `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Unapaswa kutumia mabano ya kawaida kuita njia ya `.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Unapaswa kuchapisha variable ya `character` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
