---
id: 660ef55dd468079679ee0092
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Sasa uko tayari kutangaza kigezo chako kingine. Ondoa kauli zote mbili za `console.log`, na mgawo wa thamani wa `character`.

Pia ondoa kigezo chako cha `secondCharacter`, lakini acha uanzishaji wa `character` usibadilike.

# --hints--

Haupaswi kuwa na kauli yoyote ya `console.log()` katika msimbo wako.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Haupaswi kuweka thamani tena ya `character`.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

Kigezo chako cha `character` kinapaswa kuwa na thamani ya `"Hello"`.

```js
assert.equal(character, "Hello");
```

Haupaswi kuwa na kigezo cha `secondCharacter`.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
