---
id: 660ef31a5be625914a0102cd
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Unapotangaza kigezo kwa neno kuu la `let`, unaweza <dfn>kuweka thamani tena</dfn> (au kubadilisha thamani ya) kigezo hicho baadaye. Katika mfano huu, thamani ya `programmer` inabadilishwa kutoka `"Naomi"` kuwa `"CamperChan"`.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Kumbuka kwamba unapoweka thamani tena kwa kigezo, hutatangi tena neno kuu la `let`.

Baada ya `console.log` yako, weka thamani `"World"` kwa kigezo chako cha `character`.

# --hints--

Unapaswa kutumia `character` mara mbili katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

Unapaswa kutumia opereta ya mgawo kuweka thamani tena kwa `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

Unapaswa kutumia mfuatano wa herufi `"World"` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Kigezo chako cha `character` kinapaswa kuwa na thamani `"World"`.

```js
assert.equal(character, "World");
```

Uwekaji thamani tena haupaswi kutumia `let`.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
