---
id: 660ef857f2806aa626d29d17
title: Hatua 17
challengeType: 1
dashedName: step-17
---

# --description--

Unaweza pia kufanya upunguzaji (`-`), kuzidisha (`*`), na kugawanya (`/`). Tafadhali jaribu kutumia voperesheni na nambari katika `console.log` yako. Unapokuwa tayari kuendelea, ondoa `console.log`.

# --hints--

Haupaswi kuwa na `console.log()` katika msimbo wako.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Kigezo chako cha `character` kinapaswa bado kuwa na thamani ya `"Hello"`.

```js
assert.equal(character, "Hello");
```

Kigezo chako cha `count` kinapaswa bado kuwa na thamani ya `8`.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
