---
id: 660ef0f7c4b8e68ccd1f0786
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript ina aina saba za data za msingi, ambapo Mfuatano wa herufi ni mojawapo yao. Katika JavaScript, <dfn>mfuatano wa herufi</dfn> unawakilisha mfululizo wa herufi na unaweza kufungwa kwa alama za nukuu moja (`'`) au alama za nukuu maradufu (`"`).

Kumbuka kuwa mfuatano wa herufi ni <dfn>usiyobadilika</dfn>, ambayo inamaanisha mara tu unapotengenezwa, hauwezi kubadilishwa. Hata hivyo, kigezo bado kinaweza kupewa thamani nyingine.

Badilisha mfuatano wako wa `"Hello"` kutumia alama za nukuu moja.

# --hints--

Haupaswi kuwa na alama za nukuu maradufu katika msimbo wako.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

Unapaswa kutumia alama za nukuu moja kwa mfuatano wako wa `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

Bado unapaswa kutumia `let` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Bado unapaswa kutangaza kigezo cha `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Kigezo chako cha `character` kinapaswa bado kuwa na mfuatano wa herufi `"Hello"` kama thamani yake.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
