---
id: 660f4a1472f8e63d76162ce5
title: Hatua 99
challengeType: 1
dashedName: step-99
---

# --description--

Je, unadhani vipi kama ungetengeneza piramidi yako ukiipindua, au <dfn>kuzigeuza</dfn>? Ni wakati wa kujaribu!

Anza kwa kuunda mzunguko mpya wa `for`. Tangaza mzunguko wako `i` na weka thamani yake kuwa `count`, kisha tumia thamani ya kweli au siyo kweli `false` kwa masharti na taarifa za mzunguko.

# --hints--

Msimbo wako unapaswa kuwa na mzunguko wa `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

Mzunguko wako wa `for` unapaswa kuanzisha `i` kwa thamani ya `count`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count/);
```

Mzunguko wako wa `for` unapaswa kutumia `false` kama sharti.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false/);
```


Mzunguko wako wa `for` unapaswa kutumia `false` kama mzunguko.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
