---
id: 660f165270622fd4ec0da3f7
title: Hatua 31
challengeType: 1
dashedName: step-31
---

# --description--

Kutangaza kigezo kwa neno muhimu la `let` kunaruhusu kuwekwa thamani tena. Hii inamaanisha unaweza kubadilisha `character` baadaye kuwa thamani tofauti kabisa.

Kwa mradi huu, hutaki kubadilisha thamani hizi za vigezo. Kwa hiyo badala yake, unapaswa kutumia `const` kutangaza vigezo hivyo. Vigezo vya `const` ni maalum.

Kwanza, kigezo cha `const` hakiwezi kuwekwa thamani tena kama kigezo cha `let`. Msimbo huu utatupa kosa:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

Kigezo cha `const` pia hakiwezi kutangazwa bila thamani. Msimbo huu utatupa kosa:

```js
const firstName;
```

Badilisha maneno yako muhimu ya `let` na `const`.

# --hints--

Unapaswa kutumia `const` kutangaza kigezo chako cha `character`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

Unapaswa kutumia `const` kutangaza kigezo chako cha `count`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

Unapaswa kutumia `const` kutangaza kigezo chako cha `rows`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

Hupaswi kutumia `let` katika msimbo wako.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
