---
id: 660f280dda5040f707c76b4a
title: Hatua 45
challengeType: 1
dashedName: step-45
---

# --description--

Sasa una mfuatano wa herufi wa `#` , lakini umbo la piramidi bado halipo. Kwa bahati nzuri, tofauti ya `i` inawakilisha nambari ya "safisha" ya sasa katika mzunguko wako, ikikuwezesha kuitumia kutengeneza muundo wa aina ya piramidi.

Ili kufanikisha hili, utatumia njia ya `.repeat()` inayopatikana kwa mfuatano wa herufi. Njia hii inakubali nambari kama hoja, ikibainisha idadi ya mara za kurudia mfuatano wa herufi unaolengwa. Kwa mfano, kutumia `.repeat()` kutengeneza mfuatano wa herufi `"Code! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

Tumia njia ya `.repeat()` kwenye `character` yako, na uipe `i` kama nambari.

# --hints--

Unapaswa kutumia njia ya `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

Unapaswa kutumia njia ya `.repeat()` kwenye tofauti yako ya `character`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

Unapaswa kupitisha `i` kwa njia yako ya `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

Unapaswa kutumia njia ya `.repeat()` katika njia ya `.push()`

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
