---
id: 660f17d4e9f227d86e834abd
title: Hatua 33
challengeType: 1
dashedName: step-33
---

# --description--

Ili kuunda piramidi, utahitaji kuunda safu nyingi. Unapohitaji kufanya zoezi mara kwa mara hadi hali fulani itakapokamilika, utatumia <dfn>mzunguko</dfn>. Kuna njia nyingi za kuandika mzunguko.

Utaanza na mzunguko wa msingi wa `for`. Mizunguko ya `for` hutumia sintaksia ifuatayo:

```js
for (iterator; condition; iteration) {
  logic;
}
```

Katika hatua zinazofuata, utachunguza kila sehemu ya mzunguko kwa undani. Kwa sasa, tengeneza mzunguko wa `for` unaojumuisha maneno `"iterator"`, `"condition"`, na `"iteration"` kwa sehemu tatu. Weka <dfn>mwili</dfn> wa mzunguko, sehemu ndani ya mabano ya mviringo `{}`, kuwa tupu.


# --hints--

Unapaswa kuwa na mzunguko wa `for`.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

Sehemu ya kwanza ya mzunguko wako wa `for` inapaswa kuwa mfuatano wa herufi `"iterator"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

Sehemu ya pili ya mzunguko wako wa `for` inapaswa kuwa mfuatano wa herufi `"condition"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

Sehemu ya tatu ya mzunguko wako wa `for` inapaswa kuwa mfuatano wa herufi `"iteration"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

Mwili wa mzunguko wako wa `for` unapaswa kuwa tupu.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
