---
id: 660f4f79e2a82a4e92290f44
title: Hatua 110
challengeType: 1
dashedName: step-110
---

# --description--

Wakati mwingine unaweza kutaka kurejesha msimbo uliotangulia uliotolewa maoni. Unaweza kufanya hivyo kwa kuondoa `/*` na `*/` kuzunguka msimbo huo. Hii huitwa <dfn>kuondoa maoni</dfn>.

Ondoa maoni tu kwenye mzunguko wako wa kwanza wa `for`. Acha maoni ya mstari mmoja na maoni mengine mawili ya mistari mingi mahali pake.

# --hints--

Unapaswa kuondoa maoni kwenye mzunguko wako wa kwanza wa `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

Haupasi kuondoa maoni yako ya mstari mmoja.

```js
assert.match(code, /\/\//);
```

Haupasi kuondoa maoni kwenye mzunguko wako wa `while`.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Haupasi kuondoa maoni kwenye mzunguko wako wa pili wa `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
