---
id: 660f0f980e98e8cf77f1ce31
title: Hatua 30
challengeType: 1
dashedName: step-30
---

# --description--

Badilisha tamko lako la `rows` ili liwe limepewa thamani safu ya data tupu tena.

Pia, badilisha mfuatano wako wa herufi wa `'Hello'` kutumia alama za nukuu maradufu tena. Kwa ujumla, haijalishi ni ipi kati ya hizo mbili unayopendelea, lakini utataka kuwa na uthabiti katika uchaguzi huo katika mradi wako mzima.

# --hints--

Safu yako ya data ya `rows` inapaswa kuwa tupu.

```js
assert.empty(rows);
```

Mfuatano wako wa herufi wa `"Hello"` unapaswa kutumia alama za nukuu maradufu.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
