---
id: 660f383d4c772c12ff59904b
title: Hatua 70
challengeType: 1
dashedName: step-70
---

# --description--

Unaweza kupitisha usemi mzima wa kihesabu kama hoja. Kitendakazi kitapokea matokeo ya kutathmini usemi huo. Kwa mfano, wito hizi mbili za kitendakazi hapa chini zingetoa matokeo sawa. (Kumbuka kwamba kanuni ya mpangilio wa operesheni PEMDAS—Mabano, Nambari za nguvu, Kuzidisha, Kugawanya, Kuongeza, Kutoa—inatumika):

```js
test(2 * 3 + 1);
test(7);
```

Tukitazama tena muundo huo:

```md
··#··
·###·
#####
```

Sahihisha thamani ya `character` ili irudiwe mara `2 * rowNumber - 1`.

Fungua tena konsoli kuona matokeo yaliyosasishwa.

# --hints--

Kitendakazi chako kinapaswa kupitisha `2 * rowNumber - 1` kwa wito wako wa `.repeat()`.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
