---
id: 660f46460f9c36330ebc07d8
title: Hatua 88
challengeType: 1
dashedName: step-88
---

# --description--

Opereta ya usawa inaweza kusababisha tabia za kushangaza katika JavaScript. Kwa mfano, `"0" == 0` ni kweli, ingawa moja ni mfuatano wa herufi na nyingine ni nambari.

Opereta ya <dfn>usawa mkali</dfn> `===` hutumika kukagua kama thamani mbili ni sawa na zina aina moja. Kwa kanuni ya jumla, huu ndio opereta ya usawa unayopaswa kutumia kila wakati. Kwa opereta ya usawa mkali, `"0" === 0` inakuwa uongo, kwa sababu ingawa zina thamani sawa ya sifuri, haziko katika aina moja.

Sahihisha sharti lako la `done == count` ili litumie opereta ya usawa mkali.

# --hints--

Sharti lako la `if` linapaswa kutumia usawa mkali.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
