---
id: 660f48e1d3682f39e81843c4
title: Hatua 94
challengeType: 1
dashedName: step-94
---

# --description--

Kizalishaji chako cha piramidi bado kinafanya kazi. Hata hivyo, kuna uwezekano wa kuingia katika mzunguko usioisha tena.

Kwa sababu unachunguza tu kama `done` si sawa na `count`, ikiwa `done` itakuwa **kubwa zaidi** kuliko `count` mzunguko wako utaendelea milele.

Sahihisha sharti la mzunguko wako ili likague kama `done` ni ndogo au sawa na `count`.

# --hints--

Mzunguko wako wa `while` unapaswa kukagua kama `done` ni ndogo au sawa na `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
