---
id: 660f46b9c417a8341729a3ab
title: Hatua 89
challengeType: 1
dashedName: step-89
---

# --description--

Wakati `done` imetimiza thamani ya `count`, tunataka mzunguko usiendelee kutekelezwa.

Ndani ya mwili wa `if`, weka thamani ya kweli au siyo kweli `false` kwa kigezo chako cha `continueLoop`.

# --hints--

Mwili wako wa `if` unapaswa kuweka `false` kwa kigezo chako cha `continueLoop`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{\s*continueLoop\s*=\s*false;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done === count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
