---
id: 660f1a00ac619ddc1e259a66
title: Hatua 35
challengeType: 1
dashedName: step-35
---

# --description--

<dfn>Sharti</dfn> la mzunguko wa `for` linaeleza mzunguko unapaswa kurudia mara ngapi. Mzunguko utaacha wakati `condition` itakapokuwa siyo kweli.

Katika JavaScript, thamani ya kweli au siyo kweli inaweza kuwa `true` au `false`. Hizi si mfuatano wa herufi - utajifunza zaidi kuhusu tofauti hii baadaye.

Kwa sasa, utatumia <dfn>opereta ya chini ya</dfn> (`<`). Hii inakuwezesha kuangalia kama thamani upande wa kushoto ni ndogo kuliko thamani upande wa kulia. Kwa mfano, `count < 3` itahesabiwa kuwa `true` ikiwa `count` ni `2`, na `false` ikiwa `count` ni `4`.

Badilisha mfuatano wako wa `"condition"` na sharti la kuangalia kama `i` ni ndogo kuliko `count`.

# --hints--

Unapaswa kutumia opereta ya chini ya.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

Unapaswa kutumia opereta ya chini ya kuangalia kama `i` ni ndogo kuliko `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Mzunguko wako wa `for` unapaswa kutumia `i < count` kama sharti.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
