---
id: 660f1e3f047bf4e403268713
title: Hatua 39
challengeType: 1
dashedName: step-39
---

# --description--

Kwa bahati mbaya, sasa huwezi kuona kile mzunguko wako unachofanya.

Tumia `let` kutangaza kigezo cha `result`, na weka thamani yake kuwa mfuatano wa herufi tupu. Mfuatano wa herufi tupu unawakilishwa na alama za nukuu zikiwa haina kitu kati yao, kama `""`.

# --hints--

Unapaswa kutangaza kigezo cha `result`.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

Unapaswa kutumia `let` kutangaza kigezo chako cha `result`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

Kigezo chako cha `result` kinapaswa kuwa mfuatano wa herufi tupu.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
