---
id: 660f0c34aad72dc712b97624
title: Hatua 27
challengeType: 1
dashedName: step-27
---

# --description--

Njia nyingine muhimu kwa mradi huu ni njia ya `.pop()`. Inatoa kipengele cha mwisho kutoka kwenye safu ya data na <dfn>inarejesha</dfn> kipengele hicho.

Wakati njia inarejesha thamani, unaweza kuifikiria kama inakurudishia thamani hiyo, ikifanya ipatikane kwa matumizi sehemu nyingine za msimbo wako.

Tengeneza kigezo kipya kinachoitwa `popped` na weka thamani yake matokeo ya `rows.pop()`. Kisha, andika `popped` kwenye konsoli.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `popped`.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

Unapaswa kutumia `let` kutangaza kigezo chako kinachoitwa `popped`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

Unapaswa kuita njia ya `.pop()`.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

Unapaswa kuita njia ya `.pop()` kwenye safu yako ya data ya `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

Unapaswa kuandika kigezo chako cha `popped` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
