---
id: 660f17294346b7d69e79db3d
title: Hatua 32
challengeType: 1
dashedName: step-32
---

# --description--

Sasa uko tayari kuanza kujenga kizalishaji chako cha piramidi. Kigezo chako cha `character` kitakuwa kama jiwe la msingi la piramidi.

`"Hello"` huenda kisifanye kazi vizuri sana kwa ajili hiyo. Badilisha thamani ya `character` iwe alama ya hash `"#"`.

# --hints--

Kigezo chako cha `character` kinapaswa kuwa alama ya hash.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
