---
id: 660f207334fabaeac3269c38
title: Hatua 41
challengeType: 1
dashedName: step-41
---

# --description--

Ili kudhibiti mfuatano wa herufi wa `result`, utatumia aina tofauti ya mzunguko. Hasa, mzunguko wa `for...of`, ambao hurudia kila kipengee katika kitu kinachoweza kurudiwa na kuwekea thamani kwa muda kidogo katika kigezo.

Sintaksia ya mzunguko wa `for...of` inaonekana kama ifuatavyo:

```js
for (const value of iterable) {

}
```

Kumbuka unaweza kutumia `const` kwa sababu kigezo hicho kinaishi kwa mzunguko mmoja tu, si wakati wote wa mzunguko mzima.

Tengeneza mzunguko wa `for...of` kurudia kupitia safu ya data ya `rows`, ukiweka kila thamani katika kigezo cha `row`.

# --hints--

Unapaswa kutumia neno kuu tofauti la `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

Unapaswa kutangaza kigezo cha `row`.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

Unapaswa kutumia `const` kutangaza kigezo chako cha `row`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Mzunguko wako wa `for...of` unapaswa kutangaza kigezo chako cha `row`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

Kigezo chako cha `row` kinapaswa kuchukuliwa kutoka `rows` kwa kutumia neno kuu la `of`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

Mwili wa mzunguko wako wa `for...of` unapaswa kuwa tupu.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
