---
id: 660f3ba3cceef11b6ba08b59
title: Hatua 75
challengeType: 1
dashedName: step-75
---

# --description--

Kwa bahati mbaya, sasa chini ya piramidi imepotea. Hii ni kwa sababu umeunda tena <dfn>kosa la moja</dfn>.

Mzunguko wako wa awali ulifanyika kwa thamani za `i` kutoka `0` hadi `7`, kwa sababu `count` ni `8` na sharti lako linahitaji `i` kuwa chini ya `count`. Mzunguko wako sasa unafanyika kwa thamani za `i` kutoka `1` hadi `7`.

Mzunguko wako unahitaji kusahihishwa ili ufanye kazi wakati `i` ni `8`, pia. Ukiangalia mantiki yako, hii inamaanisha mzunguko wako unapaswa kufanya kazi wakati `i` ni <dfn>chini ya au sawa na</dfn> `count`. Unaweza kutumia opereta ya chini ya au sawa na `<=` kwa hili.

Sahihisha sharti la mzunguko wako ili ufanye kazi wakati `i` ni chini ya au sawa na `count`.

# --hints--

Mzunguko wako wa `for` haupaswi kuangalia kama `i` ni chini ya `count`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Mzunguko wako wa `for` unapaswa kuangalia kama `i` ni chini ya au sawa na `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
