---
id: 660f359af3e32e0f1a6880b7
title: Hatua 68
challengeType: 1
dashedName: step-68
---

# --description--

Sasa unapaswa kuona kundi lile lile la herufi kwenye konsoli yako. Kitendakazi chako cha `padRow` kinafanya jambo lile lile ulilokuwa unalifanya awali, lakini sasa kiko katika sehemu inayoweza kutumika tena yenyewe.

Tumia opereta ya kuongeza kuunganisha nafasi moja `" "` mwanzoni na mwishoni mwa mfuatano wako wa herufi wa `character` unaorudiwa.

Kumbuka kuwa unaweza kutumia opereta ya `+` kuunganisha mfuatano wa herufi kama hivi:

```js
" " + "string"
```

# --hints--

Unapaswa kuunganisha nafasi moja mwanzoni mwa thamani unayorejesha.

```js
assert.match(padRow(1, 1), /^\s/);
```

Unapaswa kuunganisha nafasi moja mwishoni mwa thamani unayorejesha.

```js
assert.match(padRow(1, 1), /\s$/);
```

Kitendakazi chako cha `padRow()` kinapaswa kurudisha mfuatano wa `character` unaorudiwa ukiwa na nafasi kabla na baada ya mfuatano huo.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
