---
id: 660f4e74f7fd3f4a99ac2e50
title: Hatua 108
challengeType: 1
dashedName: step-108
---

# --description--

Safu za data pia zina njia ya `.shift()`. Hii itatoa kipengele cha **kwanza** cha safu ya data, tofauti na `.pop()` ambayo hutoa kipengele cha mwisho. Hapa kuna mfano wa njia ya `.shift()`:

```js
const numbers = [1, 2, 3];
numbers.shift();
```

Safu ya data ya `numbers` itakuwa `[2, 3]`.

Moja kwa moja chini ya safu yako ya data ya `numbers`, tangaza kigezo cha `shifted` na weka thamani yake kuwa matokeo ya wito wa `.shift()` kwenye safu ya data ya `numbers`. Katika mstari unaofuata, andika kigezo cha `shifted` kwenye konsoli.

# --hints--

Unapaswa kutumia `const` kutangaza kigezo cha `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

Unapaswa kuita `.shift()` kwenye safu yako ya data ya `numbers`.

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

Unapaswa kuweka thamani ya matokeo ya wito wako wa `.shift()` kwenye kigezo chako cha `shifted`.

```js
assert.equal(shifted, 1);
```

Unapaswa kuandika kigezo chako cha `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
