---
id: 660f039ff313dbb696b007ca
title: Hatua 19
challengeType: 1
dashedName: step-19
---

# --description--

Wakati safu ya data ina thamani, au <dfn>vipengele</dfn>, thamani hizo hutenganishwa kwa koma. Hapa kuna safu ya data inayoshikilia mfuatano wa herufi mbili:

```js
let array = ["first", "second"];
```

Badilisha tamko lako la `rows` kuwa safu ya data yenye mfuatano wa herufi `"Naomi"`, `"Quincy"`, na `"CamperChan"`. Mpangilio wa thamani katika safu ya data ni muhimu, kwa hivyo fuata mpangilio huo. Kumbuka kwamba mfuatano wa herufi unazingatia herufi kubwa na ndogo.

# --hints--

Kipengele cha kwanza katika safu yako ya data kinapaswa kuwa mfuatano wa herufi `"Naomi"`.

```js
assert.equal(rows[0], "Naomi");
```

Kipengele cha pili katika safu yako ya data kinapaswa kuwa mfuatano wa herufi `"Quincy"`.

```js
assert.equal(rows[1], "Quincy");
```

Kipengele cha tatu katika safu yako ya data kinapaswa kuwa mfuatano wa herufi `"CamperChan"`.

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
