---
id: 660f3ce51f70571e1c5227c8
title: Hatua 76
challengeType: 1
dashedName: step-76
---

# --description--

Maoni yanaweza kusaidia kueleza kwa nini msimbo wako unachukua njia fulani, au kuacha vidokezo vya kufanya baadaye kwa ajili yako mwenyewe.

Katika JavaScript, unaweza kutumia `//` kuacha maoni ya mstari mmoja katika msimbo wako.

Ongeza maoni ya mstari mmoja juu ya mzunguko wako wa `for` ili kujikumbusha kubadilisha msimbo kuwa aina tofauti ya mzunguko.

# --hints--

Unapaswa kuanza maoni ya mstari mmoja kwa `//`.

```js
assert.match(code, /\/\//);
```

Maoni yako ya mstari mmoja yanapaswa kuwa na angalau herufi tano.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
