---
id: 660f540c2176ea5dec01306d
title: Hatua 118
challengeType: 1
dashedName: step-118
---

# --description--

Kazi nzuri! Jaribu thamani tofauti kwa vigezo vyako vya `character`, `count`, na `inverted`.

Unapokuwa tayari kuendelea na mradi wako unaofuata, weka `character` kuwa `"!"`, `count` kuwa `10`, na `inverted` kuwa `false` ili kuendelea.

Hongera kwa kumaliza mradi wako wa kwanza wa JavaScript!

# --hints--

Unapaswa kuweka `character` kuwa `"!"`.

```js
assert.equal(character, "!");
```

Unapaswa kuweka `count` kuwa `10`.

```js
assert.equal(count, 10);
```

Unapaswa kuweka `inverted` kuwa `false`.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
