---
id: 660f033cf051ebb50ea3bf48
title: Hatua 18
challengeType: 1
dashedName: step-18
---

# --description--

Katika programu, mara nyingi utahitaji kufanya kazi na data nyingi. Kuna miundo mingi ya data inayoweza kusaidia kupanga na kusimamia data yako. Moja ya miundo ya data ya msingi ni <dfn>safu ya data</dfn>.

<dfn>Safu ya data</dfn> ni aina ya data isiyo ya msingi inayoweza kuhifadhi mfuatano wa thamani. Aina zisizo za msingi za data zinatofautiana na aina za msingi kwa kuwa zinaweza kuhifadhi data tata zaidi. Aina za msingi kama mfuatano wa herufi na nambari zinaweza kuhifadhi thamani moja kwa wakati mmoja.

Safu za data zinaonyeshwa kwa kutumia mabano ya mraba (`[]`). Hapa kuna mfano wa kigezo chenye thamani ya safu ya data tupu:

```js
let array = [];
```

Tangaza kigezo cha `rows` na weka thamani yake kuwa safu ya data tupu.

# --hints--

Unapaswa kuwa na kigezo cha `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

Unapaswa kutumia `let` kutangaza kigezo chako cha `rows`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

Unapaswa kuweka thamani ya safu ya data tupu kwa kigezo chako cha `rows`.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
