---
id: 660f0ee51d7460ce88cd248d
title: Hatua 29
challengeType: 1
dashedName: step-29
---

# --description--

Je, ulikuwa unatarajia kuona `4` kwenye konsoli? `.push()` hurudisha urefu mpya wa safu ya data, baada ya kuongeza thamani unayoiweka.

Ni muhimu kuelewa ni thamani gani njia hurudisha. Chukua muda kujaribu `.push()` na `.pop()`. Unapokuwa tayari, ondoa wito wote wa `.push()` na `.pop()`, pamoja na taarifa zako za `console.log`.

# --hints--

Haupaswi kuwa na wito wa `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

Haupaswi kuwa na wito wa `.pop()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

Haupaswi kuwa na taarifa zozote za logi.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

Haupaswi kuwa na kigezo cha `popped`.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

Haupaswi kuwa na kigezo cha `pushed`.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
