---
id: 660f45ccf4ca5c31f253005a
title: Hatua 86
challengeType: 1
dashedName: step-86
---

# --description--

Kwa sasa, ikiwa utabadilisha `continueLoop` kuwa true, mzunguko wako wa `while` utaendelea milele. Hii huitwa <dfn>mzunguko usio na mwisho</dfn>, na unapaswa kuwa makini kuepuka hizi. Mzunguko usio na mwisho unaweza kufunga mfumo wako, ukahitaji kuanzisha upya kabisa ili kutoka.

Ili kuepuka hili, anza kwa kutumia opereta ya kuongeza kuongeza thamani ya kigezo cha `done` ndani ya mzunguko wako.

# --hints--

Mzunguko wako wa `while` unapaswa kuongeza kigezo cha `done`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
