---
id: 660f2b6fd54ac1fc142804dd
title: Hatua 47
challengeType: 1
dashedName: step-47
---

# --description--

Mantiki ya kuunda muundo huu wa piramidi huenda ikawa ngumu, ambayo inamaanisha ni wakati mzuri wa kutoa msimbo huo na kuuweka katika kitendakazi.

<dfn>Kitendakazi</dfn> ni kipande cha msimbo kinachoweza kutumika tena katika programu yako yote. Vitendakazi hutangazwa kwa sintaksia ifuatayo:

```js
function name(parameter) {

}
```

Neno kuu la `function` linaambia JavaScript kwamba kigezo cha `name` kitakuwa kitendakazi. `parameter` ni kigezo kinachowakilisha thamani inayopitishwa ndani ya kitendakazi wakati kinapotumika. Kitendakazi kinaweza kuwa na vigezo vingi au vichache kama unavyotaka. Kama katika mzunguko wa `for`, nafasi kati ya mabano ya mviringo ni <dfn>mwili wa kitendakazi</dfn>.

Tangaza kitendakazi cha `padRow`. Usitengeneze vigezo vyovyote bado. Mwili wa kitendakazi unapaswa kuwa tupu. Kumbuka kwamba unahitaji kutumia mkataba wa camel case kwa majina yako.

# --hints--

Unapaswa kutumia neno kuu la `function`.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

Unapaswa kutangaza kitendakazi cha `padRow`.

```js
assert.isFunction(padRow);
```

Kitendakazi chako cha `padRow()` hakipaswi kuwa na vigezo vyovyote.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

Kitendakazi chako cha `padRow()` kinapaswa kuwa na mwili tupu.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
