---
id: 660f53ad3d39175c5d4335ac
title: Hatua 117
challengeType: 1
dashedName: step-117
---

# --description--

Kizalishaji chako cha piramidi sasa kiko katika hali ya kukamilika, na kina utendaji zaidi kuliko ulivyopanga awali! Hatua inayofuata ni kusafisha msimbo wako.

Ondoa maoni yote, ya mstari mmoja na ya mistari mingi, kutoka kwenye msimbo wako.

# --hints--

Haupaswi kuwa na maoni yoyote ya mstari mmoja katika msimbo wako.

```js
assert.notMatch(code, /\/\//);
```

Haupaswi kuwa na maoni yoyote ya mistari mingi katika msimbo wako.

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

Haupaswi kuwa na maoni yoyote katika msimbo wako.

```js
assert.equal(code, __helpers.removeJSComments(code));
```

Unapaswa kuondoa msimbo uliokuwa umewekwa maoni kwa kutumia maoni ya mistari mingi.

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
