---
id: 660f0da9bf1035c9097af20a
title: Hatua 28
challengeType: 1
dashedName: step-28
---

# --description--

Unapaswa kuwa umeona `"freeCodeCamp"` ikichapishwa kwenye konsoli. Hii ni kwa sababu `.pop()` inarudisha thamani iliyotolewa kutoka kwenye safu ya data - na ulisukuma `"freeCodeCamp"` mwishoni mwa safu ya data hapo awali.

Lakini `.push()` inarudisha nini? Weka thamani ya `rows.push()` uliyonayo kwenye kigezo kipya cha `pushed`, na kiandike kwenye kumbukumbu.

# --hints--

Unapaswa kutangaza kigezo cha `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

Unapaswa kutumia `let` kutangaza kigezo chako cha `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

Unapaswa kuweka thamani ya `rows.push("freeCodeCamp")` kwenye kigezo chako cha `pushed`.

```js
assert.equal(pushed, 4);
```

Unapaswa kuandika kwenye kumbukumbu kigezo chako cha `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
