---
id: 660f2a70ad6225fa503e71c3
title: Hatua 46
challengeType: 1
dashedName: step-46
---

# --description--

Unakaribia! Hivi sasa, unakutana na kile kinachojulikana kama <dfn>kosa la moja</dfn>, tatizo la mara kwa mara katika lugha za kuanzia sufuri kama JavaScript.

Kielezo cha kwanza cha safu yako ya `rows` ni `0`, ndiyo maana unaanza mzunguko wako wa `for` na `i = 0`. Lakini kurudia mfuatano wa herufi mara sifuri hutoa hakuna kitu cha kuchapisha.

Ili kurekebisha hili, ongeza `1` kwenye thamani ya `i` katika wito wako wa `.repeat()`. Usirudishe tena kwenye `i` kama ulivyofanya katika masharti ya mzunguko wako.

# --hints--

Unapaswa kuongeza `1` kwenye `i` katika njia yako ya `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
