---
id: 660f4455f457ef2e3ec6920f
title: Hatua 80
challengeType: 1
dashedName: step-80
---

# --description--

Sasa mfuatano wa herufi hauchapishi tena, kwa sababu `false` si `true`. Lakini je, kuhusu thamani nyingine?

Jaribu kubadilisha sharti kuwa mfuatano wa herufi `"false"`.

# --hints--

Sentensi yako ya `if` inapaswa kuwa na mfuatano wa herufi `"false"` kama sharti.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*('|")false\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (false) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
