---
id: 660f1b6e60bd9edf902c81fd
title: Hatua 36
challengeType: 1
dashedName: step-36
---

# --description--

Kauli ya <dfn>iteration</dfn> itakuambia mzunguko wako ufanye nini na mzungumzaji baada ya kila mzunguko.

Unapoweka thamani upya kwa kigezo, unaweza kutumia kigezo hicho kurejea thamani ya awali kabla ya kuweka thamani upya. Hii inakuwezesha kufanya mambo kama kuongeza tatu kwa nambari iliyopo. Kwa mfano, `bees = bees + 3` itakuongezea thamani ya `bees` kwa tatu.

Tumia sintaksia hiyo kubadilisha mfuatano wako wa herufi wa `"iteration"` kwa kauli ya kuweka thamani upya inayoongeza `i` kwa moja.

# --hints--

Unapaswa kuongeza moja kwa kigezo chako cha `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

Unapaswa kuweka thamani ya `i + 1` tena kwa kigezo chako cha `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Mzunguko wako wa `for` unapaswa kuongeza `i` kwa `1` kila mara ya iteration.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
