---
id: 660f09a2694b59c3a10ee304
title: Hatua 23
challengeType: 1
dashedName: step-23
---

# --description--

Kwa sasa, ondoa logi yako ya kwanza ya konsoli na mgawo wako wa `rows[rows.length - 1]`. Acha tamko la pili la logi la `rows` kwa baadaye.

# --hints--

Unapaswa kuondoa tamko lako la `console.log(rows[0])`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

Unapaswa kuondoa mgawo wako wa `rows[rows.length - 1]`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

Hupaswi kuondoa tamko lako la `console.log(rows)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
