---
id: 660f4efcb8068e4cb470dca1
title: Hatua 109
challengeType: 1
dashedName: step-109
---

# --description--

Sasa baada ya kujaribu njia hizi, unaweza kufanya tena mbinu ya piramidi iliyopinduliwa. Lakini kwanza unahitaji kusafisha majaribio yako.

Ondoa safu ya data ya `numbers`, pamoja na wito wa njia na wito wa logi.

# --hints--

Haupaswi kuwa na safu ya data ya `numbers`.

```js
assert.notMatch(__helpers.removeJSComments(code), /numbers/);
```

Haupaswi kuwa na kigezo cha `unshifted`.

```js
assert.notMatch(__helpers.removeJSComments(code), /unshifted/);
```

Haupaswi kuwa na kigezo cha `shifted`.

```js
assert.notMatch(__helpers.removeJSComments(code), /shifted/);
```

Haupaswi kuwa na taarifa zako za `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];
const shifted = numbers.shift();
console.log(shifted);
const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
