---
id: 660f447efc0e722f016c1be0
title: Hatua 81
challengeType: 1
dashedName: step-81
---

# --description--

Maandishi yameonekana tena! Hii ni kwa sababu `"false"` ni mfuatano wa herufi, ambao ukitathminiwa kuwa thamani ya kweli au siyo kweli unakuwa `true`. Hii inamaanisha `"false"` ni thamani ya kweli.

<dfn>Thamani ya kweli</dfn> ni thamani inayochukuliwa kuwa kweli wakati inapotathminiwa kama thamani ya kweli au siyo kweli. Thamani nyingi unazokutana nazo katika JavaScript zitakuwa thamani za kweli.

<dfn>Thamani ya uongo</dfn> ni kinyume chake - thamani inayochukuliwa kuwa siyo kweli wakati inapotathminiwa kama thamani ya kweli au siyo kweli. JavaScript ina orodha iliyobainishwa ya thamani za uongo. Baadhi yao ni pamoja na `false`, `0`, `""`, `null`, `undefined`, na `NaN`.

Jaribu kubadilisha sharti lako la `if` kuwa mfuatano wa herufi tupu `""`, ambao ni thamani ya uongo.

# --hints--

Sentensi yako ya `if` inapaswa kuwa na `""` kama sharti.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
