---
id: 660f4cffb1459d45e34902d1
title: Hatua 106
challengeType: 1
dashedName: step-106
---

# --description--

Unaweza kujenga piramidi iliyogeuzwa bila ya kuhitaji kuzunguka "nyuma" kama ulivyofanya.

Ili kufanya hivi, utahitaji kujifunza njia mpya kadhaa za safu ya data. Anza kwa kutumia `const` kutangaza variable ya `numbers`. Weka thamani yake kuwa safu ya data yenye vipengele `1`, `2`, na `3`. Kisha andika safu ya data ya `numbers` kwenye kumbukumbu.

# --hints--

Unapaswa kutumia `const` kutangaza variable ya `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

Variable yako ya `numbers` inapaswa kuwa safu ya data.

```js
assert.isArray(numbers);
```

Safu yako ya data ya `numbers` inapaswa kuwa na vipengele `1`, `2`, na `3` kwa mpangilio huo huo.

```js
assert.deepEqual(numbers, [1,2,3]);
```

Unapaswa kuandika safu yako ya data ya `numbers` kwenye kumbukumbu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
