---
id: 660f47afe4c98536715d5fa4
title: Hatua 91
challengeType: 1
dashedName: step-91
---

# --description--

Opereta ya <dfn>kutokufanana kwa ukali</dfn> `!==` inakuwezesha kuangalia kama thamani mbili hazilingani, au hazina aina sawa. Sintaksia ni sawa na opereta ya usawa: `value !== 4`.

Kwa sasa mzunguko wa `while` unaendeshwa tu ikiwa `continueLoop` ni kweli. Sahihisha hali ya mzunguko wa `while` ili kuangalia kama `done` haifani na `count`.

# --hints--

Mzunguko wako wa `while` unapaswa kuangalia kama `done` na `count` hazilingani.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
