---
id: 660f3b664421471aa595170f
title: Hatua 74
challengeType: 1
dashedName: step-74
---

# --description--

Piramidi inaonekana kidogo ya kuchekesha sasa. Kwa sababu unaanza mzunguko kwa `1` badala ya `0`, huna haja ya kuongeza moja kwa `i` unapoipitisha kwa `padRow`.

Sahihisha hoja ya kwanza ya wito wako wa `padRow` iwe `i`.

# --hints--

Haupaswi kupitisha `i + 1` kwa `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

Unapaswa kupitisha `i` kwa `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
