---
id: 660f20473aef47e9b8c9afc6
title: Hatua 40
challengeType: 1
dashedName: step-40
---

# --description--

Ongeza usemi wa kumbukumbu ili kuchapisha thamani ya `result`. Kulingana na konsoli unayotumia, huenda usione chochote kikichapishwa.

# --hints--

Unapaswa kuongeza wito wa `console.log`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Unapaswa kuchapisha thamani ya kigezo chako cha `result`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
