---
id: 660f2fbd45b520046cac68e8
title: Hatua 65
challengeType: 1
dashedName: step-65
---

# --description--

Kumbuka katika hatua ya awali, ulijifunza kuhusu thamani za return. Kitendakazi kinaweza <dfn>kurudisha</dfn> thamani kwa ajili ya programu yako kuitumia tofauti.

Katika kitendakazi, neno kuu la `return` linatumika kubainisha thamani ya return. Kwa mfano, kitendakazi hiki kingerudisha thamani iliyotolewa kwa kigezo cha kwanza:

```js
function name(parameter) {
  return parameter;
}
```

Tumia neno kuu la `return` kurudisha thamani ya variable ya `character`, ikirudiwa mara `rowNumber`.

# --hints--

Unapaswa kutumia njia ya `.repeat()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

Unapaswa kutumia njia ya `.repeat()` kwenye variable yako ya `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

Unapaswa kupitisha `rowNumber` kwa wito wako wa `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

Unapaswa kutumia neno kuu la `return`.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

Unapaswa kurudisha matokeo ya wito wako wa `.repeat()`.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
