---
id: 660f44f10ea40f300b896a5e
title: Hatua 84
challengeType: 1
dashedName: step-84
---

# --description--

Sasa umezoea kufanya kazi na tamko za `if...else if...else`, unaweza kuondoa hizo kutoka kwenye msimbo wako.

Mara tu unapomaliza hilo, tumia `let` kutangaza kigezo cha `continueLoop` na kisha weka thamani ya kweli au siyo kweli ya `false`. Kisha tumia `let` kutangaza kigezo cha `done` na kisha weka thamani ya `0`.

# --hints--

Haupaswi kuwa na tamko la `if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

Haupaswi kuwa na tamko la `else if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Haupaswi kuwa na tamko la `else`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

Unapaswa kutumia `let` kutangaza kigezo cha `continueLoop`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

Kigezo chako cha `continueLoop` kinapaswa kuwa na thamani ya `false`.

```js
assert.isFalse(continueLoop);
```

Unapaswa kutumia `let` kutangaza kigezo cha `done`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

Kigezo chako cha `done` kinapaswa kuwa na thamani ya `0`.

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
