---
id: 660f255022991ef34ed0ee88
title: Hatua 44
challengeType: 1
dashedName: step-44
---

# --description--

Kuchapisha nambari hakutatoa piramidi yenye muonekano mzuri. Sasa unapotoka maudhui yaliyopangwa ya safu yako ya data ya `rows`, ni wakati wa kusahihisha mzunguko wako wa awali.

Badala ya kusukuma `i` kwenye safu ya data, sukuma thamani ya kigezo chako cha `character`.

# --hints--

Haupaswi tena kusukuma kigezo chako cha `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Unapaswa kusukuma kigezo chako cha `character`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
