---
id: 660f415b76859a2736771607
title: Hatua 78
challengeType: 1
dashedName: step-78
---

# --description--

Piramidi yako imepotea tena. Hilo ni sawa - hilo linatarajiwa.

Kabla hujaunda mzunguko wako mpya, unahitaji kujifunza kuhusu tamko la `if`. Tamko la <dfn>`if`</dfn> linakuwezesha kuendesha kipande cha msimbo tu wakati hali fulani inatimizwa. Zinatumia sintaksia ifuatayo:

```js
if (condition) {
  logic
}
```

Unda tamko la `if` lenye thamani ya kweli au siyo kweli `true` kama hali. Katika mwili, chapisha mfuatano wa herufi `"Condition is true"`.

# --hints--

Unapaswa kuunda tamko la `if`.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

Tamko lako la `if` linapaswa kuwa na `true` kama hali.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

Mwili wa tamko lako la `if` unapaswa kuchapisha `"Condition is true"`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
