---
id: 660f374d532dc41189cc9cc2
title: Hatua 69
challengeType: 1
dashedName: step-69
---

# --description--

Sasa ni wakati wa kidogo cha hesabu. Fikiria piramidi yenye safu tatu. Ikiwa tunataka iwe katikati, itakuwa kama ifuatavyo:

```md
··#··
·###·
#####
```

Nafasi tupu zimebadilishwa na nukta za katikati, au nukta za wastani, kwa ajili ya urahisi wa kusoma. Ukizidisha muundo huo, utaona kuwa nafasi mwanzoni na mwishoni mwa safu zinafuata muundo fulani.

Sahihisha mfuatano wako wa nafasi tupu ili urudishwe `rowCount - rowNumber` mara.

Fungua konsoli kuona matokeo.

# --hints--

Unapaswa kuita `.repeat()` kwenye mfuatano wako wa `" "` ili urudishe mara `rowCount - rowNumber`.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
