---
id: 660f18f059fe0fda192ce394
title: Hatua 34
challengeType: 1
dashedName: step-34
---

# --description--

Mzunguko wako sasa unahitaji mzunguzaji (iterator) sahihi. <dfn>Iterator</dfn> ni kigezo ambacho unaweza kutangaza mahsusi katika mzunguko wako wa `for` kudhibiti jinsi mzunguko unavyorudia au kupita kupitia mantiki yako.

Ni mkataba wa kawaida kutumia `i` kama kigezo chako cha mzunguzaji katika mzunguko. Mzunguko wa `for` unakuwezesha kutangaza hili katika mabano ya kawaida ya `()`. Kwa mfano, hapa kuna mzunguko wa `for` unaotangaza kigezo cha `index` na kuweka thamani yake kuwa `100`.

```js
for (let index = 100; "second"; "third") {

}
```

Badilisha mfuatano wa herufi `"iterator"` kwa tamko la `let` kwa kigezo `i`. Weka thamani `0` kuanzia. Hii itampa kigezo cha `i` thamani ya `0` kwa **mara ya kwanza** mzunguko wako unapokimbia.

# --hints--

Unapaswa kutumia `let` kutangaza kigezo cha `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

Unapaswa kuweka thamani `0` kwa kigezo chako cha `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

Mzunguko wako wa `for` unapaswa kuanzisha mzunguzaji wa `i` kwa `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
