---
id: 660f07d231941bc11719f664
title: Hatua 21
challengeType: 1
dashedName: step-21
---

# --description--

Safu za data ni maalum kwa sababu zinachukuliwa kuwa <dfn>zinaweza kubadilika</dfn>. Hii inamaanisha unaweza kubadilisha thamani katika kielezo moja moja moja.

Kwa mfano, msimbo huu utaweka nambari `25` kwenye kipengele cha pili katika safu ya data:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

Sahihisha kipengele cha **tatu** cha safu yako ya `rows` kuwa nambari `10`. Kisha chapisha safu ya `rows` kwenye konsoli yako.

# --hints--

Unapaswa kutumia uandishi wa mabano kwenye safu ya `rows` tena.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

Unapaswa kupata upatikanaji wa kipengele cha tatu cha safu ya `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

Unapaswa kutumia opereta ya mgawo kwenye kipengele cha tatu cha safu ya `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

Unapaswa kuweka thamani `10` kwenye kipengele cha tatu cha safu yako ya `rows`.

```js
assert.equal(rows[2], 10);
```

Unapaswa kuwa na tamko la pili la `console.log` katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Unapaswa kuandika safu ya `rows` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
