---
id: 660f1bf673487ae0bb25b900
title: Hatua 37
challengeType: 1
dashedName: step-37
---

# --description--

Mzunguko wako sasa unapaswa kuendeshwa mara nane. Ndani ya mwili wa mzunguko, chapisha thamani ya mzunguko wa `i` na uone kinachotokea.

# --hints--

Unapaswa kutumia `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

Unapaswa kuandika thamani ya `i`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

Unapaswa kuandika thamani ya `i` katika mzunguko wako wa `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
