---
id: 660f2eccfe3f820304af1b39
title: Hatua 64
challengeType: 1
dashedName: step-64
---

# --description--

Ili kujua jinsi ya kuunda muundo wa safu, kitendakazi chako cha `padRow` kitahitaji kujua ni nambari gani ya safu ulipo, na ni safu ngapi kwa jumla zinazozalishwa.

Njia bora ya kufanya hivi ni kwa kuunda vigezo vya kitendakazi kwa ajili yao. Toa kitendakazi chako cha `padRow` kigezo cha `rowNumber` na `rowCount`. Vigezo vingi hutenganishwa kwa koma:

```js
function name(first, second) {

}
```

# --hints--

Kitendakazi chako cha `padRow` kinapaswa kuwa na kigezo cha `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

Unapaswa kuongeza koma baada ya kigezo chako cha `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

Kitendakazi chako cha `padRow` kinapaswa kuwa na kigezo cha `rowCount`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
