---
id: 660f34e99571070d56d2f231
title: Hatua 67
challengeType: 1
dashedName: step-67
---

# --description--

Kitendakazi chako cha `padRow` kina vigezo viwili ulivyovieleza. Thamani hutolewa kwa vigezo hivyo wakati kitendakazi kinapoitwa.

Thamani unazotoa kwa wito wa kitendakazi huitwa <dfn>hoja</dfn>, na unazipitia hoja kwa wito wa kitendakazi. Hapa kuna wito wa kitendakazi ukiwa na `"Hello"` imetumiwa kama hoja:

```js
test("Hello");
```

Pitia `i + 1` na `count` kama hoja kwa wito lako la `padRow`. Kama vigezo, hoja hutenganishwa kwa koma.

# --hints--

Unapaswa kupitisha `i + 1` kwa wito lako la `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

Unapaswa kuwa na koma baada ya hoja yako ya `i + 1`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

Unapaswa kupitisha `count` kama hoja yako ya pili.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
