---
id: 660f4de78f775e480ba2e451
title: Hatua 107
challengeType: 1
dashedName: step-107
---

# --description--

Njia ya `.unshift()` ya safu ya data inakuwezesha kuongeza thamani mwanzoni mwa safu ya data, tofauti na `.push()` ambayo huongeza thamani mwishoni mwa safu ya data. `.unshift()` hurudisha urefu mpya wa safu ya data ambayo ilitumiwa.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

Tumia `const` kutangaza variable ya `unshifted`, na weka thamani yake kuwa matokeo ya wito wa `.unshift()` kwenye safu yako ya data ya `numbers`. Pitia `5` kama hoja. Kisha chapisha variable yako ya `unshifted`.

# --hints--

Unapaswa kutangaza variable yako ya `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

Unapaswa kuita njia ya `.unshift()` kwenye safu yako ya data ya `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

Unapaswa kupitisha `5` kama hoja kwa wito wako wa `.unshift()`.

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

Unapaswa kuweka matokeo ya wito wako wa `.unshift()` kwenye variable yako ya `unshifted`.

```js
assert.equal(unshifted, 4);
```

Unapaswa kuchapisha variable yako ya `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
