---
id: 660f39b444fd6f16d1e49c1f
title: Hatua 73
challengeType: 1
dashedName: step-73
---

# --description--

Badala ya kuhitaji kupitisha `i + 1` kwa wito wako wa `padRow`, unaweza badala yake kuanza mzunguko wako kwenye `1`. Hii itakuwezesha kuunda mzunguko unaoanza na kielezo cha moja.

Sahihisha mzunguko wako ili uanze kwa `1` badala ya `0`.

# --hints--

Mzunguko wako wa `for` unapaswa kuanzisha `i` kwa `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
