---
id: 660f061d259bbebc37461080
title: Hatua 20
challengeType: 1
dashedName: step-20
---

# --description--

Unaweza kupata thamani ndani ya safu ya data kwa kutumia <dfn>kielezo</dfn> cha thamani hiyo. Kielezo ni nambari inayowakilisha nafasi ya thamani katika safu ya data, kuanzia `0` kwa thamani ya kwanza.

Unaweza kupata thamani kwa kutumia <dfn>uandishi wa mabano</dfn>, kama `array[0]`.

Tumia `console.log` na uandishi wa mabano kuchapisha thamani ya kwanza katika safu yako ya `rows`.

# --hints--

Unapaswa kuwa na tamko la `console.log()` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Unapaswa kupata safu yako ya `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

Unapaswa kutumia uandishi wa mabano na safu yako ya `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

Unapaswa kutumia uandishi wa mabano kupata kipengele cha kwanza cha safu yako ya `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

Unapaswa kuandika kipengele cha kwanza cha safu yako ya `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
