---
id: 660f3915b41a441537ec9f5e
title: Hatua 72
challengeType: 1
dashedName: step-72
---

# --description--

Kwa sababu unakuongeza `i` kwa `1` tu, unaweza kutumia <dfn>opereta ya kuongeza</dfn> `++`. Opereta hii huongeza thamani ya kigezo kwa 1, ikisasisha mgawo wa thamani kwa kigezo hicho. Kwa mfano, `test` itageuka kuwa `8` hapa:

```js
let test = 7;
test++;
```

Badilisha mgawo wako wa kuongeza na opereta ya kuongeza kwa mzunguko wako wa kurudia.

# --hints--

Mzunguko wako wa `for` haupaswi kutumia mgawo wa kuongeza na `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

Mzunguko wako wa `for` unapaswa kutumia opereta ya kuongeza kwenye `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
