---
id: 660f229d2dbe09ef2954a4a1
title: Hatua 42
challengeType: 1
dashedName: step-42
---

# --description--

Kumbuka katika mzunguko wako wa awali ulitumia opereta ya kuongeza `+` kuongeza thamani ya `i` kwa `1`.

Unaweza kufanya jambo kama hilo kwa thamani ya mfuatano wa herufi, kwa kuunganisha mfuatano mpya wa herufi kwenye mfuatano uliopo. Kwa mfano, `hello = hello + " World";` itauongeza mfuatano wa herufi `" World"` kwenye mfuatano ulihifadhiwa katika variable ya `hello`. Hii huitwa <dfn>kuunganisha</dfn>.

Katika mzunguko wako wa `for...of`, tumia opereta ya kuongeza kuunganisha thamani ya `row` kwenye thamani ya `result`.

# --hints--

Unapaswa kutumia opereta ya kuunganisha kwenye variable yako ya `result`.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

Unapaswa kuunganisha `row` kwenye variable yako ya `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

Unapaswa kuweka thamani ya matokeo ya kuunganisha tena kwenye variable ya `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
