---
id: 660f3dd626be3a1ffe27e5d1
title: Hatua 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript pia ina msaada kwa maoni ya mistari mingi. Maoni ya mistari mingi huanza na `/*` na kumalizika na `*/`.

Tofauti na maoni ya mstari mmoja, maoni ya mistari mingi yatafunika mistari mingi.

Tumia `/*` na `*/` kubadilisha mzunguko wako wa sasa wa `for`, ikiwa ni pamoja na mwili wake, kuwa maoni ya mistari mingi.

# --hints--

Unapaswa kuanza maoni ya mistari mingi na `/*`.

```js
assert.match(code, /\/\*/);
```

Unapaswa kumaliza maoni ya mistari mingi na `*/`.

```js
assert.match(code, /\*\//);
```

Mzunguko wako mzima wa `for` unapaswa kuwekwa kama maoni.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
