---
id: 660f0a55847d6cc485f29ba5
title: Hatua 26
challengeType: 1
dashedName: step-26
---

# --description--

Njia (method) katika JavaScript ni kitendakazi kinachohusishwa na baadhi ya thamani au vitu fulani. Mfano ambao tayari umeukutana naye ni njia ya `.log()`, ambayo ni sehemu ya kitu cha `console`.

Safu za data zina njia zao, na njia ya kwanza utakayochunguza ni `.push()`. Hii inakuwezesha "kusukuma" thamani mwishoni mwa safu ya data. Hapa kuna mfano wa kuongeza nambari `12` mwishoni mwa safu ya data:

```js
array.push(12);
```

Tumia `.push()` kuongeza mfuatano wa herufi `"freeCodeCamp"` mwishoni mwa safu yako ya `rows`. Ongeza msimbo huu kabla ya `console.log` ili uone mabadiliko uliyoyafanya kwenye safu yako ya data.

# --hints--

Unapaswa kutumia `.push()` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

Unapaswa kutumia njia ya `.push()` ya safu yako ya `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Unapaswa kupitisha mfuatano wa herufi `"freeCodeCamp"` kwa njia yako ya `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
