---
id: 660f38c34a4de6141c0c369f
title: Hatua 71
challengeType: 1
dashedName: step-71
---

# --description--

Kizalishaji chako cha piramidi sasa kinafanya kazi kama inavyotarajiwa. Lakini hii ni fursa nzuri ya kuchunguza zaidi msimbo uliouandika.

Kiongezaji siyo njia pekee ya kuongeza thamani kwenye kigezo. Kiongezaji cha mgawo wa thamani kinaweza kutumika kama njia fupi ya kusema "chukua thamani ya awali ya kigezo, ongeza thamani hii, na weka matokeo tena kwenye kigezo." Kwa mfano, sentensi hizi mbili zingetoa matokeo sawa:

```js
test = test + 1;
test += 1;
```

Sahihisha sentensi yako ya kurudia katika mzunguko wa `for` ili itumie mgawo wa thamani wa kuongeza.

# --hints--

Mzunguko wako wa `for` haupaswi kutumia `i = i + 1`;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Mzunguko wako wa `for` unapaswa kutumia mgawo wa thamani wa kuongeza na `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
