---
id: 660f1cedf3676fe26122ebf6
title: Hatua 38
challengeType: 1
dashedName: step-38
---

# --description--

Unapaswa kuona nambari sifuri hadi saba zikichapishwa kwenye konsoli yako, moja kwa mstari. Hii itatumika kama msingi wa kuunda piramidi yako.

Badilisha kauli yako ya log na kauli ya kusukuma `i` kwenye safu ya data ya `rows`.

# --hints--

Haupaswi kuwa na wito la `console.log`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Unapaswa kuita `.push()` kwenye safu ya data ya `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Unapaswa kusukuma `i` kwenye safu ya data ya `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

`.push()` yako inapaswa kutokea katika mzunguko wako wa `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
