---
id: 660f455b044d3230ed971e98
title: Hatua 85
challengeType: 1
dashedName: step-85
---

# --description--

Mzunguko wa <dfn>`while`</dfn> utaendelea kurudiwa tena na tena hadi `condition` iliyotangazwa isiwe kweli tena. Una sintaksia ifuatayo:

```js
while (condition) {
  logic;
}
```

Tumia sintaksia hiyo kutangaza mzunguko wa `while` ukiwa na `continueLoop` kama sharti. Mwili wake unapaswa kuwa tupu.

# --hints--

Unapaswa kutumia mzunguko wa `while`.

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

Mzunguko wako wa `while` unapaswa kutumia `continueLoop` kama sharti.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
