---
id: 6610b9f7619764fad5fd516d
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Unaweza pia kuweka thamani ya kigezo katika kigezo kingine. Kwa mfano:

```js
let first = "One";
let second = "Two";
second = first;
```

Kigezo cha `second` sasa kingekuwa na thamani ya `"One"`.

Ili kuona hili likifanyika, badilisha mgawo wa `secondCharacter` kutoka `"Test"` kwenda kigezo chako cha `character`.

Kisha fungua konsoli kuona kinachorekodiwa.

# --hints--

Haupaswi kuweka thamani ya `"Test"` katika kigezo chako cha `secondCharacter`.

```js
assert.notEqual(secondCharacter, "Test");
```

Unapaswa kuweka thamani ya kigezo cha `character` katika kigezo chako cha `secondCharacter`. Usisahau alama ya nukta.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

Kigezo chako cha `secondCharacter` sasa kinapaswa kuwa na thamani ya `"World"`.

```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
