---
id: 6610c2d8d67563174fcf96dc
title: Hatua 52
challengeType: 1
dashedName: step-52
---

# --description--

Unapokuwa na thamani ambayo imeandikwa waziwazi katika msimbo wako, kama mfuatano wa herufi `"Hello!"` katika kitendakazi chako, inachukuliwa kuwa <dfn>imeandikwa ngumu</dfn>. Kuandika thamani ngumu ndani ya kitendakazi kunaweza kuifanya isitumike tena kama unavyotaka.

Badala yake, unaweza kufafanua <dfn>vigezo</dfn> kwa kitendakazi. Vigezo ni vigezo maalum vinavyopokea thamani wakati unapofanya wito wa kitendakazi, na vinaweza kutumika ndani ya kitendakazi chako kubadilisha matokeo ya msimbo wa kitendakazi kwa njia ya mabadiliko.

Ili kuongeza kigezo kwenye kitendakazi chako, unahitaji kuongeza jina la kigezo ndani ya mabano ya kawaida. Kwa mfano, kitendakazi hiki `demo` kina kigezo `name`:

```js
function demo(name) {

}
```

`name` inaonekana ni kigezo muhimu, kwa hivyo endelea na uiongeze kwenye kitendakazi chako cha `padRow`.

# --hints--

Kitendakazi chako cha `padRow` kinapaswa kuwa na kigezo `name`.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
