---
id: 6610c8cfe4cf4d278e35c156
title: Hatua 63
challengeType: 1
dashedName: step-63
---

# --description--

Kama ilivyotarajiwa, kitendakazi chako sasa kinarudisha `undefined` tena. Kigezo chako cha `call` hakihitajiki tena, kwa hivyo ondoa tamko la `call` na `console.log` kwa kigezo cha `call`.

# --hints--

Haupaswi kuwa na tamko la `call`.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

Haupaswi kuandika kumbukumbu ya kigezo chako cha `call`.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
