---
id: 6610c48c4ea0891afa7c4696
title: Hatua 54
challengeType: 1
dashedName: step-54
---

# --description--

Ukifungua tena konsoli yako, utaona kuwa kitendakazi chako cha `padRow` kinarejesha `undefined`, ingawa ulitaja thamani ya return! Hii ni kwa sababu vigezo vinahitaji kupewa thamani unapotumia wito wa kitendakazi.

Unapopita thamani kwa wito wa kitendakazi, thamani hiyo huitwa <dfn>hoja</dfn>. Hapa kuna mfano wa kutumia kitendakazi cha `demo` na kupitisha `"Naomi"` kama hoja kwa kigezo cha `name`.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

Pita jina lako kama hoja kwa kigezo cha `name` katika wito wako wa `padRow`. Kumbuka kuwa jina lako ni mfuatano wa herufi, hivyo utahitaji kutumia nukuu.

# --hints--

Unapaswa kupitisha mfuatano wa herufi kwa wito wako wa `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
