---
id: 6610bf6fa14d700beed1b109
title: Hatua 87
challengeType: 1
dashedName: step-87
---

# --description--

Opereta ya <dfn>usawa</dfn> `==` hutumika kuangalia kama thamani mbili ni sawa. Ili kulinganisha thamani mbili, utatumia usemi kama `value == 8`.

Chini ya `done++` ndani ya mzunguko wako, ongeza usemi wa `if`. Usemi huo unapaswa kuangalia kama `done` ni sawa na `count` kwa kutumia opereta ya usawa.

# --hints--

Unapaswa kutumia usemi wa `if` ndani ya mzunguko wako. Unapaswa kuongezwa baada ya `done++`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

Usemi wako wa `if` unapaswa kutumia opereta ya usawa kulinganisha `done` na `count` katika sharti.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
