---
id: 6610c16c4fa0df12c0e30675
title: Hatua 49
challengeType: 1
dashedName: step-49
---

# --description--

Unawito kitendakazi chako cha `padRow`, lakini hufanyi chochote na wito huo wa kitendakazi. Vitendakazi vyote katika JavaScript <dfn>hurejesha</dfn> thamani, ikimaanisha hutoa matokeo yaliyotangazwa ya kuitwa kwao ili utumie mahali pengine.

Ili kuona matokeo ya kuitwa kwa kitendakazi chako cha `padRow`, tangaza kigezo cha `call` na weka thamani ya wito wako wa `padRow` katika kigezo hicho.

# --hints--

Unapaswa kutangaza kigezo cha `call`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

Unapaswa kutumia `const` kutangaza kigezo chako cha `call`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

Unapaswa kuweka thamani ya `call` kuwa matokeo ya wito wako wa `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
