---
id: 6610c77d50636722e5b6be17
title: Hatua 60
challengeType: 1
dashedName: step-60
---

# --description--

Chini ya tamko la `return`, andika mfuatano wa herufi `"This works!"` kwenye konsoli.

Baada ya kufanya hivyo, utaona kuwa mfuatano wa herufi `"This works!"` hauonekani kwenye konsoli, na mstari wa `console.log("This works!")` umeandikwa kwa rangi ya kijivu.

Nakili kumbukumbu ya konsoli na ubandike juu ya tamko la `return`. Sasa, mfuatano wa herufi `"This works!"` unapaswa kuonekana kwenye konsoli.

Jambo muhimu la kujua kuhusu neno kuu la `return` ni kwamba halitumi tu kufafanua thamani itakayotolewa kutoka kwenye kitendakazi chako, bali pia linasimamisha utekelezaji wa msimbo wako ndani ya kitendakazi au tamko la block. Hii inamaanisha msimbo wowote baada ya tamko la `return` hautatekelezwa.

# --hints--

Kitendakazi chako cha `padRow` kinapaswa kurudisha kigezo cha `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

`console.log` yako ya kwanza inapaswa kuja baada ya neno kuu la `return`.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

`console.log` yako ya pili inapaswa kuja kabla ya neno kuu la `return`.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
