---
id: 6610b8017d1671f2814e8c77
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Unapotangaza mabadiliko bila kuanzisha thamani yake, huchukuliwa kuwa <dfn>haijaanzishwa</dfn>. Kwa sasa, mabadiliko yako ya `secondCharacter` hayajaanzishwa.

Ongeza `console.log()` ili kuona ni thamani gani mabadiliko yako ya `secondCharacter` yana.

# --hints--

Unapaswa kuongeza `console.log` ya pili.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Unapaswa kuandika mabadiliko yako ya `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
