---
id: 6610c87eac0f0b256d7b037e
title: Hatua 62
challengeType: 1
dashedName: step-62
---

# --description--

Kwa sababu kitendakazi chako hakitumi tena kigezo, kubadilisha hoja hakikuathiri.

Endelea na ondoa tamko la `test` na tamko la `return` kutoka kwenye kitendakazi chako cha `padRow`, ili kitendakazi kiwe tupu tena.

# --hints--

Kitendakazi chako cha `padRow` hakipaswi kuwa na kigezo cha `test`.

```js
assert.notMatch(padRow.toString(), /test/);
```

Kitendakazi chako cha `padRow` hakipaswi kurudisha thamani.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
