---
id: 6610c71600966a2191d3a64a
title: Hatua 59
challengeType: 1
dashedName: step-59
---

# --description--

Thamani zinazorejeshwa kutoka kwenye kitendakazi hutumika kwa kuita kitendakazi hicho. Unaweza kutumia wito wa kitendakazi moja kwa moja kama thamani inayorejeshwa, au kukamata thamani iliyorejeshwa katika kigezo. Kwa njia hii, unaweza kutumia thamani iliyopewa thamani katika kigezo chenye wigo wa ndani, nje ya kitendakazi ambacho kiliumbwa.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

Ili kutumia thamani yako ya `"Testing"`, irudishe kutoka kwenye kitendakazi cha `padRow` kwa kusahihisha tamko lako la `return` ili lirudishe tu kigezo cha `test`.


# --hints--

Kitendakazi chako cha `padRow` kinapaswa kurudisha kigezo cha `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
