---
id: 6610b741b54b90f0c0fb3d58
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Wakati majina ya vigezo ni maneno zaidi ya moja, kuna mikataba maalum ya jinsi unavyopaswa kuweka herufi kubwa katika maneno hayo. Katika JavaScript, mkataba unaotumika ni <dfn>camel case</dfn>.

Camel case ina maana kwamba neno la kwanza katika jina linaandikwa kwa herufi ndogo kabisa, lakini maneno yanayofuata yote huandikwa kwa herufi kubwa mwanzoni mwa kila neno. Hapa kuna mifano ya camel case:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

Tumia camel case kutangaza kigezo kipya cha `secondCharacter`.

# --hints--

Unapaswa kutangaza kigezo cha `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

Unapaswa kutumia `let` kutangaza kigezo chako cha `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

Hupaswi kuweka thamani kwa kigezo chako cha `secondCharacter`. Usisahau alama ya semicolon.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
