---
id: 6610c424b7119919b62932f4
title: Hatua 53
challengeType: 1
dashedName: step-53
---

# --description--

Kitendakazi hakihitaji kurudisha thamani iliyowekwa moja kwa moja. Kinaweza kurudisha thamani iliyohifadhiwa katika kigezo. Vigezo ni vigezo maalum vya kitendakazi, kwa hivyo vinaweza pia kurudishwa.

Badilisha kitendakazi chako cha `padRow` ili `return` kigezo cha `name` moja kwa moja.

# --hints--

Kitendakazi chako cha `padRow` kinapaswa kurudisha thamani ya kigezo cha `name`.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
