---
id: 6610c538372aa61cc0f5b122
title: Hatua 57
challengeType: 1
dashedName: step-57
---

# --description--

Mabadiliko katika JavaScript yanapatikana katika <dfn>wigo</dfn> maalum. Kwa maneno mengine, mahali ambapo mabadiliko yanatangazwa huamua mahali katika msimbo wako yanapoweza kutumika.

Wigo wa kwanza ni wigo wa ulimwengu (global scope). Mabadiliko yanayotangazwa nje ya "kizuizi" chochote kama kitendakazi au `for` loop yako yako katika <dfn>wigo wa ulimwengu</dfn>. Mabadiliko yako ya `character`, `count`, na `rows` yote yako katika wigo wa ulimwengu.

Wakati mabadiliko yako yako katika wigo wa ulimwengu, kitendakazi kinaweza kuyapata katika ufafanuzi wake. Hapa kuna mfano wa kitendakazi kinachotumia mabadiliko ya `title` ya ulimwengu:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

Mfano huu utarudisha `"Professor Naomi"`. Sasisha kitendakazi chako cha `padRow` kurudisha thamani ya kuunganisha mabadiliko yako ya `character` mwanzoni mwa kigezo cha `name`.

# --hints--

Kitendakazi chako cha `padRow` kinapaswa kuunganisha `character` mwanzoni mwa `name`.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

Kitendakazi chako cha `padRow` kinapaswa kurudisha matokeo ya `character + name`.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
