---
id: 6610c105bbdacc114d6cdc44
title: Hatua 48
challengeType: 1
dashedName: step-48
---

# --description--

Ili kutumia kitendakazi, unahitaji kuitaja. <dfn>Wito wa kitendakazi</dfn> unaambia programu yako kuendesha msimbo kutoka kwa kitendakazi popote unapotaka kuitaja. Sintaksia ya wito wa kitendakazi ni jina la kitendakazi likifuatiwa na mabano ya kawaida. Kwa mfano, msimbo huu unaeleza na kuita kitendakazi cha `test`.

```js
function test() {

}
test();
```

Itaje kitendakazi chako cha `padRow`.

# --hints--

Unapaswa kuita kitendakazi cha `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Mwili wa kitendakazi chako cha `padRow` unapaswa kuwa tupu.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
