---
id: 6610c21b3ef82015573ffbbe
title: Hatua 51
challengeType: 1
dashedName: step-51
---

# --description--

Kigezo chako cha `call` kina thamani ya `undefined`, ingawa umeitaja! Hii ni kwa sababu kitendakazi chako cha `padRow` kwa sasa hakirudishi thamani yoyote. Kwa chaguo-msingi, vitendakazi hurudisha `undefined` kama thamani yao.

Ili kurudisha kitu kingine, unahitaji kutumia neno kuu la `return`. Hapa kuna mfano wa kitendakazi kinachorudisha mfuatano wa herufi `"Functions are cool!"`:

```js
function demo() {
  return "Functions are cool!";
}
```

Tumia neno kuu la `return` ili kitendakazi chako kirudishe mfuatano wa herufi `"Hello!"`.

# --hints--

Kitendakazi chako cha `padRow` kinapaswa kutumia neno kuu la `return`.

```js
assert.match(padRow.toString(), /return/);
```

Kitendakazi chako cha `padRow` kinapaswa kurudisha mfuatano wa herufi `"Hello!"`.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
