---
id: 6610b8f6a98d25f4d485a94d
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Thamani ya msingi ya kinabadilisho kisichouanzishwa ni `undefined`. Hii ni aina maalum ya data inayowakilisha thamani ambayo bado haina ufafanuzi.

Bado unaweza kuweka thamani kwa kinabadilisho kisichouanzishwa. Hapa kuna mfano:

```js
let uninitialized;
uninitialized = "assigned";
```

Weka mfuatano wa herufi `"Test"` kwa kinabadilisho chako cha `secondCharacter` chini ya tamko lako. Fungua konsoli kuona jinsi kumbukumbu yako imebadilika.

# --hints--

Haupaswi kuanzisha `secondCharacter`. Kumbuka kuwa uanzishaji ni kuweka thamani wakati unapotangaza kinabadilisho.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

Unapaswa kutumia opereta ya mgawo kwenye `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

Unapaswa kuweka mfuatano wa herufi `"Test"` kwa kinabadilisho chako cha `secondCharacter`.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
