---
id: 66458f0a05df478aa627629e
title: Hatua 82
challengeType: 1
dashedName: step-82
---

# --description--

Maandishi yamepotea tena! Mfuatano wa herufi tupu huhesabiwa kuwa `false`, na kufanya kuwa <dfn>thamani ya uongo</dfn>. Utafundishwa zaidi kuhusu thamani chanya na thamani ya uongo katika miradi ijayo.

Mbali na tamko za `if`, JavaScript pia ina tamko za <dfn>else if</dfn>. Tamko za `else if` hukuruhusu kuangalia masharti mengi ndani ya kipande kimoja cha msimbo.

Hapa ni sintaksia ya tamko la `else if`:

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

Kama sharti la kwanza ni `false`, JavaScript itakagua sharti linalofuata katika mnyororo. Ikiwa sharti la pili ni `false`, JavaScript itakagua sharti la tatu, na kadhalika.

Chini ya tamko lako la `if`, ongeza tamko la `else if` ambalo linakagua kama `5` ni ndogo kuliko `10`. Kisha ndani ya mwili wa tamko la `else if`, andika mfuatano wa herufi `"5 is less than 10"` kwenye konsoli.

Kagua konsoli kuona matokeo.

# --hints--

Unapaswa kuwa na tamko la `else if`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Tamko lako la `else if` linapaswa kukagua kama `5` ni ndogo kuliko `10`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

Unapaswa kuandika mfuatano wa herufi `"5 is less than 10"` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
