---
id: 664599653fcd6e97104f9261
title: Hatua 83
challengeType: 1
dashedName: step-83
---

# --description--

Wakati mwingine utataka kuendesha msimbo tofauti wakati masharti yote ya `if...else if` ni `false`. Unaweza kufanya hivi kwa kuongeza kipande cha msimbo cha `else`.

Kipande cha msimbo cha `else` kitaangaliwa tu ikiwa masharti katika vipande vya msimbo vya `if` na `else if` hayatajwi.

Hapa kipande cha msimbo cha `else` kimeongezwa kwenye kipande cha msimbo cha `else if`.

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

Ongeza kipande cha msimbo cha `else` kwenye kipande cha msimbo cha `else if`. Ndani ya kipande cha msimbo cha `else`, andika mfuatano wa herufi `"This is the else block"` kwenye konsoli.

Ili kuona matokeo kwenye konsoli, unaweza kubadilisha kwa mkono `<` katika tamko la `else if` kuwa `>`. Hilo litatengeneza hali ya `false` na kipande cha msimbo cha `else` kitaendeshwa.

# --hints--

Unapaswa kuwa na kipande cha msimbo cha `else`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

Kipande chako cha msimbo cha `else` kinapaswa kuandika mfuatano wa herufi `"This is the else block"` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
