---
id: 664ed97c55a99e5ffab759dc
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Kabla ya kuendelea, unapaswa kuchukua muda kidogo kukagua dhana ulizojifunza.

Tumia neno kuu la `let` kutangaza variable ya `profession` na variable ya `age`. Anzisha `profession` na mfuatano wa herufi `"teacher"`, lakini usianzishe `age` na thamani yoyote.

Andika zote variable zako kwenye konsoli ili kuona matokeo.

# --hints--

Unapaswa kutangaza variable ya `profession`.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

Unapaswa kutumia `let` kutangaza variable ya `profession`.

```js
assert.match(code, /let\s+profession/);
```

Unapaswa kuweka thamani ya mfuatano wa herufi `"teacher"` kwa variable ya `profession`.

```js
assert.strictEqual(profession, "teacher");
```

Unapaswa kutangaza variable ya `age`.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

Unapaswa kutumia `let` kutangaza variable ya `age`.

```js
assert.match(code, /let\s+age/);
```

Haupaswi kuweka thamani kwa variable ya `age`.

```js
assert.isUndefined(age);
```

Unapaswa kuandika `age` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

Unapaswa kuandika `profession` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
