---
id: 664ed9344ba3495e1f67f8ea
title: Hatua 14
challengeType: 1
dashedName: step-14
---

# --description--

Sasa baada ya kukagua tamko na uanzishaji, ondoa msimbo uliouandika kwa ajili ya ukaguzi wako.

Usiondoe variable yako ya `character`.

# --hints--

Haupaswi kuwa na variable ya `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

Haupaswi kuwa na variable ya `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

Haupaswi kuwa na tamko la `console.log` kwa ajili ya `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

Haupaswi kuwa na tamko la `console.log` kwa ajili ya `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

Bado unapaswa kuwa na variable yako ya `character`.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
