---
id: 665732da4815b70bb083915e
title: Hatua 24
challengeType: 1
dashedName: step-24
---


# --description--

Katika hatua chache zilizopita, umejifunza yote kuhusu kufanya kazi na safu za data. Chukua muda kidogo ukague kile ulichojifunza.

Anza kwa kutangaza variable ya `cities` na kuianzisha kama safu ya data ya mfuatano wa herufi `"London"`, `"New York"`, na `"Mumbai"`. Kisha andika variable hiyo kwenye konsoli.

Baada ya kuandika, badilisha kipengele cha mwisho cha `cities` kuwa mfuatano wa herufi `"Mexico City"`, kisha andika variable ya `cities` tena.

Ukipata kwa usahihi, unapaswa kuona matokeo haya kwenye konsoli.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

Unapaswa kutumia `let` kutangaza variable ya `cities`.

```js
assert.match(code, /let\s+cities/);
```

Unapaswa kuweka safu ya data ya mfuatano wa herufi `"London"`, `"New York"`, na `"Mumbai"` kwenye variable ya `cities`.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

Unapaswa kutumia `console.log()` kuandika safu ya data yote ya `cities` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

Unapaswa kusahihisha kipengele cha mwisho cha safu ya data ya `cities` kuwa mfuatano wa herufi `"Mexico City"`. Kumbuka kuwa unaweza kupata kipengele cha mwisho cha safu ya data kwa kutumia `array[array.length - 1]`.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

Unapaswa kuwa na kauli mbili za `console.log(cities)` katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
