---
id: 66ea7e7399244c6d4b6052a1
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

La parte successiva del workshop rivedrà come funziona il metodo `Math.random()`.

Inizia aggiungendo un altro `console.log()` che scriva nella console la stringa `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."`.
 
# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere un `console.log()` che scriva nella console la stringa `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."`.

```js
assert.equal(
  getLogs()[1],
  "The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."
);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
