---
id: 66ea8239b9eeb76f12299030
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

La parte successiva del workshop è rivedere come generare un numero casuale tra due valori.

Inizia aggiungendo un altro `console.log()` che registra il messaggio `"Now, generate a random number between two values."` nella console.

Poi, crea una variabile chiamata `min` e assegnale il valore di `1` e crea una variabile chiamata `max` e assegnale il valore di `100`.

Nel passaggio successivo, genererai un numero casuale tra questi due valori.

# --hints--

Dovresti avere un `console.log()` che registra il messaggio `"Now, generate a random number between two values."`.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

Dovresti avere una variabile chiamata `min`.

```js
assert.isNotNull(min);
```

La tua variabile `min` dovrebbe essere un numero.

```js
assert.isNumber(min);
```

Dovresti assegnare il numero `1` alla tua variabile `min`.

```js
assert.equal(min, 1);
```

Dovresti avere una variabile chiamata `max`.

```js
assert.isNotNull(max);
```

La tua variabile `max` dovrebbe essere un numero.

```js
assert.isNumber(max);
```

Dovresti assegnare il numero `100` alla tua variabile `max`.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
