---
id: 66ea7fc048a6e26e33c598fe
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Ora è il momento di generare un numero casuale usando il metodo `Math.random()`.

Crea una variabile chiamata `randomNum` e assegnale il valore del risultato della chiamata al metodo `Math.random()`.

Poi, stampa la variabile `randomNum` nella console.

Prova ad aggiungere uno spazio nel codice per rieseguire il bot e vedere numeri casuali diversi generati.

# --hints--

Dovresti avere una variabile chiamata `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomNum);
```

Dovresti assegnare il risultato della chiamata al metodo `Math.random()` alla variabile `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
const randomNum = explorer.variables.randomNum;
assert.equal(randomNum?.value?.toString(), "Math.random()");
```

Dovresti stampare la variabile `randomNum` nella console.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
