---
id: 66ea845f54195c7020cc521b
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

La formula per generare un numero casuale tra due valori è la seguente:

```js
Math.random() * (maximum - minimum) + minimum;
```

Questo produrrà un risultato che è un numero a virgola mobile compreso tra due valori.

Crea una variabile chiamata `randomNum2` e assegnale il risultato della generazione di un valore tra i valori `min` e `max`.

Poi, stampa nella console la variabile `randomNum2` per vedere il risultato. Prova ad aggiornare la pagina per vedere risultati diversi.

# --hints--

Dovresti avere una variabile chiamata `randomNum2`.

```js
assert.isNotNull(randomNum2);
```

La tua variabile `randomNum2` dovrebbe essere un numero.

```js
assert.isNumber(randomNum2);
```

Dovresti assegnare il risultato della generazione di un valore tra i valori `min` e `max`. Assicurati di non inserire un valore fisso. Dovresti usare la formula fornita sopra.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

Dovresti stampare nella console la variabile `randomNum2`.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
