---
id: 66eb3ec3b07dfe8f898646c0
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Come hai imparato nelle lezioni precedenti, il metodo `Math.floor()` arrotonda il valore per difetto all'intero più vicino.

```js
const price = 10.99;
Math.floor(price); // 10
```

Crea una variabile chiamata `numRoundedDown` e assegnale il risultato dell'arrotondamento per difetto del numero a virgola mobile `6.7` all'intero più vicino.

Poi, stampa la variabile `numRoundedDown` nella console per vedere il risultato.

# --hints--

Dovresti avere una variabile `numRoundedDown`.

```js
assert.isNotNull(numRoundedDown);
```

La variabile `numRoundedDown` dovrebbe essere un numero.

```js
assert.isNumber(numRoundedDown);
```

La tua variabile `numRoundedDown` dovrebbe essere uguale al risultato dell'arrotondamento per difetto del numero a virgola mobile `6.7` all'intero più vicino.

```js
assert.equal(numRoundedDown, 6);
```

Non dovresti inserire manualmente il valore di `6` per la variabile `numRoundedDown`.

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

Dovresti stampare la variabile `numRoundedDown` nella console.

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
