---
id: 66eb488c0129e192542f0865
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Nelle lezioni precedenti, hai imparato a lavorare con il metodo `Math.ceil()` in questo modo:

```js
const price = 10.01;
Math.ceil(price); // 11
```

Il metodo `Math.ceil()` arrotonda il valore per eccesso all'intero più vicino.

Crea una variabile chiamata `numRoundedUp` e assegnale il risultato dell'arrotondamento per eccesso del numero a virgola mobile `3.2` all'intero più vicino.

Poi, stampa la variabile `numRoundedUp` nella console per vedere il risultato.

# --hints--

Dovresti avere una variabile `numRoundedUp`.

```js
assert.isNotNull(numRoundedUp);
```

La variabile `numRoundedUp` dovrebbe essere un numero.

```js
assert.isNumber(numRoundedUp);
```

La tua variabile `numRoundedUp` dovrebbe essere uguale al risultato dell'arrotondamento per difetto del numero a virgola mobile `3.2` all'intero più vicino.

```js
assert.equal(numRoundedUp, 4);
```

Non dovresti inserire manualmente il valore di `4` per la variabile `numRoundedUp`.

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

Dovresti stampare la variabile `numRoundedUp` nella console.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
