---
id: 66ece147323c91f654bf8dd9
title: Passaggio 15
challengeType: 1
dashedName: step-15
---

# --description--

Nelle lezioni precedenti, hai imparato come ottenere i valori massimo e minimo da un intervallo di numeri in questo modo:

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

Crea una variabile chiamata `maxNum` e assegnale il risultato di trovare il numero massimo tra i numeri `3, 125, 55, 24`. Poi, stampa il valore di `maxNum` nella console.

Sotto, crea una variabile chiamata `minNum` e assegnale il risultato di trovare il numero minimo tra i numeri `6, 90, 14, 90, 2`. Poi, stampa il valore di `minNum` nella console.

# --hints--

Dovresti avere una variabile chiamata `maxNum`.

```js
assert.isNotNull(maxNum);
```

La tua variabile `maxNum` dovrebbe essere assegnata al risultato di trovare il numero massimo tra i numeri `3, 125, 55, 24`.

```js
assert.equal(maxNum, 125);
```

Non dovresti inserire manualmente il valore di `125` nella variabile `maxNum`. Assicurati di usare invece il metodo `Math.max()`.

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

Dovresti stampare il valore di `maxNum` nella console.

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

Dovresti avere una variabile chiamata `minNum`.

```js
assert.isNotNull(minNum);
```

La tua variabile `minNum` dovrebbe essere assegnata al risultato di trovare il numero minimo tra i numeri `6, 90, 14, 90, 2`.

```js
assert.equal(minNum, 2);
```

Non dovresti inserire manualmente il valore di `2` nella variabile `minNum`. Assicurati di usare invece il metodo `Math.min()`.

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

Dovresti stampare il valore di `minNum` nella console.

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
