---
id: 69218e46f4490b33f2481a1a
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

La formula per generare un numero intero casuale tra due valori è la seguente:

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

Questo produrrà un risultato che è un numero intero compreso tra due valori.

Crea una variabile chiamata `randomInt` e assegnale il risultato della generazione di un valore tra i valori `min` e `max`.

Poi, stampa nella console la variabile `randomInt` per vedere il risultato. Prova ad aggiornare la pagina per vedere risultati diversi.

# --hints--

Dovresti avere una variabile chiamata `randomInt`.

```js
assert.isNotNull(randomInt);
```

La tua variabile `randomInt` dovrebbe essere un numero.

```js
assert.isNumber(randomInt);
```

Dovresti assegnare il risultato della generazione di un valore tra i valori `min` e `max`. Assicurati di non inserire un valore fisso. Dovresti usare la formula fornita sopra.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

Dovresti stampare nella console la variabile `randomInt`.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
