---
id: 642db8c409d9991d0b3b2f0d
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

あなたのプロジェクトは基本的なHTMLコンテナとそれに対応するCSSから始まります。最初のタスクは、スプレッドシートのセルをプログラムで生成することです。

グローバルな`window`オブジェクトはブラウザのウィンドウ（またはタブ）を表します。これは`onload`プロパティを持っており、ページ全体（スタイルシートやスクリプトを含む）が読み込まれたときの動作を定義できます。

まず、`onload`の`window`プロパティをパラメータなしのアロー関数に設定してください。その関数内で、`container`変数を宣言し、`id`の`"container"`で要素を取得した値を代入してください。

# --hints--

`onload`オブジェクトの`window`プロパティにアクセスしてください。

```js
assert.match(code, /window\.onload/);
```

`onload`プロパティを関数に設定してください。

```js
assert.isFunction(window.onload);
```

アロー構文を使ってください。

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

`onload`関数はパラメータを取らないようにしてください。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

`container`関数内で`onload`変数を宣言してください。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

`container`変数は`const`で宣言してください。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

`document.getElementById()`を使ってください。

```js
assert.match(code, /document\.getElementById\(/);
```

`id`の`"container"`で要素を取得してください。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

`#container`に`container`要素を代入してください。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
