---
id: 642df9df4b5216350de7b0d2
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

数値の範囲を生成する関数が必要です。

空の`range`関数を宣言し、`start`と`end`のパラメータを受け取ってください。`Array()`コンストラクタを使い、空の配列を暗黙的に返してください。

# --hints--

`range`変数を宣言してください。

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

`const`を使って`range`変数を宣言してください。

```js
assert.match(code, /const\s+range/);
```

`range`変数は関数であるべきです。

```js
assert.isFunction(range);
```

`range`関数はアロー構文を使うべきです。

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

`range`関数は最初に`start`パラメータを受け取るべきです。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

`range`関数は2番目に`end`パラメータを受け取るべきです。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

`range`関数は暗黙の返り値を使うべきです。これは中括弧を使わないことを意味します。

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

`range`関数は`Array()`コンストラクタを使うべきです。プリミティブコンストラクタは`new`キーワードを必要としません。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

`Array()`コンストラクタには何も渡さないでください。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
