---
id: 642ddfdea4200e313f80a4b6
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`document` オブジェクトには、新しい HTML 要素を動的に作成できる `.createElement()` メソッドがあることを覚えておいてください。

`createLabel` 関数内で、`label` 変数を宣言し、新しい `div` 要素を代入してください。

# --hints--

`label` 関数内で `createLabel` 変数を宣言する必要があります。

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

`label` 変数は `const` で宣言してください。

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

`.createElement()` オブジェクトの `document` メソッドを使うべきです。

```js
assert.match(code, /document\.createElement\(/);
```

`"div"` メソッドに文字列 `.createElement()` を渡してください。

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

新しい `div` 要素を `label` に代入してください。

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
