---
id: 642e02be7845f13b014cd2b0
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`range` 関数ができたので、それを使って文字の範囲を作成することもできます。

`charRange` とアロー構文を使って `const` 関数を宣言してください。`start` と `end` のパラメータを受け取る必要があります。この関数は `range()` と `start` を引数として `end` を呼び出した結果を暗黙的に返すようにしてください。

# --hints--

`charRange` 変数を宣言してください。

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

`charRange` 変数は `const` で宣言してください。

```js
assert.match(code, /const\s+charRange/);
```

`charRange` 変数は関数である必要があります。

```js
assert.isFunction(charRange);
```

`charRange` 関数はアロー構文を使ってください。

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

`charRange` 関数は `start` を最初のパラメータとして受け取る必要があります。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

`charRange` 関数は `end` を2番目のパラメータとして受け取る必要があります。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

`charRange` 関数は暗黙の返り値を使ってください。

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

`charRange` 関数は `range` 関数を呼び出す必要があります。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

`start` と `end` を `range` の呼び出しに引数として渡してください。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
