---
id: 642e0011c45c893845842058
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`Array()` コンストラクタには、配列を特定の値で埋めるために使える `.fill()` メソッドがあります。これを使って配列を `start` の値で埋めることができます。

`.fill()` コンストラクタに `Array()` メソッドを連結し、`start` の値を渡してください。

# --hints--

`.fill()` メソッドを使うべきです。

```js
assert.match(code, /\.fill\(/);
```

`.fill()` コンストラクタに `Array()` メソッドを呼び出すべきです。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(/);
```

`start` メソッドに `.fill()` を渡すべきです。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(\s*start\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
const range = (start, end) => Array(end - start + 1);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
