---
id: 64345c560591891f64976f7a
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

ヘルパー関数が完成したので、`onload` イベントハンドラ内で `letters` 変数を宣言してください。`charRange()` と `"A"` の文字を引数として `"J"` を呼び出した結果を代入します。

# --hints--

`letters` 変数を宣言してください。

```js
assert.match(code, /(?:let|const|var)\s+letters/);
```

`const` を使って `letters` 変数を宣言してください。

```js
assert.match(code, /const\s+letters/);
```

`charRange()` の呼び出しを `letters` 変数に代入してください。

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(/);
```

`"A"` を `charRange()` 呼び出しの最初の引数として渡してください。

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1/);
```

`"J"` を `charRange()` 呼び出しの2番目の引数として渡してください。

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1\s*,\s*('|"|`)J\2\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
