---
id: 643475e13dc727231acd0f72
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

ネストされた `.forEach()` 呼び出しの中で、`input` 変数を宣言してください。`.createElement()` オブジェクトの `document` メソッドを使って `input` 要素を作成します。`type` 属性を `"text"` に、`id` 属性を `letter + number` に設定してください。

# --hints--

`input` 変数を宣言する必要があります。

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

`const` を使って `input` 変数を宣言してください。

```js
assert.match(code, /const\s+input/)
```

`.createElement()` オブジェクトの `document` メソッドを呼び出してください。

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

`"input"` メソッドに文字列 `.createElement()` を渡してください。

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

新しい `input` 要素を `input` 変数に代入してください。

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

`type` 要素の `input` プロパティにアクセスしてください。

```js
assert.match(code, /input\.type/);
```

`type` 要素の `input` 属性を `"text"` に設定してください。

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

`id` 要素の `input` プロパティにアクセスしてください。

```js
assert.match(code, /input\.id/);
```

`id` 要素の `input` 属性を `letter + number` に設定してください。

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
