---
id: 643498755d54c6279ba09078
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

ほとんどのスプレッドシートプログラムには計算用の組み込み関数が含まれています。

`sum` 関数を宣言してください。この関数は `nums` パラメータを受け取ります。このパラメータは数値の配列です。配列のすべての数値を合計するために `reduce` を呼び出した結果を返すようにしてください。

# --hints--

`sum` 変数を宣言する必要があります。

```js
assert.match(code, /(?:let|const|var)\s+sum/);
```

`const` を使って `sum` 変数を宣言してください。

```js
assert.match(code, /const\s+sum/);
```

`sum` 変数は関数である必要があります。

```js
assert.isFunction(sum);
```

`sum` 関数はアロー構文を使うべきです。

```js
assert.match(code, /const\s+sum\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

`sum` 関数は `nums` パラメータを持つべきです。

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>/);
```

`sum` 関数は暗黙の戻り値を使うべきです。

```js
assert.notMatch(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

`sum` 関数は `.reduce()` を `nums` に呼び出した結果を返すべきです。

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*nums\.reduce\(/);
```

`sum` 関数は `nums` のすべての数値の合計を返すべきです。

```js
const numbers = [1, 2, 3, 4, 5];
assert.equal(sum(numbers), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
