---
id: 6434750c53db16218f41e6e1
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

`range()` は配列を返すので、関数呼び出しに直接配列メソッドを連結できることを覚えておいてください。

`range()` を `1` と `99` を引数として呼び出し、`.forEach()` メソッドを連結してください。`.forEach()` メソッドには、`number` をパラメータに取る空のコールバック関数を渡してください。

# --hints--

`range()` 関数を呼び出す必要があります。

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

`1` を `range()` 呼び出しの最初の引数として渡す必要があります。

```js
assert.match(code, /range\(\s*1/);
```

`99` を `range()` 呼び出しの2番目の引数として渡す必要があります。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

`.forEach()` メソッドを `range()` 呼び出しに連結する必要があります。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

`.forEach()` にアロー構文を使ったコールバック関数を渡す必要があります。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

コールバック関数は `number` を唯一のパラメータにする必要があります。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
