---
id: 6434552bcc0a951a0a99df3b
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

あなたの`range`関数は数値を期待していますが、`start`と`end`の値は文字列（具体的には`A`のような単一文字）になります。

`start`の呼び出しで、`end`と`range()`の値を`.charCodeAt()`メソッドを使って数値に変換してください。そのメソッドには引数として数値の`0`を渡します。

# --hints--

`.charCodeAt()`メソッドを使うべきです。

```js
assert.match(code, /\.charCodeAt\(/);
```

`.charCodeAt()`に対して`start`メソッドを呼び出すべきです。

```js
assert.match(code, /start\.charCodeAt\(/);
```

`0`の`.charCodeAt()`メソッドに`start`を渡すべきです。

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

`.charCodeAt()`に対して`end`メソッドを呼び出すべきです。

```js
assert.match(code, /end\.charCodeAt\(/);
```

`0`の`.charCodeAt()`メソッドに`end`を渡すべきです。

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

`.charCodeAt()`の呼び出しで`range`メソッドを直接使うべきです。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
