---
id: 643715013330824ecaa70442
title: ステップ 23
challengeType: 0
dashedName: step-23
---

# --description--

次の関数では、数値の配列の中央値を計算します。まず、`median` アロー関数を宣言し、`nums` パラメータを受け取るようにしてください。

関数内で、`sorted` 変数を宣言し、`nums` 配列のコピーをソートした値を代入してください。

配列の浅いコピーを作成するには、`slice()` メソッドを使うべきです。

# --hints--

`median` 変数を宣言してください。

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

`const` を使って `median` 変数を宣言してください。

```js
assert.match(code, /const\s+median/);
```

`median` 変数は関数であるべきです。

```js
assert.isFunction(median);
```

`median` 関数はアロー構文を使うべきです。

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

`median` 関数は `nums` パラメータを持つべきです。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

`median` 関数は暗黙の戻り値を使わないでください。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

`median` 関数は `sorted` 変数を持つべきです。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

`const` を使って `sorted` 変数を宣言してください。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

`.slice()` を使って `nums` 配列のコピーを `sorted` に代入してください。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

`.sort()` メソッドを `.slice()` メソッドにチェーンしてください。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

`sort` メソッドにコールバック関数を渡して、数値を正しく昇順にソートしてください。明確にするために暗黙の戻り値を使ってください。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
