---
id: 6437133052eaf04d7300e622
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

`average` 関数を宣言してください。この関数は `nums` 配列を `nums` パラメータとして受け取ります。配列内のすべての数値の平均を返す必要があります。

平均は、配列内のすべての数値の合計を配列の長さで割ることで計算できます。使用できる `sum` 関数があることを覚えておいてください。

# --hints--

`average` 変数を宣言してください。

```js
assert.match(code, /(?:let|const|var)\s+average/);
```

`const` を使って `average` 変数を宣言してください。

```js
assert.match(code, /const\s+average/);
```

`average` 変数は関数であるべきです。

```js
assert.isFunction(average);
```

`average` 関数はアロー構文を使うべきです。

```js
assert.match(code, /const\s+average\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

`average` 関数は `nums` パラメータを持つべきです。

```js
assert.match(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)/);
```

`average` 関数は暗黙の戻り値を使うべきです。

```js
assert.notMatch(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

`average` 関数は `nums` 配列の平均値を返すべきです。

```js
assert.equal(average([1,2,3]), 2);
assert.equal(average([1,2,3,4,5]), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
