---
id: 6437124c4c03dd4c8fb35d56
title: ステップ 21
challengeType: 0
dashedName: step-21
---

# --description--

`isEven` 関数を宣言してください。この関数は `num` パラメータを受け取り、数が偶数なら `true` を返し、そうでなければ `false` を返します。数が偶数か奇数かを判定するにはモジュラス演算子 `%` を使ってください。

# --hints--

`isEven` 変数を宣言する必要があります。

```js
assert.match(code, /(?:let|const|var)\s+isEven/);
```

`const` を使って `isEven` 変数を宣言してください。

```js
assert.match(code, /const\s+isEven/);
```

`isEven` 変数は関数であるべきです。

```js
assert.isFunction(isEven);
```

`isEven` 関数はアロー構文を使うべきです。

```js
assert.match(code, /const\s+isEven\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

`isEven` 関数は `num` パラメータを持つべきです。

```js
assert.match(code, /const\s+isEven\s*=\s*(\(\s*num\s*\)|num)\s*=>/);
```

`isEven` 関数はモジュラス演算子 `%` を使うべきです。

```js
assert.match(isEven.toString(),  /%/);
```

`isEven` 関数は偶数の場合に `true` を返すべきです。

```js
assert.isTrue(isEven(2));
assert.isTrue(isEven(1000));
assert.isTrue(isEven(42));
```

`isEven` 関数は奇数の場合に `false` を返すべきです。

```js
assert.isFalse(isEven(1));
assert.isFalse(isEven(333));
assert.isFalse(isEven(777777777));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
const sum = nums => nums.reduce((acc, el) => acc + el, 0);

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
