---
id: 64496df724dd3716a71fe971
title: ステップ 26
challengeType: 0
dashedName: step-26
---

# --description--

オブジェクトのプロパティはキーと値のペアで構成されます。オブジェクトリテラルを宣言するとき、省略形プロパティ名を使うことができます。省略形プロパティ名の構文を使うと、変数名がプロパティのキーになり、その値がプロパティの値になります。

次の例では、`user` オブジェクトを `userId`、`firstName`、`loggedIn` のプロパティで宣言しています。

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

スプレッドシートのすべての関数を管理するために、`spreadsheetFunctions` オブジェクトを宣言してください。省略形構文を使って、`sum`、`average`、`median` を `spreadsheetFunctions` オブジェクトのプロパティとして設定してください。

# --hints--

`spreadsheetFunctions` 変数を宣言してください。

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

`const` を使って `spreadsheetFunctions` 変数を宣言してください。

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

`spreadsheetFunctions` 変数はオブジェクトである必要があります。

```js
assert.isObject(spreadsheetFunctions);
```

`spreadsheetFunctions` オブジェクトは `sum` プロパティを持つ必要があります。

```js
assert.property(spreadsheetFunctions, "sum");
```

`sum` プロパティは `sum` 関数である必要があります。

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

`spreadsheetFunctions` オブジェクトは `average` プロパティを持つ必要があります。

```js
assert.property(spreadsheetFunctions, "average");
```

`average` プロパティは `average` 関数である必要があります。

```js
assert.equal(spreadsheetFunctions?.average, average);
```

`spreadsheetFunctions` オブジェクトは `median` プロパティを持つ必要があります。

```js
assert.property(spreadsheetFunctions, "median");
```

`median` プロパティは `median` 関数である必要があります。

```js
assert.equal(spreadsheetFunctions?.median, median);
```

分割代入構文を使ってプロパティを代入してください。

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
