---
id: 65ae458e23954c3469e0c209
title: ステップ 18
challengeType: 0
dashedName: step-18
---


# --description--

これまでのプロジェクトで`setAttribute`メソッドについて学びました。JavaScriptで属性を更新する別の方法は、次の構文を使うことです。

```js
el.attribute = value;
```

ハイフン付きのHTML属性値、例えば`aria-label`のプロパティ名はキャメルケースに従い、`ariaLabel`のようになります。

```js
el.ariaLabel = "Aria Label Value";
```

`aria-label`要素の`input`属性を、`id`属性と同じ値に設定してください。


# --hints--

`input.ariaLabel`があるはずです。

```js
assert.match(code, /input\.ariaLabel\s*=/);
```


`input.ariaLabel`に`input.id`と同じ値を設定してください。

```js
assert.match(code, /input\.ariaLabel\s*=\s*((letter\s*\+\s*number)|(input.id))/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
     createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {
     const input = document.createElement("input");
     input.type = "text";
     input.id = letter + number;
 
    })   
--fcc-editable-region--
  })
}
```
