---
id: 68f890b73bbfc9efa7379e7b
title: バグのあるアプリで型変換エラーをデバッグする
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

あなたは地元のウェブ開発会社に入社したばかりで、最初の課題は前の開発者が残したバグのあるコードを修正することです。

彼らは簡単な算術演算を試みていましたが、何かがうまくいかず、結果がおかしくなっています。あなたの仕事はコードをレビューし、デバッグして修正し、期待通りに動作し読みやすくすることです。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `firstResult` という名前の変数があり、`5` と `10` を正しく加算して `15` の値を出すこと。
2. `secondResult` という名前の変数があり、`5` から `8` を正しく減算して `3` の値を出すこと。
3. `thirdResult` という名前の変数があり、2つの数字を加算して `6` の値を出すこと。現在式で使われているブール値を数字に置き換えてください。
4. `fourthResult` という名前の変数があり、2つの数字を加算して `8` の値を出すこと。現在式で使われているブール値を数字に置き換えてください。
5. `fifthResult` という名前の変数があり、2つの数字を正しく掛け算して `20` の積を得ること。現在式で使われている文字列を数字に置き換えてください。
6. `sixthResult` という名前の変数があり、2つの数字を正しく加算して `22` の値を出すこと。現在式で使われている `null` を数字に置き換えてください。

# --hints--

`firstResult` という名前の変数が必要です。

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` は数値 `15` を持つべきです。

```js
assert.strictEqual(firstResult, 15);
```

`secondResult` という名前の変数が必要です。

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` は数値 `3` を持つべきです。

```js
assert.strictEqual(secondResult, 3);
```

`thirdResult` という名前の変数が必要です。

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` は数値 `6` を持つべきです。

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` はブール値（`true` または `false`）を含んではいけません。

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

`fourthResult` という名前の変数が必要です。

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` は数値 `8` を持つべきです。

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` はブール値（`true` または `false`）を含んではいけません。

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

`fifthResult` という名前の変数が必要です。

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` は数値 `20` を持つべきです。

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` は文字列の掛け算（例: "10" * 2）を含んではいけません。

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

`sixthResult` という名前の変数が必要です。

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` は数値 `22` を持つべきです。

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` は式の中に `null` を含んではいけません。

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
