---
id: 587d7faf367417b2b2512be9
title: Tuma data kwa njia ya njia ya JavaScript XMLHttpRequest
challengeType: 6
forumTopicId: 301504
dashedName: post-data-with-the-javascript-xmlhttprequest-method
---

# --description--

Katika mifano iliyopita, ulipokea data kutoka kwa rasilimali ya nje. Pia unaweza kutuma data kwa rasilimali ya nje, mradi tu rasilimali hiyo inaunga mkono maombi ya AJAX na unajua URL.

Njia ya JavaScript ya `XMLHttpRequest` pia hutumika kutuma data kwa seva. Hapa kuna mfano:

```js
const xhr = new XMLHttpRequest();
xhr.open('POST', url, true);
xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
xhr.onreadystatechange = function () {
  if (xhr.readyState === 4 && xhr.status === 201) {
    const serverResponse = JSON.parse(xhr.response);
    document.getElementsByClassName('message')[0].textContent =
      serverResponse.userName + serverResponse.suffix;
  }
};
const body = JSON.stringify({ userName: userName, suffix: ' loves cats!' });
xhr.send(body);
```

Umeshuhudia baadhi ya njia hizi hapo awali. Hapa njia ya `open` huanzisha ombi kama `POST` kwa URL iliyotolewa ya rasilimali ya nje, na hupitisha `true` kama kigezo cha tatu - ikionyesha kufanya operesheni hiyo isiyo ya wakati mmoja.

Njia ya `setRequestHeader` huweka thamani ya kichwa cha ombi la HTTP, ambacho kina taarifa kuhusu mtumaji na ombi. Lazima iitwe baada ya njia ya `open`, lakini kabla ya njia ya `send`. Vigezo viwili ni jina la kichwa na thamani ya kuweka kama mwili wa kichwa hicho.

Ifuatayo, mshughulikiaji wa tukio la `onreadystatechange` hushughulikia mabadiliko ya hali ya ombi. `readyState` ya `4` ina maana operesheni imekamilika, na `status` ya `201` ina maana ombi lilifanikiwa. Kwa hiyo, HTML ya hati inaweza kusasishwa.

Mwishowe, njia ya `send` hutuma ombi na thamani ya `body`. `body` inajumuisha `userName` na funguo ya `suffix`.

# --instructions--

Sahihisha msimbo ili ufanye ombi la `POST` kwa kiishio cha API. Kisha andika jina lako katika sehemu ya ingizo na bofya `Send Message`. Kitendakazi chako cha AJAX kinapaswa kubadilisha `Reply from Server will be here.` na data kutoka kwa seva. Panga jibu ili kuonyesha jina lako likiambatanishwa na maandishi ya ` loves cats`.

# --hints--

Msimbo wako unapaswa kuunda `XMLHttpRequest` mpya.

```js
assert.match(code, /new\s*XMLHttpRequest\(\s*\)/g);
```

Msimbo wako unapaswa kutumia njia ya `open` kuanzisha ombi la `POST` kwa seva.

```js
assert.match(code, /\.open\(\s*('|")POST\1\s*,\s*url\s*,\s*true\s*\)/g);
```

Msimbo wako unapaswa kutumia njia ya `setRequestHeader`.

```js
assert.match(
  code,
  /\.setRequestHeader\(\s*('|")Content-Type\1\s*,\s*('|")application\/json;\s*charset=UTF-8\2\s*\)/g
);
```

Msimbo wako unapaswa kuwa na mshughulikiaji wa tukio la `onreadystatechange` uliowekwa kwa kitendakazi.

```js
assert.match(code, /\.onreadystatechange\s*=/g);
```

Msimbo wako unapaswa kupata kipengele chenye darasa la `message` na kubadilisha `textContent` yake kuwa `userName loves cats`.

```js
assert.match(
  code,
  /document\.getElementsByClassName\(\s*('|")message\1\s*\)\[0\]\.textContent\s*=\s*.+?\.userName\s*\+\s*.+?\.suffix/g
);
```

Msimbo wako unapaswa kutumia njia ya `send`.

```js
assert.match(code, /\.send\(\s*body\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line

      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message box">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line
      const xhr = new XMLHttpRequest();
      xhr.open('POST', url, true);
      xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
      xhr.onreadystatechange = function () {
        if (xhr.readyState === 4 && xhr.status === 201) {
          const serverResponse = JSON.parse(xhr.response);
          document.getElementsByClassName('message')[0].textContent =
            serverResponse.userName + serverResponse.suffix;
        }
      };
      const body = JSON.stringify({
        userName: userName,
        suffix: ' loves cats!'
      });
      xhr.send(body);
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```
