---
id: 587d7faf367417b2b2512be8
title: Pata data ya eneo la mtumizi ili kupata kuratibu za GPS za mtumizi
challengeType: 6
forumTopicId: 18188
dashedName: get-geolocation-data-to-find-a-users-gps-coordinates
---

# --description--

Jambo jingine la kuvutia unaloweza kufanya ni kupata upatikanaji wa eneo la mtumizi wako kwa sasa. Kivinjari kila kimoja kina kivinjari kilicho jumuishwa ambacho kinaweza kukupa taarifa hii.

Kivinjari kitapata mrefu na upana wa mtumizi kwa sasa.

Utaona onyo la kuomba ruhusa au kuzuia tovuti hii kujua eneo lako la sasa. Changamoto inaweza kukamilika kwa njia yoyote, mradi tu msimbo ni sahihi.

Kwa kuchagua ruhusa, utaona maandishi kwenye simu ya matokeo yanabadilika kuwa upana na mrefu wako.

Hapa kuna msimbo unaofanya hivi:

```js
if (navigator.geolocation) {
  navigator.geolocation.getCurrentPosition(function (position) {
    document.getElementById('data').innerHTML =
      'latitude: ' +
      position.coords.latitude +
      '<br>longitude: ' +
      position.coords.longitude;
  });
}
```

Kwanza, unakagua kama kitu cha `navigator.geolocation` kipo. Ikiwa kiko, njia ya `getCurrentPosition` kwenye kitu hicho inaitwa, ambayo huanzisha ombi isiyo ya wakati mmoja la nafasi ya mtumizi. Ikiwa ombi linafanikiwa, kitendakazi cha mwitiko katika njia hiyo hufanya kazi. Kitendakazi hiki kinapata thamani za kitu cha `position` kwa upana na mrefu kwa kutumia uandishi wa nukta na kusasisha HTML.

# --instructions--

Ongeza msimbo wa mfano ndani ya lebo za `script` ili ukague eneo la mtumizi kwa sasa na uliweke kwenye HTML.

# --hints--

Msimbo wako unapaswa kutumia `navigator.geolocation` kupata eneo la mtumizi kwa sasa.

```js
assert.match(code, /navigator\.geolocation\.getCurrentPosition/g);
```

Msimbo wako unapaswa kutumia `position.coords.latitude` kuonyesha eneo la upana la mtumizi.

```js
assert.match(code, /position\.coords\.latitude/g);
```

Msimbo wako unapaswa kutumia `position.coords.longitude` kuonyesha eneo la mrefu la mtumizi.

```js
assert.match(code, /position\.coords\.longitude/g);
```

Unapaswa kuonyesha nafasi ya mtumizi ndani ya kipengele cha `div` kwa kutumia `id="data"`.

```js
assert.match(
  code,
  /document\.getElementById\(\s*('|")data\1\s*\)\.innerHTML/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  // Add your code below this line

  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data"></div>
```

# --solutions--

```html
<script>
  // Add your code below this line
  if (navigator.geolocation) {
    navigator.geolocation.getCurrentPosition(function(position) {
      document.getElementById('data').innerHTML = "latitude: " + position.coords.latitude + "<br>longitude: " + position.coords.longitude;
    });
  }
  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data">

</div>

</section>
```
