---
id: 587d7fae367417b2b2512be5
title: Badilisha data ya JSON kuwa HTML
challengeType: 6
forumTopicId: 16807
dashedName: convert-json-data-to-html
---

# --description--

Sasa unapoanza kupata data kutoka kwa API ya JSON, unaweza kuionyesha katika HTML.

Unaweza kutumia njia ya `forEach` kuzunguka data kwa sababu vitu vya picha za paka vimehifadhiwa katika safu ya data. Unapofika kwenye kila kipengee, unaweza kubadilisha vipengele vya HTML.

Kwanza, tangaza kigezo cha html na `let html = "";`.

Kisha, zunguka JSON, ukiongeza HTML kwenye kigezo kinachozunguka majina ya funguo katika lebo za `strong`, ikifuatiwa na thamani. Unapomaliza kuzunguka, onyesha.

Huu ni msimbo unaofanya hivi:

```js
let html = '';
json.forEach(function (val) {
  const keys = Object.keys(val);
  html += "<div class = 'cat'>";
  keys.forEach(function (key) {
    html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
  });
  html += '</div><br>';
});
```

**Kumbuka:** Kwa changamoto hii, unahitaji kuongeza vipengele vipya vya HTML kwenye ukurasa, kwa hivyo huwezi kutegemea `textContent`. Badala yake, unahitaji kutumia `innerHTML`, ambayo inaweza kufanya tovuti kuwa hatarini kwa mashambulizi ya cross-site scripting.

# --instructions--

Ongeza njia ya `forEach` kuzunguka data ya JSON na kuunda vipengele vya HTML kuionyesha.

Hapa kuna mfano wa JSON:

```json
[
  {
    "id": 0,
    "imageLink": "https://cdn.freecodecamp.org/curriculum/legacy-json-apis-ajax/funny-cat.jpg",
    "altText": "A white cat wearing a green helmet shaped melon on its head. ",
    "codeNames": ["Juggernaut", "Mrs. Wallace", "Buttercup"]
  }
]
```

# --hints--

Msimbo wako unapaswa kuhifadhi data katika kigezo cha `html`

```js
assert.match(__helpers.removeWhiteSpace(code), /html(\+=|=html\+)/g);
```

Msimbo wako unapaswa kutumia njia ya `forEach` kuzunguka data ya JSON kutoka API.

```js
assert.match(code, /json\.forEach/g);
```

Msimbo wako unapaswa kuzunguka majina ya funguo katika lebo za `strong`.

```js
assert.match(code, /<strong>.+<\/strong>/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line

        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line
        json.forEach(function (val) {
          var keys = Object.keys(val);
          html += "<div class = 'cat'>";
          keys.forEach(function (key) {
            html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
          });
          html += '</div><br>';
        });
        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
