---
id: 587d7fad367417b2b2512be1
title: Shughulikia matukio ya bofya kwa JavaScript ukitumia sifa ya onclick
challengeType: 6
forumTopicId: 301503
dashedName: handle-click-events-with-javascript-using-the-onclick-property
---

# --description--

Unataka msimbo wako utekelezwe mara tu ukurasa wako unapomalizika kupakia. Kwa madhumuni hayo, unaweza kuambatanisha tukio la JavaScript kwenye hati linaloitwa `DOMContentLoaded`. Huu ndio msimbo unaofanya hivyo:

```js
document.addEventListener('DOMContentLoaded', function () {});
```

Unaweza kutekeleza wasimamizi wa tukio wanaoingia ndani ya kitendakazi cha `DOMContentLoaded`. Unaweza kutekeleza mshughulikiaji wa tukio la `onclick` unaochochewa wakati mtumizi anapobofya kipengee cha `#getMessage`, kwa kuongeza msimbo ufuatao:

```js
document.getElementById('getMessage').onclick = function () {};
```

# --instructions--

Ongeza mshughulikiaji wa tukio la bofya ndani ya kitendakazi cha `DOMContentLoaded` kwa kipengee chenye kitambulisho cha `getMessage`.

# --hints--

Msimbo wako unapaswa kutumia njia ya `document.getElementById` kuchagua kipengee chenye kitambulisho cha `getMessage`.

```js
assert.match(code, /document\s*\.getElementById\(\s*('|")getMessage\1\s*\)/g);
```

Msimbo wako unapaswa kuongeza mshughulikiaji wa tukio la `onclick`.

```js
assert.isFunction(document.getElementById('getMessage').onclick);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    document.getElementById('getMessage').onclick = function () {};
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
