---
id: 587d7fad367417b2b2512be2
title: Badilisha maandishi kwa matukio ya bofya
challengeType: 6
forumTopicId: 301500
dashedName: change-text-with-click-events
---

# --description--

Wakati tukio la bofya linapotokea, unaweza kutumia JavaScript kusahihisha kipengele cha HTML.

Kwa mfano, wakati mtumizi anapobofya kitufe cha `Get Message`, hubadilisha maandishi ya kipengele chenye darasa la vitu `message` kuwa `Here is the message`.

Hii hufanyika kwa kuongeza msimbo ufuatao ndani ya tukio la bofya:

```js
document.getElementsByClassName('message')[0].textContent =
  'Here is the message';
```

# --instructions--

Ongeza msimbo ndani ya mshughulikiaji wa tukio la `onclick` kubadilisha maandishi ndani ya kipengele cha `message` kuwa `Here is the message`.

# --hints--

Msimbo wako unapaswa kutumia njia ya `document.getElementsByClassName` kuchagua kipengele chenye darasa la vitu `message` na kuweka `textContent` yake kuwa mfuatano wa herufi uliotolewa.

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.textContent\s*?=\s*?('|")Here is the message\2/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      document.getElementsByClassName('message')[0].textContent =
        'Here is the message';
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
