---
id: 686b9720ee1d032bd77a480a
title: ISBNバリデーターのデバッグ
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

ISBN（国際標準図書番号）は商業用の書籍に割り当てられる一意の識別子です。長さは10文字または13文字です。最後の文字はチェックディジットで、前の数字（メインの数字）から計算されます。

Camperbotは独自のISBNバリデーターを作ろうとしましたが、いくつかのミスをしてしまいました。

このラボでは、既存のコードを修正して正しく動作するようにします。

期待される動作：
ユーザーがプログラムを実行すると、`Enter ISBN and length: `というプロンプトが表示されます。ユーザーは検証したいISBNコードを`ISBN,length`形式で入力できます。ハイフンなしのISBNを入力し、カンマとその長さ（`10`または`13`）を続けて入力してください。

入力例：
ISBN-10コードの場合は`1530051126,10`。
ISBN-13コードの場合は`9781530051120,13`。

チェックディジットの求め方：

- メインの数字はISBN-10コードの場合は最初の9桁、ISBN-13コードの場合は最初の12桁です。すべて数字でなければなりません。
- 最後の文字は与えられたチェックディジットです。期待されるチェックディジットはメインの数字から計算され、両者が一致した場合にのみISBNコードは有効です。
- 例えば、ISBN-10コード`1530051126`ではメインの数字は`153005112`で、与えられたチェックディジットは`6`です。ISBN-13コード`9781530051120`ではメインの数字は`978153005112`で、与えられたチェックディジットは`0`です。
- このラボでは詳細な計算ロジックを知る必要はありません。`calculate_check_digit_10`関数と`calculate_check_digit_13`関数が計算を行い、期待されるチェックディジットを文字列で返します。
- ISBN-10コードのチェックディジットは`0`から`9`の数字か大文字の`X`です。
- ISBN-13コードのチェックディジットは`0`から`9`の数字です。

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. 現在のコードの`IndentationError`を修正してください。
1. ユーザーがカンマ区切りの値を入力しなくても、プログラムは`IndexError`をクラッシュせずに処理できるようにしてください。
1. ユーザーがカンマ区切りの値を入力しなかった場合、コンソールに`Enter comma-separated values.`というメッセージを表示し、プログラムを終了してください。
1. ユーザーが長さに数値以外の値を入力しても、プログラムは`ValueError`をクラッシュせずに処理できるようにしてください。
1. ユーザーが長さに数値以外の値を入力した場合、コンソールに`Length must be a number.`というメッセージを表示し、プログラムを終了してください。
1. `validate_isbn`関数のオフバイワンエラーを修正してください。
1. ユーザーが有効なISBNコードを入力したときに発生する`TypeError`を修正してください。
1. ユーザーが有効なISBNコードを入力したときに発生する`IndexError`を修正してください。
1. ユーザーがメインの数字に数字以外の文字を含む誤ったISBNコードを入力しても、発生する`ValueError`をクラッシュせずに処理できるようにしてください。
1. ユーザーがメインの数字に数字以外の文字を含む誤ったISBNコードを入力した場合、コンソールに`Invalid character was found.`というメッセージを表示してください。
1. ユーザーが`1530051126,10`を入力した場合、`Valid ISBN Code.`というメッセージを表示してください。
1. ユーザーが`9781530051120,13`を入力した場合、`Valid ISBN Code.`というメッセージを表示してください。

**重要**：テストを正しく実行するために、グローバルスペースの`main()`呼び出しをコメントアウトする必要があります。

プロジェクトを完了すると、ユーザーが入力した値に応じて以下のメッセージが表示されます。

| ISBNコード | 長さ | メッセージ | 入力例 |
|-----------|--------|---------| ------------- |
| 有効 | 有効 | `Valid ISBN Code.` | `1530051126,10` |
| 無効な数字 | 有効 | `Invalid ISBN Code.` | `1530051125,10` |
| 指定された長さと一致しないか空欄 | 有効 | `ISBN-10 code should be 10 digits long.` または `ISBN-13 code should be 13 digits long.`（入力された長さによる） | `9781530051120,10` または `1530051126,13` |
| メインの数字に数字以外の文字が含まれる | 有効 | `Invalid character was found.` | `15-0051126,10` |
| どれでも | 無効な数字 | `Length should be 10 or 13.` | `1530051126,9` |
| どれでも | 数字以外の文字が含まれるか空欄 | `Length must be a number.` | `1530051125,A` |
| カンマ区切りでない | カンマ区切りでない | `Enter comma-separated values.` | `1530051125` |

手動でコードをテストしたい場合は、以下の値を使えます。

有効なISBN-10コードの入力例：
`1530051126,10`
`9971502100,10`
`080442957X,10`

有効なISBN-13コードの入力例：
`9781530051120,13`
`9781947172104,13`

# --hints--

テストを正しく動作させるために、`main`関数の呼び出しをコメントアウトしてください。

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

`validate_isbn`関数を用意してください。

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

`calculate_check_digit_10`関数を用意してください。

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

`calculate_check_digit_13`関数を用意してください。

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

`main`関数を用意してください。

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

ユーザーがカンマ区切りでない値を入力した場合、コンソールに`Enter comma-separated values.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

ユーザーが長さに数値以外の値を入力した場合、コンソールに`Length must be a number.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

ユーザーがメインの数字に数字以外の文字を含む誤ったISBNコードを入力した場合、コンソールに`Invalid character was found.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

ユーザーが`1530051126,10`を入力した場合、コンソールに`Valid ISBN Code.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

ユーザーが`9781530051120,13`を入力した場合、`Valid ISBN Code.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

ユーザーが`1530051125,10`を入力した場合、`Invalid ISBN Code.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

ユーザーが`9781530051120,10`を入力した場合、`ISBN-10 code should be 10 digits long.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

ユーザーが`1530051126,13`を入力した場合、`ISBN-13 code should be 13 digits long.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

ユーザーが`15-0051126,10`を入力した場合、`Invalid character was found.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

ユーザーが`1530051126,9`を入力した場合、`Length should be 10 or 13.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

ユーザーが`1530051125,A`を入力した場合、`Length must be a number.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

ユーザーが`1530051125`を入力した場合、`Enter comma-separated values.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

ユーザーが`9971502100,10`を入力した場合、`Valid ISBN Code.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

ユーザーが`080442957X,10`を入力した場合、`Valid ISBN Code.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

ユーザーが`9781947172104,13`を入力した場合、`Valid ISBN Code.`というメッセージが表示されます。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
