---
id: afcc8d540bea9ea2669306b6
title: Tengeneza kitendakazi cha kurudia mfuatano wa herufi
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

Katika maabara hii, utaunda kitendakazi kinachorudia mfuatano wa herufi uliotolewa mara fulani. Kwa madhumuni ya maabara hii, usitumie njia iliyo jengwa ndani ya `.repeat()`.

**Lengo:** Timiza hadithi za mtumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za mtumizi:**

1. Unapaswa kuunda kitendakazi kinachoitwa `repeatStringNumTimes` kinachopokea vigezo viwili: mfuatano wa herufi na nambari.
2. Kitendakazi kinapaswa kurudisha mfuatano wa herufi ukirudiwa mara nambari iliyotajwa.
3. Ikiwa nambari ni ndogo au sawa na sifuri, kitendakazi kinapaswa kurudisha mfuatano wa herufi tupu.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `repeatStringNumTimes`.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` inapaswa kupokea vigezo viwili.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

Kitendakazi `repeatStringNumTimes` kinapaswa kila mara kurudisha mfuatano wa herufi.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` inapaswa kurudisha mfuatano wa herufi `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` inapaswa kurudisha mfuatano wa herufi `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` inapaswa kurudisha mfuatano wa herufi `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` inapaswa kurudisha mfuatano wa herufi `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` inapaswa kurudisha mfuatano wa herufi `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` inapaswa kurudisha mfuatano wa herufi tupu (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` inapaswa kurudisha `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

Njia iliyo jengwa ndani ya `repeat()` haipaswi kutumika.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
