---
id: 5b7d72c338cd7e35b63f3e14
title: 브라우저 대체 수단으로 호환성 향상
challengeType: 0
forumTopicId: 301087
dashedName: improve-compatibility-with-browser-fallbacks
---

# --description--

CSS 작업을 하다 보면 언젠가 브라우저 호환성 문제를 만나게 될 가능성이 큽니다. 그래서 잠재적인 문제를 피하려면 브라우저 대체 수단을 제공하는 것이 중요합니다.

브라우저가 웹페이지의 CSS를 파싱할 때, 인식하거나 지원하지 않는 속성은 무시합니다. 예를 들어, 사이트에서 CSS 변수를 사용해 배경색을 할당하면 Internet Explorer는 CSS 변수를 지원하지 않기 때문에 배경색을 무시합니다. 이 경우 브라우저는 해당 속성에 대해 가지고 있는 값을 사용합니다. 만약 다른 값이 설정되어 있지 않으면 기본값으로 되돌아가는데, 보통 이 기본값은 이상적이지 않습니다.

즉, 브라우저 대체 수단을 제공하려면 기존 선언 바로 앞에 더 널리 지원되는 값을 하나 더 제공하면 됩니다. 이렇게 하면 오래된 브라우저는 대체 수단을 사용하고, 최신 브라우저는 캐스케이드에서 나중에 오는 선언을 해석합니다.

# --instructions--

`.red-box` 클래스의 배경색을 설정하는 데 변수가 사용된 것 같습니다. 기존 선언 바로 앞에 다른 `background` 선언을 추가하고 그 값을 `red`로 설정하여 브라우저 호환성을 향상하세요.

# --hints--

`.red-box` 규칙에 기존 `background` 선언 바로 앞에 `red`를 `background`로 설정한 대체 수단이 포함되어야 합니다.

```js
const spacelessCode = __helpers.removeWhiteSpace(__helpers.removeCssComments(code));
assert.match(
  spacelessCode,
  /\.red-box{background:(red|#ff0000|#f00|rgb\(255,0,0\)|rgb\(100%,0%,0%\)|hsl\(0,100%,50%\));background:var\(--red-color\);height:200px;width:200px;}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  :root {
    --red-color: red;
  }
  .red-box {

    background: var(--red-color);
    height: 200px;
    width:200px;
  }
</style>
<div class="red-box"></div>
```

# --solutions--

```html
<style>
  :root {
    --red-color: red;
  }
  .red-box {
    background: red;
    background: var(--red-color);
    height: 200px;
    width:200px;
  }
</style>
<div class="red-box"></div>
```
