---
id: bad87fee1348bd9aedf08807
title: 구글 폰트 가져오기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9MRsJ'
forumTopicId: 18200
dashedName: import-a-google-font
---

# --description--

대부분 운영체제에서 찾을 수 있는 일반 폰트를 지정하는 것 외에도, 웹사이트에서 사용할 비표준 맞춤 웹 폰트를 지정할 수 있습니다. 인터넷에는 웹 폰트의 다양한 출처가 있습니다. 이 예제에서는 Google Fonts 라이브러리에 집중하겠습니다.

Google Fonts는 폰트 URL을 참조하여 CSS에서 사용할 수 있는 무료 웹 폰트 라이브러리입니다.

그럼 구글 폰트를 가져와 적용해 보겠습니다(만약 사용자의 국가에서 Google이 차단되어 있다면 이 도전 과제를 건너뛰어야 합니다).

Google 폰트를 가져오려면 Google Fonts 라이브러리에서 폰트 URL을 복사한 다음 HTML에 붙여넣으면 됩니다. 이번 도전 과제에서는 `Lobster` 폰트를 가져올 것입니다. 이렇게 하려면 다음 코드 스니펫을 복사하여 코드 편집기 상단(열리는 `style` 요소 앞)에 붙여넣으세요.

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
```

이제 CSS에서 FAMILY_NAME으로 `Lobster`를 사용하여 `Lobster` 폰트를 사용할 수 있습니다. 다음 예제를 참고하세요.

```css
font-family: FAMILY_NAME, GENERIC_NAME;
```

GENERIC_NAME은 선택 사항이며, 지정한 다른 폰트를 사용할 수 없을 때 대체 수단으로 사용되는 폰트입니다. 이 내용은 다음 도전 과제에서 다룹니다.

패밀리 이름은 대소문자를 구분하며, 이름에 공백이 있을 경우 따옴표로 감싸야 합니다. 예를 들어 `"Open Sans"` 폰트를 사용하려면 따옴표가 필요하지만, `Lobster` 폰트는 필요하지 않습니다.

# --instructions--

`Lobster` 폰트를 웹 페이지에 가져오세요. 그런 다음 요소 선택자를 사용하여 `Lobster` 요소의 `font-family`를 `h2`로 설정하세요.

# --hints--

`Lobster` 폰트를 가져와야 합니다.

```js
assert.exists(document.querySelector('link[href*="googleapis" i]'));
```

`h2` 요소는 `Lobster` 폰트를 사용해야 합니다.

```js
const h2 = document.querySelector('h2'); 
const fontFamily = window.getComputedStyle(h2)['font-family']; 
assert.match(fontFamily, /lobster/i);
```

폰트를 변경하려면 `h2` 요소 선택자만 사용하세요.

```js
assert.match(__helpers.removeHtmlComments(code), /\s*[^\.]h2\s*\{\s*font-family\s*:\s*('|"|)Lobster\1\s*(,\s*('|"|)[a-z -]+\3\s*)?(;\s*\}|\})/gi);
```

`p` 요소는 여전히 `monospace` 폰트를 사용해야 합니다.

```js
const paragraphElement = document.querySelector('p');
const fontFamily = window.getComputedStyle(paragraphElement)['font-family']; 
assert.match(fontFamily, /monospace/i);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }

  h2 {
    font-family: Lobster;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
