---
id: bad87fee1348bd9aedf06756
title: 인라인 스타일로 클래스 선언 재정의하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDRha'
forumTopicId: 18252
dashedName: override-class-declarations-with-inline-styles
---

# --description--

id 선언이 `style` 요소 CSS 어디에 선언되든 클래스 선언을 재정의한다는 것을 증명했습니다.

CSS를 재정의하는 다른 방법도 있습니다. 인라인 스타일을 기억하시나요?

# --instructions--

인라인 스타일을 사용해 `h1` 요소를 흰색으로 만들어 보세요. 인라인 스타일은 다음과 같이 생겼습니다:

```html
<h1 style="color: green;">
```

`blue-text` 요소에 `pink-text`와 `h1` 클래스를 그대로 두세요.

# --hints--

`h1` 요소는 `pink-text` 클래스를 가져야 합니다.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

`h1` 요소는 `blue-text` 클래스를 가져야 합니다.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

`h1` 요소는 `orange-text` 아이디를 가져야 합니다.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

`h1` 요소는 인라인 스타일을 가져야 합니다.

```js
assert.exists(document.querySelector('h1[style]'));
```

`h1` 요소는 흰색이어야 합니다.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
