---
id: bad87fee1348bd9aedf08726
title: 특정 색상을 위한 16진수 코드 사용
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

CSS에서 색상을 나타내는 다른 방법이 있다는 것을 알고 계셨나요? 그중 하나가 16진수 코드, 줄여서 헥스 코드라고 합니다.

우리는 보통 각 자리 숫자에 0부터 9까지의 기호를 사용하는 <dfn>10진수</dfn> 또는 10진법 숫자를 사용합니다. <dfn>16진수</dfn>(또는 <dfn>헥스</dfn>)는 16진법 숫자입니다. 즉, 16개의 서로 다른 기호를 사용합니다. 10진수처럼 0-9 기호는 0부터 9까지의 값을 나타냅니다. 그리고 A, B, C, D, E, F는 10부터 15까지의 값을 나타냅니다. 총합해서 0부터 F까지가 16진수 자리 숫자를 나타내며, 총 16가지 가능한 값을 가집니다. <a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">16진수에 관한 자세한 정보는 여기</a>에서 확인하세요.

CSS에서는 빨강(R), 초록(G), 파랑(B) 컴포넌트 각각에 대해 2자리씩, 총 6자리 16진수 숫자를 사용해 색상을 나타낼 수 있습니다. 예를 들어, `#000000`는 검은색이며 가장 낮은 값이기도 합니다. <a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">RGB 색상 시스템에 관한 자세한 정보는 여기</a>에서 확인하세요.

```css
body {
  color: #000000;
}
```

# --instructions--

`black` 요소의 background-color에서 `body`이라는 단어를 16진수 코드 표현인 `#000000`로 바꾸세요.

# --hints--

`body` 요소는 검은색 `background-color`를 가져야 합니다.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

검은색의 16진수 코드를 `black`라는 단어 대신 사용해야 합니다.

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
