---
id: bad87fee1348bd9aedf08756
title: 한 스타일을 다른 스타일보다 우선시하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ8wnHv'
forumTopicId: 18258
dashedName: prioritize-one-style-over-another
---

# --description--

때때로 HTML 요소에 서로 충돌하는 여러 스타일이 적용될 수 있습니다.

예를 들어, `h1` 요소가 동시에 초록색과 분홍색일 수는 없습니다.

텍스트를 분홍색으로 만드는 클래스를 만든 다음 요소에 적용하면 어떻게 되는지 살펴보겠습니다. 우리의 클래스가 `body` 요소의 `color: green;` CSS 속성을 *재정의*할까요?

# --instructions--

요소에 분홍색을 지정하는 `pink-text`라는 CSS 클래스를 만드세요.

`h1` 요소에 `pink-text` 클래스를 지정하세요.

# --hints--

`h1` 요소에 `pink-text` 클래스를 지정해야 합니다.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

`<style>`에 `pink-text` CSS 클래스가 있어 `color`을 변경해야 합니다.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*\{\s*color\s*:\s*.+\s*;?\s*\}/g);
```

`h1` 요소가 분홍색이어야 합니다.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
</style>
<h1>Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```
