---
id: bad88fee1348bd9aedf08825
title: 요소의 패딩 조절하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cED8ZC2'
forumTopicId: 301083
dashedName: adjust-the-padding-of-an-element
---

# --description--

이제 잠시 Cat Photo App을 내려놓고 HTML 스타일링에 대해 더 배워보겠습니다.

이미 눈치채셨겠지만, 모든 HTML 요소는 기본적으로 작은 직사각형입니다.

세 가지 중요한 속성이 각 HTML 요소를 둘러싼 공간을 제어합니다: `padding`, `border`, 그리고 `margin`입니다.

요소의 `padding`는 요소의 내용과 `border` 사이의 공간 크기를 조절합니다.

여기서 파란 상자와 빨간 상자가 노란 상자 안에 중첩되어 있는 것을 볼 수 있습니다. 빨간 상자가 파란 상자보다 더 많은 `padding`를 가지고 있다는 점에 주목하세요.

파란 상자의 `padding`를 늘리면 텍스트와 그 주위 테두리 사이의 거리(`padding`)가 늘어납니다.

# --instructions--

파란 상자의 `padding`를 빨간 상자의 것과 같게 바꾸세요.

# --hints--

`blue-box` 클래스는 요소에 `20px`의 `padding`를 주어야 합니다.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
