---
id: bad87fee1348bd9aedf07756
title: Important를 사용해 모든 다른 스타일 재정의하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cm24rcp'
forumTopicId: 18249
dashedName: override-all-other-styles-by-using-important
---

# --description--

좋아요! 인라인 스타일이 `style` 요소 내 모든 CSS 선언을 재정의한다는 것을 방금 증명했습니다.

하지만 잠깐만요. CSS를 재정의하는 마지막 방법이 하나 더 있습니다. 이것이 가장 강력한 방법입니다. 하지만 그 전에 왜 CSS를 재정의해야 하는지 이야기해 보겠습니다.

많은 상황에서 CSS 라이브러리를 사용하게 됩니다. 이 라이브러리가 실수로 자신의 CSS를 재정의할 수 있습니다. 그래서 특정 요소에 반드시 특정 CSS가 적용되어야 할 때 `!important`를 사용할 수 있습니다.

우리의 `pink-text` 클래스 선언으로 다시 돌아가 봅시다. `pink-text` 클래스가 이후의 클래스 선언, id 선언, 인라인 스타일에 의해 재정의되었던 것을 기억하세요.

# --instructions--

핑크 텍스트 요소의 색상 선언에 `!important` 키워드를 추가해 `h1` 요소가 100% 핑크색이 되도록 하세요.

이렇게 하는 예시는 다음과 같습니다:

```css
color: red !important;
```

# --hints--

`h1` 요소는 `pink-text` 클래스를 가져야 합니다.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

`h1` 요소는 `blue-text` 클래스를 가져야 합니다.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

`h1` 요소는 `id`가 `orange-text`여야 합니다.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

`h1` 요소는 `color: white` 인라인 스타일을 가져야 합니다.

```js
const commentessCode = __helpers.removeHtmlComments(code);
assert.match(commentessCode, /<h1.*style/gi);
assert.match(commentessCode, /<h1.*style.*color\s*?:/gi);
```

`pink-text` 클래스 선언에는 모든 다른 선언을 재정의하기 위해 `!important` 키워드가 있어야 합니다.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*?\{[\s\S]*?color:.*pink.*!important\s*;?[^\.]*\}/g);
```

`h1` 요소는 핑크색이어야 합니다.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink !important;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
