---
id: bad87fee1348bd9aedf04756
title: 이후 CSS에서 스타일 재정의하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDQug'
forumTopicId: 18253
dashedName: override-styles-in-subsequent-css
---

# --description--

우리 `pink-text` 클래스가 `body` 요소의 CSS 선언을 재정의했습니다!

우리는 클래스가 `body` 요소의 CSS를 재정의할 수 있다는 것을 증명했습니다. 그렇다면 다음 논리적인 질문은 `pink-text` 클래스를 어떻게 재정의할 수 있느냐입니다.

# --instructions--

요소에 파란색을 지정하는 `blue-text`라는 추가 CSS 클래스를 만드세요. `pink-text` 클래스 선언 아래에 위치하도록 하세요.

`blue-text` 클래스를 `h1` 요소에 `pink-text` 클래스와 함께 적용하고, 어떤 클래스가 우선하는지 확인해 보세요.

HTML 요소에 여러 클래스 속성을 적용할 때는 다음과 같이 공백으로 구분합니다:

```html
class="class1 class2"
```

**참고:** HTML 요소에 클래스가 나열된 순서는 중요하지 않습니다.

하지만 `class` 섹션 내 `<style>` 선언 순서는 중요합니다. 두 번째 선언이 항상 첫 번째 선언보다 우선합니다. `.blue-text`이 두 번째로 선언되었기 때문에 `.pink-text`의 속성을 재정의합니다.

# --hints--

`h1` 요소에 `pink-text` 클래스를 지정하세요.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

`h1` 요소에 `blue-text` 클래스를 지정하세요.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

`blue-text`와 `pink-text` 모두 같은 `h1` 요소에 속해야 합니다.

```js
assert.isTrue(document.querySelector('.pink-text').classList.contains('blue-text'));
```

`h1` 요소는 파란색이어야 합니다.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }

  .blue-text {
    color: blue;
  }  
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```
