---
id: bad87fee1348bd9aedf08746
title: 본문 요소에서 스타일 상속하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bmdtR'
forumTopicId: 18204
dashedName: inherit-styles-from-the-body-element
---

# --description--

이제 모든 HTML 페이지에 `body` 요소가 있고, 그 `body` 요소도 CSS로 스타일을 적용할 수 있다는 것을 증명했습니다.

기억하세요, `body` 요소는 다른 HTML 요소처럼 스타일을 지정할 수 있으며, 다른 모든 요소는 `body` 요소의 스타일을 상속받습니다.

# --instructions--

먼저, 텍스트가 `h1`인 `Hello World` 요소를 만드세요.

그다음, 페이지의 모든 요소에 `green` 색상을 주기 위해 `color: green;` 요소의 스타일 선언에 `body`을 추가하세요.

마지막으로, `body` 요소의 스타일 선언에 `monospace`을 추가하여 `font-family: monospace;` 요소에 `body` 폰트 패밀리를 지정하세요.

# --hints--

`h1` 요소를 만들어야 합니다.

```js
assert.isNotEmpty(document.querySelectorAll('h1'));
```

`h1` 요소는 텍스트가 `Hello World`여야 합니다.

```js
assert.match(
  document.querySelector('h1').textContent,
  /hello world/i
);
```

`h1` 요소는 닫는 태그가 있어야 합니다.

```js
const commentlessCode = __helpers.removeHtmlComments(code);
assert.match(commentlessCode, /<\/h1>/g);
assert.match(commentlessCode, /<h1/g);
assert.lengthOf(commentlessCode.match(/<\/h1>/g), commentlessCode.match(/<h1/g).length);
```

`body` 요소는 `color` 속성이 `green`여야 합니다.

```js
const bodyElement = document.querySelector('body');
const color = window.getComputedStyle(bodyElement)['color']; 
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

`body` 요소는 `font-family` 속성이 `monospace`여야 합니다.

```js
const bodyElement = document.querySelector('body');
const fontFamily = window.getComputedStyle(bodyElement)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

`h1` 요소는 `monospace` 요소에서 `body` 폰트를 상속받아야 합니다.

```js
const h1Element = document.querySelector('h1');
const fontFamily = window.getComputedStyle(h1Element)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

`h1` 요소는 `green` 요소에서 `body` 색상을 상속받아야 합니다.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }

</style>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }

</style>
<h1>Hello World!</h1>
```
