---
id: bad87fee1348bd9aedf08826
title: 요소의 패딩을 지정할 때 시계 방향 표기법 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB7mBS9'
forumTopicId: 18346
dashedName: use-clockwise-notation-to-specify-the-padding-of-an-element
---

# --description--

요소의 `padding-top`, `padding-right`, `padding-bottom`, `padding-left` 프로퍼티를 각각 지정하는 대신, 다음과 같이 한 줄로 모두 지정할 수 있습니다:

```css
padding: 10px 20px 10px 20px;
```

이 네 값은 시계 방향처럼 작동합니다: 위, 오른쪽, 아래, 왼쪽 순서이며, 각 방향별 패딩 지시문을 사용하는 것과 정확히 같은 결과를 만듭니다.

# --instructions--

시계 방향 표기법을 사용해 `.blue-box` 클래스에 위와 왼쪽에는 `padding`의 `40px`를, 아래와 오른쪽에는 `20px`만 주도록 하세요.

# --hints--

`blue-box` 클래스는 요소의 위쪽에 `40px`의 `padding`를 주어야 합니다.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '40px');
```

`blue-box` 클래스는 요소의 오른쪽에 `20px`의 `padding`를 주어야 합니다.

```js
const blueBox = document.querySelector('.blue-box');
const paddingRight = window.getComputedStyle(blueBox)['padding-right'];
assert.strictEqual(paddingRight, '20px');
```

`blue-box` 클래스는 요소의 아래쪽에 `20px`의 `padding`를 주어야 합니다.

```js
const blueBox = document.querySelector('.blue-box');
const paddingBottom = window.getComputedStyle(blueBox)['padding-bottom'];
assert.strictEqual(paddingBottom, '20px');
```

`blue-box` 클래스는 요소의 왼쪽에 `40px`의 `padding`를 주어야 합니다.

```js
const blueBox = document.querySelector('.blue-box');
const paddingLeft = window.getComputedStyle(blueBox)['padding-left'];
assert.strictEqual(paddingLeft, '40px');
```

`blue-box` 클래스의 패딩을 설정할 때 시계 방향 표기법을 사용해야 합니다.

```js
const css =  __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(css, /\.blue-box\s*{[\s\S]*padding\s*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
