---
id: bad87fee1348bd9aedf08719
title: 축약된 16진수 코드 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpKAm'
forumTopicId: 18338
dashedName: use-abbreviated-hex-code
---

# --description--

1600만 가지가 넘는 색상 가능성에 압도당하는 사람이 많습니다. 그리고 16진수 코드를 기억하기도 어렵습니다. 다행히도 축약할 수 있습니다.

예를 들어, 빨간색의 16진수 코드는 `#FF0000`인데 `#F00`로 축약할 수 있습니다. 이 축약된 형태는 빨강, 초록, 파랑 각각 한 자리 숫자를 제공합니다.

이렇게 하면 가능한 색상 수가 약 4,000가지로 줄어듭니다. 하지만 브라우저는 `#FF0000`와 `#F00`를 정확히 같은 색상으로 해석합니다.

# --instructions--

축약된 16진수 코드를 사용해 올바른 요소에 색상을 지정해 보세요.

<table><tbody><tr><th>색상</th><th>축약 16진수 코드</th></tr><tr><td>시안</td><td><code>#0FF</code></td></tr><tr><td>초록</td><td><code>#0F0</code></td></tr><tr><td>빨강</td><td><code>#F00</code></td></tr><tr><td>자홍</td><td><code>#F0F</code></td></tr></tbody></table>

# --hints--

텍스트가 `h1`인 `I am red!` 요소에 `color` 빨간색을 지정하세요.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

빨간색의 축약 16진수 코드를 `#FF0000` 16진수 코드 대신 사용하세요.

```js
assert.match(code,/\.red-text\s*?{\s*?color\s*:\s*?#F00\s*?;?\s*?}/gi);
```

텍스트가 `h1`인 `I am green!` 요소에 `color` 초록색을 지정하세요.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

초록색의 축약 16진수 코드를 `#00FF00` 16진수 코드 대신 사용하세요.

```js
assert.match(__helpers.removeCssComments(code), /\.green-text\s*?{\s*?color\s*:\s*?#0F0\s*?;?\s*?}/gi);
```

텍스트가 `h1`인 `I am cyan!` 요소에 `color` 시안색을 지정하세요.

```js
const cyanText = document.querySelector('.cyan-text');
const color = window.getComputedStyle(cyanText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 255)');
```

시안색의 축약 16진수 코드를 `#00FFFF` 16진수 코드 대신 사용하세요.

```js
assert.match(__helpers.removeCssComments(code), /\.cyan-text\s*?{\s*?color\s*:\s*?#0FF\s*?;?\s*?}/gi);
```

텍스트가 `h1`인 `I am fuchsia!` 요소에 `color` 자홍색을 지정하세요.

```js
const fuchsiaText = document.querySelector('.fuchsia-text');
const color = window.getComputedStyle(fuchsiaText)['color']; 
assert.strictEqual(color,'rgb(255, 0, 255)');
```

자홍색의 축약 16진수 코드를 `#FF00FF` 16진수 코드 대신 사용하세요.

```js
assert.match(__helpers.removeCssComments(code), /\.fuchsia-text\s*?{\s*?color\s*:\s*?#F0F\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .fuchsia-text {
    color: #000000;
  }
  .cyan-text {
    color: #000000;
  }
  .green-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #F00;
  }
  .fuchsia-text {
    color: #F0F;
  }
  .cyan-text {
    color: #0FF;
  }
  .green-text {
    color: #0F0;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```
