---
id: bad87fee1348bd9aedf08721
title: 16진수 코드를 사용해 색상 혼합하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cK89PhP'
forumTopicId: 18359
dashedName: use-hex-code-to-mix-colors
---

# --description--

복습하자면, 16진수 코드는 6개의 16진수 숫자를 사용해 색상을 나타내며, 각각 빨강(R), 초록(G), 파랑(B) 컴포넌트에 두 자리씩 할당합니다.

이 세 가지 순수한 색상(빨강, 초록, 파랑)에서 각 색상의 양을 조절해 1,600만 가지 이상의 다른 색상을 만들 수 있습니다!

예를 들어, 주황색은 순수한 빨강에 약간의 초록이 섞이고 파랑은 전혀 없는 색입니다. 16진수 코드로는 `#FFA500`로 나타냅니다.

`0` 숫자는 16진수 코드에서 가장 낮은 숫자로, 색상이 전혀 없음을 나타냅니다.

`F` 숫자는 16진수 코드에서 가장 높은 숫자로, 최대 밝기를 나타냅니다.

# --instructions--

`style` 요소 안의 색상 단어들을 올바른 16진수 코드로 바꾸세요.

<table><tbody><tr><th>색상</th><th>16진수 코드</th></tr><tr><td>Dodger Blue</td><td><code>#1E90FF</code></td></tr><tr><td>Green</td><td><code>#00FF00</code></td></tr><tr><td>Orange</td><td><code>#FFA500</code></td></tr><tr><td>Red</td><td><code>#FF0000</code></td></tr></tbody></table>

# --hints--

텍스트가 `h1`인 `I am red!` 요소에 `color` 빨강색을 지정하세요.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

빨강색을 나타내는 `hex code`를 단어 `red` 대신 사용하세요.

```js
assert.match(code, /\.red-text\s*?{\s*?color\s*:\s*?(#FF0000|#F00)\s*?;?\s*?}/gi);
```

텍스트가 `h1`인 `I am green!` 요소에 `color` 초록색을 지정하세요.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

초록색을 나타내는 `hex code`를 단어 `green` 대신 사용하세요.

```js
assert.match(code, /\.green-text\s*?{\s*?color\s*:\s*?(#00FF00|#0F0)\s*?;?\s*?}/gi);
```

텍스트가 `h1`인 `I am dodger blue!` 요소에 `color` 도저 블루 색상을 지정하세요.

```js
const blueText = document.querySelector('.dodger-blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(30, 144, 255)');
```

도저 블루 색상을 나타내는 `hex code`를 단어 `dodgerblue` 대신 사용하세요.

```js
assert.match(code, /\.dodger-blue-text\s*?{\s*?color\s*:\s*?#1E90FF\s*?;?\s*?}/gi);
```

텍스트가 `h1`인 `I am orange!` 요소에 `color` 주황색을 지정하세요.

```js
const orangeText = document.querySelector('.orange-text');
const color = window.getComputedStyle(orangeText)['color']; 

assert.strictEqual(color, 'rgb(255, 165, 0)');
```

주황색을 나타내는 `hex code`를 단어 `orange` 대신 사용하세요.

```js
assert.match(code, /\.orange-text\s*?{\s*?color\s*:\s*?#FFA500\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: black;
  }
  .green-text {
    color: black;
  }
  .dodger-blue-text {
    color: black;
  }
  .orange-text {
    color: black;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #FF0000;
  }
  .green-text {
    color: #00FF00;
  }
  .dodger-blue-text {
    color: #1E90FF;
  }
  .orange-text {
    color: #FFA500;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```
