---
id: bad87fee1348bd9aedf08805
title: 요소 스타일링을 위한 CSS 선택자 사용
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKMBT2'
forumTopicId: 18349
dashedName: use-css-selectors-to-style-elements
---

# --description--

CSS를 사용하면 페이지에서 요소의 모양을 바꾸는 데 사용할 수 있는 수백 가지 CSS 속성이 있습니다.

`<h2 style="color: red;">CatPhotoApp</h2>`에 입력했을 때, 인라인 CSS로 개별 `h2` 요소를 스타일링하고 있었습니다. CSS는 Cascading Style Sheets의 약자입니다.

이것도 요소 스타일을 지정하는 한 방법이지만, CSS를 적용하는 더 좋은 방법이 있습니다.

코드 상단에 다음과 같은 `style` 블록을 만드세요:

```html
<style>
</style>
```

그 스타일 블록 안에서 모든 `h2` 요소에 대한 <dfn>CSS 선택자</dfn>를 만들 수 있습니다. 예를 들어, 모든 `h2` 요소를 빨간색으로 만들고 싶다면 다음과 같은 스타일 규칙을 추가하면 됩니다:

```html
<style>
  h2 {
    color: red;
  }
</style>
```

각 요소의 스타일 규칙 주위에 여는 중괄호(`{`)와 닫는 중괄호(`}`)가 모두 있어야 한다는 점을 기억하세요. 또한 요소의 스타일 정의가 여는 태그와 닫는 스타일 태그 사이에 있어야 합니다. 마지막으로, 각 스타일 규칙 끝에 세미콜론을 꼭 추가하세요.

# --instructions--

`h2` 요소의 스타일 속성을 삭제하고 대신 CSS `style` 블록을 만드세요. 모든 `h2` 요소를 파란색으로 바꾸는 데 필요한 CSS를 추가하세요.

# --hints--

`style` 속성은 `h2` 요소에서 삭제해야 합니다.

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

`style` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
assert.isAtLeast(document.querySelectorAll('style:not(.fcc-hide-header)').length, 1);
```

`h2` 요소는 파란색이어야 합니다.

```js
const h2Element = document.querySelector('h2');
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

스타일시트 `h2` 선언은 세미콜론과 닫는 중괄호가 포함되어 유효해야 합니다.

```js
assert.match(__helpers.removeCssComments(code), /h2\s*\{\s*color\s*:.*;\s*\}/g);
```

모든 `style` 요소는 유효하고 닫는 태그가 있어야 합니다.

```js
assert.match(__helpers.removeHtmlComments(code), /<\/style>/g);

const closingElementLength = __helpers.removeHtmlComments(code).match(/<\/style>/g).length;
const openingElementsLength = __helpers.removeHtmlComments(code).match(/<style((\s)*((type|media|scoped|title|disabled)="[^"]*")?(\s)*)*>/g).length;
assert.strictEqual(closingElementLength, openingElementsLength);
```

# --seed--

## --seed-contents--

```html
<h2 style="color: red;">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  h2 {
    color: blue;
  }
</style>
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
