---
id: bad87fee1348bd9aedf08736
title: HTML 본문 요소 스타일 지정
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB77PHW'
forumTopicId: 18313
dashedName: style-the-html-body-element
---

# --description--

이제 새로 시작해서 CSS 상속에 대해 이야기해 보겠습니다.

모든 HTML 페이지에는 `body` 요소가 있습니다.

# --instructions--

`body` 요소가 여기 있다는 것을 `background-color`를 검은색으로 지정해서 증명할 수 있습니다.

다음 내용을 `style` 요소에 추가하면 됩니다:

```css
body {
  background-color: black;
}
```

# --hints--

`body` 요소에 검은색 `background-color`가 있어야 합니다.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color'];

assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

CSS 규칙은 여는 중괄호와 닫는 중괄호가 모두 있어야 올바르게 작성됩니다.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

CSS 규칙은 세미콜론으로 끝나야 합니다.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

# --seed--

## --seed-contents--

```html
<style>

</style>
```

# --solutions--

```html
<style>
body {
  background-color: black;
}
</style>
```
