---
id: bad82fee1322bd9aedf08721
title: 절대 단위와 상대 단위 이해하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN66JSL'
forumTopicId: 301089
dashedName: understand-absolute-versus-relative-units
---

# --description--

최근 도전 과제들은 모두 요소의 여백이나 패딩을 픽셀(`px`) 단위로 설정했습니다. 픽셀은 길이 단위의 한 종류로, 브라우저가 항목의 크기나 간격을 정하는 데 사용합니다. `px` 외에도 CSS에는 사용할 수 있는 다양한 길이 단위 옵션이 있습니다.

길이 단위의 두 가지 주요 유형은 절대 단위와 상대 단위입니다. 절대 단위는 물리적 길이 단위에 연결됩니다. 예를 들어, `in`와 `mm`는 각각 인치와 밀리미터를 나타냅니다. 절대 길이 단위는 화면에서 실제 측정값에 근사하지만, 화면 해상도에 따라 약간의 차이가 있을 수 있습니다.

상대 단위는 `em`나 `rem`처럼 다른 길이 값에 상대적입니다. 예를 들어, `em`는 요소의 폰트 크기를 기준으로 합니다. 만약 `font-size` 속성 자체를 설정할 때 사용하면, 부모 요소의 `font-size`에 상대적입니다.

**참고:** 뷰포트 크기에 연결된 여러 상대 단위 옵션이 있습니다. 이는 반응형 웹 디자인 원칙 섹션에서 다룹니다.

# --instructions--

클래스 `padding`인 요소에 `red-box` 속성을 추가하고 `1.5em`로 설정하세요.

# --hints--

클래스 `red-box`에는 `padding` 속성이 있어야 합니다.

```js
const redBox =document.querySelector('.red-box'); 
const style = window.getComputedStyle(redBox); 

assert.notEqual(style['padding-top'], '0px');
assert.notEqual(style['padding-right'], '0px');
assert.notEqual(style['padding-bottom'], '0px');
assert.notEqual(style['padding-left'], '0px');
```

클래스 `red-box`는 요소에 1.5em의 `padding`를 주어야 합니다.

```js
assert.match(__helpers.removeCssComments(code), /\.red-box\s*?{[\s\S]*padding\s*:\s*?1\.5em/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;

  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;
    padding: 1.5em;
  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```
