---
id: bad87fee1348bd9afdf08726
title: 요소의 여백을 지정할 때 시계 방향 표기법 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpybAd'
forumTopicId: 18345
dashedName: use-clockwise-notation-to-specify-the-margin-of-an-element
---

# --description--

이번에는 `margin`를 사용해서 다시 시도해 보겠습니다.

요소의 `margin-top`, `margin-right`, `margin-bottom`, `margin-left` 속성을 각각 지정하는 대신, 이렇게 한 줄로 모두 지정할 수 있습니다:

```css
margin: 10px 20px 10px 20px;
```

이 네 값은 시계 방향처럼 작동합니다: 위, 오른쪽, 아래, 왼쪽 순서이며, 각 방향별 여백 지시문을 사용하는 것과 정확히 같은 결과를 만듭니다.

# --instructions--

시계 방향 표기법을 사용해서 `blue-box` 클래스를 가진 요소에 위와 왼쪽에는 `40px`의 여백을, 아래와 오른쪽에는 `20px`의 여백을 주세요.

# --hints--

`blue-box` 클래스는 요소의 위쪽에 `40px`만큼 `margin`을 주어야 합니다.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '40px');
```

`blue-box` 클래스는 요소의 오른쪽에 `20px`만큼 `margin`을 주어야 합니다.

```js
const blueBox = document.querySelector('.blue-box');
const marginRight = window.getComputedStyle(blueBox)['margin-right'];
assert.strictEqual(marginRight, '20px');
```

`blue-box` 클래스는 요소의 아래쪽에 `20px`만큼 `margin`을 주어야 합니다.

```js
const blueBox = document.querySelector('.blue-box');
const marginBottom = window.getComputedStyle(blueBox)['margin-bottom'];
assert.strictEqual(marginBottom, '20px');
```

`blue-box` 클래스는 요소의 왼쪽에 `40px`만큼 `margin`을 주어야 합니다.

```js
const blueBox = document.querySelector('.blue-box');
const marginLeft = window.getComputedStyle(blueBox)['margin-left'];
assert.strictEqual(marginLeft, '40px');
```

여백을 설정할 때 시계 방향 표기법을 사용해야 합니다 `blue-box` 클래스에.

```js
const cssCode = __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(cssCode, /\.blue-box\s*{[\s\S]*margin[\s]*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    margin: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
