---
id: bad82fee1348bd9aedf08721
title: RGB를 사용해 색상 혼합하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cm24JU6'
forumTopicId: 18368
dashedName: use-rgb-to-mix-colors
---

# --description--

16진수 코드와 마찬가지로, 서로 다른 값 조합을 사용해 RGB에서 색상을 혼합할 수 있습니다.

# --instructions--

`style` 요소에 있는 16진수 코드를 올바른 RGB 값으로 바꾸세요.

<table><tbody><tr><th>색상</th><th>RGB</th></tr><tr><td>파랑</td><td><code>rgb(0, 0, 255)</code></td></tr><tr><td>빨강</td><td><code>rgb(255, 0, 0)</code></td></tr><tr><td>오키드</td><td><code>rgb(218, 112, 214)</code></td></tr><tr><td>시에나</td><td><code>rgb(160, 82, 45)</code></td></tr></tbody></table>

# --hints--

텍스트가 `h1`인 `I am red!` 요소는 `color` 빨강이어야 합니다.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

빨간색에는 `rgb`를 사용하세요.

```js
assert.match(code, /\.red-text\s*{\s*color\s*:\s*rgb\(\s*255\s*,\s*0\s*,\s*0\s*\)\s*;?\s*}/gi);
```

텍스트가 `h1`인 `I am orchid!` 요소는 `color` 오키드이어야 합니다.

```js
const orchidText = document.querySelector('.orchid-text');
const color = window.getComputedStyle(orchidText)['color']; 
assert.strictEqual(color, 'rgb(218, 112, 214)');
```

오키드 색상에는 `rgb`를 사용하세요.

```js
assert.match(__helpers.removeCssComments(code), /\.orchid-text\s*{\s*color\s*:\s*rgb\(\s*218\s*,\s*112\s*,\s*214\s*\)\s*;?\s*}/gi);
```

텍스트가 `h1`인 `I am blue!` 요소는 `color` 파랑이어야 합니다.

```js
const blueText = document.querySelector('.blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

파란색에는 `rgb`를 사용하세요.

```js
assert.match(__helpers.removeCssComments(code), /\.blue-text\s*{\s*color\s*:\s*rgb\(\s*0\s*,\s*0\s*,\s*255\s*\)\s*;?\s*}/gi);
```

텍스트가 `h1`인 `I am sienna!` 요소는 `color` 시에나이어야 합니다.

```js
const siennaText = document.querySelector('.sienna-text');
const color = window.getComputedStyle(siennaText)['color']; 
assert.strictEqual(color, 'rgb(160, 82, 45)');
```

시에나 색상에는 `rgb`를 사용하세요.

```js
assert.match(__helpers.removeCssComments(code), /\.sienna-text\s*{\s*color\s*:\s*rgb\(\s*160\s*,\s*82\s*,\s*45\s*\)\s*;?\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .orchid-text {
    color: #000000;
  }
  .sienna-text {
    color: #000000;
  }
  .blue-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="orchid-text">I am orchid!</h1>

<h1 class="sienna-text">I am sienna!</h1>

<h1 class="blue-text">I am blue!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: rgb(255, 0, 0);
  }
  .orchid-text {
    color: rgb(218, 112, 214);
  }
  .sienna-text {
    color: rgb(160, 82, 45);
  }
  .blue-text {
    color:rgb(0, 0, 255);
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="orchid-text">I am orchid!</h1>

<h1 class="sienna-text">I am sienna!</h1>

<h1 class="blue-text">I am blue!</h1>
```
