---
id: bad87fee1348bd9aede08718
title: 요소 색상을 지정할 때 RGB 값을 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

CSS에서 색상을 나타내는 또 다른 방법은 `RGB` 값을 사용하는 것입니다.

검은색의 `RGB` 값은 다음과 같습니다:

```css
rgb(0, 0, 0)
```

흰색의 `RGB` 값은 다음과 같습니다:

```css
rgb(255, 255, 255)
```

16진수 코드를 사용할 때처럼 6자리 16진수 숫자를 사용하는 대신, `RGB`에서는 각 색상의 밝기를 0부터 255 사이의 숫자로 지정합니다.

계산해 보면, 한 색상의 두 자리 숫자는 16 곱하기 16으로 256개의 총 값을 나타냅니다. 그래서 `RGB`는 0부터 세기 시작하며, 16진수 코드와 정확히 같은 수의 가능한 값을 가집니다.

다음은 `body` 배경을 RGB 코드로 주황색으로 바꾸는 예시입니다.

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

`body` 요소의 배경색에 있는 16진수 코드를 검은색의 RGB 값인 `rgb(0, 0, 0)`로 바꾸세요.

# --hints--

`body` 요소는 검은색 배경이어야 합니다.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

`rgb`를 사용해서 `body` 요소에 검은색 배경을 지정해야 합니다.

```js
assert.match(__helpers.removeCssComments(code), /rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
