---
id: bad87fee1348bd9aedf08808
title: 폰트가 어떻게 대체되어야 하는지 지정하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cpVKBfQ'
forumTopicId: 18304
dashedName: specify-how-fonts-should-degrade
---

# --description--

모든 브라우저에서 사용할 수 있는 기본 폰트가 여러 개 있습니다. 이러한 일반 폰트 패밀리에는 `monospace`, `serif` 그리고 `sans-serif`가 포함됩니다.

하나의 폰트를 사용할 수 없을 때, 브라우저가 다른 폰트로 "대체"하도록 지정할 수 있습니다.

예를 들어, 어떤 요소에 `Helvetica` 폰트를 사용하되 `sans-serif`를 사용할 수 없을 때 `Helvetica` 폰트로 대체되도록 하려면 다음과 같이 지정하세요:

```css
p {
  font-family: Helvetica, sans-serif;
}
```

일반 폰트 패밀리 이름은 대소문자를 구분하지 않습니다. 또한 CSS 키워드이기 때문에 따옴표가 필요하지 않습니다.

# --instructions--

먼저, `monospace` 요소에 `h2` 폰트를 적용하여 이제 `Lobster`과 `monospace` 두 가지 폰트를 가지도록 하세요.

지난 도전 과제에서 `Lobster` 태그를 사용해 `link` 폰트를 가져왔습니다. 이제 Google Fonts에서 `Lobster` 폰트를 가져오는 부분을 HTML 주석 처리하여 더 이상 사용할 수 없게 하세요. 그러면 `h2` 요소가 `monospace` 폰트로 대체되는 것을 확인할 수 있습니다.

**참고:** 컴퓨터에 `Lobster` 폰트가 설치되어 있으면 브라우저가 해당 폰트를 찾을 수 있기 때문에 대체 현상이 보이지 않을 수 있습니다.

# --hints--

h2 요소는 `Lobster` 폰트를 사용해야 합니다.

```js
const h2Element = document.querySelector('h2');
const fontFamily = window.getComputedStyle(h2Element)['font-family']; 
assert.match(fontFamily, /^"?lobster/i);
```

h2 요소는 `monospace` 폰트를 사용할 수 없을 때 `Lobster` 폰트로 대체되어야 합니다.

```js
assert.match(__helpers.removeCssComments(code), /\s*h2\s*\{\s*font-family\s*\:\s*(\'|"|)Lobster\1\s*,\s*monospace\s*;?\s*\}/gi);
```

Google에서 `Lobster` 폰트를 가져오는 호출 앞에 `<!--`를 넣어 주석 처리해야 합니다.

```js
assert.match(code, /<!--[^fc]/gi);
```

주석은 `-->`를 추가하여 닫아야 합니다.

```js
assert.match(code, /[^fc]-->/gi);
```

# --seed--

## --seed-contents--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<!--<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">-->
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster, monospace;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
