---
id: bad87fee1348bd8aedf06756
title: ID 속성 스타일링으로 클래스 선언 재정의하기
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpDhB'
forumTopicId: 18251
dashedName: override-class-declarations-by-styling-id-attributes
---

# --description--

브라우저가 CSS를 선언 순서대로 위에서 아래로 읽는다는 것을 방금 증명했습니다. 즉, 충돌이 발생하면 브라우저는 가장 나중에 선언된 CSS를 사용합니다. `h1` 요소의 클래스에서 `blue-text`를 `pink-text` 앞에 두었더라도 선언 순서를 보고 사용 순서를 보지 않는다는 점을 기억하세요!

하지만 아직 끝나지 않았습니다. CSS를 재정의하는 다른 방법도 있습니다. `id` 속성을 기억하시나요?

`h1` 요소에 `id` 속성을 주고 그 속성을 스타일링해서 `pink-text`와 `blue-text` 클래스를 재정의하고 `h1` 요소를 주황색으로 만들어 보세요.

# --instructions--

`h1` 요소에 `id` 속성으로 `orange-text`를 주세요. `id` 스타일은 다음과 같이 생겼습니다:

```html
<h1 id="orange-text">
```

`h1` 요소에는 `blue-text`와 `pink-text` 클래스를 그대로 두세요.

`style` 요소에 `orange-text` id에 대한 CSS 선언을 만드세요. 예시는 다음과 같습니다:

```css
#brown-text {
  color: brown;
}
```

**참고:** `pink-text` 클래스 위나 아래에 이 CSS를 선언해도 상관없습니다. `id` 속성이 항상 우선하기 때문입니다.

# --hints--

`h1` 요소는 `pink-text` 클래스를 가져야 합니다.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

`h1` 요소는 `blue-text` 클래스를 가져야 합니다.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

`h1` 요소는 `orange-text` id를 가져야 합니다.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'),'orange-text');
```

`h1` 요소는 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`orange-text` id에 CSS 선언이 있어야 합니다.

```js
assert.match(__helpers.removeCssComments(code), /#orange-text\s*{/gi);
```

`h1` 요소에는 `style` 속성이 없어야 합니다.

```js
assert.notMatch(__helpers.removeHtmlComments(code), /<h1.*style.*>/gi);
```

`h1` 요소는 주황색이어야 합니다.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 165, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
  #orange-text {
    color: orange;
  }  
</style>
<h1 id="orange-text"  class="pink-text blue-text">Hello World!</h1>
```
