---
id: 66f53dc2c5bd6a11d6c3282f
title: Eine Passwortgenerator-App erstellen
challengeType: 26
dashedName: lab-password-generator
---

# --description--

In diesem Lab üben Sie den Umgang mit Funktionen, indem Sie einen zufälligen Passwortgenerator erstellen.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `generatePassword` erstellen, die einen Parameter entgegennimmt, der die Länge des generierten Passworts angibt. Sie können den Parameter beliebig benennen.
2. Ihre Funktion sollte einen String zurückgeben, der ein zufällig generiertes Passwort darstellt. Sie sollten den folgenden String und verschiedene `Math`-Methoden verwenden, um einen neuen String mit zufälligen Zeichen zurückzugeben: `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`.
3. Sie sollten eine Variable namens `password` definieren und ihr das Ergebnis des Aufrufs der `generatePassword`-Funktion mit einem numerischen Argument zuweisen, das die gewünschte Passwortlänge darstellt.
4. Sie sollten eine `console.log` haben, die einen einzelnen String ausgibt, der aus der Verkettung der Nachricht `Generated password:` und der `password`-Variablen besteht, getrennt durch ein Leerzeichen.

# --hints--

Sie sollten eine `generatePassword`-Funktion mit einem Parameter haben. Sie können den Parameter beliebig benennen.

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

Ihre `generatePassword`-Funktion sollte einen String zurückgeben.

```js
const result = generatePassword(5);
assert.isString(result);
```

Ihre `generatePassword`-Funktion sollte einen neuen String zurückgeben, der die korrekte Länge hat.

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

Ihre Funktion sollte ein zufällig generiertes Passwort mit gültigen Zeichen zurückgeben.

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

Ihre Funktion sollte ein zufällig generiertes Passwort zurückgeben, das mehr als ein einzigartiges Zeichen enthält.

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

Ihre Funktion sollte jedes Mal, wenn sie aufgerufen wird, einen neuen zufälligen String zurückgeben.

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

Sie sollten eine `password`-Variable haben.

```js
assert.isDefined(password);
```

Ihre `password`-Variable sollte ein String sein.

```js
assert.isString(password);
```

Sie sollten die `generatePassword`-Funktion mit einem numerischen Argument aufrufen und das zurückgegebene Passwort in der `password`-Variablen speichern.

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

Sie sollten einen einzelnen String ausgeben, der `Generated password:` und `password` durch ein einzelnes Leerzeichen getrennt kombiniert, und zwar mit `+` oder einem Template Literal.

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
