---
id: 6551f3bfed58796ebb1268dc
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Jedes Zeichen eines Strings kann über einen numerischen Index referenziert werden. Die Zählung der Indizes beginnt bei null. Das erste Zeichen eines Strings hat also den Index `0`. Zum Beispiel befindet sich in dem String `'Hello World'` `'H'` an Index `0`, `'e'` an Index `1` und so weiter.

Auf jedes Zeichen eines Strings kann mit der Klammernotation zugegriffen werden. Sie müssen den Variablennamen gefolgt von eckigen Klammern schreiben und den Index des Zeichens zwischen die Klammern setzen:

```py
text = 'Hello World'
r = text[8]
```

Geben Sie nun statt `text` nur das Zeichen an Index `6` aus.

# --hints--

Sie sollten weiterhin die `print()`-Funktion aufrufen.

```js
assert.match(code, /print\s*\(.*\)/)
```

Sie sollten `text[6]` an die `print()`-Funktion übergeben, indem Sie es zwischen die runden Klammern setzen. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
