---
id: 6551fb531ce7f773071a0b84
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Wie Sie sehen, ergibt die Ausgabe von `type(text)` `<class 'str'>`, was bedeutet, dass Ihre Variable ein String ist, angezeigt als `str`.

Gehen Sie nun in eine neue Zeile und erstellen Sie eine weitere Variable namens `shift` und weisen Sie dieser Variable den Wert `3` zu.

# --hints--

Sie sollten eine Variable namens `shift` deklarieren. Achten Sie darauf, den Variablennamen am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^shift\s*=/m)
```

Sie sollten den Wert `3` Ihrer Variable `shift` zuweisen.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
