---
id: 6551f4259d66026ee28cb205
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Sie können auch auf Zeichen eines Strings zugreifen, indem Sie vom Ende des Strings aus zählen. Das letzte Zeichen hat den Index `-1`, das vorletzte `-2` und so weiter.

Ändern Sie nun Ihren bestehenden `print()`-Aufruf so, dass das letzte Zeichen Ihres Strings ausgegeben wird.

# --hints--

Sie sollten weiterhin die `print()`-Funktion aufrufen.

```js
assert.match(code, /print\s*\(.*\)/)
```

Sie sollten `text[-1]` an die `print()`-Funktion übergeben, indem Sie es zwischen die runden Klammern setzen. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
