---
id: 6551f22653e6ce6e4f1d1370
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Ändern Sie Ihren `print(shift)`-Aufruf so, dass der Datentyp Ihrer `shift`-Variablen ausgegeben wird.

# --hints--

Sie sollten Ihren bestehenden `print(type(text))`-Aufruf beibehalten.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Sie sollten `shift` an die `type()`-Funktion übergeben.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

Sie sollten `type(shift)` an Ihre `print()`-Funktion übergeben. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

Sie sollten kein `print(shift)` in Ihrem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
