---
id: 6551f1e8737dbe6e0cb08ea4
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Jetzt drucke deine neue Variable.

# --hints--

Du solltet deinen bestehenden `print(type(text))`-Aufruf nicht entfernen.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Du solltest die `print()` Funktion aufrufen, indem du die `shift`Variable übergibst. Achte darauf, dass der Funktionsaufruf am Anfang der Zeile steht.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

Dein Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
