---
id: 6551f6f9ec6db270eb83f3db
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Sie sehen `11` im Terminal ausgegeben, weil `'Hello World'` 11 Zeichen enthält.

Eine weitere nützliche eingebaute Funktion ist `type()`, die den Datentyp einer Variablen zurückgibt. Ändern Sie Ihren `print()`-Aufruf so, dass der Datentyp von `text` ausgegeben wird.

# --hints--

Sie sollten die `type()`-Funktion aufrufen.

```js
assert.match(code, /type\s*\(.*\)/)
```

Sie sollten `text` an die `type()`-Funktion übergeben, indem Sie es zwischen die runden Klammern setzen.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

Sie sollten `type(text)` an die `print()`-Funktion übergeben, indem Sie es zwischen die runden Klammern setzen. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Sie sollten kein `print(len(text))` in Ihrem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückungen.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
