---
id: 6551f22653e6ce6e4f1d1370
title: Schitt 11
challengeType: 20
dashedName: step-11
---

# --description--

Ändere den `print(shift)` Aufruf, um den Datentyp deiner `shift` -Variable anzuzeigen.

# --hints--

Du solltest deinen bestehenden `print(type(text))`-Aufruf behalten.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Du solltest `shift` an die Funktion `type()` übergeben.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

Du solltest `type(shift)` an die Funktion `print()` übergeben. Achte darauf, dass der Funktionsaufruf am Anfang der Zeile steht.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

Du solltest kein `print(shift)` in deinem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

Dein Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
