---
id: 6551ef504d91936d2d4e54f8
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Ein *Argument* ist ein Objekt oder ein Ausdruck, der an eine Funktion übergeben wird – eingefügt zwischen den öffnenden und schließenden runden Klammern – wenn sie aufgerufen wird:

```py
greet = 'Hello!'
print(greet)
```

Der Code im obigen Beispiel würde den String `'Hello!'` ausgeben, welcher der Wert der Variablen `greet` ist, die als Argument an `print()` übergeben wurde.

Geben Sie Ihre `text`-Variable auf dem Bildschirm aus, indem Sie die `text`-Variable als Argument an die `print()`-Funktion übergeben.

# --hints--

Sie sollten `text` an die `print()`-Funktion übergeben, indem Sie den Namen dieser Variablen innerhalb der öffnenden und schließenden runden Klammern hinzufügen. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
