---
id: 6551f1e8737dbe6e0cb08ea4
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Geben Sie jetzt Ihre neue Variable aus.

# --hints--

Sie sollten Ihren bestehenden `print(type(text))` Aufruf nicht entfernen.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Sie sollten die `print()` Funktion aufrufen und dabei die `shift` Variable übergeben. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
