---
id: 6551ecf597038c36637f3759
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Variablen sind ein wesentlicher Bestandteil von Python und jeder Programmiersprache. Eine Variable ist ein Name, der auf ein Objekt verweist oder zeigt. Sie können eine Variable deklarieren, indem Sie den Variablennamen auf der linken Seite des Zuweisungsoperators `=` schreiben und den Wert, der dieser Variable zugewiesen werden soll, auf der rechten Seite des Zuweisungsoperators angeben:

```py
variable_name = value
```

Erstellen Sie eine Variable namens `number` und weisen Sie Ihrer neuen Variable den Wert `5` zu.

# --hints--

Sie sollten eine Variable namens `number` deklarieren. Achten Sie darauf, den Variablennamen am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^number\s*=/m)
```

Sie sollten der Variablen `5` den Wert `number` zuweisen.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
