---
id: 6551fb531ce7f773071a0b84
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Wie du sehehn kannst, ist die Ausgabe des `type(text)` `<class 'str'>`, was bedeutet, dass die Variable ein String ist, der als `str` angezeigt wird.

Gehe nun in eine neue Zeile und erstelle eine andere Variable namens `shift` und weise ihr den Wert `3` zu.

# --hints--

Du solltest eine Variable mit dem Namen `shift` deklarieren. Achte darauf, dass der Variablenname am Anfang der Zeile steht.

```js
assert.match(code, /^shift\s*=/m)
```

Du solltest der `shift`-Variablen den Wert `3` zuweisen.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

Dein Code enth채lt ung체ltige Syntax und/oder ung체ltige Einr체ckung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
