---
id: 6551f5756c20146fc94f8675
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Sie können mit der eingebauten `len()`-Funktion auf die Anzahl der Zeichen in einem String zugreifen.

Ändern Sie Ihren bestehenden `print()`-Aufruf, indem Sie `len(text)` anstelle von `text[-1]` übergeben.

# --hints--


Sie sollten die `len()`-Funktion aufrufen.

```js
assert.match(code, /len\s*\(.*\)/)
```

Sie sollten die Variable `text` an die `len()`-Funktion übergeben, indem Sie sie zwischen die runden Klammern setzen.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

Sie sollten `len(text)` an die `print()`-Funktion übergeben, indem Sie sie zwischen die runden Klammern setzen. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

Sie sollten `print(text[-1])` nicht in Ihrem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückungen.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
