---
id: 65524d2a1a253b8bb5197ae2
title: Schritt 28
challengeType: 20
dashedName: step-28
---

# --description--

`find` gibt wieder `-1` für Großbuchstaben zurück und auch für das Leerzeichen. Um das Leerzeichen kümmern Sie sich später.

Ändern Sie vorerst die `text`-Schleife so, dass sie nicht über `for`, sondern über `text.lower()` durchläuft.

# --hints--

Sie sollten Ihre `for`-Schleife so anpassen, dass sie über `text.lower()` und nicht über `text` iteriert.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
