---
id: 6552008c0d9d9075cbec9772
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Sie werden die `.find()`-Methode verwenden, um die Position jedes Buchstabens Ihrer Nachricht im Alphabet zu finden. Eine Methode ist ähnlich wie eine Funktion, gehört aber zu einem Objekt.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

Oben wird die `.find()`-Methode *auf* `sentence` (dem String, in dem gesucht wird) aufgerufen, und `'r'` (das zu findende Zeichen) wird als Argument übergeben. Der `sentence.find('r')`-Aufruf gibt `4` zurück, das der Index des ersten Vorkommens von `'r'` in `sentence` ist.

Am Ende Ihres Codes rufen Sie `.find()` auf `alphabet` auf und übergeben `'z'` als Argument an die Methode.

# --hints--

Sie sollten die `.find()`-Methode aufrufen.

```js
assert.match(code, /find\s*\(.*\)/)
```

Sie sollten die `.find()`-Methode auf der Variablen `alphabet` aufrufen.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


Sie sollten die `.find()`-Methode auf der Variablen `alphabet` aufrufen und `'z'` an die Methode übergeben. Achten Sie darauf, den Methodenaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


Ihr Code enthält ungültige Syntax und/oder ungültige Einrückungen.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
