---
id: 655208d59b131e7816f18c96
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Die `print()`-Funktion gibt Ihnen nur eine Ausgabe in der Konsole, aber Funktionen und Methoden können einen Rückgabewert haben, den Sie in Ihrem Code verwenden können.

Weisen Sie nun `alphabet.find(text[0])` einer Variablen namens `index` zu. Auf diese Weise speichert `index` den von `alphabet.find(text[0])` zurückgegebenen Wert.

# --hints--

Sie sollten eine Variable namens `index` deklarieren. Achten Sie darauf, den Variablennamen am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^index\s*=/m)
```

Sie sollten den von `alphabet.find(text[0])` zurückgegebenen Wert Ihrer Variablen `index` zuweisen.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
