---
id: 6552487e689f6e895f658717
title: Schritt 33
challengeType: 20
dashedName: step-33
---

# --description--

Jetzt müssen Sie am Ende Ihres Schleifenkörpers eine `new_char`-Variable erstellen. Setzen Sie ihren Wert auf `alphabet[new_index]`.

# --hints--

Sie sollten eine `new_char`-Variable innerhalb Ihrer `for`-Schleife erstellen.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_char\s*=/));
```

Sie sollten Ihre `new_char`-Variable am Ende Ihres Schleifenkörpers auf `alphabet[new_index]` setzen.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
