---
id: 655251308f31958d06cdf267
title: Schritt 36
challengeType: 20
dashedName: step-36
---

# --description--

Im Moment wird das verschlüsselte Zeichen in jeder Iteration aktualisiert. Es wäre besser, den verschlüsselten String in einer neuen Variable zu speichern. Deklarieren Sie vor Ihrer `for`-Schleife eine Variable namens `encrypted_text` und weisen Sie dieser Variable einen leeren String (`''`) zu.

# --hints--

Sie sollten vor Ihrer `encrypted_text`-Schleife eine Variable namens `for` deklarieren.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

Sie sollten Ihrer `encrypted_text`-Variable einen leeren String zuweisen.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
