---
id: 65521fc818947e800bffe48a
title: Schritt 26
challengeType: 20
dashedName: step-26
---

# --description--

Innerhalb der `for`-Schleife, bevor Sie das aktuelle Zeichen ausgeben, deklarieren Sie eine Variable namens `index` und weisen Sie dieser Variable den Wert zu, den `alphabet.find(char)` zurückgibt.

# --hints--

Sie sollten eine neue Variable namens `index` am Anfang Ihrer `for`-Schleife deklarieren.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

Sie sollten `alphabet.find(char)` Ihrer neuen `index`-Variable zuweisen.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
