---
id: 65520f8282faf57a0db4f7fe
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Wie Sie an der Ausgabe sehen können, befindet sich `'h'` an Index `7` im `alphabet`-String. Nun müssen Sie den Buchstaben an Index `7` plus dem Wert von `shift` finden. Dafür können Sie den Additionsoperator `+` genauso verwenden, wie Sie ihn für eine mathematische Addition nutzen würden.

Ändern Sie Ihre `shifted`-Variable so, dass sie den Wert von `alphabet` an Index `index + shift` speichert.

# --hints--

Sie sollten `alphabet[index + shift]` Ihrer `shifted`-Variable zuweisen.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
