---
id: 65522255d5b9cd80f335c6f2
title: Schritt 29
challengeType: 20
dashedName: step-29
---

# --description--

Deklarieren Sie am Ende Ihres Schleifenkörpers eine Variable namens `new_index` und weisen Sie dieser Variable den Wert von `index + shift` zu.

# --hints--

Sie sollten eine Variable namens `new_index` innerhalb Ihrer `for`-Schleife deklarieren.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

Sie sollten `index + shift` am Ende Ihres `for`-Schleifenkörpers der neuen Variable zuweisen.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
