---
id: 65520e6f2b9678799977f24d
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Entfernen Sie den letzten `print()`-Aufruf. Geben Sie dann statt `text[0]` `text[0].lower()` als Argument an Ihren `.find()`-Aufruf weiter und sehen Sie sich die Ausgabe an.

# --hints--

Sie sollten `print(text.lower())` nicht in Ihrem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

Sie sollten weiterhin Ihre `index`-Variable ausgeben.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

Sie sollten Ihren `alphabet.find(text[0])`-Aufruf aktualisieren, um `text[0].lower()` als Argument zu verwenden. Achten Sie darauf, den Methodenaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
