---
id: 65525e359ca28d938baa82c5
title: Schritt 39
challengeType: 20
dashedName: step-39
---

# --description--

Sie können denselben Effekt wie `a = a + b` erzielen, indem Sie den Additionszuweisungsoperator verwenden:

```py
a += b
```

Der Additionszuweisungsoperator ermöglicht es Ihnen, einen Wert zu einer Variablen hinzuzufügen und das Ergebnis dann dieser Variablen zuzuweisen.

Verwenden Sie den `+=`-Operator, um einen Wert hinzuzufügen und ihn gleichzeitig `encrypted_text` zuzuweisen.

# --hints--

Sie sollten den Additionszuweisungsoperator verwenden, um `alphabet[new_index]` zum aktuellen Wert von `encrypted_text` innerhalb Ihres Schleifenkörpers hinzuzufügen.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
