---
id: 655261b2e1f2c197093f3993
title: Schritt 41
challengeType: 20
dashedName: step-41
---

# --description--

Derzeit werden Leerzeichen als `'c'` verschlüsselt. Um den ursprünglichen Abstand in der Klartextnachricht beizubehalten, benötigen Sie eine bedingte `if`-Anweisung. Diese besteht aus dem Schlüsselwort `if`, einer Bedingung und einem Doppelpunkt `:`.

```py
if x != 0:
    print(x)
```

Im obigen Beispiel ist die Bedingung der `if`-Anweisung `x != 0`. Der Code `print(x)` im Körper der `if`-Anweisung wird nur ausgeführt, wenn die Bedingung zu `True` ausgewertet wird (in diesem Beispiel bedeutet das, dass `x` ungleich null ist).

Ersetzen Sie am Anfang Ihrer for-Schleife `print(char == ' ')` durch eine `if`-Anweisung. Die Bedingung dieser `if`-Anweisung sollte zu `True` ausgewertet werden, wenn `char` ein Leerzeichen ist, und sonst zu `False`. Geben Sie im Körper der `if`-Anweisung den String `'space!'` aus. Denken Sie daran, diese Zeile einzurücken.

# --hints--

Sie sollten `print(char == ' ')` nicht in Ihrem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

Sie sollten `print(char == ' ')` durch eine `if`-Anweisung ersetzen, die ausgelöst wird, wenn `char == ' '`. Verwenden Sie keine runden Klammern, um die `if`-Bedingung einzuschließen, und denken Sie daran, den abschließenden Doppelpunkt einzufügen.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

Sie sollten den String `'space!'` innerhalb Ihrer neuen `if`-Anweisung ausgeben.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
