---
id: 655220953ba90d80514d7ee2
title: Schritt 27
challengeType: 20
dashedName: step-27
---

# --description--

Derzeit nimmt die `print()`-Funktion ein einzelnes Argument `char` entgegen, aber sie kann mehrere Argumente annehmen, die durch ein Komma getrennt sind.

Fügen Sie ein zweites Argument zu `print(char)` hinzu, damit es das Zeichen und seinen Index im Alphabet ausgibt.

# --hints--

Sie sollten `index` als zweites Argument zu Ihrem bestehenden `print(char)`-Aufruf hinzufügen. Vergessen Sie nicht, die Argumente mit einem Komma zu trennen.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
