---
id: 655209a4a27dd37873c4cac3
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` gibt den Index des passenden Zeichens innerhalb des Strings zurück. Wenn das Zeichen nicht gefunden wird, gibt es `-1` zurück. Wie Sie sehen, wird das erste Zeichen in `text`, Großbuchstabe `'H'`, nicht gefunden, da `alphabet` nur Kleinbuchstaben enthält.

Sie können einen String mit der Methode `.lower()` in die Kleinbuchstaben-Entsprechung umwandeln. Fügen Sie einen weiteren `print()`-Aufruf hinzu, um `text.lower()` auszugeben und das Ergebnis zu sehen.

# --hints--

Sie sollten `print(index)` weiterhin in Ihrem Code haben. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Sie sollten `text.lower()` ausgeben. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
