---
id: 6553a44b1801991847d8cc69
title: Schritt 47
challengeType: 20
dashedName: step-47
---

# --description--

Ändern Sie als Nächstes Ihren `print()`-Aufruf so, dass `'encrypted text:', encrypted_text` ausgegeben wird, und platzieren Sie ihn außerhalb der `for`-Schleife, damit der verschlüsselte String nur einmal ausgegeben wird.

# --hints--

Sie sollten `'encrypted text:', encrypted_text` nach Ihrer for-Schleife ausgeben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/m) })
```

In Ihrem Code sollte `print('char:', char, 'encrypted text:', encrypted_text)` nicht enthalten sein.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.notMatch(commentless_code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
