---
id: 6553efd6ada3f42aa2d75448
title: Schritt 53
challengeType: 20
dashedName: step-53
---

# --description--

Jetzt sollten Sie die Ausgabe wieder sehen. Obwohl dieser Ansatz funktioniert, verbessert er die Wiederverwendbarkeit des Codes nicht wesentlich. Wenn Sie Ihre Funktion wiederholt aufrufen, erhalten Sie immer dasselbe Ergebnis. Funktionen können jedoch mit *Parametern* deklariert werden, um Vielseitigkeit und Anpassungsmöglichkeiten einzuführen:

```py
def function_name(param_1, param_2):
    <code>
```

Parameter sind Variablen, die Sie innerhalb Ihrer Funktion verwenden können. Eine Funktion kann mit unterschiedlicher Anzahl von Parametern deklariert werden. Im obigen Beispiel sind `param_1` und `param_2` Parameter.

Ändern Sie Ihre Funktionsdeklaration so, dass sie zwei Parameter namens `message` und `offset` entgegennimmt.

Danach wird im Terminal ein Fehler angezeigt. Wie Sie diesen beheben, sehen Sie in den nächsten Schritten.

# --hints--

Ihre `caesar`-Funktion sollte `message` und `offset` als Parameter annehmen. Denken Sie daran, die Parameter mit einem Komma zu trennen.

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
