---
id: 6553a7d8d05cbb1ae335a665
title: Schritt 45
challengeType: 20
dashedName: step-45
---

# --description--

Wenn die Schleife den Buchstaben `Z` erreicht, überschreitet die Summe `index + shift` den letzten Index des Strings `alphabet`. Daher versucht `alphabet[new_index]`, einen ungültigen Index zu verwenden, was dazu führt, dass ein `IndexError` geworfen wird.

Sie können bemerken, dass die Ausgabe im Terminal direkt vor dem `Z` stoppt, dem letzten `print` vor dem Auftreten des Fehlers.

In diesem Fall kann der Modulo-Operator (`%`) verwendet werden, um den Rest der Division zwischen zwei Zahlen zurückzugeben. Zum Beispiel ist `5 % 2` gleich `1`, weil 5 geteilt durch 2 einen Quotienten von 2 und einen Rest von 1 hat.

Fassen Sie `index + shift` mit runden Klammern ein und wenden Sie Modulo auf den Ausdruck mit `26` an, der die Länge von `alphabet` ist.

# --hints--

Sie sollten `new_index = (index + shift) % 26` in Ihrer `else` Anweisung haben.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
