---
id: 6553f8c570f9982e013a8886
title: Schritt 57
challengeType: 20
dashedName: step-57
---

# --description--

Derzeit wird jeder einzelne Buchstabe immer mit demselben Buchstaben verschlüsselt, abhängig vom angegebenen Versatz. Was wäre, wenn der Versatz für jeden Buchstaben unterschiedlich wäre? Das wäre viel schwieriger zu entschlüsseln. Dieser Algorithmus wird als Vigenère-Chiffre bezeichnet, bei der der Versatz für jeden Buchstaben durch einen anderen Text bestimmt wird, den man *Schlüssel* nennt.

Beginnen Sie damit, Ihre Caesar-Chiffre in eine Vigenère-Chiffre zu verwandeln, indem Sie die beiden Funktionsaufrufe entfernen.

# --hints--

Sie sollten Ihre beiden `caesar()` Funktionsaufrufe entfernen.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
