---
id: 6553ed69ece88d29594748aa
title: Schritt 52
challengeType: 20
dashedName: step-52
---

# --description--

Um eine Funktion auszuführen, muss sie aufgerufen (oder invokiert) werden, indem man ein Paar runder Klammern direkt hinter ihren Namen setzt, so:

```py
function_name()
```

Rufen Sie am Ende Ihres Codes Ihre `caesar`-Funktion auf. Achten Sie auf die Einrückung.

# --hints--

Sie sollten Ihre `caesar`-Funktion aufrufen. Achten Sie darauf, den Funktionsaufruf am Zeilenanfang zu schreiben.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
