---
id: 6553f6086add4b2cbb99fd78
title: Schritt 56
challengeType: 20
dashedName: step-56
---

# --description--

Rufen Sie am Ende Ihres Codes, nach Ihrem bestehenden `caesar(text, shift)`-Aufruf, Ihre Funktion erneut auf.

Übergeben Sie diesmal die Variable `text` und die Ganzzahl `13` als Argumente.

# --hints--

Sie sollten Ihre Funktion erneut aufrufen und diesmal `text` und `13` als Argumente übergeben.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

Sie sollten den bestehenden Funktionsaufruf beibehalten.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
