---
id: 6553f4f66099802c6ae94613
title: Schritt 55
challengeType: 20
dashedName: step-55
---

# --description--

Derzeit löst Ihr Code einen `TypeError` aus, weil die `caesar`-Funktion mit zwei Parametern (`message` und `offset`) definiert ist und daher erwartet, mit zwei *Argumenten* aufgerufen zu werden.

Das Aufrufen von `caesar()` ohne die erforderlichen Argumente stoppt die Ausführung des Codes.

Geben Sie `message` und `offset` Werte, indem Sie `text` und `shift` als Argumente an den `caesar`-Funktionsaufruf übergeben.

# --hints--

Sie sollten `text` und `shift` als Argumente an Ihren Funktionsaufruf übergeben, indem Sie sie in die runden Klammern einschließen. Vergessen Sie nicht, die Argumente mit einem Komma zu trennen.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
