---
id: 6553a755879b131a445e664c
title: Schritt 44
challengeType: 20
dashedName: step-44
---

# --description--

Versuchen Sie, den String `'Hello Zaira'` Ihrer `text`-Variable zuzuweisen, und sehen Sie, was im Terminal passiert.

Sie werden eine `string index out of range`-Ausnahme sehen. Keine Sorge, Sie werden bald herausfinden, wie Sie das beheben!

# --hints--

Sie sollten eine `text`-Variable haben.

```js
assert.match(code, /^text\s*=/m)
```

Sie sollten den String `'Hello Zaira'` Ihrer `text`-Variable zuweisen.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
