---
id: 6553a572f7a65718f1e42e18
title: Schritt 49
challengeType: 20
dashedName: step-49
---

# --description--

Eine Funktion ist im Wesentlichen ein wiederverwendbarer Codeblock. Sie sind bereits einigen eingebauten Funktionen begegnet, wie `print()`, `find()` und `len()`. Aber Sie können auch eigene Funktionen definieren, und zwar so:

```py
def function_name():
    <code>
```

Eine Funktionsdeklaration beginnt mit dem Schlüsselwort `def`, gefolgt vom Funktionsnamen — einem gültigen Variablennamen — und einem Paar runder Klammern. Die Deklaration endet mit einem Doppelpunkt.

Direkt nach Ihrer `shift`-Variablen deklarieren Sie eine Funktion namens `caesar` und rücken alle folgenden Zeilen ein, um Ihrer neuen Funktion einen Rumpf zu geben.

# --hints--

Sie sollten das Schlüsselwort `def` verwenden, um eine neue Funktion zu deklarieren.

```js
assert.match(code, /^def\s+/m)
```

Sie sollten `caesar` als Funktionsnamen nach dem Schlüsselwort `def` schreiben. Denken Sie daran, nach `def` ein Leerzeichen zu setzen.

```js
assert.match(code, /^def\s+caesar/m)
```

Sie sollten ein Paar runder Klammern nach dem Funktionsnamen hinzufügen. Vergessen Sie nicht den abschließenden Doppelpunkt.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

Sie sollten alle Zeilen nach `shift = 3` einrücken, damit sie zum Rumpf Ihrer neuen Funktion werden.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
