---
id: 6553f6b4ec51112d44d737c8
title: Schritt 48
challengeType: 20
dashedName: step-48
---

# --description--

Fügen Sie direkt vor dem `print`-Aufruf einen weiteren hinzu und übergeben Sie `'plain text:', text` als Argumente an `print()`. Verwenden Sie die gleiche Einrückung.

# --hints--

Sie sollten `'plain text:', text` nach Ihrer for-Schleife ausgeben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*^print/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
