---
id: 6553980e0527fa115c705646
title: Schritt 42
challengeType: 20
dashedName: step-42
---

# --description--

Verwenden Sie jetzt statt des Ausgebens von `'space!'` den Additionszuweisungsoperator, um das Leerzeichen (derzeit gespeichert in `char`) zum aktuellen Wert von `encrypted_text` hinzuzufügen.

# --hints--

Sie sollten `print('space!')` nicht in Ihrem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

Sie sollten den `+=`-Operator verwenden, um `char` zum aktuellen Wert von `encrypted_text` innerhalb Ihrer neuen `if`-Anweisung hinzuzufügen.

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
