---
id: 6554d0332949b133a0b35eaa
title: Schritt 68
challengeType: 20
dashedName: step-68
---

# --description--

Rufen Sie Ihre Funktion auf und übergeben Sie `text` und `custom_key` als Argumente. Speichern Sie den Rückgabewert des Funktionsaufrufs in einer Variablen namens `encryption`.

# --hints--

Sie sollten `vigenere` aufrufen und `text` sowie `custom_key` als Argumente übergeben.

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

Sie sollten eine `encryption`-Variable haben.

```js
({ test: () => assert(__userGlobals.has("encryption")) })
```

Sie sollten Ihren Funktionsaufruf der `encryption`-Variablen zuweisen.

```js
({ test: () => assert.match(code, /encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--

--fcc-editable-region--
```
