---
id: 6554a334a40edb1fb4eff827
title: Schritt 62
challengeType: 20
dashedName: step-62
---

# --description--

Deklarieren Sie als Nächstes innerhalb des `else`-Blocks eine Variable namens `key_char` und weisen Sie ihr den Wert von `key` am Index `key_index` mod(`%`) der Länge von `key` zu.

# --hints--

Sie sollten am Anfang Ihres `key_char`-Klauselkörpers eine Variable namens `else` deklarieren.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");
assert(block_body.match(/^\s*key_char\s*=/));
```

Sie sollten `key[key_index % len(key)]` Ihrer Variablen `key_char` zuweisen.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");    
assert(block_body.match(/key_char\s*=\s*key\s*\[\s*key_index\s*%\s*len\s*\(\s*key\s*\)\s*\]/));
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--        
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
--fcc-editable-region--        
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
