---
id: 655491bd5b98b813fa5bedca
title: Schritt 58
challengeType: 20
dashedName: step-58
---

# --description--

Ändern Sie jetzt Ihre Funktionsdeklaration: Ändern Sie den Funktionsnamen in `vigenere` und den zweiten Parameter in `key`.

# --hints--

Sie sollten Ihren Funktionsnamen in `vigenere` ändern.

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

Ihre `vigenere`-Funktion sollte `message` und `key` als Parameter übernehmen.

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
