---
id: 65549561463f0016876e852c
title: Schritt 60
challengeType: 20
dashedName: step-60
---

# --description--

Da Ihr Schlüssel kürzer ist als der Text, den Sie verschlüsseln müssen, müssen Sie ihn wiederholen, um den gesamten verschlüsselten Text zu erzeugen.  
Deklarieren Sie zu Beginn des Funktionskörpers eine `key_index`-Variable und setzen Sie sie auf null.

# --hints--

Sie sollten zu Beginn des Funktionskörpers eine Variable namens `key_index` deklarieren.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

Sie sollten `0` Ihrer `key_index`-Variable zuweisen.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
