---
id: 6554dfce1683be3c0c9609a6
title: Schritt 72
challengeType: 20
dashedName: step-72
---

# --description--

Jetzt können Sie die letzten beiden Zeilen auskommentieren und Ihren Funktionsaufruf so ändern, dass `1` als drittes Argument übergeben wird.

# --hints--

Sie sollten die letzten beiden Zeilen Ihres Codes wiederherstellen und `1` als drittes Argument an Ihren Funktionsaufruf übergeben.

```js
({ test: () => assert.match(code, /^encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*,\s*1\s*\)\s*print\s*\(\s*encryption\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
--fcc-editable-region--
#encryption = vigenere(text, custom_key)
#print(encryption)
--fcc-editable-region--
```
