---
id: 6554a88d5af937226f4a9121
title: Schritt 65
challengeType: 20
dashedName: step-65
---

# --description--

Die `.index()`-Methode ist identisch mit der `.find()`-Methode, aber sie wirft eine `ValueError`-Ausnahme, wenn sie die Teilzeichenkette nicht finden kann.

Eine `ValueError` ist eine eingebaute Ausnahme, die ausgelöst wird, wenn einer Funktion ein Argument mit dem richtigen Typ, aber einem ungeeigneten Wert übergeben wird.

Nachdem Sie `key_index` inkrementiert haben, deklarieren Sie eine Variable namens `offset`. Finden Sie den Index, den `key_char` in `alphabet` hat, und weisen Sie ihn `offset` zu. Verwenden Sie die `.index()`, um den Index zu finden.

# --hints--

Sie sollten eine Variable namens `offset` deklarieren.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

Ihre `offset`-Variable sollte den Wert von `alphabet.index(key_char)` speichern.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
