---
id: 65549f90cf78131c96ebcf28
title: Schritt 61
challengeType: 20
dashedName: step-61
---

# --description--

Beim Programmieren ist Lesbarkeit entscheidend. Kommentare dienen als hilfreiche Notizen, die die Logik hinter Ihrem Code erklären. Sie sind wertvoll, wenn Sie nach einiger Zeit zu einem Projekt zurückkehren, und helfen auch Ihren Kolleginnen und Kollegen, den Code zu verstehen.

In Python können Sie einen Kommentar mit einem `#` schreiben. Alles, was nach dem `#` folgt, wird nicht ausgeführt.

Fügen Sie vor Ihrer `if`-Anweisung einen Kommentar hinzu, der `Append space to the message` sagt.

# --hints--

Sie sollten einen Kommentar hinzufügen, der `Append space to the message` sagt. Vergessen Sie nicht das `#`-Zeichen am Anfang.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
