---
id: 6554a57ec0a2c52106e7ee50
title: Schritt 64
challengeType: 20
dashedName: step-64
---

# --description--

Schreiben Sie innerhalb der `else`-Klausel einen Kommentar mit dem Text `Find the right key character to encode`.

# --hints--

Sie sollten einen Kommentar mit dem Text `Find the right key character to encode` erstellen. Vergessen Sie nicht das `#` am Anfang.

```js
({ test: () => assert.match(code, /#\s*Find\sthe\sright\skey\scharacter\sto\sencode/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
            
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
