---
id: 6554ad2463b8892748f8efdd
title: Schritt 67
challengeType: 20
dashedName: step-67
---

# --description--

Im Moment gibt Ihre Funktion einige Strings aus, aber diese Werte können von anderen Teilen des Codes nicht verwendet werden, um Aktionen auszuführen.

Dafür müssen Sie eine `return`-Anweisung verwenden:

```python
def foo():
    return 'spam'
```

Sie müssen `return` gefolgt von einem Leerzeichen und dem Wert schreiben, den die Funktion zurückgeben soll. Sobald die `return`-Anweisung gefunden wird, wird dieser Wert zurückgegeben und die Ausführung der Funktion stoppt, wobei mit der nächsten Codezeile nach dem Funktionsaufruf fortgefahren wird. Im obigen Beispiel gibt die `foo`-Funktion den String `'spam'` zurück.

Entfernen Sie die beiden `print()`-Aufrufe aus Ihrer Funktion und geben Sie `encrypted_text` zurück.

# --hints--

Sie sollten die beiden `print()`-Aufrufe aus Ihrer Funktion entfernen.

```js
({ 
  test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/print\s*\(\s*("|')plain\stext:\1\s*,\s*message\s*\)/.test(commentless_code))
    assert.isFalse(/print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/.test(commentless_code))
  }
})
```

Ihre Funktion sollte `encrypted_text` zurückgeben.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*/);
    const regex = /return\s+encrypted_text/
    assert(function_body.match(regex));
    assert.notMatch(block_body, regex);
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
