---
id: 6554930320d70414e7b6acc6
title: Schritt 59
challengeType: 20
dashedName: step-59
---

# --description--

Löschen Sie Ihre `shift`-Variable. Deklarieren Sie dann eine neue Variable namens `custom_key` und weisen Sie dieser Variable den Wert des Strings `'python'` zu.

# --hints--

Sie sollten die `shift`-Variable und ihren Wert löschen.

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

Sie sollten eine Variable namens `custom_key` deklarieren.

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

Sie sollten den String `'python'` Ihrer `custom_key`-Variable zuweisen.

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
