---
id: 6554e0adc7bb193cbfdb36d5
title: Schritt 73
challengeType: 20
dashedName: step-73
---

# --description--

Überprüfen Sie, ob die Funktion den String wieder in den Klartext entschlüsseln kann.

Deklarieren Sie eine weitere Variable namens `decryption` und weisen Sie ihr `vigenere(encryption, custom_key, -1)` zu.

# --hints--

Sie sollten `vigenere` aufrufen und `encryption`, `custom_key` sowie `-1` als Argumente übergeben.

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/) })
```

Sie sollten eine `decryption`-Variable deklarieren.

```js
({ test: () => assert(__userGlobals.has("decryption")) })
```

Ihre `decryption`-Variable sollte den Wert von `vigenere(encryption, custom_key, -1)` haben.

```js
({ test: () => assert.match(code, /^decryption\s*=\s*vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--    
encryption = vigenere(text, custom_key, 1)
print(encryption)

--fcc-editable-region--
```
