---
id: 6554ac937a49be2701af4f2f
title: Schritt 66
challengeType: 20
dashedName: step-66
---

# --description--

Über der `offset`-Variable erstellen Sie einen weiteren Kommentar mit dem Text `Define the offset and the encrypted letter`.

# --hints--

Sie sollten einen Kommentar mit dem Text `Define the offset and the encrypted letter` erstellen. Vergessen Sie nicht das `#` am Anfang.

```js
({ test: () => assert.match(code, /#\s*Define\sthe\soffset\sand\sthe\sencrypted\sletter/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
