---
id: 6555d6e12c95701172b55709
title: Schritt 77
challengeType: 20
dashedName: step-77
---

# --description--

Funktionen können mit standardmäßigen Argumenten aufgerufen werden. Ein standardmäßiges Argument gibt den Wert an, den die Funktion annehmen soll, wenn das Argument nicht übergeben wird. Zum Beispiel akzeptiert die untenstehende Funktion drei Argumente, aber Sie können sie mit zwei Argumenten aufrufen. Das dritte nimmt dann standardmäßig den angegebenen Wert an:

```py
def foo(a, b, c=value):
    <code>
```

Ändern Sie die `vigenere`-Funktion so, dass ihr `direction`-Parameter einen Standardwert von `1` hat.

# --hints--

Der `direction`-Parameter Ihrer `vigenere`-Funktion sollte einen Standardwert von `1` haben.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

--fcc-editable-region--
def vigenere(message, key, direction):
--fcc-editable-region--
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
