---
id: 6555201d9b7fc917399f9f0b
title: Schritt 80
challengeType: 20
dashedName: step-80
---

# --description--

Die `.isalpha()`-Methode gibt `True` zurück, wenn alle Zeichen des Strings, auf dem sie aufgerufen wird, Buchstaben sind. Zum Beispiel gibt der folgende Code `True` zurück:

```py
'freeCodeCamp'.isalpha()
# True
```

Löschen Sie die Bedingung `char == ' '` und ersetzen Sie sie durch einen Aufruf der `.isalpha()`-Methode auf `char`.

# --hints--

Sie sollten `if char.isalpha():` in Ihrem Code haben.

```js
({ test: () => assert.match(code, /if\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--    
        # Append space to the message
        if char == ' ':
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
