---
id: 6555e1bfcd374e18c6be8e58
title: Schritt 91
challengeType: 20
dashedName: step-91
---

# --description--

In Python gibt es eine einfache Möglichkeit, Strings zu formatieren. *f-strings* ermöglichen es Ihnen, Werte in Ihre Strings zu interpolieren.

Interpolation bedeutet, Platzhalter zu schreiben, die beim Ausführen des Programms durch die angegebenen Werte ersetzt werden. Zum Beispiel können Sie mit `'Encrypted text: ' + text` dasselbe Ergebnis wie mit `f'Encrypted text: {text}'` erzielen. Sie müssen ein `f` vor die Anführungszeichen setzen, um den f-string zu erstellen, und die Variablen oder Ausdrücke, die Sie interpolieren möchten, zwischen geschweifte Klammern schreiben.

Ändern Sie den ersten `print()`-Aufruf so, dass er einen f-string verwendet.

# --hints--

Sie sollten den `print('Encrypted text: ' + text)`-Aufruf so ändern, dass er dasselbe Ergebnis mit einem f-string ausgibt.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
