---
id: 6555231eac4d9f19bd3d44b5
title: Schritt 83
challengeType: 20
dashedName: step-83
---

# --description--

Das Schlüsselwort `pass` kann als Platzhalter für zukünftigen Code verwendet werden. Es hat keine Auswirkung auf Ihren Code, kann Sie aber vor Fehlern bewahren, die bei unvollständigem Code auftreten würden:

```py
def foo():
    pass
```

Das Aufrufen von `vigenere` mit `1` zum Verschlüsseln und `-1` zum Entschlüsseln ist in Ordnung, aber es könnte etwas kryptisch sein. Erstellen Sie eine neue Funktion namens `encrypt`, die die Parameter `message` und `key` übernimmt, und verwenden Sie `pass`, um den Funktionskörper zu füllen.


# --hints--

Sie sollten eine neue Funktion namens `encrypt` deklarieren. Denken Sie daran, das Schlüsselwort `pass` im Funktionskörper zu verwenden.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

Ihre `encrypt`-Funktion sollte `message` und `key` als Parameter übernehmen. Die Reihenfolge ist wichtig.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
