---
id: 655520c20cb1e6177b0641d6
title: Schritt 81
challengeType: 20
dashedName: step-81
---

# --description--

Der `not`-Operator wird verwendet, um einen Ausdruck zu negieren. Wenn er vor einem truthy-Wert steht – einem Wert, der zu `True` ausgewertet wird – gibt er `False` zurück und umgekehrt.

Fügen Sie den `not`-Operator am Anfang der `if`-Bedingung hinzu, um zu prüfen, ob das Zeichen nicht alphabetisch ist.

# --hints--

Sie sollten den `not`-Operator in der Bedingung Ihrer `if`-Anweisung verwenden.

```js
({ test: () => assert.match(code, /if\s+not\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char.isalpha():
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
