---
id: 6555db4a7b788e15795674e5
title: Schritt 89
challengeType: 20
dashedName: step-89
---

# --description--

Zwei oder mehr Strings können durch Verwendung des `+`-Operators verkettet werden. Zum Beispiel ergibt `'Hello' + ' there!'` `'Hello there!'`.

Rufen Sie die `print()`-Funktion auf und verwenden Sie den `+`-Operator, um die `text`-Variable an den String `'Encrypted text: '` anzuhängen. Achten Sie auf die Leerzeichen.

# --hints--

Sie sollten `print()` aufrufen und `'Encrypted text: ' + text` als Argumente übergeben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^print\s*\(\s*("|')Encrypted\stext:\s\1\s*\+\s*text\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückungen.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--    

--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
