---
id: 655a2a7210094920069b117c
title: Schritt 40
challengeType: 20
dashedName: step-40
---

# --description--

Vergleichsoperatoren ermöglichen es Ihnen, zwei Objekte anhand ihrer Werte zu vergleichen. Sie können einen Vergleichsoperator verwenden, indem Sie ihn zwischen die Objekte setzen, die Sie vergleichen möchten.
Sie geben einen *Boolean*-Wert zurück – nämlich `True` oder `False` – abhängig von der Wahrheit des Ausdrucks.

Python hat die folgenden Vergleichsoperatoren:

<table>
  <thead>
    <tr>
      <th>Operator</th>
      <th>Bedeutung</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Gleich</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Nicht gleich</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Größer als</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Kleiner als</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>Größer oder gleich</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>Kleiner oder gleich</td>
    </tr>
  </tbody>
</table>

Geben Sie zu Beginn Ihres Schleifenkörpers das Ergebnis des Vergleichs von `char` mit einem Leerzeichen (`' '`) aus. Verwenden Sie dafür den Gleichheitsoperator `==`.

# --hints--

Sie sollten `char` mit einem Leerzeichen mit dem Gleichheitsoperator innerhalb Ihrer `for`-Schleife vergleichen.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

Sie sollten das Ergebnis des Vergleichs von `char` mit einem Leerzeichen innerhalb Ihrer `for`-Schleife ausgeben.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

Sie sollten das Ergebnis des Vergleichs von `char` mit einem Leerzeichen zu Beginn Ihrer Schleife ausgeben.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
