---
id: 655a2a7210094920069b117c
title: Schritt 40
challengeType: 20
dashedName: step-40
---

# --description--

Vergleichsoperatoren ermöglichen dir zwei Objekte anhand ihrer Werte zu vergleichen. Du kannst einen Vergleichsoperator verwenden, indem du ihn zwischen die Objekte, die du vergleichen möchtest, platzierst. Sie geben einen *booleschen* Wert zurück — nämlich `True` oder `False` — abhängig vom Zutreffen des Ausdrucks.

In Python gibt es die folgenden Vergleichsoperatoren:

<table>
  <thead>
    <tr>
      <th>Operator</th>
      <th>Bedeutung</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Gleich</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Ungleich</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Größer als</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Kleiner als</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>Größer oder gleich</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>Kleiner oder gleich</td>
    </tr>
  </tbody>
</table>

Am Anfang deines Schleifenkörpers gebe das Ergebnis des Vergleichs von `char` mit einem Leerzeichen (`' '`). Verwende dafür den Gleichheits-Operator `==`.

# --hints--

Du solltest `char` mit einem Leerzeichen vergleichen und dafür den Gleichheits-Operator innerhalb deiner `for` Schleife verwenden.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

Du solltest das Ergebnis des Vergleichs von `char` mit einem Leerzeichen innerhalb deiner `for`Schleife mit print ausgeben.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

Du solltest das Ergebnis des Vergleichs von `char` mit einem Leerzeichen am Anfang der Schleife ausgeben.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
