---
id: 65688a50e6c998a21d8e41d3
title: Schritt 30
challengeType: 20
dashedName: step-30
---

# --description--

Strings sind unveränderlich, das heißt, sie können nach ihrer Erstellung nicht mehr geändert werden. Zum Beispiel könnten Sie denken, dass der folgende Code den Wert von `my_string` in den String `'train'` ändert, aber das ist nicht gültig:

```py
my_string = 'brain'
my_string[0] = 't'
```

Bestätigen Sie, dass Sie mit der Klammernotation auf den ersten Buchstaben in `text` zugreifen und versuchen, ihn in ein Zeichen Ihrer Wahl zu ändern. Sie werden sehen, dass die Ausgabe verschwindet und ein Fehler erscheint.

# --hints--

`text = 'Hello World'` sollte weiterhin in Ihrem Code vorhanden sein.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

Sie sollten mit `text` auf den ersten Buchstaben in `text[0]` zugreifen.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

Sie sollten den Operator `=` verwenden, um ein Zeichen Ihrer Wahl `text[0]` zuzuweisen. Vergessen Sie nicht, das Zeichen entweder in einfache oder doppelte Anführungszeichen zu setzen.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
