---
id: 65688b5a1655a7a6caede847
title: Schritt 31
challengeType: 20
dashedName: step-31
---

# --description--

Wenn Sie versuchen, die einzelnen Zeichen eines Strings zu ändern, wie Sie es im vorherigen Schritt getan haben, erhalten Sie einen `TypeError`, der auftritt, wenn ein Objekt eines ungeeigneten Typs in Ihrem Code verwendet wird.

Wie Sie aus der Fehlermeldung sehen können, unterstützen Strings keine Item-Zuweisung, da sie unveränderlich sind. Eine Variable kann jedoch einem anderen String neu zugewiesen werden:

```py
message = 'Hello World'
message = 'Hello there!'
```

Löschen Sie die `text[0]`-Zeile und weisen Sie `text` den String `'Albatross'` neu zu.

# --hints--

Sie sollten keine `text[0] = 'C'` in Ihrem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

Sie sollten `text` den String `'Albatross'` neu zuweisen. Ändern Sie `text = 'Hello World'` nicht.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
