---
id: 6568b7a40c7ba0ccbc4b4425
title: Schritt 32
challengeType: 20
dashedName: step-32
---

# --description--

Wie Sie sehen, wird jeder Buchstabe des neu zugewiesenen Strings innerhalb der Schleife ausgegeben.

Gehen Sie zum ursprünglichen String zurück, indem Sie die `text` Zuweisung löschen.

# --hints--

Sie sollten `text = 'Albatross'` nicht in Ihrem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

Sie sollten weiterhin eine `text` Variable mit dem Wert `'Hello World'` haben.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
