---
id: 656886ec4e4d9b96ded3c565
title: Schritt 25
challengeType: 20
dashedName: step-25
---

# --description--

Die Iterationsvariable kann jeden gültigen Namen haben, aber es ist praktisch, ihr einen aussagekräftigen Namen zu geben.

Benen Sie Ihre `i`-Variable in `char` um.

# --hints--

Sie sollten `i` durch `char` ersetzen. 

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
