---
id: 65687d2f8c7ee27b0446cef3
title: Schritt 24
challengeType: 20
dashedName: step-24
---

# --description--

Der Code, der bei jeder Iteration ausgeführt wird — platziert nach dem `:` — bildet den Körper der Schleife. Dieser Code muss eingerückt sein. In Python wird empfohlen, pro Einrückungsebene 4 Leerzeichen zu verwenden. Diese eingerückte Ebene ist ein Codeblock.

```py
for i in text:
    print(i)
```

Python verwendet Einrückungen, um Codeblöcke anzuzeigen. Ein Doppelpunkt am Ende einer Zeile signalisiert, dass ein neuer eingerückter Codeblock folgt.

Wenn also nach dem abschließenden Doppelpunkt kein eingerückter Block gefunden wird, stoppt die Codeausführung und ein `IndentationError` wird geworfen. Dieser Code zeigt keine Ausgabe an, sondern löst stattdessen einen `IndentationError` aus:

```py
for i in text:
print(i)
```

Geben Sie Ihrer `for`-Schleife einen Körper, indem Sie einen Aufruf von `print(i)` hinzufügen. Denken Sie daran, den Schleifenkörper einzurücken.

# --hints--

Sie sollten `print(i)` zu Ihrem `for`-Schleifenkörper hinzufügen. Achten Sie auf die Einrückung.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
