---
id: 65951b637ddec63611c482c0
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

You can use the built-in function `print()` to print the output of your code on the terminal.

Functions are reusable code blocks that you can call, or invoke, to run their code when you need them. To call a function, you just need to write a pair of parentheses next to its name. You will learn more about functions very soon.

For now, go to a new line and add an empty call to the `print()` function. You should not see any output yet.

# --hints--

You should have `print()` in your code. Achte darauf, dass der Funktionsaufruf am Anfang der Zeile steht.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

Dein Code enth채lt ung체ltige Syntax und/oder ung체ltige Einr체ckung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
