---
id: 659530ca823d6c47f61daa8a
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Geben Sie Ihre `shifted`-Variable aus.

# --hints--

Sie sollten Ihre `shifted`-Variable ausgeben. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
