---
id: 65951b637ddec63611c482c0
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Sie können die eingebaute Funktion `print()` verwenden, um die Ausgabe Ihres Codes im Terminal auszugeben.

Funktionen sind wiederverwendbare Codeblöcke, die Sie aufrufen können, um ihren Code auszuführen, wenn Sie ihn benötigen. Um eine Funktion aufzurufen, müssen Sie nur ein Paar runder Klammern neben ihrem Namen schreiben. Sie werden sehr bald mehr über Funktionen lernen.

Gehen Sie jetzt auf eine neue Zeile und fügen Sie einen leeren Aufruf der Funktion `print()` hinzu. Sie sollten noch keine Ausgabe sehen.

# --hints--

Sie sollten `print()` in Ihrem Code haben. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
