---
id: 65952b70924edb449faf1950
title: Schritt 19
challengeType: 20
dashedName: step-19
---

# --description--

Deklarieren Sie eine neue Variable mit dem Namen `shifted`. Verwenden Sie die Klammernotation, um auf den Wert von `alphabet` am Index `index` zuzugreifen, und weisen Sie ihn Ihrer neuen Variable zu.

# --hints--

Sie sollten eine Variable mit dem Namen `shifted` deklarieren. Achten Sie darauf, den Variablennamen am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^shifted\s*=/m)
```

Sie sollten `alphabet[index]` Ihrer Variablen `shifted` zuweisen.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
