---
id: 659526f3e5424b3eac4b6aa7
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Geben Sie als Nächstes die Variable `index` in der Konsole aus.

# --hints--

Sie sollten Ihre Variable `index` ausgeben. Achten Sie darauf, den Funktionsaufruf am Anfang der Zeile zu platzieren.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
