---
id: 659db3e3670d3e712be82593
title: Schritt 22
challengeType: 20
dashedName: step-22
---

# --description--

Den Vorgang, für jeden Buchstaben in `text` die Position im Alphabet zu finden und den verschobenen Buchstaben zu bestimmen, zu wiederholen, kann mühsam sein. Zum Glück können Sie das mit einer Schleife vereinfachen.

Entfernen Sie vorerst alle Codezeilen unterhalb der Deklaration der `alphabet`-Variable.

# --hints--

`text = 'Hello World'` sollte weiterhin in Ihrem Code vorhanden sein.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

`shift = 3` sollte weiterhin in Ihrem Code vorhanden sein.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

`alphabet = 'abcdefghijklmnopqrstuvwxyz'` sollte weiterhin in Ihrem Code vorhanden sein.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

Sie sollten die `index`-Variable und ihren Wert löschen.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

`print(index)` sollte nicht in Ihrem Code vorhanden sein.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

Sie sollten die `shifted`-Variable und ihren Wert löschen.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

`print(shifted)` sollte nicht in Ihrem Code vorhanden sein.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

Ihr Code enthält ungültige Syntax und/oder ungültige Einrückungen.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
