---
id: 6624e0ec0bf47017eed84b1f
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Die erste Art von Chiffre, die Sie erstellen werden, heißt *Caesar*-Chiffre. Genauer gesagt nehmen Sie jeden Buchstaben in Ihrer Nachricht, finden seine Position im Alphabet, nehmen den Buchstaben, der 3 Positionen weiter im Alphabet steht, und ersetzen den ursprünglichen Buchstaben durch den neuen Buchstaben.

Um dies zu implementieren, verwenden Sie die in dem vorherigen Schritt besprochene `.find()`-Methode. Ändern Sie Ihren bestehenden `.find()`-Aufruf, indem Sie `text[0]` als Argument übergeben, anstatt `'z'`.

# --hints--

Sie sollten Ihren bestehenden `alphabet.find('z')`-Aufruf ändern, indem Sie `text[0]` an die Methode übergeben.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
