---
id: 666064e915eba7aa1de03f6b
title: Schritt 51
challengeType: 20
dashedName: step-51
---

# --description--

Beheben Sie jetzt den Fehler, indem Sie die Zeile entfernen, die versucht, die Variable `alphabet` außerhalb der Funktion `caesar` auszugeben.

# --hints--

Sie sollten die Zeile `print(alphabet)` entfernen.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--
print(alphabet)
--fcc-editable-region--
```
