---
id: 66605a64a9822f01ec5ae8ce
title: Schritt 50
challengeType: 20
dashedName: step-50
---

# --description--

In Python bestimmt der *Gültigkeitsbereich* einer Variablen, wo auf diese Variable zugegriffen werden kann:

- Variablen, die außerhalb einer Funktion definiert sind, haben einen *globalen* Gültigkeitsbereich und können von jedem Teil Ihres Codes aus aufgerufen werden.

- Variablen, die innerhalb einer Funktion definiert sind, sind *lokal* für diese Funktion und können außerhalb nicht aufgerufen werden.

Um dies in der Praxis zu sehen, versuchen Sie, die `alphabet` Variable am Ende Ihres Codes auszugeben. Dies wird eine `NameError` Ausnahme auslösen.

Sie sollten eine Fehlermeldung sehen, die anzeigt, dass `alphabet` nicht definiert ist. Das liegt daran, dass `alphabet` innerhalb der `caesar` Funktion definiert ist und außerhalb nicht zugänglich ist.

# --hints--

Sie sollten versuchen, die `alphabet` Variable außerhalb der caesar-Funktion auszugeben.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
