---
id: 64dc963d866fff3dd0329993
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

이전 단계에서 선언한 변수를 출력하고, 원판 개수를 변경해 요구되는 최솟값 이동 횟수가 얼마나 빠르게 증가하는지 확인하세요.

# --hints--

`number_of_moves` 변수를 출력해야 합니다.

```js
({ test: () => assert.match(code, /print\s*\(\s*number_of_moves\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move():
    pass
```
