---
id: 64dc922df2919e3c38cead61
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

코드 맨 위에 `NUMBER_OF_DISKS`라는 변수를 선언하여 디스크 개수를 저장하고 `3` 값을 할당하세요. 그런 다음 `range()` 함수에 전달된 첫 번째 인수를 새 변수로 바꾸세요.

# --hints--

`NUMBER_OF_DISKS`라는 변수를 선언해야 합니다.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

`NUMBER_OF_DISKS` 변수는 `3` 값을 가져야 합니다.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

`range()` 함수의 첫 번째 인수 값을 `NUMBER_OF_DISKS` 변수로 바꿔야 합니다.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
