---
id: 64dcd9bbc2268127e7898d77
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

마지막으로, 이동 번호를 3으로 나누었을 때 나머지가 없으면 `'B'`와 `'C'` 사이의 이동이 허용됩니다.

그 조건에 대한 `elif` 문장을 추가하세요. 그런 다음 조건이 충족되면 적절한 문자열을 출력하세요.

# --hints--

`elif`일 때 실행할 `remainder == 0` 문장이 있어야 합니다.

```js
({test: () => assert.match(code, /elif\s+remainder\s*==\s*0\s*:/)})
```

문자열 `f'Move {i + 1} allowed between {auxiliary} and {target}'`를 출력해야 합니다.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*auxiliary\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
