---
id: 64dc916138dbe03bdfe6cec6
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

하노이의 탑 목표는 모든 원판을 마지막 막대기로 옮기는 것입니다. 이를 위해 세 가지 간단한 규칙을 따라야 합니다:

1. 가장 위에 있는 원판만 옮길 수 있습니다
2. 한 번에 한 개의 원판만 옮길 수 있습니다
3. 더 큰 원판을 더 작은 원판 위에 놓을 수 없습니다

기존 코드 아래에 `move`라는 이름의 빈 함수를 선언하세요. 나중에 이 함수를 사용해 원판을 막대기 사이에 옮길 것입니다. 지금은 오류를 피하기 위해 함수 본문 안에 `pass` 키워드를 사용하세요.

# --hints--

`move`라는 이름의 빈 함수를 선언해야 합니다. 함수 본문 안에 올바른 들여쓰기로 `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
