---
id: 64dc985e6720a23edac63b51
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

막대 사이에서 허용되는 디스크 이동은 3번의 이동마다 반복되는 패턴을 보입니다. 예를 들어, 막대 `A`와 막대 `C` 사이의 이동은 첫 번째, 네 번째, 일곱 번째 이동에서 허용되며, 이때 이동 번호를 3으로 나눈 나머지가 1과 같습니다.

이전에 만든 `for` 반복문 안에서 기존 `print()` 호출을 `if`일 때 실행되는 `(i + 1) % 3 == 1` 문장으로 바꾸세요. 이 `if` 문장 안에서 f-문자열을 사용해 `f'Move {i + 1} allowed between {source} and {target}'`을 출력하세요. 참고로 `i + 1`은 이동 번호이며, 첫 번째 반복에서는 `i`가 0입니다.

# --hints--

`if`일 때 실행되어야 하는 `(i + 1) % 3 == 1` 문장이 있어야 합니다.

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

f-문자열을 사용해 `f'Move {i + 1} allowed between {source} and {target}'` 문자열을 출력해야 합니다.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
