---
id: 64dc92a9718fb53ca3d1bad7
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

하노이의 탑 퍼즐은 디스크 개수를 n이라 할 때 2<sup>n</sup> - 1번의 이동으로 해결할 수 있습니다. `number_of_moves`라는 변수를 선언하고 총 이동 횟수를 이 변수에 할당하세요.

Python에서 거듭제곱 연산자는 `**`입니다.

# --hints--

`number_of_moves`라는 변수를 선언해야 합니다.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

`number_of_moves`의 값은 이동 횟수를 계산하는 표현식이어야 합니다.

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
