---
id: 64dcdee07a14f0299ec30559
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

막대 사이에서 허용된 이동을 찾는 코드를 작성했지만, 실제 이동은 양방향일 수 있습니다.
`print()` 호출 후에 `forward`이라는 변수를 선언하고 `False`로 설정하세요. 이 변수를 사용해 디스크를 어느 방향으로 이동해야 하는지 확인할 것입니다.

# --hints--

`forward`라는 변수를 선언해야 합니다.

```js
({ test: () => assert.match(code, /forward\s*=\s*./) })
```

`forward` 변수는 `False`로 설정해야 합니다.

```js
({ test: () => assert.match(code, /forward\s*=\s*False/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
