---
id: 64dc97ab25730b3e9de86ea5
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

이 프로젝트가 끝나면 하노이 탑 퍼즐에 대한 재귀적 해결책을 만들겠지만, 지금은 이 문제에 대한 반복적 접근법을 탐구할 것입니다.

먼저 함수에 `for` 반복문을 추가하여 `number_of_moves`를 순차적으로 반복하고 현재 반복 횟수를 출력하세요.

# --hints--

이동 횟수를 순차적으로 반복하려면 `for` 반복문을 작성하세요. 이를 위해 `range()` 함수를 사용하세요.

```js
({ test: () => {
    const rgs = [
        /for\s+\w+\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

각 반복마다 현재 이동 번호를 출력해야 합니다.

```js
({ test: () => {
    const rgs = [
        /for\s+(\w+)\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:\s+print\s*\(\s*\1\s*\)/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
