---
id: 64dce7ee2494e82bf4205a3f
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

`source`에서 `target`으로 디스크를 반드시 옮겨야 하는 또 다른 경우는, 출발 리스트가 **비어 있지 않고** `source`의 마지막 디스크가 `target`의 마지막 디스크보다 작을 때입니다.

이 조건을 확인하는 `elif` 문장을 추가하세요. 그런 다음, 조건이 충족되면 `forward` 변수를 `True`으로 설정하세요.

# --hints--

`elif`가 비어 있지 않고 `rods[source]`인지 확인하는 `rods[source][-1] < rods[target][-1]` 문장이 있어야 합니다.

```js
const allowedRes = [
    "rods\\s*\\[\\s*source\\s*\\]",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>\\s*0",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>=\\s*1",
];
const re = new RegExp(`elif\\s+((${allowedRes.join(")|(")}))\\s+and\\s+(rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*<\\s*rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]|rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*>\\s*rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\])\\s*:`);
assert.match(code, re);
```

새로운 `forward` 문장 안에서 `True`을 `elif`로 설정해야 합니다. 이전에 할당한 `forward`은 변경하지 마세요.

```js
const allowedRes = [
    "rods\\s*\\[\\s*source\\s*\\]",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>\\s*0",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>=\\s*1",
];
const re = new RegExp(`elif\\s+((${allowedRes.join(")|(")}))\\s+and\\s+(rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*<\\s*rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]|rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*>\\s*rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\])\\s*:\\s+forward\\s*=\\s*True`);
assert.match(code, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
--fcc-editable-region--            
            if not rods[target]:
                forward = True
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
