---
id: 64dcd3d61c448e2676501f43
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

이동 번호를 3으로 나눈 나머지가 2일 때, `'A'`와 `'B'` (출발 막대와 보조 막대) 사이의 이동이 허용됩니다.

그 조건에 대해 `elif` 문장을 추가하세요. 그런 다음 조건이 충족되면 적절한 문자열을 출력하세요.

# --hints--

`elif`일 때 실행할 `remainder == 2` 문장이 있어야 합니다.

```js
({test: () => assert.match(code, /elif\s+remainder\s*==\s*2:/)})
```

문자열 `f'Move {i + 1} allowed between {source} and {auxiliary}'`를 출력해야 합니다.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*auxiliary\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
