---
id: 64dc90e837670a3ba26a1e61
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

이제 타입이 요구된 대로 `list`이 되었으니, `print()` 호출을 제거할 수 있습니다.

# --hints--

코드에서 `print(type(rods['A']))`를 제거하세요.

```js
({ test: () => assert.match(code, /\}(?!\s+print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")A\1\s*\]\s*\)\s*\))/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
