---
id: 64dc8ea01436383a88256d7f
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

퍼즐은 첫 번째 막대에 디스크가 크기 순서대로 쌓여 있으며, 가장 작은 디스크가 위에 있고 가장 큰 디스크가 아래에 있습니다. 첫 번째 리스트에 다양한 디스크 크기를 나타내는 숫자를 채우세요.

첫 번째 리스트에 항목을 수동으로 추가하는 대신, `3` 함수를 사용해 `1`부터 `range()`까지 숫자가 감소하는 수열을 생성하고 `rods['A']`에 할당하세요. 여기서 `3`는 더미 아래쪽에 있는 가장 큰 디스크를 나타내고 `1`는 더미 위쪽에 있는 가장 작은 디스크를 나타냅니다.

구문은 `range(x, y, h)`이며, `x`은 시작 정수(포함), `y`은 마지막 정수(미포함), `h`는 수열에서 한 숫자와 다음 숫자 사이의 차이입니다.

# --hints--

`range()` 함수를 사용해 `rods['A']`에 숫자 수열을 할당하세요. `range()` 함수를 호출하는 구문은 `range(x, y, h)`입니다.

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

`range()` 함수를 사용해 `3`부터 `1`까지의 숫자 수열을 `rods['A']`에 할당하세요.

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
