---
id: 64dc8b1faf9c60304ca8a9b6
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 프로젝트에서는 하노이의 탑으로 알려진 수학 퍼즐을 해결할 것입니다. 이 퍼즐은 세 개의 막대와 서로 다른 직경을 가진 여러 개의 원판으로 구성되어 있습니다.

이 퍼즐의 목표는 특정 규칙을 따라 첫 번째 막대에서 세 번째 막대로 원판을 옮기는 것입니다. 이 규칙은 더 큰 원판을 더 작은 원판 위에 놓는 것을 제한합니다.

막대를 나타내기 위해 `rods`라는 빈 딕셔너리를 만드는 것부터 시작하세요.

# --hints--

`rods`라는 변수를 만들어야 합니다.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

`rods` 변수는 빈 딕셔너리여야 합니다.

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
