---
id: 64dc97005bc7943e2978df0a
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

이제 `print()` 호출을 제거하세요. 그런 다음 `move()` 함수 안에서 `pass` 키워드를 제거하고 `rods` 사전의 내용을 출력하세요.

# --hints--

코드에 `print(number_of_moves)`가 없어야 합니다.

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

`pass` 키워드를 제거하고 `rods` 사전을 출력하세요.

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
