---
id: 64dc976bf864693e668d67e8
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

이제 함수를 호출하고 터미널에서 출력 결과를 확인하세요.

# --hints--

`move()` 함수를 호출해야 합니다.

```js
({test: () => assert.match(code, /^move\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    print(rods)
--fcc-editable-region--

--fcc-editable-region--
```
