---
id: 64dc9004f658183afa85f4f1
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이제 `'A'` 키의 데이터 유형을 `type()` 함수에 전달하여 확인하고 결과를 터미널에 출력하세요.

# --hints--

`'A'` 키를 `type()` 함수에 전달해야 합니다.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

`rods['A']`의 유형을 `print()` 함수를 사용해 출력해야 합니다.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
