---
id: 64dc90949ff85e3b37be40d0
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`range()` 함수는 변경할 수 없는 수열을 반환합니다. 보시다시피 `rods['A']`의 데이터 유형은 `range`이지만, 리스트로 만들고 싶습니다.

그렇게 하려면 `range()` 호출을 `list()` 함수에 전달하세요.

# --hints--

`range()` 호출을 `list()` 함수에 전달해야 합니다.

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
