---
id: 64ddd280d69fa818ca782d50
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

보시다시피, 디스크 1은 세 번 움직일 때마다 왔다 갔다 하고 있습니다. 이는 다른 막대들 사이의 이동도 처리해야 하기 때문입니다.
이전 단계들에서 작성한 코드를 반복해서 막대만 바꾸는 대신, 그 코드를 함수 안으로 옮겨 각 조건문에서 호출하는 것이 더 좋습니다. `make_allowed_move()`라는 빈 함수를 선언하고 `pass` 키워드를 잊지 마세요.

# --hints--

`make_allowed_move`라는 빈 함수를 선언해야 합니다. 함수 본문 안에 올바른 들여쓰기로 `pass` 키워드를 사용하는 것을 기억하세요.

```js
({ test: () => assert(runPython(`
      import inspect
      inspect.isfunction(make_allowed_move)
  `)) 
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if not rods[target]:
                forward = True
            elif rods[source] and rods[source][-1] < rods[target][-1]:
                forward = True              
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
            else:
                print(f'Moving disk {rods[target][-1]} from {target} to {source}')
                rods[source].append(rods[target].pop())
            
            # display our progress
            print(rods)
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
