---
id: 64ddcc2ee182e0164c3539f8
title: 26단계
challengeType: 20
dashedName: step-26
---

# --description--

이동을 출력한 후에는 원본 막대에서 마지막 요소를 제거하고 대상 막대에 추가해야 합니다. 이를 위해 `.pop()` 메서드와 `.append()` 메서드를 사용하세요.

# --hints--

`rods[source]`에서 마지막 요소를 제거해야 합니다.

```js
({ test: () => assert.match(code, /rods\s*\[\s*source\s*\]\s*\.pop\s*\(\s*\)/) })
```

코드에 `rods[target].append(rods[source].pop())`가 있어야 합니다.

```js
({ test: () => assert.match(code, /rods\s*\[\s*target\s*\]\s*\.append\s*\(\s*rods\s*\[\s*source\s*\]\s*\.pop\s*\(\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if not rods[target]:
                forward = True
            elif rods[source] and rods[source][-1] < rods[target][-1]:
                forward = True
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
