---
id: 64de773f81facd14653f49c8
title: 36단계
challengeType: 20
dashedName: step-36
---

# --description--

잘 작동하는 것 같아요! 하지만 출력 결과가 그리 읽기 쉽지 않습니다. 막대를 출력한 후에 새 줄 문자를 출력해서 이 문제를 고치세요.

# --hints--

두 개의 `'\n'` 호출에서 두 번째 인수로 새 줄 문자(`print(rods)`)가 포함된 문자열을 전달해야 합니다.

```js
({ test: () => assert.match(code, /print\(\s*rods\s*,\s*('|")\\n\1\s*\)(?=.*print\(\s*rods\s*,\s*('|")\\n\2\s*\))/s) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def make_allowed_move(rod1, rod2):    
    forward = False
    if not rods[rod2]:
        forward = True
    elif rods[rod1] and rods[rod1][-1] < rods[rod2][-1]:
        forward = True
                      
    if forward:
        print(f'Moving disk {rods[rod1][-1]} from {rod1} to {rod2}')
        rods[rod2].append(rods[rod1].pop())
    else:
        print(f'Moving disk {rods[rod2][-1]} from {rod2} to {rod1}')
        rods[rod1].append(rods[rod2].pop())
--fcc-editable-region--        
    # display our progress
    print(rods)

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
--fcc-editable-region--    
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            make_allowed_move(source, target)
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
            make_allowed_move(source, auxiliary)
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')
            make_allowed_move(auxiliary, target)

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
