---
id: 64de1469fccce61940f498c4
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

`rods` 사전은 원판이 있는 세 개의 막대를 나타낼 것입니다. `'A'`, `'B'`, `'C'` 문자열을 키로 주고 각각을 빈 리스트로 설정하세요.

# --hints--

`rods` 사전에는 `'A'` 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']`는 빈 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

`rods` 사전에는 `'B'` 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']`는 빈 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

`rods` 사전에는 `'C'` 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']`는 빈 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
