---
id: 64de7be06eb689161dd63cf0
title: 38단계
challengeType: 20
dashedName: step-38
---

# --description--

이전에 작성한 조건문은 디스크 개수가 홀수일 때만 유효합니다.

`if`가 홀수일 때 실행할 중첩된 `n`를 추가하고, `print()`와 `make_allowed_move()` 호출에 들여쓰기 한 단계씩 추가하세요.

# --hints--

`if`가 홀수일 때 실행할 `n` 문을 중첩하세요. `print()`와 `make_allowed_move()` 호출에 들여쓰기를 잊지 마세요.

```js
({ test: () => assert.match(code, /if\s+n\s*%\s*2\s*(!=\s*0|==\s*1)\s*:\s+print.*\s+make_allowed_move\(.*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def make_allowed_move(rod1, rod2):    
    forward = False
    if not rods[rod2]:
        forward = True
    elif rods[rod1] and rods[rod1][-1] < rods[rod2][-1]:
        forward = True
                      
    if forward:
        print(f'Moving disk {rods[rod1][-1]} from {rod1} to {rod2}')
        rods[rod2].append(rods[rod1].pop())
    else:
        print(f'Moving disk {rods[rod2][-1]} from {rod2} to {rod1}')
        rods[rod1].append(rods[rod2].pop())
    
    # display our progress
    print(rods, '\n')

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--     
        if remainder == 1:
            
            print(f'Move {i + 1} allowed between {source} and {target}')
            make_allowed_move(source, target)
--fcc-editable-region--            
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
            make_allowed_move(source, auxiliary)
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')
            make_allowed_move(auxiliary, target)
           
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
