---
id: 64df45a3ad4f8719e5355244
title: 46단계
challengeType: 20
dashedName: step-46
---

# --description--

`n - 1` 원반을 옮기는 단계를 더 세분화하면 단일 원반만 고려하는 단계가 됩니다. 이것이 첫 번째 이동이 될 것입니다. 첫 번째 이동이 발생한 후에는 재귀 호출이 풀리면서 다음 이동들이 생성됩니다. 각 재귀 단계에서 막대가 차지하는 역할이 출발지, 도착지, 보조지 사이에서 바뀐다는 점을 기억하세요.

지금은 각 재귀 호출이 리스트에 아무 변경도 가하지 않고 `rods` 사전을 출력합니다. `print()` 호출 전에 `rods[source]` 리스트에서 마지막 요소를 제거하고 `rods[target]` 리스트에 추가하세요.

# --hints--

`rods[source]` 호출 전에 `rods[target]` 리스트에서 마지막 요소를 제거하고 `print` 리스트에 추가해야 합니다.

```js
({ test: () => assert.isTrue(runPython(`
_log = []
def capture_print(*args):
    _log.append(repr(args))

old_print = print
print = capture_print

old_rods = rods
rods2 = {
    'A': list(range(5, 0, -1)),
    'B': [],
    'C': []
}
rods = rods2

move(5, 'A', 'B', 'C')
rods = old_rods
print = old_print

_expected_prints = [
r"({'A': [5, 4, 3, 2], 'B': [], 'C': [1]}, '\\n')",
r"({'A': [5, 4, 3], 'B': [], 'C': [1, 2]}, '\\n')",
r"({'A': [5, 4], 'B': [], 'C': [1, 2, 3]}, '\\n')",
r"({'A': [5], 'B': [], 'C': [1, 2, 3, 4]}, '\\n')",
r"({'A': [], 'B': [], 'C': [1, 2, 3, 4, 5]}, '\\n')",
]

for args in _log:
    if not _expected_prints:
        break
    if args == _expected_prints[0]:
        _expected_prints.pop(0)

not rods2['A'] and rods2['C'] == list(range(1, 6)) and not _expected_prints
`))})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # display starting configuration
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
