---
id: 64df496c6a8ddf1b38db1ed6
title: 50단계
challengeType: 20
dashedName: step-50
---

# --description--

이전 단계에서, 하위 문제의 가장 큰 원판을 목표 막대기로 옮기는 코드를 작성했었습니다.

이제 해야 할 일은 이미 옮긴 `n - 1`개의 원판을 이동하기 위해 또 다른 재귀 호출을 추가하는 것입니다. 첫 번째 재귀 호출을 복사해서 `if` 블록의 끝에 붙여넣으세요.

함수 인수의 순서가 올바르지 않다는 점에 주의하세요. 올바른 순서를 찾아보세요.

# --hints--

`move(n - 1,  auxiliary, source, target)`를 `move` 함수의 끝에서 호출해야 합니다.

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
