---
id: 64df3f1011888113fbd3d81b
title: 44단계
challengeType: 20
dashedName: step-44
---

# --description--

재귀로 퍼즐을 해결하려면, 먼저 원래 문제를 더 작은 하위 문제로 나누는 것이 중요합니다.

`n`개의 원판을 세 번째 막대에 크기 순서대로 쌓는 최종 구성을 얻으려면 다음과 같이 이동해야 합니다:

- `n - 1`개의 원판을 출발지에서 보조 막대로 옮기고
- 가장 큰 원판을 출발지에서 목표 막대로 옮기고
- 그리고 `n - 1`개의 원판을 보조 막대에서 목표 막대로 옮깁니다.

따라서 `move` 함수가 가장 먼저 해야 할 일은 `n - 1`를 첫 번째 인수로 하여 자기 자신을 호출하는 것입니다. 하지만 기본 사례를 정의하지 않고 이렇게 하면 `RecursionError`가 발생합니다. 이는 함수가 무한히 자기 자신을 호출하기 때문입니다.

주석과 `print()` 호출 전에, `n - 1`를 첫 번째 인수로 하여 재귀 함수 호출을 추가하고 함수 본문이 `n`가 0보다 클 때만 실행되도록 하세요. 지금은 다른 인수들은 같은 순서로 두세요.

# --hints--

`move` 함수 본문은 `if`가 0보다 클 때 실행되는 `n` 문으로 시작해야 합니다.

```js
({ test: () => assert.match(code, /def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:\s+if\s+n\s*>\s*0\s*:/) })
```

`move(n - 1, source, auxiliary, target)` 문 안에서 `if`를 호출해야 합니다. `print()` 호출은 들여쓰기 하세요.

```js
const transformedCode = "\n" + code.replace(/\r/g, '');
const move = __helpers.python.getDef(transformedCode, "move");
const { function_parameters, function_body } = move;

assert.match(function_parameters, /n\s*,\s*source\s*,\s*auxiliary\s*,\s*target/);
const functionBodyIndent = function_body.match(/ +/)[0];
const re = new RegExp(`${functionBodyIndent}if\\s+n\\s*>\\s*0\\s*:\\s+^${functionBodyIndent}( +)move\\(\\s*n\\s*-\\s*1\\s*,\\s*source\\s*,\\s*auxiliary\\s*,\\s*target\\s*\\).+?^${functionBodyIndent}\\1print\\s*\\(\\s*rods\\s*,\\s*("|')\\\\n\\2\\s*\\)`, "ms");
assert.match(function_body, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
