---
id: 64df47b32b92301a815d5ef8
title: 49단계
challengeType: 20
dashedName: step-49
---

# --description--

처음에, 방금 추가한 재귀 호출은 `n - 1`개의 원반을 두 번째 막대로 옮기는 하위 문제를 다룹니다.

그 이유로, `target` 인수는 두 번째 막대에 해당하고, `auxiliary` 인수는 세 번째 막대에 해당합니다. 재귀가 진행됨에 따라 이 값들은 계속 바뀔 것임을 기억하세요.

재귀 호출에서 `target`과 `auxiliary`의 인수 순서를 바꿔 고치세요.

# --hints--

`move(n - 1, source, auxiliary, target)` 호출에서 인수 순서를 수정해야 합니다.

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
