---
id: 64e33a488b4a2717fa22ebf3
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

`print()` 호출 전에 `display starting configuration`라고 주석을 하나 더 추가하세요.

# --hints--

제공된 텍스트를 사용해 `print()` 호출 전에 주석을 추가하세요.

```js
({ test: () => assert.match(code, /#\s*display\sstarting\sconfiguration\s+print\s*\(\s*rods\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    
    print(rods)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
