---
id: 64e34146860065146733883b
title: 53단계
challengeType: 20
dashedName: step-53
---

# --description--

재귀는 때때로 이해하기 어려울 수 있지만, 더 간결한 코드를 작성할 수 있는 힘을 줍니다. 이 경우에는 원판의 짝수와 홀수를 구분할 필요조차 없습니다.

`NUMBER_OF_DISKS`를 `5`로 설정하고 출력 결과를 확인하세요.

# --hints--

`NUMBER_OF_DISKS` 변수를 5로 설정하세요.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 5) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
A = list(range(NUMBER_OF_DISKS, 0, -1))
B = []
C = []

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        target.append(source.pop())
        
        # display our progress
        print(A, B, C, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, A, B, C)
--fcc-editable-region--
```
