---
id: 650c6082e5586f9e3acfcd3b
title: 52단계
challengeType: 20
dashedName: step-52
---

# --description--

이제 `rods` 사전을 삭제하고 그 키들을 각각 `A`, `B`, `C` 변수로 바꾸되, 값은 그대로 유지하세요. 코드를 이 변경 사항에 맞게 리팩터링하세요. 터미널에 출력이 보인다면 이 단계를 올바르게 수행한 것입니다.

# --hints--

`rods['A']`를 `A`라는 변수로 바꾸되 값을 유지하세요.

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

`rods['B']`를 `B`라는 변수로 바꾸되 값을 유지하세요.

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

`rods['C']`를 `C`라는 변수로 바꾸되 값을 유지하세요.

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

`print()` 호출을 수정하여 `A` 객체 대신 `B`, `C`, `rods`을 출력하세요. `print()` 호출에서 줄바꿈 문자는 유지하세요.

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

코드에 `target.append(source.pop())`가 있어야 합니다.

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```
