---
id: 65782342850feb3b8d62f936
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

`(i + 1) % 3` 표현식을 여러 번 사용할 예정이므로, 변수를 만들어 저장하는 것이 편리합니다.

`if` 문장 바로 위에 `remainder` 변수를 선언하고 이 변수에 `(i + 1) % 3` 값을 할당하세요.

# --hints--

`remainder` 블록 바로 위에 `if` 변수를 선언해야 합니다.

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=/m) })
```

`remainder` 변수의 값은 `(i + 1) % 3` 표현식이어야 합니다.

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=\s*\(\s*i\s*\+\s*1\s*\)\s*%\s*3/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
