---
id: 657b667a772ed53e82962c81
title: 54단계
challengeType: 20
dashedName: step-54
---

# --description--

코드 가독성을 개선할 수 있는 한 가지 방법이 더 있습니다.

`if`이 0 이하일 때 `n`이 실행되도록 수정하고, 함수 실행을 중단하는 `return` 문을 추가하세요.

# --hints--

기존 `if`을 수정하여 `n`가 `0` 이하인지 확인해야 합니다.

```js
({ test: () =>
  {
    const hanoi = __helpers.python.getDef(code, "move");
    const {function_body} = hanoi;    
    assert(function_body.match(/^\s{4}if\s+n\s*<=\s*0\s*:/m));
  }
})
```

`return` 블록의 첫 줄에 `if` 문을 추가해야 합니다.

```js
({ test: () =>
  {
    const hanoi = __helpers.python.getDef(code, "move");
    const {function_body} = hanoi;    
    assert(function_body.match(/^(\s{4})if\s+n\s*<=\s*0\s*:\s*\1\1return/m));
  }
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 5
A = list(range(NUMBER_OF_DISKS, 0, -1))
B = []
C = []
--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:

--fcc-editable-region--        
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        target.append(source.pop())
        
        # display our progress
        print(A, B, C, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, A, B, C)

```
