---
id: 657b6a4a49faec5c600287ba
title: 55단계
challengeType: 20
dashedName: step-55
---

# --description--

마지막 단계로, `return` 문장 이후 모든 코드의 들여쓰기 수준을 줄이세요.

잘하셨습니다. Tower of Hanoi 연습 프로젝트를 완료했습니다.

# --hints--

`return` 문장 이후 모든 코드의 들여쓰기 수준을 줄이세요.

```js
({ test: () => assert(runPython(`
    a, b, c = [3, 2, 1], [], []
    move(3, a, b, c)
    a == [] and b == [] and c == [3, 2, 1]
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 5
A = list(range(NUMBER_OF_DISKS, 0, -1))
B = []
C = []

def move(n, source, auxiliary, target):
    if n <= 0:
        return
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        target.append(source.pop())
        
        # display our progress
        print(A, B, C, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, A, B, C)
--fcc-editable-region--
```

# --solutions--

```py
NUMBER_OF_DISKS = 5
A = list(range(NUMBER_OF_DISKS, 0, -1))
B = []
C = []

def move(n, source, auxiliary, target):
    if n <= 0:
        return
    # move n - 1 disks from source to auxiliary, so they are out of the way
    move(n - 1, source, target, auxiliary)
        
    # move the nth disk from source to target
    target.append(source.pop())
        
    # display our progress
    print(A, B, C, '\n')
        
    # move the n - 1 disks that we left on auxiliary onto target
    move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, A, B, C)
```
