---
id: 66ed41f912d0bb1dc62da5dd
title: Erstellen Sie einen JavaScript-Trivia-Bot
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten `"Hello! I'm your coding fun fact guide!"` als Begrüßungsnachricht für den Benutzer in der Konsole ausgeben.
1. Sie sollten drei Variablen erstellen: `botName`, `botLocation` und `favoriteLanguage`, die den Namen des Bots, seinen Herkunftsort und seine bevorzugte Programmiersprache speichern.
1. Sie sollten `"My name is (botName) and I live on (botLocation)."` in der Konsole ausgeben.
1. Sie sollten `"My favorite programming language is (favoriteLanguage)."` in der Konsole ausgeben.
1. Sie sollten `let` verwenden, um eine `codingFact`-Variable zu erstellen und ihr einen String zuzuweisen, der eine interessante Tatsache über die Lieblingsprogrammiersprache Ihres Bots enthält und die Verwendung der `favoriteLanguage`-Variable einschließt.
1. Sie sollten `codingFact` in der Konsole ausgeben.
1. Sie sollten die `codingFact`-Variable mit einer neuen Tatsache über die Lieblingssprache des Bots unter Verwendung der `favoriteLanguage`-Variable erneut zuweisen.
1. Sie sollten `codingFact` erneut in der Konsole ausgeben.
1. Sie sollten die `codingFact`-Variable erneut einer weiteren neuen Tatsache über die Lieblingssprache des Bots unter Verwendung der `favoriteLanguage`-Variable zuweisen.
1. Sie sollten `codingFact` zum dritten Mal in der Konsole ausgeben.
1. Sie sollten `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` als Abschiedsaussage des Bots in der Konsole ausgeben.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten `"Hello! I'm your coding fun fact guide!"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

Sie sollten eine `botName`-Variable deklarieren. Überprüfen Sie sorgfältig auf Rechtschreib- oder Groß-/Kleinschreibfehler.

```js
assert.exists(botName);
```

Ihre `botName`-Variable sollte ein String sein.

```js
assert.isString(botName);
```

Sie sollten eine `botLocation`-Variable deklarieren. Überprüfen Sie sorgfältig auf Rechtschreib- oder Groß-/Kleinschreibfehler.

```js
assert.exists(botLocation);
```

Ihre `botLocation`-Variable sollte ein String sein.

```js
assert.isString(botLocation);
```

Sie sollten eine `favoriteLanguage`-Variable deklarieren. Überprüfen Sie sorgfältig auf Rechtschreib- oder Groß-/Kleinschreibfehler.

```js
assert.exists(favoriteLanguage);
```

Ihre `favoriteLanguage`-Variable sollte ein String sein.

```js
assert.isString(favoriteLanguage);
```

Sie sollten `"My name is (botName) and I live on (botLocation)."` in der Konsole ausgeben und die Variablen in den String einfügen.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

Sie sollten `"My favorite programming language is (favoriteLanguage)."` in der Konsole ausgeben und die Variable in den String einfügen.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

Sie sollten `let` verwenden, um eine neue Variable `codingFact` zu deklarieren.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

Sie sollten `codingFact` einen Wert zuweisen, der `favoriteLanguage` enthält.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

Sie sollten `codingFact` in der Konsole ausgeben.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

Sie sollten `codingFact` einen neuen Wert zuweisen, der ebenfalls `favoriteLanguage` enthält, und ihn in der Konsole ausgeben.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

Sie sollten `codingFact` zum dritten Mal einen Wert zuweisen, der ebenfalls `favoriteLanguage` enthält, und ihn in der Konsole ausgeben.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

Sie sollten `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` in der Konsole ausgeben und die Werte der Variablen hinzufügen.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
