---
id: a97fd23d9b809dac9921074f
title: 建置一個選用引數加總函式
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

在這個實驗中，你將建置一個函式（程式），該函式最多接受兩個引數，並將它們相加，但如果只有一個引數，則傳回一個等待第二個數字以進行相加的函式（程式）。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 創建一個函式（程式）`addTogether`。
2. 如果函式接收兩個引數，函式應該傳回這兩個引數的總和。
3. 如果只提供一個引數，`addTogether` 應該傳回一個函式。
   - 當以單一引數呼叫傳回的函式時，應該傳回總和。

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. 如果任一引數不是有效的數字，則傳回 `undefined`。

# --hints--

你應該擁有 `addTogether` 函式（程式）。

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` 應該傳回 5。

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` 應該傳回 53.4。

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` 應該傳回一個函式（程式）。

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` 應該傳回 12。

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` 應該傳回 `undefined`。

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
