---
id: 68ce56d54d3cad1149a1c4cf
title: Tengeneza ukurasa wa kuchagua kipengele
challengeType: 25
dashedName: design-a-feature-selection-page
demoType: onClick
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Ukurasa wako unapaswa kuwa na kipengele cha `h1` chenye maandishi ya `Feature Selection`.
2. Ukurasa wako unapaswa kuwa na kipengele cha `div` chenye darasa la vitu la `feature-card-container`.
3. Ukurasa wako unapaswa kuwa na angalau vipengele viwili vya `label` kila kikiwa na darasa la vitu la `feature-card` ndani ya `div`.
4. Kila kipengele chako cha `label` kinapaswa kuwa na maandishi ya lebo na kipengele cha `input` cha aina ya `checkbox` ndani yake.
5. Tengeneza kichaguzi kinacholenga visanduku vya kuchagua, na tumia mtindo ufuatao:
    - Visanduku vyote vya kuchagua vinapaswa kuwekwa kuwa `appearance: none;` katika CSS yako.
    - Visanduku vyote vya kuchagua vinapaswa kuwa na upana wa mpaka, rangi na mtindo wa mpaka unayochagua.
6. Wakati kisanduku cha kuchagua kimechaguliwa, kinapaswa kuonyesha alama ya kuangalia `✓` katikati ya kisanduku.
7. Wakati kisanduku cha kuchagua kimechaguliwa, rangi ya nyuma ya kisanduku inapaswa kubadilika kuwa rangi unayochagua.
8. Wakati kisanduku cha kuchagua kimechaguliwa, rangi ya mpaka wa kisanduku inapaswa kubadilika kuwa rangi unayochagua.

**Kumbuka:** Hakikisha umeunganisha faili lako la CSS katika HTML yako.

# --hints--

Unapaswa kuwa na kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1'));
```

`h1` yako inapaswa kuwa na maandishi ya `Feature Selection`.

```js
assert.equal(document.querySelector('h1').innerText.trim(), `Feature Selection`);
```

Unapaswa kuwa na kipengele cha `div` chenye darasa la vitu la `feature-card-container`.

```js
assert.exists(document.querySelector('div.feature-card-container'));
```

Unapaswa kuwa na vipengele vya angalau viwili vya `label` ndani ya `.feature-card-container`.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isAtLeast(labels.length, 2);
```

Kila kipengele cha `label` ndani ya `.feature-card-container` kinapaswa kuwa na darasa la vitu la `feature-card`.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isTrue(label.classList.contains('feature-card')));
```

Ndani ya kila kipengele cha `label` unapaswa kuwa na maandishi ya lebo.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isNotEmpty(label.innerText.trim()));
```

Ndani ya kila kipengele cha lebo kunapaswa kuwa na kipengele kimoja cha `input` cha aina ya `checkbox`.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.exists(label.querySelector('input[type="checkbox"]')));
```

Unapaswa kuwa na kichaguzi cha visanduku vya kuchagua.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Visanduku vyako vya kuchagua vinapaswa kuwekwa kuwa `appearance: none`.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.equal(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('appearance'), 'none');
```

Visanduku vyako vya kuchagua vinapaswa kuwa na upana wa mpaka unayochagua.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-width'));
```

Visanduku vyako vya kuchagua vinapaswa kuwa na rangi ya mpaka unayochagua.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-color'));
```

Visanduku vyako vya kuchagua vinapaswa kuwa na mtindo wa mpaka unayochagua.

```js
const selectors = [  
  '.feature-card [type="checkbox"]',
  '.feature-card input[type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input', // there aren't other input elements, so this is a valid approach
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-style'));
```

Wakati visanduku vya kuchagua vimechaguliwa, inapaswa kuwepo alama ya kuangalia `✓` katikati ya kisanduku.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;

    checkbox.checked = true;
    const after = window.getComputedStyle(checkbox, ':after');
    assert.equal(after.content, '"✓"');
});
```

Wakati kisanduku cha kuchagua kimechaguliwa, rangi ya nyuma ya kisanduku inapaswa kubadilika kuwa rangi unayochagua.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBG = window.getComputedStyle(checkbox).backgroundColor;
    checkbox.checked = true;
    const checkedBG = window.getComputedStyle(checkbox).backgroundColor;

    assert.notEqual(checkedBG, uncheckedBG);
});
```

Wakati kisanduku cha kuchagua kimechaguliwa, rangi ya mpaka wa kisanduku inapaswa kubadilika kuwa rangi unayochagua.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBC = window.getComputedStyle(checkbox).borderColor;
    checkbox.checked = true;
    const checkedBC = window.getComputedStyle(checkbox).borderColor;

    assert.notEqual(checkedBC, uncheckedBC);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Selection Feature Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Feature Selection</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Feature Selection</h1>
    <div class="feature-card-container">
      <label class="feature-card">
        <input type="checkbox" />
        <h4>Cloud Storage</h4>
        <p>100 Gigabyte secure storage</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Dedicated Support</h4>
        <p>24/7 customer help</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Advanced Analytics</h4>
        <p>Insights & reports</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Custom User Themes</h4>
        <p>Personalized dashboard design</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Multi-User Collab</h4>
        <p>Team access and sharing</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>API Access</h4>
        <p>Integrate with your custom built tools</p>
      </label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  flex-direction: column;
  height: 100vh;
}

.feature-card-container {
  display: flex;
  flex-wrap: wrap;
  gap: 20px;
  max-width: 600px;
  justify-content: center;
}

.feature-card-container > * {
  flex: 1 1 180px;
}

.feature-card {
  border: 2px solid #ccc;
  border-radius: 12px;
  padding: 20px;
  text-align: center;
  position: relative;
  cursor: pointer;
}

.feature-card:hover {
  border-color: #f1be32;
  box-shadow: 0 4px 12px rgba(255, 201, 53, 0.6);
}

.feature-card input[type='checkbox'] {
  position: absolute;
  top: 5px;
  right: 5px;
  width: 24px;
  height: 24px;
  cursor: pointer;
  appearance: none;
  -webkit-appearance: none;
  border: 2px solid #ccc;
  border-radius: 4px;
  background-color: white;
}

.feature-card input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #f1be32;
}

.feature-card input[type='checkbox']:checked::after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  line-height: 20px;
}

.feature-card h4 {
  margin: 10px 0 5px;
  font-size: 1.2rem;
}

@media (max-width: 600px) {
  body {
    height: auto;
  }
}
```
