---
id: 66a83bdcf425e7446900b7c4
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이전 수업에서 `form` 요소를 다루는 방법을 배웠습니다:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

`action` 속성은 폼이 제출될 때 폼 데이터를 어디로 보낼지 지정하는 데 사용됩니다.

`method` 속성은 폼 데이터를 보낼 때 사용할 HTTP 메서드를 지정하는 데 사용됩니다. 가장 일반적인 메서드는 `GET`와 `POST`입니다.

**참고**: HTTP 메서드가 어떻게 작동하는지는 이후 모듈에서 배우게 됩니다.

`main` 요소 안에 `form` 요소를 추가하고, `action` 속성을 `"https://hotel-feedback.freecodecamp.org"`로, `method` 속성을 `"POST"`으로 설정하세요.

# --hints--

열린 `form` 태그가 있어야 합니다.

```js
assert.match(code, /<form.*>/i);
```

닫힌 `form` 태그가 있어야 합니다.

```js
assert.match(code, /<\/form>/i);
```

`form` 요소에 `action` 속성이 `"https://hotel-feedback.freecodecamp.org"`로 설정되어 있어야 합니다.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

`form` 요소에 `method` 속성이 `"POST"`로 설정되어 있어야 합니다.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
