---
id: 66a83601cd819e37f0dccd14
title: 1단계
challengeType: 0
dashedName: step-1
---

# --description--

이 워크숍에서는 호텔 피드백 폼을 만들어 HTML 폼 작업을 연습합니다.

소개 텍스트로는 메인 제목과 피드백을 남기는 것에 대한 간단한 안내 문구를 표시하세요.

`body` 요소 안에 `header` 요소를 추가하세요.

`header` 요소 안에 `h1` 요소를 추가하세요. `h1` 요소의 텍스트는 `Hotel Feedback Form` 여야 합니다.

`h1` 요소 아래에 `p` 요소를 추가하세요. `p` 요소의 텍스트는 `Thank you for staying with us. Please provide feedback on your recent stay.` 여야 합니다.

# --hints--

`header` 여는 태그가 있어야 합니다.

```js
assert.match(code, /<header>/i);
```

`header` 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/header>/i);
```

`h1` 여는 태그가 있어야 합니다.

```js
assert.match(code, /<h1>/i);
```

`h1` 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/h1>/i);
```

`h1` 요소의 텍스트는 `Hotel Feedback Form` 여야 합니다.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

`h1` 요소는 `header` 요소 안에 있어야 합니다.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

`p` 여는 태그가 있어야 합니다.

```js
assert.match(code, /<p>/i);
```

`p` 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/p>/i);
```

단락 요소의 텍스트는 `Thank you for staying with us. Please provide feedback on your recent stay.` 여야 합니다.

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

단락 요소는 `header` 안에 있어야 합니다.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

단락 요소는 `h1` 다음에 와야 합니다.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
