---
id: 66a84111965a0c46df6bbd0a
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

이전 수업에서 `label` 요소를 `input`와 이렇게 연결하는 방법을 배웠습니다:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

`for` 요소의 `label` 속성은 `id` 요소의 `input` 속성과 일치해야 합니다. 이것을 명시적 연결이라고 합니다.

`legend` 요소 아래에 `label` 요소를 추가하고 `Name (required):` 텍스트를 넣으세요. `for` 속성은 `"full-name"` 값으로 설정하세요.

그다음 `label` 요소 아래에 속성이 없는 `input` 요소를 추가하세요. 다음 단계에서 필요한 속성을 추가할 것입니다.

# --hints--

열린 `label` 태그가 있어야 합니다.

```js
assert.match(code, /<label.*>/);
```

닫힌 `label` 태그가 있어야 합니다.

```js
assert.match(code, /<\/label>/);
```

`label` 요소 안에 `legend` 아래에 `fieldset` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

`label` 요소는 `for` 값을 가진 `"full-name"` 속성을 가져야 합니다.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

`label` 요소는 `Name (required):` 텍스트를 가져야 합니다.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

`input` 요소 아래에 `label` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
