---
id: 66a8380d911e3f4270d5cadc
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

이제 페이지의 주요 내용을 나타내는 `main` 요소를 추가할 차례입니다.

# --hints--

`main` 여는 태그가 있어야 합니다.

```js
assert.match(code, /<main>/i);
```

`main` 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/main>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
