---
id: 66a83fec026a7a4631e084d2
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`fieldset` 요소를 다룰 때, 입력 그룹을 설명하는 캡션을 사용하는 것이 일반적입니다. 이를 위해 `legend` 요소를 사용할 수 있습니다.

다음은 `legend` 요소를 사용하는 예시입니다:

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

`fieldset` 요소 안에 `legend` 요소를 추가하고 텍스트로 `Personal Information`를 넣으세요.

# --hints--

열리는 `legend` 태그가 있어야 합니다.

```js
assert.match(code, /<legend>/);
```

닫히는 `legend` 태그가 있어야 합니다.

```js
assert.match(code, /<\/legend>/);
```

`legend` 요소는 `fieldset` 요소 안에 있어야 합니다.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

`legend` 요소는 텍스트 `Personal Information`를 포함해야 합니다.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
