---
id: 66a83e5e491625454b6f62c3
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

폼은 사용자가 데이터를 입력할 수 있는 `inputs`로 구성됩니다. 관련 입력을 함께 묶으려면 `fieldset` 요소를 사용하세요.

다음은 `fieldset` 요소를 사용하는 예입니다:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

`form` 요소 안에 `fieldset` 요소를 추가하세요.

# --hints--

열린 `fieldset` 태그가 있어야 합니다.

```js
assert.match(code, /<fieldset>/i);
```

닫힌 `fieldset` 태그가 있어야 합니다.

```js
assert.match(code, /<\/fieldset>/i);
```

`fieldset` 요소는 `form` 요소 안에 있어야 합니다.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
