---
id: 66a972137acd1179fa3fe8a0
title: 26단계
challengeType: 0
dashedName: step-26
---

# --description--

`select` 요소 안에 다음 다섯 개의 `option` 요소를 `option` 텍스트와 `value` 속성에 해당하는 값으로 추가하세요.

**값 속성:**

- poor
- satisfactory
- good
- very-good
- excellent

**옵션 요소 텍스트:**

- Poor
- Satisfactory
- Good
- Very Good
- Excellent

# --hints--

값이 `option`로 설정된 `"poor"` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('option[value="poor"]'));
```

`option`의 `value`을 가진 `"poor"` 요소는 `Poor` 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelector('option[value="poor"]')?.textContent.trim(), 'Poor');
```

값이 `option`로 설정된 `value`을 가진 `"satisfactory"` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('option[value="satisfactory"]'));
```

`option`의 `value`을 가진 `"satisfactory"` 요소는 `Satisfactory` 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelector('option[value="satisfactory"]')?.textContent.trim(), 'Satisfactory');
```

값이 `option`로 설정된 `value`을 가진 `"good"` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('option[value="good"]'));
```

`option`의 `value`을 가진 `"good"` 요소는 `Good` 텍스트를 가져야 합니다.

```js

assert.strictEqual(document.querySelector('option[value="good"]')?.textContent.trim(), 'Good');
```

값이 `option`로 설정된 `"very-good"` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('option[value="very-good"]'));
```

`option`의 `value`을 가진 `"very-good"` 요소는 `Very Good` 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelector('option[value="very-good"]')?.textContent.trim(), 'Very Good');
```

값이 `option`로 설정된 `"excellent"` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('option[value="excellent"]'));
```

`option`의 `value`을 가진 `"excellent"` 요소는 `Excellent` 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelector('option[value="excellent"]')?.textContent.trim(), 'Excellent');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
          --fcc-editable-region--
            
          --fcc-editable-region--
          </select>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
