---
id: 66a954b2bcddba72076c1857
title: 17단계
challengeType: 0
dashedName: step-17
---

# --description--

`label` 요소 아래에 `radio` 버튼을 추가하고 `id`를 `"no-option"`로, `name` 속성을 `"hotel-stay"`로, `value` 속성을 `"no"`로 설정하세요.

새로운 `radio` 버튼 아래에 또 다른 `label` 요소를 추가하고 `for` 속성을 `"no-option"`로 설정하세요. `label`의 텍스트는 `No`이어야 합니다.

완료하면 라디오 버튼을 하나씩 선택해 보면서 작동하는지 확인할 수 있습니다.

# --hints--

첫 번째 `input` 요소 아래에 `label` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

`input` 요소는 `type`인 `radio` 속성을 가져야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

`input` 요소는 `id`인 `"no-option"` 속성을 가져야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

`input` 요소는 `name`인 `"hotel-stay"` 속성을 가져야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

입력 요소는 `value`로 설정된 `"no"` 속성을 가져야 합니다.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) label + input')?.value,
  'no'
);
```

두 번째 `label` 요소 아래에 두 번째 `input` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

`label` 요소는 `for`로 설정된 `"no-option"` 속성을 가져야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

`label` 요소는 `No` 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
