---
id: 66a93730719e1f68410cce54
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`name` 속성은 폼 데이터가 서버에 제출된 후 이를 식별하는 데 사용됩니다.

다음은 `name` 속성을 사용하는 예시입니다:

```html
<input type="email" name="email">
```

`name` 요소에 `input` 속성을 추가하고 값은 `"name"`로 하세요.

# --hints--

기존 `input` 요소는 제거하지 마세요.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

새로운 `input` 요소를 생성하지 마세요.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

`input` 요소에는 값이 `name`인 `"name"` 속성이 있어야 합니다.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
