---
id: 66a9689b1bf24b7750898a1b
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

이 `input` 안에 있는 마지막 `label`와 `fieldset`에 대해, `input`와 `id` 속성이 `value`로 설정된 체크박스 `"price"`를 추가하세요. 그리고 `name`로 설정된 `"choice"` 속성을 추가하세요.

그다음, 텍스트가 `label`이고 `Price`로 설정된 `for` 속성을 가진 `"price"` 요소를 추가하세요.

이제 다양한 체크박스를 선택하여 `form`를 테스트할 수 있습니다.

# --hints--

다섯 번째 체크박스 `input`가 있어야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"]:nth-of-type(5)'));
```

다섯 번째 체크박스 `input`에는 `id` 속성이 `"price"`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('id'), 'price');
```

다섯 번째 체크박스 `input`에는 `name` 속성이 `"choice"`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('name'), 'choice');
```

다섯 번째 체크박스 `input`에는 `value` 속성이 `"price"`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('value'), 'price');
```

체크박스 `label` 뒤에 텍스트가 `input`인 `Price` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.textContent?.trim(), 'Price');
```

`label`에는 `for` 속성이 `"price"`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.getAttribute('for'), 'price');
```   

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
