---
id: 66a9419e2d18476c645ce693
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

숫자 `input`는 숫자 입력 필드를 만들 때 사용합니다.

다음은 숫자 입력 필드의 예입니다:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

`min`과 `max` 속성은 입력 필드에 입력할 수 있는 최솟값과 최댓값을 설정하는 데 사용합니다.

`label` 요소 아래에 `type` 속성이 `"number"`로 설정되고 `id`가 `"age"`인 입력을 추가하세요.

`name` 속성은 `"age"`로, `min` 속성은 `"3"`로, `max` 속성은 `"100"`로 설정해야 합니다.

# --hints--

`input` 요소 아래에 숫자 `label`가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

숫자 `input`에는 `id`가 `age`로 설정되어야 합니다.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

숫자 `input`에는 `name` 속성이 `age`로 설정되어야 합니다.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

숫자 `input`에는 `min` 속성이 `"3"`로 설정되어야 합니다.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

숫자 `input`에는 `max` 속성이 `"100"`로 설정되어야 합니다.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
