---
id: 66a93c95bc58e26a8fe95818
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

다음으로, 이메일 `input` 아래에 `"email"` 유형의 `label`를 추가하세요. 이 `input`는 `id` 속성이 `"email"` 값으로 설정되어야 합니다. 또한, `name` 속성을 `"email"` 값으로 설정하세요.

이 `input`도 필수이므로 `required` 속성을 반드시 추가하세요.

마지막으로, `placeholder` 속성을 `"example@email.com"` 값으로 설정하세요.

# --hints--

이메일 `input` 아래에 `label` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

`input` 요소는 `type`가 `"email"`이어야 합니다.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

`input` 요소는 `id` 속성이 `"email"`로 설정되어야 합니다.

```js
assert.isNotNull(document.querySelector('input#email'));
```

`input` 요소는 `name` 속성이 `"email"`로 설정되어야 합니다.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

`input` 요소는 `required` 속성을 가져야 합니다.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

`input` 요소는 `placeholder` 속성이 `"example@email.com"`로 설정되어야 합니다.
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
