---
id: 66a962954f4e0d76223b37ed
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

다음으로, `input` 체크박스를 하나 더 추가하고 `id`와 `value` 속성을 `"location"`로 설정하며, `name` 속성은 `"choice"`로 설정하세요.

`label` 요소에 대해서는 `Location`의 텍스트와 `for` 속성을 `"location"`로 설정해야 합니다.

그 아래 `label` 요소 아래에 `input` 체크박스를 하나 더 추가하고 `id`와 `value` 속성을 `"reputation"`로 설정하며, `name` 속성은 `"choice"`로 설정하세요.

`label` 요소에 대해서는 `Reputation`의 텍스트와 `for` 속성을 `"reputation"`로 설정해야 합니다.

# --hints--

세 번째 체크박스 `input`가 있어야 합니다.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

세 번째 체크박스 `input`는 `id` 속성을 `"location"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

세 번째 체크박스 `input`는 `name` 속성을 `"choice"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'choice');
```

세 번째 체크박스 `input`는 `value` 속성을 `"location"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

체크박스 `label` 뒤에 텍스트 `input`가 있는 `Location` 요소가 있어야 합니다.

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

`label`는 `for` 속성을 `"location"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

`input` 아래에 네 번째 체크박스 `label`가 있어야 합니다.

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

네 번째 체크박스 `input`는 `id` 속성을 `"reputation"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

네 번째 체크박스 `input`는 `name` 속성을 `"choice"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'choice');
```

네 번째 체크박스 `input`는 `value` 속성을 `"reputation"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

평판 체크박스 아래에 텍스트 `label`가 있는 `Reputation` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

`label`는 `for` 속성을 `"reputation"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
