---
id: 66a9401c9d660d6bb15993e2
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

호텔 피드백 양식에 사용자가 나이를 입력할 수 있는 옵션이 있어야 합니다.

먼저, 양식에 텍스트 `label`가 있는 `Age (optional):` 요소를 추가하세요.

`for` 속성은 `"age"`로 설정해야 합니다.

# --hints--

`label` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

`label` 요소에 `for` 속성이 `"age"`로 설정되어야 합니다.

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

`label` 요소에 텍스트 `Age (optional):`가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
