---
id: 66a9521bc70162712caf118d
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

사용자가 여러 옵션 중 하나를 선택하게 하려면 라디오 버튼 세트를 사용할 수 있습니다.

다음은 두 개의 라디오 버튼 예시입니다:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

이 예시에서 라디오 버튼들은 같은 `name` 속성 값을 사용해 그룹화되어 있습니다. 이는 한 번에 하나의 라디오 버튼만 선택할 수 있음을 의미합니다.

`legend` 요소 아래에 `radio` 버튼을 추가하세요. 이 버튼은 `id`가 `"yes-option"`로 설정되어 있고, `name` 속성이 `"hotel-stay"`로, `value` 속성이 `"yes"`로 설정되어야 합니다.

`radio` 버튼 아래에 `label` 요소를 추가하고, 텍스트는 `Yes`이며 `for` 속성이 `"yes-option"`로 설정되어야 합니다.

# --hints--

`radio` 요소 아래에 `legend` 버튼이 있어야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

`radio` 버튼은 `id` 속성이 `"yes-option"`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

`radio` 버튼은 `name` 속성이 `"hotel-stay"`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

라디오 버튼은 `value` 속성이 `"yes"`로 설정되어야 합니다.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

`label` 버튼 아래에 `radio` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

`label` 요소는 `for` 속성이 `"yes-option"`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

`label` 요소는 텍스트 `Yes`을 포함해야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
