---
id: 66a9577022877d72d8f43b4f
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

다음 섹션에서는 사용자가 특정 호텔을 선택한 이유를 묻습니다. 사용자는 여러 옵션을 선택할 수 있습니다.

먼저 `fieldset` 요소를 하나 더 추가하세요.

`fieldset` 요소 안에 `legend` 요소를 추가하고 텍스트로 `Why did you choose to stay at our hotel? (Check all that apply)`를 넣으세요.

# --hints--

`fieldset` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 3);
```

`legend` 요소 안에 `fieldset` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(3) legend'), 1);
```

`legend` 요소에는 `Why did you choose to stay at our hotel? (Check all that apply)` 텍스트가 있어야 합니다. 철자 오류와 띄어쓰기 문제를 다시 확인하세요.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) legend')?.innerText.trim(), 'Why did you choose to stay at our hotel? (Check all that apply)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        --fcc-editable-region--
        
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
