---
id: 66a937e74920ba68ebe5e86d
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

이전 수업에서 `placeholder`와 `required` 속성을 이렇게 다루는 방법을 배웠습니다:

```html
<input type="text" placeholder="e.g., John Doe" required>
```

기존 `input` 요소에 `placeholder` 속성을 값 `"e.g., John Doe"`로 추가하세요.

또한 `required` 요소에 `input` 속성을 추가하세요.

# --hints--

기존 `input` 요소는 제거하지 마세요.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

`input`에는 값이 `placeholder`인 `"e.g., John Doe"` 속성이 있어야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

`input`에는 `required` 속성이 있어야 합니다.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
