---
id: 66a930b20f589b6664c51cb0
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

사용자가 전체 이름을 입력하면 `input`는 일반 텍스트를 허용합니다.

이전 수업에서 `type` 속성을 이렇게 다루는 방법을 배웠습니다:

```html
<input type="text">
```

기존 `input` 요소에 `type` 속성을 `"text"`로 추가하세요.

또한 `input` 요소에 `id` 속성을 값 `"full-name"`로 지정하세요.

# --hints--

기존 `input` 요소를 제거하지 마세요.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

새로운 `input` 요소를 만들지 마세요.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

`input` 요소에 `type` 속성을 값 `"text"`로 지정해야 합니다.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

`input` 요소에 `id` 속성을 값 `"full-name"`로 지정해야 합니다.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
