---
id: 66a94ea5df66236ebed260e8
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

다음 섹션에서는 사용자가 이전에 호텔에 머문 적이 있는지 묻는 역할을 합니다.

`fieldset` 요소를 추가하는 것부터 시작하세요.

`fieldset` 요소 안에 `legend` 텍스트가 있는 `Was this your first time at our hotel?` 요소를 추가하세요.

# --hints--

`fieldset` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 2);
```

`legend` 요소 안에 `fieldset` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('fieldset legend'), 2);
```

`legend` 요소에는 `Was this your first time at our hotel?` 텍스트가 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('fieldset legend')[1]?.textContent.trim(), 'Was this your first time at our hotel?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
