---
id: 66a93bbe65a26169dbf3bc39
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

호텔 피드백 양식에 사용자 이메일 주소도 수집해야 합니다.

양식에 텍스트가 `label`인 새 `Email address (required):` 요소를 추가하는 것부터 시작하세요. `label` 요소에는 `for` 값을 가진 `"email"` 속성이 있어야 합니다.

# --hints--

`label` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

`label` 요소에는 텍스트 `Email address (required):`가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

`label` 요소에는 `for` 값을 가진 `"email"` 속성이 있어야 합니다.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
