---
id: 67a51d3a8a6fe123b77b0c6e
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

`input` 요소는 `size` 속성을 가질 수 있습니다. 이 속성은 사용자가 입력할 때 보여야 하는 문자 수를 정의합니다. `size`의 값은 0보다 큰 음이 아닌 정수여야 합니다. `size`이 지정되지 않았거나 잘못된 값으로 지정된 경우, 입력은 브라우저가 설정한 기본 너비를 갖게 됩니다.

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

name과 email 입력에 `size` 속성을 값 `"20"`로 지정하세요.

# --hints--

name `input`에 `size` 속성을 값 `"20"`로 지정해야 합니다.

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

email `input`에 `size` 속성을 값 `"20"`로 지정해야 합니다.

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
