---
id: 66cc26907ff6908402af0149
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

要測試你的 `calculateSquare` 函式， 你需要呼叫該函式數次。

為你的程式碼添加一個 `console.log`，該 `console.log` 呼叫 `calculateSquare` 函式並帶有引數 `2`。

然後，為你的程式碼添加另一個 `console.log`，該 `console.log` 呼叫 `calculateSquare` 函式並帶入引數 `9`。

# --hints--

你應該有一個 `console.log`，它以引數 `2` 呼叫 `calculateSquare` 函式（程式）。

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*2\s*\)\s*\)/);
```

你應該有一個 `console.log`，它以引數 `9` 呼叫 `calculateSquare` 函式（程式）。

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*9\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

--fcc-editable-region--

--fcc-editable-region--
```
