---
id: 66cc0f1ae40802781b2ea972
title: 步驟6
challengeType: 1
dashedName: step-6
---

# --description--

In the previous lectures, you learned how to work with function parameters.

函數參數是在函數聲明中定義的變量，用作佔位符。

下面是一個有參數的函數示例：

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

聲明一個名爲 `calculateSum` 的函數，它有 `num1` 和 `num2` 兩個參數。

# --hints--

你應聲明一個名爲 `calculateSum` 的函數。

```js
assert.isFunction(calculateSum);
```

你的 `calculateSum` 函數應包含 `num1` 和 `num2` 兩個參數,。

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
