---
id: 66cc21d23238dc8240a8a182
title: 步驟 16
challengeType: 1
dashedName: step-16
---

# --description--

如果能讓計算器計算一個數字的平方就更好了。 一個數字的平方就是這個數字乘以它自己。

要在 JavaScript 中計算一個數字的平方，可以使用 <dfn>Math.pow()</dfn> 方法。 `Math.pow()` 方法需要兩個參數：底數和指數。

你還可以使用指數運算符（`**`）來計算一個數字的平方。

這是一個例子：

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

聲明一個名爲 `calculateSquare` 的函數，接受參數 `num` 並返回 `num` 的平方。

# --hints--

你應有一個名爲 `calculateSquare` 的函數。

```js
assert.isFunction(calculateSquare);
```

你的 `calculateSquare` 函數應接受參數 `num`。

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的 `calculateSquare` 函數應返回 `num` 的平方。

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
