---
id: 66cc172af7159a7c67804544
title: 步驟 12
challengeType: 1
dashedName: step-12
---

# --description--

現在是為你的計算機添加乘法功能的時候了。

宣告一個名為 `calculateProduct` 的函式（程式），它接受兩個參數，`num1` 和 `num2`。

在函式內，傳回兩個參數的乘積。

然後，添加一個 `console.log`，用引數 `13` 和 `5` 呼叫 `calculateProduct` 函式（程式）。

# --hints--

你應該有一個名為 `calculateProduct` 的函式（程式）。

```js
assert.isFunction(calculateProduct);
```

你的 `calculateProduct` 函式應該有兩個名為 `num1` 和 `num2` 的參數。

```js
const regex = __helpers.functionRegex('calculateProduct', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的 `calculateProduct` 函式（程式）應該傳回這兩個參數的乘積。

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

你應該有一個 `console.log`，它呼叫 `calculateProduct` 函式（程式），並帶入引數 `13` 和 `5`。

```js
assert.match(code, /console\.log\s*\(\s*calculateProduct\s*\(\s*13\s*,\s*5\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
