---
id: 66cc0a9e06e00b75d6782be9
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

在之前的課程中，你學會了如何呼叫（喚起）一個函式。呼叫函式意味著執行函式內的程式碼。

這裡是如何呼叫函式（程式）的提醒：

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

為你的程式碼添加一個 `console.log` 述語，並在其中呼叫 `addTwoAndSeven` 函式（程式）。

你現在應該會看到 `2` 和 `7` 的總和已記錄到主控台。

# --hints--

你的程式碼中應該有一個 `console.log`。

```js
assert.match(code, /console\.log(.*)/);
```

你應該在 `console.log` 裡呼叫 `addTwoAndSeven` 函式（程式）。

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
