---
id: 66cc14ca14e65e7a9ebecb80
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

現在你的計算機可以將兩個數字相加，是時候創建一個函式（程式）來將兩個數字相減。

宣告一個名為 `calculateDifference` 的函式（程式），該函式接受兩個參數，`num1` 和 `num2`。

在函式內，傳回兩個參數之間的差異。

# --hints--

你應該有一個名為 `calculateDifference` 的函式（程式）。

```js
assert.isFunction(calculateDifference);
```

你的 `calculateDifference` 函式應該有兩個名為 `num1` 和 `num2` 的參數。

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的 `calculateDifference` 函式（程式）應該傳回兩個參數之間的差異。

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
