---
id: 66cc095dfe1682753d2ab030
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

你在先前課程中學到的概念之一是如何從函式傳回返回值。

這裡提醒你如何從函式傳回一個值：

```js
function myFunction() {
  return "Hello World";
}
```

在你的 `addTwoAndSeven` 函式中，傳回 `2` 和 `7` 的總和。

# --hints--

你的 `addTwoAndSeven` 函式應該傳回 `2` 和 `7` 的總和。

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

你的 `addTwoAndSeven` 函式（程式）應該使用 `+` 運算子傳回總和。

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
