---
id: 66cc12fa504b0479dac479a0
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

函式（程式）的優點是它們可以使用不同的引數進行呼叫，讓你能夠用各種數值重複使用相同的程式碼。

以下是一些使用不同字串引數的範例函式呼叫：

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

為你的程式碼添加一個 `console.log`，該 `console.log` 呼叫 `calculateSum` 函式（程式），並帶入引數 `2` 和 `5`。

# --hints--

你應該有一個 `console.log`，它呼叫 `calculateSum` 函式（程式），並帶入引數 `2` 和 `5`。

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
