---
id: 66cc10958d72e57907f898b2
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

在你的 `calculateSum` 函式中，你需要傳回兩個參數的總和。

# --hints--

你的 `calculateSum` 函式（程式）應該傳回兩個參數的總和。

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
