---
id: 66cc06012e46aa6bc9b8c001
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置計算機應用程式來複習如何使用函式（程式）。 

在之前的課程中，你學會了如何這樣宣告函式（程式）：

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

首先建立一個名為 `addTwoAndSeven` 的函式（程式）。

你可以選擇使用一般函式語法或箭頭函式語法。

# --hints--

你應該宣告一個名為 `addTwoAndSeven` 的函式（程式）。

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
