---
id: 66cc21d23238dc8240a8a182
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

如果你的計算機能計算一個數字的平方會很不錯。數字的平方是該數字乘以它本身。

要在 JavaScript 中計算一個數字的平方，你可以使用 <dfn>Math.pow()</dfn> 方法。`Math.pow()` 方法接受兩個引數：底數和指數。 

你也可以使用冪次運算子（`**`）來計算一個數字的平方。

這裡有一個範例：

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

宣告一個名為 `calculateSquare` 的函式，該函式帶有一個 `num` 參數並傳回 `num` 的平方。

# --hints--

你應該有一個名為 `calculateSquare` 的函式（程式）。

```js
assert.isFunction(calculateSquare);
```

你的 `calculateSquare` 函式（程式）應該接受一個 `num` 參數。

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的 `calculateSquare` 函式（程式）應該傳回 `num` 的平方。

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
