---
id: 66cc0ba5881acb7692cfc4de
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

現在，是時候為你的程式碼添加另一個函式（程式）。

宣告一個名為 `addThreeAndFour` 的函式（程式），傳回 `3` 和 `4` 的總和。

然後在 `console.log` 裡呼叫 `addThreeAndFour` 函式（程式）以查看結果。

# --hints--

你應該宣告一個名為 `addThreeAndFour` 的函式（程式）。

```js
assert.isFunction(addThreeAndFour);
```

你的 `addThreeAndFour` 函式（程式）應該傳回 `3` 和 `4` 的總和。

```js
assert.equal(addThreeAndFour(), 7);
```

你的 `addThreeAndFour` 函式應該使用 `+` 運算子傳回總和。

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

你應該在 `console.log` 內呼叫 `addThreeAndFour` 函式（程式）。

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
