---
id: 66cc0de2920bec775f610424
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

即使這些函式（程式）如預期運作，你的程式碼中仍有大量重複。 

由於你正在建置一個計算機，你不想為你想要相加的每一種可能的數字組合都創建一個函式（程式）。

最好創建一個可重複使用的函式（程式），能將任意兩個數字相加。

移除你目前為止所撰寫的所有程式碼。

在下一步中，你將檢視如何在函式中使用參數和引數。

# --hints--

你不應該有一個名為 `addTwoAndSeven` 的函式。

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

你不應該有對 `addTwoAndSeven` 的函式（程式）呼叫。

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

你不應該有一個名為 `addThreeAndFour` 的函式（程式）。

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

你不應該有對 `addThreeAndFour` 的函式（程式）呼叫。

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
