---
id: 66cc0de2920bec775f610424
title: 步驟5
challengeType: 1
dashedName: step-5
---

# --description--

儘管這些函數的運行符合預期，但你的代碼中存在大量重複。

由於你正在構建一個計算器，因此不必爲每種可能的數字加法組合創建一個函數。

最好是創建一個可重複使用的函數，將任意兩個數字相加。

刪除迄今爲止編寫的所有代碼。

下一步，你將學習如何在函數中使用參數。

# --hints--

你不應有名爲 `addTwoAndSeven` 的函數。

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

你不應有名爲 `addTwoAndSeven` 的函數。

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

你不應有名爲 `addThreeAndFour` 的函數。

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

你不應有名爲 `addThreeAndFour` 的函數。

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
