---
id: 66cc0ba5881acb7692cfc4de
title: 步驟4
challengeType: 1
dashedName: step-4
---

# --description--

現在，是時候添加另一個函數了。

聲明一個名爲 `addThreeAndFour` 的函數，返回 `3` 與 `4` 的和。

然後在一個 `console.log` 裏面調用 `addThreeAndFour` 函數，查看結果。

# --hints--

你應聲明一個名爲 `addThreeAndFour` 的函數。

```js
assert.isFunction(addThreeAndFour);
```

你的 `addThreeAndFour` 函數應返回 `3` 與 `4` 的和。

```js
assert.equal(addThreeAndFour(), 7);
```

你的 `addThreeAndFour` 函數應使用 `+` 運算符返回和。

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

你應在一個 `console.log` 中調用 `addThreeAndFour` 函數。

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
