---
id: 66cc0a9e06e00b75d6782be9
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

In previous lessons, you learned how to call (invoke) a function. 調用函數意味着執行函數內部的代碼。

下面是提醒大家如何調用函數：

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

添加一個 `console.log` 語句，並在其中調用 `addTwoAndSeven` 函數。

現在你應該可以看到控制檯記錄了 `2` 和 `7` 的和。

# --hints--

你的代碼中應該有一個 `console.log`。

```js
assert.match(code, /console\.log(.*)/);
```

你應在 `console.log` 中調用 `addTwoAndSeven` 函數。

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
