---
id: 66cc0f1ae40802781b2ea972
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

在之前的課程中，你學會了如何使用函式（程式）的參數。

函式（程式）的參數是在函式（程式）的宣告中定義的變數，並且充當佔位符。

這是一個具有參數的函式（程式）範例：

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

宣告一個名為 `calculateSum` 的函式（程式），它接受兩個參數，`num1` 和 `num2`。

# --hints--

你應該宣告一個名為 `calculateSum` 的函式（程式）。

```js
assert.isFunction(calculateSum);
```

你的 `calculateSum` 函式（程式）應該接受兩個參數，`num1` 和 `num2`。

```js
const explorer = await __helpers.Explorer(code);
const { calculateSum } = explorer.allFunctions;
assert.lengthOf(calculateSum?.parameters, 2);
assert.equal(calculateSum?.parameters[0].toString(), "num1");
assert.equal(calculateSum?.parameters[1].toString(), "num2");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
