---
id: 66cc16abc30b367ba1d1f9ef
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

現在是測試你的 `calculateDifference` 函式（程式） 的時候了。

首先為你的程式碼添加一個 `console.log`，該 `console.log` 呼叫 `calculateDifference` 函式（程式），並帶入引數 `22` 和 `5`。

然後，為你的程式碼添加第二個 `console.log`，該 `console.log` 呼叫 `calculateDifference` 函式（程式），並帶入引數 `12` 和 `1`。

最後，為你的第三個 `console.log` 添加一個呼叫 `calculateDifference` 函式（程式），並帶入引數 `17` 和 `9`。

# --hints--

你應該有一個 `console.log`，它呼叫 `calculateDifference` 函式（程式），並帶入引數 `22` 和 `5`。

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*22\s*,\s*5\s*\)\s*\)\s*;?/);
```

你應該有一個 `console.log`，它呼叫 `calculateDifference` 函式（程式），並帶入引數 `12` 和 `1`。

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*12\s*,\s*1\s*\)\s*\)\s*;?/);
```

你應該有一個 `console.log`，它呼叫 `calculateDifference` 函式（程式），並帶入引數 `17` 和 `9`。

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*17\s*,\s*9\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
