---
id: 66cc1a3a39aef47d6473cb2f
title: 步驟 13
challengeType: 1
dashedName: step-13
---

# --description--

下一步是為你的計算機添加除法功能。

宣告一個名為 `calculateQuotient` 的函式（程式），它接受兩個參數，`num1` 和 `num2`。

在函式內，傳回兩個參數的商。

然後，為你的程式碼添加一個 `console.log`，該 `console.log` 呼叫 `calculateQuotient` 函式（程式），並帶入引數 `7` 和 `11`。

# --hints--

你應該有一個名為 `calculateQuotient` 的函式（程式）。

```js
assert.isFunction(calculateQuotient);
```

你的 `calculateQuotient` 函式（程式）應該有兩個名為 `num1` 和 `num2` 的參數。

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的 `calculateQuotient` 函式（程式）應該傳回這兩個參數的商。

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

你應該有一個 `console.log`，它呼叫 `calculateQuotient` 函式（程式），並帶入引數 `7` 和 `11`。

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
