---
id: 66cc2739f687e484d50bb6f1
title: 步驟 18
challengeType: 1
dashedName: step-18
---

# --description--

你計算機的最後一個功能將是平方根功能。

平方是一個數字乘以它自己。因此，平方根是當該數字乘以它自己時會得到原始數字的數字。

例如，`9` 的平方根是 `3`，因為 `3 * 3 = 9`。

要取得數字的平方根，你可以使用 <dfn>Math.sqrt()</dfn> 方法。

這裡有一個範例：

```js
// result: 3
Math.sqrt(9);
```

宣告一個名為 `calculateSquareRoot` 的函式，該函式接受一個 `num` 參數並傳回 `num` 的平方根。

# --hints--

你應該有一個名為 `calculateSquareRoot` 的函式（程式）。

```js
assert.isFunction(calculateSquareRoot);
```

你的 `calculateSquareRoot` 函式（程式）應該接受一個 `num` 參數。

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的 `calculateSquareRoot` 函式（程式）應該傳回 `num` 的平方根。

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
