---
id: 66cc0ba5881acb7692cfc4de
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

現在是時候添加另一個函式（程式）了。

宣告一個名為 `addThreeAndFour` 的函式（程式），傳回 `3` 和 `4` 的總和。

然後在 `console.log` 裡呼叫 `addThreeAndFour` 函式（程式）以查看結果。

# --hints--

你應該宣告一個名為 `addThreeAndFour` 的函式（程式）。

```js
assert.isFunction(addThreeAndFour);
```

你的 `addThreeAndFour` 函式（程式）應該傳回 `3` 和 `4` 的總和。

```js
assert.strictEqual(addThreeAndFour(), 7);
```

你的 `addThreeAndFour` 函式應該使用 `+` 運算子傳回總和。

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

你應該在 `console.log` 內呼叫 `addThreeAndFour` 函式（程式）。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
