---
id: 66cdf76685e4cb5a8726e27b
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

為了進一步測試你的 `calculateSum` 函式， 你需要用不同的引數來呼叫它。

為你的程式碼添加另一個 `console.log`，該 `console.log` 呼叫 `calculateSum` 函式（程式），並帶有引數 `10` 和 `10`。

在那個 `console.log` 下面，為你的程式碼添加另一個 `console.log`，該 `console.log` 呼叫 `calculateSum` 函式（程式），並帶入引數 `5` 和 `5`。

# --hints--

你應該有一個 `console.log`，它呼叫 `calculateSum` 函式（程式），並帶入引數 `10` 和 `10`。

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

你應該有一個 `console.log`，它呼叫 `calculateSum` 函式（程式），並帶入引數 `5` 和 `5`。

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--

--fcc-editable-region--
```
