---
id: 6718d2d59337c822ecb697ff
title: Desarrollar un Programa de Gestión de Cuentas Bancarias
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio. 

**Historias de usuario:**

1. Debes definir una clase llamada `BankAccount` con un constructor. El constructor debe establecer el valor predeterminado del `balance` a `0` e inicializar un array vacío llamado `transactions` para almacenar registros de transacciones como objetos.

2. Cada transacción almacenada en el array `transactions` debe ser un objeto con dos propiedades: `type` y `amount`. La propiedad `type` debería ser `deposit` o `withdraw`, y la propiedad `amount` debería ser la cantidad depositada o retirada.

3. You should define a method named `deposit` that takes the deposit amount as a parameter. When the deposit amount is greater than `0`, it should:

   - Empujar un nuevo objeto al array transacciones con el tipo `deposit` y la cantidad depositada.
   - Actualiza el balance.
   - Devuelve <code>"Se han depositado exitosamente $[amount]. Nuevo balance: $[balance]"</code>.

4. Si la cantidad es menos o igual a `0`, el método `deposit` debería devolver `"Deposit amount must be greater than zero."`.

5. You should define a method named `withdraw` that takes an amount as a parameter. This method should update the current balance according to withdrawals. When the amount to be withdrawn is greater than `0` and less than or equal to the current balance, it should:

   - Empujar un nuevo objeto al array de transacciones con el tipo `withdraw` y la cantidad retirada.
   - Actualiza el balance.
   - Devuelve <code>"Se han retirado exitosamente $[amount]. Nuevo balance: $[balance]"</code>.

6. Si la cantidad a retirar es menos o igual a `0` o mayor al balance actual, el método `withdraw` debe devolver `"Insufficient balance or invalid amount."`.


7. Debes definir un método llamado `checkBalance` que devuelve el balance actual en el formato `"Current balance: $[balance]"`.

8. Debes definir un método llamado `listAllDeposits` que iterará a través del array de transacciones y devuelve todos los depósitos en el formato `"Deposits: amount,amount,..."`.

9. Debes definir un método llamado `listAllWithdrawals` que iterará a través del array de transacciones y deuelve todos los retiros en el formato `"Withdrawals: amount,amount,..."`.

10. Debes crear una nueva instancia de `BankAccount` llamada `myAccount`.

11. Tu cuenta bancaria `myAccount` debe tener al menos cinco transacciones.

12. Tu cuenta bancaria `myAccount` debe tener al menos dos depositos.

13. Tu cuenta bancaria `myAccount` debe tener al menos dos retiros.

14. Tu cuenta bancaria `myAccount` debe tener un balance mayor que `$100`.

# --hints--

Debes definir una clase llamada `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

El objeto `BankAccount` debe tener inicialmente un `balance` de `0` y un array vacío `transactions` para almacenar el registro de transacciones.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

Debes tener un método `deposit` que tome la cantidad del depósito como parámetro.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

Debes tener un método `withdraw` que tome la cantidad retirada como parámetro.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

Debes tener un método `checkBalance` que compruebe el balance actual.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

Debes tener un método `listAllDeposits` que enliste todos los depósitos.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

Debes tener un método `listAllWithdrawals` que enliste todos los retiros.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

<code>BankAccount.deposit(100)` debe devolver `"Succesfuly deposited $100. New balance: $100"</code>.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` debe devolver `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` debe devolver `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

Cuando el balance de la cuenta es `100`,`BankAccount.withdraw(150)` debe devolver `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` debe devolver `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` debe devolver `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

Cuando el balance es <code>200`, `BankAccount.withdraw(150)` debe devolver `"Succesfuly withdrew $150. New balance: $50"</code>.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

Cuando el balance de la cuenta es `200`, `BankAccount.checkBalance()` debe devolver `"Current balance: $200"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

Cuando depositas `10`, `35`, `90`, el método `listAllDeposits` debe devolver `"Deposits: 10,35,90"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

Cuando retires `20`,`50`,`100`, el método `listAllWithdrawals` debe devolver `"Withdrawals: 20,50,100"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

Debes tener una instancia de `BankAccount` llamada `myAccount`.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

Tu cuenta bancaria `myAccount` debe tener al menos cinco transacciones.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

Tu cuenta bancaria `myAccount` debe tener al menos dos depositos.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

Tu cuenta bancaria `myAccount` debe tener al menos dos retiros.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```


Tu cuenta bancaria `myAccount` debe tener un balance mayor que `$100`.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
