---
id: 68f890b73bbfc9efa7379e7b
title: "Eseguire il debug degli errori di conversione implicita in un'app con bug"
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Hai appena iniziato a lavorare in un'agenzia locale di sviluppo web, e il tuo primo compito è sistemare del codice con bug lasciato dagli sviluppatori precedenti.

Stavano tentando semplici operazioni aritmetiche, ma qualcosa è andato storto, i risultati non hanno senso. La tua attività è rivedere, eseguire il debug e correggere il codice in modo che funzioni come previsto e sia più facile da leggere.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una variabile chiamata `firstResult` che somma correttamente i numeri `5` e `10` per produrre il valore `15`.
2. Dovresti avere una variabile chiamata `secondResult` che sottrae correttamente `5` da `8` per produrre il valore `3`.
3. Dovresti avere una variabile chiamata `thirdResult` che produce il valore `6` sommando due numeri. Sostituisci il booleano attualmente usato nell'espressione con un numero.
4. Dovresti avere una variabile chiamata `fourthResult` che produce il valore `8` sommando due numeri. Sostituisci il booleano attualmente usato nell'espressione con un numero.
5. Dovresti avere una variabile chiamata `fifthResult` che moltiplica correttamente due numeri per ottenere il prodotto `20`. Sostituisci la stringa attualmente usata nell'espressione con un numero.
6. Dovresti avere una variabile chiamata `sixthResult` che somma correttamente due numeri per produrre il valore `22`. Sostituisci il `null` attualmente usato nell'espressione con un numero.

# --hints--

Dovresti avere una variabile chiamata `firstResult`.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` dovrebbe avere il valore numerico `15`.

```js
assert.strictEqual(firstResult, 15);
```

Dovresti avere una variabile chiamata `secondResult`.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` dovrebbe avere il valore numerico `3`.

```js
assert.strictEqual(secondResult, 3);
```

Dovresti avere una variabile chiamata `thirdResult`.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` dovrebbe avere il valore numerico `6`.

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` non dovrebbe includere un valore booleano (`true` o `false`).

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Dovresti avere una variabile chiamata `fourthResult`.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` dovrebbe avere il valore numerico `8`.

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` non dovrebbe includere un valore booleano (`true` o `false`).

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Dovresti avere una variabile chiamata `fifthResult`.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` dovrebbe avere il valore numerico `20`.

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` non dovrebbe includere una moltiplicazione di stringhe (es. "10" * 2).

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

Dovresti avere una variabile chiamata `sixthResult`.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` dovrebbe avere il valore numerico `22`.

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` non dovrebbe includere `null` nella sua espressione.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
