---
id: 6909ed80e1c195f49564edac
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, aprenderás a trabajar con sentencias condicionales y operadores de comparación.

Comienza declarando una variable llamada `truthyOrFalsy` y asígnale el valor booleano `true`.

# --hints--

Debes tener una variable llamada `truthyOrFalsy`.

```js
assert.exists(truthyOrFalsy);
```

Debes asignar `true` a tu variable `truthyOrFalsy`.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
