---
id: 690a562fd611e8155599ce98
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Cambia la asignación de `truthyOrFalsy` estableciendo su valor en `"freeCodeCamp"`. Aún verás que se registra `true` en la consola porque una cadena no vacía es un valor truthy.

# --hints--

Tu variable `truthyOrFalsy` debe tener el valor `"freeCodeCamp"`.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
