---
id: 690a0f7a5c22c15222a6aa5c
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Ahora crea otra estructura `if` que use `isTimmyAGamer` como condición. Dentro del cuerpo de tu estructura `if`, registra `"Timmy loves to play World of Warcraft."` en la consola.

Ten en cuenta que no verás nada registrado en la consola porque la condición de tu nueva estructura `if` es falsa.

# --hints--

Debes tener dos sentencias `if` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

Tu segunda estructura `if` debe usar `isTimmyAGamer` como condición.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

Debes registrar `"Timmy loves to play World of Warcraft."` en la consola dentro de tu segundo `if`.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--
const isTimmyAGamer = false;


--fcc-editable-region--
```
