---
id: 690a110a76d0496d9e7c28eb
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Declara otra variable llamada `timmyAge` y asígnale el valor `18`.

# --hints--

Debes tener una variable llamada `timmyAge`.

```js
assert.exists(timmyAge);
```

Debes asignar `18` a tu variable `timmyAge`.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
