---
id: 690a57fef66aac39dffc61a2
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora que sabes qué son los valores truthy y falsy, elimina tanto la declaración de tu variable como la instrucción `console.log` de tu código.

# --hints--

No deberías tener una variable `truthyOrFalsy`.

```js
const explorer = await __helpers.Explorer(code);
assert.notExists(explorer.variables.truthyOrFalsy);
```

No deberías tener `console.log(Boolean(truthyOrFalsy))` en tu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
