---
id: 690a06833d3e83b5226ade78
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Una estructura condicional se usa para ejecutar código dependiendo de una condición. Está compuesta por la palabra clave `if`, seguida de una condición entre paréntesis y un cuerpo entre llaves. Aquí tienes un ejemplo de una estructura `if`:

```js
if (condition) {
  console.log("condition is truthy");
}
```

Cuando `condition` es un valor verdadero (cadena no vacía, número distinto de cero, `true`, etc.) se ejecuta el código dentro del cuerpo de la sentencia `if`.

Debajo de tu variable, crea una estructura `if` que use `hasDeveloperJob` como condición. Dentro del cuerpo de tu estructura `if`, muestra en la consola `"Timmy is employed as a developer."`.

# --hints--

Debes tener una sentencia `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

Tu sentencia `if` debe usar `hasDeveloperJob` como condición.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

Debes registrar `"Timmy is employed as a developer."` en la consola dentro de tu sentencia `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const hasDeveloperJob = true;


--fcc-editable-region--
```
