---
id: 690a0df27ee9ee36d935d530
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Declara otra variable llamada `isTimmyAGamer` y asígnale el booleano `false`.

# --hints--

Debes declarar una variable llamada `isTimmyAGamer`.

```js
assert.exists(isTimmyAGamer);
```

Debes asignar el valor booleano `false` a tu variable `isTimmyAGamer`.

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
