---
id: 690a3ad44874a928a2fe3077
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Los operadores de comparación te permiten comparar valores y escribir declaraciones condicionales más complejas.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

Por ejemplo, el código anterior muestra `"Happy Birthday!"` en la consola solo cuando los valores de las variables `today` y `birthday` son iguales.

Crea una tercera estructura `if`. Para su condición, usa el operador `>=` para verificar si `timmyAge` es mayor o igual a `16`.

Dentro del cuerpo de tu nueva estructura `if`, registra `"Timmy is old enough to drive."` en la consola.

# --hints--

Debes tener una tercera estructura `if` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

Tu tercer `if` debe usar `timmyAge >= 16` como condición.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s* 16\s*\)\s*\{.+\}/s);
```

Debes registrar `"Timmy is old enough to drive."` en la consola dentro de tu tercer `if`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}
--fcc-editable-region--
const timmyAge = 18;


--fcc-editable-region--
```
