---
id: 690a51806fb743bc896c3af7
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Como aprendiste en una lección anterior, los booleanos son un tipo de dato que solo tiene los valores `true` y `false`.

Un valor truthy es un valor que se evalúa como `true` en un contexto booleano. Ejemplos de valores truthy son:

- Cadenas no vacías
- Números distintos de cero
- El booleano `true`

Y otros valores que encontrarás más adelante.

Puedes usar la función `Boolean()` para verificar la veracidad de un valor. Por ejemplo, `Boolean("Hello World!")` devolverá `true` porque `"Hello World!"` es verdadero.

Ahora, llama a `Boolean(truthyOrFalsy)` y registra el resultado en la consola.

# --hints--

Debes registrar `Boolean(truthyOrFalsy)` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
