---
id: 690a3f884e083a6a33b81fcf
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

Finalmente, cambia la declaración de `timmyAge` y asígnale un número menor que `16`. Después de eso verás `"Timmy is not old enough to drive."` registrado en la consola.

# --hints--

Tu variable `timmyAge` debe tener un valor menor que `16`.

```js
assert.isAtMost(timmyAge, 15);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}
--fcc-editable-region--
const timmyAge = 18;
--fcc-editable-region--
if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```

# --solutions--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 15;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```
