---
id: 690a58cc4f828a49eab1b53a
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Declara una variable llamada `hasDeveloperJob` y asígnale el valor booleano `true`.

# --hints--

Debes tener una variable llamada `hasDeveloperJob`.

```js
assert.exists(hasDeveloperJob);
```

Debes asignar `true` a tu variable `hasDeveloperJob`.

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
