---
id: 690a57095c4afb27bbe751b3
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Los valores falsy son valores que se evalúan como `false` en un contexto booleano. Ejemplos de valores falsy son:

- `""` (cadena vacía)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

Te encontrarás con todos estos al avanzar en este curso.

Por ahora, cambia la asignación de `truthyOrFalsy` estableciendo su valor a una cadena vacía. Después de eso, verás `false` registrado en la consola.

# --hints--

Debes asignar una cadena vacía a `truthyOrFalsy`.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
