---
id: 66c8ba975ee7230e29f6c4ac
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Cuando el usuario es elegible para un préstamo, querrás mostrar un mensaje en la consola.

Para eso, construirás una función dentro de la cual tendrás algunas verificaciones que devolverán qué préstamo es elegible para el solicitante.

Crea una función vacía `getLoanMessage` con parámetros `annualIncome` y `creditScore`.

# --hints--

Deberías crear una función `getLoanMessage`

```js
assert.isFunction(getLoanMessage)
```

Tu función `getLoanMessage` debería tener `annualIncome` y `creditScore` como parámetros.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` debería ser una función vacía.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
