---
id: ac6993d51946422351508a41
title: Einen Algorithmus zum Kürzen eines Strings implementieren
challengeType: 26
dashedName: truncate-a-string
---

# --description--

In diesem Labor üben Sie, einen String auf eine bestimmte Länge zu kürzen.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion `truncateString` haben, die zwei Argumente akzeptiert, das erste ein String, das zweite eine Zahl.
1. Wenn die Länge des Strings größer als die angegebene Zahl ist, soll der String so gekürzt werden, dass seine Länge der angegebenen Zahl entspricht, und `...` soll am Ende des gekürzten Strings angehängt werden.
1. Wenn die Länge des Strings gleich oder kleiner als die angegebene Zahl ist, soll der String unverändert zurückgegeben werden.

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` sollte den String `A-tisket...` zurückgeben.

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` sollte den String `Peter Piper...` zurückgeben.

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` sollte den String `A-tisket a-tasket A green and yellow basket` zurückgeben.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` sollte den String `A-tisket a-tasket A green and yellow basket` zurückgeben.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` sollte den String `A...` zurückgeben.

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` sollte den String `Ab...` zurückgeben.

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
