---
id: a39963a4c10bc8b4d4f06d7e
title: Tekeleza kitendakazi cha kuondoa thamani
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

Katika maabara hii, utaunda kitendakazi kinachopokea safu ya data ya awali kama hoja yake ya kwanza, ikifuatiwa na hoja moja au zaidi zinazowakilisha thamani za kuondoa.

**Lengo**: Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi cha `destroyer` kinachokubali safu ya data na hoja moja au zaidi za ziada.
1. Kitendakazi cha `destroyer` kinapaswa kurudisha safu mpya isiyo na vipengele vyote kutoka kwenye hoja ya kwanza vinavyolingana na hoja zozote za baadaye.
1. Kitendakazi lazima kikubali idadi isiyojulikana ya hoja.

# --hints--

Unapaswa kuwa na kitendakazi cha `destroyer`.

```js
assert.isFunction(destroyer);
```

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` inapaswa kurudisha `[1, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` inapaswa kurudisha `[1, 5, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` inapaswa kurudisha `[1]`.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` inapaswa kurudisha `[]`.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` inapaswa kurudisha `["hamburger"]`.

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

`destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )` inapaswa kurudisha `[12, 92, 65]`.

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```
