---
id: 66adb844118ba74107ce771f
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

A próxima mensagem do bot será sobre a localização do bot.

Crie uma variável chamada `botLocationSentence`.

Depois, use a concatenação de strings com o operador `+` para unir a string `"I live in "` com a variável `botLocation`, seguida por um ponto (`.`).

Atribua este valor à variável `botLocationSentence`.

Depois, registre o valor de `botLocationSentence` no console.

# --hints--

Você deve ter uma variável `botLocationSentence`.

```js
assert.isNotNull(botLocationSentence);
```

A variável `botLocationSentence` deve conter o valor de uma string.

```js
assert.isString(botLocationSentence);
```

Você deve usar a concatenação de strings com o operador `+` para unir a string `"I live in "` com a variável `botLocation` seguida por um ponto (`.`). Preste atenção aos espaços.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

Você deve ter uma instrução de `console` no código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

A instrução `console` deve produzir a variável `botLocationSentence`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
