---
id: 66adcf383276814776aba3ca
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Em seguida, crie uma variável chamada `favoriteSubjectSentence`.

Use a concatenação de strings para unir a string `"My favorite subject is "` com a variável `favoriteSubject`, seguida por um ponto.

Atribua o resultado à variável `favoriteSubjectSentence`.

Então, registre o valor de `favoriteSubjectSentence` no `console`.

# --hints--

Você deve ter uma variável chamada `favoriteSubjectSentence`.

```js
assert.isNotNull(favoriteSubjectSentence);
```

A variável `favoriteSubjectSentence` deve ser uma string.

```js
assert.isString(favoriteSubjectSentence);
```

Você deve usar concatenação de strings para unir a string `"My favorite subject is "` com a variável `favoriteSubject`, seguida por um ponto.Cuidado com os espaços.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

Você deve ter uma instrução de `console` no código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

Você deve registrar o valor de `favoriteSubjectSentence` no `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
