---
id: 66ad984e92ce213a66f7e51c
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora, é hora de atribuir um valor à sua variável `bot`.

Nas lições anteriores, você aprendeu como atribuir valores a variáveis assim:

```js
variableName = "Here is the value";
```

Lembre-se: o que está do lado direito do operador de atribuição `=` é o valor que você está atribuindo à variável do lado esquerdo.

Atribua a string `"teacherBot"` à variável `bot`.

# --hints--

Você deve atribuir uma string à variável `bot`.

```js
assert.isString(bot);
```

Você deve atribuir à variável `bot` a string `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
