---
id: 66ad8ab945d266383f318cbf
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Agora você deve ver a primeira mensagem do bot no console.

É hora de adicionar uma segunda mensagem do bot.

Adicione outra instrução de `console.log` para emitir a mensagem `"I am excited to talk to you."` no console.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve ter uma segunda declaração `console.log()` no seu código.

```js
assert.lengthOf(logSpy.calls, 2);
```

A instrução `console` deve ter a mensagem `"I am excited to talk to you."`. Não se esqueça das aspas.

```js
assert.deepInclude(logSpy.calls, ["I am excited to talk to you."]);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
