---
id: 66adc82170acb1464bd348eb
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

As últimas mensagens do bot falarão sobre o assunto favorito dele.

Comece criando uma variável chamada `favoriteSubject` e atribua a ela a string `"Computer Science"`.

# --hints--

Você deve ter uma variável chamada `favoriteSubject`.

```js
assert.isNotNull(favoriteSubject);
```

A variável `favoriteSubject` deve ser uma string.

```js
assert.isString(favoriteSubject);
```

Você deve atribuir a string `"Computer Science"` à variável `favoriteSubject`.

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
