---
id: 66adc42868cab843ccee87d9
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Agora é hora de ver o novo valor do `bot`.

Comece criando uma variável chamada `nicknameIntroduction`.

Utilize a concatenação de string para unir a string `"My nickname is "` com a variável `bot` seguida por um ponto (`.`).

Atribua a string resultante à variável `nicknameIntroduction`.

Em seguida, registre o valor de `nicknameIntroduction` no console. 

# --hints--

Você deve ter uma variável chamada `nicknameIntroduction`.

```js
assert.isNotNull(nicknameIntroduction);
```

A variável `nicknameIntroduction` deve ser uma string.

```js
assert.isString(nicknameIntroduction);
```

Você deve usar a concatenação de strings para unir a string `"My nickname is "` com a variável `bot` seguida por um ponto (`.`).Cuidado com os espaços.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

Você deve ter uma instrução de `console` no código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

Você deve registrar o valor de `nicknameIntroduction` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
