---
id: 66ad8d150264db3926eccfeb
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Nas lições anteriores, você aprendeu sobre a palavra-chave `let` e como declarar variáveis em JavaScript. 

Aqui está um lembrete de como declarar uma variável usando a palavra-chave `let`:

```js
let greeting;
```

Use a palavra-chave `let` para declarar uma variável chamada `bot`.

**NOTA**: Você está usando `let` aqui porque mais adiante no workshop, você estará alterando o valor da variável `bot`.

# --hints--

Você deve ter a palavra-chave `let` no código.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

Você deve declarar uma variável chamada `bot`.

```js
assert.isNotNull(bot);
```

Você deve usar a palavra-chave `let` para declarar a variável.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
