---
id: 66ad984e92ce213a66f7e51c
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Quando você precisa declarar variáveis com múltiplas palavras, você pode usar a convenção de nomenclatura <dfn>camelCase</dfn>.

Ao usar `camelCase`, a primeira palavra é toda em minúscula e a primeira letra de cada palavra subsequente é capitalizada.

Aqui está um exemplo:

```js
let thisIsCamelCase;
```

Use `let` para declarar uma variável chamada `botLocation`.

# --hints--

Você deve usar `let` para declarar a variável.

```js
assert.lengthOf(code.match(/let/g), 2);
```

A variável deve ser nomeada `botLocation`.

```js
assert.isNotNull(botLocation);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
