---
id: 66ad998285df023b23dacdd3
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora, é hora de inicializar a variável `botLocation`.

Quando você precisa declarar variáveis com múltiplas palavras, você pode usar a convenção de nomenclatura <dfn>camelCase</dfn>.

Ao usar `camelCase`, a primeira palavra é toda em minúscula e a primeira letra de cada palavra subsequente é capitalizada. 

Aqui está um exemplo:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

Declare e atribua a string `"the universe"` à variável `botLocation` na mesma linha usando a palavra-chave `let`.

# --hints--

Use a palavra-chave `let` para declarar uma variável.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

A variável deve ser nomeada `botLocation`.

```js
assert.isNotNull(botLocation);
```

Você deve atribuir uma string à variável `botLocation`.

```js
assert.isString(botLocation);
```

Você deve atribuir à variável `botLocation` a string `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
