---
id: 66ad9c633fa26d3c1475eae3
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Agora, é hora de adicionar outra mensagem do bot.

Adicione outra instrução de `console` ao código que registra a mensagem `"Allow me to introduce myself."`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve ter uma terceira declaração `console.log()` no seu código.

```js
assert.lengthOf(logSpy.calls, 3);
```

A instrução de `console` deve ter a mensagem `"Allow me to introduce myself."`. Não se esqueça das aspas.

```js
assert.deepInclude(logSpy.calls, ["Allow me to introduce myself."]);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
