---
id: 669fcb5106eeb13ab7afcd09
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

En este taller, practicará trabajar con HTML semántico construyendo una página de blog dedicada al gato Sr. Whiskers. 

La primera sección que construirá es el <dfn>encabezado</dfn> de la página. 

El elemento `header` se usa para representar contenido introductorio como la navegación de la página y otra información introductoria. 

Aquí hay un ejemplo utilizando el elemento `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Dentro del elemento `body`, agregue un elemento `header`.

# --hints--

Debe tener una etiqueta de apertura `header`.

```js
assert.match(code, /<header>/i);
```

Debe tener una etiqueta de cierre `header`.

```js
assert.match(code, /<\/header>/i);
```

Su elemento `header` debe estar dentro del elemento `body`.

```js
assert.match(code, /<\/header>[.\n\s]*<\/body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  --fcc-editable-region--
  <body>
    
  </body>
  --fcc-editable-region--
</html>
```
