---
id: 669fdfc9a5e3da42d2376609
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Ahora que ha añadido la sección de "about", pruebe haciendo clic en el enlace `About` para ver cómo la página salta hacia esa sección.

La siguiente sección en la página del blog será una lista de publicaciones que hablan sobre el Sr. Whiskers.

Agregue otro elemento `section` con un `id` configurado como `"posts"`.

Dentro del elemento `section`, añada un elemento `h2` con el texto de `Posts`.

# --hints--

Su código debe tener una etiqueta de apertura `section`.

```js
assert.match(code, /<section.*>/i);
```

Su código debe tener una etiqueta de cierre `section`.

```js
assert.match(code, /<\/section>/i);
```

Su elemento `section` debería tener un `id` configurado como `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

Tu elemento `section` debe estar dentro de tu elemento `main`.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Su elemento `h2` debería estar dentro del elemento `section`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

Su elemento `h2` debería tener el texto de `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
