---
id: 669fdc11c9b0263fe0814a7a
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Dentro de cada uno de los elementos `li`, necesitarás tener un elemento de anclaje.

Para el primer elemento de anclaje, el texto debe ser `About` y el valor del atributo `href` debe ser `"#about"`. El símbolo de almohadilla delante de `about` representa un `id` que se añadirá más adelante en el proyecto.

Para el segundo elemento de anclaje, el texto debe ser `Posts` y el valor del atributo `href` debe ser `"#posts"`.

Para el tercer elemento de anclaje, el texto debe ser `Contact` y el valor del atributo `href` debe ser `"#contact"`.

# --hints--

Tu primer elemento de anclaje debe tener el texto `About` dentro de tu primer elemento `li`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

Tu primer elemento de anclaje debe tener un atributo `href` establecido en `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Tu segundo elemento de anclaje debe tener el texto `Posts` dentro de tu segundo elemento `li`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

Tu segundo elemento de anclaje debe tener un atributo `href` establecido en `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Tu tercer elemento de anclaje debe tener el texto `Contact` dentro de tu tercer elemento `li`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

Tu tercer elemento de anclaje debe tener un atributo `href` establecido en `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      <nav>
        <ul>
          <li></li>
          <li></li>
          <li></li>
        </ul>
      </nav>
--fcc-editable-region--
    </header>
  </body>
</html>
```
