---
id: 66a49f685961e997e337cab1
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

El elemento `header` será responsable de mostrar el título principal, la imagen y la navegación de la página del blog.

Dentro del elemento `header`, añade un `h1` con el texto `Welcome to Mr. Whiskers' Blog Page!`.

# --hints--

Debes tener una etiqueta `h1` de apertura.

```js
assert.match(code, /<h1>/i);
```

Debes tener una etiqueta `h1` de cierre.

```js
assert.match(code, /<\/h1>/i);
```

Tu elemento `h1` debe tener el texto `Welcome to Mr. Whiskers' Blog Page!`. Verifica dos veces que no haya errores ortográficos.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

El elemento `h1` debería estar anidado dentro del elemento `header`.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
  </body>
</html>
```
