---
id: 68de40f53fcaf106f51cc20f
title: بناء مشغل صوت وفيديو باستخدام HTML
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر `h1` لعنوان الصفحة الرئيسي.
2. يجب أن يكون لديك عنصران `section`.
3. داخل أول عنصر `section`، يجب أن يكون لديك عنصر `h2` لعنوان الفيديو المشغل.
4. أسفل عنصر `h2`، يجب أن يكون لديك عنصر `video` مع خاصيتي `controls` و`width`. يجب تعيين خاصية `width` إلى `640`.
5. داخل عنصر `video`، يجب أن يكون لديك عنصر `source` مع خاصية `src` تشير إلى ملف فيديو وخاصية `type`.
   - يمكنك استخدام `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.
6. داخل العنصر الثاني `section`، يجب أن يكون لديك عنصر `h2` لعنوان الأغنية المشغلة.
7. أسفل عنصر `h2`، يجب أن يكون لديك عنصر `audio` مع خاصيتي `controls` و`loop`، وخاصية `src` تشير إلى ملف صوتي.
   - يمكنك استخدام `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`.
   - أو `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`.

# --hints--

يجب أن يكون لديك عنصر `h1` لعنوان الصفحة الرئيسي.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

يجب ألا يكون عنصر `h1` فارغًا.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

يجب أن يكون لديك عنصر `h1` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

يجب أن يكون لديك عنصران من نوع `section`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

داخل أول عنصر `section`، يجب أن يكون لديك عنصر `h2` لعنوان الفيديو المشغل.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

أسفل عنصر `h2`، يجب أن يكون لديك عنصر `video`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

يجب أن يحتوي عنصر `video` على خاصية `controls`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

يجب أن يحتوي عنصر `video` على خاصية `width` مع القيمة `640`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

داخل عنصر `video`، يجب أن يكون لديك عنصر `source`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

يجب أن يحتوي عنصر `source` على خاصية `src` تشير إلى ملف فيديو.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

يجب أن يحتوي عنصر `source` على خاصية `type`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

يجب أن تتطابق خاصية `type` مع نوع الوسائط لامتداد ملف خاصية `src`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

داخل العنصر الثاني `section`، يجب أن يكون لديك عنصر `h2` لعنوان الأغنية المشغلة.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

يجب أن تحتوي جميع عناصر `h2` على نص ما.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

أسفل عنصر `h2`، يجب أن يكون لديك عنصر `audio`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

يجب أن يحتوي عنصر `audio` على خاصية `controls`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

يجب أن يحتوي عنصر `audio` على خاصية `loop`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

يجب أن يحتوي عنصر `audio` على خاصية `src` تشير إلى ملف صوتي.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
