---
id: 6ac0a4463662d9738c393e21
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

У попередньому кроці ви створили **рядок**: послідовність символів, укладених у одинарні або подвійні лапки. У програмуванні рядки використовують, щоб представити текст.

**Термінал** — це місце, де ви можете побачити вихідні дані свого коду.

Щоб побачити вихідні дані, перетворіть свій існуючий рядок коду на:

```py
print('Hello')
```

Переконайтеся, що починаєте з початку лінії.

# --hints--

У вашому коді має бути `print('Hello')`. Переконайтеся, що починаєте з початку лінії.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
