---
id: 6ac0b35f3b29a27f7e207295
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

У попередньому кроці ви створили змінну `sentence`. Тепер виведіть її значення в термінал, набравши `print(sentence)` на виділеній лінії.

Таким чином, ви завершили свою першу програму на Python.

# --hints--

У вашому коді має бути `print(sentence)`. Переконайтеся, що починаєте з початку лінії.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
