---
id: 6ac09f8420b4b67e075d2cd2
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите свою першу програму на Python.

Це ваш редактор коду, де ви писатимете код на Python. Лінія 1 виділена. Перейдіть туди і введіть:

```py
'Hello'
```

Переконайтеся, що починаєте з початку лінії і не забувайте лапки.

Коли закінчите, натисніть на кнопку «перевірити свій код», щоб дізнатися, чи він правильний.

# --hints--

Ви повинні ввести `'Hello'`. Переконайтеся, що починаєте з початку лінії і не забувайте лапки.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
