---
id: 6ac0ae563f6de70930f47ef7
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Пишучи код, ви часто захочете повторно використовувати або посилатися на існуючі значення у своєму коді. Змінні дозволяють це робити.

Щоб призначити значення змінній, напишіть ім’я змінної, потім оператор присвоєння (`=`) і значення, яке хочете призначити цій змінній:

```py
sentence = 'I am learning Python'
```

Перейдіть до виділеної лінії, створіть змінну з ім’ям `sentence` і призначте їй рядок `I am learning Python`.

# --hints--

У вашому коді має бути `sentence = 'I am learning Python'`. Переконайтеся, що починаєте з початку лінії.

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
