---
id: 6ac0aaba914d81b32567c076
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` — це вбудована функція, яку ви можете використовувати, щоб відправити вихідні дані до терміналу. Як ви бачите, текст `Hello` тепер видно в терміналі.

Тепер виведіть текст `freeCodeCamp` у термінал, перейшовши до виділеної лінії та ввівши:

```py
print('freeCodeCamp')
```

Переконайтеся, що починаєте з початку лінії.

# --hints--

У вашому коді має бути `print('freeCodeCamp')`. Переконайтеся, що починаєте з початку лінії.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
