---
id: 6989f2e35e6941e13e445156
title: passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Agora você vai implementar o loop principal do BFS. Crie um loop `while` que continue enquanto o `queue` não estiver vazio, ou seja, enquanto `queue.length` for maior que `0`.

Dentro do loop, registre `queue` no console.

# --hints--

Você deve criar um loop `while` com a condição `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

Você deve registrar `queue` no console dentro do seu loop `while`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
