---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Agora você precisa lidar com a adição dos parênteses de fechamento. A regra principal é que você só pode adicionar um parêntese de fechamento se isso mantiver o equilíbrio, ou seja, deve haver mais parênteses de abertura usados do que parênteses de fechamento.

Ainda dentro do bloco `else`, adicione uma segunda declaração `if` para verificar se `closesUsed` é menor que `opensUsed`.

Se essa condição for verdadeira, adicione outro novo array ao `queue`: `[current + ')', opensUsed, closesUsed + 1]`. Isso representa o estado após adicionar um parêntese de fechamento.

# --hints--

Você deve ter uma segunda declaração `if` com a condição `closesUsed < opensUsed` dentro do seu bloco `else`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

Você deve adicionar `[current + ')', opensUsed, closesUsed + 1]` à fila quando a condição for verdadeira.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
