---
id: 6989f48c3f0292c84f8d8938
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Agora você precisa identificar quando construiu uma combinação completa de parênteses. Uma combinação completa tem comprimento igual ao dobro do número de pares (já que cada par contribui com um parêntese de abertura e um de fechamento).

Dentro do seu loop `while`, adicione uma declaração `if` para verificar se `current.length === 2 * pairs`. Quando essa condição for verdadeira, adicione `current` ao array `result`.

# --hints--

Você deve ter uma declaração `if` com a condição `current.length === 2 * pairs` dentro do seu loop `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

Você deve adicionar `current` ao array `result` dentro da sua nova declaração `if`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
