---
id: 6989ef885db3a3488a13b426
title: passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Antes de implementar o algoritmo principal, você precisa validar a entrada. O parâmetro `pairs` deve ser um inteiro, pois representa o número de pares de parênteses a gerar.

Adicione uma declaração `if` no início da sua função para verificar se `pairs` não é um inteiro. Use o método `Number.isInteger()` para isso.

Se a condição for verdadeira, retorne a string `The number of pairs should be an integer`.

# --hints--

Você deve ter uma instrução if com a condição `!Number.isInteger(pairs)` dentro da função `genParentheses`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

Você deve retornar uma string quando a entrada não for um inteiro.

```js
let result = genParentheses("i");
assert.isString(result);
```

Você deve retornar a string `The number of pairs should be an integer` quando a entrada for uma string.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
