---
id: 6989f555728e63ad8b349d84
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Se a string atual ainda não estiver completa, você precisa explorar os próximos estados possíveis. Adicione uma cláusula `else` para a declaração `if` atual.

Dentro do bloco `else`, você vai lidar com a adição de parênteses de abertura. Adicione uma declaração `if` para verificar se `opensUsed` é menor que `pairs`. Isso garante que você não use mais parênteses de abertura do que o permitido.

Se essa condição for `true`, insira um novo array na `queue`: `[current + '(', opensUsed + 1, closesUsed]`. Isso representa o estado após adicionar um parêntese de abertura.

# --hints--

Você deve ter uma cláusula `else` para sua declaração `if` dentro do loop `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

Você deve ter uma declaração `if` dentro do bloco `else` com a condição `opensUsed < pairs`.

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

Você deve inserir `[current + '(', opensUsed + 1, closesUsed]` na fila quando a condição for `true`.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
