---
id: 6989eb65324d97775c4965fe
title: passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai implementar uma função que gera todas as combinações válidas de parênteses usando uma abordagem de busca em largura (BFS). Por exemplo, as combinações válidas de dois pares de parênteses são `(())` e `()()`.

Comece criando uma função chamada `genParentheses` com um único parâmetro `pairs`. Por enquanto, retorne um array vazio da função.

# --hints--

Você deve definir uma função chamada `genParentheses`.

```js
assert.isFunction(genParentheses);
```

Sua função `genParentheses` deve ter um único parâmetro chamado `pairs`.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

Sua função deve retornar um array vazio.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
