---
id: 6989f1ba358acba651d5ee6b
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora você vai configurar a estrutura de dados para armazenar seus resultados. Crie uma variável chamada `result` e inicialize-a com um array vazio. Esse array vai armazenar todas as combinações válidas de parênteses que você gerar.

Atualize sua instrução return para retornar `result` em vez de um array vazio.

# --hints--

Você deve declarar uma variável chamada `result` na sua função `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.result);
```

Você deve inicializar `result` com um array vazio.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { result } = genParentheses?.variables ?? {};
assert.isFalse(result?.value.isEmpty());
const resultValue = eval(result?.value.toString());
assert.isArray(resultValue);
assert.isEmpty(resultValue);
```

Você deve retornar `result` no final da sua função.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.isTrue(genParentheses?.hasReturn('result'));
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
