---
id: 6989f10a72ad424d4ad6d514
title: passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Em seguida, você precisa validar que o número de pares é pelo menos um, já que não é possível gerar combinações de parênteses com zero ou pares negativos.

Adicione outra declaração `if` para verificar se `pairs` é menor que `1`. Se essa condição for verdadeira, retorne a string `The number of pairs should be at least 1`.

# --hints--

Você deve ter uma segunda declaração `if` com a condição `pairs < 1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

Você deve retornar a string `The number of pairs should be at least 1` quando `pairs` for menor que `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
