---
id: 6989f3632f8b8bdf168aa321
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Dentro do seu `while` loop, use `queue.shift()` para remover e obter o primeiro elemento da fila. Isso implementa o comportamento primeiro a entrar, primeiro a sair (FIFO) característico do BFS.

Desestruture esse array em três variáveis: `current`, `opensUsed` e `closesUsed`.

# --hints--

Você deve usar `queue.shift()` para obter o primeiro elemento da fila.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

Você deve desestruturar o resultado em três variáveis: `current`, `opensUsed` e `closesUsed`.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return result;
}
```
