---
id: 6989f207306332c06b0add39
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Para a abordagem de busca em largura, você vai usar uma fila para acompanhar diferentes estados enquanto constrói as combinações de parênteses. Cada estado será representado como um array contendo três elementos:

- A string atual que está sendo construída
- O número de parênteses de abertura usados até agora
- O número de parênteses de fechamento usados até agora

Crie uma variável chamada `queue` e inicialize-a com um array contendo um array: `['', 0, 0]`. Isso representa o estado inicial com uma string vazia e zero parênteses usados.

# --hints--

Você deve declarar uma variável chamada `queue` na sua função `genParentheses`.

```js
assert.match(genParentheses.toString(), /queue/);
```

Você deve inicializar `queue` com um array contendo o array `['', 0, 0]`.

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  return result;
}
```
