---
id: 6989eb65324d97775c4965fe
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você implementará uma função que gera todas as combinações válidas de parênteses usando uma abordagem de busca em largura (BFS). Por exemplo, as combinações válidas de dois pares de parênteses são `(())` e `()()`.

Comece criando uma função chamada `genParentheses` com um único parâmetro `pairs`. Por enquanto, retorne um array vazio da função.

# --hints--

Você deve definir uma função chamada `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
```

Sua função `genParentheses` deve ter um único parâmetro chamado `pairs`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.equal(genParentheses?.parameters.length, 1);
assert.equal(genParentheses?.parameters[0].toString(), 'pairs');
```

Sua função deve retornar um array vazio.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
