---
id: 6989f1ba358acba651d5ee6b
title: passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora você vai configurar a estrutura de dados para armazenar seus resultados. Crie uma variável chamada `result` e inicialize-a com um array vazio. Esse array vai armazenar todas as combinações válidas de parênteses que você gerar.

Atualize sua instrução return para retornar `result` em vez de um array vazio.

# --hints--

Você deve declarar uma variável chamada `result` na sua função `genParentheses`.

```js
assert.match(genParentheses.toString(), /result/);
```

Você deve inicializar `result` com um array vazio.

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

Você deve retornar `result` no final da sua função.

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
